# Command: eprover --tstp-format --proof-object --output-level=4 simple.in
fof(c_0_1, axiom, (![X1]:f(X1,g(X1))=g(X1)), file('simple.in', clause_1)).
fof(c_0_2, axiom, (![X2]:![X3]:h(f(g(X2),X3))=g(X3)), file('simple.in', clause_1)).
fof(c_0_3, conjecture, (h(g(g(c)))=g(g(g(c)))), file('simple.in', clause_1)).
fof(c_0_4, negated_conjecture, (~(h(g(g(c)))=g(g(g(c))))),inference(assume_negation, [status(cth)],[c_0_3])).
fof(c_0_5, negated_conjecture, (h(g(g(c)))!=g(g(g(c)))),inference(fof_simplification, [status(thm)],[c_0_4])).
fof(c_0_6, plain, (![X2]:f(X2,g(X2))=g(X2)),inference(variable_rename, [status(thm)],[c_0_1])).
cnf(c_0_7,plain,(f(X1,g(X1))=g(X1)),inference(split_conjunct, [status(thm)],[c_0_6])).
fof(c_0_8, plain, (![X4]:![X5]:h(f(g(X4),X5))=g(X5)),inference(variable_rename, [status(thm)],[c_0_2])).
cnf(c_0_9,plain,(h(f(g(X1),X2))=g(X2)),inference(split_conjunct, [status(thm)],[c_0_8])).
cnf(c_0_10,negated_conjecture,(h(g(g(c)))!=g(g(g(c)))),inference(split_conjunct, [status(thm)],[c_0_5])).
# Initializing proof state
# Scanning for AC axioms
cnf(c_0_11,plain,(h(g(g(X1)))=g(g(g(X1)))),inference(pm,[status(thm)],[c_0_9,c_0_7])).
cnf(c_0_12,plain,(f(g(g(X1)),h(g(g(X1))))=g(g(g(X1)))),inference(pm,[status(thm)],[c_0_7,c_0_11])).
cnf(c_0_13,plain,(h(f(h(g(g(X1))),X2))=g(X2)),inference(pm,[status(thm)],[c_0_9,c_0_11])).
cnf(c_0_14,plain,(g(g(h(g(g(X1)))))=h(g(g(g(g(X1)))))),inference(pm,[status(thm)],[c_0_11,c_0_11])).
cnf(c_0_15,plain,(g(h(g(g(X1))))=h(g(g(g(X1))))),inference(pm,[status(thm)],[c_0_11,c_0_11])).
cnf(c_0_16,plain,(f(g(g(X1)),h(g(g(X1))))=h(g(g(X1)))),inference(rw, [status(thm)],[c_0_12,c_0_11])).
cnf(c_0_17,plain,(g(g(h(g(g(X1)))))=h(g(h(g(g(X1)))))),inference(rw, [status(thm)],[c_0_14,c_0_11])).
cnf(c_0_18,plain,(g(h(g(g(X1))))=h(h(g(g(X1))))),inference(rw, [status(thm)],[c_0_15,c_0_11])).
cnf(c_0_19,negated_conjecture,($false),inference(rw, [status(thm)],[c_0_10,c_0_11])).
cnf(c_0_20,negated_conjecture,($false),inference(cn,[status(thm)],[c_0_19])).
cnf(c_0_21,negated_conjecture,($false), c_0_20,[proof]).

