fof(c_0_1, axiom, (![X1]:![X2]:![X3]:![X4]:(accept_team(X1,X4,X2,X3)<=>((accept_city(X1,X2)&accept_leader(X1,X4))&accept_number(X1,X3)))), file('Axioms/AGT001+0.ax', a1_1)).
fof(c_0_2, axiom, (![X1]:![X3]:![X5]:((accept_number(X1,X3)&less(X5,X3))=>accept_number(X1,X5))), file('Axioms/AGT001+0.ax', a1_2)).
fof(c_0_3, axiom, (![X1]:![X3]:![X5]:![X6]:((accept_population(X1,X6,X3)&less(X5,X3))=>accept_population(X1,X6,X5))), file('Axioms/AGT001+0.ax', a1_3)).
fof(c_0_4, axiom, (![X1]:![X4]:![X2]:(the_agent_in_all_proposed_teams(X1,X4,X2)=>(accept_leader(X1,X4)&accept_city(X1,X2)))), file('Axioms/AGT001+0.ax', a1_4)).
fof(c_0_5, axiom, (![X1]:![X4]:![X2]:(any_agent_in_all_proposed_teams(X1,X4,X2)=>accept_leader(X1,X4))), file('Axioms/AGT001+0.ax', a1_5)).
fof(c_0_6, axiom, (![X1]:![X4]:![X2]:(the_agent_not_in_any_proposed_teams(X1,X4,X2)=>~((accept_city(X1,X2)&accept_leader(X1,X4))))), file('Axioms/AGT001+0.ax', a1_6)).
fof(c_0_7, axiom, (![X1]:![X3]:(min_number_of_proposed_agents(X1,X3)=>accept_number(X1,X3))), file('Axioms/AGT001+0.ax', a1_7)).
fof(c_0_8, axiom, (![X1]:(((((accept_population(X1,atheist,n65)&accept_population(X1,christian,n20))&accept_population(X1,muslim,n7))&accept_population(X1,native,n4))&accept_population(X1,other,n4))<=>accept_city(X1,suffertown))), file('Axioms/AGT001+0.ax', a2_1)).
fof(c_0_9, axiom, (![X1]:(((((accept_population(X1,atheist,n54)&accept_population(X1,christian,n23))&accept_population(X1,muslim,n3))&accept_population(X1,native,n1))&accept_population(X1,other,n9))<=>accept_city(X1,centraltown))), file('Axioms/AGT001+0.ax', a2_2)).
fof(c_0_10, axiom, (![X1]:(((((accept_population(X1,atheist,n30)&accept_population(X1,christian,n8))&accept_population(X1,muslim,n60))&accept_population(X1,native,n1))&accept_population(X1,other,n1))<=>accept_city(X1,sunnysideport))), file('Axioms/AGT001+0.ax', a2_3)).
fof(c_0_11, axiom, (![X1]:(((((accept_population(X1,atheist,n70)&accept_population(X1,christian,n15))&accept_population(X1,muslim,n1))&accept_population(X1,native,n10))&accept_population(X1,other,n4))<=>accept_city(X1,centrallakecity))), file('Axioms/AGT001+0.ax', a2_4)).
fof(c_0_12, axiom, (![X1]:(((((accept_population(X1,atheist,n68)&accept_population(X1,christian,n16))&accept_population(X1,muslim,n1))&accept_population(X1,native,n11))&accept_population(X1,other,n4))<=>accept_city(X1,stjosephburgh))), file('Axioms/AGT001+0.ax', a2_5)).
fof(c_0_13, axiom, (![X1]:(((((accept_population(X1,atheist,n70)&accept_population(X1,christian,n13))&accept_population(X1,muslim,n0))&accept_population(X1,native,n15))&accept_population(X1,other,n2))<=>accept_city(X1,northport))), file('Axioms/AGT001+0.ax', a2_6)).
fof(c_0_14, axiom, (![X1]:(((((accept_population(X1,atheist,n12)&accept_population(X1,christian,n3))&accept_population(X1,muslim,n0))&accept_population(X1,native,n85))&accept_population(X1,other,n0))<=>accept_city(X1,coastvillage))), file('Axioms/AGT001+0.ax', a2_7)).
fof(c_0_15, axiom, (![X1]:(((((accept_population(X1,atheist,n0)&accept_population(X1,christian,n0))&accept_population(X1,muslim,n0))&accept_population(X1,native,n100))&accept_population(X1,other,n0))<=>accept_city(X1,sunsetpoint))), file('Axioms/AGT001+0.ax', a2_8)).
fof(c_0_16, axiom, (![X1]:(((((accept_population(X1,atheist,n75)&accept_population(X1,christian,n24))&accept_population(X1,muslim,n1))&accept_population(X1,native,n0))&accept_population(X1,other,n0))<=>accept_city(X1,towna))), file('Axioms/AGT001+0.ax', a2_9)).
fof(c_0_17, axiom, (![X1]:(((((accept_population(X1,atheist,n75)&accept_population(X1,christian,n25))&accept_population(X1,muslim,n0))&accept_population(X1,native,n0))&accept_population(X1,other,n0))<=>accept_city(X1,citya))), file('Axioms/AGT001+0.ax', a2_10)).
fof(c_0_18, axiom, (![X1]:(((((accept_population(X1,atheist,n70)&accept_population(X1,christian,n20))&accept_population(X1,muslim,n8))&accept_population(X1,native,n0))&accept_population(X1,other,n2))<=>accept_city(X1,townb))), file('Axioms/AGT001+0.ax', a2_11)).
fof(c_0_19, axiom, (![X1]:(((((accept_population(X1,atheist,n78)&accept_population(X1,christian,n20))&accept_population(X1,muslim,n1))&accept_population(X1,native,n0))&accept_population(X1,other,n1))<=>accept_city(X1,cityb))), file('Axioms/AGT001+0.ax', a2_12)).
fof(c_0_20, axiom, (![X1]:(((((accept_population(X1,atheist,n30)&accept_population(X1,christian,n0))&accept_population(X1,muslim,n65))&accept_population(X1,native,n0))&accept_population(X1,other,n5))<=>accept_city(X1,townc))), file('Axioms/AGT001+0.ax', a2_13)).
fof(c_0_21, axiom, (accept_team(countryahumanitarianorganization,countryacivilorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_1)).
fof(c_0_22, axiom, (accept_team(countryahumanitarianorganization,countryacivilorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_2)).
fof(c_0_23, axiom, (accept_team(countryahumanitarianorganization,countryacivilorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_3)).
fof(c_0_24, axiom, (accept_team(countryahumanitarianorganization,countryafirstaidorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_4)).
fof(c_0_25, axiom, (the_agent_in_all_proposed_teams(countryahumanitarianorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_5)).
fof(c_0_26, axiom, (any_agent_in_all_proposed_teams(countryahumanitarianorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_6)).
fof(c_0_27, axiom, (accept_team(countryccivilorganization,countrybhumanitarianorganization,cityb,n2)), file('Axioms/AGT001+1.ax', event_7)).
fof(c_0_28, axiom, (accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,cityb,n4)), file('Axioms/AGT001+1.ax', event_8)).
fof(c_0_29, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n4))), file('Axioms/AGT001+1.ax', event_9)).
fof(c_0_30, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n5))), file('Axioms/AGT001+1.ax', event_10)).
fof(c_0_31, axiom, (the_agent_not_in_any_proposed_teams(muslimcountrybhumanitarianorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_11)).
fof(c_0_32, axiom, (any_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_12)).
fof(c_0_33, axiom, (accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,coastvillage,n5)), file('Axioms/AGT001+1.ax', event_13)).
fof(c_0_34, axiom, (accept_team(countryamedicalorganization,countryacivilorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_14)).
fof(c_0_35, axiom, (accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,cityb,n4)), file('Axioms/AGT001+1.ax', event_15)).
fof(c_0_36, axiom, (accept_team(countryafirstaidorganization,countryacivilorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_16)).
fof(c_0_37, axiom, (accept_number(countrybhumanitarianorganization,n4)), file('Axioms/AGT001+1.ax', event_17)).
fof(c_0_38, axiom, (accept_team(countrybcivilorganization,countrybhumanitarianorganization,cityb,n4)), file('Axioms/AGT001+1.ax', event_18)).
fof(c_0_39, axiom, (accept_team(christiansufferterrahumanitarianorganization,countryacivilorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_19)).
fof(c_0_40, axiom, (accept_team(sufferterragovernment,countryacivilorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_20)).
fof(c_0_41, axiom, (accept_number(countrybhumanitarianorganization,n1)), file('Axioms/AGT001+1.ax', event_21)).
fof(c_0_42, axiom, (the_agent_in_all_proposed_teams(countrybhumanitarianorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_22)).
fof(c_0_43, axiom, (any_agent_in_all_proposed_teams(countrybhumanitarianorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_23)).
fof(c_0_44, axiom, (accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,coastvillage,n5)), file('Axioms/AGT001+1.ax', event_24)).
fof(c_0_45, axiom, (accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_25)).
fof(c_0_46, axiom, (accept_team(christiansufferterrahumanitarianorganization,countryacivilorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_26)).
fof(c_0_47, axiom, (accept_team(countryamedicalorganization,countryacivilorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_27)).
fof(c_0_48, axiom, (accept_number(countryahumanitarianorganization,n2)), file('Axioms/AGT001+1.ax', event_28)).
fof(c_0_49, axiom, (~(accept_team(countryamedicalorganization,countryahumanitarianorganization,coastvillage,n2))), file('Axioms/AGT001+1.ax', event_29)).
fof(c_0_50, axiom, (~(accept_team(countryamedicalorganization,countryafirstaidorganization,coastvillage,n2))), file('Axioms/AGT001+1.ax', event_30)).
fof(c_0_51, axiom, (~(accept_team(countryamedicalorganization,countryacivilorganization,coastvillage,n2))), file('Axioms/AGT001+1.ax', event_31)).
fof(c_0_52, axiom, (~(accept_team(countryamedicalorganization,christiansufferterrahumanitarianorganization,coastvillage,n2))), file('Axioms/AGT001+1.ax', event_32)).
fof(c_0_53, axiom, (~(accept_team(countryamedicalorganization,sufferterragovernment,coastvillage,n2))), file('Axioms/AGT001+1.ax', event_33)).
fof(c_0_54, axiom, (~(accept_team(countryamedicalorganization,countryahumanitarianorganization,coastvillage,n6))), file('Axioms/AGT001+1.ax', event_34)).
fof(c_0_55, axiom, (~(accept_team(countryamedicalorganization,countryacivilorganization,coastvillage,n6))), file('Axioms/AGT001+1.ax', event_35)).
fof(c_0_56, axiom, (~(accept_team(countryamedicalorganization,countryafirstaidorganization,coastvillage,n6))), file('Axioms/AGT001+1.ax', event_36)).
fof(c_0_57, axiom, (the_agent_not_in_any_proposed_teams(countryamedicalorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_37)).
fof(c_0_58, axiom, (any_agent_in_all_proposed_teams(countryamedicalorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_38)).
fof(c_0_59, axiom, (accept_team(countryccivilorganization,countrycmedicalorganization,towna,n4)), file('Axioms/AGT001+1.ax', event_39)).
fof(c_0_60, axiom, (~(accept_team(countryccivilorganization,countrycmedicalorganization,towna,n5))), file('Axioms/AGT001+1.ax', event_40)).
fof(c_0_61, axiom, (the_agent_in_all_proposed_teams(countryccivilorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_41)).
fof(c_0_62, axiom, (any_agent_in_all_proposed_teams(countryccivilorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_42)).
fof(c_0_63, axiom, (accept_team(countrybcivilorganization,countrycmedicalorganization,towna,n4)), file('Axioms/AGT001+1.ax', event_43)).
fof(c_0_64, axiom, (accept_team(countrybcivilorganization,countrycmedicalorganization,towna,n5)), file('Axioms/AGT001+1.ax', event_44)).
fof(c_0_65, axiom, (accept_number(countrybcivilorganization,n5)), file('Axioms/AGT001+1.ax', event_45)).
fof(c_0_66, axiom, (the_agent_in_all_proposed_teams(countrybcivilorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_46)).
fof(c_0_67, axiom, (any_agent_in_all_proposed_teams(countrybcivilorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_47)).
fof(c_0_68, axiom, (accept_team(sufferterragovernment,countryacivilorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_48)).
fof(c_0_69, axiom, (the_agent_in_all_proposed_teams(countrycmedicalorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_49)).
fof(c_0_70, axiom, (any_agent_in_all_proposed_teams(countrycmedicalorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_50)).
fof(c_0_71, axiom, (accept_team(countrycmedicalorganization,christiancountrychumanitarianorganization,coastvillage,n5)), file('Axioms/AGT001+1.ax', event_51)).
fof(c_0_72, axiom, (accept_number(countrycmedicalorganization,n4)), file('Axioms/AGT001+1.ax', event_52)).
fof(c_0_73, axiom, (accept_team(christiancountrychumanitarianorganization,countrycmedicalorganization,towna,n4)), file('Axioms/AGT001+1.ax', event_53)).
fof(c_0_74, axiom, (accept_number(countrycmedicalorganization,n5)), file('Axioms/AGT001+1.ax', event_54)).
fof(c_0_75, axiom, (accept_team(christiancountrychumanitarianorganization,countrycmedicalorganization,towna,n5)), file('Axioms/AGT001+1.ax', event_55)).
fof(c_0_76, axiom, (the_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_56)).
fof(c_0_77, axiom, (any_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,countryacivilorganization,towna)), file('Axioms/AGT001+1.ax', event_57)).
fof(c_0_78, axiom, (accept_team(countryafirstaidorganization,countryacivilorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_58)).
fof(c_0_79, axiom, (accept_number(countryacivilorganization,n2)), file('Axioms/AGT001+1.ax', event_59)).
fof(c_0_80, axiom, (accept_team(countryacivilorganization,countryahumanitarianorganization,coastvillage,n5)), file('Axioms/AGT001+1.ax', event_60)).
fof(c_0_81, axiom, (accept_team(countryacivilorganization,countryahumanitarianorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_61)).
fof(c_0_82, axiom, (accept_team(countryacivilorganization,countryafirstaidorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_62)).
fof(c_0_83, axiom, (the_agent_in_all_proposed_teams(countryacivilorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_63)).
fof(c_0_84, axiom, (any_agent_in_all_proposed_teams(countryacivilorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_64)).
fof(c_0_85, axiom, (accept_number(countryafirstaidorganization,n2)), file('Axioms/AGT001+1.ax', event_65)).
fof(c_0_86, axiom, (accept_team(countryafirstaidorganization,countryahumanitarianorganization,coastvillage,n5)), file('Axioms/AGT001+1.ax', event_66)).
fof(c_0_87, axiom, (accept_team(countryafirstaidorganization,countryahumanitarianorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_67)).
fof(c_0_88, axiom, (accept_team(countryafirstaidorganization,countryacivilorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_68)).
fof(c_0_89, axiom, (the_agent_in_all_proposed_teams(countryafirstaidorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_69)).
fof(c_0_90, axiom, (any_agent_in_all_proposed_teams(countryafirstaidorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_70)).
fof(c_0_91, axiom, (~(accept_team(countryccivilorganization,christiancountrychumanitarianorganization,coastvillage,n5))), file('Axioms/AGT001+1.ax', event_71)).
fof(c_0_92, axiom, (~(accept_team(countryccivilorganization,christiancountrychumanitarianorganization,coastvillage,n6))), file('Axioms/AGT001+1.ax', event_72)).
fof(c_0_93, axiom, (accept_number(christiancountrychumanitarianorganization,n5)), file('Axioms/AGT001+1.ax', event_73)).
fof(c_0_94, axiom, (accept_team(countrybcivilorganization,christiancountrychumanitarianorganization,coastvillage,n5)), file('Axioms/AGT001+1.ax', event_74)).
fof(c_0_95, axiom, (accept_team(countrybcivilorganization,christiancountrychumanitarianorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_75)).
fof(c_0_96, axiom, (accept_team(countrybcivilorganization,muslimcountrybhumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_76)).
fof(c_0_97, axiom, (accept_team(countrybcivilorganization,countrybhumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_77)).
fof(c_0_98, axiom, (accept_team(countrybcivilorganization,christiancountrychumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_78)).
fof(c_0_99, axiom, (accept_team(countrybcivilorganization,countrycmedicalorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_79)).
fof(c_0_100, axiom, (the_agent_in_all_proposed_teams(countrybcivilorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_80)).
fof(c_0_101, axiom, (any_agent_in_all_proposed_teams(countrybcivilorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_81)).
fof(c_0_102, axiom, (accept_team(sufferterragovernment,countryahumanitarianorganization,coastvillage,n5)), file('Axioms/AGT001+1.ax', event_82)).
fof(c_0_103, axiom, (accept_team(sufferterragovernment,countryahumanitarianorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_83)).
fof(c_0_104, axiom, (accept_team(sufferterragovernment,countryacivilorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_84)).
fof(c_0_105, axiom, (accept_team(sufferterragovernment,countryafirstaidorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_85)).
fof(c_0_106, axiom, (the_agent_in_all_proposed_teams(sufferterragovernment,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_86)).
fof(c_0_107, axiom, (any_agent_in_all_proposed_teams(sufferterragovernment,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_87)).
fof(c_0_108, axiom, (accept_number(christiansufferterrahumanitarianorganization,n2)), file('Axioms/AGT001+1.ax', event_88)).
fof(c_0_109, axiom, (accept_number(countryahumanitarianorganization,n5)), file('Axioms/AGT001+1.ax', event_89)).
fof(c_0_110, axiom, (accept_team(christiansufferterrahumanitarianorganization,countryahumanitarianorganization,coastvillage,n5)), file('Axioms/AGT001+1.ax', event_90)).
fof(c_0_111, axiom, (accept_team(christiansufferterrahumanitarianorganization,countryahumanitarianorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_91)).
fof(c_0_112, axiom, (accept_team(christiansufferterrahumanitarianorganization,countryacivilorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_92)).
fof(c_0_113, axiom, (accept_team(christiansufferterrahumanitarianorganization,countryafirstaidorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_93)).
fof(c_0_114, axiom, (the_agent_in_all_proposed_teams(christiansufferterrahumanitarianorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_94)).
fof(c_0_115, axiom, (any_agent_in_all_proposed_teams(christiansufferterrahumanitarianorganization,christiancountrychumanitarianorganization,coastvillage)), file('Axioms/AGT001+1.ax', event_95)).
fof(c_0_116, axiom, (accept_team(countrycmedicalorganization,christiancountrychumanitarianorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_96)).
fof(c_0_117, axiom, (accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,coastvillage,n6)), file('Axioms/AGT001+1.ax', event_97)).
fof(c_0_118, axiom, (accept_number(muslimcountrybhumanitarianorganization,n3)), file('Axioms/AGT001+1.ax', event_98)).
fof(c_0_119, axiom, (accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,townc,n4)), file('Axioms/AGT001+1.ax', event_99)).
fof(c_0_120, axiom, (accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_100)).
fof(c_0_121, axiom, (accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_101)).
fof(c_0_122, axiom, (accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_102)).
fof(c_0_123, axiom, (accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_103)).
fof(c_0_124, axiom, (the_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_104)).
fof(c_0_125, axiom, (any_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_105)).
fof(c_0_126, axiom, (accept_team(countryacivilorganization,countryahumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_106)).
fof(c_0_127, axiom, (accept_team(countryafirstaidorganization,countryahumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_107)).
fof(c_0_128, axiom, (accept_team(christiancountrychumanitarianorganization,countryccivilorganization,townc,n4)), file('Axioms/AGT001+1.ax', event_108)).
fof(c_0_129, axiom, (accept_team(christiancountrychumanitarianorganization,muslimcountrybhumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_109)).
fof(c_0_130, axiom, (accept_team(christiancountrychumanitarianorganization,countrybcivilorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_110)).
fof(c_0_131, axiom, (accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_111)).
fof(c_0_132, axiom, (accept_team(christiancountrychumanitarianorganization,countrycmedicalorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_112)).
fof(c_0_133, axiom, (the_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_113)).
fof(c_0_134, axiom, (any_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_114)).
fof(c_0_135, axiom, (~(accept_team(countryccivilorganization,muslimcountrybhumanitarianorganization,townc,n6))), file('Axioms/AGT001+1.ax', event_115)).
fof(c_0_136, axiom, (~(accept_team(countryccivilorganization,countrybcivilorganization,townc,n6))), file('Axioms/AGT001+1.ax', event_116)).
fof(c_0_137, axiom, (~(accept_team(countryccivilorganization,countrybhumanitarianorganization,townc,n6))), file('Axioms/AGT001+1.ax', event_117)).
fof(c_0_138, axiom, (~(accept_team(countryccivilorganization,christiancountrychumanitarianorganization,townc,n6))), file('Axioms/AGT001+1.ax', event_118)).
fof(c_0_139, axiom, (~(accept_team(countryccivilorganization,countrycmedicalorganization,townc,n6))), file('Axioms/AGT001+1.ax', event_119)).
fof(c_0_140, axiom, (the_agent_in_all_proposed_teams(countryccivilorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_120)).
fof(c_0_141, axiom, (any_agent_in_all_proposed_teams(countryccivilorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_121)).
fof(c_0_142, axiom, (accept_team(countrybhumanitarianorganization,countrybcivilorganization,townc,n2)), file('Axioms/AGT001+1.ax', event_122)).
fof(c_0_143, axiom, (accept_team(countrybhumanitarianorganization,muslimcountrybhumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_123)).
fof(c_0_144, axiom, (accept_team(sufferterragovernment,countryahumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_124)).
fof(c_0_145, axiom, (accept_team(christiansufferterrahumanitarianorganization,countryahumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_125)).
fof(c_0_146, axiom, (accept_team(countryamedicalorganization,countryahumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_126)).
fof(c_0_147, axiom, (accept_team(countrycmedicalorganization,countryccivilorganization,townc,n4)), file('Axioms/AGT001+1.ax', event_127)).
fof(c_0_148, axiom, (accept_team(countrycmedicalorganization,muslimcountrybhumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_128)).
fof(c_0_149, axiom, (accept_team(countrycmedicalorganization,countrybcivilorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_129)).
fof(c_0_150, axiom, (accept_team(countrycmedicalorganization,countrybhumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_130)).
fof(c_0_151, axiom, (accept_team(countrycmedicalorganization,christiancountrychumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_131)).
fof(c_0_152, axiom, (the_agent_in_all_proposed_teams(countrycmedicalorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_132)).
fof(c_0_153, axiom, (any_agent_in_all_proposed_teams(countrycmedicalorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_133)).
fof(c_0_154, axiom, (accept_team(countrybhumanitarianorganization,countrybcivilorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_134)).
fof(c_0_155, axiom, (accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_135)).
fof(c_0_156, axiom, (accept_team(countrybhumanitarianorganization,countrycmedicalorganization,townc,n6)), file('Axioms/AGT001+1.ax', event_136)).
fof(c_0_157, axiom, (the_agent_in_all_proposed_teams(countrybhumanitarianorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_137)).
fof(c_0_158, axiom, (any_agent_in_all_proposed_teams(countrybhumanitarianorganization,countryahumanitarianorganization,townc)), file('Axioms/AGT001+1.ax', event_138)).
fof(c_0_159, axiom, (accept_team(christiansufferterrahumanitarianorganization,countryahumanitarianorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_139)).
fof(c_0_160, axiom, (accept_number(christiansufferterrahumanitarianorganization,n6)), file('Axioms/AGT001+1.ax', event_140)).
fof(c_0_161, axiom, (the_agent_in_all_proposed_teams(christiansufferterrahumanitarianorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_141)).
fof(c_0_162, axiom, (any_agent_in_all_proposed_teams(christiansufferterrahumanitarianorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_142)).
fof(c_0_163, axiom, (~(accept_team(countryccivilorganization,countrybhumanitarianorganization,cityb,n5))), file('Axioms/AGT001+1.ax', event_143)).
fof(c_0_164, axiom, (~(accept_team(countryccivilorganization,countrybhumanitarianorganization,cityb,n6))), file('Axioms/AGT001+1.ax', event_144)).
fof(c_0_165, axiom, (accept_team(sufferterragovernment,countryahumanitarianorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_145)).
fof(c_0_166, axiom, (the_agent_in_all_proposed_teams(sufferterragovernment,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_146)).
fof(c_0_167, axiom, (any_agent_in_all_proposed_teams(sufferterragovernment,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_147)).
fof(c_0_168, axiom, (accept_team(countryafirstaidorganization,countryahumanitarianorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_148)).
fof(c_0_169, axiom, (accept_number(countryafirstaidorganization,n6)), file('Axioms/AGT001+1.ax', event_149)).
fof(c_0_170, axiom, (the_agent_in_all_proposed_teams(countryafirstaidorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_150)).
fof(c_0_171, axiom, (any_agent_in_all_proposed_teams(countryafirstaidorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_151)).
fof(c_0_172, axiom, (accept_team(countryafirstaidorganization,sufferterragovernment,towna,n5)), file('Axioms/AGT001+1.ax', event_152)).
fof(c_0_173, axiom, (accept_team(countryafirstaidorganization,sufferterragovernment,towna,n6)), file('Axioms/AGT001+1.ax', event_153)).
fof(c_0_174, axiom, (accept_team(countrycmedicalorganization,countrybhumanitarianorganization,cityb,n5)), file('Axioms/AGT001+1.ax', event_154)).
fof(c_0_175, axiom, (accept_team(countrycmedicalorganization,countrybhumanitarianorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_155)).
fof(c_0_176, axiom, (accept_team(countryamedicalorganization,countryahumanitarianorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_156)).
fof(c_0_177, axiom, (accept_number(countryamedicalorganization,n6)), file('Axioms/AGT001+1.ax', event_157)).
fof(c_0_178, axiom, (the_agent_in_all_proposed_teams(countryamedicalorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_158)).
fof(c_0_179, axiom, (any_agent_in_all_proposed_teams(countryamedicalorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_159)).
fof(c_0_180, axiom, (accept_team(countrybcivilorganization,countrybhumanitarianorganization,cityb,n5)), file('Axioms/AGT001+1.ax', event_160)).
fof(c_0_181, axiom, (accept_team(countrybcivilorganization,countrybhumanitarianorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_161)).
fof(c_0_182, axiom, (accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,cityb,n5)), file('Axioms/AGT001+1.ax', event_162)).
fof(c_0_183, axiom, (accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_163)).
fof(c_0_184, axiom, (the_agent_in_all_proposed_teams(countryahumanitarianorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_164)).
fof(c_0_185, axiom, (any_agent_in_all_proposed_teams(countryahumanitarianorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_165)).
fof(c_0_186, axiom, (accept_number(countryahumanitarianorganization,n6)), file('Axioms/AGT001+1.ax', event_166)).
fof(c_0_187, axiom, (accept_team(countryacivilorganization,countryahumanitarianorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_167)).
fof(c_0_188, axiom, (accept_number(countryacivilorganization,n6)), file('Axioms/AGT001+1.ax', event_168)).
fof(c_0_189, axiom, (the_agent_in_all_proposed_teams(countryacivilorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_169)).
fof(c_0_190, axiom, (any_agent_in_all_proposed_teams(countryacivilorganization,countrybhumanitarianorganization,cityb)), file('Axioms/AGT001+1.ax', event_170)).
fof(c_0_191, axiom, (accept_number(countrybhumanitarianorganization,n5)), file('Axioms/AGT001+1.ax', event_171)).
fof(c_0_192, axiom, (accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,cityb,n5)), file('Axioms/AGT001+1.ax', event_172)).
fof(c_0_193, axiom, (accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,cityb,n6)), file('Axioms/AGT001+1.ax', event_173)).
fof(c_0_194, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,towna,n2))), file('Axioms/AGT001+1.ax', event_174)).
fof(c_0_195, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n2))), file('Axioms/AGT001+1.ax', event_175)).
fof(c_0_196, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n2))), file('Axioms/AGT001+1.ax', event_176)).
fof(c_0_197, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,towna,n2))), file('Axioms/AGT001+1.ax', event_177)).
fof(c_0_198, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,towna,n3))), file('Axioms/AGT001+1.ax', event_178)).
fof(c_0_199, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n3))), file('Axioms/AGT001+1.ax', event_179)).
fof(c_0_200, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n3))), file('Axioms/AGT001+1.ax', event_180)).
fof(c_0_201, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,towna,n3))), file('Axioms/AGT001+1.ax', event_181)).
fof(c_0_202, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,towna,n2))), file('Axioms/AGT001+1.ax', event_182)).
fof(c_0_203, axiom, (accept_number(muslimcountrybhumanitarianorganization,n2)), file('Axioms/AGT001+1.ax', event_183)).
fof(c_0_204, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,towna,n4))), file('Axioms/AGT001+1.ax', event_184)).
fof(c_0_205, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,towna,n6))), file('Axioms/AGT001+1.ax', event_185)).
fof(c_0_206, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,towna,n6))), file('Axioms/AGT001+1.ax', event_186)).
fof(c_0_207, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n6))), file('Axioms/AGT001+1.ax', event_187)).
fof(c_0_208, axiom, (accept_team(countryamedicalorganization,sufferterragovernment,towna,n5)), file('Axioms/AGT001+1.ax', event_188)).
fof(c_0_209, axiom, (accept_team(countryamedicalorganization,sufferterragovernment,towna,n6)), file('Axioms/AGT001+1.ax', event_189)).
fof(c_0_210, axiom, (accept_team(countryahumanitarianorganization,sufferterragovernment,towna,n5)), file('Axioms/AGT001+1.ax', event_190)).
fof(c_0_211, axiom, (accept_team(countryahumanitarianorganization,sufferterragovernment,towna,n6)), file('Axioms/AGT001+1.ax', event_191)).
fof(c_0_212, axiom, (accept_number(christiancountrychumanitarianorganization,n2)), file('Axioms/AGT001+1.ax', event_192)).
fof(c_0_213, axiom, (accept_team(countrybhumanitarianorganization,countrybcivilorganization,towna,n3)), file('Axioms/AGT001+1.ax', event_193)).
fof(c_0_214, axiom, (accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n3)), file('Axioms/AGT001+1.ax', event_194)).
fof(c_0_215, axiom, (accept_team(countrybhumanitarianorganization,countrycmedicalorganization,towna,n3)), file('Axioms/AGT001+1.ax', event_195)).
fof(c_0_216, axiom, (accept_team(countrybhumanitarianorganization,countryccivilorganization,towna,n3)), file('Axioms/AGT001+1.ax', event_196)).
fof(c_0_217, axiom, (accept_number(countrybhumanitarianorganization,n2)), file('Axioms/AGT001+1.ax', event_197)).
fof(c_0_218, axiom, (accept_team(countrybhumanitarianorganization,countrybcivilorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_198)).
fof(c_0_219, axiom, (accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_199)).
fof(c_0_220, axiom, (accept_team(countrybhumanitarianorganization,countrycmedicalorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_200)).
fof(c_0_221, axiom, (the_agent_in_all_proposed_teams(countrybhumanitarianorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_201)).
fof(c_0_222, axiom, (any_agent_in_all_proposed_teams(countrybhumanitarianorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_202)).
fof(c_0_223, axiom, (accept_number(countrybcivilorganization,n2)), file('Axioms/AGT001+1.ax', event_203)).
fof(c_0_224, axiom, (accept_number(countrybcivilorganization,n3)), file('Axioms/AGT001+1.ax', event_204)).
fof(c_0_225, axiom, (accept_team(countrybcivilorganization,countryccivilorganization,towna,n4)), file('Axioms/AGT001+1.ax', event_205)).
fof(c_0_226, axiom, (accept_team(countrybcivilorganization,countrybhumanitarianorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_206)).
fof(c_0_227, axiom, (accept_team(countrybcivilorganization,christiancountrychumanitarianorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_207)).
fof(c_0_228, axiom, (accept_team(countrybcivilorganization,countrycmedicalorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_208)).
fof(c_0_229, axiom, (the_agent_in_all_proposed_teams(countrybcivilorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_209)).
fof(c_0_230, axiom, (any_agent_in_all_proposed_teams(countrybcivilorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_210)).
fof(c_0_231, axiom, (accept_number(sufferterragovernment,n2)), file('Axioms/AGT001+1.ax', event_211)).
fof(c_0_232, axiom, (~(accept_team(christiansufferterrahumanitarianorganization,sufferterragovernment,towna,n2))), file('Axioms/AGT001+1.ax', event_212)).
fof(c_0_233, axiom, (~(accept_team(christiansufferterrahumanitarianorganization,sufferterragovernment,towna,n6))), file('Axioms/AGT001+1.ax', event_213)).
fof(c_0_234, axiom, (accept_number(countryccivilorganization,n2)), file('Axioms/AGT001+1.ax', event_214)).
fof(c_0_235, axiom, (accept_number(countryccivilorganization,n3)), file('Axioms/AGT001+1.ax', event_215)).
fof(c_0_236, axiom, (~(accept_team(countryccivilorganization,countrybcivilorganization,towna,n6))), file('Axioms/AGT001+1.ax', event_216)).
fof(c_0_237, axiom, (~(accept_team(countryccivilorganization,countrybhumanitarianorganization,towna,n6))), file('Axioms/AGT001+1.ax', event_217)).
fof(c_0_238, axiom, (~(accept_team(countryccivilorganization,christiancountrychumanitarianorganization,towna,n6))), file('Axioms/AGT001+1.ax', event_218)).
fof(c_0_239, axiom, (~(accept_team(countryccivilorganization,countrycmedicalorganization,towna,n6))), file('Axioms/AGT001+1.ax', event_219)).
fof(c_0_240, axiom, (the_agent_in_all_proposed_teams(countryccivilorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_220)).
fof(c_0_241, axiom, (any_agent_in_all_proposed_teams(countryccivilorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_221)).
fof(c_0_242, axiom, (accept_number(sufferterragovernment,n5)), file('Axioms/AGT001+1.ax', event_222)).
fof(c_0_243, axiom, (accept_team(countryacivilorganization,sufferterragovernment,towna,n5)), file('Axioms/AGT001+1.ax', event_223)).
fof(c_0_244, axiom, (accept_number(sufferterragovernment,n6)), file('Axioms/AGT001+1.ax', event_224)).
fof(c_0_245, axiom, (accept_team(countryacivilorganization,sufferterragovernment,towna,n6)), file('Axioms/AGT001+1.ax', event_225)).
fof(c_0_246, axiom, (accept_number(countrycmedicalorganization,n2)), file('Axioms/AGT001+1.ax', event_226)).
fof(c_0_247, axiom, (accept_number(countrycmedicalorganization,n3)), file('Axioms/AGT001+1.ax', event_227)).
fof(c_0_248, axiom, (accept_team(countrycmedicalorganization,countryccivilorganization,towna,n4)), file('Axioms/AGT001+1.ax', event_228)).
fof(c_0_249, axiom, (accept_team(countrycmedicalorganization,countrybcivilorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_229)).
fof(c_0_250, axiom, (accept_team(countrycmedicalorganization,countrybhumanitarianorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_230)).
fof(c_0_251, axiom, (accept_team(countrycmedicalorganization,christiancountrychumanitarianorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_231)).
fof(c_0_252, axiom, (the_agent_in_all_proposed_teams(countrycmedicalorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_232)).
fof(c_0_253, axiom, (any_agent_in_all_proposed_teams(countrycmedicalorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_233)).
fof(c_0_254, axiom, (accept_number(christiancountrychumanitarianorganization,n3)), file('Axioms/AGT001+1.ax', event_234)).
fof(c_0_255, axiom, (accept_number(countryccivilorganization,n4)), file('Axioms/AGT001+1.ax', event_235)).
fof(c_0_256, axiom, (accept_team(christiancountrychumanitarianorganization,countryccivilorganization,towna,n4)), file('Axioms/AGT001+1.ax', event_236)).
fof(c_0_257, axiom, (accept_number(countrybcivilorganization,n6)), file('Axioms/AGT001+1.ax', event_237)).
fof(c_0_258, axiom, (accept_team(christiancountrychumanitarianorganization,countrybcivilorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_238)).
fof(c_0_259, axiom, (accept_number(countrybhumanitarianorganization,n6)), file('Axioms/AGT001+1.ax', event_239)).
fof(c_0_260, axiom, (accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_240)).
fof(c_0_261, axiom, (accept_team(christiancountrychumanitarianorganization,countrycmedicalorganization,towna,n6)), file('Axioms/AGT001+1.ax', event_241)).
fof(c_0_262, axiom, (accept_number(christiancountrychumanitarianorganization,n6)), file('Axioms/AGT001+1.ax', event_242)).
fof(c_0_263, axiom, (the_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_243)).
fof(c_0_264, axiom, (any_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_244)).
fof(c_0_265, axiom, (accept_number(countrycmedicalorganization,n6)), file('Axioms/AGT001+1.ax', event_245)).
fof(c_0_266, axiom, (~(accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n6))), file('Axioms/AGT001+1.ax', event_246)).
fof(c_0_267, axiom, (accept_number(muslimcountrybhumanitarianorganization,n6)), file('Axioms/AGT001+1.ax', event_247)).
fof(c_0_268, axiom, (the_agent_not_in_any_proposed_teams(muslimcountrybhumanitarianorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_248)).
fof(c_0_269, axiom, (any_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,sufferterragovernment,towna)), file('Axioms/AGT001+1.ax', event_249)).
fof(c_0_270, axiom, (rdn_translate(n0,rdn_pos(rdnn(n0)))), file('Axioms/NUM005+0.ax', rdn0)).
fof(c_0_271, axiom, (rdn_translate(n1,rdn_pos(rdnn(n1)))), file('Axioms/NUM005+0.ax', rdn1)).
fof(c_0_272, axiom, (rdn_translate(n2,rdn_pos(rdnn(n2)))), file('Axioms/NUM005+0.ax', rdn2)).
fof(c_0_273, axiom, (rdn_translate(n3,rdn_pos(rdnn(n3)))), file('Axioms/NUM005+0.ax', rdn3)).
fof(c_0_274, axiom, (rdn_translate(n4,rdn_pos(rdnn(n4)))), file('Axioms/NUM005+0.ax', rdn4)).
fof(c_0_275, axiom, (rdn_translate(n5,rdn_pos(rdnn(n5)))), file('Axioms/NUM005+0.ax', rdn5)).
fof(c_0_276, axiom, (rdn_translate(n6,rdn_pos(rdnn(n6)))), file('Axioms/NUM005+0.ax', rdn6)).
fof(c_0_277, axiom, (rdn_translate(n7,rdn_pos(rdnn(n7)))), file('Axioms/NUM005+0.ax', rdn7)).
fof(c_0_278, axiom, (rdn_translate(n8,rdn_pos(rdnn(n8)))), file('Axioms/NUM005+0.ax', rdn8)).
fof(c_0_279, axiom, (rdn_translate(n9,rdn_pos(rdnn(n9)))), file('Axioms/NUM005+0.ax', rdn9)).
fof(c_0_280, axiom, (rdn_translate(n10,rdn_pos(rdn(rdnn(n0),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn10)).
fof(c_0_281, axiom, (rdn_translate(n11,rdn_pos(rdn(rdnn(n1),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn11)).
fof(c_0_282, axiom, (rdn_translate(n12,rdn_pos(rdn(rdnn(n2),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn12)).
fof(c_0_283, axiom, (rdn_translate(n13,rdn_pos(rdn(rdnn(n3),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn13)).
fof(c_0_284, axiom, (rdn_translate(n14,rdn_pos(rdn(rdnn(n4),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn14)).
fof(c_0_285, axiom, (rdn_translate(n15,rdn_pos(rdn(rdnn(n5),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn15)).
fof(c_0_286, axiom, (rdn_translate(n16,rdn_pos(rdn(rdnn(n6),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn16)).
fof(c_0_287, axiom, (rdn_translate(n17,rdn_pos(rdn(rdnn(n7),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn17)).
fof(c_0_288, axiom, (rdn_translate(n18,rdn_pos(rdn(rdnn(n8),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn18)).
fof(c_0_289, axiom, (rdn_translate(n19,rdn_pos(rdn(rdnn(n9),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdn19)).
fof(c_0_290, axiom, (rdn_translate(n20,rdn_pos(rdn(rdnn(n0),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn20)).
fof(c_0_291, axiom, (rdn_translate(n21,rdn_pos(rdn(rdnn(n1),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn21)).
fof(c_0_292, axiom, (rdn_translate(n22,rdn_pos(rdn(rdnn(n2),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn22)).
fof(c_0_293, axiom, (rdn_translate(n23,rdn_pos(rdn(rdnn(n3),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn23)).
fof(c_0_294, axiom, (rdn_translate(n24,rdn_pos(rdn(rdnn(n4),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn24)).
fof(c_0_295, axiom, (rdn_translate(n25,rdn_pos(rdn(rdnn(n5),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn25)).
fof(c_0_296, axiom, (rdn_translate(n26,rdn_pos(rdn(rdnn(n6),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn26)).
fof(c_0_297, axiom, (rdn_translate(n27,rdn_pos(rdn(rdnn(n7),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn27)).
fof(c_0_298, axiom, (rdn_translate(n28,rdn_pos(rdn(rdnn(n8),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn28)).
fof(c_0_299, axiom, (rdn_translate(n29,rdn_pos(rdn(rdnn(n9),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdn29)).
fof(c_0_300, axiom, (rdn_translate(n30,rdn_pos(rdn(rdnn(n0),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn30)).
fof(c_0_301, axiom, (rdn_translate(n31,rdn_pos(rdn(rdnn(n1),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn31)).
fof(c_0_302, axiom, (rdn_translate(n32,rdn_pos(rdn(rdnn(n2),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn32)).
fof(c_0_303, axiom, (rdn_translate(n33,rdn_pos(rdn(rdnn(n3),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn33)).
fof(c_0_304, axiom, (rdn_translate(n34,rdn_pos(rdn(rdnn(n4),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn34)).
fof(c_0_305, axiom, (rdn_translate(n35,rdn_pos(rdn(rdnn(n5),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn35)).
fof(c_0_306, axiom, (rdn_translate(n36,rdn_pos(rdn(rdnn(n6),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn36)).
fof(c_0_307, axiom, (rdn_translate(n37,rdn_pos(rdn(rdnn(n7),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn37)).
fof(c_0_308, axiom, (rdn_translate(n38,rdn_pos(rdn(rdnn(n8),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn38)).
fof(c_0_309, axiom, (rdn_translate(n39,rdn_pos(rdn(rdnn(n9),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdn39)).
fof(c_0_310, axiom, (rdn_translate(n40,rdn_pos(rdn(rdnn(n0),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn40)).
fof(c_0_311, axiom, (rdn_translate(n41,rdn_pos(rdn(rdnn(n1),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn41)).
fof(c_0_312, axiom, (rdn_translate(n42,rdn_pos(rdn(rdnn(n2),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn42)).
fof(c_0_313, axiom, (rdn_translate(n43,rdn_pos(rdn(rdnn(n3),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn43)).
fof(c_0_314, axiom, (rdn_translate(n44,rdn_pos(rdn(rdnn(n4),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn44)).
fof(c_0_315, axiom, (rdn_translate(n45,rdn_pos(rdn(rdnn(n5),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn45)).
fof(c_0_316, axiom, (rdn_translate(n46,rdn_pos(rdn(rdnn(n6),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn46)).
fof(c_0_317, axiom, (rdn_translate(n47,rdn_pos(rdn(rdnn(n7),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn47)).
fof(c_0_318, axiom, (rdn_translate(n48,rdn_pos(rdn(rdnn(n8),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn48)).
fof(c_0_319, axiom, (rdn_translate(n49,rdn_pos(rdn(rdnn(n9),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdn49)).
fof(c_0_320, axiom, (rdn_translate(n50,rdn_pos(rdn(rdnn(n0),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn50)).
fof(c_0_321, axiom, (rdn_translate(n51,rdn_pos(rdn(rdnn(n1),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn51)).
fof(c_0_322, axiom, (rdn_translate(n52,rdn_pos(rdn(rdnn(n2),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn52)).
fof(c_0_323, axiom, (rdn_translate(n53,rdn_pos(rdn(rdnn(n3),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn53)).
fof(c_0_324, axiom, (rdn_translate(n54,rdn_pos(rdn(rdnn(n4),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn54)).
fof(c_0_325, axiom, (rdn_translate(n55,rdn_pos(rdn(rdnn(n5),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn55)).
fof(c_0_326, axiom, (rdn_translate(n56,rdn_pos(rdn(rdnn(n6),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn56)).
fof(c_0_327, axiom, (rdn_translate(n57,rdn_pos(rdn(rdnn(n7),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn57)).
fof(c_0_328, axiom, (rdn_translate(n58,rdn_pos(rdn(rdnn(n8),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn58)).
fof(c_0_329, axiom, (rdn_translate(n59,rdn_pos(rdn(rdnn(n9),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdn59)).
fof(c_0_330, axiom, (rdn_translate(n60,rdn_pos(rdn(rdnn(n0),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn60)).
fof(c_0_331, axiom, (rdn_translate(n61,rdn_pos(rdn(rdnn(n1),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn61)).
fof(c_0_332, axiom, (rdn_translate(n62,rdn_pos(rdn(rdnn(n2),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn62)).
fof(c_0_333, axiom, (rdn_translate(n63,rdn_pos(rdn(rdnn(n3),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn63)).
fof(c_0_334, axiom, (rdn_translate(n64,rdn_pos(rdn(rdnn(n4),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn64)).
fof(c_0_335, axiom, (rdn_translate(n65,rdn_pos(rdn(rdnn(n5),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn65)).
fof(c_0_336, axiom, (rdn_translate(n66,rdn_pos(rdn(rdnn(n6),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn66)).
fof(c_0_337, axiom, (rdn_translate(n67,rdn_pos(rdn(rdnn(n7),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn67)).
fof(c_0_338, axiom, (rdn_translate(n68,rdn_pos(rdn(rdnn(n8),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn68)).
fof(c_0_339, axiom, (rdn_translate(n69,rdn_pos(rdn(rdnn(n9),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdn69)).
fof(c_0_340, axiom, (rdn_translate(n70,rdn_pos(rdn(rdnn(n0),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn70)).
fof(c_0_341, axiom, (rdn_translate(n71,rdn_pos(rdn(rdnn(n1),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn71)).
fof(c_0_342, axiom, (rdn_translate(n72,rdn_pos(rdn(rdnn(n2),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn72)).
fof(c_0_343, axiom, (rdn_translate(n73,rdn_pos(rdn(rdnn(n3),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn73)).
fof(c_0_344, axiom, (rdn_translate(n74,rdn_pos(rdn(rdnn(n4),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn74)).
fof(c_0_345, axiom, (rdn_translate(n75,rdn_pos(rdn(rdnn(n5),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn75)).
fof(c_0_346, axiom, (rdn_translate(n76,rdn_pos(rdn(rdnn(n6),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn76)).
fof(c_0_347, axiom, (rdn_translate(n77,rdn_pos(rdn(rdnn(n7),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn77)).
fof(c_0_348, axiom, (rdn_translate(n78,rdn_pos(rdn(rdnn(n8),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn78)).
fof(c_0_349, axiom, (rdn_translate(n79,rdn_pos(rdn(rdnn(n9),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdn79)).
fof(c_0_350, axiom, (rdn_translate(n80,rdn_pos(rdn(rdnn(n0),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn80)).
fof(c_0_351, axiom, (rdn_translate(n81,rdn_pos(rdn(rdnn(n1),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn81)).
fof(c_0_352, axiom, (rdn_translate(n82,rdn_pos(rdn(rdnn(n2),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn82)).
fof(c_0_353, axiom, (rdn_translate(n83,rdn_pos(rdn(rdnn(n3),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn83)).
fof(c_0_354, axiom, (rdn_translate(n84,rdn_pos(rdn(rdnn(n4),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn84)).
fof(c_0_355, axiom, (rdn_translate(n85,rdn_pos(rdn(rdnn(n5),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn85)).
fof(c_0_356, axiom, (rdn_translate(n86,rdn_pos(rdn(rdnn(n6),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn86)).
fof(c_0_357, axiom, (rdn_translate(n87,rdn_pos(rdn(rdnn(n7),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn87)).
fof(c_0_358, axiom, (rdn_translate(n88,rdn_pos(rdn(rdnn(n8),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn88)).
fof(c_0_359, axiom, (rdn_translate(n89,rdn_pos(rdn(rdnn(n9),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdn89)).
fof(c_0_360, axiom, (rdn_translate(n90,rdn_pos(rdn(rdnn(n0),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn90)).
fof(c_0_361, axiom, (rdn_translate(n91,rdn_pos(rdn(rdnn(n1),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn91)).
fof(c_0_362, axiom, (rdn_translate(n92,rdn_pos(rdn(rdnn(n2),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn92)).
fof(c_0_363, axiom, (rdn_translate(n93,rdn_pos(rdn(rdnn(n3),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn93)).
fof(c_0_364, axiom, (rdn_translate(n94,rdn_pos(rdn(rdnn(n4),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn94)).
fof(c_0_365, axiom, (rdn_translate(n95,rdn_pos(rdn(rdnn(n5),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn95)).
fof(c_0_366, axiom, (rdn_translate(n96,rdn_pos(rdn(rdnn(n6),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn96)).
fof(c_0_367, axiom, (rdn_translate(n97,rdn_pos(rdn(rdnn(n7),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn97)).
fof(c_0_368, axiom, (rdn_translate(n98,rdn_pos(rdn(rdnn(n8),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn98)).
fof(c_0_369, axiom, (rdn_translate(n99,rdn_pos(rdn(rdnn(n9),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdn99)).
fof(c_0_370, axiom, (rdn_translate(n100,rdn_pos(rdn(rdnn(n0),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn100)).
fof(c_0_371, axiom, (rdn_translate(n101,rdn_pos(rdn(rdnn(n1),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn101)).
fof(c_0_372, axiom, (rdn_translate(n102,rdn_pos(rdn(rdnn(n2),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn102)).
fof(c_0_373, axiom, (rdn_translate(n103,rdn_pos(rdn(rdnn(n3),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn103)).
fof(c_0_374, axiom, (rdn_translate(n104,rdn_pos(rdn(rdnn(n4),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn104)).
fof(c_0_375, axiom, (rdn_translate(n105,rdn_pos(rdn(rdnn(n5),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn105)).
fof(c_0_376, axiom, (rdn_translate(n106,rdn_pos(rdn(rdnn(n6),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn106)).
fof(c_0_377, axiom, (rdn_translate(n107,rdn_pos(rdn(rdnn(n7),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn107)).
fof(c_0_378, axiom, (rdn_translate(n108,rdn_pos(rdn(rdnn(n8),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn108)).
fof(c_0_379, axiom, (rdn_translate(n109,rdn_pos(rdn(rdnn(n9),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn109)).
fof(c_0_380, axiom, (rdn_translate(n110,rdn_pos(rdn(rdnn(n0),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn110)).
fof(c_0_381, axiom, (rdn_translate(n111,rdn_pos(rdn(rdnn(n1),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn111)).
fof(c_0_382, axiom, (rdn_translate(n112,rdn_pos(rdn(rdnn(n2),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn112)).
fof(c_0_383, axiom, (rdn_translate(n113,rdn_pos(rdn(rdnn(n3),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn113)).
fof(c_0_384, axiom, (rdn_translate(n114,rdn_pos(rdn(rdnn(n4),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn114)).
fof(c_0_385, axiom, (rdn_translate(n115,rdn_pos(rdn(rdnn(n5),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn115)).
fof(c_0_386, axiom, (rdn_translate(n116,rdn_pos(rdn(rdnn(n6),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn116)).
fof(c_0_387, axiom, (rdn_translate(n117,rdn_pos(rdn(rdnn(n7),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn117)).
fof(c_0_388, axiom, (rdn_translate(n118,rdn_pos(rdn(rdnn(n8),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn118)).
fof(c_0_389, axiom, (rdn_translate(n119,rdn_pos(rdn(rdnn(n9),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn119)).
fof(c_0_390, axiom, (rdn_translate(n120,rdn_pos(rdn(rdnn(n0),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn120)).
fof(c_0_391, axiom, (rdn_translate(n121,rdn_pos(rdn(rdnn(n1),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn121)).
fof(c_0_392, axiom, (rdn_translate(n122,rdn_pos(rdn(rdnn(n2),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn122)).
fof(c_0_393, axiom, (rdn_translate(n123,rdn_pos(rdn(rdnn(n3),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn123)).
fof(c_0_394, axiom, (rdn_translate(n124,rdn_pos(rdn(rdnn(n4),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn124)).
fof(c_0_395, axiom, (rdn_translate(n125,rdn_pos(rdn(rdnn(n5),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn125)).
fof(c_0_396, axiom, (rdn_translate(n126,rdn_pos(rdn(rdnn(n6),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn126)).
fof(c_0_397, axiom, (rdn_translate(n127,rdn_pos(rdn(rdnn(n7),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdn127)).
fof(c_0_398, axiom, (rdn_translate(nn1,rdn_neg(rdnn(n1)))), file('Axioms/NUM005+0.ax', rdnn1)).
fof(c_0_399, axiom, (rdn_translate(nn2,rdn_neg(rdnn(n2)))), file('Axioms/NUM005+0.ax', rdnn2)).
fof(c_0_400, axiom, (rdn_translate(nn3,rdn_neg(rdnn(n3)))), file('Axioms/NUM005+0.ax', rdnn3)).
fof(c_0_401, axiom, (rdn_translate(nn4,rdn_neg(rdnn(n4)))), file('Axioms/NUM005+0.ax', rdnn4)).
fof(c_0_402, axiom, (rdn_translate(nn5,rdn_neg(rdnn(n5)))), file('Axioms/NUM005+0.ax', rdnn5)).
fof(c_0_403, axiom, (rdn_translate(nn6,rdn_neg(rdnn(n6)))), file('Axioms/NUM005+0.ax', rdnn6)).
fof(c_0_404, axiom, (rdn_translate(nn7,rdn_neg(rdnn(n7)))), file('Axioms/NUM005+0.ax', rdnn7)).
fof(c_0_405, axiom, (rdn_translate(nn8,rdn_neg(rdnn(n8)))), file('Axioms/NUM005+0.ax', rdnn8)).
fof(c_0_406, axiom, (rdn_translate(nn9,rdn_neg(rdnn(n9)))), file('Axioms/NUM005+0.ax', rdnn9)).
fof(c_0_407, axiom, (rdn_translate(nn10,rdn_neg(rdn(rdnn(n0),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn10)).
fof(c_0_408, axiom, (rdn_translate(nn11,rdn_neg(rdn(rdnn(n1),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn11)).
fof(c_0_409, axiom, (rdn_translate(nn12,rdn_neg(rdn(rdnn(n2),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn12)).
fof(c_0_410, axiom, (rdn_translate(nn13,rdn_neg(rdn(rdnn(n3),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn13)).
fof(c_0_411, axiom, (rdn_translate(nn14,rdn_neg(rdn(rdnn(n4),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn14)).
fof(c_0_412, axiom, (rdn_translate(nn15,rdn_neg(rdn(rdnn(n5),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn15)).
fof(c_0_413, axiom, (rdn_translate(nn16,rdn_neg(rdn(rdnn(n6),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn16)).
fof(c_0_414, axiom, (rdn_translate(nn17,rdn_neg(rdn(rdnn(n7),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn17)).
fof(c_0_415, axiom, (rdn_translate(nn18,rdn_neg(rdn(rdnn(n8),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn18)).
fof(c_0_416, axiom, (rdn_translate(nn19,rdn_neg(rdn(rdnn(n9),rdnn(n1))))), file('Axioms/NUM005+0.ax', rdnn19)).
fof(c_0_417, axiom, (rdn_translate(nn20,rdn_neg(rdn(rdnn(n0),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn20)).
fof(c_0_418, axiom, (rdn_translate(nn21,rdn_neg(rdn(rdnn(n1),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn21)).
fof(c_0_419, axiom, (rdn_translate(nn22,rdn_neg(rdn(rdnn(n2),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn22)).
fof(c_0_420, axiom, (rdn_translate(nn23,rdn_neg(rdn(rdnn(n3),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn23)).
fof(c_0_421, axiom, (rdn_translate(nn24,rdn_neg(rdn(rdnn(n4),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn24)).
fof(c_0_422, axiom, (rdn_translate(nn25,rdn_neg(rdn(rdnn(n5),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn25)).
fof(c_0_423, axiom, (rdn_translate(nn26,rdn_neg(rdn(rdnn(n6),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn26)).
fof(c_0_424, axiom, (rdn_translate(nn27,rdn_neg(rdn(rdnn(n7),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn27)).
fof(c_0_425, axiom, (rdn_translate(nn28,rdn_neg(rdn(rdnn(n8),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn28)).
fof(c_0_426, axiom, (rdn_translate(nn29,rdn_neg(rdn(rdnn(n9),rdnn(n2))))), file('Axioms/NUM005+0.ax', rdnn29)).
fof(c_0_427, axiom, (rdn_translate(nn30,rdn_neg(rdn(rdnn(n0),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn30)).
fof(c_0_428, axiom, (rdn_translate(nn31,rdn_neg(rdn(rdnn(n1),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn31)).
fof(c_0_429, axiom, (rdn_translate(nn32,rdn_neg(rdn(rdnn(n2),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn32)).
fof(c_0_430, axiom, (rdn_translate(nn33,rdn_neg(rdn(rdnn(n3),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn33)).
fof(c_0_431, axiom, (rdn_translate(nn34,rdn_neg(rdn(rdnn(n4),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn34)).
fof(c_0_432, axiom, (rdn_translate(nn35,rdn_neg(rdn(rdnn(n5),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn35)).
fof(c_0_433, axiom, (rdn_translate(nn36,rdn_neg(rdn(rdnn(n6),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn36)).
fof(c_0_434, axiom, (rdn_translate(nn37,rdn_neg(rdn(rdnn(n7),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn37)).
fof(c_0_435, axiom, (rdn_translate(nn38,rdn_neg(rdn(rdnn(n8),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn38)).
fof(c_0_436, axiom, (rdn_translate(nn39,rdn_neg(rdn(rdnn(n9),rdnn(n3))))), file('Axioms/NUM005+0.ax', rdnn39)).
fof(c_0_437, axiom, (rdn_translate(nn40,rdn_neg(rdn(rdnn(n0),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn40)).
fof(c_0_438, axiom, (rdn_translate(nn41,rdn_neg(rdn(rdnn(n1),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn41)).
fof(c_0_439, axiom, (rdn_translate(nn42,rdn_neg(rdn(rdnn(n2),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn42)).
fof(c_0_440, axiom, (rdn_translate(nn43,rdn_neg(rdn(rdnn(n3),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn43)).
fof(c_0_441, axiom, (rdn_translate(nn44,rdn_neg(rdn(rdnn(n4),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn44)).
fof(c_0_442, axiom, (rdn_translate(nn45,rdn_neg(rdn(rdnn(n5),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn45)).
fof(c_0_443, axiom, (rdn_translate(nn46,rdn_neg(rdn(rdnn(n6),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn46)).
fof(c_0_444, axiom, (rdn_translate(nn47,rdn_neg(rdn(rdnn(n7),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn47)).
fof(c_0_445, axiom, (rdn_translate(nn48,rdn_neg(rdn(rdnn(n8),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn48)).
fof(c_0_446, axiom, (rdn_translate(nn49,rdn_neg(rdn(rdnn(n9),rdnn(n4))))), file('Axioms/NUM005+0.ax', rdnn49)).
fof(c_0_447, axiom, (rdn_translate(nn50,rdn_neg(rdn(rdnn(n0),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn50)).
fof(c_0_448, axiom, (rdn_translate(nn51,rdn_neg(rdn(rdnn(n1),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn51)).
fof(c_0_449, axiom, (rdn_translate(nn52,rdn_neg(rdn(rdnn(n2),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn52)).
fof(c_0_450, axiom, (rdn_translate(nn53,rdn_neg(rdn(rdnn(n3),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn53)).
fof(c_0_451, axiom, (rdn_translate(nn54,rdn_neg(rdn(rdnn(n4),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn54)).
fof(c_0_452, axiom, (rdn_translate(nn55,rdn_neg(rdn(rdnn(n5),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn55)).
fof(c_0_453, axiom, (rdn_translate(nn56,rdn_neg(rdn(rdnn(n6),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn56)).
fof(c_0_454, axiom, (rdn_translate(nn57,rdn_neg(rdn(rdnn(n7),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn57)).
fof(c_0_455, axiom, (rdn_translate(nn58,rdn_neg(rdn(rdnn(n8),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn58)).
fof(c_0_456, axiom, (rdn_translate(nn59,rdn_neg(rdn(rdnn(n9),rdnn(n5))))), file('Axioms/NUM005+0.ax', rdnn59)).
fof(c_0_457, axiom, (rdn_translate(nn60,rdn_neg(rdn(rdnn(n0),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn60)).
fof(c_0_458, axiom, (rdn_translate(nn61,rdn_neg(rdn(rdnn(n1),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn61)).
fof(c_0_459, axiom, (rdn_translate(nn62,rdn_neg(rdn(rdnn(n2),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn62)).
fof(c_0_460, axiom, (rdn_translate(nn63,rdn_neg(rdn(rdnn(n3),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn63)).
fof(c_0_461, axiom, (rdn_translate(nn64,rdn_neg(rdn(rdnn(n4),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn64)).
fof(c_0_462, axiom, (rdn_translate(nn65,rdn_neg(rdn(rdnn(n5),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn65)).
fof(c_0_463, axiom, (rdn_translate(nn66,rdn_neg(rdn(rdnn(n6),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn66)).
fof(c_0_464, axiom, (rdn_translate(nn67,rdn_neg(rdn(rdnn(n7),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn67)).
fof(c_0_465, axiom, (rdn_translate(nn68,rdn_neg(rdn(rdnn(n8),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn68)).
fof(c_0_466, axiom, (rdn_translate(nn69,rdn_neg(rdn(rdnn(n9),rdnn(n6))))), file('Axioms/NUM005+0.ax', rdnn69)).
fof(c_0_467, axiom, (rdn_translate(nn70,rdn_neg(rdn(rdnn(n0),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn70)).
fof(c_0_468, axiom, (rdn_translate(nn71,rdn_neg(rdn(rdnn(n1),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn71)).
fof(c_0_469, axiom, (rdn_translate(nn72,rdn_neg(rdn(rdnn(n2),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn72)).
fof(c_0_470, axiom, (rdn_translate(nn73,rdn_neg(rdn(rdnn(n3),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn73)).
fof(c_0_471, axiom, (rdn_translate(nn74,rdn_neg(rdn(rdnn(n4),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn74)).
fof(c_0_472, axiom, (rdn_translate(nn75,rdn_neg(rdn(rdnn(n5),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn75)).
fof(c_0_473, axiom, (rdn_translate(nn76,rdn_neg(rdn(rdnn(n6),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn76)).
fof(c_0_474, axiom, (rdn_translate(nn77,rdn_neg(rdn(rdnn(n7),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn77)).
fof(c_0_475, axiom, (rdn_translate(nn78,rdn_neg(rdn(rdnn(n8),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn78)).
fof(c_0_476, axiom, (rdn_translate(nn79,rdn_neg(rdn(rdnn(n9),rdnn(n7))))), file('Axioms/NUM005+0.ax', rdnn79)).
fof(c_0_477, axiom, (rdn_translate(nn80,rdn_neg(rdn(rdnn(n0),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn80)).
fof(c_0_478, axiom, (rdn_translate(nn81,rdn_neg(rdn(rdnn(n1),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn81)).
fof(c_0_479, axiom, (rdn_translate(nn82,rdn_neg(rdn(rdnn(n2),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn82)).
fof(c_0_480, axiom, (rdn_translate(nn83,rdn_neg(rdn(rdnn(n3),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn83)).
fof(c_0_481, axiom, (rdn_translate(nn84,rdn_neg(rdn(rdnn(n4),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn84)).
fof(c_0_482, axiom, (rdn_translate(nn85,rdn_neg(rdn(rdnn(n5),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn85)).
fof(c_0_483, axiom, (rdn_translate(nn86,rdn_neg(rdn(rdnn(n6),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn86)).
fof(c_0_484, axiom, (rdn_translate(nn87,rdn_neg(rdn(rdnn(n7),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn87)).
fof(c_0_485, axiom, (rdn_translate(nn88,rdn_neg(rdn(rdnn(n8),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn88)).
fof(c_0_486, axiom, (rdn_translate(nn89,rdn_neg(rdn(rdnn(n9),rdnn(n8))))), file('Axioms/NUM005+0.ax', rdnn89)).
fof(c_0_487, axiom, (rdn_translate(nn90,rdn_neg(rdn(rdnn(n0),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn90)).
fof(c_0_488, axiom, (rdn_translate(nn91,rdn_neg(rdn(rdnn(n1),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn91)).
fof(c_0_489, axiom, (rdn_translate(nn92,rdn_neg(rdn(rdnn(n2),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn92)).
fof(c_0_490, axiom, (rdn_translate(nn93,rdn_neg(rdn(rdnn(n3),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn93)).
fof(c_0_491, axiom, (rdn_translate(nn94,rdn_neg(rdn(rdnn(n4),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn94)).
fof(c_0_492, axiom, (rdn_translate(nn95,rdn_neg(rdn(rdnn(n5),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn95)).
fof(c_0_493, axiom, (rdn_translate(nn96,rdn_neg(rdn(rdnn(n6),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn96)).
fof(c_0_494, axiom, (rdn_translate(nn97,rdn_neg(rdn(rdnn(n7),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn97)).
fof(c_0_495, axiom, (rdn_translate(nn98,rdn_neg(rdn(rdnn(n8),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn98)).
fof(c_0_496, axiom, (rdn_translate(nn99,rdn_neg(rdn(rdnn(n9),rdnn(n9))))), file('Axioms/NUM005+0.ax', rdnn99)).
fof(c_0_497, axiom, (rdn_translate(nn100,rdn_neg(rdn(rdnn(n0),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn100)).
fof(c_0_498, axiom, (rdn_translate(nn101,rdn_neg(rdn(rdnn(n1),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn101)).
fof(c_0_499, axiom, (rdn_translate(nn102,rdn_neg(rdn(rdnn(n2),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn102)).
fof(c_0_500, axiom, (rdn_translate(nn103,rdn_neg(rdn(rdnn(n3),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn103)).
fof(c_0_501, axiom, (rdn_translate(nn104,rdn_neg(rdn(rdnn(n4),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn104)).
fof(c_0_502, axiom, (rdn_translate(nn105,rdn_neg(rdn(rdnn(n5),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn105)).
fof(c_0_503, axiom, (rdn_translate(nn106,rdn_neg(rdn(rdnn(n6),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn106)).
fof(c_0_504, axiom, (rdn_translate(nn107,rdn_neg(rdn(rdnn(n7),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn107)).
fof(c_0_505, axiom, (rdn_translate(nn108,rdn_neg(rdn(rdnn(n8),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn108)).
fof(c_0_506, axiom, (rdn_translate(nn109,rdn_neg(rdn(rdnn(n9),rdn(rdnn(n0),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn109)).
fof(c_0_507, axiom, (rdn_translate(nn110,rdn_neg(rdn(rdnn(n0),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn110)).
fof(c_0_508, axiom, (rdn_translate(nn111,rdn_neg(rdn(rdnn(n1),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn111)).
fof(c_0_509, axiom, (rdn_translate(nn112,rdn_neg(rdn(rdnn(n2),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn112)).
fof(c_0_510, axiom, (rdn_translate(nn113,rdn_neg(rdn(rdnn(n3),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn113)).
fof(c_0_511, axiom, (rdn_translate(nn114,rdn_neg(rdn(rdnn(n4),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn114)).
fof(c_0_512, axiom, (rdn_translate(nn115,rdn_neg(rdn(rdnn(n5),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn115)).
fof(c_0_513, axiom, (rdn_translate(nn116,rdn_neg(rdn(rdnn(n6),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn116)).
fof(c_0_514, axiom, (rdn_translate(nn117,rdn_neg(rdn(rdnn(n7),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn117)).
fof(c_0_515, axiom, (rdn_translate(nn118,rdn_neg(rdn(rdnn(n8),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn118)).
fof(c_0_516, axiom, (rdn_translate(nn119,rdn_neg(rdn(rdnn(n9),rdn(rdnn(n1),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn119)).
fof(c_0_517, axiom, (rdn_translate(nn120,rdn_neg(rdn(rdnn(n0),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn120)).
fof(c_0_518, axiom, (rdn_translate(nn121,rdn_neg(rdn(rdnn(n1),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn121)).
fof(c_0_519, axiom, (rdn_translate(nn122,rdn_neg(rdn(rdnn(n2),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn122)).
fof(c_0_520, axiom, (rdn_translate(nn123,rdn_neg(rdn(rdnn(n3),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn123)).
fof(c_0_521, axiom, (rdn_translate(nn124,rdn_neg(rdn(rdnn(n4),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn124)).
fof(c_0_522, axiom, (rdn_translate(nn125,rdn_neg(rdn(rdnn(n5),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn125)).
fof(c_0_523, axiom, (rdn_translate(nn126,rdn_neg(rdn(rdnn(n6),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn126)).
fof(c_0_524, axiom, (rdn_translate(nn127,rdn_neg(rdn(rdnn(n7),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn127)).
fof(c_0_525, axiom, (rdn_translate(nn128,rdn_neg(rdn(rdnn(n8),rdn(rdnn(n2),rdnn(n1)))))), file('Axioms/NUM005+0.ax', rdnn128)).
fof(c_0_526, axiom, (rdn_non_zero_digit(rdnn(n1))), file('Axioms/NUM005+1.ax', rdn_digit1)).
fof(c_0_527, axiom, (rdn_non_zero_digit(rdnn(n2))), file('Axioms/NUM005+1.ax', rdn_digit2)).
fof(c_0_528, axiom, (rdn_non_zero_digit(rdnn(n3))), file('Axioms/NUM005+1.ax', rdn_digit3)).
fof(c_0_529, axiom, (rdn_non_zero_digit(rdnn(n4))), file('Axioms/NUM005+1.ax', rdn_digit4)).
fof(c_0_530, axiom, (rdn_non_zero_digit(rdnn(n5))), file('Axioms/NUM005+1.ax', rdn_digit5)).
fof(c_0_531, axiom, (rdn_non_zero_digit(rdnn(n6))), file('Axioms/NUM005+1.ax', rdn_digit6)).
fof(c_0_532, axiom, (rdn_non_zero_digit(rdnn(n7))), file('Axioms/NUM005+1.ax', rdn_digit7)).
fof(c_0_533, axiom, (rdn_non_zero_digit(rdnn(n8))), file('Axioms/NUM005+1.ax', rdn_digit8)).
fof(c_0_534, axiom, (rdn_non_zero_digit(rdnn(n9))), file('Axioms/NUM005+1.ax', rdn_digit9)).
fof(c_0_535, axiom, (rdn_positive_less(rdnn(n0),rdnn(n1))), file('Axioms/NUM005+1.ax', rdn_positive_less01)).
fof(c_0_536, axiom, (rdn_positive_less(rdnn(n1),rdnn(n2))), file('Axioms/NUM005+1.ax', rdn_positive_less12)).
fof(c_0_537, axiom, (rdn_positive_less(rdnn(n2),rdnn(n3))), file('Axioms/NUM005+1.ax', rdn_positive_less23)).
fof(c_0_538, axiom, (rdn_positive_less(rdnn(n3),rdnn(n4))), file('Axioms/NUM005+1.ax', rdn_positive_less34)).
fof(c_0_539, axiom, (rdn_positive_less(rdnn(n4),rdnn(n5))), file('Axioms/NUM005+1.ax', rdn_positive_less45)).
fof(c_0_540, axiom, (rdn_positive_less(rdnn(n5),rdnn(n6))), file('Axioms/NUM005+1.ax', rdn_positive_less56)).
fof(c_0_541, axiom, (rdn_positive_less(rdnn(n6),rdnn(n7))), file('Axioms/NUM005+1.ax', rdn_positive_less67)).
fof(c_0_542, axiom, (rdn_positive_less(rdnn(n7),rdnn(n8))), file('Axioms/NUM005+1.ax', rdn_positive_less78)).
fof(c_0_543, axiom, (rdn_positive_less(rdnn(n8),rdnn(n9))), file('Axioms/NUM005+1.ax', rdn_positive_less89)).
fof(c_0_544, axiom, (![X7]:![X8]:![X9]:((rdn_positive_less(rdnn(X7),rdnn(X8))&rdn_positive_less(rdnn(X8),rdnn(X9)))=>rdn_positive_less(rdnn(X7),rdnn(X9)))), file('Axioms/NUM005+1.ax', rdn_positive_less_transitivity)).
fof(c_0_545, axiom, (![X10]:![X11]:![X12]:![X13]:(rdn_positive_less(X11,X13)=>rdn_positive_less(rdn(rdnn(X10),X11),rdn(rdnn(X12),X13)))), file('Axioms/NUM005+1.ax', rdn_positive_less_multi_digit_high)).
fof(c_0_546, axiom, (![X10]:![X14]:![X12]:((rdn_positive_less(rdnn(X10),rdnn(X12))&rdn_non_zero(X14))=>rdn_positive_less(rdn(rdnn(X10),X14),rdn(rdnn(X12),X14)))), file('Axioms/NUM005+1.ax', rdn_positive_less_multi_digit_low)).
fof(c_0_547, axiom, (![X15]:![X12]:![X13]:(rdn_non_zero(X13)=>rdn_positive_less(rdnn(X15),rdn(rdnn(X12),X13)))), file('Axioms/NUM005+1.ax', rdn_extra_digits_positive_less)).
fof(c_0_548, axiom, (![X7]:(rdn_non_zero_digit(rdnn(X7))=>rdn_non_zero(rdnn(X7)))), file('Axioms/NUM005+1.ax', rdn_non_zero_by_digit)).
fof(c_0_549, axiom, (![X15]:![X14]:(rdn_non_zero(X14)=>rdn_non_zero(rdn(rdnn(X15),X14)))), file('Axioms/NUM005+1.ax', rdn_non_zero_by_structure)).
fof(c_0_550, axiom, (![X7]:![X8]:![X16]:![X17]:(((rdn_translate(X7,rdn_pos(X16))&rdn_translate(X8,rdn_pos(X17)))&rdn_positive_less(X16,X17))=>less(X7,X8))), file('Axioms/NUM005+1.ax', less_entry_point_pos_pos)).
fof(c_0_551, axiom, (![X7]:![X8]:![X16]:![X17]:((rdn_translate(X7,rdn_neg(X16))&rdn_translate(X8,rdn_pos(X17)))=>less(X7,X8))), file('Axioms/NUM005+1.ax', less_entry_point_neg_pos)).
fof(c_0_552, axiom, (![X7]:![X8]:![X16]:![X17]:(((rdn_translate(X7,rdn_neg(X16))&rdn_translate(X8,rdn_neg(X17)))&rdn_positive_less(X17,X16))=>less(X7,X8))), file('Axioms/NUM005+1.ax', less_entry_point_neg_neg)).
fof(c_0_553, axiom, (![X7]:![X8]:(less(X7,X8)<=>(~(less(X8,X7))&X8!=X7))), file('Axioms/NUM005+1.ax', less_property)).
fof(c_0_554, axiom, (![X7]:![X8]:(less_or_equal(X7,X8)<=>(less(X7,X8)|X7=X8))), file('Axioms/NUM005+1.ax', less_or_equal)).
fof(c_0_555, axiom, (![X7]:![X8]:![X9]:((sum(X7,n1,X8)&less(X9,X8))=>less_or_equal(X9,X7))), file('Axioms/NUM005+1.ax', less_successor)).
fof(c_0_556, conjecture, (accept_team(countryamedicalorganization,countryacivilorganization,towna,n6)), file('<stdin>', query_1)).
fof(c_0_557, negated_conjecture, (~(accept_team(countryamedicalorganization,countryacivilorganization,towna,n6))),inference(assume_negation, [status(cth)],[c_0_556])).
fof(c_0_558, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n4)),inference(fof_simplification, [status(thm)],[c_0_29])).
fof(c_0_559, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n5)),inference(fof_simplification, [status(thm)],[c_0_30])).
fof(c_0_560, plain, (~accept_team(countryamedicalorganization,countryahumanitarianorganization,coastvillage,n2)),inference(fof_simplification, [status(thm)],[c_0_49])).
fof(c_0_561, plain, (~accept_team(countryamedicalorganization,countryafirstaidorganization,coastvillage,n2)),inference(fof_simplification, [status(thm)],[c_0_50])).
fof(c_0_562, plain, (~accept_team(countryamedicalorganization,countryacivilorganization,coastvillage,n2)),inference(fof_simplification, [status(thm)],[c_0_51])).
fof(c_0_563, plain, (~accept_team(countryamedicalorganization,christiansufferterrahumanitarianorganization,coastvillage,n2)),inference(fof_simplification, [status(thm)],[c_0_52])).
fof(c_0_564, plain, (~accept_team(countryamedicalorganization,sufferterragovernment,coastvillage,n2)),inference(fof_simplification, [status(thm)],[c_0_53])).
fof(c_0_565, plain, (~accept_team(countryamedicalorganization,countryahumanitarianorganization,coastvillage,n6)),inference(fof_simplification, [status(thm)],[c_0_54])).
fof(c_0_566, plain, (~accept_team(countryamedicalorganization,countryacivilorganization,coastvillage,n6)),inference(fof_simplification, [status(thm)],[c_0_55])).
fof(c_0_567, plain, (~accept_team(countryamedicalorganization,countryafirstaidorganization,coastvillage,n6)),inference(fof_simplification, [status(thm)],[c_0_56])).
fof(c_0_568, plain, (~accept_team(countryccivilorganization,countrycmedicalorganization,towna,n5)),inference(fof_simplification, [status(thm)],[c_0_60])).
fof(c_0_569, plain, (~accept_team(countryccivilorganization,christiancountrychumanitarianorganization,coastvillage,n5)),inference(fof_simplification, [status(thm)],[c_0_91])).
fof(c_0_570, plain, (~accept_team(countryccivilorganization,christiancountrychumanitarianorganization,coastvillage,n6)),inference(fof_simplification, [status(thm)],[c_0_92])).
fof(c_0_571, plain, (~accept_team(countryccivilorganization,muslimcountrybhumanitarianorganization,townc,n6)),inference(fof_simplification, [status(thm)],[c_0_135])).
fof(c_0_572, plain, (~accept_team(countryccivilorganization,countrybcivilorganization,townc,n6)),inference(fof_simplification, [status(thm)],[c_0_136])).
fof(c_0_573, plain, (~accept_team(countryccivilorganization,countrybhumanitarianorganization,townc,n6)),inference(fof_simplification, [status(thm)],[c_0_137])).
fof(c_0_574, plain, (~accept_team(countryccivilorganization,christiancountrychumanitarianorganization,townc,n6)),inference(fof_simplification, [status(thm)],[c_0_138])).
fof(c_0_575, plain, (~accept_team(countryccivilorganization,countrycmedicalorganization,townc,n6)),inference(fof_simplification, [status(thm)],[c_0_139])).
fof(c_0_576, plain, (~accept_team(countryccivilorganization,countrybhumanitarianorganization,cityb,n5)),inference(fof_simplification, [status(thm)],[c_0_163])).
fof(c_0_577, plain, (~accept_team(countryccivilorganization,countrybhumanitarianorganization,cityb,n6)),inference(fof_simplification, [status(thm)],[c_0_164])).
fof(c_0_578, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,towna,n2)),inference(fof_simplification, [status(thm)],[c_0_194])).
fof(c_0_579, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n2)),inference(fof_simplification, [status(thm)],[c_0_195])).
fof(c_0_580, plain, (~accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n2)),inference(fof_simplification, [status(thm)],[c_0_196])).
fof(c_0_581, plain, (~accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,towna,n2)),inference(fof_simplification, [status(thm)],[c_0_197])).
fof(c_0_582, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,towna,n3)),inference(fof_simplification, [status(thm)],[c_0_198])).
fof(c_0_583, plain, (~accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n3)),inference(fof_simplification, [status(thm)],[c_0_199])).
fof(c_0_584, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n3)),inference(fof_simplification, [status(thm)],[c_0_200])).
fof(c_0_585, plain, (~accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,towna,n3)),inference(fof_simplification, [status(thm)],[c_0_201])).
fof(c_0_586, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,towna,n2)),inference(fof_simplification, [status(thm)],[c_0_202])).
fof(c_0_587, plain, (~accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,towna,n4)),inference(fof_simplification, [status(thm)],[c_0_204])).
fof(c_0_588, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_205])).
fof(c_0_589, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_206])).
fof(c_0_590, plain, (~accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_207])).
fof(c_0_591, plain, (~accept_team(christiansufferterrahumanitarianorganization,sufferterragovernment,towna,n2)),inference(fof_simplification, [status(thm)],[c_0_232])).
fof(c_0_592, plain, (~accept_team(christiansufferterrahumanitarianorganization,sufferterragovernment,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_233])).
fof(c_0_593, plain, (~accept_team(countryccivilorganization,countrybcivilorganization,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_236])).
fof(c_0_594, plain, (~accept_team(countryccivilorganization,countrybhumanitarianorganization,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_237])).
fof(c_0_595, plain, (~accept_team(countryccivilorganization,christiancountrychumanitarianorganization,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_238])).
fof(c_0_596, plain, (~accept_team(countryccivilorganization,countrycmedicalorganization,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_239])).
fof(c_0_597, plain, (~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_266])).
fof(c_0_598, plain, (![X7]:![X8]:(less(X7,X8)<=>(~less(X8,X7)&X8!=X7))),inference(fof_simplification, [status(thm)],[c_0_553])).
fof(c_0_599, negated_conjecture, (~accept_team(countryamedicalorganization,countryacivilorganization,towna,n6)),inference(fof_simplification, [status(thm)],[c_0_557])).
fof(c_0_600, plain, (![X1]:![X2]:![X3]:![X4]:((~accept_team(X1,X4,X2,X3)|((accept_city(X1,X2)&accept_leader(X1,X4))&accept_number(X1,X3)))&(((~accept_city(X1,X2)|~accept_leader(X1,X4))|~accept_number(X1,X3))|accept_team(X1,X4,X2,X3)))),inference(fof_nnf, [status(thm)],[c_0_1])).
fof(c_0_601, plain, ((![X1]:![X2]:![X3]:![X4]:(~accept_team(X1,X4,X2,X3)|((accept_city(X1,X2)&accept_leader(X1,X4))&accept_number(X1,X3)))&![X1]:![X2]:![X3]:![X4]:(((~accept_city(X1,X2)|~accept_leader(X1,X4))|~accept_number(X1,X3))|accept_team(X1,X4,X2,X3)))),inference(shift_quantors, [status(thm)],[c_0_600])).
fof(c_0_602, plain, ((![X5]:![X6]:![X7]:![X8]:(~accept_team(X5,X8,X6,X7)|((accept_city(X5,X6)&accept_leader(X5,X8))&accept_number(X5,X7)))&![X9]:![X10]:![X11]:![X12]:(((~accept_city(X9,X10)|~accept_leader(X9,X12))|~accept_number(X9,X11))|accept_team(X9,X12,X10,X11)))),inference(variable_rename, [status(thm)],[c_0_601])).
fof(c_0_603, plain, (![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((~accept_team(X5,X8,X6,X7)|((accept_city(X5,X6)&accept_leader(X5,X8))&accept_number(X5,X7)))&(((~accept_city(X9,X10)|~accept_leader(X9,X12))|~accept_number(X9,X11))|accept_team(X9,X12,X10,X11)))),inference(shift_quantors, [status(thm)],[c_0_602])).
fof(c_0_604, plain, (![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((((accept_city(X5,X6)|~accept_team(X5,X8,X6,X7))&(accept_leader(X5,X8)|~accept_team(X5,X8,X6,X7)))&(accept_number(X5,X7)|~accept_team(X5,X8,X6,X7)))&(((~accept_city(X9,X10)|~accept_leader(X9,X12))|~accept_number(X9,X11))|accept_team(X9,X12,X10,X11)))),inference(distribute, [status(thm)],[c_0_603])).
cnf(c_0_605,plain,(accept_team(X1,X2,X3,X4)|~accept_number(X1,X4)|~accept_leader(X1,X2)|~accept_city(X1,X3)),inference(split_conjunct, [status(thm)],[c_0_604])).
cnf(c_0_606,plain,(accept_number(X1,X4)|~accept_team(X1,X2,X3,X4)),inference(split_conjunct, [status(thm)],[c_0_604])).
cnf(c_0_607,plain,(accept_leader(X1,X2)|~accept_team(X1,X2,X3,X4)),inference(split_conjunct, [status(thm)],[c_0_604])).
cnf(c_0_608,plain,(accept_city(X1,X3)|~accept_team(X1,X2,X3,X4)),inference(split_conjunct, [status(thm)],[c_0_604])).
fof(c_0_609, plain, (![X1]:![X3]:![X5]:((~accept_number(X1,X3)|~less(X5,X3))|accept_number(X1,X5))),inference(fof_nnf, [status(thm)],[c_0_2])).
fof(c_0_610, plain, (![X6]:![X7]:![X8]:((~accept_number(X6,X7)|~less(X8,X7))|accept_number(X6,X8))),inference(variable_rename, [status(thm)],[c_0_609])).
cnf(c_0_611,plain,(accept_number(X1,X2)|~less(X2,X3)|~accept_number(X1,X3)),inference(split_conjunct, [status(thm)],[c_0_610])).
fof(c_0_612, plain, (![X1]:![X3]:![X5]:![X6]:((~accept_population(X1,X6,X3)|~less(X5,X3))|accept_population(X1,X6,X5))),inference(fof_nnf, [status(thm)],[c_0_3])).
fof(c_0_613, plain, (![X7]:![X8]:![X9]:![X10]:((~accept_population(X7,X10,X8)|~less(X9,X8))|accept_population(X7,X10,X9))),inference(variable_rename, [status(thm)],[c_0_612])).
cnf(c_0_614,plain,(accept_population(X1,X2,X3)|~less(X3,X4)|~accept_population(X1,X2,X4)),inference(split_conjunct, [status(thm)],[c_0_613])).
fof(c_0_615, plain, (![X1]:![X4]:![X2]:(~the_agent_in_all_proposed_teams(X1,X4,X2)|(accept_leader(X1,X4)&accept_city(X1,X2)))),inference(fof_nnf, [status(thm)],[c_0_4])).
fof(c_0_616, plain, (![X5]:![X6]:![X7]:(~the_agent_in_all_proposed_teams(X5,X6,X7)|(accept_leader(X5,X6)&accept_city(X5,X7)))),inference(variable_rename, [status(thm)],[c_0_615])).
fof(c_0_617, plain, (![X5]:![X6]:![X7]:((accept_leader(X5,X6)|~the_agent_in_all_proposed_teams(X5,X6,X7))&(accept_city(X5,X7)|~the_agent_in_all_proposed_teams(X5,X6,X7)))),inference(distribute, [status(thm)],[c_0_616])).
cnf(c_0_618,plain,(accept_city(X1,X3)|~the_agent_in_all_proposed_teams(X1,X2,X3)),inference(split_conjunct, [status(thm)],[c_0_617])).
cnf(c_0_619,plain,(accept_leader(X1,X2)|~the_agent_in_all_proposed_teams(X1,X2,X3)),inference(split_conjunct, [status(thm)],[c_0_617])).
fof(c_0_620, plain, (![X1]:![X4]:![X2]:(~any_agent_in_all_proposed_teams(X1,X4,X2)|accept_leader(X1,X4))),inference(fof_nnf, [status(thm)],[c_0_5])).
fof(c_0_621, plain, (![X1]:![X4]:(![X2]:~any_agent_in_all_proposed_teams(X1,X4,X2)|accept_leader(X1,X4))),inference(shift_quantors, [status(thm)],[c_0_620])).
fof(c_0_622, plain, (![X5]:![X6]:(![X7]:~any_agent_in_all_proposed_teams(X5,X6,X7)|accept_leader(X5,X6))),inference(variable_rename, [status(thm)],[c_0_621])).
fof(c_0_623, plain, (![X5]:![X6]:![X7]:(~any_agent_in_all_proposed_teams(X5,X6,X7)|accept_leader(X5,X6))),inference(shift_quantors, [status(thm)],[c_0_622])).
cnf(c_0_624,plain,(accept_leader(X1,X2)|~any_agent_in_all_proposed_teams(X1,X2,X3)),inference(split_conjunct, [status(thm)],[c_0_623])).
fof(c_0_625, plain, (![X1]:![X4]:![X2]:(~the_agent_not_in_any_proposed_teams(X1,X4,X2)|(~accept_city(X1,X2)|~accept_leader(X1,X4)))),inference(fof_nnf, [status(thm)],[c_0_6])).
fof(c_0_626, plain, (![X5]:![X6]:![X7]:(~the_agent_not_in_any_proposed_teams(X5,X6,X7)|(~accept_city(X5,X7)|~accept_leader(X5,X6)))),inference(variable_rename, [status(thm)],[c_0_625])).
cnf(c_0_627,plain,(~accept_leader(X1,X2)|~accept_city(X1,X3)|~the_agent_not_in_any_proposed_teams(X1,X2,X3)),inference(split_conjunct, [status(thm)],[c_0_626])).
fof(c_0_628, plain, (![X1]:![X3]:(~min_number_of_proposed_agents(X1,X3)|accept_number(X1,X3))),inference(fof_nnf, [status(thm)],[c_0_7])).
fof(c_0_629, plain, (![X4]:![X5]:(~min_number_of_proposed_agents(X4,X5)|accept_number(X4,X5))),inference(variable_rename, [status(thm)],[c_0_628])).
cnf(c_0_630,plain,(accept_number(X1,X2)|~min_number_of_proposed_agents(X1,X2)),inference(split_conjunct, [status(thm)],[c_0_629])).
fof(c_0_631, plain, (![X1]:((((((~accept_population(X1,atheist,n65)|~accept_population(X1,christian,n20))|~accept_population(X1,muslim,n7))|~accept_population(X1,native,n4))|~accept_population(X1,other,n4))|accept_city(X1,suffertown))&(~accept_city(X1,suffertown)|((((accept_population(X1,atheist,n65)&accept_population(X1,christian,n20))&accept_population(X1,muslim,n7))&accept_population(X1,native,n4))&accept_population(X1,other,n4))))),inference(fof_nnf, [status(thm)],[c_0_8])).
fof(c_0_632, plain, ((![X1]:(((((~accept_population(X1,atheist,n65)|~accept_population(X1,christian,n20))|~accept_population(X1,muslim,n7))|~accept_population(X1,native,n4))|~accept_population(X1,other,n4))|accept_city(X1,suffertown))&![X1]:(~accept_city(X1,suffertown)|((((accept_population(X1,atheist,n65)&accept_population(X1,christian,n20))&accept_population(X1,muslim,n7))&accept_population(X1,native,n4))&accept_population(X1,other,n4))))),inference(shift_quantors, [status(thm)],[c_0_631])).
fof(c_0_633, plain, ((![X2]:(((((~accept_population(X2,atheist,n65)|~accept_population(X2,christian,n20))|~accept_population(X2,muslim,n7))|~accept_population(X2,native,n4))|~accept_population(X2,other,n4))|accept_city(X2,suffertown))&![X3]:(~accept_city(X3,suffertown)|((((accept_population(X3,atheist,n65)&accept_population(X3,christian,n20))&accept_population(X3,muslim,n7))&accept_population(X3,native,n4))&accept_population(X3,other,n4))))),inference(variable_rename, [status(thm)],[c_0_632])).
fof(c_0_634, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n65)|~accept_population(X2,christian,n20))|~accept_population(X2,muslim,n7))|~accept_population(X2,native,n4))|~accept_population(X2,other,n4))|accept_city(X2,suffertown))&(~accept_city(X3,suffertown)|((((accept_population(X3,atheist,n65)&accept_population(X3,christian,n20))&accept_population(X3,muslim,n7))&accept_population(X3,native,n4))&accept_population(X3,other,n4))))),inference(shift_quantors, [status(thm)],[c_0_633])).
fof(c_0_635, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n65)|~accept_population(X2,christian,n20))|~accept_population(X2,muslim,n7))|~accept_population(X2,native,n4))|~accept_population(X2,other,n4))|accept_city(X2,suffertown))&(((((accept_population(X3,atheist,n65)|~accept_city(X3,suffertown))&(accept_population(X3,christian,n20)|~accept_city(X3,suffertown)))&(accept_population(X3,muslim,n7)|~accept_city(X3,suffertown)))&(accept_population(X3,native,n4)|~accept_city(X3,suffertown)))&(accept_population(X3,other,n4)|~accept_city(X3,suffertown))))),inference(distribute, [status(thm)],[c_0_634])).
cnf(c_0_636,plain,(accept_population(X1,other,n4)|~accept_city(X1,suffertown)),inference(split_conjunct, [status(thm)],[c_0_635])).
cnf(c_0_637,plain,(accept_population(X1,native,n4)|~accept_city(X1,suffertown)),inference(split_conjunct, [status(thm)],[c_0_635])).
cnf(c_0_638,plain,(accept_population(X1,muslim,n7)|~accept_city(X1,suffertown)),inference(split_conjunct, [status(thm)],[c_0_635])).
cnf(c_0_639,plain,(accept_population(X1,christian,n20)|~accept_city(X1,suffertown)),inference(split_conjunct, [status(thm)],[c_0_635])).
cnf(c_0_640,plain,(accept_population(X1,atheist,n65)|~accept_city(X1,suffertown)),inference(split_conjunct, [status(thm)],[c_0_635])).
cnf(c_0_641,plain,(accept_city(X1,suffertown)|~accept_population(X1,other,n4)|~accept_population(X1,native,n4)|~accept_population(X1,muslim,n7)|~accept_population(X1,christian,n20)|~accept_population(X1,atheist,n65)),inference(split_conjunct, [status(thm)],[c_0_635])).
fof(c_0_642, plain, (![X1]:((((((~accept_population(X1,atheist,n54)|~accept_population(X1,christian,n23))|~accept_population(X1,muslim,n3))|~accept_population(X1,native,n1))|~accept_population(X1,other,n9))|accept_city(X1,centraltown))&(~accept_city(X1,centraltown)|((((accept_population(X1,atheist,n54)&accept_population(X1,christian,n23))&accept_population(X1,muslim,n3))&accept_population(X1,native,n1))&accept_population(X1,other,n9))))),inference(fof_nnf, [status(thm)],[c_0_9])).
fof(c_0_643, plain, ((![X1]:(((((~accept_population(X1,atheist,n54)|~accept_population(X1,christian,n23))|~accept_population(X1,muslim,n3))|~accept_population(X1,native,n1))|~accept_population(X1,other,n9))|accept_city(X1,centraltown))&![X1]:(~accept_city(X1,centraltown)|((((accept_population(X1,atheist,n54)&accept_population(X1,christian,n23))&accept_population(X1,muslim,n3))&accept_population(X1,native,n1))&accept_population(X1,other,n9))))),inference(shift_quantors, [status(thm)],[c_0_642])).
fof(c_0_644, plain, ((![X2]:(((((~accept_population(X2,atheist,n54)|~accept_population(X2,christian,n23))|~accept_population(X2,muslim,n3))|~accept_population(X2,native,n1))|~accept_population(X2,other,n9))|accept_city(X2,centraltown))&![X3]:(~accept_city(X3,centraltown)|((((accept_population(X3,atheist,n54)&accept_population(X3,christian,n23))&accept_population(X3,muslim,n3))&accept_population(X3,native,n1))&accept_population(X3,other,n9))))),inference(variable_rename, [status(thm)],[c_0_643])).
fof(c_0_645, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n54)|~accept_population(X2,christian,n23))|~accept_population(X2,muslim,n3))|~accept_population(X2,native,n1))|~accept_population(X2,other,n9))|accept_city(X2,centraltown))&(~accept_city(X3,centraltown)|((((accept_population(X3,atheist,n54)&accept_population(X3,christian,n23))&accept_population(X3,muslim,n3))&accept_population(X3,native,n1))&accept_population(X3,other,n9))))),inference(shift_quantors, [status(thm)],[c_0_644])).
fof(c_0_646, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n54)|~accept_population(X2,christian,n23))|~accept_population(X2,muslim,n3))|~accept_population(X2,native,n1))|~accept_population(X2,other,n9))|accept_city(X2,centraltown))&(((((accept_population(X3,atheist,n54)|~accept_city(X3,centraltown))&(accept_population(X3,christian,n23)|~accept_city(X3,centraltown)))&(accept_population(X3,muslim,n3)|~accept_city(X3,centraltown)))&(accept_population(X3,native,n1)|~accept_city(X3,centraltown)))&(accept_population(X3,other,n9)|~accept_city(X3,centraltown))))),inference(distribute, [status(thm)],[c_0_645])).
cnf(c_0_647,plain,(accept_population(X1,other,n9)|~accept_city(X1,centraltown)),inference(split_conjunct, [status(thm)],[c_0_646])).
cnf(c_0_648,plain,(accept_population(X1,native,n1)|~accept_city(X1,centraltown)),inference(split_conjunct, [status(thm)],[c_0_646])).
cnf(c_0_649,plain,(accept_population(X1,muslim,n3)|~accept_city(X1,centraltown)),inference(split_conjunct, [status(thm)],[c_0_646])).
cnf(c_0_650,plain,(accept_population(X1,christian,n23)|~accept_city(X1,centraltown)),inference(split_conjunct, [status(thm)],[c_0_646])).
cnf(c_0_651,plain,(accept_population(X1,atheist,n54)|~accept_city(X1,centraltown)),inference(split_conjunct, [status(thm)],[c_0_646])).
cnf(c_0_652,plain,(accept_city(X1,centraltown)|~accept_population(X1,other,n9)|~accept_population(X1,native,n1)|~accept_population(X1,muslim,n3)|~accept_population(X1,christian,n23)|~accept_population(X1,atheist,n54)),inference(split_conjunct, [status(thm)],[c_0_646])).
fof(c_0_653, plain, (![X1]:((((((~accept_population(X1,atheist,n30)|~accept_population(X1,christian,n8))|~accept_population(X1,muslim,n60))|~accept_population(X1,native,n1))|~accept_population(X1,other,n1))|accept_city(X1,sunnysideport))&(~accept_city(X1,sunnysideport)|((((accept_population(X1,atheist,n30)&accept_population(X1,christian,n8))&accept_population(X1,muslim,n60))&accept_population(X1,native,n1))&accept_population(X1,other,n1))))),inference(fof_nnf, [status(thm)],[c_0_10])).
fof(c_0_654, plain, ((![X1]:(((((~accept_population(X1,atheist,n30)|~accept_population(X1,christian,n8))|~accept_population(X1,muslim,n60))|~accept_population(X1,native,n1))|~accept_population(X1,other,n1))|accept_city(X1,sunnysideport))&![X1]:(~accept_city(X1,sunnysideport)|((((accept_population(X1,atheist,n30)&accept_population(X1,christian,n8))&accept_population(X1,muslim,n60))&accept_population(X1,native,n1))&accept_population(X1,other,n1))))),inference(shift_quantors, [status(thm)],[c_0_653])).
fof(c_0_655, plain, ((![X2]:(((((~accept_population(X2,atheist,n30)|~accept_population(X2,christian,n8))|~accept_population(X2,muslim,n60))|~accept_population(X2,native,n1))|~accept_population(X2,other,n1))|accept_city(X2,sunnysideport))&![X3]:(~accept_city(X3,sunnysideport)|((((accept_population(X3,atheist,n30)&accept_population(X3,christian,n8))&accept_population(X3,muslim,n60))&accept_population(X3,native,n1))&accept_population(X3,other,n1))))),inference(variable_rename, [status(thm)],[c_0_654])).
fof(c_0_656, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n30)|~accept_population(X2,christian,n8))|~accept_population(X2,muslim,n60))|~accept_population(X2,native,n1))|~accept_population(X2,other,n1))|accept_city(X2,sunnysideport))&(~accept_city(X3,sunnysideport)|((((accept_population(X3,atheist,n30)&accept_population(X3,christian,n8))&accept_population(X3,muslim,n60))&accept_population(X3,native,n1))&accept_population(X3,other,n1))))),inference(shift_quantors, [status(thm)],[c_0_655])).
fof(c_0_657, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n30)|~accept_population(X2,christian,n8))|~accept_population(X2,muslim,n60))|~accept_population(X2,native,n1))|~accept_population(X2,other,n1))|accept_city(X2,sunnysideport))&(((((accept_population(X3,atheist,n30)|~accept_city(X3,sunnysideport))&(accept_population(X3,christian,n8)|~accept_city(X3,sunnysideport)))&(accept_population(X3,muslim,n60)|~accept_city(X3,sunnysideport)))&(accept_population(X3,native,n1)|~accept_city(X3,sunnysideport)))&(accept_population(X3,other,n1)|~accept_city(X3,sunnysideport))))),inference(distribute, [status(thm)],[c_0_656])).
cnf(c_0_658,plain,(accept_population(X1,other,n1)|~accept_city(X1,sunnysideport)),inference(split_conjunct, [status(thm)],[c_0_657])).
cnf(c_0_659,plain,(accept_population(X1,native,n1)|~accept_city(X1,sunnysideport)),inference(split_conjunct, [status(thm)],[c_0_657])).
cnf(c_0_660,plain,(accept_population(X1,muslim,n60)|~accept_city(X1,sunnysideport)),inference(split_conjunct, [status(thm)],[c_0_657])).
cnf(c_0_661,plain,(accept_population(X1,christian,n8)|~accept_city(X1,sunnysideport)),inference(split_conjunct, [status(thm)],[c_0_657])).
cnf(c_0_662,plain,(accept_population(X1,atheist,n30)|~accept_city(X1,sunnysideport)),inference(split_conjunct, [status(thm)],[c_0_657])).
cnf(c_0_663,plain,(accept_city(X1,sunnysideport)|~accept_population(X1,other,n1)|~accept_population(X1,native,n1)|~accept_population(X1,muslim,n60)|~accept_population(X1,christian,n8)|~accept_population(X1,atheist,n30)),inference(split_conjunct, [status(thm)],[c_0_657])).
fof(c_0_664, plain, (![X1]:((((((~accept_population(X1,atheist,n70)|~accept_population(X1,christian,n15))|~accept_population(X1,muslim,n1))|~accept_population(X1,native,n10))|~accept_population(X1,other,n4))|accept_city(X1,centrallakecity))&(~accept_city(X1,centrallakecity)|((((accept_population(X1,atheist,n70)&accept_population(X1,christian,n15))&accept_population(X1,muslim,n1))&accept_population(X1,native,n10))&accept_population(X1,other,n4))))),inference(fof_nnf, [status(thm)],[c_0_11])).
fof(c_0_665, plain, ((![X1]:(((((~accept_population(X1,atheist,n70)|~accept_population(X1,christian,n15))|~accept_population(X1,muslim,n1))|~accept_population(X1,native,n10))|~accept_population(X1,other,n4))|accept_city(X1,centrallakecity))&![X1]:(~accept_city(X1,centrallakecity)|((((accept_population(X1,atheist,n70)&accept_population(X1,christian,n15))&accept_population(X1,muslim,n1))&accept_population(X1,native,n10))&accept_population(X1,other,n4))))),inference(shift_quantors, [status(thm)],[c_0_664])).
fof(c_0_666, plain, ((![X2]:(((((~accept_population(X2,atheist,n70)|~accept_population(X2,christian,n15))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n10))|~accept_population(X2,other,n4))|accept_city(X2,centrallakecity))&![X3]:(~accept_city(X3,centrallakecity)|((((accept_population(X3,atheist,n70)&accept_population(X3,christian,n15))&accept_population(X3,muslim,n1))&accept_population(X3,native,n10))&accept_population(X3,other,n4))))),inference(variable_rename, [status(thm)],[c_0_665])).
fof(c_0_667, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n70)|~accept_population(X2,christian,n15))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n10))|~accept_population(X2,other,n4))|accept_city(X2,centrallakecity))&(~accept_city(X3,centrallakecity)|((((accept_population(X3,atheist,n70)&accept_population(X3,christian,n15))&accept_population(X3,muslim,n1))&accept_population(X3,native,n10))&accept_population(X3,other,n4))))),inference(shift_quantors, [status(thm)],[c_0_666])).
fof(c_0_668, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n70)|~accept_population(X2,christian,n15))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n10))|~accept_population(X2,other,n4))|accept_city(X2,centrallakecity))&(((((accept_population(X3,atheist,n70)|~accept_city(X3,centrallakecity))&(accept_population(X3,christian,n15)|~accept_city(X3,centrallakecity)))&(accept_population(X3,muslim,n1)|~accept_city(X3,centrallakecity)))&(accept_population(X3,native,n10)|~accept_city(X3,centrallakecity)))&(accept_population(X3,other,n4)|~accept_city(X3,centrallakecity))))),inference(distribute, [status(thm)],[c_0_667])).
cnf(c_0_669,plain,(accept_population(X1,other,n4)|~accept_city(X1,centrallakecity)),inference(split_conjunct, [status(thm)],[c_0_668])).
cnf(c_0_670,plain,(accept_population(X1,native,n10)|~accept_city(X1,centrallakecity)),inference(split_conjunct, [status(thm)],[c_0_668])).
cnf(c_0_671,plain,(accept_population(X1,muslim,n1)|~accept_city(X1,centrallakecity)),inference(split_conjunct, [status(thm)],[c_0_668])).
cnf(c_0_672,plain,(accept_population(X1,christian,n15)|~accept_city(X1,centrallakecity)),inference(split_conjunct, [status(thm)],[c_0_668])).
cnf(c_0_673,plain,(accept_population(X1,atheist,n70)|~accept_city(X1,centrallakecity)),inference(split_conjunct, [status(thm)],[c_0_668])).
cnf(c_0_674,plain,(accept_city(X1,centrallakecity)|~accept_population(X1,other,n4)|~accept_population(X1,native,n10)|~accept_population(X1,muslim,n1)|~accept_population(X1,christian,n15)|~accept_population(X1,atheist,n70)),inference(split_conjunct, [status(thm)],[c_0_668])).
fof(c_0_675, plain, (![X1]:((((((~accept_population(X1,atheist,n68)|~accept_population(X1,christian,n16))|~accept_population(X1,muslim,n1))|~accept_population(X1,native,n11))|~accept_population(X1,other,n4))|accept_city(X1,stjosephburgh))&(~accept_city(X1,stjosephburgh)|((((accept_population(X1,atheist,n68)&accept_population(X1,christian,n16))&accept_population(X1,muslim,n1))&accept_population(X1,native,n11))&accept_population(X1,other,n4))))),inference(fof_nnf, [status(thm)],[c_0_12])).
fof(c_0_676, plain, ((![X1]:(((((~accept_population(X1,atheist,n68)|~accept_population(X1,christian,n16))|~accept_population(X1,muslim,n1))|~accept_population(X1,native,n11))|~accept_population(X1,other,n4))|accept_city(X1,stjosephburgh))&![X1]:(~accept_city(X1,stjosephburgh)|((((accept_population(X1,atheist,n68)&accept_population(X1,christian,n16))&accept_population(X1,muslim,n1))&accept_population(X1,native,n11))&accept_population(X1,other,n4))))),inference(shift_quantors, [status(thm)],[c_0_675])).
fof(c_0_677, plain, ((![X2]:(((((~accept_population(X2,atheist,n68)|~accept_population(X2,christian,n16))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n11))|~accept_population(X2,other,n4))|accept_city(X2,stjosephburgh))&![X3]:(~accept_city(X3,stjosephburgh)|((((accept_population(X3,atheist,n68)&accept_population(X3,christian,n16))&accept_population(X3,muslim,n1))&accept_population(X3,native,n11))&accept_population(X3,other,n4))))),inference(variable_rename, [status(thm)],[c_0_676])).
fof(c_0_678, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n68)|~accept_population(X2,christian,n16))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n11))|~accept_population(X2,other,n4))|accept_city(X2,stjosephburgh))&(~accept_city(X3,stjosephburgh)|((((accept_population(X3,atheist,n68)&accept_population(X3,christian,n16))&accept_population(X3,muslim,n1))&accept_population(X3,native,n11))&accept_population(X3,other,n4))))),inference(shift_quantors, [status(thm)],[c_0_677])).
fof(c_0_679, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n68)|~accept_population(X2,christian,n16))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n11))|~accept_population(X2,other,n4))|accept_city(X2,stjosephburgh))&(((((accept_population(X3,atheist,n68)|~accept_city(X3,stjosephburgh))&(accept_population(X3,christian,n16)|~accept_city(X3,stjosephburgh)))&(accept_population(X3,muslim,n1)|~accept_city(X3,stjosephburgh)))&(accept_population(X3,native,n11)|~accept_city(X3,stjosephburgh)))&(accept_population(X3,other,n4)|~accept_city(X3,stjosephburgh))))),inference(distribute, [status(thm)],[c_0_678])).
cnf(c_0_680,plain,(accept_population(X1,other,n4)|~accept_city(X1,stjosephburgh)),inference(split_conjunct, [status(thm)],[c_0_679])).
cnf(c_0_681,plain,(accept_population(X1,native,n11)|~accept_city(X1,stjosephburgh)),inference(split_conjunct, [status(thm)],[c_0_679])).
cnf(c_0_682,plain,(accept_population(X1,muslim,n1)|~accept_city(X1,stjosephburgh)),inference(split_conjunct, [status(thm)],[c_0_679])).
cnf(c_0_683,plain,(accept_population(X1,christian,n16)|~accept_city(X1,stjosephburgh)),inference(split_conjunct, [status(thm)],[c_0_679])).
cnf(c_0_684,plain,(accept_population(X1,atheist,n68)|~accept_city(X1,stjosephburgh)),inference(split_conjunct, [status(thm)],[c_0_679])).
cnf(c_0_685,plain,(accept_city(X1,stjosephburgh)|~accept_population(X1,other,n4)|~accept_population(X1,native,n11)|~accept_population(X1,muslim,n1)|~accept_population(X1,christian,n16)|~accept_population(X1,atheist,n68)),inference(split_conjunct, [status(thm)],[c_0_679])).
fof(c_0_686, plain, (![X1]:((((((~accept_population(X1,atheist,n70)|~accept_population(X1,christian,n13))|~accept_population(X1,muslim,n0))|~accept_population(X1,native,n15))|~accept_population(X1,other,n2))|accept_city(X1,northport))&(~accept_city(X1,northport)|((((accept_population(X1,atheist,n70)&accept_population(X1,christian,n13))&accept_population(X1,muslim,n0))&accept_population(X1,native,n15))&accept_population(X1,other,n2))))),inference(fof_nnf, [status(thm)],[c_0_13])).
fof(c_0_687, plain, ((![X1]:(((((~accept_population(X1,atheist,n70)|~accept_population(X1,christian,n13))|~accept_population(X1,muslim,n0))|~accept_population(X1,native,n15))|~accept_population(X1,other,n2))|accept_city(X1,northport))&![X1]:(~accept_city(X1,northport)|((((accept_population(X1,atheist,n70)&accept_population(X1,christian,n13))&accept_population(X1,muslim,n0))&accept_population(X1,native,n15))&accept_population(X1,other,n2))))),inference(shift_quantors, [status(thm)],[c_0_686])).
fof(c_0_688, plain, ((![X2]:(((((~accept_population(X2,atheist,n70)|~accept_population(X2,christian,n13))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n15))|~accept_population(X2,other,n2))|accept_city(X2,northport))&![X3]:(~accept_city(X3,northport)|((((accept_population(X3,atheist,n70)&accept_population(X3,christian,n13))&accept_population(X3,muslim,n0))&accept_population(X3,native,n15))&accept_population(X3,other,n2))))),inference(variable_rename, [status(thm)],[c_0_687])).
fof(c_0_689, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n70)|~accept_population(X2,christian,n13))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n15))|~accept_population(X2,other,n2))|accept_city(X2,northport))&(~accept_city(X3,northport)|((((accept_population(X3,atheist,n70)&accept_population(X3,christian,n13))&accept_population(X3,muslim,n0))&accept_population(X3,native,n15))&accept_population(X3,other,n2))))),inference(shift_quantors, [status(thm)],[c_0_688])).
fof(c_0_690, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n70)|~accept_population(X2,christian,n13))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n15))|~accept_population(X2,other,n2))|accept_city(X2,northport))&(((((accept_population(X3,atheist,n70)|~accept_city(X3,northport))&(accept_population(X3,christian,n13)|~accept_city(X3,northport)))&(accept_population(X3,muslim,n0)|~accept_city(X3,northport)))&(accept_population(X3,native,n15)|~accept_city(X3,northport)))&(accept_population(X3,other,n2)|~accept_city(X3,northport))))),inference(distribute, [status(thm)],[c_0_689])).
cnf(c_0_691,plain,(accept_population(X1,other,n2)|~accept_city(X1,northport)),inference(split_conjunct, [status(thm)],[c_0_690])).
cnf(c_0_692,plain,(accept_population(X1,native,n15)|~accept_city(X1,northport)),inference(split_conjunct, [status(thm)],[c_0_690])).
cnf(c_0_693,plain,(accept_population(X1,muslim,n0)|~accept_city(X1,northport)),inference(split_conjunct, [status(thm)],[c_0_690])).
cnf(c_0_694,plain,(accept_population(X1,christian,n13)|~accept_city(X1,northport)),inference(split_conjunct, [status(thm)],[c_0_690])).
cnf(c_0_695,plain,(accept_population(X1,atheist,n70)|~accept_city(X1,northport)),inference(split_conjunct, [status(thm)],[c_0_690])).
cnf(c_0_696,plain,(accept_city(X1,northport)|~accept_population(X1,other,n2)|~accept_population(X1,native,n15)|~accept_population(X1,muslim,n0)|~accept_population(X1,christian,n13)|~accept_population(X1,atheist,n70)),inference(split_conjunct, [status(thm)],[c_0_690])).
fof(c_0_697, plain, (![X1]:((((((~accept_population(X1,atheist,n12)|~accept_population(X1,christian,n3))|~accept_population(X1,muslim,n0))|~accept_population(X1,native,n85))|~accept_population(X1,other,n0))|accept_city(X1,coastvillage))&(~accept_city(X1,coastvillage)|((((accept_population(X1,atheist,n12)&accept_population(X1,christian,n3))&accept_population(X1,muslim,n0))&accept_population(X1,native,n85))&accept_population(X1,other,n0))))),inference(fof_nnf, [status(thm)],[c_0_14])).
fof(c_0_698, plain, ((![X1]:(((((~accept_population(X1,atheist,n12)|~accept_population(X1,christian,n3))|~accept_population(X1,muslim,n0))|~accept_population(X1,native,n85))|~accept_population(X1,other,n0))|accept_city(X1,coastvillage))&![X1]:(~accept_city(X1,coastvillage)|((((accept_population(X1,atheist,n12)&accept_population(X1,christian,n3))&accept_population(X1,muslim,n0))&accept_population(X1,native,n85))&accept_population(X1,other,n0))))),inference(shift_quantors, [status(thm)],[c_0_697])).
fof(c_0_699, plain, ((![X2]:(((((~accept_population(X2,atheist,n12)|~accept_population(X2,christian,n3))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n85))|~accept_population(X2,other,n0))|accept_city(X2,coastvillage))&![X3]:(~accept_city(X3,coastvillage)|((((accept_population(X3,atheist,n12)&accept_population(X3,christian,n3))&accept_population(X3,muslim,n0))&accept_population(X3,native,n85))&accept_population(X3,other,n0))))),inference(variable_rename, [status(thm)],[c_0_698])).
fof(c_0_700, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n12)|~accept_population(X2,christian,n3))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n85))|~accept_population(X2,other,n0))|accept_city(X2,coastvillage))&(~accept_city(X3,coastvillage)|((((accept_population(X3,atheist,n12)&accept_population(X3,christian,n3))&accept_population(X3,muslim,n0))&accept_population(X3,native,n85))&accept_population(X3,other,n0))))),inference(shift_quantors, [status(thm)],[c_0_699])).
fof(c_0_701, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n12)|~accept_population(X2,christian,n3))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n85))|~accept_population(X2,other,n0))|accept_city(X2,coastvillage))&(((((accept_population(X3,atheist,n12)|~accept_city(X3,coastvillage))&(accept_population(X3,christian,n3)|~accept_city(X3,coastvillage)))&(accept_population(X3,muslim,n0)|~accept_city(X3,coastvillage)))&(accept_population(X3,native,n85)|~accept_city(X3,coastvillage)))&(accept_population(X3,other,n0)|~accept_city(X3,coastvillage))))),inference(distribute, [status(thm)],[c_0_700])).
cnf(c_0_702,plain,(accept_population(X1,other,n0)|~accept_city(X1,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_701])).
cnf(c_0_703,plain,(accept_population(X1,native,n85)|~accept_city(X1,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_701])).
cnf(c_0_704,plain,(accept_population(X1,muslim,n0)|~accept_city(X1,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_701])).
cnf(c_0_705,plain,(accept_population(X1,christian,n3)|~accept_city(X1,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_701])).
cnf(c_0_706,plain,(accept_population(X1,atheist,n12)|~accept_city(X1,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_701])).
cnf(c_0_707,plain,(accept_city(X1,coastvillage)|~accept_population(X1,other,n0)|~accept_population(X1,native,n85)|~accept_population(X1,muslim,n0)|~accept_population(X1,christian,n3)|~accept_population(X1,atheist,n12)),inference(split_conjunct, [status(thm)],[c_0_701])).
fof(c_0_708, plain, (![X1]:((((((~accept_population(X1,atheist,n0)|~accept_population(X1,christian,n0))|~accept_population(X1,muslim,n0))|~accept_population(X1,native,n100))|~accept_population(X1,other,n0))|accept_city(X1,sunsetpoint))&(~accept_city(X1,sunsetpoint)|((((accept_population(X1,atheist,n0)&accept_population(X1,christian,n0))&accept_population(X1,muslim,n0))&accept_population(X1,native,n100))&accept_population(X1,other,n0))))),inference(fof_nnf, [status(thm)],[c_0_15])).
fof(c_0_709, plain, ((![X1]:(((((~accept_population(X1,atheist,n0)|~accept_population(X1,christian,n0))|~accept_population(X1,muslim,n0))|~accept_population(X1,native,n100))|~accept_population(X1,other,n0))|accept_city(X1,sunsetpoint))&![X1]:(~accept_city(X1,sunsetpoint)|((((accept_population(X1,atheist,n0)&accept_population(X1,christian,n0))&accept_population(X1,muslim,n0))&accept_population(X1,native,n100))&accept_population(X1,other,n0))))),inference(shift_quantors, [status(thm)],[c_0_708])).
fof(c_0_710, plain, ((![X2]:(((((~accept_population(X2,atheist,n0)|~accept_population(X2,christian,n0))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n100))|~accept_population(X2,other,n0))|accept_city(X2,sunsetpoint))&![X3]:(~accept_city(X3,sunsetpoint)|((((accept_population(X3,atheist,n0)&accept_population(X3,christian,n0))&accept_population(X3,muslim,n0))&accept_population(X3,native,n100))&accept_population(X3,other,n0))))),inference(variable_rename, [status(thm)],[c_0_709])).
fof(c_0_711, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n0)|~accept_population(X2,christian,n0))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n100))|~accept_population(X2,other,n0))|accept_city(X2,sunsetpoint))&(~accept_city(X3,sunsetpoint)|((((accept_population(X3,atheist,n0)&accept_population(X3,christian,n0))&accept_population(X3,muslim,n0))&accept_population(X3,native,n100))&accept_population(X3,other,n0))))),inference(shift_quantors, [status(thm)],[c_0_710])).
fof(c_0_712, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n0)|~accept_population(X2,christian,n0))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n100))|~accept_population(X2,other,n0))|accept_city(X2,sunsetpoint))&(((((accept_population(X3,atheist,n0)|~accept_city(X3,sunsetpoint))&(accept_population(X3,christian,n0)|~accept_city(X3,sunsetpoint)))&(accept_population(X3,muslim,n0)|~accept_city(X3,sunsetpoint)))&(accept_population(X3,native,n100)|~accept_city(X3,sunsetpoint)))&(accept_population(X3,other,n0)|~accept_city(X3,sunsetpoint))))),inference(distribute, [status(thm)],[c_0_711])).
cnf(c_0_713,plain,(accept_population(X1,other,n0)|~accept_city(X1,sunsetpoint)),inference(split_conjunct, [status(thm)],[c_0_712])).
cnf(c_0_714,plain,(accept_population(X1,native,n100)|~accept_city(X1,sunsetpoint)),inference(split_conjunct, [status(thm)],[c_0_712])).
cnf(c_0_715,plain,(accept_population(X1,muslim,n0)|~accept_city(X1,sunsetpoint)),inference(split_conjunct, [status(thm)],[c_0_712])).
cnf(c_0_716,plain,(accept_population(X1,christian,n0)|~accept_city(X1,sunsetpoint)),inference(split_conjunct, [status(thm)],[c_0_712])).
cnf(c_0_717,plain,(accept_population(X1,atheist,n0)|~accept_city(X1,sunsetpoint)),inference(split_conjunct, [status(thm)],[c_0_712])).
cnf(c_0_718,plain,(accept_city(X1,sunsetpoint)|~accept_population(X1,other,n0)|~accept_population(X1,native,n100)|~accept_population(X1,muslim,n0)|~accept_population(X1,christian,n0)|~accept_population(X1,atheist,n0)),inference(split_conjunct, [status(thm)],[c_0_712])).
fof(c_0_719, plain, (![X1]:((((((~accept_population(X1,atheist,n75)|~accept_population(X1,christian,n24))|~accept_population(X1,muslim,n1))|~accept_population(X1,native,n0))|~accept_population(X1,other,n0))|accept_city(X1,towna))&(~accept_city(X1,towna)|((((accept_population(X1,atheist,n75)&accept_population(X1,christian,n24))&accept_population(X1,muslim,n1))&accept_population(X1,native,n0))&accept_population(X1,other,n0))))),inference(fof_nnf, [status(thm)],[c_0_16])).
fof(c_0_720, plain, ((![X1]:(((((~accept_population(X1,atheist,n75)|~accept_population(X1,christian,n24))|~accept_population(X1,muslim,n1))|~accept_population(X1,native,n0))|~accept_population(X1,other,n0))|accept_city(X1,towna))&![X1]:(~accept_city(X1,towna)|((((accept_population(X1,atheist,n75)&accept_population(X1,christian,n24))&accept_population(X1,muslim,n1))&accept_population(X1,native,n0))&accept_population(X1,other,n0))))),inference(shift_quantors, [status(thm)],[c_0_719])).
fof(c_0_721, plain, ((![X2]:(((((~accept_population(X2,atheist,n75)|~accept_population(X2,christian,n24))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n0))|~accept_population(X2,other,n0))|accept_city(X2,towna))&![X3]:(~accept_city(X3,towna)|((((accept_population(X3,atheist,n75)&accept_population(X3,christian,n24))&accept_population(X3,muslim,n1))&accept_population(X3,native,n0))&accept_population(X3,other,n0))))),inference(variable_rename, [status(thm)],[c_0_720])).
fof(c_0_722, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n75)|~accept_population(X2,christian,n24))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n0))|~accept_population(X2,other,n0))|accept_city(X2,towna))&(~accept_city(X3,towna)|((((accept_population(X3,atheist,n75)&accept_population(X3,christian,n24))&accept_population(X3,muslim,n1))&accept_population(X3,native,n0))&accept_population(X3,other,n0))))),inference(shift_quantors, [status(thm)],[c_0_721])).
fof(c_0_723, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n75)|~accept_population(X2,christian,n24))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n0))|~accept_population(X2,other,n0))|accept_city(X2,towna))&(((((accept_population(X3,atheist,n75)|~accept_city(X3,towna))&(accept_population(X3,christian,n24)|~accept_city(X3,towna)))&(accept_population(X3,muslim,n1)|~accept_city(X3,towna)))&(accept_population(X3,native,n0)|~accept_city(X3,towna)))&(accept_population(X3,other,n0)|~accept_city(X3,towna))))),inference(distribute, [status(thm)],[c_0_722])).
cnf(c_0_724,plain,(accept_population(X1,other,n0)|~accept_city(X1,towna)),inference(split_conjunct, [status(thm)],[c_0_723])).
cnf(c_0_725,plain,(accept_population(X1,native,n0)|~accept_city(X1,towna)),inference(split_conjunct, [status(thm)],[c_0_723])).
cnf(c_0_726,plain,(accept_population(X1,muslim,n1)|~accept_city(X1,towna)),inference(split_conjunct, [status(thm)],[c_0_723])).
cnf(c_0_727,plain,(accept_population(X1,christian,n24)|~accept_city(X1,towna)),inference(split_conjunct, [status(thm)],[c_0_723])).
cnf(c_0_728,plain,(accept_population(X1,atheist,n75)|~accept_city(X1,towna)),inference(split_conjunct, [status(thm)],[c_0_723])).
cnf(c_0_729,plain,(accept_city(X1,towna)|~accept_population(X1,other,n0)|~accept_population(X1,native,n0)|~accept_population(X1,muslim,n1)|~accept_population(X1,christian,n24)|~accept_population(X1,atheist,n75)),inference(split_conjunct, [status(thm)],[c_0_723])).
fof(c_0_730, plain, (![X1]:((((((~accept_population(X1,atheist,n75)|~accept_population(X1,christian,n25))|~accept_population(X1,muslim,n0))|~accept_population(X1,native,n0))|~accept_population(X1,other,n0))|accept_city(X1,citya))&(~accept_city(X1,citya)|((((accept_population(X1,atheist,n75)&accept_population(X1,christian,n25))&accept_population(X1,muslim,n0))&accept_population(X1,native,n0))&accept_population(X1,other,n0))))),inference(fof_nnf, [status(thm)],[c_0_17])).
fof(c_0_731, plain, ((![X1]:(((((~accept_population(X1,atheist,n75)|~accept_population(X1,christian,n25))|~accept_population(X1,muslim,n0))|~accept_population(X1,native,n0))|~accept_population(X1,other,n0))|accept_city(X1,citya))&![X1]:(~accept_city(X1,citya)|((((accept_population(X1,atheist,n75)&accept_population(X1,christian,n25))&accept_population(X1,muslim,n0))&accept_population(X1,native,n0))&accept_population(X1,other,n0))))),inference(shift_quantors, [status(thm)],[c_0_730])).
fof(c_0_732, plain, ((![X2]:(((((~accept_population(X2,atheist,n75)|~accept_population(X2,christian,n25))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n0))|~accept_population(X2,other,n0))|accept_city(X2,citya))&![X3]:(~accept_city(X3,citya)|((((accept_population(X3,atheist,n75)&accept_population(X3,christian,n25))&accept_population(X3,muslim,n0))&accept_population(X3,native,n0))&accept_population(X3,other,n0))))),inference(variable_rename, [status(thm)],[c_0_731])).
fof(c_0_733, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n75)|~accept_population(X2,christian,n25))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n0))|~accept_population(X2,other,n0))|accept_city(X2,citya))&(~accept_city(X3,citya)|((((accept_population(X3,atheist,n75)&accept_population(X3,christian,n25))&accept_population(X3,muslim,n0))&accept_population(X3,native,n0))&accept_population(X3,other,n0))))),inference(shift_quantors, [status(thm)],[c_0_732])).
fof(c_0_734, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n75)|~accept_population(X2,christian,n25))|~accept_population(X2,muslim,n0))|~accept_population(X2,native,n0))|~accept_population(X2,other,n0))|accept_city(X2,citya))&(((((accept_population(X3,atheist,n75)|~accept_city(X3,citya))&(accept_population(X3,christian,n25)|~accept_city(X3,citya)))&(accept_population(X3,muslim,n0)|~accept_city(X3,citya)))&(accept_population(X3,native,n0)|~accept_city(X3,citya)))&(accept_population(X3,other,n0)|~accept_city(X3,citya))))),inference(distribute, [status(thm)],[c_0_733])).
cnf(c_0_735,plain,(accept_population(X1,other,n0)|~accept_city(X1,citya)),inference(split_conjunct, [status(thm)],[c_0_734])).
cnf(c_0_736,plain,(accept_population(X1,native,n0)|~accept_city(X1,citya)),inference(split_conjunct, [status(thm)],[c_0_734])).
cnf(c_0_737,plain,(accept_population(X1,muslim,n0)|~accept_city(X1,citya)),inference(split_conjunct, [status(thm)],[c_0_734])).
cnf(c_0_738,plain,(accept_population(X1,christian,n25)|~accept_city(X1,citya)),inference(split_conjunct, [status(thm)],[c_0_734])).
cnf(c_0_739,plain,(accept_population(X1,atheist,n75)|~accept_city(X1,citya)),inference(split_conjunct, [status(thm)],[c_0_734])).
cnf(c_0_740,plain,(accept_city(X1,citya)|~accept_population(X1,other,n0)|~accept_population(X1,native,n0)|~accept_population(X1,muslim,n0)|~accept_population(X1,christian,n25)|~accept_population(X1,atheist,n75)),inference(split_conjunct, [status(thm)],[c_0_734])).
fof(c_0_741, plain, (![X1]:((((((~accept_population(X1,atheist,n70)|~accept_population(X1,christian,n20))|~accept_population(X1,muslim,n8))|~accept_population(X1,native,n0))|~accept_population(X1,other,n2))|accept_city(X1,townb))&(~accept_city(X1,townb)|((((accept_population(X1,atheist,n70)&accept_population(X1,christian,n20))&accept_population(X1,muslim,n8))&accept_population(X1,native,n0))&accept_population(X1,other,n2))))),inference(fof_nnf, [status(thm)],[c_0_18])).
fof(c_0_742, plain, ((![X1]:(((((~accept_population(X1,atheist,n70)|~accept_population(X1,christian,n20))|~accept_population(X1,muslim,n8))|~accept_population(X1,native,n0))|~accept_population(X1,other,n2))|accept_city(X1,townb))&![X1]:(~accept_city(X1,townb)|((((accept_population(X1,atheist,n70)&accept_population(X1,christian,n20))&accept_population(X1,muslim,n8))&accept_population(X1,native,n0))&accept_population(X1,other,n2))))),inference(shift_quantors, [status(thm)],[c_0_741])).
fof(c_0_743, plain, ((![X2]:(((((~accept_population(X2,atheist,n70)|~accept_population(X2,christian,n20))|~accept_population(X2,muslim,n8))|~accept_population(X2,native,n0))|~accept_population(X2,other,n2))|accept_city(X2,townb))&![X3]:(~accept_city(X3,townb)|((((accept_population(X3,atheist,n70)&accept_population(X3,christian,n20))&accept_population(X3,muslim,n8))&accept_population(X3,native,n0))&accept_population(X3,other,n2))))),inference(variable_rename, [status(thm)],[c_0_742])).
fof(c_0_744, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n70)|~accept_population(X2,christian,n20))|~accept_population(X2,muslim,n8))|~accept_population(X2,native,n0))|~accept_population(X2,other,n2))|accept_city(X2,townb))&(~accept_city(X3,townb)|((((accept_population(X3,atheist,n70)&accept_population(X3,christian,n20))&accept_population(X3,muslim,n8))&accept_population(X3,native,n0))&accept_population(X3,other,n2))))),inference(shift_quantors, [status(thm)],[c_0_743])).
fof(c_0_745, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n70)|~accept_population(X2,christian,n20))|~accept_population(X2,muslim,n8))|~accept_population(X2,native,n0))|~accept_population(X2,other,n2))|accept_city(X2,townb))&(((((accept_population(X3,atheist,n70)|~accept_city(X3,townb))&(accept_population(X3,christian,n20)|~accept_city(X3,townb)))&(accept_population(X3,muslim,n8)|~accept_city(X3,townb)))&(accept_population(X3,native,n0)|~accept_city(X3,townb)))&(accept_population(X3,other,n2)|~accept_city(X3,townb))))),inference(distribute, [status(thm)],[c_0_744])).
cnf(c_0_746,plain,(accept_population(X1,other,n2)|~accept_city(X1,townb)),inference(split_conjunct, [status(thm)],[c_0_745])).
cnf(c_0_747,plain,(accept_population(X1,native,n0)|~accept_city(X1,townb)),inference(split_conjunct, [status(thm)],[c_0_745])).
cnf(c_0_748,plain,(accept_population(X1,muslim,n8)|~accept_city(X1,townb)),inference(split_conjunct, [status(thm)],[c_0_745])).
cnf(c_0_749,plain,(accept_population(X1,christian,n20)|~accept_city(X1,townb)),inference(split_conjunct, [status(thm)],[c_0_745])).
cnf(c_0_750,plain,(accept_population(X1,atheist,n70)|~accept_city(X1,townb)),inference(split_conjunct, [status(thm)],[c_0_745])).
cnf(c_0_751,plain,(accept_city(X1,townb)|~accept_population(X1,other,n2)|~accept_population(X1,native,n0)|~accept_population(X1,muslim,n8)|~accept_population(X1,christian,n20)|~accept_population(X1,atheist,n70)),inference(split_conjunct, [status(thm)],[c_0_745])).
fof(c_0_752, plain, (![X1]:((((((~accept_population(X1,atheist,n78)|~accept_population(X1,christian,n20))|~accept_population(X1,muslim,n1))|~accept_population(X1,native,n0))|~accept_population(X1,other,n1))|accept_city(X1,cityb))&(~accept_city(X1,cityb)|((((accept_population(X1,atheist,n78)&accept_population(X1,christian,n20))&accept_population(X1,muslim,n1))&accept_population(X1,native,n0))&accept_population(X1,other,n1))))),inference(fof_nnf, [status(thm)],[c_0_19])).
fof(c_0_753, plain, ((![X1]:(((((~accept_population(X1,atheist,n78)|~accept_population(X1,christian,n20))|~accept_population(X1,muslim,n1))|~accept_population(X1,native,n0))|~accept_population(X1,other,n1))|accept_city(X1,cityb))&![X1]:(~accept_city(X1,cityb)|((((accept_population(X1,atheist,n78)&accept_population(X1,christian,n20))&accept_population(X1,muslim,n1))&accept_population(X1,native,n0))&accept_population(X1,other,n1))))),inference(shift_quantors, [status(thm)],[c_0_752])).
fof(c_0_754, plain, ((![X2]:(((((~accept_population(X2,atheist,n78)|~accept_population(X2,christian,n20))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n0))|~accept_population(X2,other,n1))|accept_city(X2,cityb))&![X3]:(~accept_city(X3,cityb)|((((accept_population(X3,atheist,n78)&accept_population(X3,christian,n20))&accept_population(X3,muslim,n1))&accept_population(X3,native,n0))&accept_population(X3,other,n1))))),inference(variable_rename, [status(thm)],[c_0_753])).
fof(c_0_755, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n78)|~accept_population(X2,christian,n20))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n0))|~accept_population(X2,other,n1))|accept_city(X2,cityb))&(~accept_city(X3,cityb)|((((accept_population(X3,atheist,n78)&accept_population(X3,christian,n20))&accept_population(X3,muslim,n1))&accept_population(X3,native,n0))&accept_population(X3,other,n1))))),inference(shift_quantors, [status(thm)],[c_0_754])).
fof(c_0_756, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n78)|~accept_population(X2,christian,n20))|~accept_population(X2,muslim,n1))|~accept_population(X2,native,n0))|~accept_population(X2,other,n1))|accept_city(X2,cityb))&(((((accept_population(X3,atheist,n78)|~accept_city(X3,cityb))&(accept_population(X3,christian,n20)|~accept_city(X3,cityb)))&(accept_population(X3,muslim,n1)|~accept_city(X3,cityb)))&(accept_population(X3,native,n0)|~accept_city(X3,cityb)))&(accept_population(X3,other,n1)|~accept_city(X3,cityb))))),inference(distribute, [status(thm)],[c_0_755])).
cnf(c_0_757,plain,(accept_population(X1,other,n1)|~accept_city(X1,cityb)),inference(split_conjunct, [status(thm)],[c_0_756])).
cnf(c_0_758,plain,(accept_population(X1,native,n0)|~accept_city(X1,cityb)),inference(split_conjunct, [status(thm)],[c_0_756])).
cnf(c_0_759,plain,(accept_population(X1,muslim,n1)|~accept_city(X1,cityb)),inference(split_conjunct, [status(thm)],[c_0_756])).
cnf(c_0_760,plain,(accept_population(X1,christian,n20)|~accept_city(X1,cityb)),inference(split_conjunct, [status(thm)],[c_0_756])).
cnf(c_0_761,plain,(accept_population(X1,atheist,n78)|~accept_city(X1,cityb)),inference(split_conjunct, [status(thm)],[c_0_756])).
cnf(c_0_762,plain,(accept_city(X1,cityb)|~accept_population(X1,other,n1)|~accept_population(X1,native,n0)|~accept_population(X1,muslim,n1)|~accept_population(X1,christian,n20)|~accept_population(X1,atheist,n78)),inference(split_conjunct, [status(thm)],[c_0_756])).
fof(c_0_763, plain, (![X1]:((((((~accept_population(X1,atheist,n30)|~accept_population(X1,christian,n0))|~accept_population(X1,muslim,n65))|~accept_population(X1,native,n0))|~accept_population(X1,other,n5))|accept_city(X1,townc))&(~accept_city(X1,townc)|((((accept_population(X1,atheist,n30)&accept_population(X1,christian,n0))&accept_population(X1,muslim,n65))&accept_population(X1,native,n0))&accept_population(X1,other,n5))))),inference(fof_nnf, [status(thm)],[c_0_20])).
fof(c_0_764, plain, ((![X1]:(((((~accept_population(X1,atheist,n30)|~accept_population(X1,christian,n0))|~accept_population(X1,muslim,n65))|~accept_population(X1,native,n0))|~accept_population(X1,other,n5))|accept_city(X1,townc))&![X1]:(~accept_city(X1,townc)|((((accept_population(X1,atheist,n30)&accept_population(X1,christian,n0))&accept_population(X1,muslim,n65))&accept_population(X1,native,n0))&accept_population(X1,other,n5))))),inference(shift_quantors, [status(thm)],[c_0_763])).
fof(c_0_765, plain, ((![X2]:(((((~accept_population(X2,atheist,n30)|~accept_population(X2,christian,n0))|~accept_population(X2,muslim,n65))|~accept_population(X2,native,n0))|~accept_population(X2,other,n5))|accept_city(X2,townc))&![X3]:(~accept_city(X3,townc)|((((accept_population(X3,atheist,n30)&accept_population(X3,christian,n0))&accept_population(X3,muslim,n65))&accept_population(X3,native,n0))&accept_population(X3,other,n5))))),inference(variable_rename, [status(thm)],[c_0_764])).
fof(c_0_766, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n30)|~accept_population(X2,christian,n0))|~accept_population(X2,muslim,n65))|~accept_population(X2,native,n0))|~accept_population(X2,other,n5))|accept_city(X2,townc))&(~accept_city(X3,townc)|((((accept_population(X3,atheist,n30)&accept_population(X3,christian,n0))&accept_population(X3,muslim,n65))&accept_population(X3,native,n0))&accept_population(X3,other,n5))))),inference(shift_quantors, [status(thm)],[c_0_765])).
fof(c_0_767, plain, (![X2]:![X3]:((((((~accept_population(X2,atheist,n30)|~accept_population(X2,christian,n0))|~accept_population(X2,muslim,n65))|~accept_population(X2,native,n0))|~accept_population(X2,other,n5))|accept_city(X2,townc))&(((((accept_population(X3,atheist,n30)|~accept_city(X3,townc))&(accept_population(X3,christian,n0)|~accept_city(X3,townc)))&(accept_population(X3,muslim,n65)|~accept_city(X3,townc)))&(accept_population(X3,native,n0)|~accept_city(X3,townc)))&(accept_population(X3,other,n5)|~accept_city(X3,townc))))),inference(distribute, [status(thm)],[c_0_766])).
cnf(c_0_768,plain,(accept_population(X1,other,n5)|~accept_city(X1,townc)),inference(split_conjunct, [status(thm)],[c_0_767])).
cnf(c_0_769,plain,(accept_population(X1,native,n0)|~accept_city(X1,townc)),inference(split_conjunct, [status(thm)],[c_0_767])).
cnf(c_0_770,plain,(accept_population(X1,muslim,n65)|~accept_city(X1,townc)),inference(split_conjunct, [status(thm)],[c_0_767])).
cnf(c_0_771,plain,(accept_population(X1,christian,n0)|~accept_city(X1,townc)),inference(split_conjunct, [status(thm)],[c_0_767])).
cnf(c_0_772,plain,(accept_population(X1,atheist,n30)|~accept_city(X1,townc)),inference(split_conjunct, [status(thm)],[c_0_767])).
cnf(c_0_773,plain,(accept_city(X1,townc)|~accept_population(X1,other,n5)|~accept_population(X1,native,n0)|~accept_population(X1,muslim,n65)|~accept_population(X1,christian,n0)|~accept_population(X1,atheist,n30)),inference(split_conjunct, [status(thm)],[c_0_767])).
cnf(c_0_774,plain,(accept_team(countryahumanitarianorganization,countryacivilorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_21])).
cnf(c_0_775,plain,(accept_team(countryahumanitarianorganization,countryacivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_22])).
cnf(c_0_776,plain,(accept_team(countryahumanitarianorganization,countryacivilorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_23])).
cnf(c_0_777,plain,(accept_team(countryahumanitarianorganization,countryafirstaidorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_24])).
cnf(c_0_778,plain,(the_agent_in_all_proposed_teams(countryahumanitarianorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_25])).
cnf(c_0_779,plain,(any_agent_in_all_proposed_teams(countryahumanitarianorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_26])).
cnf(c_0_780,plain,(accept_team(countryccivilorganization,countrybhumanitarianorganization,cityb,n2)),inference(split_conjunct, [status(thm)],[c_0_27])).
cnf(c_0_781,plain,(accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,cityb,n4)),inference(split_conjunct, [status(thm)],[c_0_28])).
cnf(c_0_782,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n4)),inference(split_conjunct, [status(thm)],[c_0_558])).
cnf(c_0_783,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n5)),inference(split_conjunct, [status(thm)],[c_0_559])).
cnf(c_0_784,plain,(the_agent_not_in_any_proposed_teams(muslimcountrybhumanitarianorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_31])).
cnf(c_0_785,plain,(any_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_32])).
cnf(c_0_786,plain,(accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,coastvillage,n5)),inference(split_conjunct, [status(thm)],[c_0_33])).
cnf(c_0_787,plain,(accept_team(countryamedicalorganization,countryacivilorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_34])).
cnf(c_0_788,plain,(accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,cityb,n4)),inference(split_conjunct, [status(thm)],[c_0_35])).
cnf(c_0_789,plain,(accept_team(countryafirstaidorganization,countryacivilorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_36])).
cnf(c_0_790,plain,(accept_number(countrybhumanitarianorganization,n4)),inference(split_conjunct, [status(thm)],[c_0_37])).
cnf(c_0_791,plain,(accept_team(countrybcivilorganization,countrybhumanitarianorganization,cityb,n4)),inference(split_conjunct, [status(thm)],[c_0_38])).
cnf(c_0_792,plain,(accept_team(christiansufferterrahumanitarianorganization,countryacivilorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_39])).
cnf(c_0_793,plain,(accept_team(sufferterragovernment,countryacivilorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_40])).
cnf(c_0_794,plain,(accept_number(countrybhumanitarianorganization,n1)),inference(split_conjunct, [status(thm)],[c_0_41])).
cnf(c_0_795,plain,(the_agent_in_all_proposed_teams(countrybhumanitarianorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_42])).
cnf(c_0_796,plain,(any_agent_in_all_proposed_teams(countrybhumanitarianorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_43])).
cnf(c_0_797,plain,(accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,coastvillage,n5)),inference(split_conjunct, [status(thm)],[c_0_44])).
cnf(c_0_798,plain,(accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_45])).
cnf(c_0_799,plain,(accept_team(christiansufferterrahumanitarianorganization,countryacivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_46])).
cnf(c_0_800,plain,(accept_team(countryamedicalorganization,countryacivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_47])).
cnf(c_0_801,plain,(accept_number(countryahumanitarianorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_48])).
cnf(c_0_802,plain,(~accept_team(countryamedicalorganization,countryahumanitarianorganization,coastvillage,n2)),inference(split_conjunct, [status(thm)],[c_0_560])).
cnf(c_0_803,plain,(~accept_team(countryamedicalorganization,countryafirstaidorganization,coastvillage,n2)),inference(split_conjunct, [status(thm)],[c_0_561])).
cnf(c_0_804,plain,(~accept_team(countryamedicalorganization,countryacivilorganization,coastvillage,n2)),inference(split_conjunct, [status(thm)],[c_0_562])).
cnf(c_0_805,plain,(~accept_team(countryamedicalorganization,christiansufferterrahumanitarianorganization,coastvillage,n2)),inference(split_conjunct, [status(thm)],[c_0_563])).
cnf(c_0_806,plain,(~accept_team(countryamedicalorganization,sufferterragovernment,coastvillage,n2)),inference(split_conjunct, [status(thm)],[c_0_564])).
cnf(c_0_807,plain,(~accept_team(countryamedicalorganization,countryahumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_565])).
cnf(c_0_808,plain,(~accept_team(countryamedicalorganization,countryacivilorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_566])).
cnf(c_0_809,plain,(~accept_team(countryamedicalorganization,countryafirstaidorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_567])).
cnf(c_0_810,plain,(the_agent_not_in_any_proposed_teams(countryamedicalorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_57])).
cnf(c_0_811,plain,(any_agent_in_all_proposed_teams(countryamedicalorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_58])).
cnf(c_0_812,plain,(accept_team(countryccivilorganization,countrycmedicalorganization,towna,n4)),inference(split_conjunct, [status(thm)],[c_0_59])).
cnf(c_0_813,plain,(~accept_team(countryccivilorganization,countrycmedicalorganization,towna,n5)),inference(split_conjunct, [status(thm)],[c_0_568])).
cnf(c_0_814,plain,(the_agent_in_all_proposed_teams(countryccivilorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_61])).
cnf(c_0_815,plain,(any_agent_in_all_proposed_teams(countryccivilorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_62])).
cnf(c_0_816,plain,(accept_team(countrybcivilorganization,countrycmedicalorganization,towna,n4)),inference(split_conjunct, [status(thm)],[c_0_63])).
cnf(c_0_817,plain,(accept_team(countrybcivilorganization,countrycmedicalorganization,towna,n5)),inference(split_conjunct, [status(thm)],[c_0_64])).
cnf(c_0_818,plain,(accept_number(countrybcivilorganization,n5)),inference(split_conjunct, [status(thm)],[c_0_65])).
cnf(c_0_819,plain,(the_agent_in_all_proposed_teams(countrybcivilorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_66])).
cnf(c_0_820,plain,(any_agent_in_all_proposed_teams(countrybcivilorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_67])).
cnf(c_0_821,plain,(accept_team(sufferterragovernment,countryacivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_68])).
cnf(c_0_822,plain,(the_agent_in_all_proposed_teams(countrycmedicalorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_69])).
cnf(c_0_823,plain,(any_agent_in_all_proposed_teams(countrycmedicalorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_70])).
cnf(c_0_824,plain,(accept_team(countrycmedicalorganization,christiancountrychumanitarianorganization,coastvillage,n5)),inference(split_conjunct, [status(thm)],[c_0_71])).
cnf(c_0_825,plain,(accept_number(countrycmedicalorganization,n4)),inference(split_conjunct, [status(thm)],[c_0_72])).
cnf(c_0_826,plain,(accept_team(christiancountrychumanitarianorganization,countrycmedicalorganization,towna,n4)),inference(split_conjunct, [status(thm)],[c_0_73])).
cnf(c_0_827,plain,(accept_number(countrycmedicalorganization,n5)),inference(split_conjunct, [status(thm)],[c_0_74])).
cnf(c_0_828,plain,(accept_team(christiancountrychumanitarianorganization,countrycmedicalorganization,towna,n5)),inference(split_conjunct, [status(thm)],[c_0_75])).
cnf(c_0_829,plain,(the_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_76])).
cnf(c_0_830,plain,(any_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,countryacivilorganization,towna)),inference(split_conjunct, [status(thm)],[c_0_77])).
cnf(c_0_831,plain,(accept_team(countryafirstaidorganization,countryacivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_78])).
cnf(c_0_832,plain,(accept_number(countryacivilorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_79])).
cnf(c_0_833,plain,(accept_team(countryacivilorganization,countryahumanitarianorganization,coastvillage,n5)),inference(split_conjunct, [status(thm)],[c_0_80])).
cnf(c_0_834,plain,(accept_team(countryacivilorganization,countryahumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_81])).
cnf(c_0_835,plain,(accept_team(countryacivilorganization,countryafirstaidorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_82])).
cnf(c_0_836,plain,(the_agent_in_all_proposed_teams(countryacivilorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_83])).
cnf(c_0_837,plain,(any_agent_in_all_proposed_teams(countryacivilorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_84])).
cnf(c_0_838,plain,(accept_number(countryafirstaidorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_85])).
cnf(c_0_839,plain,(accept_team(countryafirstaidorganization,countryahumanitarianorganization,coastvillage,n5)),inference(split_conjunct, [status(thm)],[c_0_86])).
cnf(c_0_840,plain,(accept_team(countryafirstaidorganization,countryahumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_87])).
cnf(c_0_841,plain,(accept_team(countryafirstaidorganization,countryacivilorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_88])).
cnf(c_0_842,plain,(the_agent_in_all_proposed_teams(countryafirstaidorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_89])).
cnf(c_0_843,plain,(any_agent_in_all_proposed_teams(countryafirstaidorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_90])).
cnf(c_0_844,plain,(~accept_team(countryccivilorganization,christiancountrychumanitarianorganization,coastvillage,n5)),inference(split_conjunct, [status(thm)],[c_0_569])).
cnf(c_0_845,plain,(~accept_team(countryccivilorganization,christiancountrychumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_570])).
cnf(c_0_846,plain,(accept_number(christiancountrychumanitarianorganization,n5)),inference(split_conjunct, [status(thm)],[c_0_93])).
cnf(c_0_847,plain,(accept_team(countrybcivilorganization,christiancountrychumanitarianorganization,coastvillage,n5)),inference(split_conjunct, [status(thm)],[c_0_94])).
cnf(c_0_848,plain,(accept_team(countrybcivilorganization,christiancountrychumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_95])).
cnf(c_0_849,plain,(accept_team(countrybcivilorganization,muslimcountrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_96])).
cnf(c_0_850,plain,(accept_team(countrybcivilorganization,countrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_97])).
cnf(c_0_851,plain,(accept_team(countrybcivilorganization,christiancountrychumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_98])).
cnf(c_0_852,plain,(accept_team(countrybcivilorganization,countrycmedicalorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_99])).
cnf(c_0_853,plain,(the_agent_in_all_proposed_teams(countrybcivilorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_100])).
cnf(c_0_854,plain,(any_agent_in_all_proposed_teams(countrybcivilorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_101])).
cnf(c_0_855,plain,(accept_team(sufferterragovernment,countryahumanitarianorganization,coastvillage,n5)),inference(split_conjunct, [status(thm)],[c_0_102])).
cnf(c_0_856,plain,(accept_team(sufferterragovernment,countryahumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_103])).
cnf(c_0_857,plain,(accept_team(sufferterragovernment,countryacivilorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_104])).
cnf(c_0_858,plain,(accept_team(sufferterragovernment,countryafirstaidorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_105])).
cnf(c_0_859,plain,(the_agent_in_all_proposed_teams(sufferterragovernment,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_106])).
cnf(c_0_860,plain,(any_agent_in_all_proposed_teams(sufferterragovernment,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_107])).
cnf(c_0_861,plain,(accept_number(christiansufferterrahumanitarianorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_108])).
cnf(c_0_862,plain,(accept_number(countryahumanitarianorganization,n5)),inference(split_conjunct, [status(thm)],[c_0_109])).
cnf(c_0_863,plain,(accept_team(christiansufferterrahumanitarianorganization,countryahumanitarianorganization,coastvillage,n5)),inference(split_conjunct, [status(thm)],[c_0_110])).
cnf(c_0_864,plain,(accept_team(christiansufferterrahumanitarianorganization,countryahumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_111])).
cnf(c_0_865,plain,(accept_team(christiansufferterrahumanitarianorganization,countryacivilorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_112])).
cnf(c_0_866,plain,(accept_team(christiansufferterrahumanitarianorganization,countryafirstaidorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_113])).
cnf(c_0_867,plain,(the_agent_in_all_proposed_teams(christiansufferterrahumanitarianorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_114])).
cnf(c_0_868,plain,(any_agent_in_all_proposed_teams(christiansufferterrahumanitarianorganization,christiancountrychumanitarianorganization,coastvillage)),inference(split_conjunct, [status(thm)],[c_0_115])).
cnf(c_0_869,plain,(accept_team(countrycmedicalorganization,christiancountrychumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_116])).
cnf(c_0_870,plain,(accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,coastvillage,n6)),inference(split_conjunct, [status(thm)],[c_0_117])).
cnf(c_0_871,plain,(accept_number(muslimcountrybhumanitarianorganization,n3)),inference(split_conjunct, [status(thm)],[c_0_118])).
cnf(c_0_872,plain,(accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,townc,n4)),inference(split_conjunct, [status(thm)],[c_0_119])).
cnf(c_0_873,plain,(accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_120])).
cnf(c_0_874,plain,(accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_121])).
cnf(c_0_875,plain,(accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_122])).
cnf(c_0_876,plain,(accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_123])).
cnf(c_0_877,plain,(the_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_124])).
cnf(c_0_878,plain,(any_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_125])).
cnf(c_0_879,plain,(accept_team(countryacivilorganization,countryahumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_126])).
cnf(c_0_880,plain,(accept_team(countryafirstaidorganization,countryahumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_127])).
cnf(c_0_881,plain,(accept_team(christiancountrychumanitarianorganization,countryccivilorganization,townc,n4)),inference(split_conjunct, [status(thm)],[c_0_128])).
cnf(c_0_882,plain,(accept_team(christiancountrychumanitarianorganization,muslimcountrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_129])).
cnf(c_0_883,plain,(accept_team(christiancountrychumanitarianorganization,countrybcivilorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_130])).
cnf(c_0_884,plain,(accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_131])).
cnf(c_0_885,plain,(accept_team(christiancountrychumanitarianorganization,countrycmedicalorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_132])).
cnf(c_0_886,plain,(the_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_133])).
cnf(c_0_887,plain,(any_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_134])).
cnf(c_0_888,plain,(~accept_team(countryccivilorganization,muslimcountrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_571])).
cnf(c_0_889,plain,(~accept_team(countryccivilorganization,countrybcivilorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_572])).
cnf(c_0_890,plain,(~accept_team(countryccivilorganization,countrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_573])).
cnf(c_0_891,plain,(~accept_team(countryccivilorganization,christiancountrychumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_574])).
cnf(c_0_892,plain,(~accept_team(countryccivilorganization,countrycmedicalorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_575])).
cnf(c_0_893,plain,(the_agent_in_all_proposed_teams(countryccivilorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_140])).
cnf(c_0_894,plain,(any_agent_in_all_proposed_teams(countryccivilorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_141])).
cnf(c_0_895,plain,(accept_team(countrybhumanitarianorganization,countrybcivilorganization,townc,n2)),inference(split_conjunct, [status(thm)],[c_0_142])).
cnf(c_0_896,plain,(accept_team(countrybhumanitarianorganization,muslimcountrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_143])).
cnf(c_0_897,plain,(accept_team(sufferterragovernment,countryahumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_144])).
cnf(c_0_898,plain,(accept_team(christiansufferterrahumanitarianorganization,countryahumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_145])).
cnf(c_0_899,plain,(accept_team(countryamedicalorganization,countryahumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_146])).
cnf(c_0_900,plain,(accept_team(countrycmedicalorganization,countryccivilorganization,townc,n4)),inference(split_conjunct, [status(thm)],[c_0_147])).
cnf(c_0_901,plain,(accept_team(countrycmedicalorganization,muslimcountrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_148])).
cnf(c_0_902,plain,(accept_team(countrycmedicalorganization,countrybcivilorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_149])).
cnf(c_0_903,plain,(accept_team(countrycmedicalorganization,countrybhumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_150])).
cnf(c_0_904,plain,(accept_team(countrycmedicalorganization,christiancountrychumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_151])).
cnf(c_0_905,plain,(the_agent_in_all_proposed_teams(countrycmedicalorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_152])).
cnf(c_0_906,plain,(any_agent_in_all_proposed_teams(countrycmedicalorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_153])).
cnf(c_0_907,plain,(accept_team(countrybhumanitarianorganization,countrybcivilorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_154])).
cnf(c_0_908,plain,(accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_155])).
cnf(c_0_909,plain,(accept_team(countrybhumanitarianorganization,countrycmedicalorganization,townc,n6)),inference(split_conjunct, [status(thm)],[c_0_156])).
cnf(c_0_910,plain,(the_agent_in_all_proposed_teams(countrybhumanitarianorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_157])).
cnf(c_0_911,plain,(any_agent_in_all_proposed_teams(countrybhumanitarianorganization,countryahumanitarianorganization,townc)),inference(split_conjunct, [status(thm)],[c_0_158])).
cnf(c_0_912,plain,(accept_team(christiansufferterrahumanitarianorganization,countryahumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_159])).
cnf(c_0_913,plain,(accept_number(christiansufferterrahumanitarianorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_160])).
cnf(c_0_914,plain,(the_agent_in_all_proposed_teams(christiansufferterrahumanitarianorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_161])).
cnf(c_0_915,plain,(any_agent_in_all_proposed_teams(christiansufferterrahumanitarianorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_162])).
cnf(c_0_916,plain,(~accept_team(countryccivilorganization,countrybhumanitarianorganization,cityb,n5)),inference(split_conjunct, [status(thm)],[c_0_576])).
cnf(c_0_917,plain,(~accept_team(countryccivilorganization,countrybhumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_577])).
cnf(c_0_918,plain,(accept_team(sufferterragovernment,countryahumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_165])).
cnf(c_0_919,plain,(the_agent_in_all_proposed_teams(sufferterragovernment,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_166])).
cnf(c_0_920,plain,(any_agent_in_all_proposed_teams(sufferterragovernment,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_167])).
cnf(c_0_921,plain,(accept_team(countryafirstaidorganization,countryahumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_168])).
cnf(c_0_922,plain,(accept_number(countryafirstaidorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_169])).
cnf(c_0_923,plain,(the_agent_in_all_proposed_teams(countryafirstaidorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_170])).
cnf(c_0_924,plain,(any_agent_in_all_proposed_teams(countryafirstaidorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_171])).
cnf(c_0_925,plain,(accept_team(countryafirstaidorganization,sufferterragovernment,towna,n5)),inference(split_conjunct, [status(thm)],[c_0_172])).
cnf(c_0_926,plain,(accept_team(countryafirstaidorganization,sufferterragovernment,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_173])).
cnf(c_0_927,plain,(accept_team(countrycmedicalorganization,countrybhumanitarianorganization,cityb,n5)),inference(split_conjunct, [status(thm)],[c_0_174])).
cnf(c_0_928,plain,(accept_team(countrycmedicalorganization,countrybhumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_175])).
cnf(c_0_929,plain,(accept_team(countryamedicalorganization,countryahumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_176])).
cnf(c_0_930,plain,(accept_number(countryamedicalorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_177])).
cnf(c_0_931,plain,(the_agent_in_all_proposed_teams(countryamedicalorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_178])).
cnf(c_0_932,plain,(any_agent_in_all_proposed_teams(countryamedicalorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_179])).
cnf(c_0_933,plain,(accept_team(countrybcivilorganization,countrybhumanitarianorganization,cityb,n5)),inference(split_conjunct, [status(thm)],[c_0_180])).
cnf(c_0_934,plain,(accept_team(countrybcivilorganization,countrybhumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_181])).
cnf(c_0_935,plain,(accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,cityb,n5)),inference(split_conjunct, [status(thm)],[c_0_182])).
cnf(c_0_936,plain,(accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_183])).
cnf(c_0_937,plain,(the_agent_in_all_proposed_teams(countryahumanitarianorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_184])).
cnf(c_0_938,plain,(any_agent_in_all_proposed_teams(countryahumanitarianorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_185])).
cnf(c_0_939,plain,(accept_number(countryahumanitarianorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_186])).
cnf(c_0_940,plain,(accept_team(countryacivilorganization,countryahumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_187])).
cnf(c_0_941,plain,(accept_number(countryacivilorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_188])).
cnf(c_0_942,plain,(the_agent_in_all_proposed_teams(countryacivilorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_189])).
cnf(c_0_943,plain,(any_agent_in_all_proposed_teams(countryacivilorganization,countrybhumanitarianorganization,cityb)),inference(split_conjunct, [status(thm)],[c_0_190])).
cnf(c_0_944,plain,(accept_number(countrybhumanitarianorganization,n5)),inference(split_conjunct, [status(thm)],[c_0_191])).
cnf(c_0_945,plain,(accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,cityb,n5)),inference(split_conjunct, [status(thm)],[c_0_192])).
cnf(c_0_946,plain,(accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,cityb,n6)),inference(split_conjunct, [status(thm)],[c_0_193])).
cnf(c_0_947,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,towna,n2)),inference(split_conjunct, [status(thm)],[c_0_578])).
cnf(c_0_948,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n2)),inference(split_conjunct, [status(thm)],[c_0_579])).
cnf(c_0_949,plain,(~accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n2)),inference(split_conjunct, [status(thm)],[c_0_580])).
cnf(c_0_950,plain,(~accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,towna,n2)),inference(split_conjunct, [status(thm)],[c_0_581])).
cnf(c_0_951,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,towna,n3)),inference(split_conjunct, [status(thm)],[c_0_582])).
cnf(c_0_952,plain,(~accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n3)),inference(split_conjunct, [status(thm)],[c_0_583])).
cnf(c_0_953,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n3)),inference(split_conjunct, [status(thm)],[c_0_584])).
cnf(c_0_954,plain,(~accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,towna,n3)),inference(split_conjunct, [status(thm)],[c_0_585])).
cnf(c_0_955,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,towna,n2)),inference(split_conjunct, [status(thm)],[c_0_586])).
cnf(c_0_956,plain,(accept_number(muslimcountrybhumanitarianorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_203])).
cnf(c_0_957,plain,(~accept_team(muslimcountrybhumanitarianorganization,countryccivilorganization,towna,n4)),inference(split_conjunct, [status(thm)],[c_0_587])).
cnf(c_0_958,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrybcivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_588])).
cnf(c_0_959,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_589])).
cnf(c_0_960,plain,(~accept_team(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_590])).
cnf(c_0_961,plain,(accept_team(countryamedicalorganization,sufferterragovernment,towna,n5)),inference(split_conjunct, [status(thm)],[c_0_208])).
cnf(c_0_962,plain,(accept_team(countryamedicalorganization,sufferterragovernment,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_209])).
cnf(c_0_963,plain,(accept_team(countryahumanitarianorganization,sufferterragovernment,towna,n5)),inference(split_conjunct, [status(thm)],[c_0_210])).
cnf(c_0_964,plain,(accept_team(countryahumanitarianorganization,sufferterragovernment,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_211])).
cnf(c_0_965,plain,(accept_number(christiancountrychumanitarianorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_212])).
cnf(c_0_966,plain,(accept_team(countrybhumanitarianorganization,countrybcivilorganization,towna,n3)),inference(split_conjunct, [status(thm)],[c_0_213])).
cnf(c_0_967,plain,(accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n3)),inference(split_conjunct, [status(thm)],[c_0_214])).
cnf(c_0_968,plain,(accept_team(countrybhumanitarianorganization,countrycmedicalorganization,towna,n3)),inference(split_conjunct, [status(thm)],[c_0_215])).
cnf(c_0_969,plain,(accept_team(countrybhumanitarianorganization,countryccivilorganization,towna,n3)),inference(split_conjunct, [status(thm)],[c_0_216])).
cnf(c_0_970,plain,(accept_number(countrybhumanitarianorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_217])).
cnf(c_0_971,plain,(accept_team(countrybhumanitarianorganization,countrybcivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_218])).
cnf(c_0_972,plain,(accept_team(countrybhumanitarianorganization,christiancountrychumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_219])).
cnf(c_0_973,plain,(accept_team(countrybhumanitarianorganization,countrycmedicalorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_220])).
cnf(c_0_974,plain,(the_agent_in_all_proposed_teams(countrybhumanitarianorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_221])).
cnf(c_0_975,plain,(any_agent_in_all_proposed_teams(countrybhumanitarianorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_222])).
cnf(c_0_976,plain,(accept_number(countrybcivilorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_223])).
cnf(c_0_977,plain,(accept_number(countrybcivilorganization,n3)),inference(split_conjunct, [status(thm)],[c_0_224])).
cnf(c_0_978,plain,(accept_team(countrybcivilorganization,countryccivilorganization,towna,n4)),inference(split_conjunct, [status(thm)],[c_0_225])).
cnf(c_0_979,plain,(accept_team(countrybcivilorganization,countrybhumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_226])).
cnf(c_0_980,plain,(accept_team(countrybcivilorganization,christiancountrychumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_227])).
cnf(c_0_981,plain,(accept_team(countrybcivilorganization,countrycmedicalorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_228])).
cnf(c_0_982,plain,(the_agent_in_all_proposed_teams(countrybcivilorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_229])).
cnf(c_0_983,plain,(any_agent_in_all_proposed_teams(countrybcivilorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_230])).
cnf(c_0_984,plain,(accept_number(sufferterragovernment,n2)),inference(split_conjunct, [status(thm)],[c_0_231])).
cnf(c_0_985,plain,(~accept_team(christiansufferterrahumanitarianorganization,sufferterragovernment,towna,n2)),inference(split_conjunct, [status(thm)],[c_0_591])).
cnf(c_0_986,plain,(~accept_team(christiansufferterrahumanitarianorganization,sufferterragovernment,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_592])).
cnf(c_0_987,plain,(accept_number(countryccivilorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_234])).
cnf(c_0_988,plain,(accept_number(countryccivilorganization,n3)),inference(split_conjunct, [status(thm)],[c_0_235])).
cnf(c_0_989,plain,(~accept_team(countryccivilorganization,countrybcivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_593])).
cnf(c_0_990,plain,(~accept_team(countryccivilorganization,countrybhumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_594])).
cnf(c_0_991,plain,(~accept_team(countryccivilorganization,christiancountrychumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_595])).
cnf(c_0_992,plain,(~accept_team(countryccivilorganization,countrycmedicalorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_596])).
cnf(c_0_993,plain,(the_agent_in_all_proposed_teams(countryccivilorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_240])).
cnf(c_0_994,plain,(any_agent_in_all_proposed_teams(countryccivilorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_241])).
cnf(c_0_995,plain,(accept_number(sufferterragovernment,n5)),inference(split_conjunct, [status(thm)],[c_0_242])).
cnf(c_0_996,plain,(accept_team(countryacivilorganization,sufferterragovernment,towna,n5)),inference(split_conjunct, [status(thm)],[c_0_243])).
cnf(c_0_997,plain,(accept_number(sufferterragovernment,n6)),inference(split_conjunct, [status(thm)],[c_0_244])).
cnf(c_0_998,plain,(accept_team(countryacivilorganization,sufferterragovernment,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_245])).
cnf(c_0_999,plain,(accept_number(countrycmedicalorganization,n2)),inference(split_conjunct, [status(thm)],[c_0_246])).
cnf(c_0_1000,plain,(accept_number(countrycmedicalorganization,n3)),inference(split_conjunct, [status(thm)],[c_0_247])).
cnf(c_0_1001,plain,(accept_team(countrycmedicalorganization,countryccivilorganization,towna,n4)),inference(split_conjunct, [status(thm)],[c_0_248])).
cnf(c_0_1002,plain,(accept_team(countrycmedicalorganization,countrybcivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_249])).
cnf(c_0_1003,plain,(accept_team(countrycmedicalorganization,countrybhumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_250])).
cnf(c_0_1004,plain,(accept_team(countrycmedicalorganization,christiancountrychumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_251])).
cnf(c_0_1005,plain,(the_agent_in_all_proposed_teams(countrycmedicalorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_252])).
cnf(c_0_1006,plain,(any_agent_in_all_proposed_teams(countrycmedicalorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_253])).
cnf(c_0_1007,plain,(accept_number(christiancountrychumanitarianorganization,n3)),inference(split_conjunct, [status(thm)],[c_0_254])).
cnf(c_0_1008,plain,(accept_number(countryccivilorganization,n4)),inference(split_conjunct, [status(thm)],[c_0_255])).
cnf(c_0_1009,plain,(accept_team(christiancountrychumanitarianorganization,countryccivilorganization,towna,n4)),inference(split_conjunct, [status(thm)],[c_0_256])).
cnf(c_0_1010,plain,(accept_number(countrybcivilorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_257])).
cnf(c_0_1011,plain,(accept_team(christiancountrychumanitarianorganization,countrybcivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_258])).
cnf(c_0_1012,plain,(accept_number(countrybhumanitarianorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_259])).
cnf(c_0_1013,plain,(accept_team(christiancountrychumanitarianorganization,countrybhumanitarianorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_260])).
cnf(c_0_1014,plain,(accept_team(christiancountrychumanitarianorganization,countrycmedicalorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_261])).
cnf(c_0_1015,plain,(accept_number(christiancountrychumanitarianorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_262])).
cnf(c_0_1016,plain,(the_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_263])).
cnf(c_0_1017,plain,(any_agent_in_all_proposed_teams(christiancountrychumanitarianorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_264])).
cnf(c_0_1018,plain,(accept_number(countrycmedicalorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_265])).
cnf(c_0_1019,plain,(~accept_team(muslimcountrybhumanitarianorganization,countrycmedicalorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_597])).
cnf(c_0_1020,plain,(accept_number(muslimcountrybhumanitarianorganization,n6)),inference(split_conjunct, [status(thm)],[c_0_267])).
cnf(c_0_1021,plain,(the_agent_not_in_any_proposed_teams(muslimcountrybhumanitarianorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_268])).
cnf(c_0_1022,plain,(any_agent_in_all_proposed_teams(muslimcountrybhumanitarianorganization,sufferterragovernment,towna)),inference(split_conjunct, [status(thm)],[c_0_269])).
cnf(c_0_1023,plain,(rdn_translate(n0,rdn_pos(rdnn(n0)))),inference(split_conjunct, [status(thm)],[c_0_270])).
cnf(c_0_1024,plain,(rdn_translate(n1,rdn_pos(rdnn(n1)))),inference(split_conjunct, [status(thm)],[c_0_271])).
cnf(c_0_1025,plain,(rdn_translate(n2,rdn_pos(rdnn(n2)))),inference(split_conjunct, [status(thm)],[c_0_272])).
cnf(c_0_1026,plain,(rdn_translate(n3,rdn_pos(rdnn(n3)))),inference(split_conjunct, [status(thm)],[c_0_273])).
cnf(c_0_1027,plain,(rdn_translate(n4,rdn_pos(rdnn(n4)))),inference(split_conjunct, [status(thm)],[c_0_274])).
cnf(c_0_1028,plain,(rdn_translate(n5,rdn_pos(rdnn(n5)))),inference(split_conjunct, [status(thm)],[c_0_275])).
cnf(c_0_1029,plain,(rdn_translate(n6,rdn_pos(rdnn(n6)))),inference(split_conjunct, [status(thm)],[c_0_276])).
cnf(c_0_1030,plain,(rdn_translate(n7,rdn_pos(rdnn(n7)))),inference(split_conjunct, [status(thm)],[c_0_277])).
cnf(c_0_1031,plain,(rdn_translate(n8,rdn_pos(rdnn(n8)))),inference(split_conjunct, [status(thm)],[c_0_278])).
cnf(c_0_1032,plain,(rdn_translate(n9,rdn_pos(rdnn(n9)))),inference(split_conjunct, [status(thm)],[c_0_279])).
cnf(c_0_1033,plain,(rdn_translate(n10,rdn_pos(rdn(rdnn(n0),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_280])).
cnf(c_0_1034,plain,(rdn_translate(n11,rdn_pos(rdn(rdnn(n1),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_281])).
cnf(c_0_1035,plain,(rdn_translate(n12,rdn_pos(rdn(rdnn(n2),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_282])).
cnf(c_0_1036,plain,(rdn_translate(n13,rdn_pos(rdn(rdnn(n3),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_283])).
cnf(c_0_1037,plain,(rdn_translate(n14,rdn_pos(rdn(rdnn(n4),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_284])).
cnf(c_0_1038,plain,(rdn_translate(n15,rdn_pos(rdn(rdnn(n5),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_285])).
cnf(c_0_1039,plain,(rdn_translate(n16,rdn_pos(rdn(rdnn(n6),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_286])).
cnf(c_0_1040,plain,(rdn_translate(n17,rdn_pos(rdn(rdnn(n7),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_287])).
cnf(c_0_1041,plain,(rdn_translate(n18,rdn_pos(rdn(rdnn(n8),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_288])).
cnf(c_0_1042,plain,(rdn_translate(n19,rdn_pos(rdn(rdnn(n9),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_289])).
cnf(c_0_1043,plain,(rdn_translate(n20,rdn_pos(rdn(rdnn(n0),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_290])).
cnf(c_0_1044,plain,(rdn_translate(n21,rdn_pos(rdn(rdnn(n1),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_291])).
cnf(c_0_1045,plain,(rdn_translate(n22,rdn_pos(rdn(rdnn(n2),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_292])).
cnf(c_0_1046,plain,(rdn_translate(n23,rdn_pos(rdn(rdnn(n3),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_293])).
cnf(c_0_1047,plain,(rdn_translate(n24,rdn_pos(rdn(rdnn(n4),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_294])).
cnf(c_0_1048,plain,(rdn_translate(n25,rdn_pos(rdn(rdnn(n5),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_295])).
cnf(c_0_1049,plain,(rdn_translate(n26,rdn_pos(rdn(rdnn(n6),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_296])).
cnf(c_0_1050,plain,(rdn_translate(n27,rdn_pos(rdn(rdnn(n7),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_297])).
cnf(c_0_1051,plain,(rdn_translate(n28,rdn_pos(rdn(rdnn(n8),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_298])).
cnf(c_0_1052,plain,(rdn_translate(n29,rdn_pos(rdn(rdnn(n9),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_299])).
cnf(c_0_1053,plain,(rdn_translate(n30,rdn_pos(rdn(rdnn(n0),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_300])).
cnf(c_0_1054,plain,(rdn_translate(n31,rdn_pos(rdn(rdnn(n1),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_301])).
cnf(c_0_1055,plain,(rdn_translate(n32,rdn_pos(rdn(rdnn(n2),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_302])).
cnf(c_0_1056,plain,(rdn_translate(n33,rdn_pos(rdn(rdnn(n3),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_303])).
cnf(c_0_1057,plain,(rdn_translate(n34,rdn_pos(rdn(rdnn(n4),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_304])).
cnf(c_0_1058,plain,(rdn_translate(n35,rdn_pos(rdn(rdnn(n5),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_305])).
cnf(c_0_1059,plain,(rdn_translate(n36,rdn_pos(rdn(rdnn(n6),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_306])).
cnf(c_0_1060,plain,(rdn_translate(n37,rdn_pos(rdn(rdnn(n7),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_307])).
cnf(c_0_1061,plain,(rdn_translate(n38,rdn_pos(rdn(rdnn(n8),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_308])).
cnf(c_0_1062,plain,(rdn_translate(n39,rdn_pos(rdn(rdnn(n9),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_309])).
cnf(c_0_1063,plain,(rdn_translate(n40,rdn_pos(rdn(rdnn(n0),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_310])).
cnf(c_0_1064,plain,(rdn_translate(n41,rdn_pos(rdn(rdnn(n1),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_311])).
cnf(c_0_1065,plain,(rdn_translate(n42,rdn_pos(rdn(rdnn(n2),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_312])).
cnf(c_0_1066,plain,(rdn_translate(n43,rdn_pos(rdn(rdnn(n3),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_313])).
cnf(c_0_1067,plain,(rdn_translate(n44,rdn_pos(rdn(rdnn(n4),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_314])).
cnf(c_0_1068,plain,(rdn_translate(n45,rdn_pos(rdn(rdnn(n5),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_315])).
cnf(c_0_1069,plain,(rdn_translate(n46,rdn_pos(rdn(rdnn(n6),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_316])).
cnf(c_0_1070,plain,(rdn_translate(n47,rdn_pos(rdn(rdnn(n7),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_317])).
cnf(c_0_1071,plain,(rdn_translate(n48,rdn_pos(rdn(rdnn(n8),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_318])).
cnf(c_0_1072,plain,(rdn_translate(n49,rdn_pos(rdn(rdnn(n9),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_319])).
cnf(c_0_1073,plain,(rdn_translate(n50,rdn_pos(rdn(rdnn(n0),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_320])).
cnf(c_0_1074,plain,(rdn_translate(n51,rdn_pos(rdn(rdnn(n1),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_321])).
cnf(c_0_1075,plain,(rdn_translate(n52,rdn_pos(rdn(rdnn(n2),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_322])).
cnf(c_0_1076,plain,(rdn_translate(n53,rdn_pos(rdn(rdnn(n3),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_323])).
cnf(c_0_1077,plain,(rdn_translate(n54,rdn_pos(rdn(rdnn(n4),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_324])).
cnf(c_0_1078,plain,(rdn_translate(n55,rdn_pos(rdn(rdnn(n5),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_325])).
cnf(c_0_1079,plain,(rdn_translate(n56,rdn_pos(rdn(rdnn(n6),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_326])).
cnf(c_0_1080,plain,(rdn_translate(n57,rdn_pos(rdn(rdnn(n7),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_327])).
cnf(c_0_1081,plain,(rdn_translate(n58,rdn_pos(rdn(rdnn(n8),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_328])).
cnf(c_0_1082,plain,(rdn_translate(n59,rdn_pos(rdn(rdnn(n9),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_329])).
cnf(c_0_1083,plain,(rdn_translate(n60,rdn_pos(rdn(rdnn(n0),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_330])).
cnf(c_0_1084,plain,(rdn_translate(n61,rdn_pos(rdn(rdnn(n1),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_331])).
cnf(c_0_1085,plain,(rdn_translate(n62,rdn_pos(rdn(rdnn(n2),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_332])).
cnf(c_0_1086,plain,(rdn_translate(n63,rdn_pos(rdn(rdnn(n3),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_333])).
cnf(c_0_1087,plain,(rdn_translate(n64,rdn_pos(rdn(rdnn(n4),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_334])).
cnf(c_0_1088,plain,(rdn_translate(n65,rdn_pos(rdn(rdnn(n5),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_335])).
cnf(c_0_1089,plain,(rdn_translate(n66,rdn_pos(rdn(rdnn(n6),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_336])).
cnf(c_0_1090,plain,(rdn_translate(n67,rdn_pos(rdn(rdnn(n7),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_337])).
cnf(c_0_1091,plain,(rdn_translate(n68,rdn_pos(rdn(rdnn(n8),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_338])).
cnf(c_0_1092,plain,(rdn_translate(n69,rdn_pos(rdn(rdnn(n9),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_339])).
cnf(c_0_1093,plain,(rdn_translate(n70,rdn_pos(rdn(rdnn(n0),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_340])).
cnf(c_0_1094,plain,(rdn_translate(n71,rdn_pos(rdn(rdnn(n1),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_341])).
cnf(c_0_1095,plain,(rdn_translate(n72,rdn_pos(rdn(rdnn(n2),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_342])).
cnf(c_0_1096,plain,(rdn_translate(n73,rdn_pos(rdn(rdnn(n3),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_343])).
cnf(c_0_1097,plain,(rdn_translate(n74,rdn_pos(rdn(rdnn(n4),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_344])).
cnf(c_0_1098,plain,(rdn_translate(n75,rdn_pos(rdn(rdnn(n5),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_345])).
cnf(c_0_1099,plain,(rdn_translate(n76,rdn_pos(rdn(rdnn(n6),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_346])).
cnf(c_0_1100,plain,(rdn_translate(n77,rdn_pos(rdn(rdnn(n7),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_347])).
cnf(c_0_1101,plain,(rdn_translate(n78,rdn_pos(rdn(rdnn(n8),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_348])).
cnf(c_0_1102,plain,(rdn_translate(n79,rdn_pos(rdn(rdnn(n9),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_349])).
cnf(c_0_1103,plain,(rdn_translate(n80,rdn_pos(rdn(rdnn(n0),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_350])).
cnf(c_0_1104,plain,(rdn_translate(n81,rdn_pos(rdn(rdnn(n1),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_351])).
cnf(c_0_1105,plain,(rdn_translate(n82,rdn_pos(rdn(rdnn(n2),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_352])).
cnf(c_0_1106,plain,(rdn_translate(n83,rdn_pos(rdn(rdnn(n3),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_353])).
cnf(c_0_1107,plain,(rdn_translate(n84,rdn_pos(rdn(rdnn(n4),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_354])).
cnf(c_0_1108,plain,(rdn_translate(n85,rdn_pos(rdn(rdnn(n5),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_355])).
cnf(c_0_1109,plain,(rdn_translate(n86,rdn_pos(rdn(rdnn(n6),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_356])).
cnf(c_0_1110,plain,(rdn_translate(n87,rdn_pos(rdn(rdnn(n7),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_357])).
cnf(c_0_1111,plain,(rdn_translate(n88,rdn_pos(rdn(rdnn(n8),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_358])).
cnf(c_0_1112,plain,(rdn_translate(n89,rdn_pos(rdn(rdnn(n9),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_359])).
cnf(c_0_1113,plain,(rdn_translate(n90,rdn_pos(rdn(rdnn(n0),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_360])).
cnf(c_0_1114,plain,(rdn_translate(n91,rdn_pos(rdn(rdnn(n1),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_361])).
cnf(c_0_1115,plain,(rdn_translate(n92,rdn_pos(rdn(rdnn(n2),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_362])).
cnf(c_0_1116,plain,(rdn_translate(n93,rdn_pos(rdn(rdnn(n3),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_363])).
cnf(c_0_1117,plain,(rdn_translate(n94,rdn_pos(rdn(rdnn(n4),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_364])).
cnf(c_0_1118,plain,(rdn_translate(n95,rdn_pos(rdn(rdnn(n5),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_365])).
cnf(c_0_1119,plain,(rdn_translate(n96,rdn_pos(rdn(rdnn(n6),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_366])).
cnf(c_0_1120,plain,(rdn_translate(n97,rdn_pos(rdn(rdnn(n7),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_367])).
cnf(c_0_1121,plain,(rdn_translate(n98,rdn_pos(rdn(rdnn(n8),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_368])).
cnf(c_0_1122,plain,(rdn_translate(n99,rdn_pos(rdn(rdnn(n9),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_369])).
cnf(c_0_1123,plain,(rdn_translate(n100,rdn_pos(rdn(rdnn(n0),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_370])).
cnf(c_0_1124,plain,(rdn_translate(n101,rdn_pos(rdn(rdnn(n1),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_371])).
cnf(c_0_1125,plain,(rdn_translate(n102,rdn_pos(rdn(rdnn(n2),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_372])).
cnf(c_0_1126,plain,(rdn_translate(n103,rdn_pos(rdn(rdnn(n3),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_373])).
cnf(c_0_1127,plain,(rdn_translate(n104,rdn_pos(rdn(rdnn(n4),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_374])).
cnf(c_0_1128,plain,(rdn_translate(n105,rdn_pos(rdn(rdnn(n5),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_375])).
cnf(c_0_1129,plain,(rdn_translate(n106,rdn_pos(rdn(rdnn(n6),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_376])).
cnf(c_0_1130,plain,(rdn_translate(n107,rdn_pos(rdn(rdnn(n7),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_377])).
cnf(c_0_1131,plain,(rdn_translate(n108,rdn_pos(rdn(rdnn(n8),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_378])).
cnf(c_0_1132,plain,(rdn_translate(n109,rdn_pos(rdn(rdnn(n9),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_379])).
cnf(c_0_1133,plain,(rdn_translate(n110,rdn_pos(rdn(rdnn(n0),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_380])).
cnf(c_0_1134,plain,(rdn_translate(n111,rdn_pos(rdn(rdnn(n1),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_381])).
cnf(c_0_1135,plain,(rdn_translate(n112,rdn_pos(rdn(rdnn(n2),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_382])).
cnf(c_0_1136,plain,(rdn_translate(n113,rdn_pos(rdn(rdnn(n3),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_383])).
cnf(c_0_1137,plain,(rdn_translate(n114,rdn_pos(rdn(rdnn(n4),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_384])).
cnf(c_0_1138,plain,(rdn_translate(n115,rdn_pos(rdn(rdnn(n5),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_385])).
cnf(c_0_1139,plain,(rdn_translate(n116,rdn_pos(rdn(rdnn(n6),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_386])).
cnf(c_0_1140,plain,(rdn_translate(n117,rdn_pos(rdn(rdnn(n7),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_387])).
cnf(c_0_1141,plain,(rdn_translate(n118,rdn_pos(rdn(rdnn(n8),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_388])).
cnf(c_0_1142,plain,(rdn_translate(n119,rdn_pos(rdn(rdnn(n9),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_389])).
cnf(c_0_1143,plain,(rdn_translate(n120,rdn_pos(rdn(rdnn(n0),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_390])).
cnf(c_0_1144,plain,(rdn_translate(n121,rdn_pos(rdn(rdnn(n1),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_391])).
cnf(c_0_1145,plain,(rdn_translate(n122,rdn_pos(rdn(rdnn(n2),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_392])).
cnf(c_0_1146,plain,(rdn_translate(n123,rdn_pos(rdn(rdnn(n3),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_393])).
cnf(c_0_1147,plain,(rdn_translate(n124,rdn_pos(rdn(rdnn(n4),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_394])).
cnf(c_0_1148,plain,(rdn_translate(n125,rdn_pos(rdn(rdnn(n5),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_395])).
cnf(c_0_1149,plain,(rdn_translate(n126,rdn_pos(rdn(rdnn(n6),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_396])).
cnf(c_0_1150,plain,(rdn_translate(n127,rdn_pos(rdn(rdnn(n7),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_397])).
cnf(c_0_1151,plain,(rdn_translate(nn1,rdn_neg(rdnn(n1)))),inference(split_conjunct, [status(thm)],[c_0_398])).
cnf(c_0_1152,plain,(rdn_translate(nn2,rdn_neg(rdnn(n2)))),inference(split_conjunct, [status(thm)],[c_0_399])).
cnf(c_0_1153,plain,(rdn_translate(nn3,rdn_neg(rdnn(n3)))),inference(split_conjunct, [status(thm)],[c_0_400])).
cnf(c_0_1154,plain,(rdn_translate(nn4,rdn_neg(rdnn(n4)))),inference(split_conjunct, [status(thm)],[c_0_401])).
cnf(c_0_1155,plain,(rdn_translate(nn5,rdn_neg(rdnn(n5)))),inference(split_conjunct, [status(thm)],[c_0_402])).
cnf(c_0_1156,plain,(rdn_translate(nn6,rdn_neg(rdnn(n6)))),inference(split_conjunct, [status(thm)],[c_0_403])).
cnf(c_0_1157,plain,(rdn_translate(nn7,rdn_neg(rdnn(n7)))),inference(split_conjunct, [status(thm)],[c_0_404])).
cnf(c_0_1158,plain,(rdn_translate(nn8,rdn_neg(rdnn(n8)))),inference(split_conjunct, [status(thm)],[c_0_405])).
cnf(c_0_1159,plain,(rdn_translate(nn9,rdn_neg(rdnn(n9)))),inference(split_conjunct, [status(thm)],[c_0_406])).
cnf(c_0_1160,plain,(rdn_translate(nn10,rdn_neg(rdn(rdnn(n0),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_407])).
cnf(c_0_1161,plain,(rdn_translate(nn11,rdn_neg(rdn(rdnn(n1),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_408])).
cnf(c_0_1162,plain,(rdn_translate(nn12,rdn_neg(rdn(rdnn(n2),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_409])).
cnf(c_0_1163,plain,(rdn_translate(nn13,rdn_neg(rdn(rdnn(n3),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_410])).
cnf(c_0_1164,plain,(rdn_translate(nn14,rdn_neg(rdn(rdnn(n4),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_411])).
cnf(c_0_1165,plain,(rdn_translate(nn15,rdn_neg(rdn(rdnn(n5),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_412])).
cnf(c_0_1166,plain,(rdn_translate(nn16,rdn_neg(rdn(rdnn(n6),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_413])).
cnf(c_0_1167,plain,(rdn_translate(nn17,rdn_neg(rdn(rdnn(n7),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_414])).
cnf(c_0_1168,plain,(rdn_translate(nn18,rdn_neg(rdn(rdnn(n8),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_415])).
cnf(c_0_1169,plain,(rdn_translate(nn19,rdn_neg(rdn(rdnn(n9),rdnn(n1))))),inference(split_conjunct, [status(thm)],[c_0_416])).
cnf(c_0_1170,plain,(rdn_translate(nn20,rdn_neg(rdn(rdnn(n0),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_417])).
cnf(c_0_1171,plain,(rdn_translate(nn21,rdn_neg(rdn(rdnn(n1),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_418])).
cnf(c_0_1172,plain,(rdn_translate(nn22,rdn_neg(rdn(rdnn(n2),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_419])).
cnf(c_0_1173,plain,(rdn_translate(nn23,rdn_neg(rdn(rdnn(n3),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_420])).
cnf(c_0_1174,plain,(rdn_translate(nn24,rdn_neg(rdn(rdnn(n4),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_421])).
cnf(c_0_1175,plain,(rdn_translate(nn25,rdn_neg(rdn(rdnn(n5),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_422])).
cnf(c_0_1176,plain,(rdn_translate(nn26,rdn_neg(rdn(rdnn(n6),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_423])).
cnf(c_0_1177,plain,(rdn_translate(nn27,rdn_neg(rdn(rdnn(n7),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_424])).
cnf(c_0_1178,plain,(rdn_translate(nn28,rdn_neg(rdn(rdnn(n8),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_425])).
cnf(c_0_1179,plain,(rdn_translate(nn29,rdn_neg(rdn(rdnn(n9),rdnn(n2))))),inference(split_conjunct, [status(thm)],[c_0_426])).
cnf(c_0_1180,plain,(rdn_translate(nn30,rdn_neg(rdn(rdnn(n0),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_427])).
cnf(c_0_1181,plain,(rdn_translate(nn31,rdn_neg(rdn(rdnn(n1),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_428])).
cnf(c_0_1182,plain,(rdn_translate(nn32,rdn_neg(rdn(rdnn(n2),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_429])).
cnf(c_0_1183,plain,(rdn_translate(nn33,rdn_neg(rdn(rdnn(n3),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_430])).
cnf(c_0_1184,plain,(rdn_translate(nn34,rdn_neg(rdn(rdnn(n4),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_431])).
cnf(c_0_1185,plain,(rdn_translate(nn35,rdn_neg(rdn(rdnn(n5),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_432])).
cnf(c_0_1186,plain,(rdn_translate(nn36,rdn_neg(rdn(rdnn(n6),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_433])).
cnf(c_0_1187,plain,(rdn_translate(nn37,rdn_neg(rdn(rdnn(n7),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_434])).
cnf(c_0_1188,plain,(rdn_translate(nn38,rdn_neg(rdn(rdnn(n8),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_435])).
cnf(c_0_1189,plain,(rdn_translate(nn39,rdn_neg(rdn(rdnn(n9),rdnn(n3))))),inference(split_conjunct, [status(thm)],[c_0_436])).
cnf(c_0_1190,plain,(rdn_translate(nn40,rdn_neg(rdn(rdnn(n0),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_437])).
cnf(c_0_1191,plain,(rdn_translate(nn41,rdn_neg(rdn(rdnn(n1),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_438])).
cnf(c_0_1192,plain,(rdn_translate(nn42,rdn_neg(rdn(rdnn(n2),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_439])).
cnf(c_0_1193,plain,(rdn_translate(nn43,rdn_neg(rdn(rdnn(n3),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_440])).
cnf(c_0_1194,plain,(rdn_translate(nn44,rdn_neg(rdn(rdnn(n4),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_441])).
cnf(c_0_1195,plain,(rdn_translate(nn45,rdn_neg(rdn(rdnn(n5),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_442])).
cnf(c_0_1196,plain,(rdn_translate(nn46,rdn_neg(rdn(rdnn(n6),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_443])).
cnf(c_0_1197,plain,(rdn_translate(nn47,rdn_neg(rdn(rdnn(n7),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_444])).
cnf(c_0_1198,plain,(rdn_translate(nn48,rdn_neg(rdn(rdnn(n8),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_445])).
cnf(c_0_1199,plain,(rdn_translate(nn49,rdn_neg(rdn(rdnn(n9),rdnn(n4))))),inference(split_conjunct, [status(thm)],[c_0_446])).
cnf(c_0_1200,plain,(rdn_translate(nn50,rdn_neg(rdn(rdnn(n0),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_447])).
cnf(c_0_1201,plain,(rdn_translate(nn51,rdn_neg(rdn(rdnn(n1),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_448])).
cnf(c_0_1202,plain,(rdn_translate(nn52,rdn_neg(rdn(rdnn(n2),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_449])).
cnf(c_0_1203,plain,(rdn_translate(nn53,rdn_neg(rdn(rdnn(n3),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_450])).
cnf(c_0_1204,plain,(rdn_translate(nn54,rdn_neg(rdn(rdnn(n4),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_451])).
cnf(c_0_1205,plain,(rdn_translate(nn55,rdn_neg(rdn(rdnn(n5),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_452])).
cnf(c_0_1206,plain,(rdn_translate(nn56,rdn_neg(rdn(rdnn(n6),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_453])).
cnf(c_0_1207,plain,(rdn_translate(nn57,rdn_neg(rdn(rdnn(n7),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_454])).
cnf(c_0_1208,plain,(rdn_translate(nn58,rdn_neg(rdn(rdnn(n8),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_455])).
cnf(c_0_1209,plain,(rdn_translate(nn59,rdn_neg(rdn(rdnn(n9),rdnn(n5))))),inference(split_conjunct, [status(thm)],[c_0_456])).
cnf(c_0_1210,plain,(rdn_translate(nn60,rdn_neg(rdn(rdnn(n0),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_457])).
cnf(c_0_1211,plain,(rdn_translate(nn61,rdn_neg(rdn(rdnn(n1),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_458])).
cnf(c_0_1212,plain,(rdn_translate(nn62,rdn_neg(rdn(rdnn(n2),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_459])).
cnf(c_0_1213,plain,(rdn_translate(nn63,rdn_neg(rdn(rdnn(n3),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_460])).
cnf(c_0_1214,plain,(rdn_translate(nn64,rdn_neg(rdn(rdnn(n4),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_461])).
cnf(c_0_1215,plain,(rdn_translate(nn65,rdn_neg(rdn(rdnn(n5),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_462])).
cnf(c_0_1216,plain,(rdn_translate(nn66,rdn_neg(rdn(rdnn(n6),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_463])).
cnf(c_0_1217,plain,(rdn_translate(nn67,rdn_neg(rdn(rdnn(n7),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_464])).
cnf(c_0_1218,plain,(rdn_translate(nn68,rdn_neg(rdn(rdnn(n8),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_465])).
cnf(c_0_1219,plain,(rdn_translate(nn69,rdn_neg(rdn(rdnn(n9),rdnn(n6))))),inference(split_conjunct, [status(thm)],[c_0_466])).
cnf(c_0_1220,plain,(rdn_translate(nn70,rdn_neg(rdn(rdnn(n0),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_467])).
cnf(c_0_1221,plain,(rdn_translate(nn71,rdn_neg(rdn(rdnn(n1),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_468])).
cnf(c_0_1222,plain,(rdn_translate(nn72,rdn_neg(rdn(rdnn(n2),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_469])).
cnf(c_0_1223,plain,(rdn_translate(nn73,rdn_neg(rdn(rdnn(n3),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_470])).
cnf(c_0_1224,plain,(rdn_translate(nn74,rdn_neg(rdn(rdnn(n4),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_471])).
cnf(c_0_1225,plain,(rdn_translate(nn75,rdn_neg(rdn(rdnn(n5),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_472])).
cnf(c_0_1226,plain,(rdn_translate(nn76,rdn_neg(rdn(rdnn(n6),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_473])).
cnf(c_0_1227,plain,(rdn_translate(nn77,rdn_neg(rdn(rdnn(n7),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_474])).
cnf(c_0_1228,plain,(rdn_translate(nn78,rdn_neg(rdn(rdnn(n8),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_475])).
cnf(c_0_1229,plain,(rdn_translate(nn79,rdn_neg(rdn(rdnn(n9),rdnn(n7))))),inference(split_conjunct, [status(thm)],[c_0_476])).
cnf(c_0_1230,plain,(rdn_translate(nn80,rdn_neg(rdn(rdnn(n0),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_477])).
cnf(c_0_1231,plain,(rdn_translate(nn81,rdn_neg(rdn(rdnn(n1),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_478])).
cnf(c_0_1232,plain,(rdn_translate(nn82,rdn_neg(rdn(rdnn(n2),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_479])).
cnf(c_0_1233,plain,(rdn_translate(nn83,rdn_neg(rdn(rdnn(n3),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_480])).
cnf(c_0_1234,plain,(rdn_translate(nn84,rdn_neg(rdn(rdnn(n4),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_481])).
cnf(c_0_1235,plain,(rdn_translate(nn85,rdn_neg(rdn(rdnn(n5),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_482])).
cnf(c_0_1236,plain,(rdn_translate(nn86,rdn_neg(rdn(rdnn(n6),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_483])).
cnf(c_0_1237,plain,(rdn_translate(nn87,rdn_neg(rdn(rdnn(n7),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_484])).
cnf(c_0_1238,plain,(rdn_translate(nn88,rdn_neg(rdn(rdnn(n8),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_485])).
cnf(c_0_1239,plain,(rdn_translate(nn89,rdn_neg(rdn(rdnn(n9),rdnn(n8))))),inference(split_conjunct, [status(thm)],[c_0_486])).
cnf(c_0_1240,plain,(rdn_translate(nn90,rdn_neg(rdn(rdnn(n0),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_487])).
cnf(c_0_1241,plain,(rdn_translate(nn91,rdn_neg(rdn(rdnn(n1),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_488])).
cnf(c_0_1242,plain,(rdn_translate(nn92,rdn_neg(rdn(rdnn(n2),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_489])).
cnf(c_0_1243,plain,(rdn_translate(nn93,rdn_neg(rdn(rdnn(n3),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_490])).
cnf(c_0_1244,plain,(rdn_translate(nn94,rdn_neg(rdn(rdnn(n4),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_491])).
cnf(c_0_1245,plain,(rdn_translate(nn95,rdn_neg(rdn(rdnn(n5),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_492])).
cnf(c_0_1246,plain,(rdn_translate(nn96,rdn_neg(rdn(rdnn(n6),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_493])).
cnf(c_0_1247,plain,(rdn_translate(nn97,rdn_neg(rdn(rdnn(n7),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_494])).
cnf(c_0_1248,plain,(rdn_translate(nn98,rdn_neg(rdn(rdnn(n8),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_495])).
cnf(c_0_1249,plain,(rdn_translate(nn99,rdn_neg(rdn(rdnn(n9),rdnn(n9))))),inference(split_conjunct, [status(thm)],[c_0_496])).
cnf(c_0_1250,plain,(rdn_translate(nn100,rdn_neg(rdn(rdnn(n0),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_497])).
cnf(c_0_1251,plain,(rdn_translate(nn101,rdn_neg(rdn(rdnn(n1),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_498])).
cnf(c_0_1252,plain,(rdn_translate(nn102,rdn_neg(rdn(rdnn(n2),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_499])).
cnf(c_0_1253,plain,(rdn_translate(nn103,rdn_neg(rdn(rdnn(n3),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_500])).
cnf(c_0_1254,plain,(rdn_translate(nn104,rdn_neg(rdn(rdnn(n4),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_501])).
cnf(c_0_1255,plain,(rdn_translate(nn105,rdn_neg(rdn(rdnn(n5),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_502])).
cnf(c_0_1256,plain,(rdn_translate(nn106,rdn_neg(rdn(rdnn(n6),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_503])).
cnf(c_0_1257,plain,(rdn_translate(nn107,rdn_neg(rdn(rdnn(n7),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_504])).
cnf(c_0_1258,plain,(rdn_translate(nn108,rdn_neg(rdn(rdnn(n8),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_505])).
cnf(c_0_1259,plain,(rdn_translate(nn109,rdn_neg(rdn(rdnn(n9),rdn(rdnn(n0),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_506])).
cnf(c_0_1260,plain,(rdn_translate(nn110,rdn_neg(rdn(rdnn(n0),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_507])).
cnf(c_0_1261,plain,(rdn_translate(nn111,rdn_neg(rdn(rdnn(n1),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_508])).
cnf(c_0_1262,plain,(rdn_translate(nn112,rdn_neg(rdn(rdnn(n2),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_509])).
cnf(c_0_1263,plain,(rdn_translate(nn113,rdn_neg(rdn(rdnn(n3),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_510])).
cnf(c_0_1264,plain,(rdn_translate(nn114,rdn_neg(rdn(rdnn(n4),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_511])).
cnf(c_0_1265,plain,(rdn_translate(nn115,rdn_neg(rdn(rdnn(n5),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_512])).
cnf(c_0_1266,plain,(rdn_translate(nn116,rdn_neg(rdn(rdnn(n6),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_513])).
cnf(c_0_1267,plain,(rdn_translate(nn117,rdn_neg(rdn(rdnn(n7),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_514])).
cnf(c_0_1268,plain,(rdn_translate(nn118,rdn_neg(rdn(rdnn(n8),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_515])).
cnf(c_0_1269,plain,(rdn_translate(nn119,rdn_neg(rdn(rdnn(n9),rdn(rdnn(n1),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_516])).
cnf(c_0_1270,plain,(rdn_translate(nn120,rdn_neg(rdn(rdnn(n0),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_517])).
cnf(c_0_1271,plain,(rdn_translate(nn121,rdn_neg(rdn(rdnn(n1),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_518])).
cnf(c_0_1272,plain,(rdn_translate(nn122,rdn_neg(rdn(rdnn(n2),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_519])).
cnf(c_0_1273,plain,(rdn_translate(nn123,rdn_neg(rdn(rdnn(n3),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_520])).
cnf(c_0_1274,plain,(rdn_translate(nn124,rdn_neg(rdn(rdnn(n4),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_521])).
cnf(c_0_1275,plain,(rdn_translate(nn125,rdn_neg(rdn(rdnn(n5),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_522])).
cnf(c_0_1276,plain,(rdn_translate(nn126,rdn_neg(rdn(rdnn(n6),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_523])).
cnf(c_0_1277,plain,(rdn_translate(nn127,rdn_neg(rdn(rdnn(n7),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_524])).
cnf(c_0_1278,plain,(rdn_translate(nn128,rdn_neg(rdn(rdnn(n8),rdn(rdnn(n2),rdnn(n1)))))),inference(split_conjunct, [status(thm)],[c_0_525])).
cnf(c_0_1279,plain,(rdn_non_zero_digit(rdnn(n1))),inference(split_conjunct, [status(thm)],[c_0_526])).
cnf(c_0_1280,plain,(rdn_non_zero_digit(rdnn(n2))),inference(split_conjunct, [status(thm)],[c_0_527])).
cnf(c_0_1281,plain,(rdn_non_zero_digit(rdnn(n3))),inference(split_conjunct, [status(thm)],[c_0_528])).
cnf(c_0_1282,plain,(rdn_non_zero_digit(rdnn(n4))),inference(split_conjunct, [status(thm)],[c_0_529])).
cnf(c_0_1283,plain,(rdn_non_zero_digit(rdnn(n5))),inference(split_conjunct, [status(thm)],[c_0_530])).
cnf(c_0_1284,plain,(rdn_non_zero_digit(rdnn(n6))),inference(split_conjunct, [status(thm)],[c_0_531])).
cnf(c_0_1285,plain,(rdn_non_zero_digit(rdnn(n7))),inference(split_conjunct, [status(thm)],[c_0_532])).
cnf(c_0_1286,plain,(rdn_non_zero_digit(rdnn(n8))),inference(split_conjunct, [status(thm)],[c_0_533])).
cnf(c_0_1287,plain,(rdn_non_zero_digit(rdnn(n9))),inference(split_conjunct, [status(thm)],[c_0_534])).
cnf(c_0_1288,plain,(rdn_positive_less(rdnn(n0),rdnn(n1))),inference(split_conjunct, [status(thm)],[c_0_535])).
cnf(c_0_1289,plain,(rdn_positive_less(rdnn(n1),rdnn(n2))),inference(split_conjunct, [status(thm)],[c_0_536])).
cnf(c_0_1290,plain,(rdn_positive_less(rdnn(n2),rdnn(n3))),inference(split_conjunct, [status(thm)],[c_0_537])).
cnf(c_0_1291,plain,(rdn_positive_less(rdnn(n3),rdnn(n4))),inference(split_conjunct, [status(thm)],[c_0_538])).
cnf(c_0_1292,plain,(rdn_positive_less(rdnn(n4),rdnn(n5))),inference(split_conjunct, [status(thm)],[c_0_539])).
cnf(c_0_1293,plain,(rdn_positive_less(rdnn(n5),rdnn(n6))),inference(split_conjunct, [status(thm)],[c_0_540])).
cnf(c_0_1294,plain,(rdn_positive_less(rdnn(n6),rdnn(n7))),inference(split_conjunct, [status(thm)],[c_0_541])).
cnf(c_0_1295,plain,(rdn_positive_less(rdnn(n7),rdnn(n8))),inference(split_conjunct, [status(thm)],[c_0_542])).
cnf(c_0_1296,plain,(rdn_positive_less(rdnn(n8),rdnn(n9))),inference(split_conjunct, [status(thm)],[c_0_543])).
fof(c_0_1297, plain, (![X7]:![X8]:![X9]:((~rdn_positive_less(rdnn(X7),rdnn(X8))|~rdn_positive_less(rdnn(X8),rdnn(X9)))|rdn_positive_less(rdnn(X7),rdnn(X9)))),inference(fof_nnf, [status(thm)],[c_0_544])).
fof(c_0_1298, plain, (![X10]:![X11]:![X12]:((~rdn_positive_less(rdnn(X10),rdnn(X11))|~rdn_positive_less(rdnn(X11),rdnn(X12)))|rdn_positive_less(rdnn(X10),rdnn(X12)))),inference(variable_rename, [status(thm)],[c_0_1297])).
cnf(c_0_1299,plain,(rdn_positive_less(rdnn(X1),rdnn(X2))|~rdn_positive_less(rdnn(X3),rdnn(X2))|~rdn_positive_less(rdnn(X1),rdnn(X3))),inference(split_conjunct, [status(thm)],[c_0_1298])).
fof(c_0_1300, plain, (![X10]:![X11]:![X12]:![X13]:(~rdn_positive_less(X11,X13)|rdn_positive_less(rdn(rdnn(X10),X11),rdn(rdnn(X12),X13)))),inference(fof_nnf, [status(thm)],[c_0_545])).
fof(c_0_1301, plain, (![X14]:![X15]:![X16]:![X17]:(~rdn_positive_less(X15,X17)|rdn_positive_less(rdn(rdnn(X14),X15),rdn(rdnn(X16),X17)))),inference(variable_rename, [status(thm)],[c_0_1300])).
cnf(c_0_1302,plain,(rdn_positive_less(rdn(rdnn(X1),X2),rdn(rdnn(X3),X4))|~rdn_positive_less(X2,X4)),inference(split_conjunct, [status(thm)],[c_0_1301])).
fof(c_0_1303, plain, (![X10]:![X14]:![X12]:((~rdn_positive_less(rdnn(X10),rdnn(X12))|~rdn_non_zero(X14))|rdn_positive_less(rdn(rdnn(X10),X14),rdn(rdnn(X12),X14)))),inference(fof_nnf, [status(thm)],[c_0_546])).
fof(c_0_1304, plain, (![X15]:![X16]:![X17]:((~rdn_positive_less(rdnn(X15),rdnn(X17))|~rdn_non_zero(X16))|rdn_positive_less(rdn(rdnn(X15),X16),rdn(rdnn(X17),X16)))),inference(variable_rename, [status(thm)],[c_0_1303])).
cnf(c_0_1305,plain,(rdn_positive_less(rdn(rdnn(X1),X2),rdn(rdnn(X3),X2))|~rdn_non_zero(X2)|~rdn_positive_less(rdnn(X1),rdnn(X3))),inference(split_conjunct, [status(thm)],[c_0_1304])).
fof(c_0_1306, plain, (![X15]:![X12]:![X13]:(~rdn_non_zero(X13)|rdn_positive_less(rdnn(X15),rdn(rdnn(X12),X13)))),inference(fof_nnf, [status(thm)],[c_0_547])).
fof(c_0_1307, plain, (![X16]:![X17]:![X18]:(~rdn_non_zero(X18)|rdn_positive_less(rdnn(X16),rdn(rdnn(X17),X18)))),inference(variable_rename, [status(thm)],[c_0_1306])).
cnf(c_0_1308,plain,(rdn_positive_less(rdnn(X1),rdn(rdnn(X2),X3))|~rdn_non_zero(X3)),inference(split_conjunct, [status(thm)],[c_0_1307])).
fof(c_0_1309, plain, (![X7]:(~rdn_non_zero_digit(rdnn(X7))|rdn_non_zero(rdnn(X7)))),inference(fof_nnf, [status(thm)],[c_0_548])).
fof(c_0_1310, plain, (![X8]:(~rdn_non_zero_digit(rdnn(X8))|rdn_non_zero(rdnn(X8)))),inference(variable_rename, [status(thm)],[c_0_1309])).
cnf(c_0_1311,plain,(rdn_non_zero(rdnn(X1))|~rdn_non_zero_digit(rdnn(X1))),inference(split_conjunct, [status(thm)],[c_0_1310])).
fof(c_0_1312, plain, (![X15]:![X14]:(~rdn_non_zero(X14)|rdn_non_zero(rdn(rdnn(X15),X14)))),inference(fof_nnf, [status(thm)],[c_0_549])).
fof(c_0_1313, plain, (![X16]:![X17]:(~rdn_non_zero(X17)|rdn_non_zero(rdn(rdnn(X16),X17)))),inference(variable_rename, [status(thm)],[c_0_1312])).
cnf(c_0_1314,plain,(rdn_non_zero(rdn(rdnn(X1),X2))|~rdn_non_zero(X2)),inference(split_conjunct, [status(thm)],[c_0_1313])).
fof(c_0_1315, plain, (![X7]:![X8]:![X16]:![X17]:(((~rdn_translate(X7,rdn_pos(X16))|~rdn_translate(X8,rdn_pos(X17)))|~rdn_positive_less(X16,X17))|less(X7,X8))),inference(fof_nnf, [status(thm)],[c_0_550])).
fof(c_0_1316, plain, (![X7]:![X8]:(![X16]:![X17]:((~rdn_translate(X7,rdn_pos(X16))|~rdn_translate(X8,rdn_pos(X17)))|~rdn_positive_less(X16,X17))|less(X7,X8))),inference(shift_quantors, [status(thm)],[c_0_1315])).
fof(c_0_1317, plain, (![X18]:![X19]:(![X20]:![X21]:((~rdn_translate(X18,rdn_pos(X20))|~rdn_translate(X19,rdn_pos(X21)))|~rdn_positive_less(X20,X21))|less(X18,X19))),inference(variable_rename, [status(thm)],[c_0_1316])).
fof(c_0_1318, plain, (![X18]:![X19]:![X20]:![X21]:(((~rdn_translate(X18,rdn_pos(X20))|~rdn_translate(X19,rdn_pos(X21)))|~rdn_positive_less(X20,X21))|less(X18,X19))),inference(shift_quantors, [status(thm)],[c_0_1317])).
cnf(c_0_1319,plain,(less(X1,X2)|~rdn_positive_less(X3,X4)|~rdn_translate(X2,rdn_pos(X4))|~rdn_translate(X1,rdn_pos(X3))),inference(split_conjunct, [status(thm)],[c_0_1318])).
fof(c_0_1320, plain, (![X7]:![X8]:![X16]:![X17]:((~rdn_translate(X7,rdn_neg(X16))|~rdn_translate(X8,rdn_pos(X17)))|less(X7,X8))),inference(fof_nnf, [status(thm)],[c_0_551])).
fof(c_0_1321, plain, (![X7]:![X8]:((![X16]:~rdn_translate(X7,rdn_neg(X16))|![X17]:~rdn_translate(X8,rdn_pos(X17)))|less(X7,X8))),inference(shift_quantors, [status(thm)],[c_0_1320])).
fof(c_0_1322, plain, (![X18]:![X19]:((![X20]:~rdn_translate(X18,rdn_neg(X20))|![X21]:~rdn_translate(X19,rdn_pos(X21)))|less(X18,X19))),inference(variable_rename, [status(thm)],[c_0_1321])).
fof(c_0_1323, plain, (![X18]:![X19]:![X20]:![X21]:((~rdn_translate(X18,rdn_neg(X20))|~rdn_translate(X19,rdn_pos(X21)))|less(X18,X19))),inference(shift_quantors, [status(thm)],[c_0_1322])).
cnf(c_0_1324,plain,(less(X1,X2)|~rdn_translate(X2,rdn_pos(X3))|~rdn_translate(X1,rdn_neg(X4))),inference(split_conjunct, [status(thm)],[c_0_1323])).
fof(c_0_1325, plain, (![X7]:![X8]:![X16]:![X17]:(((~rdn_translate(X7,rdn_neg(X16))|~rdn_translate(X8,rdn_neg(X17)))|~rdn_positive_less(X17,X16))|less(X7,X8))),inference(fof_nnf, [status(thm)],[c_0_552])).
fof(c_0_1326, plain, (![X7]:![X8]:(![X16]:![X17]:((~rdn_translate(X7,rdn_neg(X16))|~rdn_translate(X8,rdn_neg(X17)))|~rdn_positive_less(X17,X16))|less(X7,X8))),inference(shift_quantors, [status(thm)],[c_0_1325])).
fof(c_0_1327, plain, (![X18]:![X19]:(![X20]:![X21]:((~rdn_translate(X18,rdn_neg(X20))|~rdn_translate(X19,rdn_neg(X21)))|~rdn_positive_less(X21,X20))|less(X18,X19))),inference(variable_rename, [status(thm)],[c_0_1326])).
fof(c_0_1328, plain, (![X18]:![X19]:![X20]:![X21]:(((~rdn_translate(X18,rdn_neg(X20))|~rdn_translate(X19,rdn_neg(X21)))|~rdn_positive_less(X21,X20))|less(X18,X19))),inference(shift_quantors, [status(thm)],[c_0_1327])).
cnf(c_0_1329,plain,(less(X1,X2)|~rdn_positive_less(X3,X4)|~rdn_translate(X2,rdn_neg(X3))|~rdn_translate(X1,rdn_neg(X4))),inference(split_conjunct, [status(thm)],[c_0_1328])).
fof(c_0_1330, plain, (![X7]:![X8]:((~less(X7,X8)|(~less(X8,X7)&X8!=X7))&((less(X8,X7)|X8=X7)|less(X7,X8)))),inference(fof_nnf, [status(thm)],[c_0_598])).
fof(c_0_1331, plain, ((![X7]:![X8]:(~less(X7,X8)|(~less(X8,X7)&X8!=X7))&![X7]:![X8]:((less(X8,X7)|X8=X7)|less(X7,X8)))),inference(shift_quantors, [status(thm)],[c_0_1330])).
fof(c_0_1332, plain, ((![X9]:![X10]:(~less(X9,X10)|(~less(X10,X9)&X10!=X9))&![X11]:![X12]:((less(X12,X11)|X12=X11)|less(X11,X12)))),inference(variable_rename, [status(thm)],[c_0_1331])).
fof(c_0_1333, plain, (![X9]:![X10]:![X11]:![X12]:((~less(X9,X10)|(~less(X10,X9)&X10!=X9))&((less(X12,X11)|X12=X11)|less(X11,X12)))),inference(shift_quantors, [status(thm)],[c_0_1332])).
fof(c_0_1334, plain, (![X9]:![X10]:![X11]:![X12]:(((~less(X10,X9)|~less(X9,X10))&(X10!=X9|~less(X9,X10)))&((less(X12,X11)|X12=X11)|less(X11,X12)))),inference(distribute, [status(thm)],[c_0_1333])).
cnf(c_0_1335,plain,(less(X1,X2)|X2=X1|less(X2,X1)),inference(split_conjunct, [status(thm)],[c_0_1334])).
cnf(c_0_1336,plain,(~less(X1,X2)|X2!=X1),inference(split_conjunct, [status(thm)],[c_0_1334])).
cnf(c_0_1337,plain,(~less(X1,X2)|~less(X2,X1)),inference(split_conjunct, [status(thm)],[c_0_1334])).
fof(c_0_1338, plain, (![X7]:![X8]:((~less_or_equal(X7,X8)|(less(X7,X8)|X7=X8))&((~less(X7,X8)&X7!=X8)|less_or_equal(X7,X8)))),inference(fof_nnf, [status(thm)],[c_0_554])).
fof(c_0_1339, plain, ((![X7]:![X8]:(~less_or_equal(X7,X8)|(less(X7,X8)|X7=X8))&![X7]:![X8]:((~less(X7,X8)&X7!=X8)|less_or_equal(X7,X8)))),inference(shift_quantors, [status(thm)],[c_0_1338])).
fof(c_0_1340, plain, ((![X9]:![X10]:(~less_or_equal(X9,X10)|(less(X9,X10)|X9=X10))&![X11]:![X12]:((~less(X11,X12)&X11!=X12)|less_or_equal(X11,X12)))),inference(variable_rename, [status(thm)],[c_0_1339])).
fof(c_0_1341, plain, (![X9]:![X10]:![X11]:![X12]:((~less_or_equal(X9,X10)|(less(X9,X10)|X9=X10))&((~less(X11,X12)&X11!=X12)|less_or_equal(X11,X12)))),inference(shift_quantors, [status(thm)],[c_0_1340])).
fof(c_0_1342, plain, (![X9]:![X10]:![X11]:![X12]:((~less_or_equal(X9,X10)|(less(X9,X10)|X9=X10))&((~less(X11,X12)|less_or_equal(X11,X12))&(X11!=X12|less_or_equal(X11,X12))))),inference(distribute, [status(thm)],[c_0_1341])).
cnf(c_0_1343,plain,(less_or_equal(X1,X2)|X1!=X2),inference(split_conjunct, [status(thm)],[c_0_1342])).
cnf(c_0_1344,plain,(less_or_equal(X1,X2)|~less(X1,X2)),inference(split_conjunct, [status(thm)],[c_0_1342])).
cnf(c_0_1345,plain,(X1=X2|less(X1,X2)|~less_or_equal(X1,X2)),inference(split_conjunct, [status(thm)],[c_0_1342])).
fof(c_0_1346, plain, (![X7]:![X8]:![X9]:((~sum(X7,n1,X8)|~less(X9,X8))|less_or_equal(X9,X7))),inference(fof_nnf, [status(thm)],[c_0_555])).
fof(c_0_1347, plain, (![X10]:![X11]:![X12]:((~sum(X10,n1,X11)|~less(X12,X11))|less_or_equal(X12,X10))),inference(variable_rename, [status(thm)],[c_0_1346])).
cnf(c_0_1348,plain,(less_or_equal(X1,X2)|~less(X1,X3)|~sum(X2,n1,X3)),inference(split_conjunct, [status(thm)],[c_0_1347])).
cnf(c_0_1349,negated_conjecture,(~accept_team(countryamedicalorganization,countryacivilorganization,towna,n6)),inference(split_conjunct, [status(thm)],[c_0_599])).
# Initializing proof state
# Scanning for AC axioms
cnf(c_0_1350,plain,(accept_city(sufferterragovernment,cityb)),inference(pm,[status(thm)],[c_0_618,c_0_919,theory(equality)])).
cnf(c_0_1351,plain,(accept_city(sufferterragovernment,coastvillage)),inference(pm,[status(thm)],[c_0_618,c_0_859,theory(equality)])).
cnf(c_0_1352,plain,(accept_city(christiansufferterrahumanitarianorganization,cityb)),inference(pm,[status(thm)],[c_0_618,c_0_914,theory(equality)])).
cnf(c_0_1353,plain,(accept_city(christiansufferterrahumanitarianorganization,coastvillage)),inference(pm,[status(thm)],[c_0_618,c_0_867,theory(equality)])).
cnf(c_0_1354,plain,(accept_city(countrybcivilorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_982,theory(equality)])).
cnf(c_0_1355,plain,(accept_city(countrybcivilorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_819,theory(equality)])).
cnf(c_0_1356,plain,(accept_city(countrybcivilorganization,townc)),inference(pm,[status(thm)],[c_0_618,c_0_853,theory(equality)])).
cnf(c_0_1357,plain,(accept_city(countryamedicalorganization,cityb)),inference(pm,[status(thm)],[c_0_618,c_0_931,theory(equality)])).
cnf(c_0_1358,plain,(accept_city(countrycmedicalorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_1005,theory(equality)])).
cnf(c_0_1359,plain,(accept_city(countrycmedicalorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_822,theory(equality)])).
cnf(c_0_1360,plain,(accept_city(countrycmedicalorganization,townc)),inference(pm,[status(thm)],[c_0_618,c_0_905,theory(equality)])).
cnf(c_0_1361,plain,(accept_city(muslimcountrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_618,c_0_877,theory(equality)])).
cnf(c_0_1362,plain,(accept_city(countrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_974,theory(equality)])).
cnf(c_0_1363,plain,(accept_city(countrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_795,theory(equality)])).
cnf(c_0_1364,plain,(accept_city(countrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_618,c_0_910,theory(equality)])).
cnf(c_0_1365,plain,(accept_city(countryccivilorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_993,theory(equality)])).
cnf(c_0_1366,plain,(accept_city(countryccivilorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_814,theory(equality)])).
cnf(c_0_1367,plain,(accept_city(countryccivilorganization,townc)),inference(pm,[status(thm)],[c_0_618,c_0_893,theory(equality)])).
cnf(c_0_1368,plain,(accept_city(christiancountrychumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_1016,theory(equality)])).
cnf(c_0_1369,plain,(accept_city(christiancountrychumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_618,c_0_829,theory(equality)])).
cnf(c_0_1370,plain,(accept_city(christiancountrychumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_618,c_0_886,theory(equality)])).
cnf(c_0_1371,plain,(accept_city(countryafirstaidorganization,cityb)),inference(pm,[status(thm)],[c_0_618,c_0_923,theory(equality)])).
cnf(c_0_1372,plain,(accept_city(countryafirstaidorganization,coastvillage)),inference(pm,[status(thm)],[c_0_618,c_0_842,theory(equality)])).
cnf(c_0_1373,plain,(accept_city(countryacivilorganization,cityb)),inference(pm,[status(thm)],[c_0_618,c_0_942,theory(equality)])).
cnf(c_0_1374,plain,(accept_city(countryacivilorganization,coastvillage)),inference(pm,[status(thm)],[c_0_618,c_0_836,theory(equality)])).
cnf(c_0_1375,plain,(accept_city(countryahumanitarianorganization,cityb)),inference(pm,[status(thm)],[c_0_618,c_0_937,theory(equality)])).
cnf(c_0_1376,plain,(accept_city(countryahumanitarianorganization,coastvillage)),inference(pm,[status(thm)],[c_0_618,c_0_778,theory(equality)])).
cnf(c_0_1377,plain,(~accept_leader(countryamedicalorganization,christiancountrychumanitarianorganization)|~accept_city(countryamedicalorganization,coastvillage)),inference(pm,[status(thm)],[c_0_627,c_0_810,theory(equality)])).
cnf(c_0_1378,plain,(~accept_leader(muslimcountrybhumanitarianorganization,sufferterragovernment)|~accept_city(muslimcountrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_627,c_0_1021,theory(equality)])).
cnf(c_0_1379,plain,(~accept_leader(muslimcountrybhumanitarianorganization,countryacivilorganization)|~accept_city(muslimcountrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_627,c_0_784,theory(equality)])).
cnf(c_0_1380,plain,(accept_leader(sufferterragovernment,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_919,theory(equality)])).
cnf(c_0_1381,plain,(accept_leader(sufferterragovernment,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_859,theory(equality)])).
cnf(c_0_1382,plain,(accept_leader(christiansufferterrahumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_914,theory(equality)])).
cnf(c_0_1383,plain,(accept_leader(christiansufferterrahumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_867,theory(equality)])).
cnf(c_0_1384,plain,(accept_leader(countrybcivilorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_619,c_0_982,theory(equality)])).
cnf(c_0_1385,plain,(accept_leader(countrybcivilorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_619,c_0_819,theory(equality)])).
cnf(c_0_1386,plain,(accept_leader(countrybcivilorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_853,theory(equality)])).
cnf(c_0_1387,plain,(accept_leader(countryamedicalorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_931,theory(equality)])).
cnf(c_0_1388,plain,(accept_leader(countrycmedicalorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_619,c_0_1005,theory(equality)])).
cnf(c_0_1389,plain,(accept_leader(countrycmedicalorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_619,c_0_822,theory(equality)])).
cnf(c_0_1390,plain,(accept_leader(countrycmedicalorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_905,theory(equality)])).
cnf(c_0_1391,plain,(accept_leader(muslimcountrybhumanitarianorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_877,theory(equality)])).
cnf(c_0_1392,plain,(accept_leader(countrybhumanitarianorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_619,c_0_974,theory(equality)])).
cnf(c_0_1393,plain,(accept_leader(countrybhumanitarianorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_619,c_0_795,theory(equality)])).
cnf(c_0_1394,plain,(accept_leader(countrybhumanitarianorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_910,theory(equality)])).
cnf(c_0_1395,plain,(accept_leader(countryccivilorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_619,c_0_993,theory(equality)])).
cnf(c_0_1396,plain,(accept_leader(countryccivilorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_619,c_0_814,theory(equality)])).
cnf(c_0_1397,plain,(accept_leader(countryccivilorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_893,theory(equality)])).
cnf(c_0_1398,plain,(accept_leader(christiancountrychumanitarianorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_619,c_0_1016,theory(equality)])).
cnf(c_0_1399,plain,(accept_leader(christiancountrychumanitarianorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_619,c_0_829,theory(equality)])).
cnf(c_0_1400,plain,(accept_leader(christiancountrychumanitarianorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_886,theory(equality)])).
cnf(c_0_1401,plain,(accept_leader(countryafirstaidorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_923,theory(equality)])).
cnf(c_0_1402,plain,(accept_leader(countryafirstaidorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_842,theory(equality)])).
cnf(c_0_1403,plain,(accept_leader(countryacivilorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_942,theory(equality)])).
cnf(c_0_1404,plain,(accept_leader(countryacivilorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_836,theory(equality)])).
cnf(c_0_1405,plain,(accept_leader(countryahumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_937,theory(equality)])).
cnf(c_0_1406,plain,(accept_leader(countryahumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_619,c_0_778,theory(equality)])).
cnf(c_0_1407,plain,(accept_leader(sufferterragovernment,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_920,theory(equality)])).
cnf(c_0_1408,plain,(accept_leader(sufferterragovernment,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_860,theory(equality)])).
cnf(c_0_1409,plain,(accept_leader(christiansufferterrahumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_915,theory(equality)])).
cnf(c_0_1410,plain,(accept_leader(christiansufferterrahumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_868,theory(equality)])).
cnf(c_0_1411,plain,(accept_leader(countrybcivilorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_624,c_0_983,theory(equality)])).
cnf(c_0_1412,plain,(accept_leader(countrybcivilorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_624,c_0_820,theory(equality)])).
cnf(c_0_1413,plain,(accept_leader(countrybcivilorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_854,theory(equality)])).
cnf(c_0_1414,plain,(accept_leader(countryamedicalorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_932,theory(equality)])).
cnf(c_0_1415,plain,(accept_leader(countryamedicalorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_811,theory(equality)])).
cnf(c_0_1416,plain,(accept_leader(countrycmedicalorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_624,c_0_1006,theory(equality)])).
cnf(c_0_1417,plain,(accept_leader(countrycmedicalorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_624,c_0_823,theory(equality)])).
cnf(c_0_1418,plain,(accept_leader(countrycmedicalorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_906,theory(equality)])).
cnf(c_0_1419,plain,(accept_leader(muslimcountrybhumanitarianorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_624,c_0_1022,theory(equality)])).
cnf(c_0_1420,plain,(accept_leader(muslimcountrybhumanitarianorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_624,c_0_785,theory(equality)])).
cnf(c_0_1421,plain,(accept_leader(muslimcountrybhumanitarianorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_878,theory(equality)])).
cnf(c_0_1422,plain,(accept_leader(countrybhumanitarianorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_624,c_0_975,theory(equality)])).
cnf(c_0_1423,plain,(accept_leader(countrybhumanitarianorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_624,c_0_796,theory(equality)])).
cnf(c_0_1424,plain,(accept_leader(countrybhumanitarianorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_911,theory(equality)])).
cnf(c_0_1425,plain,(accept_leader(countryccivilorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_624,c_0_994,theory(equality)])).
cnf(c_0_1426,plain,(accept_leader(countryccivilorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_624,c_0_815,theory(equality)])).
cnf(c_0_1427,plain,(accept_leader(countryccivilorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_894,theory(equality)])).
cnf(c_0_1428,plain,(accept_leader(christiancountrychumanitarianorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_624,c_0_1017,theory(equality)])).
cnf(c_0_1429,plain,(accept_leader(christiancountrychumanitarianorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_624,c_0_830,theory(equality)])).
cnf(c_0_1430,plain,(accept_leader(christiancountrychumanitarianorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_887,theory(equality)])).
cnf(c_0_1431,plain,(accept_leader(countryafirstaidorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_924,theory(equality)])).
cnf(c_0_1432,plain,(accept_leader(countryafirstaidorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_843,theory(equality)])).
cnf(c_0_1433,plain,(accept_leader(countryacivilorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_943,theory(equality)])).
cnf(c_0_1434,plain,(accept_leader(countryacivilorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_837,theory(equality)])).
cnf(c_0_1435,plain,(accept_leader(countryahumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_938,theory(equality)])).
cnf(c_0_1436,plain,(accept_leader(countryahumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_624,c_0_779,theory(equality)])).
cnf(c_0_1437,plain,(~accept_city(countryamedicalorganization,coastvillage)),inference(cn,[status(thm)],[c_0_1377, theory(equality,[symmetry])])).
cnf(c_0_1438,plain,(~accept_city(muslimcountrybhumanitarianorganization,towna)),inference(cn,[status(thm)],[c_0_1378, theory(equality,[symmetry])])).
cnf(c_0_1439,plain,(~accept_city(muslimcountrybhumanitarianorganization,towna)),inference(cn,[status(thm)],[c_0_1379, theory(equality,[symmetry])])).
cnf(c_0_1440,plain,(rdn_non_zero(rdnn(n6))),inference(pm,[status(thm)],[c_0_1311,c_0_1284,theory(equality)])).
cnf(c_0_1441,plain,(rdn_non_zero(rdnn(n5))),inference(pm,[status(thm)],[c_0_1311,c_0_1283,theory(equality)])).
cnf(c_0_1442,plain,(rdn_non_zero(rdnn(n2))),inference(pm,[status(thm)],[c_0_1311,c_0_1280,theory(equality)])).
cnf(c_0_1443,plain,(rdn_non_zero(rdnn(n8))),inference(pm,[status(thm)],[c_0_1311,c_0_1286,theory(equality)])).
cnf(c_0_1444,plain,(rdn_non_zero(rdnn(n9))),inference(pm,[status(thm)],[c_0_1311,c_0_1287,theory(equality)])).
cnf(c_0_1445,plain,(rdn_non_zero(rdnn(n1))),inference(pm,[status(thm)],[c_0_1311,c_0_1279,theory(equality)])).
cnf(c_0_1446,plain,(rdn_non_zero(rdnn(n3))),inference(pm,[status(thm)],[c_0_1311,c_0_1281,theory(equality)])).
cnf(c_0_1447,plain,(rdn_non_zero(rdnn(n4))),inference(pm,[status(thm)],[c_0_1311,c_0_1282,theory(equality)])).
cnf(c_0_1448,plain,(rdn_non_zero(rdnn(n7))),inference(pm,[status(thm)],[c_0_1311,c_0_1285,theory(equality)])).
cnf(c_0_1449,plain,(accept_leader(countryahumanitarianorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_776,theory(equality)])).
cnf(c_0_1450,plain,(accept_city(countryahumanitarianorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_776,theory(equality)])).
cnf(c_0_1451,plain,(accept_number(countryahumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_776,theory(equality)])).
cnf(c_0_1452,plain,(accept_leader(countryahumanitarianorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_775,theory(equality)])).
cnf(c_0_1453,plain,(accept_city(countryahumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_775,theory(equality)])).
cnf(c_0_1454,plain,(accept_number(countryahumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_775,theory(equality)])).
cnf(c_0_1455,plain,(accept_leader(countryahumanitarianorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_774,theory(equality)])).
cnf(c_0_1456,plain,(accept_city(countryahumanitarianorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_774,theory(equality)])).
cnf(c_0_1457,plain,(accept_number(countryahumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_774,theory(equality)])).
cnf(c_0_1458,plain,(accept_leader(countryahumanitarianorganization,countryafirstaidorganization)),inference(pm,[status(thm)],[c_0_607,c_0_777,theory(equality)])).
cnf(c_0_1459,plain,(accept_city(countryahumanitarianorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_777,theory(equality)])).
cnf(c_0_1460,plain,(accept_number(countryahumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_777,theory(equality)])).
cnf(c_0_1461,plain,(accept_leader(countryahumanitarianorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_607,c_0_963,theory(equality)])).
cnf(c_0_1462,plain,(accept_city(countryahumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_963,theory(equality)])).
cnf(c_0_1463,plain,(accept_number(countryahumanitarianorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_963,theory(equality)])).
cnf(c_0_1464,plain,(accept_leader(countryahumanitarianorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_607,c_0_964,theory(equality)])).
cnf(c_0_1465,plain,(accept_city(countryahumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_964,theory(equality)])).
cnf(c_0_1466,plain,(accept_number(countryahumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_964,theory(equality)])).
cnf(c_0_1467,plain,(accept_leader(countryacivilorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_833,theory(equality)])).
cnf(c_0_1468,plain,(accept_city(countryacivilorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_833,theory(equality)])).
cnf(c_0_1469,plain,(accept_number(countryacivilorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_833,theory(equality)])).
cnf(c_0_1470,plain,(accept_leader(countryacivilorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_834,theory(equality)])).
cnf(c_0_1471,plain,(accept_city(countryacivilorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_834,theory(equality)])).
cnf(c_0_1472,plain,(accept_number(countryacivilorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_834,theory(equality)])).
cnf(c_0_1473,plain,(accept_leader(countryacivilorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_940,theory(equality)])).
cnf(c_0_1474,plain,(accept_city(countryacivilorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_940,theory(equality)])).
cnf(c_0_1475,plain,(accept_number(countryacivilorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_940,theory(equality)])).
cnf(c_0_1476,plain,(accept_leader(countryacivilorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_879,theory(equality)])).
cnf(c_0_1477,plain,(accept_city(countryacivilorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_879,theory(equality)])).
cnf(c_0_1478,plain,(accept_number(countryacivilorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_879,theory(equality)])).
cnf(c_0_1479,plain,(accept_leader(countryacivilorganization,countryafirstaidorganization)),inference(pm,[status(thm)],[c_0_607,c_0_835,theory(equality)])).
cnf(c_0_1480,plain,(accept_city(countryacivilorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_835,theory(equality)])).
cnf(c_0_1481,plain,(accept_number(countryacivilorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_835,theory(equality)])).
cnf(c_0_1482,plain,(accept_leader(countryacivilorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_607,c_0_996,theory(equality)])).
cnf(c_0_1483,plain,(accept_city(countryacivilorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_996,theory(equality)])).
cnf(c_0_1484,plain,(accept_number(countryacivilorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_996,theory(equality)])).
cnf(c_0_1485,plain,(accept_leader(countryacivilorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_607,c_0_998,theory(equality)])).
cnf(c_0_1486,plain,(accept_city(countryacivilorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_998,theory(equality)])).
cnf(c_0_1487,plain,(accept_number(countryacivilorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_998,theory(equality)])).
cnf(c_0_1488,plain,(accept_leader(countryafirstaidorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_839,theory(equality)])).
cnf(c_0_1489,plain,(accept_city(countryafirstaidorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_839,theory(equality)])).
cnf(c_0_1490,plain,(accept_number(countryafirstaidorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_839,theory(equality)])).
cnf(c_0_1491,plain,(accept_leader(countryafirstaidorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_840,theory(equality)])).
cnf(c_0_1492,plain,(accept_city(countryafirstaidorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_840,theory(equality)])).
cnf(c_0_1493,plain,(accept_number(countryafirstaidorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_840,theory(equality)])).
cnf(c_0_1494,plain,(accept_leader(countryafirstaidorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_921,theory(equality)])).
cnf(c_0_1495,plain,(accept_city(countryafirstaidorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_921,theory(equality)])).
cnf(c_0_1496,plain,(accept_number(countryafirstaidorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_921,theory(equality)])).
cnf(c_0_1497,plain,(accept_leader(countryafirstaidorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_880,theory(equality)])).
cnf(c_0_1498,plain,(accept_city(countryafirstaidorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_880,theory(equality)])).
cnf(c_0_1499,plain,(accept_number(countryafirstaidorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_880,theory(equality)])).
cnf(c_0_1500,plain,(accept_leader(countryafirstaidorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_841,theory(equality)])).
cnf(c_0_1501,plain,(accept_city(countryafirstaidorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_841,theory(equality)])).
cnf(c_0_1502,plain,(accept_number(countryafirstaidorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_841,theory(equality)])).
cnf(c_0_1503,plain,(accept_leader(countryafirstaidorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_831,theory(equality)])).
cnf(c_0_1504,plain,(accept_city(countryafirstaidorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_831,theory(equality)])).
cnf(c_0_1505,plain,(accept_number(countryafirstaidorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_831,theory(equality)])).
cnf(c_0_1506,plain,(accept_leader(countryafirstaidorganization,countryacivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_789,theory(equality)])).
cnf(c_0_1507,plain,(accept_city(countryafirstaidorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_789,theory(equality)])).
cnf(c_0_1508,plain,(accept_number(countryafirstaidorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_789,theory(equality)])).
cnf(c_0_1509,plain,(accept_leader(countryafirstaidorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_607,c_0_925,theory(equality)])).
cnf(c_0_1510,plain,(accept_city(countryafirstaidorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_925,theory(equality)])).
cnf(c_0_1511,plain,(accept_number(countryafirstaidorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_925,theory(equality)])).
cnf(c_0_1512,plain,(accept_leader(countryafirstaidorganization,sufferterragovernment)),inference(pm,[status(thm)],[c_0_607,c_0_926,theory(equality)])).
cnf(c_0_1513,plain,(accept_city(countryafirstaidorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_926,theory(equality)])).
cnf(c_0_1514,plain,(accept_number(countryafirstaidorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_926,theory(equality)])).
cnf(c_0_1515,plain,(accept_leader(christiancountrychumanitarianorganization,countryccivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_1009,theory(equality)])).
cnf(c_0_1516,plain,(accept_city(christiancountrychumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_1009,theory(equality)])).
cnf(c_0_1517,plain,(accept_number(christiancountrychumanitarianorganization,n4)),inference(pm,[status(thm)],[c_0_606,c_0_1009,theory(equality)])).
cnf(c_0_1518,plain,(accept_leader(christiancountrychumanitarianorganization,countryccivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_881,theory(equality)])).
cnf(c_0_1519,plain,(accept_city(christiancountrychumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_881,theory(equality)])).
cnf(c_0_1520,plain,(accept_number(christiancountrychumanitarianorganization,n4)),inference(pm,[status(thm)],[c_0_606,c_0_881,theory(equality)])).
cnf(c_0_1521,plain,(accept_leader(christiancountrychumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_1013,theory(equality)])).
cnf(c_0_1522,plain,(accept_city(christiancountrychumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_1013,theory(equality)])).
cnf(c_0_1523,plain,(accept_number(christiancountrychumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_1013,theory(equality)])).
cnf(c_0_1524,plain,(accept_leader(christiancountrychumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_788,theory(equality)])).
cnf(c_0_1525,plain,(accept_city(christiancountrychumanitarianorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_788,theory(equality)])).
cnf(c_0_1526,plain,(accept_number(christiancountrychumanitarianorganization,n4)),inference(pm,[status(thm)],[c_0_606,c_0_788,theory(equality)])).
cnf(c_0_1527,plain,(accept_leader(christiancountrychumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_935,theory(equality)])).
cnf(c_0_1528,plain,(accept_city(christiancountrychumanitarianorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_935,theory(equality)])).
cnf(c_0_1529,plain,(accept_number(christiancountrychumanitarianorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_935,theory(equality)])).
cnf(c_0_1530,plain,(accept_leader(christiancountrychumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_936,theory(equality)])).
cnf(c_0_1531,plain,(accept_city(christiancountrychumanitarianorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_936,theory(equality)])).
cnf(c_0_1532,plain,(accept_number(christiancountrychumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_936,theory(equality)])).
cnf(c_0_1533,plain,(accept_leader(christiancountrychumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_884,theory(equality)])).
cnf(c_0_1534,plain,(accept_city(christiancountrychumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_884,theory(equality)])).
cnf(c_0_1535,plain,(accept_number(christiancountrychumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_884,theory(equality)])).
cnf(c_0_1536,plain,(accept_leader(christiancountrychumanitarianorganization,muslimcountrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_882,theory(equality)])).
cnf(c_0_1537,plain,(accept_city(christiancountrychumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_882,theory(equality)])).
cnf(c_0_1538,plain,(accept_number(christiancountrychumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_882,theory(equality)])).
cnf(c_0_1539,plain,(accept_leader(christiancountrychumanitarianorganization,countrycmedicalorganization)),inference(pm,[status(thm)],[c_0_607,c_0_826,theory(equality)])).
cnf(c_0_1540,plain,(accept_city(christiancountrychumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_826,theory(equality)])).
cnf(c_0_1541,plain,(accept_number(christiancountrychumanitarianorganization,n4)),inference(pm,[status(thm)],[c_0_606,c_0_826,theory(equality)])).
cnf(c_0_1542,plain,(accept_leader(christiancountrychumanitarianorganization,countrycmedicalorganization)),inference(pm,[status(thm)],[c_0_607,c_0_828,theory(equality)])).
cnf(c_0_1543,plain,(accept_city(christiancountrychumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_828,theory(equality)])).
cnf(c_0_1544,plain,(accept_number(christiancountrychumanitarianorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_828,theory(equality)])).
cnf(c_0_1545,plain,(accept_leader(christiancountrychumanitarianorganization,countrycmedicalorganization)),inference(pm,[status(thm)],[c_0_607,c_0_1014,theory(equality)])).
cnf(c_0_1546,plain,(accept_city(christiancountrychumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_1014,theory(equality)])).
cnf(c_0_1547,plain,(accept_number(christiancountrychumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_1014,theory(equality)])).
cnf(c_0_1548,plain,(accept_leader(christiancountrychumanitarianorganization,countrycmedicalorganization)),inference(pm,[status(thm)],[c_0_607,c_0_885,theory(equality)])).
cnf(c_0_1549,plain,(accept_city(christiancountrychumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_885,theory(equality)])).
cnf(c_0_1550,plain,(accept_number(christiancountrychumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_885,theory(equality)])).
cnf(c_0_1551,plain,(accept_leader(christiancountrychumanitarianorganization,countrybcivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_1011,theory(equality)])).
cnf(c_0_1552,plain,(accept_city(christiancountrychumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_1011,theory(equality)])).
cnf(c_0_1553,plain,(accept_number(christiancountrychumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_1011,theory(equality)])).
cnf(c_0_1554,plain,(accept_leader(christiancountrychumanitarianorganization,countrybcivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_883,theory(equality)])).
cnf(c_0_1555,plain,(accept_city(christiancountrychumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_883,theory(equality)])).
cnf(c_0_1556,plain,(accept_number(christiancountrychumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_883,theory(equality)])).
cnf(c_0_1557,plain,(accept_leader(countryccivilorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_780,theory(equality)])).
cnf(c_0_1558,plain,(accept_city(countryccivilorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_780,theory(equality)])).
cnf(c_0_1559,plain,(accept_number(countryccivilorganization,n2)),inference(pm,[status(thm)],[c_0_606,c_0_780,theory(equality)])).
cnf(c_0_1560,plain,(accept_leader(countryccivilorganization,countrycmedicalorganization)),inference(pm,[status(thm)],[c_0_607,c_0_812,theory(equality)])).
cnf(c_0_1561,plain,(accept_city(countryccivilorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_812,theory(equality)])).
cnf(c_0_1562,plain,(accept_number(countryccivilorganization,n4)),inference(pm,[status(thm)],[c_0_606,c_0_812,theory(equality)])).
cnf(c_0_1563,plain,(accept_leader(countrybhumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_797,theory(equality)])).
cnf(c_0_1564,plain,(accept_city(countrybhumanitarianorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_797,theory(equality)])).
cnf(c_0_1565,plain,(accept_number(countrybhumanitarianorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_797,theory(equality)])).
cnf(c_0_1566,plain,(accept_leader(countrybhumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_798,theory(equality)])).
cnf(c_0_1567,plain,(accept_city(countrybhumanitarianorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_798,theory(equality)])).
cnf(c_0_1568,plain,(accept_number(countrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_798,theory(equality)])).
cnf(c_0_1569,plain,(accept_leader(countrybhumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_967,theory(equality)])).
cnf(c_0_1570,plain,(accept_city(countrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_967,theory(equality)])).
cnf(c_0_1571,plain,(accept_number(countrybhumanitarianorganization,n3)),inference(pm,[status(thm)],[c_0_606,c_0_967,theory(equality)])).
cnf(c_0_1572,plain,(accept_leader(countrybhumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_972,theory(equality)])).
cnf(c_0_1573,plain,(accept_city(countrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_972,theory(equality)])).
cnf(c_0_1574,plain,(accept_number(countrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_972,theory(equality)])).
cnf(c_0_1575,plain,(accept_leader(countrybhumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_908,theory(equality)])).
cnf(c_0_1576,plain,(accept_city(countrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_908,theory(equality)])).
cnf(c_0_1577,plain,(accept_number(countrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_908,theory(equality)])).
cnf(c_0_1578,plain,(accept_leader(countrybhumanitarianorganization,countryccivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_969,theory(equality)])).
cnf(c_0_1579,plain,(accept_city(countrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_969,theory(equality)])).
cnf(c_0_1580,plain,(accept_number(countrybhumanitarianorganization,n3)),inference(pm,[status(thm)],[c_0_606,c_0_969,theory(equality)])).
cnf(c_0_1581,plain,(accept_leader(countrybhumanitarianorganization,muslimcountrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_896,theory(equality)])).
cnf(c_0_1582,plain,(accept_city(countrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_896,theory(equality)])).
cnf(c_0_1583,plain,(accept_number(countrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_896,theory(equality)])).
cnf(c_0_1584,plain,(accept_leader(countrybhumanitarianorganization,countrycmedicalorganization)),inference(pm,[status(thm)],[c_0_607,c_0_968,theory(equality)])).
cnf(c_0_1585,plain,(accept_city(countrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_968,theory(equality)])).
cnf(c_0_1586,plain,(accept_number(countrybhumanitarianorganization,n3)),inference(pm,[status(thm)],[c_0_606,c_0_968,theory(equality)])).
cnf(c_0_1587,plain,(accept_leader(countrybhumanitarianorganization,countrycmedicalorganization)),inference(pm,[status(thm)],[c_0_607,c_0_973,theory(equality)])).
cnf(c_0_1588,plain,(accept_city(countrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_973,theory(equality)])).
cnf(c_0_1589,plain,(accept_number(countrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_973,theory(equality)])).
cnf(c_0_1590,plain,(accept_leader(countrybhumanitarianorganization,countrycmedicalorganization)),inference(pm,[status(thm)],[c_0_607,c_0_909,theory(equality)])).
cnf(c_0_1591,plain,(accept_city(countrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_909,theory(equality)])).
cnf(c_0_1592,plain,(accept_number(countrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_909,theory(equality)])).
cnf(c_0_1593,plain,(accept_leader(countrybhumanitarianorganization,countrybcivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_966,theory(equality)])).
cnf(c_0_1594,plain,(accept_city(countrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_966,theory(equality)])).
cnf(c_0_1595,plain,(accept_number(countrybhumanitarianorganization,n3)),inference(pm,[status(thm)],[c_0_606,c_0_966,theory(equality)])).
cnf(c_0_1596,plain,(accept_leader(countrybhumanitarianorganization,countrybcivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_971,theory(equality)])).
cnf(c_0_1597,plain,(accept_city(countrybhumanitarianorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_971,theory(equality)])).
cnf(c_0_1598,plain,(accept_number(countrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_971,theory(equality)])).
cnf(c_0_1599,plain,(accept_leader(countrybhumanitarianorganization,countrybcivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_895,theory(equality)])).
cnf(c_0_1600,plain,(accept_city(countrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_895,theory(equality)])).
cnf(c_0_1601,plain,(accept_number(countrybhumanitarianorganization,n2)),inference(pm,[status(thm)],[c_0_606,c_0_895,theory(equality)])).
cnf(c_0_1602,plain,(accept_leader(countrybhumanitarianorganization,countrybcivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_907,theory(equality)])).
cnf(c_0_1603,plain,(accept_city(countrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_907,theory(equality)])).
cnf(c_0_1604,plain,(accept_number(countrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_907,theory(equality)])).
cnf(c_0_1605,plain,(accept_leader(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_786,theory(equality)])).
cnf(c_0_1606,plain,(accept_city(muslimcountrybhumanitarianorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_786,theory(equality)])).
cnf(c_0_1607,plain,(accept_number(muslimcountrybhumanitarianorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_786,theory(equality)])).
cnf(c_0_1608,plain,(accept_leader(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_870,theory(equality)])).
cnf(c_0_1609,plain,(accept_city(muslimcountrybhumanitarianorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_870,theory(equality)])).
cnf(c_0_1610,plain,(accept_number(muslimcountrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_870,theory(equality)])).
cnf(c_0_1611,plain,(accept_leader(muslimcountrybhumanitarianorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_875,theory(equality)])).
cnf(c_0_1612,plain,(accept_city(muslimcountrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_875,theory(equality)])).
cnf(c_0_1613,plain,(accept_number(muslimcountrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_875,theory(equality)])).
cnf(c_0_1614,plain,(accept_leader(muslimcountrybhumanitarianorganization,countryccivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_872,theory(equality)])).
cnf(c_0_1615,plain,(accept_city(muslimcountrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_872,theory(equality)])).
cnf(c_0_1616,plain,(accept_number(muslimcountrybhumanitarianorganization,n4)),inference(pm,[status(thm)],[c_0_606,c_0_872,theory(equality)])).
cnf(c_0_1617,plain,(accept_leader(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_781,theory(equality)])).
cnf(c_0_1618,plain,(accept_city(muslimcountrybhumanitarianorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_781,theory(equality)])).
cnf(c_0_1619,plain,(accept_number(muslimcountrybhumanitarianorganization,n4)),inference(pm,[status(thm)],[c_0_606,c_0_781,theory(equality)])).
cnf(c_0_1620,plain,(accept_leader(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_945,theory(equality)])).
cnf(c_0_1621,plain,(accept_city(muslimcountrybhumanitarianorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_945,theory(equality)])).
cnf(c_0_1622,plain,(accept_number(muslimcountrybhumanitarianorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_945,theory(equality)])).
cnf(c_0_1623,plain,(accept_leader(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_946,theory(equality)])).
cnf(c_0_1624,plain,(accept_city(muslimcountrybhumanitarianorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_946,theory(equality)])).
cnf(c_0_1625,plain,(accept_number(muslimcountrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_946,theory(equality)])).
cnf(c_0_1626,plain,(accept_leader(muslimcountrybhumanitarianorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_874,theory(equality)])).
cnf(c_0_1627,plain,(accept_city(muslimcountrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_874,theory(equality)])).
cnf(c_0_1628,plain,(accept_number(muslimcountrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_874,theory(equality)])).
cnf(c_0_1629,plain,(accept_leader(muslimcountrybhumanitarianorganization,countrycmedicalorganization)),inference(pm,[status(thm)],[c_0_607,c_0_876,theory(equality)])).
cnf(c_0_1630,plain,(accept_city(muslimcountrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_876,theory(equality)])).
cnf(c_0_1631,plain,(accept_number(muslimcountrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_876,theory(equality)])).
cnf(c_0_1632,plain,(accept_leader(muslimcountrybhumanitarianorganization,countrybcivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_873,theory(equality)])).
cnf(c_0_1633,plain,(accept_city(muslimcountrybhumanitarianorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_873,theory(equality)])).
cnf(c_0_1634,plain,(accept_number(muslimcountrybhumanitarianorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_873,theory(equality)])).
cnf(c_0_1635,plain,(accept_leader(countrycmedicalorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_824,theory(equality)])).
cnf(c_0_1636,plain,(accept_city(countrycmedicalorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_824,theory(equality)])).
cnf(c_0_1637,plain,(accept_number(countrycmedicalorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_824,theory(equality)])).
cnf(c_0_1638,plain,(accept_leader(countrycmedicalorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_869,theory(equality)])).
cnf(c_0_1639,plain,(accept_city(countrycmedicalorganization,coastvillage)),inference(pm,[status(thm)],[c_0_608,c_0_869,theory(equality)])).
cnf(c_0_1640,plain,(accept_number(countrycmedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_869,theory(equality)])).
cnf(c_0_1641,plain,(accept_leader(countrycmedicalorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_1004,theory(equality)])).
cnf(c_0_1642,plain,(accept_city(countrycmedicalorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_1004,theory(equality)])).
cnf(c_0_1643,plain,(accept_number(countrycmedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_1004,theory(equality)])).
cnf(c_0_1644,plain,(accept_leader(countrycmedicalorganization,christiancountrychumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_904,theory(equality)])).
cnf(c_0_1645,plain,(accept_city(countrycmedicalorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_904,theory(equality)])).
cnf(c_0_1646,plain,(accept_number(countrycmedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_904,theory(equality)])).
cnf(c_0_1647,plain,(accept_leader(countrycmedicalorganization,countryccivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_1001,theory(equality)])).
cnf(c_0_1648,plain,(accept_city(countrycmedicalorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_1001,theory(equality)])).
cnf(c_0_1649,plain,(accept_number(countrycmedicalorganization,n4)),inference(pm,[status(thm)],[c_0_606,c_0_1001,theory(equality)])).
cnf(c_0_1650,plain,(accept_leader(countrycmedicalorganization,countryccivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_900,theory(equality)])).
cnf(c_0_1651,plain,(accept_city(countrycmedicalorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_900,theory(equality)])).
cnf(c_0_1652,plain,(accept_number(countrycmedicalorganization,n4)),inference(pm,[status(thm)],[c_0_606,c_0_900,theory(equality)])).
cnf(c_0_1653,plain,(accept_leader(countrycmedicalorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_1003,theory(equality)])).
cnf(c_0_1654,plain,(accept_city(countrycmedicalorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_1003,theory(equality)])).
cnf(c_0_1655,plain,(accept_number(countrycmedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_1003,theory(equality)])).
cnf(c_0_1656,plain,(accept_leader(countrycmedicalorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_927,theory(equality)])).
cnf(c_0_1657,plain,(accept_city(countrycmedicalorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_927,theory(equality)])).
cnf(c_0_1658,plain,(accept_number(countrycmedicalorganization,n5)),inference(pm,[status(thm)],[c_0_606,c_0_927,theory(equality)])).
cnf(c_0_1659,plain,(accept_leader(countrycmedicalorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_928,theory(equality)])).
cnf(c_0_1660,plain,(accept_city(countrycmedicalorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_928,theory(equality)])).
cnf(c_0_1661,plain,(accept_number(countrycmedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_928,theory(equality)])).
cnf(c_0_1662,plain,(accept_leader(countrycmedicalorganization,countrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_903,theory(equality)])).
cnf(c_0_1663,plain,(accept_city(countrycmedicalorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_903,theory(equality)])).
cnf(c_0_1664,plain,(accept_number(countrycmedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_903,theory(equality)])).
cnf(c_0_1665,plain,(accept_leader(countrycmedicalorganization,muslimcountrybhumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_901,theory(equality)])).
cnf(c_0_1666,plain,(accept_city(countrycmedicalorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_901,theory(equality)])).
cnf(c_0_1667,plain,(accept_number(countrycmedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_901,theory(equality)])).
cnf(c_0_1668,plain,(accept_leader(countrycmedicalorganization,countrybcivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_1002,theory(equality)])).
cnf(c_0_1669,plain,(accept_city(countrycmedicalorganization,towna)),inference(pm,[status(thm)],[c_0_608,c_0_1002,theory(equality)])).
cnf(c_0_1670,plain,(accept_number(countrycmedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_1002,theory(equality)])).
cnf(c_0_1671,plain,(accept_leader(countrycmedicalorganization,countrybcivilorganization)),inference(pm,[status(thm)],[c_0_607,c_0_902,theory(equality)])).
cnf(c_0_1672,plain,(accept_city(countrycmedicalorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_902,theory(equality)])).
cnf(c_0_1673,plain,(accept_number(countrycmedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_902,theory(equality)])).
cnf(c_0_1674,plain,(accept_leader(countryamedicalorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_929,theory(equality)])).
cnf(c_0_1675,plain,(accept_city(countryamedicalorganization,cityb)),inference(pm,[status(thm)],[c_0_608,c_0_929,theory(equality)])).
cnf(c_0_1676,plain,(accept_number(countryamedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_929,theory(equality)])).
cnf(c_0_1677,plain,(accept_leader(countryamedicalorganization,countryahumanitarianorganization)),inference(pm,[status(thm)],[c_0_607,c_0_899,theory(equality)])).
cnf(c_0_1678,plain,(accept_city(countryamedicalorganization,townc)),inference(pm,[status(thm)],[c_0_608,c_0_899,theory(equality)])).
cnf(c_0_1679,plain,(accept_number(countryamedicalorganization,n6)),inference(pm,[status(thm)],[c_0_606,c_0_899,theory(equality)])).
cnf(c_0_1680,plain,($false),inference(sr,[status(thm)],[c_0_800,c_0_1349,theory(equality)])).
cnf(c_0_1681,plain,($false), c_0_1680,['proof']).

