% Problem    : BOO007-2 : TPTP v6.1.0. Released v1.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n100.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:13:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C02_F1_PI_AE_Q12_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(multiply(X1,add(X2,X3))=add(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', distributivity4)).
cnf(c_0_1,axiom,(multiply(X1,inverse(X1))=additive_identity), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', multiplicative_inverse1)).
cnf(c_0_2,axiom,(add(X1,additive_identity)=X1), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', additive_id1)).
cnf(c_0_3,axiom,(add(additive_identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', additive_id2)).
cnf(c_0_4,axiom,(add(X1,multiply(X2,X3))=multiply(add(X1,X2),add(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', distributivity2)).
cnf(c_0_5,axiom,(add(X1,inverse(X1))=multiplicative_identity), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', additive_inverse1)).
cnf(c_0_6,axiom,(multiply(X1,multiplicative_identity)=X1), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', multiplicative_id1)).
cnf(c_0_7,axiom,(multiply(X1,X2)=multiply(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', commutativity_of_multiply)).
cnf(c_0_8,axiom,(add(X1,X2)=add(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', commutativity_of_add)).
cnf(c_0_9,axiom,(multiply(multiplicative_identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/BOO003-0.ax', multiplicative_id2)).
cnf(c_0_10,negated_conjecture,(multiply(a,multiply(b,c))!=multiply(multiply(a,b),c)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_associativity)).
cnf(c_0_11,plain,(multiply(X1,add(X2,inverse(X1)))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_12,plain,(multiply(X1,additive_identity)=additive_identity), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_3]), c_0_1])).
cnf(c_0_13,plain,(add(X1,multiply(X2,inverse(X1)))=add(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_5]), c_0_6])).
cnf(c_0_14,plain,(multiply(X1,add(X1,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_2]), c_0_7]), c_0_12]), c_0_2])).
cnf(c_0_15,plain,(add(X1,X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_1]), c_0_2])).
cnf(c_0_16,plain,(multiply(X1,add(X2,X1))=X1), inference(spm,[status(thm)],[c_0_14, c_0_8])).
cnf(c_0_17,plain,(add(X1,multiply(X1,X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_15]), c_0_14])).
cnf(c_0_18,plain,(multiply(X1,add(inverse(X1),X2))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_3])).
cnf(c_0_19,plain,(multiply(X1,multiply(X1,X2))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_7])).
cnf(c_0_20,plain,(add(X1,multiply(X2,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_15]), c_0_7]), c_0_14])).
cnf(c_0_21,plain,(multiply(X1,inverse(inverse(X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_5]), c_0_6])).
cnf(c_0_22,plain,(multiply(X1,add(multiply(X1,X2),X3))=multiply(X1,add(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_19]), c_0_0])).
cnf(c_0_23,plain,(add(multiply(X1,X2),multiply(X2,X3))=multiply(X2,add(X1,X3))), inference(spm,[status(thm)],[c_0_0, c_0_7])).
cnf(c_0_24,plain,(add(X1,multiply(inverse(X1),X2))=add(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_5]), c_0_9])).
cnf(c_0_25,plain,(add(X1,inverse(inverse(X1)))=inverse(inverse(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_8])).
cnf(c_0_26,plain,(add(X1,multiply(multiply(X2,X1),X3))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_20]), c_0_14])).
cnf(c_0_27,plain,(multiply(X1,multiply(X2,add(X1,X3)))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_17])).
cnf(c_0_28,plain,(add(X1,add(X2,X1))=add(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_16]), c_0_8])).
cnf(c_0_29,plain,(add(X1,multiply(multiply(X2,inverse(X1)),X3))=add(X1,multiply(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_13]), c_0_4])).
cnf(c_0_30,plain,(inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_1]), c_0_2]), c_0_25])).
cnf(c_0_31,plain,(multiply(X1,multiply(multiply(X2,X1),X3))=multiply(multiply(X2,X1),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_26]), c_0_7])).
cnf(c_0_32,plain,(add(X1,multiply(X2,multiply(X3,X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_20]), c_0_7]), c_0_14])).
cnf(c_0_33,plain,(multiply(X1,multiply(X2,add(X3,X1)))=multiply(X1,X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34,plain,(multiply(multiply(X1,X2),X3)=multiply(X2,multiply(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_29]), c_0_18]), c_0_30]), c_0_31])).
cnf(c_0_35,plain,(multiply(X1,multiply(X2,multiply(X3,X1)))=multiply(X2,multiply(X3,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_32]), c_0_7])).
cnf(c_0_36,negated_conjecture,(multiply(c,multiply(a,b))!=multiply(a,multiply(b,c))), inference(rw,[status(thm)],[c_0_10, c_0_7])).
cnf(c_0_37,plain,(multiply(X1,multiply(X2,X3))=multiply(X3,multiply(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_20]), c_0_34]), c_0_35]), c_0_34])).
cnf(c_0_38,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_37])]), ['proof']).
# SZS output end CNFRefutation.
EOF
