% Problem    : COL006-6 : TPTP v6.1.0. Released v2.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n134.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:21:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___213_C18_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___107_B42_F1_PI_SE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(strong_fixed_point=apply(apply(s,apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', strong_fixed_point)).
cnf(c_0_1,axiom,(apply(apply(apply(s,X1),X2),X3)=apply(apply(X1,X3),apply(X2,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', s_definition)).
cnf(c_0_2,axiom,(apply(apply(k,X1),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', k_definition)).
cnf(c_0_3,negated_conjecture,(apply(strong_fixed_point,fixed_pt)!=apply(fixed_pt,apply(strong_fixed_point,fixed_pt))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_strong_fixed_point)).
cnf(c_0_4,plain,(strong_fixed_point=apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k))))), inference(rw,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_5,plain,(apply(apply(apply(s,k),X1),X2)=X2), inference(spm,[status(thm)],[c_0_2, c_0_1])).
cnf(c_0_6,negated_conjecture,(apply(fixed_pt,apply(apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt))!=apply(apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_4]), c_0_4])).
cnf(c_0_7,plain,(apply(apply(apply(s,apply(apply(s,k),X1)),X2),X3)=apply(X3,apply(X2,X3))), inference(spm,[status(thm)],[c_0_1, c_0_5])).
cnf(c_0_8,plain,(apply(apply(apply(s,apply(k,X1)),X2),X3)=apply(X1,apply(X2,X3))), inference(spm,[status(thm)],[c_0_1, c_0_2])).
cnf(c_0_9,negated_conjecture,(apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k))))),fixed_pt)!=apply(apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10,plain,(apply(apply(apply(apply(s,s),X1),apply(k,X2)),X3)=apply(X2,apply(apply(X1,apply(k,X2)),X3))), inference(spm,[status(thm)],[c_0_8, c_0_1])).
cnf(c_0_11,negated_conjecture,(apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k))))),fixed_pt)!=apply(apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt))), inference(rw,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12,negated_conjecture,(apply(fixed_pt,apply(apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt)))!=apply(apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_1]), c_0_5]), c_0_10])).
cnf(c_0_13,plain,(apply(apply(apply(s,apply(apply(s,k),X1)),X2),X3)=apply(apply(apply(s,apply(apply(s,k),X4)),X2),X3)), inference(spm,[status(thm)],[c_0_7, c_0_7])).
cnf(c_0_14,plain,(apply(apply(apply(s,apply(X1,X2)),X3),apply(X4,X2))=apply(apply(apply(apply(s,X1),X4),X2),apply(X3,apply(X4,X2)))), inference(spm,[status(thm)],[c_0_1, c_0_1])).
cnf(c_0_15,negated_conjecture,(apply(fixed_pt,apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt)))!=apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16,plain,(apply(apply(apply(s,apply(X1,apply(X2,X3))),X4),apply(apply(apply(s,X5),X2),X3))=apply(apply(apply(apply(s,X1),apply(X5,X3)),apply(X2,X3)),apply(X4,apply(apply(apply(s,X5),X2),X3)))), inference(spm,[status(thm)],[c_0_14, c_0_1])).
cnf(c_0_17,plain,(apply(apply(apply(s,apply(s,k)),X1),X2)=apply(X1,X2)), inference(spm,[status(thm)],[c_0_1, c_0_5])).
cnf(c_0_18,negated_conjecture,(apply(fixed_pt,apply(fixed_pt,apply(apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt))))!=apply(fixed_pt,apply(apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_8]), c_0_2]), c_0_17]), c_0_5]), c_0_8]), c_0_8]), c_0_2]), c_0_17]), c_0_5]), c_0_8])).
cnf(c_0_19,plain,(apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(s,apply(k,X2)),X3)),X4)=apply(X4,apply(X2,apply(X3,X4)))), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_20,negated_conjecture,(apply(fixed_pt,apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k))))),fixed_pt))!=apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k))))),fixed_pt)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_1]), c_0_1])).
cnf(c_0_21,negated_conjecture,(apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(s,apply(apply(s,k),X2)),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))))),fixed_pt)!=apply(apply(apply(s,apply(apply(s,k),X2)),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k))))),fixed_pt)), inference(spm,[status(thm)],[c_0_20, c_0_7])).
cnf(c_0_22,plain,(apply(apply(apply(s,X1),apply(apply(s,apply(apply(s,k),X2)),X3)),X4)=apply(apply(X1,X4),apply(X4,apply(X3,X4)))), inference(spm,[status(thm)],[c_0_1, c_0_7])).
cnf(c_0_23,plain,(apply(apply(apply(s,apply(s,apply(s,k))),X1),X2)=apply(X2,apply(X1,X2))), inference(spm,[status(thm)],[c_0_1, c_0_17])).
cnf(c_0_24,negated_conjecture,(apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k))))),fixed_pt)!=apply(apply(apply(s,apply(s,apply(s,k))),apply(apply(s,apply(s,apply(s,k))),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))))),fixed_pt)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_5]), c_0_23]), c_0_23])).
cnf(c_0_25,negated_conjecture,(apply(fixed_pt,apply(apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt)))!=apply(apply(apply(s,apply(s,apply(s,k))),apply(apply(s,apply(s,apply(s,k))),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))))),fixed_pt)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_1]), c_0_5]), c_0_10])).
cnf(c_0_26,negated_conjecture,(apply(fixed_pt,apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt)))!=apply(apply(apply(s,apply(s,apply(s,k))),apply(apply(s,apply(s,apply(s,k))),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))))),fixed_pt)), inference(spm,[status(thm)],[c_0_25, c_0_13])).
cnf(c_0_27,negated_conjecture,(apply(fixed_pt,apply(fixed_pt,apply(apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)),apply(apply(apply(s,apply(apply(s,apply(k,s)),k)),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))),fixed_pt))))!=apply(apply(apply(s,apply(s,apply(s,k))),apply(apply(s,apply(s,apply(s,k))),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))))),fixed_pt)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_16]), c_0_8]), c_0_2]), c_0_17]), c_0_5]), c_0_8])).
cnf(c_0_28,negated_conjecture,(apply(fixed_pt,apply(apply(apply(s,apply(apply(s,k),X1)),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k))))),fixed_pt))!=apply(apply(apply(s,apply(s,apply(s,k))),apply(apply(s,apply(s,apply(s,k))),apply(apply(apply(s,s),apply(s,apply(apply(s,apply(k,s)),k))),apply(k,apply(apply(s,apply(apply(s,k),k)),apply(apply(s,k),k)))))),fixed_pt)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_19]), c_0_1])).
cnf(c_0_29,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_7]), c_0_23]), c_0_23])]), ['proof']).
# SZS output end CNFRefutation.
EOF
