% Problem    : GRP497-1 : TPTP v6.1.0. Released v2.6.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n180.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:07:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(identity=double_divide(X1,inverse(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', identity)).
cnf(c_0_1,axiom,(inverse(X1)=double_divide(X1,identity)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', inverse)).
cnf(c_0_2,axiom,(double_divide(double_divide(identity,double_divide(X1,double_divide(X2,identity))),double_divide(double_divide(X2,double_divide(X3,X1)),identity))=X3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', single_axiom)).
cnf(c_0_3,negated_conjecture,(multiply(identity,a2)!=a2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_these_axioms_2)).
cnf(c_0_4,axiom,(multiply(X1,X2)=double_divide(double_divide(X2,X1),identity)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', multiply)).
cnf(c_0_5,plain,(double_divide(X1,double_divide(X1,identity))=identity), inference(rw,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_6,plain,(double_divide(identity,identity)=identity), inference(spm,[status(thm)],[c_0_5, c_0_2])).
cnf(c_0_7,plain,(double_divide(double_divide(identity,double_divide(identity,double_divide(X1,identity))),identity)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_5]), c_0_6])).
cnf(c_0_8,plain,(double_divide(double_divide(identity,double_divide(identity,double_divide(X1,identity))),X1)=identity), inference(spm,[status(thm)],[c_0_5, c_0_7])).
cnf(c_0_9,plain,(double_divide(double_divide(identity,double_divide(X1,double_divide(X2,X1))),identity)=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_8]), c_0_7]), c_0_6])).
cnf(c_0_10,plain,(double_divide(double_divide(identity,double_divide(double_divide(X1,X2),identity)),X1)=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_9]), c_0_6])).
cnf(c_0_11,plain,(double_divide(identity,X1)=double_divide(X1,identity)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_5]), c_0_6]), c_0_6])).
cnf(c_0_12,negated_conjecture,(double_divide(double_divide(a2,identity),identity)!=a2), inference(rw,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_13,plain,(double_divide(double_divide(identity,double_divide(identity,double_divide(X1,X2))),X1)=X2), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14,plain,(double_divide(identity,double_divide(identity,double_divide(X1,double_divide(X2,X1))))=X2), inference(rw,[status(thm)],[c_0_9, c_0_11])).
cnf(c_0_15,negated_conjecture,(double_divide(identity,double_divide(identity,a2))!=a2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_11]), c_0_11])).
cnf(c_0_16,plain,(double_divide(identity,double_divide(identity,X1))=double_divide(X2,double_divide(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_11])).
cnf(c_0_17,plain,(double_divide(double_divide(identity,double_divide(double_divide(double_divide(X1,double_divide(X2,X3)),identity),double_divide(X4,identity))),double_divide(double_divide(X4,X2),identity))=double_divide(identity,double_divide(X3,double_divide(X1,identity)))), inference(spm,[status(thm)],[c_0_2, c_0_2])).
cnf(c_0_18,negated_conjecture,(double_divide(X1,double_divide(a2,X1))!=a2), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(double_divide(double_divide(identity,double_divide(X1,double_divide(X2,X1))),X2)=identity), inference(spm,[status(thm)],[c_0_5, c_0_9])).
cnf(c_0_20,plain,(double_divide(double_divide(identity,double_divide(double_divide(identity,double_divide(X1,double_divide(X2,X3))),double_divide(X4,identity))),double_divide(identity,double_divide(X4,X2)))=double_divide(identity,double_divide(X3,double_divide(X1,identity)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_11]), c_0_11])).
cnf(c_0_21,negated_conjecture,(double_divide(double_divide(X1,a2),double_divide(identity,double_divide(identity,X1)))!=a2), inference(spm,[status(thm)],[c_0_18, c_0_16])).
cnf(c_0_22,plain,(double_divide(double_divide(identity,double_divide(double_divide(identity,double_divide(X1,X2)),double_divide(identity,double_divide(X3,double_divide(X4,identity))))),double_divide(identity,double_divide(double_divide(identity,double_divide(X4,double_divide(X2,X3))),double_divide(X1,identity))))=identity), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_23,plain,(double_divide(double_divide(identity,double_divide(identity,X1)),double_divide(identity,double_divide(identity,double_divide(X2,X1))))=X2), inference(spm,[status(thm)],[c_0_13, c_0_13])).
cnf(c_0_24,negated_conjecture,(double_divide(double_divide(X1,a2),double_divide(identity,double_divide(X1,identity)))!=a2), inference(spm,[status(thm)],[c_0_21, c_0_11])).
cnf(c_0_25,plain,(double_divide(double_divide(identity,X1),X1)=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_6]), c_0_6]), c_0_11]), c_0_14])).
cnf(c_0_26,plain,(double_divide(identity,double_divide(identity,double_divide(identity,double_divide(identity,X1))))=X1), inference(spm,[status(thm)],[c_0_11, c_0_13])).
cnf(c_0_27,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_11]), c_0_26])]), ['proof']).
# SZS output end CNFRefutation.
EOF
