% Problem    : GRP181-3 : TPTP v6.1.0. Bugfixed v1.2.1.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n185.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:15:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C02_F1_PI_AE_Q12_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', associativity)).
cnf(c_0_1,axiom,(multiply(inverse(X1),X1)=identity), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_inverse)).
cnf(c_0_2,axiom,(multiply(identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_identity)).
cnf(c_0_3,axiom,(multiply(least_upper_bound(X1,X2),X3)=least_upper_bound(multiply(X1,X3),multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_lub2)).
cnf(c_0_4,axiom,(least_upper_bound(X1,X2)=least_upper_bound(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', symmetry_of_lub)).
cnf(c_0_5,hypothesis,(least_upper_bound(a,c)=least_upper_bound(b,c)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p12x_2)).
cnf(c_0_6,axiom,(multiply(X1,greatest_lower_bound(X2,X3))=greatest_lower_bound(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_glb1)).
cnf(c_0_7,hypothesis,(inverse(least_upper_bound(X1,X2))=greatest_lower_bound(inverse(X1),inverse(X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p12x_4)).
cnf(c_0_8,axiom,(greatest_lower_bound(X1,X2)=greatest_lower_bound(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', symmetry_of_glb)).
cnf(c_0_9,axiom,(multiply(X1,least_upper_bound(X2,X3))=least_upper_bound(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_lub1)).
cnf(c_0_10,hypothesis,(inverse(greatest_lower_bound(X1,X2))=least_upper_bound(inverse(X1),inverse(X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p12x_3)).
cnf(c_0_11,hypothesis,(greatest_lower_bound(a,c)=greatest_lower_bound(b,c)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p12x_1)).
cnf(c_0_12,negated_conjecture,(a!=b), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_p12x)).
cnf(c_0_13,plain,(multiply(inverse(X1),multiply(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_14,plain,(multiply(inverse(inverse(X1)),identity)=X1), inference(spm,[status(thm)],[c_0_13, c_0_1])).
cnf(c_0_15,plain,(multiply(inverse(inverse(X1)),X2)=multiply(X1,X2)), inference(spm,[status(thm)],[c_0_13, c_0_13])).
cnf(c_0_16,plain,(multiply(X1,identity)=X1), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_17,plain,(inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_16])).
cnf(c_0_18,plain,(multiply(inverse(identity),X1)=X1), inference(spm,[status(thm)],[c_0_13, c_0_2])).
cnf(c_0_19,plain,(multiply(X1,inverse(X1))=identity), inference(spm,[status(thm)],[c_0_1, c_0_17])).
cnf(c_0_20,plain,(multiply(inverse(inverse(identity)),X1)=X1), inference(spm,[status(thm)],[c_0_13, c_0_18])).
cnf(c_0_21,plain,(multiply(X1,multiply(X2,inverse(multiply(X1,X2))))=identity), inference(spm,[status(thm)],[c_0_0, c_0_19])).
cnf(c_0_22,plain,(inverse(identity)=identity), inference(spm,[status(thm)],[c_0_1, c_0_20])).
cnf(c_0_23,plain,(multiply(X1,multiply(inverse(X1),X2))=X2), inference(spm,[status(thm)],[c_0_13, c_0_17])).
cnf(c_0_24,plain,(multiply(X1,inverse(multiply(X2,X1)))=inverse(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_21]), c_0_16])).
cnf(c_0_25,plain,(multiply(least_upper_bound(X1,X2),inverse(X2))=least_upper_bound(identity,multiply(X1,inverse(X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_3, c_0_19]), c_0_4])).
cnf(c_0_26,hypothesis,(least_upper_bound(c,a)=least_upper_bound(b,c)), inference(rw,[status(thm)],[c_0_5, c_0_4])).
cnf(c_0_27,plain,(greatest_lower_bound(X1,multiply(X1,X2))=multiply(X1,greatest_lower_bound(identity,X2))), inference(spm,[status(thm)],[c_0_6, c_0_16])).
cnf(c_0_28,hypothesis,(greatest_lower_bound(identity,inverse(X1))=inverse(least_upper_bound(identity,X1))), inference(spm,[status(thm)],[c_0_7, c_0_22])).
cnf(c_0_29,plain,(inverse(multiply(X1,inverse(X2)))=multiply(X2,inverse(X1))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30,hypothesis,(multiply(least_upper_bound(b,c),inverse(a))=least_upper_bound(identity,multiply(c,inverse(a)))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31,plain,(multiply(X1,inverse(least_upper_bound(identity,X1)))=greatest_lower_bound(X1,identity)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_19]), c_0_28])).
cnf(c_0_32,hypothesis,(inverse(least_upper_bound(identity,multiply(c,inverse(a))))=multiply(a,inverse(least_upper_bound(b,c)))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33,plain,(greatest_lower_bound(identity,multiply(inverse(X1),X2))=multiply(inverse(X1),greatest_lower_bound(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_1]), c_0_8])).
cnf(c_0_34,hypothesis,(greatest_lower_bound(identity,multiply(c,inverse(a)))=multiply(c,inverse(least_upper_bound(b,c)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_0]), c_0_13]), c_0_8])).
cnf(c_0_35,plain,(multiply(least_upper_bound(X1,X2),inverse(X1))=least_upper_bound(identity,multiply(X2,inverse(X1)))), inference(spm,[status(thm)],[c_0_3, c_0_19])).
cnf(c_0_36,plain,(multiply(inverse(X1),greatest_lower_bound(identity,X1))=inverse(least_upper_bound(identity,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_1]), c_0_8]), c_0_28])).
cnf(c_0_37,hypothesis,(greatest_lower_bound(identity,multiply(a,inverse(c)))=multiply(a,inverse(least_upper_bound(b,c)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_29]), c_0_0]), c_0_16]), c_0_29]), c_0_0]), c_0_13])).
cnf(c_0_38,hypothesis,(least_upper_bound(identity,multiply(a,inverse(c)))=least_upper_bound(identity,multiply(b,inverse(c)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_26]), c_0_25])).
cnf(c_0_39,plain,(multiply(X1,least_upper_bound(inverse(X1),X2))=least_upper_bound(identity,multiply(X1,X2))), inference(spm,[status(thm)],[c_0_9, c_0_19])).
cnf(c_0_40,plain,(multiply(X1,least_upper_bound(X2,inverse(X1)))=least_upper_bound(identity,multiply(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_19]), c_0_4])).
cnf(c_0_41,hypothesis,(inverse(least_upper_bound(identity,multiply(b,inverse(c))))=multiply(c,inverse(least_upper_bound(b,c)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_29]), c_0_0]), c_0_13]), c_0_38])).
cnf(c_0_42,hypothesis,(multiply(X1,inverse(greatest_lower_bound(X1,X2)))=least_upper_bound(identity,multiply(X1,inverse(X2)))), inference(spm,[status(thm)],[c_0_39, c_0_10])).
cnf(c_0_43,hypothesis,(greatest_lower_bound(c,a)=greatest_lower_bound(b,c)), inference(rw,[status(thm)],[c_0_11, c_0_8])).
cnf(c_0_44,hypothesis,(multiply(X1,inverse(greatest_lower_bound(X2,X1)))=least_upper_bound(identity,multiply(X1,inverse(X2)))), inference(spm,[status(thm)],[c_0_40, c_0_10])).
cnf(c_0_45,hypothesis,(greatest_lower_bound(identity,multiply(b,inverse(c)))=multiply(b,inverse(least_upper_bound(b,c)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_41]), c_0_0]), c_0_13]), c_0_8])).
cnf(c_0_46,hypothesis,(least_upper_bound(identity,multiply(c,inverse(a)))=least_upper_bound(identity,multiply(c,inverse(b)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_47,hypothesis,(greatest_lower_bound(identity,multiply(c,inverse(b)))=multiply(c,inverse(least_upper_bound(b,c)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_45]), c_0_29]), c_0_0]), c_0_16]), c_0_29]), c_0_0]), c_0_13])).
cnf(c_0_48,plain,(multiply(inverse(multiply(X1,X2)),X1)=inverse(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_24]), c_0_17])).
cnf(c_0_49,hypothesis,(multiply(least_upper_bound(b,c),inverse(a))=least_upper_bound(identity,multiply(c,inverse(b)))), inference(rw,[status(thm)],[c_0_30, c_0_46])).
cnf(c_0_50,hypothesis,(inverse(least_upper_bound(identity,multiply(c,inverse(b))))=multiply(b,inverse(least_upper_bound(b,c)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_47]), c_0_29]), c_0_0]), c_0_13])).
cnf(c_0_51,hypothesis,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_0]), c_0_1]), c_0_16]), c_0_17]), c_0_12]), ['proof']).
# SZS output end CNFRefutation.
EOF
