% Problem    : LAT096-1 : TPTP v6.1.0. Released v2.6.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n106.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:32:56 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____102_C18_F1_PI_AE_Q4_CS_SP_S1S
# and selection function SelectComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)),meet(join(meet(X2,meet(meet(join(X2,X4),join(X5,X2)),X2)),meet(X6,join(X2,meet(meet(join(X2,X4),join(X5,X2)),X2)))),join(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', single_axiom)).
cnf(c_0_1,negated_conjecture,(meet(meet(join(a,b),join(c,b)),b)!=b), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_wal_axioms_5)).
cnf(c_0_2,plain,(join(meet(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4),meet(join(meet(X1,join(join(meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X5),meet(X6,join(meet(X2,X3),meet(X3,join(X2,X3))))),join(meet(X2,X3),meet(X3,join(X2,X3))))),meet(X3,join(X1,join(join(meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X5),meet(X6,join(meet(X2,X3),meet(X3,join(X2,X3))))),join(meet(X2,X3),meet(X3,join(X2,X3))))))),join(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4)))=join(meet(X2,X3),meet(X3,join(X2,X3)))), inference(spm,[status(thm)],[c_0_0, c_0_0])).
cnf(c_0_3,plain,(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)),meet(join(meet(X6,join(join(meet(X2,X7),meet(X8,X2)),X2)),meet(join(meet(X2,meet(meet(join(X2,X7),join(X8,X2)),X2)),meet(X9,join(X2,meet(meet(join(X2,X7),join(X8,X2)),X2)))),join(X6,join(join(meet(X2,X7),meet(X8,X2)),X2)))),join(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), inference(spm,[status(thm)],[c_0_2, c_0_0])).
cnf(c_0_4,plain,(join(meet(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3),meet(X2,join(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3)))=X2), inference(spm,[status(thm)],[c_0_3, c_0_0])).
cnf(c_0_5,plain,(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)),meet(X2,join(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), inference(spm,[status(thm)],[c_0_2, c_0_4])).
cnf(c_0_6,plain,(join(meet(X1,X2),meet(X1,join(X1,X2)))=X1), inference(spm,[status(thm)],[c_0_4, c_0_4])).
cnf(c_0_7,plain,(join(meet(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)),meet(X2,join(X1,join(join(meet(X2,X3),meet(X4,X2)),X2))))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_5]), c_0_5]), c_0_5]), c_0_6])).
cnf(c_0_8,plain,(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), inference(rw,[status(thm)],[c_0_5, c_0_7])).
cnf(c_0_9,plain,(join(meet(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)),meet(join(meet(X2,meet(meet(join(X2,X3),join(X4,X2)),X2)),meet(X5,join(X2,meet(meet(join(X2,X3),join(X4,X2)),X2)))),join(X1,join(join(meet(X2,X3),meet(X4,X2)),X2))))=X2), inference(spm,[status(thm)],[c_0_0, c_0_7])).
cnf(c_0_10,plain,(join(meet(X1,meet(meet(join(X1,X2),join(X3,X1)),X1)),meet(X4,join(X1,meet(meet(join(X1,X2),join(X3,X1)),X1))))=X1), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_11,plain,(join(meet(X1,X2),meet(X2,join(X1,X2)))=X2), inference(spm,[status(thm)],[c_0_8, c_0_10])).
cnf(c_0_12,plain,(meet(meet(join(X1,X2),join(X3,X1)),X1)=X1), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_13,plain,(join(meet(X1,X1),meet(X2,join(X1,X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_12]), c_0_12])).
cnf(c_0_14,plain,(join(join(meet(X1,X2),meet(X3,X1)),X1)=X1), inference(spm,[status(thm)],[c_0_7, c_0_13])).
cnf(c_0_15,plain,(meet(meet(join(X1,X2),X1),X1)=X1), inference(spm,[status(thm)],[c_0_12, c_0_14])).
cnf(c_0_16,plain,(meet(meet(X1,meet(X1,X2)),meet(X1,X2))=meet(X1,X2)), inference(spm,[status(thm)],[c_0_15, c_0_6])).
cnf(c_0_17,plain,(meet(meet(X1,meet(X2,X1)),meet(X2,X1))=meet(X2,X1)), inference(spm,[status(thm)],[c_0_15, c_0_11])).
cnf(c_0_18,plain,(meet(X1,X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_12]), c_0_12])).
cnf(c_0_19,plain,(join(meet(X1,X1),join(X1,X1))=X1), inference(spm,[status(thm)],[c_0_13, c_0_12])).
cnf(c_0_20,plain,(meet(meet(X1,X1),X1)=X1), inference(spm,[status(thm)],[c_0_17, c_0_12])).
cnf(c_0_21,plain,(join(join(meet(X1,X2),X1),X1)=X1), inference(spm,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_22,plain,(meet(join(X1,X1),X1)=X1), inference(spm,[status(thm)],[c_0_12, c_0_18])).
cnf(c_0_23,plain,(join(X1,join(X1,X1))=X1), inference(rw,[status(thm)],[c_0_19, c_0_18])).
cnf(c_0_24,plain,(join(X1,meet(X1,join(meet(X1,X1),X1)))=X1), inference(spm,[status(thm)],[c_0_11, c_0_20])).
cnf(c_0_25,plain,(join(X1,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_23])).
cnf(c_0_26,plain,(meet(meet(X1,join(X2,X1)),X1)=X1), inference(spm,[status(thm)],[c_0_12, c_0_24])).
cnf(c_0_27,plain,(join(X1,meet(X2,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_18]), c_0_25])).
cnf(c_0_28,plain,(meet(meet(meet(X1,X2),X2),meet(X1,X2))=meet(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_29,plain,(join(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,X2),X1))=meet(X1,X2)), inference(spm,[status(thm)],[c_0_6, c_0_6])).
cnf(c_0_30,plain,(meet(X1,meet(join(X1,X2),X1))=meet(join(X1,X2),X1)), inference(spm,[status(thm)],[c_0_28, c_0_15])).
cnf(c_0_31,plain,(meet(join(X1,X2),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_27]), c_0_18]), c_0_15]), c_0_15]), c_0_25])).
cnf(c_0_32,plain,(meet(X1,join(meet(X1,X2),X1))=join(meet(X1,X2),X1)), inference(spm,[status(thm)],[c_0_31, c_0_21])).
cnf(c_0_33,plain,(meet(join(meet(X1,X2),X1),X1)=X1), inference(spm,[status(thm)],[c_0_26, c_0_32])).
cnf(c_0_34,plain,(join(meet(X1,X2),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_33]), c_0_21]), c_0_33]), c_0_25])).
cnf(c_0_35,plain,(join(meet(X1,meet(X2,X1)),X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_27]), c_0_18])).
cnf(c_0_36,plain,(meet(X1,meet(X2,X1))=meet(X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_11])).
cnf(c_0_37,plain,(meet(meet(X1,X2),X1)=meet(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_34]), c_0_25])).
cnf(c_0_38,plain,(join(meet(X1,X2),X2)=X2), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39,plain,(meet(X1,join(X2,X1))=X1), inference(rw,[status(thm)],[c_0_26, c_0_37])).
cnf(c_0_40,plain,(join(X1,join(X2,X1))=join(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_41,plain,(meet(meet(join(X1,X2),join(X3,X2)),X2)=X2), inference(spm,[status(thm)],[c_0_12, c_0_40])).
cnf(c_0_42,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1, c_0_41])]), ['proof']).
# SZS output end CNFRefutation.
EOF
