% Problem    : BOO031-1 : TPTP v6.1.0. Released v2.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n130.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:17:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(multiply(add(X1,inverse(X1)),X2)=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', property3)).
cnf(c_0_1,axiom,(add(X1,inverse(X1))=n1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', additive_inverse)).
cnf(c_0_2,axiom,(multiply(multiply(add(X1,X2),add(X2,X3)),X2)=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', l4)).
cnf(c_0_3,axiom,(multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', associativity_of_multiply)).
cnf(c_0_4,axiom,(multiply(X1,add(X2,add(X1,X3)))=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', l2)).
cnf(c_0_5,axiom,(add(X1,multiply(X2,multiply(X1,X3)))=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', l1)).
cnf(c_0_6,axiom,(add(multiply(X1,inverse(X1)),X2)=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', property3_dual)).
cnf(c_0_7,axiom,(multiply(X1,inverse(X1))=n0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', multiplicative_inverse)).
cnf(c_0_8,axiom,(add(multiply(X1,X2),add(multiply(X2,X3),multiply(X3,X1)))=multiply(add(X1,X2),multiply(add(X2,X3),add(X3,X1)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', distributivity)).
cnf(c_0_9,negated_conjecture,(multiply(a,add(b,c))!=add(multiply(b,a),multiply(c,a))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_multiply_add_property)).
cnf(c_0_10,axiom,(add(add(X1,X2),X3)=add(X1,add(X2,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', associativity_of_add)).
cnf(c_0_11,plain,(multiply(n1,X1)=X1), inference(rw,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_12,plain,(multiply(add(X1,X2),multiply(add(X2,X3),X2))=X2), inference(rw,[status(thm)],[c_0_2, c_0_3])).
cnf(c_0_13,plain,(multiply(X1,add(X2,X1))=X1), inference(spm,[status(thm)],[c_0_4, c_0_5])).
cnf(c_0_14,plain,(add(X1,multiply(X1,X2))=X1), inference(spm,[status(thm)],[c_0_5, c_0_11])).
cnf(c_0_15,plain,(multiply(X1,add(X2,n1))=X1), inference(spm,[status(thm)],[c_0_4, c_0_1])).
cnf(c_0_16,plain,(add(n0,X1)=X1), inference(rw,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_17,plain,(multiply(add(X1,X2),X2)=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_1]), c_0_11])).
cnf(c_0_18,plain,(add(X1,multiply(X2,X1))=X1), inference(spm,[status(thm)],[c_0_5, c_0_4])).
cnf(c_0_19,plain,(multiply(X1,multiply(X2,X1))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_3])).
cnf(c_0_20,plain,(multiply(X1,n1)=X1), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21,plain,(multiply(X1,X1)=X1), inference(spm,[status(thm)],[c_0_13, c_0_16])).
cnf(c_0_22,plain,(multiply(X1,X2)=multiply(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_23,plain,(add(n1,X1)=n1), inference(spm,[status(thm)],[c_0_18, c_0_20])).
cnf(c_0_24,plain,(add(X1,n1)=n1), inference(spm,[status(thm)],[c_0_11, c_0_15])).
cnf(c_0_25,plain,(multiply(X1,multiply(X1,X2))=multiply(X1,X2)), inference(spm,[status(thm)],[c_0_3, c_0_21])).
cnf(c_0_26,plain,(multiply(X1,multiply(X2,X3))=multiply(X2,multiply(X3,X1))), inference(spm,[status(thm)],[c_0_3, c_0_22])).
cnf(c_0_27,plain,(multiply(X1,multiply(add(X2,X1),X3))=multiply(X1,X3)), inference(spm,[status(thm)],[c_0_3, c_0_13])).
cnf(c_0_28,plain,(add(X1,X2)=add(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_23]), c_0_24]), c_0_11]), c_0_11]), c_0_20]), c_0_11]), c_0_14])).
cnf(c_0_29,plain,(add(multiply(X1,X2),add(multiply(X2,X3),multiply(X2,multiply(X3,X1))))=multiply(X2,add(multiply(X2,X3),X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_14]), c_0_3]), c_0_25]), c_0_26]), c_0_22]), c_0_27])).
cnf(c_0_30,plain,(add(multiply(X1,X2),multiply(X1,multiply(X2,X3)))=multiply(X1,X2)), inference(spm,[status(thm)],[c_0_14, c_0_3])).
cnf(c_0_31,plain,(add(multiply(X1,X2),add(multiply(X3,X1),multiply(X1,multiply(X2,X3))))=multiply(X1,add(multiply(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_14]), c_0_3]), c_0_25]), c_0_28]), c_0_22]), c_0_27])).
cnf(c_0_32,plain,(add(multiply(X1,X2),multiply(X2,X3))=multiply(X2,add(multiply(X2,X3),X1))), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33,plain,(add(X1,multiply(X2,multiply(X3,X1)))=X1), inference(spm,[status(thm)],[c_0_18, c_0_3])).
cnf(c_0_34,plain,(add(multiply(X1,X2),multiply(X1,X3))=multiply(X1,add(multiply(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32]), c_0_28]), c_0_33])).
cnf(c_0_35,plain,(multiply(X1,add(multiply(X1,X2),X3))=multiply(X1,add(multiply(X1,X3),X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_22]), c_0_34])).
cnf(c_0_36,negated_conjecture,(add(multiply(a,b),multiply(a,c))!=multiply(a,add(b,c))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_22]), c_0_22])).
cnf(c_0_37,plain,(multiply(X1,add(multiply(X1,X2),X3))=multiply(X1,add(X2,multiply(X1,X3)))), inference(spm,[status(thm)],[c_0_35, c_0_28])).
cnf(c_0_38,plain,(multiply(X1,add(X1,X2))=X1), inference(spm,[status(thm)],[c_0_4, c_0_16])).
cnf(c_0_39,plain,(multiply(add(X1,X2),multiply(add(X2,X3),add(X3,X1)))=add(multiply(X1,X2),multiply(X3,add(multiply(X3,X1),X2)))), inference(rw,[status(thm)],[c_0_8, c_0_32])).
cnf(c_0_40,plain,(add(X1,add(multiply(X1,X2),X3))=add(X1,X3)), inference(spm,[status(thm)],[c_0_10, c_0_14])).
cnf(c_0_41,plain,(multiply(X1,add(X2,add(X3,X1)))=X1), inference(spm,[status(thm)],[c_0_13, c_0_10])).
cnf(c_0_42,negated_conjecture,(multiply(a,add(c,multiply(a,b)))!=multiply(a,add(b,c))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_34]), c_0_28])).
cnf(c_0_43,plain,(multiply(X1,add(X2,multiply(X1,X3)))=multiply(X1,add(X3,multiply(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_37]), c_0_37])).
cnf(c_0_44,plain,(multiply(X1,multiply(add(X1,X2),X3))=multiply(X1,X3)), inference(spm,[status(thm)],[c_0_3, c_0_38])).
cnf(c_0_45,plain,(multiply(add(X1,X2),multiply(add(X2,X3),add(X3,X1)))=add(multiply(X1,X2),multiply(X3,add(X1,multiply(X3,X2))))), inference(rw,[status(thm)],[c_0_39, c_0_37])).
cnf(c_0_46,plain,(multiply(X1,multiply(X2,add(X1,X3)))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_40]), c_0_3])).
cnf(c_0_47,plain,(multiply(X1,multiply(X2,add(X3,X1)))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_14]), c_0_3])).
cnf(c_0_48,negated_conjecture,(multiply(a,add(b,multiply(a,c)))!=multiply(a,add(b,c))), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49,plain,(multiply(X1,add(X2,multiply(X1,X3)))=multiply(X1,add(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_37]), c_0_46]), c_0_47])).
cnf(c_0_50,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49])]), ['proof']).
# SZS output end CNFRefutation.
EOF
