% Problem    : GRP179-3 : TPTP v6.1.0. Bugfixed v1.2.1.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n102.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:13:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C02_F1_PI_AE_Q12_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,hypothesis,(inverse(multiply(X1,X2))=multiply(inverse(X2),inverse(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p18_3)).
cnf(c_0_1,hypothesis,(inverse(identity)=identity), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p18_1)).
cnf(c_0_2,axiom,(multiply(identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_identity)).
cnf(c_0_3,hypothesis,(inverse(inverse(X1))=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p18_2)).
cnf(c_0_4,axiom,(multiply(X1,greatest_lower_bound(X2,X3))=greatest_lower_bound(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_glb1)).
cnf(c_0_5,axiom,(multiply(inverse(X1),X1)=identity), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_inverse)).
cnf(c_0_6,axiom,(least_upper_bound(X1,greatest_lower_bound(X1,X2))=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', lub_absorbtion)).
cnf(c_0_7,axiom,(multiply(greatest_lower_bound(X1,X2),X3)=greatest_lower_bound(multiply(X1,X3),multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_glb2)).
cnf(c_0_8,axiom,(greatest_lower_bound(X1,X2)=greatest_lower_bound(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', symmetry_of_glb)).
cnf(c_0_9,axiom,(multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', associativity)).
cnf(c_0_10,axiom,(least_upper_bound(X1,X2)=least_upper_bound(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', symmetry_of_lub)).
cnf(c_0_11,axiom,(greatest_lower_bound(X1,least_upper_bound(X1,X2))=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', glb_absorbtion)).
cnf(c_0_12,axiom,(greatest_lower_bound(X1,greatest_lower_bound(X2,X3))=greatest_lower_bound(greatest_lower_bound(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', associativity_of_glb)).
cnf(c_0_13,axiom,(multiply(least_upper_bound(X1,X2),X3)=least_upper_bound(multiply(X1,X3),multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_lub2)).
cnf(c_0_14,axiom,(multiply(X1,least_upper_bound(X2,X3))=least_upper_bound(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_lub1)).
cnf(c_0_15,axiom,(least_upper_bound(X1,least_upper_bound(X2,X3))=least_upper_bound(least_upper_bound(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', associativity_of_lub)).
cnf(c_0_16,negated_conjecture,(least_upper_bound(inverse(a),identity)!=inverse(greatest_lower_bound(a,identity))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_p18)).
cnf(c_0_17,hypothesis,(multiply(inverse(X1),identity)=inverse(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_18,hypothesis,(multiply(X1,identity)=X1), inference(spm,[status(thm)],[c_0_17, c_0_3])).
cnf(c_0_19,hypothesis,(greatest_lower_bound(X1,multiply(X1,X2))=multiply(X1,greatest_lower_bound(identity,X2))), inference(spm,[status(thm)],[c_0_4, c_0_18])).
cnf(c_0_20,hypothesis,(multiply(X1,inverse(X1))=identity), inference(spm,[status(thm)],[c_0_5, c_0_3])).
cnf(c_0_21,hypothesis,(least_upper_bound(X1,multiply(X1,greatest_lower_bound(identity,X2)))=X1), inference(spm,[status(thm)],[c_0_6, c_0_19])).
cnf(c_0_22,plain,(greatest_lower_bound(X1,multiply(X2,X1))=multiply(greatest_lower_bound(X2,identity),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_2]), c_0_8])).
cnf(c_0_23,hypothesis,(multiply(X1,multiply(inverse(X1),X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_20]), c_0_2])).
cnf(c_0_24,hypothesis,(least_upper_bound(identity,inverse(greatest_lower_bound(identity,X1)))=inverse(greatest_lower_bound(identity,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_5]), c_0_10])).
cnf(c_0_25,hypothesis,(greatest_lower_bound(inverse(X1),inverse(multiply(X1,X2)))=multiply(greatest_lower_bound(identity,inverse(X2)),inverse(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_0]), c_0_8])).
cnf(c_0_26,hypothesis,(multiply(X1,greatest_lower_bound(identity,multiply(inverse(X1),X2)))=greatest_lower_bound(X1,X2)), inference(spm,[status(thm)],[c_0_19, c_0_23])).
cnf(c_0_27,hypothesis,(greatest_lower_bound(identity,inverse(greatest_lower_bound(identity,X1)))=identity), inference(spm,[status(thm)],[c_0_11, c_0_24])).
cnf(c_0_28,hypothesis,(greatest_lower_bound(inverse(X1),inverse(greatest_lower_bound(X1,X2)))=inverse(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_2])).
cnf(c_0_29,plain,(multiply(inverse(X1),multiply(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_5]), c_0_2])).
cnf(c_0_30,hypothesis,(greatest_lower_bound(X1,inverse(greatest_lower_bound(inverse(X1),X2)))=X1), inference(spm,[status(thm)],[c_0_28, c_0_3])).
cnf(c_0_31,hypothesis,(multiply(X1,inverse(multiply(X2,X1)))=inverse(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_0]), c_0_3])).
cnf(c_0_32,plain,(greatest_lower_bound(X1,least_upper_bound(X2,X1))=X1), inference(spm,[status(thm)],[c_0_11, c_0_10])).
cnf(c_0_33,hypothesis,(multiply(X1,greatest_lower_bound(inverse(X1),X2))=greatest_lower_bound(identity,multiply(X1,X2))), inference(spm,[status(thm)],[c_0_4, c_0_20])).
cnf(c_0_34,hypothesis,(greatest_lower_bound(X1,inverse(greatest_lower_bound(X2,inverse(X1))))=X1), inference(spm,[status(thm)],[c_0_30, c_0_8])).
cnf(c_0_35,plain,(multiply(greatest_lower_bound(identity,inverse(X1)),X1)=greatest_lower_bound(X1,identity)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_5]), c_0_8])).
cnf(c_0_36,hypothesis,(multiply(inverse(multiply(X1,X2)),X1)=inverse(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_31]), c_0_3])).
cnf(c_0_37,hypothesis,(greatest_lower_bound(identity,multiply(X1,inverse(X2)))=multiply(greatest_lower_bound(X1,X2),inverse(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_20]), c_0_8])).
cnf(c_0_38,plain,(greatest_lower_bound(X1,greatest_lower_bound(X2,X1))=greatest_lower_bound(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_6]), c_0_12])).
cnf(c_0_39,hypothesis,(greatest_lower_bound(identity,multiply(X1,inverse(greatest_lower_bound(X2,X1))))=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_20]), c_0_3])).
cnf(c_0_40,plain,(least_upper_bound(X1,greatest_lower_bound(X2,X1))=X1), inference(spm,[status(thm)],[c_0_6, c_0_8])).
cnf(c_0_41,plain,(least_upper_bound(X1,multiply(X2,X1))=multiply(least_upper_bound(X2,identity),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_2]), c_0_10])).
cnf(c_0_42,hypothesis,(least_upper_bound(X1,multiply(X1,X2))=multiply(X1,least_upper_bound(identity,X2))), inference(spm,[status(thm)],[c_0_14, c_0_18])).
cnf(c_0_43,hypothesis,(multiply(greatest_lower_bound(identity,X1),inverse(X1))=greatest_lower_bound(identity,inverse(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_3]), c_0_8])).
cnf(c_0_44,hypothesis,(least_upper_bound(inverse(X1),multiply(inverse(multiply(X2,X1)),X3))=multiply(inverse(multiply(X2,X1)),least_upper_bound(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_36]), c_0_10])).
cnf(c_0_45,hypothesis,(multiply(greatest_lower_bound(X1,X2),inverse(greatest_lower_bound(X2,X1)))=identity), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_46,plain,(least_upper_bound(X1,least_upper_bound(X2,X1))=least_upper_bound(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_11]), c_0_15])).
cnf(c_0_47,hypothesis,(multiply(X1,inverse(least_upper_bound(X1,multiply(X2,X1))))=inverse(least_upper_bound(X2,identity))), inference(spm,[status(thm)],[c_0_31, c_0_41])).
cnf(c_0_48,hypothesis,(least_upper_bound(inverse(multiply(X1,X2)),multiply(inverse(X2),X3))=multiply(inverse(X2),least_upper_bound(inverse(X1),X3))), inference(spm,[status(thm)],[c_0_14, c_0_0])).
cnf(c_0_49,hypothesis,(inverse(multiply(inverse(X1),X2))=multiply(inverse(X2),X1)), inference(spm,[status(thm)],[c_0_0, c_0_3])).
cnf(c_0_50,hypothesis,(multiply(inverse(X1),multiply(inverse(X2),X3))=multiply(inverse(multiply(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_9, c_0_0])).
cnf(c_0_51,plain,(multiply(least_upper_bound(X1,inverse(X2)),X2)=least_upper_bound(identity,multiply(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_5]), c_0_10])).
cnf(c_0_52,hypothesis,(multiply(greatest_lower_bound(identity,X1),least_upper_bound(identity,inverse(X1)))=least_upper_bound(greatest_lower_bound(identity,X1),greatest_lower_bound(identity,inverse(X1)))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_53,hypothesis,(least_upper_bound(greatest_lower_bound(X1,X2),X3)=least_upper_bound(X3,greatest_lower_bound(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_3]), c_0_1]), c_0_2]), c_0_1]), c_0_2])).
cnf(c_0_54,plain,(greatest_lower_bound(least_upper_bound(X1,X2),least_upper_bound(X2,X1))=least_upper_bound(X1,X2)), inference(spm,[status(thm)],[c_0_32, c_0_46])).
cnf(c_0_55,plain,(greatest_lower_bound(X1,greatest_lower_bound(least_upper_bound(X1,X2),X3))=greatest_lower_bound(X1,X3)), inference(spm,[status(thm)],[c_0_12, c_0_11])).
cnf(c_0_56,hypothesis,(multiply(inverse(least_upper_bound(identity,X1)),X1)=inverse(least_upper_bound(identity,inverse(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50]), c_0_51]), c_0_29]), c_0_10])).
cnf(c_0_57,hypothesis,(multiply(greatest_lower_bound(identity,inverse(X1)),least_upper_bound(identity,X1))=least_upper_bound(greatest_lower_bound(identity,X1),greatest_lower_bound(identity,inverse(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_3]), c_0_53]), c_0_8])).
cnf(c_0_58,hypothesis,(greatest_lower_bound(inverse(X1),multiply(inverse(multiply(X2,X1)),X3))=multiply(inverse(multiply(X2,X1)),greatest_lower_bound(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_36]), c_0_8])).
cnf(c_0_59,hypothesis,(multiply(least_upper_bound(X1,X2),inverse(least_upper_bound(X2,X1)))=identity), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_54]), c_0_54])).
cnf(c_0_60,plain,(least_upper_bound(greatest_lower_bound(X1,X2),greatest_lower_bound(X2,least_upper_bound(X1,X3)))=greatest_lower_bound(least_upper_bound(X1,X3),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_55]), c_0_53])).
cnf(c_0_61,hypothesis,(greatest_lower_bound(least_upper_bound(identity,X1),least_upper_bound(identity,inverse(X1)))=least_upper_bound(greatest_lower_bound(identity,X1),greatest_lower_bound(identity,inverse(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_56]), c_0_3]), c_0_3]), c_0_3]), c_0_57])).
cnf(c_0_62,plain,(least_upper_bound(X1,least_upper_bound(greatest_lower_bound(X1,X2),X3))=least_upper_bound(X1,X3)), inference(spm,[status(thm)],[c_0_15, c_0_6])).
cnf(c_0_63,hypothesis,(greatest_lower_bound(least_upper_bound(X1,X2),X3)=greatest_lower_bound(X3,least_upper_bound(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_3]), c_0_1]), c_0_2]), c_0_1]), c_0_2])).
cnf(c_0_64,plain,(multiply(inverse(X1),least_upper_bound(X1,X2))=least_upper_bound(identity,multiply(inverse(X1),X2))), inference(spm,[status(thm)],[c_0_14, c_0_5])).
cnf(c_0_65,hypothesis,(least_upper_bound(greatest_lower_bound(identity,X1),greatest_lower_bound(identity,inverse(X1)))=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_11]), c_0_62]), c_0_6]), c_0_63]), c_0_10]), c_0_61])).
cnf(c_0_66,hypothesis,(multiply(inverse(greatest_lower_bound(identity,X1)),greatest_lower_bound(identity,inverse(X1)))=inverse(X1)), inference(spm,[status(thm)],[c_0_29, c_0_43])).
cnf(c_0_67,negated_conjecture,(least_upper_bound(identity,inverse(a))!=inverse(greatest_lower_bound(identity,a))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_8]), c_0_10])).
cnf(c_0_68,hypothesis,(least_upper_bound(identity,inverse(X1))=inverse(greatest_lower_bound(identity,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_18]), c_0_66])).
cnf(c_0_69,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68])]), ['proof']).
# SZS output end CNFRefutation.
EOF
