% Problem    : COL044-1 : TPTP v6.1.0. Released v1.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n188.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:19:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(apply(apply(apply(b,X1),X2),X3)=apply(X1,apply(X2,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', b_definition)).
cnf(c_0_1,axiom,(apply(apply(apply(n,X1),X2),X3)=apply(apply(apply(X1,X3),X2),X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', n_definition)).
cnf(c_0_2,negated_conjecture,(apply(X1,f(X1))!=apply(f(X1),apply(X1,f(X1)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_fixed_point)).
cnf(c_0_3,plain,(apply(apply(apply(n,b),X1),X2)=apply(X2,apply(X1,X2))), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_4,plain,(apply(apply(apply(n,apply(b,X1)),X2),X3)=apply(apply(X1,apply(X3,X2)),X3)), inference(spm,[status(thm)],[c_0_1, c_0_0])).
cnf(c_0_5,negated_conjecture,(apply(apply(apply(n,b),X1),f(X1))!=apply(X1,f(X1))), inference(spm,[status(thm)],[c_0_2, c_0_3])).
cnf(c_0_6,plain,(apply(apply(apply(n,apply(b,apply(b,X1))),X2),X3)=apply(X1,apply(apply(X3,X2),X3))), inference(spm,[status(thm)],[c_0_0, c_0_4])).
cnf(c_0_7,negated_conjecture,(apply(apply(apply(apply(n,apply(b,apply(b,apply(n,b)))),X1),X2),f(apply(apply(X2,X1),X2)))!=apply(apply(apply(X2,X1),X2),f(apply(apply(X2,X1),X2)))), inference(spm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_8,negated_conjecture,($false), inference(er,[status(thm)],[c_0_7]), ['proof']).
# SZS output end CNFRefutation.
EOF
