% Problem    : GRP172-1 : TPTP v6.1.0. Bugfixed v1.2.1.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n117.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:01:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C02_F1_PI_AE_Q12_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', associativity)).
cnf(c_0_1,axiom,(multiply(inverse(X1),X1)=identity), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_inverse)).
cnf(c_0_2,axiom,(multiply(identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_identity)).
cnf(c_0_3,axiom,(greatest_lower_bound(X1,greatest_lower_bound(X2,X3))=greatest_lower_bound(greatest_lower_bound(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', associativity_of_glb)).
cnf(c_0_4,hypothesis,(greatest_lower_bound(identity,a)=identity), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p04b_1)).
cnf(c_0_5,axiom,(multiply(X1,greatest_lower_bound(X2,X3))=greatest_lower_bound(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_glb1)).
cnf(c_0_6,hypothesis,(greatest_lower_bound(identity,b)=identity), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p04b_2)).
cnf(c_0_7,negated_conjecture,(greatest_lower_bound(identity,multiply(a,b))!=identity), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_p04b)).
cnf(c_0_8,plain,(multiply(inverse(X1),multiply(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_9,plain,(multiply(inverse(inverse(X1)),identity)=X1), inference(spm,[status(thm)],[c_0_8, c_0_1])).
cnf(c_0_10,plain,(multiply(inverse(inverse(X1)),X2)=multiply(X1,X2)), inference(spm,[status(thm)],[c_0_8, c_0_8])).
cnf(c_0_11,plain,(multiply(X1,identity)=X1), inference(rw,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12,hypothesis,(greatest_lower_bound(identity,greatest_lower_bound(a,X1))=greatest_lower_bound(identity,X1)), inference(spm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_13,plain,(greatest_lower_bound(X1,multiply(X1,X2))=multiply(X1,greatest_lower_bound(identity,X2))), inference(spm,[status(thm)],[c_0_5, c_0_11])).
cnf(c_0_14,hypothesis,(greatest_lower_bound(identity,multiply(a,greatest_lower_bound(identity,X1)))=greatest_lower_bound(identity,multiply(a,X1))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_15,hypothesis,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_6]), c_0_11]), c_0_4]), c_0_7]), ['proof']).
# SZS output end CNFRefutation.
EOF
