% Problem    : GRP772-1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n116.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:29:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___103_C18_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(eta(X1)=product(i(X1),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos12)).
cnf(c_0_1,axiom,(i(X1)=difference(X1,one)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos09)).
cnf(c_0_2,axiom,(product(i(X1),X1)=product(X1,j(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos11)).
cnf(c_0_3,axiom,(j(X1)=quotient(one,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos10)).
cnf(c_0_4,axiom,(product(X1,product(eta(X1),X2))=product(j(j(X1)),X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos14)).
cnf(c_0_5,axiom,(product(a(X1,X2,X3),product(X4,X5))=product(product(a(X1,X2,X3),X4),X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos27)).
cnf(c_0_6,axiom,(product(product(X1,X2),a(X3,X4,X5))=product(X1,product(X2,a(X3,X4,X5)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos28)).
cnf(c_0_7,axiom,(product(X1,difference(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos03)).
cnf(c_0_8,axiom,(difference(X1,product(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos04)).
cnf(c_0_9,axiom,(product(a(X1,X2,X3),difference(X3,product(a(X3,X1,X2),X3)))=one), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos29)).
cnf(c_0_10,axiom,(product(one,X1)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos02)).
cnf(c_0_11,axiom,(product(quotient(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos06)).
cnf(c_0_12,axiom,(difference(X1,product(product(X1,X2),X3))=quotient(product(X2,product(X3,X1)),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos07)).
cnf(c_0_13,axiom,(product(X1,one)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos01)).
cnf(c_0_14,axiom,(a(i(X1),X2,X3)=a(j(X1),X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos31)).
cnf(c_0_15,axiom,(a(j(X1),X2,X3)=a(X2,X3,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos32)).
cnf(c_0_16,axiom,(a(X1,i(X2),X3)=a(X1,j(X2),X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos30)).
cnf(c_0_17,axiom,(quotient(product(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos05)).
cnf(c_0_18,negated_conjecture,(a(x0,x1,x1)!=a(x1,x1,x0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
cnf(c_0_19,plain,(product(difference(X1,one),X1)=eta(X1)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_20,plain,(product(difference(X1,one),X1)=product(X1,quotient(one,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_1]), c_0_3])).
cnf(c_0_21,plain,(product(X1,product(product(difference(X1,one),X1),X2))=product(quotient(one,quotient(one,X1)),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_3]), c_0_3]), c_0_19])).
cnf(c_0_22,plain,(product(a(X1,X2,X3),product(X4,quotient(one,product(a(X1,X2,X3),X4))))=product(difference(product(a(X1,X2,X3),X4),one),product(a(X1,X2,X3),X4))), inference(spm,[status(thm)],[c_0_20, c_0_5])).
cnf(c_0_23,plain,(product(X1,product(difference(X1,X2),a(X3,X4,X5)))=product(X2,a(X3,X4,X5))), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_24,plain,(difference(X1,product(a(X1,X2,X3),X1))=difference(a(X2,X3,X1),one)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_25,plain,(product(difference(a(X1,X2,X3),one),a(X1,X2,X3))=one), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_10]), c_0_23]), c_0_10]), c_0_23]), c_0_10]), c_0_11])).
cnf(c_0_26,plain,(product(X1,difference(a(X2,X3,X1),one))=product(a(X1,X2,X3),X1)), inference(spm,[status(thm)],[c_0_7, c_0_24])).
cnf(c_0_27,plain,(product(X1,difference(a(X2,X3,X4),one))=quotient(X1,a(X2,X3,X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_25]), c_0_13]), c_0_5]), c_0_8])).
cnf(c_0_28,plain,(a(quotient(one,X1),X2,X3)=a(difference(X1,one),X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_1]), c_0_3])).
cnf(c_0_29,plain,(a(quotient(one,X1),X2,X3)=a(X2,X3,X1)), inference(rw,[status(thm)],[c_0_15, c_0_3])).
cnf(c_0_30,plain,(difference(quotient(X1,X2),X1)=X2), inference(spm,[status(thm)],[c_0_8, c_0_11])).
cnf(c_0_31,plain,(quotient(X1,a(X2,X3,X1))=product(a(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32,plain,(a(difference(X1,one),X2,X3)=a(X2,X3,X1)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33,plain,(difference(product(a(X1,X2,X3),X1),X1)=a(X2,X3,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34,plain,(a(X1,X2,quotient(one,X3))=a(X3,X1,X2)), inference(spm,[status(thm)],[c_0_32, c_0_30])).
cnf(c_0_35,plain,(a(X1,quotient(one,X2),X3)=a(X1,difference(X2,one),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_1]), c_0_3])).
cnf(c_0_36,plain,(difference(product(a(X1,X2,X3),X2),X2)=a(X3,difference(X1,one),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])).
cnf(c_0_37,plain,(product(a(X1,X2,X3),product(X1,a(X2,X3,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_31]), c_0_5])).
cnf(c_0_38,plain,(a(X1,difference(X2,one),X2)=a(X2,X1,X2)), inference(spm,[status(thm)],[c_0_33, c_0_36])).
cnf(c_0_39,plain,(quotient(X1,difference(X2,X1))=X2), inference(spm,[status(thm)],[c_0_17, c_0_7])).
cnf(c_0_40,plain,(quotient(X1,product(X1,a(X2,X3,X1)))=a(X1,X2,X3)), inference(spm,[status(thm)],[c_0_17, c_0_37])).
cnf(c_0_41,plain,(a(X1,difference(X2,one),X1)=a(X1,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_38]), c_0_35]), c_0_32])).
cnf(c_0_42,plain,(a(X1,X2,difference(X3,one))=a(X3,X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_39])).
cnf(c_0_43,plain,(a(X1,X1,X2)=a(X2,X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_40]), c_0_42])).
cnf(c_0_44,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_43])]), ['proof']).
# SZS output end CNFRefutation.
EOF
