% Problem    : LAT021-1 : TPTP v6.1.0. Released v2.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n119.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:27:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C48_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(meet(join(X1,meet(X2,X3)),join(X1,X2))=join(X1,meet(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/LAT004-0.ax', quasi_lattice2)).
cnf(c_0_1,axiom,(meet(X1,X2)=meet(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/LAT004-0.ax', commutativity_of_meet)).
cnf(c_0_2,axiom,(join(X1,X2)=join(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/LAT004-0.ax', commutativity_of_join)).
cnf(c_0_3,axiom,(meet(meet(X1,X2),X3)=meet(X1,meet(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/LAT004-0.ax', associativity_of_meet)).
cnf(c_0_4,axiom,(meet(X1,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/LAT004-0.ax', idempotence_of_meet)).
cnf(c_0_5,axiom,(join(meet(X1,join(X2,X3)),meet(X1,X2))=meet(X1,join(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/LAT004-0.ax', quasi_lattice1)).
cnf(c_0_6,axiom,(join(join(X1,X2),X3)=join(X1,join(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/LAT004-0.ax', associativity_of_join)).
cnf(c_0_7,axiom,(join(X1,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/LAT004-0.ax', idempotence_of_join)).
cnf(c_0_8,hypothesis,(join(join(X1,meet(X2,X3)),meet(join(X1,X2),X3))=join(X1,meet(X2,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bowden)).
cnf(c_0_9,negated_conjecture,(meet(a,join(b,c))!=join(meet(a,b),meet(a,c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_distributivity)).
cnf(c_0_10,plain,(meet(join(X1,X2),join(X1,meet(X2,X3)))=join(X1,meet(X2,X3))), inference(rw,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_11,plain,(meet(join(X1,X2),join(meet(X2,X3),X1))=join(meet(X2,X3),X1)), inference(spm,[status(thm)],[c_0_10, c_0_2])).
cnf(c_0_12,plain,(meet(X1,meet(X2,X3))=meet(X3,meet(X1,X2))), inference(spm,[status(thm)],[c_0_1, c_0_3])).
cnf(c_0_13,plain,(meet(join(X1,X2),join(X2,X1))=join(X2,X1)), inference(spm,[status(thm)],[c_0_11, c_0_4])).
cnf(c_0_14,plain,(join(meet(X1,X2),meet(X1,join(X2,X3)))=meet(X1,join(X2,X3))), inference(rw,[status(thm)],[c_0_5, c_0_2])).
cnf(c_0_15,plain,(join(X1,join(X1,X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_16,plain,(meet(X1,meet(X2,X3))=meet(X3,meet(X2,X1))), inference(spm,[status(thm)],[c_0_12, c_0_1])).
cnf(c_0_17,plain,(meet(join(X1,X2),meet(join(X2,X1),X3))=meet(join(X2,X1),X3)), inference(spm,[status(thm)],[c_0_3, c_0_13])).
cnf(c_0_18,plain,(join(X1,meet(X1,join(X1,X2)))=meet(X1,join(X1,X2))), inference(spm,[status(thm)],[c_0_14, c_0_4])).
cnf(c_0_19,plain,(join(X1,join(X2,X1))=join(X2,X1)), inference(spm,[status(thm)],[c_0_15, c_0_2])).
cnf(c_0_20,hypothesis,(join(X1,join(meet(X2,X3),meet(join(X1,X2),X3)))=join(X1,meet(X2,X3))), inference(rw,[status(thm)],[c_0_8, c_0_6])).
cnf(c_0_21,plain,(meet(join(X1,X2),X3)=meet(X3,join(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_13])).
cnf(c_0_22,plain,(join(meet(X1,X2),meet(X2,join(X1,X3)))=meet(X2,join(X1,X3))), inference(spm,[status(thm)],[c_0_14, c_0_1])).
cnf(c_0_23,plain,(join(X1,meet(X1,join(X2,X1)))=meet(X1,join(X2,X1))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24,hypothesis,(join(X1,meet(X2,join(X3,X1)))=join(X1,meet(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
cnf(c_0_25,plain,(meet(X1,join(X2,X1))=join(X1,meet(X2,X1))), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_26,plain,(join(X1,join(X2,X3))=join(X2,join(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_2]), c_0_6])).
cnf(c_0_27,plain,(meet(X1,meet(X2,X3))=meet(X2,meet(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_3, c_0_1]), c_0_3])).
cnf(c_0_28,plain,(meet(join(X1,meet(X2,X1)),X3)=meet(X1,meet(join(X2,X1),X3))), inference(spm,[status(thm)],[c_0_3, c_0_25])).
cnf(c_0_29,plain,(join(X1,join(X2,X3))=join(X3,join(X1,X2))), inference(spm,[status(thm)],[c_0_2, c_0_6])).
cnf(c_0_30,plain,(join(meet(X1,X2),meet(X1,join(X3,X2)))=meet(X1,join(X3,X2))), inference(spm,[status(thm)],[c_0_14, c_0_2])).
cnf(c_0_31,hypothesis,(join(X1,meet(X2,join(X1,X3)))=join(X1,meet(X3,X2))), inference(spm,[status(thm)],[c_0_24, c_0_2])).
cnf(c_0_32,plain,(meet(X1,meet(X1,X2))=meet(X1,X2)), inference(spm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_33,plain,(join(X1,join(meet(X1,join(X1,X2)),meet(X2,X3)))=join(X1,meet(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_10]), c_0_1]), c_0_26])).
cnf(c_0_34,plain,(meet(X1,join(X1,X2))=join(X1,meet(X2,X1))), inference(spm,[status(thm)],[c_0_25, c_0_2])).
cnf(c_0_35,plain,(meet(X1,meet(join(X1,meet(X1,X2)),join(X2,X3)))=meet(X1,join(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_14]), c_0_2]), c_0_27])).
cnf(c_0_36,plain,(meet(join(X1,meet(X1,X2)),X3)=meet(X1,meet(join(X2,X1),X3))), inference(spm,[status(thm)],[c_0_28, c_0_1])).
cnf(c_0_37,plain,(join(meet(X1,X2),meet(X1,join(X3,join(X4,X2))))=meet(X1,join(X3,join(X4,X2)))), inference(spm,[status(thm)],[c_0_14, c_0_29])).
cnf(c_0_38,plain,(join(meet(X1,X2),meet(join(X3,X2),X1))=meet(join(X3,X2),X1)), inference(spm,[status(thm)],[c_0_30, c_0_1])).
cnf(c_0_39,hypothesis,(join(X1,meet(join(X1,X2),X3))=join(X1,meet(X2,X3))), inference(spm,[status(thm)],[c_0_24, c_0_21])).
cnf(c_0_40,plain,(join(X1,join(meet(X2,X3),meet(join(X1,X2),join(X1,join(meet(X2,X3),X4)))))=meet(join(X1,X2),join(X1,join(meet(X2,X3),X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_10]), c_0_6]), c_0_6]), c_0_6])).
cnf(c_0_41,hypothesis,(join(X1,join(X2,meet(X3,join(X1,X4))))=join(X2,join(X1,meet(X4,X3)))), inference(spm,[status(thm)],[c_0_26, c_0_31])).
cnf(c_0_42,plain,(meet(X1,join(meet(X1,X2),X3))=join(meet(X1,X2),meet(X3,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_32]), c_0_31])).
cnf(c_0_43,plain,(join(X1,join(X2,join(X1,X3)))=join(X2,join(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_19]), c_0_6]), c_0_6])).
cnf(c_0_44,plain,(join(X1,join(meet(X2,X1),meet(X2,X3)))=join(X1,meet(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34]), c_0_6]), c_0_15])).
cnf(c_0_45,plain,(meet(X1,meet(X2,X1))=meet(X2,X1)), inference(spm,[status(thm)],[c_0_32, c_0_1])).
cnf(c_0_46,hypothesis,(join(X1,meet(join(X2,X1),X3))=join(X1,meet(X2,X3))), inference(spm,[status(thm)],[c_0_24, c_0_1])).
cnf(c_0_47,plain,(meet(join(X1,X2),join(meet(X1,X3),X2))=join(meet(X1,X3),X2)), inference(spm,[status(thm)],[c_0_11, c_0_2])).
cnf(c_0_48,plain,(meet(X1,meet(join(X2,X1),join(X2,X3)))=meet(X1,join(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36]), c_0_32])).
cnf(c_0_49,hypothesis,(meet(X1,join(X2,join(meet(X1,X2),X3)))=meet(join(X3,X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_37]), c_0_29]), c_0_38])).
cnf(c_0_50,hypothesis,(join(X1,join(meet(join(X1,X2),X3),X4))=join(X1,join(meet(X2,X3),X4))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_39]), c_0_6])).
cnf(c_0_51,plain,(meet(join(X1,X2),join(X1,join(meet(X2,X3),X4)))=join(X1,join(meet(X2,X3),meet(X4,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41]), c_0_31]), c_0_42]), c_0_43])).
cnf(c_0_52,plain,(join(X1,join(meet(X2,X1),meet(X3,X2)))=join(X1,meet(X3,X2))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53,hypothesis,(join(meet(X1,X2),X3)=join(X3,meet(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_19]), c_0_24]), c_0_3]), c_0_45])).
cnf(c_0_54,plain,(join(X1,join(X2,X3))=join(X3,join(X2,X1))), inference(spm,[status(thm)],[c_0_29, c_0_2])).
cnf(c_0_55,plain,(meet(X1,meet(join(X2,X3),join(X3,X1)))=meet(X1,join(X3,X2))), inference(spm,[status(thm)],[c_0_48, c_0_21])).
cnf(c_0_56,hypothesis,(meet(join(X1,X2),join(X2,X3))=join(X2,meet(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_57,hypothesis,(meet(X1,join(X2,meet(X3,X1)))=join(meet(X1,X3),meet(X2,X1))), inference(spm,[status(thm)],[c_0_42, c_0_53])).
cnf(c_0_58,plain,(meet(X1,meet(X2,join(X2,meet(X2,X1))))=meet(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_7]), c_0_2]), c_0_12]), c_0_1])).
cnf(c_0_59,plain,(meet(X1,join(X1,meet(X1,X2)))=join(X1,meet(X1,X2))), inference(spm,[status(thm)],[c_0_10, c_0_7])).
cnf(c_0_60,plain,(join(meet(X1,X2),join(X3,meet(X3,X1)))=join(X3,meet(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_44]), c_0_53])).
cnf(c_0_61,plain,(join(meet(X1,X2),meet(X3,X1))=meet(X1,join(X3,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
cnf(c_0_62,plain,(meet(X1,join(X2,meet(X2,X1)))=meet(X2,X1)), inference(rw,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63,plain,(join(meet(X1,X2),meet(X1,X3))=meet(X1,join(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_53]), c_0_62]), c_0_61])).
cnf(c_0_64,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_63])]), ['proof']).
# SZS output end CNFRefutation.
EOF
