% Problem    : REL031-2 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n091.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:23:56 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4d
# and selection function SelectCQIPrecWNTNp.
#
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S2U
# and selection function SelectNewComplexAHPExceptRRHorn.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(converse(composition(X1,X2))=composition(converse(X2),converse(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_multiplicativity_10)).
cnf(c_0_1,axiom,(converse(converse(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_idempotence_8)).
cnf(c_0_2,axiom,(composition(X1,one)=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_identity_6)).
cnf(c_0_3,negated_conjecture,(join(composition(converse(sk2),sk2),one)=one), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_18)).
cnf(c_0_4,axiom,(join(X1,X2)=join(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux1_join_commutativity_1)).
cnf(c_0_5,axiom,(composition(join(X1,X2),X3)=join(composition(X1,X3),composition(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_distributivity_7)).
cnf(c_0_6,negated_conjecture,(join(composition(converse(sk1),sk1),one)=one), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_17)).
cnf(c_0_7,axiom,(join(X1,join(X2,X3))=join(join(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux2_join_associativity_2)).
cnf(c_0_8,axiom,(converse(join(X1,X2))=join(converse(X1),converse(X2))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_additivity_9)).
cnf(c_0_9,axiom,(composition(X1,composition(X2,X3))=composition(composition(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_associativity_5)).
cnf(c_0_10,negated_conjecture,(join(composition(converse(composition(sk1,sk2)),composition(sk1,sk2)),one)!=one), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_19)).
cnf(c_0_11,plain,(converse(composition(converse(X1),X2))=composition(converse(X2),X1)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_12,plain,(composition(converse(one),X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_2]), c_0_1])).
cnf(c_0_13,plain,(converse(one)=one), inference(spm,[status(thm)],[c_0_2, c_0_12])).
cnf(c_0_14,plain,(composition(one,X1)=X1), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_15,negated_conjecture,(join(one,composition(converse(sk2),sk2))=one), inference(rw,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_16,plain,(join(X1,composition(X2,X1))=composition(join(one,X2),X1)), inference(spm,[status(thm)],[c_0_5, c_0_14])).
cnf(c_0_17,negated_conjecture,(join(one,composition(converse(sk1),sk1))=one), inference(rw,[status(thm)],[c_0_6, c_0_4])).
cnf(c_0_18,negated_conjecture,(join(one,join(composition(converse(sk2),sk2),X1))=join(one,X1)), inference(spm,[status(thm)],[c_0_7, c_0_15])).
cnf(c_0_19,plain,(converse(join(converse(X1),X2))=join(X1,converse(X2))), inference(spm,[status(thm)],[c_0_8, c_0_1])).
cnf(c_0_20,negated_conjecture,(join(X1,composition(converse(sk1),composition(sk1,X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_9]), c_0_14])).
cnf(c_0_21,plain,(converse(composition(X1,converse(X2)))=composition(X2,converse(X1))), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_22,negated_conjecture,(join(one,composition(join(converse(sk2),X1),sk2))=join(one,composition(X1,sk2))), inference(spm,[status(thm)],[c_0_18, c_0_5])).
cnf(c_0_23,negated_conjecture,(join(X1,composition(X1,composition(converse(sk1),sk1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_1]), c_0_11]), c_0_21]), c_0_9])).
cnf(c_0_24,plain,(composition(converse(X1),composition(converse(X2),X3))=composition(converse(composition(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_9, c_0_0])).
cnf(c_0_25,negated_conjecture,(join(one,composition(converse(composition(sk1,sk2)),composition(sk1,sk2)))!=one), inference(rw,[status(thm)],[c_0_10, c_0_4])).
cnf(c_0_26,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_15]), c_0_24]), c_0_9]), c_0_25]), ['proof']).
# SZS output end CNFRefutation.
EOF
