% Problem    : GRP572-1 : TPTP v6.1.0. Bugfixed v2.7.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n099.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:58:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(identity=double_divide(X1,inverse(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', identity)).
cnf(c_0_1,axiom,(inverse(X1)=double_divide(X1,identity)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', inverse)).
cnf(c_0_2,axiom,(double_divide(double_divide(X1,double_divide(double_divide(X2,double_divide(X1,X3)),double_divide(X3,identity))),double_divide(identity,identity))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', single_axiom)).
cnf(c_0_3,negated_conjecture,(multiply(a,b)!=multiply(b,a)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_these_axioms_4)).
cnf(c_0_4,axiom,(multiply(X1,X2)=double_divide(double_divide(X2,X1),identity)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', multiply)).
cnf(c_0_5,plain,(double_divide(X1,double_divide(X1,identity))=identity), inference(rw,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_6,plain,(double_divide(double_divide(X1,identity),double_divide(identity,identity))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_5]), c_0_5])).
cnf(c_0_7,plain,(double_divide(double_divide(identity,double_divide(X1,double_divide(identity,identity))),double_divide(identity,identity))=double_divide(X1,identity)), inference(spm,[status(thm)],[c_0_2, c_0_6])).
cnf(c_0_8,plain,(double_divide(identity,identity)=identity), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_5]), c_0_6])).
cnf(c_0_9,plain,(double_divide(double_divide(X1,identity),identity)=X1), inference(rw,[status(thm)],[c_0_6, c_0_8])).
cnf(c_0_10,plain,(double_divide(double_divide(identity,double_divide(X1,identity)),identity)=double_divide(X1,identity)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8]), c_0_8])).
cnf(c_0_11,plain,(double_divide(identity,double_divide(X1,identity))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_9])).
cnf(c_0_12,plain,(double_divide(double_divide(double_divide(X1,double_divide(double_divide(X2,double_divide(X1,X3)),double_divide(X3,identity))),double_divide(double_divide(X4,X2),double_divide(double_divide(identity,identity),identity))),double_divide(identity,identity))=X4), inference(spm,[status(thm)],[c_0_2, c_0_2])).
cnf(c_0_13,plain,(double_divide(identity,X1)=double_divide(X1,identity)), inference(spm,[status(thm)],[c_0_11, c_0_9])).
cnf(c_0_14,plain,(double_divide(identity,double_divide(double_divide(X1,double_divide(double_divide(X2,double_divide(X1,X3)),double_divide(X3,identity))),double_divide(identity,double_divide(X4,X2))))=X4), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_8]), c_0_8]), c_0_8]), c_0_13]), c_0_13])).
cnf(c_0_15,plain,(double_divide(identity,double_divide(double_divide(identity,X1),double_divide(identity,double_divide(X2,X1))))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_8]), c_0_8]), c_0_13]), c_0_11])).
cnf(c_0_16,plain,(double_divide(identity,double_divide(identity,X1))=X1), inference(spm,[status(thm)],[c_0_11, c_0_13])).
cnf(c_0_17,plain,(double_divide(identity,double_divide(double_divide(X1,X2),X1))=double_divide(identity,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_15]), c_0_16])).
cnf(c_0_18,plain,(double_divide(double_divide(X1,X2),X1)=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_16])).
cnf(c_0_19,plain,(double_divide(identity,double_divide(X1,double_divide(double_divide(X2,double_divide(X1,X3)),double_divide(X3,identity))))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_8]), c_0_13])).
cnf(c_0_20,negated_conjecture,(double_divide(double_divide(b,a),identity)!=double_divide(double_divide(a,b),identity)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_4]), c_0_4])).
cnf(c_0_21,plain,(double_divide(X1,double_divide(X2,X1))=X2), inference(spm,[status(thm)],[c_0_18, c_0_18])).
cnf(c_0_22,plain,(double_divide(double_divide(X1,X2),X2)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_16]), c_0_18]), c_0_16])).
cnf(c_0_23,negated_conjecture,(double_divide(identity,double_divide(a,b))!=double_divide(identity,double_divide(b,a))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_13]), c_0_13])).
cnf(c_0_24,plain,(double_divide(X1,X2)=double_divide(X2,X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24])]), ['proof']).
# SZS output end CNFRefutation.
EOF
