% Problem    : RNG021-7 : TPTP v6.1.0. Released v1.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n115.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:49:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_B02_F1_AE_CS_SP_PI_S0S
# and selection function SelectComplexG.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(add(X1,add(X2,X3))=add(add(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', associativity_for_addition)).
cnf(c_0_1,axiom,(add(X1,additive_inverse(X1))=additive_identity), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', right_additive_inverse)).
cnf(c_0_2,axiom,(add(additive_identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', left_additive_identity)).
cnf(c_0_3,axiom,(additive_inverse(additive_inverse(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', additive_inverse_additive_inverse)).
cnf(c_0_4,axiom,(add(X1,X2)=add(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', commutativity_for_addition)).
cnf(c_0_5,negated_conjecture,(associator(add(u,v),x,y)!=add(associator(u,x,y),associator(v,x,y))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_linearised_form3)).
cnf(c_0_6,axiom,(associator(X1,X2,X3)=add(multiply(multiply(X1,X2),X3),additive_inverse(multiply(X1,multiply(X2,X3))))), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', associator)).
cnf(c_0_7,axiom,(multiply(add(X1,X2),X3)=add(multiply(X1,X3),multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', distribute2)).
cnf(c_0_8,plain,(add(X1,add(additive_inverse(X1),X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_9,plain,(add(additive_inverse(X1),add(X1,X2))=X2), inference(spm,[status(thm)],[c_0_8, c_0_3])).
cnf(c_0_10,plain,(add(additive_inverse(X1),add(X2,X1))=X2), inference(spm,[status(thm)],[c_0_9, c_0_4])).
cnf(c_0_11,plain,(add(X1,additive_inverse(add(X1,X2)))=additive_inverse(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_10]), c_0_4])).
cnf(c_0_12,negated_conjecture,(add(add(multiply(multiply(u,x),y),additive_inverse(multiply(u,multiply(x,y)))),add(multiply(multiply(v,x),y),additive_inverse(multiply(v,multiply(x,y)))))!=add(multiply(multiply(add(u,v),x),y),additive_inverse(multiply(add(u,v),multiply(x,y))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_6]), c_0_6]), c_0_6])).
cnf(c_0_13,plain,(additive_inverse(add(additive_inverse(X1),X2))=add(X1,additive_inverse(X2))), inference(spm,[status(thm)],[c_0_8, c_0_11])).
cnf(c_0_14,plain,(add(X1,add(X2,additive_inverse(X1)))=X2), inference(spm,[status(thm)],[c_0_8, c_0_4])).
cnf(c_0_15,negated_conjecture,(add(additive_inverse(multiply(u,multiply(x,y))),add(multiply(multiply(u,x),y),add(additive_inverse(multiply(v,multiply(x,y))),multiply(multiply(v,x),y))))!=add(additive_inverse(add(multiply(u,multiply(x,y)),multiply(v,multiply(x,y)))),add(multiply(multiply(u,x),y),multiply(multiply(v,x),y)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_4]), c_0_4]), c_0_4]), c_0_7]), c_0_7]), c_0_7]), c_0_0])).
cnf(c_0_16,plain,(add(additive_inverse(X1),X2)=additive_inverse(add(X1,additive_inverse(X2)))), inference(spm,[status(thm)],[c_0_3, c_0_13])).
cnf(c_0_17,plain,(add(X1,additive_inverse(X2))=additive_inverse(add(X2,additive_inverse(X1)))), inference(spm,[status(thm)],[c_0_11, c_0_14])).
cnf(c_0_18,negated_conjecture,(additive_inverse(add(multiply(u,multiply(x,y)),additive_inverse(add(multiply(multiply(u,x),y),additive_inverse(add(multiply(v,multiply(x,y)),additive_inverse(multiply(multiply(v,x),y))))))))!=additive_inverse(add(multiply(u,multiply(x,y)),add(multiply(v,multiply(x,y)),additive_inverse(add(multiply(multiply(u,x),y),multiply(multiply(v,x),y))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16]), c_0_16]), c_0_16]), c_0_0])).
cnf(c_0_19,plain,(add(X1,additive_inverse(add(X2,additive_inverse(X3))))=additive_inverse(add(X2,additive_inverse(add(X1,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_0]), c_0_16]), c_0_0]), c_0_16]), c_0_3])).
cnf(c_0_20,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_3])]), ['proof']).
# SZS output end CNFRefutation.
EOF
