% Problem    : GRP767-1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n185.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:53:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C48_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(product(X1,difference(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos03)).
cnf(c_0_1,axiom,(difference(X1,product(product(X1,X2),X3))=quotient(product(X2,product(X3,X1)),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos07)).
cnf(c_0_2,axiom,(product(quotient(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos06)).
cnf(c_0_3,axiom,(quotient(product(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos05)).
cnf(c_0_4,axiom,(product(one,X1)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos02)).
cnf(c_0_5,axiom,(difference(X1,product(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos04)).
cnf(c_0_6,axiom,(eta(X1)=product(i(X1),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos12)).
cnf(c_0_7,axiom,(i(X1)=difference(X1,one)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos09)).
cnf(c_0_8,axiom,(product(i(i(X1)),X2)=product(eta(X1),product(X1,X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos13)).
cnf(c_0_9,axiom,(product(i(X1),X1)=product(X1,j(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos11)).
cnf(c_0_10,axiom,(j(X1)=quotient(one,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos10)).
cnf(c_0_11,axiom,(product(X1,one)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos01)).
cnf(c_0_12,axiom,(product(X1,product(eta(X1),X2))=product(j(j(X1)),X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos14)).
cnf(c_0_13,axiom,(product(eta(X1),product(X2,X3))=product(product(eta(X1),X2),X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos16)).
cnf(c_0_14,axiom,(difference(product(X1,X2),product(X1,product(X2,X3)))=quotient(quotient(product(X3,product(X1,X2)),X2),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos08)).
cnf(c_0_15,negated_conjecture,(product(j(j(x0)),j(product(x1,x0)))!=j(x1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
cnf(c_0_16,plain,(product(X1,quotient(product(X2,product(X3,X1)),X1))=product(product(X1,X2),X3)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_17,plain,(product(product(X1,quotient(X2,product(X3,X1))),X3)=product(X1,quotient(X2,X1))), inference(spm,[status(thm)],[c_0_16, c_0_2])).
cnf(c_0_18,plain,(quotient(X1,X1)=one), inference(spm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_19,plain,(quotient(product(X1,quotient(X2,X1)),X3)=product(X1,quotient(X2,product(X3,X1)))), inference(spm,[status(thm)],[c_0_3, c_0_17])).
cnf(c_0_20,plain,(product(X1,quotient(X2,product(product(X1,quotient(X2,X1)),X1)))=one), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(quotient(one,quotient(X1,product(product(X2,quotient(X1,X2)),X2)))=X2), inference(spm,[status(thm)],[c_0_3, c_0_20])).
cnf(c_0_22,plain,(difference(quotient(X1,X2),X1)=X2), inference(spm,[status(thm)],[c_0_5, c_0_2])).
cnf(c_0_23,plain,(quotient(one,quotient(product(X1,X2),product(product(X2,X1),X2)))=X2), inference(spm,[status(thm)],[c_0_21, c_0_3])).
cnf(c_0_24,plain,(quotient(product(X1,X2),product(product(X2,X1),X2))=difference(X2,one)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,plain,(product(difference(X1,one),product(product(X1,X2),X1))=product(X2,X1)), inference(spm,[status(thm)],[c_0_2, c_0_24])).
cnf(c_0_26,plain,(product(X1,quotient(product(X2,X1),product(X3,X1)))=quotient(product(X1,X2),X3)), inference(spm,[status(thm)],[c_0_19, c_0_3])).
cnf(c_0_27,plain,(product(difference(X1,one),X1)=eta(X1)), inference(rw,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_28,plain,(product(difference(X1,one),product(X2,X1))=product(difference(X1,X2),X1)), inference(spm,[status(thm)],[c_0_25, c_0_0])).
cnf(c_0_29,plain,(quotient(product(X1,product(X2,quotient(X3,X1))),quotient(X3,X1))=difference(quotient(X3,X1),product(X3,X2))), inference(spm,[status(thm)],[c_0_1, c_0_2])).
cnf(c_0_30,plain,(quotient(product(X1,X2),quotient(X3,X1))=product(X1,quotient(product(X2,X1),X3))), inference(spm,[status(thm)],[c_0_26, c_0_2])).
cnf(c_0_31,plain,(product(product(difference(X1,one),X1),product(X1,X2))=product(difference(difference(X1,one),one),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_7]), c_0_7]), c_0_27])).
cnf(c_0_32,plain,(product(difference(X1,one),X1)=product(X1,quotient(one,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_7]), c_0_10])).
cnf(c_0_33,plain,(product(difference(X1,quotient(X2,X1)),X1)=product(difference(X1,one),X2)), inference(spm,[status(thm)],[c_0_28, c_0_2])).
cnf(c_0_34,plain,(product(X1,quotient(product(product(X2,quotient(X3,X1)),X1),X3))=difference(quotient(X3,X1),product(X3,X2))), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35,plain,(quotient(X1,one)=X1), inference(spm,[status(thm)],[c_0_11, c_0_2])).
cnf(c_0_36,plain,(quotient(X1,product(product(X2,quotient(X1,X2)),X2))=difference(X2,one)), inference(spm,[status(thm)],[c_0_5, c_0_20])).
cnf(c_0_37,plain,(product(product(X1,quotient(one,X1)),product(X1,X2))=product(difference(difference(X1,one),one),X2)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38,plain,(quotient(product(difference(X1,one),X2),X1)=difference(X1,quotient(X2,X1))), inference(spm,[status(thm)],[c_0_3, c_0_33])).
cnf(c_0_39,plain,(product(X1,product(product(X2,quotient(one,X1)),X1))=difference(quotient(one,X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_4])).
cnf(c_0_40,plain,(product(product(X1,quotient(X2,X1)),X1)=difference(difference(X1,one),X2)), inference(spm,[status(thm)],[c_0_22, c_0_36])).
cnf(c_0_41,plain,(product(product(X1,quotient(one,X1)),X1)=difference(difference(X1,one),one)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_11]), c_0_11])).
cnf(c_0_42,plain,(product(quotient(one,X1),quotient(one,quotient(one,X1)))=product(X1,quotient(one,X1))), inference(spm,[status(thm)],[c_0_32, c_0_22])).
cnf(c_0_43,plain,(product(X1,product(product(difference(X1,one),X1),X2))=product(quotient(one,quotient(one,X1)),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_10]), c_0_10]), c_0_27])).
cnf(c_0_44,plain,(difference(quotient(one,X1),quotient(X2,quotient(one,X1)))=product(X1,product(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_22]), c_0_30]), c_0_35])).
cnf(c_0_45,plain,(product(X1,product(difference(X1,X2),X1))=quotient(X2,quotient(one,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_22]), c_0_5])).
cnf(c_0_46,plain,(quotient(product(difference(X1,X2),product(X3,X1)),X1)=difference(X1,product(X2,X3))), inference(spm,[status(thm)],[c_0_1, c_0_0])).
cnf(c_0_47,plain,(product(product(product(difference(X1,one),X1),X2),X3)=product(product(difference(X1,one),X1),product(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_27]), c_0_27])).
cnf(c_0_48,plain,(product(product(X1,quotient(one,X1)),quotient(one,X1))=difference(X1,one)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_22])).
cnf(c_0_49,plain,(product(X1,product(product(X1,quotient(one,X1)),X2))=product(quotient(one,quotient(one,X1)),X2)), inference(rw,[status(thm)],[c_0_43, c_0_32])).
cnf(c_0_50,plain,(product(quotient(one,X1),product(X1,product(X2,X1)))=quotient(X2,quotient(one,X1))), inference(spm,[status(thm)],[c_0_0, c_0_44])).
cnf(c_0_51,plain,(quotient(quotient(X1,X2),quotient(one,X2))=product(X2,product(difference(X2,one),X1))), inference(spm,[status(thm)],[c_0_45, c_0_33])).
cnf(c_0_52,plain,(difference(product(X1,X2),product(X1,X3))=quotient(difference(X2,product(X3,X1)),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_0]), c_0_46])).
cnf(c_0_53,plain,(product(product(product(X1,quotient(one,X1)),X2),X3)=product(product(X1,quotient(one,X1)),product(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_32]), c_0_32])).
cnf(c_0_54,plain,(product(product(X1,quotient(one,X1)),quotient(one,quotient(one,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_42]), c_0_22])).
cnf(c_0_55,plain,(product(X1,product(X1,quotient(one,X1)))=quotient(one,quotient(one,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_11]), c_0_11])).
cnf(c_0_56,plain,(product(quotient(one,X1),product(X1,X2))=product(X1,product(difference(X1,one),X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_2]), c_0_51])).
cnf(c_0_57,plain,(quotient(difference(X1,product(difference(X2,X3),X2)),X2)=difference(product(X2,X1),X3)), inference(spm,[status(thm)],[c_0_52, c_0_0])).
cnf(c_0_58,plain,(product(X1,product(X1,quotient(one,product(product(X2,X1),X1))))=quotient(one,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_32]), c_0_19]), c_0_0])).
cnf(c_0_59,plain,(quotient(one,quotient(one,product(X1,quotient(one,X1))))=product(X1,quotient(one,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_42]), c_0_55])).
cnf(c_0_60,plain,(product(product(X1,quotient(one,X2)),X2)=product(product(X1,X2),difference(X2,one))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_56]), c_0_16])).
cnf(c_0_61,plain,(quotient(difference(X1,product(X2,quotient(one,X2))),X2)=difference(product(X2,X1),one)), inference(spm,[status(thm)],[c_0_57, c_0_32])).
cnf(c_0_62,plain,(product(X1,product(X1,quotient(one,product(X2,X1))))=quotient(one,quotient(X2,X1))), inference(spm,[status(thm)],[c_0_58, c_0_2])).
cnf(c_0_63,plain,(quotient(product(product(X1,quotient(one,X1)),product(X2,product(X3,X1))),X1)=difference(X1,product(product(quotient(one,quotient(one,X1)),X2),X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1, c_0_49]), c_0_53])).
cnf(c_0_64,plain,(product(product(X1,quotient(one,X1)),quotient(one,product(X1,quotient(one,X1))))=one), inference(spm,[status(thm)],[c_0_2, c_0_59])).
cnf(c_0_65,plain,(product(X1,product(product(X2,X1),difference(X1,one)))=difference(quotient(one,X1),X2)), inference(rw,[status(thm)],[c_0_39, c_0_60])).
cnf(c_0_66,plain,(difference(X1,product(X2,quotient(one,X2)))=product(difference(product(X2,X1),one),X2)), inference(spm,[status(thm)],[c_0_2, c_0_61])).
cnf(c_0_67,plain,(product(product(X1,X2),product(X2,quotient(one,X2)))=product(X1,quotient(one,quotient(one,X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_49]), c_0_3])).
cnf(c_0_68,plain,(difference(X1,quotient(one,quotient(X2,X1)))=product(X1,quotient(one,product(X2,X1)))), inference(spm,[status(thm)],[c_0_5, c_0_62])).
cnf(c_0_69,plain,(quotient(product(X1,quotient(one,quotient(one,X2))),product(X2,quotient(one,X2)))=product(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_55]), c_0_64]), c_0_4]), c_0_59]), c_0_53]), c_0_5])).
cnf(c_0_70,plain,(product(difference(product(X1,quotient(one,X2)),one),X1)=product(X2,product(X1,quotient(one,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_53]), c_0_0]), c_0_11]), c_0_66])).
cnf(c_0_71,plain,(product(quotient(X1,X2),quotient(one,quotient(one,X2)))=product(X1,product(X2,quotient(one,X2)))), inference(spm,[status(thm)],[c_0_67, c_0_2])).
cnf(c_0_72,plain,(product(X1,quotient(one,product(product(X2,X1),X1)))=difference(X1,quotient(one,X2))), inference(spm,[status(thm)],[c_0_68, c_0_3])).
cnf(c_0_73,plain,(product(difference(product(quotient(one,quotient(one,X1)),quotient(one,X2)),one),X1)=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_42]), c_0_71]), c_0_4]), c_0_3])).
cnf(c_0_74,plain,(quotient(X1,difference(X2,X1))=X2), inference(spm,[status(thm)],[c_0_3, c_0_0])).
cnf(c_0_75,plain,(difference(X1,product(quotient(one,quotient(one,X1)),X2))=product(product(X1,quotient(one,X1)),X2)), inference(spm,[status(thm)],[c_0_5, c_0_49])).
cnf(c_0_76,plain,(product(product(X1,quotient(one,X1)),quotient(one,X2))=product(X1,quotient(one,product(X2,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74]), c_0_75])).
cnf(c_0_77,negated_conjecture,(product(quotient(one,quotient(one,x0)),quotient(one,product(x1,x0)))!=quotient(one,x1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_10]), c_0_10]), c_0_10]), c_0_10])).
cnf(c_0_78,plain,(product(quotient(one,quotient(one,X1)),quotient(one,X2))=quotient(one,quotient(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_76]), c_0_62])).
cnf(c_0_79,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_78]), c_0_3])]), ['proof']).
# SZS output end CNFRefutation.
EOF
