% Problem    : GRP560-1 : TPTP v6.1.0. Bugfixed v2.7.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n142.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:28:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___209_C18_F1_AE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(divide(X1,inverse(divide(divide(X2,X3),divide(X1,X3))))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', single_axiom)).
cnf(c_0_1,negated_conjecture,(multiply(a,b)!=multiply(b,a)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_these_axioms_4)).
cnf(c_0_2,axiom,(multiply(X1,X2)=divide(X1,inverse(X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', multiply)).
cnf(c_0_3,plain,(divide(X1,inverse(divide(X2,divide(X1,inverse(divide(divide(X2,X3),divide(X4,X3)))))))=X4), inference(spm,[status(thm)],[c_0_0, c_0_0])).
cnf(c_0_4,plain,(divide(X1,inverse(divide(X2,X2)))=X1), inference(spm,[status(thm)],[c_0_3, c_0_0])).
cnf(c_0_5,plain,(divide(X1,inverse(divide(X2,X1)))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_4]), c_0_4])).
cnf(c_0_6,plain,(divide(inverse(divide(X1,X2)),inverse(X1))=X2), inference(spm,[status(thm)],[c_0_5, c_0_5])).
cnf(c_0_7,plain,(inverse(divide(divide(X1,X1),X2))=X2), inference(spm,[status(thm)],[c_0_4, c_0_6])).
cnf(c_0_8,plain,(divide(X1,divide(X1,X2))=X2), inference(spm,[status(thm)],[c_0_0, c_0_7])).
cnf(c_0_9,plain,(inverse(divide(X1,X2))=divide(X2,X1)), inference(spm,[status(thm)],[c_0_8, c_0_5])).
cnf(c_0_10,plain,(divide(X1,divide(X2,X2))=X1), inference(spm,[status(thm)],[c_0_4, c_0_7])).
cnf(c_0_11,plain,(divide(divide(X1,X2),inverse(X2))=X1), inference(rw,[status(thm)],[c_0_6, c_0_9])).
cnf(c_0_12,plain,(divide(divide(X1,X2),X1)=inverse(X2)), inference(spm,[status(thm)],[c_0_9, c_0_8])).
cnf(c_0_13,plain,(divide(divide(X1,X1),X2)=inverse(X2)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_14,plain,(divide(inverse(X1),inverse(X2))=divide(X2,X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15,plain,(inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_13]), c_0_13])).
cnf(c_0_16,negated_conjecture,(divide(b,inverse(a))!=divide(a,inverse(b))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_1, c_0_2]), c_0_2])).
cnf(c_0_17,plain,(divide(X1,inverse(X2))=divide(X2,inverse(X1))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_18,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17])]), ['proof']).
# SZS output end CNFRefutation.
EOF
