% Problem    : RNG026-6 : TPTP v6.1.0. Released v1.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n100.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:05:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(add(X1,add(X2,X3))=add(add(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', associativity_for_addition)).
cnf(c_0_1,axiom,(add(X1,additive_inverse(X1))=additive_identity), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', right_additive_inverse)).
cnf(c_0_2,axiom,(add(additive_identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', left_additive_identity)).
cnf(c_0_3,axiom,(additive_inverse(additive_inverse(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', additive_inverse_additive_inverse)).
cnf(c_0_4,axiom,(add(X1,X2)=add(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', commutativity_for_addition)).
cnf(c_0_5,negated_conjecture,(add(add(associator(multiply(a,b),c,d),associator(a,b,multiply(c,d))),additive_inverse(add(add(associator(a,multiply(b,c),d),multiply(a,associator(b,c,d))),multiply(associator(a,b,c),d))))!=additive_identity), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_teichmuller_identity)).
cnf(c_0_6,axiom,(associator(X1,X2,X3)=add(multiply(multiply(X1,X2),X3),additive_inverse(multiply(X1,multiply(X2,X3))))), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', associator)).
cnf(c_0_7,axiom,(multiply(X1,add(X2,X3))=add(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', distribute1)).
cnf(c_0_8,axiom,(multiply(add(X1,X2),X3)=add(multiply(X1,X3),multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', distribute2)).
cnf(c_0_9,axiom,(multiply(X1,additive_identity)=additive_identity), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', right_multiplicative_zero)).
cnf(c_0_10,axiom,(multiply(additive_identity,X1)=additive_identity), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', left_multiplicative_zero)).
cnf(c_0_11,axiom,(add(X1,additive_identity)=X1), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', right_additive_identity)).
cnf(c_0_12,plain,(add(X1,add(additive_inverse(X1),X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_13,plain,(add(additive_inverse(X1),add(X1,X2))=X2), inference(spm,[status(thm)],[c_0_12, c_0_3])).
cnf(c_0_14,plain,(add(additive_inverse(X1),add(X2,X1))=X2), inference(spm,[status(thm)],[c_0_13, c_0_4])).
cnf(c_0_15,negated_conjecture,(add(add(add(multiply(multiply(multiply(a,b),c),d),additive_inverse(multiply(multiply(a,b),multiply(c,d)))),add(multiply(multiply(a,b),multiply(c,d)),additive_inverse(multiply(a,multiply(b,multiply(c,d)))))),additive_inverse(add(add(add(multiply(multiply(a,multiply(b,c)),d),additive_inverse(multiply(a,multiply(multiply(b,c),d)))),multiply(a,add(multiply(multiply(b,c),d),additive_inverse(multiply(b,multiply(c,d)))))),multiply(add(multiply(multiply(a,b),c),additive_inverse(multiply(a,multiply(b,c)))),d))))!=additive_identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_6]), c_0_6]), c_0_6]), c_0_6]), c_0_6])).
cnf(c_0_16,plain,(add(X1,additive_inverse(add(X1,X2)))=additive_inverse(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_14]), c_0_4])).
cnf(c_0_17,negated_conjecture,(add(multiply(multiply(a,b),multiply(c,d)),add(additive_inverse(multiply(a,multiply(b,multiply(c,d)))),add(multiply(multiply(multiply(a,b),c),d),add(additive_inverse(multiply(multiply(a,b),multiply(c,d))),additive_inverse(add(multiply(add(multiply(multiply(a,b),c),additive_inverse(multiply(a,multiply(b,c)))),d),add(multiply(a,add(multiply(multiply(b,c),d),additive_inverse(multiply(b,multiply(c,d))))),add(multiply(multiply(a,multiply(b,c)),d),additive_inverse(multiply(a,multiply(multiply(b,c),d)))))))))))!=additive_identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_4]), c_0_4]), c_0_4]), c_0_0]), c_0_0]), c_0_0]), c_0_0])).
cnf(c_0_18,plain,(additive_inverse(add(X1,X2))=add(additive_inverse(X1),additive_inverse(X2))), inference(spm,[status(thm)],[c_0_13, c_0_16])).
cnf(c_0_19,plain,(add(multiply(X1,add(X2,X3)),X4)=add(multiply(X1,X2),add(multiply(X1,X3),X4))), inference(spm,[status(thm)],[c_0_0, c_0_7])).
cnf(c_0_20,negated_conjecture,(add(multiply(multiply(a,b),multiply(c,d)),add(additive_inverse(multiply(a,multiply(b,multiply(c,d)))),add(multiply(multiply(multiply(a,b),c),d),add(additive_inverse(multiply(multiply(a,b),multiply(c,d))),add(additive_inverse(multiply(add(multiply(multiply(a,b),c),additive_inverse(multiply(a,multiply(b,c)))),d)),add(multiply(a,multiply(multiply(b,c),d)),add(additive_inverse(multiply(multiply(a,multiply(b,c)),d)),additive_inverse(multiply(a,add(multiply(multiply(b,c),d),additive_inverse(multiply(b,multiply(c,d)))))))))))))!=additive_identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_18]), c_0_18]), c_0_3]), c_0_4]), c_0_4]), c_0_0])).
cnf(c_0_21,plain,(add(X1,add(X2,X3))=add(X2,add(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_4]), c_0_0])).
cnf(c_0_22,plain,(add(multiply(add(X1,X2),X3),X4)=add(multiply(X1,X3),add(multiply(X2,X3),X4))), inference(spm,[status(thm)],[c_0_0, c_0_8])).
cnf(c_0_23,plain,(add(X1,add(X2,additive_inverse(X1)))=X2), inference(spm,[status(thm)],[c_0_12, c_0_4])).
cnf(c_0_24,plain,(add(multiply(X1,X2),add(multiply(X1,additive_inverse(X2)),X3))=X3), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_1]), c_0_9]), c_0_2])).
cnf(c_0_25,negated_conjecture,(add(multiply(multiply(a,b),multiply(c,d)),add(multiply(a,multiply(multiply(b,c),d)),add(multiply(multiply(multiply(a,b),c),d),add(additive_inverse(multiply(a,multiply(b,multiply(c,d)))),add(additive_inverse(multiply(multiply(a,b),multiply(c,d))),add(additive_inverse(multiply(multiply(a,multiply(b,c)),d)),add(additive_inverse(multiply(add(multiply(multiply(a,b),c),additive_inverse(multiply(a,multiply(b,c)))),d)),additive_inverse(multiply(a,add(multiply(multiply(b,c),d),additive_inverse(multiply(b,multiply(c,d)))))))))))))!=additive_identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_21]), c_0_21]), c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_26,plain,(additive_inverse(multiply(X1,add(X2,X3)))=add(additive_inverse(multiply(X1,X2)),additive_inverse(multiply(X1,X3)))), inference(spm,[status(thm)],[c_0_18, c_0_7])).
cnf(c_0_27,plain,(additive_inverse(multiply(add(X1,X2),X3))=add(additive_inverse(multiply(X1,X3)),additive_inverse(multiply(X2,X3)))), inference(spm,[status(thm)],[c_0_18, c_0_8])).
cnf(c_0_28,plain,(add(multiply(X1,X2),add(multiply(additive_inverse(X1),X2),X3))=X3), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1]), c_0_10]), c_0_2])).
cnf(c_0_29,plain,(additive_inverse(multiply(X1,X2))=multiply(X1,additive_inverse(X2))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30,negated_conjecture,(add(multiply(multiply(a,b),multiply(c,d)),add(multiply(a,multiply(multiply(b,c),d)),add(multiply(multiply(multiply(a,b),c),d),add(additive_inverse(multiply(a,multiply(b,multiply(c,d)))),add(additive_inverse(multiply(multiply(a,b),multiply(c,d))),add(additive_inverse(multiply(a,multiply(multiply(b,c),d))),add(additive_inverse(multiply(multiply(a,multiply(b,c)),d)),add(additive_inverse(multiply(multiply(multiply(a,b),c),d)),add(additive_inverse(multiply(additive_inverse(multiply(a,multiply(b,c))),d)),additive_inverse(multiply(a,additive_inverse(multiply(b,multiply(c,d))))))))))))))!=additive_identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26]), c_0_21]), c_0_21]), c_0_27]), c_0_0])).
cnf(c_0_31,plain,(multiply(additive_inverse(X1),additive_inverse(X2))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_1]), c_0_11]), c_0_29])).
cnf(c_0_32,negated_conjecture,(add(multiply(a,multiply(b,multiply(c,d))),add(multiply(multiply(a,b),multiply(c,d)),add(multiply(a,multiply(multiply(b,c),d)),add(multiply(multiply(multiply(a,b),c),d),add(multiply(a,multiply(b,multiply(c,additive_inverse(d)))),add(multiply(multiply(a,b),multiply(c,additive_inverse(d))),add(multiply(a,multiply(multiply(b,c),additive_inverse(d))),multiply(add(multiply(a,multiply(b,c)),add(multiply(multiply(a,b),c),multiply(a,multiply(b,additive_inverse(c))))),additive_inverse(d)))))))))!=additive_identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_3]), c_0_4]), c_0_21]), c_0_8]), c_0_21]), c_0_8]), c_0_21]), c_0_21]), c_0_21]), c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_33,plain,(multiply(X1,additive_inverse(X2))=multiply(additive_inverse(X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_28]), c_0_29])).
cnf(c_0_34,plain,(multiply(additive_inverse(X1),multiply(X2,additive_inverse(X3)))=multiply(X1,multiply(X2,X3))), inference(spm,[status(thm)],[c_0_31, c_0_29])).
cnf(c_0_35,negated_conjecture,(add(multiply(a,multiply(b,multiply(c,d))),add(multiply(multiply(a,b),multiply(c,d)),add(multiply(a,multiply(multiply(b,c),d)),add(multiply(multiply(multiply(a,b),c),d),add(multiply(a,multiply(b,multiply(c,additive_inverse(d)))),add(multiply(multiply(a,b),multiply(c,additive_inverse(d))),add(multiply(a,multiply(multiply(b,c),additive_inverse(d))),multiply(add(multiply(a,multiply(b,c)),add(multiply(multiply(a,b),c),multiply(a,multiply(additive_inverse(b),c)))),additive_inverse(d)))))))))!=additive_identity), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36,plain,(multiply(X1,multiply(X2,additive_inverse(X3)))=multiply(additive_inverse(X1),multiply(X2,X3))), inference(spm,[status(thm)],[c_0_33, c_0_29])).
cnf(c_0_37,plain,(multiply(additive_inverse(X1),multiply(X2,X3))=multiply(X1,multiply(additive_inverse(X2),X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_33]), c_0_3])).
cnf(c_0_38,plain,(add(multiply(X1,X2),add(X3,multiply(additive_inverse(X1),X2)))=X3), inference(spm,[status(thm)],[c_0_28, c_0_4])).
cnf(c_0_39,plain,(add(X1,add(X2,X3))=add(X3,add(X1,X2))), inference(spm,[status(thm)],[c_0_4, c_0_0])).
cnf(c_0_40,negated_conjecture,(add(multiply(a,multiply(b,multiply(c,d))),add(multiply(multiply(a,b),multiply(c,d)),add(multiply(a,multiply(multiply(b,c),d)),add(multiply(multiply(multiply(a,b),c),d),add(multiply(additive_inverse(a),multiply(multiply(b,c),d)),add(multiply(a,multiply(b,multiply(c,additive_inverse(d)))),add(multiply(multiply(a,b),multiply(c,additive_inverse(d))),multiply(multiply(multiply(a,b),c),additive_inverse(d)))))))))!=additive_identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36]), c_0_21]), c_0_21]), c_0_37]), c_0_38])).
cnf(c_0_41,plain,(add(multiply(X1,X2),add(X3,add(multiply(additive_inverse(X1),X2),X4)))=add(X3,X4)), inference(spm,[status(thm)],[c_0_21, c_0_28])).
cnf(c_0_42,plain,(add(multiply(X1,X2),add(X3,add(X4,multiply(X1,additive_inverse(X2)))))=add(X3,X4)), inference(spm,[status(thm)],[c_0_24, c_0_39])).
cnf(c_0_43,plain,(add(multiply(X1,X2),add(X3,multiply(X1,additive_inverse(X2))))=X3), inference(spm,[status(thm)],[c_0_24, c_0_4])).
cnf(c_0_44,negated_conjecture,(add(multiply(a,multiply(b,multiply(c,d))),add(multiply(multiply(a,b),multiply(c,d)),add(multiply(a,multiply(b,multiply(c,additive_inverse(d)))),multiply(multiply(a,b),multiply(c,additive_inverse(d))))))!=additive_identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_45,plain,(add(multiply(X1,multiply(X2,X3)),add(X4,multiply(X1,multiply(X2,additive_inverse(X3)))))=X4), inference(spm,[status(thm)],[c_0_43, c_0_29])).
cnf(c_0_46,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45]), c_0_7]), c_0_7]), c_0_7]), c_0_1]), c_0_9]), c_0_9]), c_0_9])]), ['proof']).
# SZS output end CNFRefutation.
EOF
