% Problem    : LCL163-1 : TPTP v6.1.0. Released v1.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n103.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:44:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C48_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(xor(X1,falsehood)=X1), file('/export/starexec/sandbox/benchmark/Axioms/LCL002-0.ax', axiom_2)).
cnf(c_0_1,axiom,(xor(X1,X2)=xor(X2,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', xor_commutativity)).
cnf(c_0_2,axiom,(xor(X1,xor(truth,X2))=xor(xor(X1,truth),X2)), file('/export/starexec/sandbox/benchmark/Axioms/LCL002-0.ax', axiom_7)).
cnf(c_0_3,axiom,(xor(X1,X1)=falsehood), file('/export/starexec/sandbox/benchmark/Axioms/LCL002-0.ax', axiom_3)).
cnf(c_0_4,axiom,(and_star(xor(and_star(xor(truth,X1),X2),truth),X2)=and_star(xor(and_star(xor(truth,X2),X1),truth),X1)), file('/export/starexec/sandbox/benchmark/Axioms/LCL002-0.ax', axiom_8)).
cnf(c_0_5,axiom,(and_star(X1,X2)=and_star(X2,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', and_star_commutativity)).
cnf(c_0_6,negated_conjecture,(implies(implies(x,y),y)!=implies(implies(y,x),x)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_wajsberg_axiom)).
cnf(c_0_7,axiom,(implies(X1,X2)=xor(truth,and_star(X1,xor(truth,X2)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', implies_definition)).
cnf(c_0_8,plain,(xor(falsehood,X1)=X1), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_9,plain,(xor(X1,xor(truth,xor(truth,X2)))=xor(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_2]), c_0_3]), c_0_0]), c_0_2])).
cnf(c_0_10,plain,(and_star(X1,xor(truth,and_star(xor(truth,X2),X1)))=and_star(X2,xor(truth,and_star(xor(truth,X1),X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_1]), c_0_5]), c_0_1]), c_0_5])).
cnf(c_0_11,plain,(xor(truth,xor(truth,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_8])).
cnf(c_0_12,plain,(xor(X1,xor(X2,truth))=xor(X2,xor(truth,X1))), inference(spm,[status(thm)],[c_0_2, c_0_1])).
cnf(c_0_13,negated_conjecture,(xor(truth,and_star(xor(truth,and_star(y,xor(truth,x))),xor(truth,x)))!=xor(truth,and_star(xor(truth,and_star(x,xor(truth,y))),xor(truth,y)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_7]), c_0_7]), c_0_7]), c_0_7])).
cnf(c_0_14,plain,(and_star(xor(truth,X1),xor(truth,and_star(xor(truth,X2),xor(truth,X1))))=and_star(X2,xor(truth,and_star(X1,X2)))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15,plain,(xor(X1,xor(truth,X2))=xor(X2,xor(truth,X1))), inference(spm,[status(thm)],[c_0_12, c_0_1])).
cnf(c_0_16,negated_conjecture,(xor(truth,and_star(xor(truth,x),xor(truth,and_star(y,xor(truth,x)))))!=xor(truth,and_star(xor(truth,y),xor(truth,and_star(x,xor(truth,y)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_5]), c_0_5])).
cnf(c_0_17,plain,(and_star(xor(truth,X1),xor(truth,and_star(X2,xor(truth,X1))))=and_star(xor(truth,X2),xor(truth,and_star(X1,xor(truth,X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_3]), c_0_0])).
cnf(c_0_18,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17])]), ['proof']).
# SZS output end CNFRefutation.
EOF
