% Problem    : LAT044-1 : TPTP v6.1.0. Released v2.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n192.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:24:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(join(X1,meet(X1,X2))=X1), file('/export/starexec/sandbox/benchmark/Axioms/LAT001-0.ax', absorption2)).
cnf(c_0_1,axiom,(meet(X1,X2)=meet(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/LAT001-0.ax', commutativity_of_meet)).
cnf(c_0_2,axiom,(join(X1,meet(complement(X1),join(X1,X2)))=join(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', orthomodular_law)).
cnf(c_0_3,axiom,(join(X1,X2)=join(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/LAT001-0.ax', commutativity_of_join)).
cnf(c_0_4,axiom,(join(join(X1,X2),X3)=join(X1,join(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/LAT001-0.ax', associativity_of_join)).
cnf(c_0_5,axiom,(complement(join(X1,X2))=meet(complement(X1),complement(X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', compatibility1)).
cnf(c_0_6,axiom,(complement(complement(X1))=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', invertability3)).
cnf(c_0_7,axiom,(complement(meet(X1,X2))=join(complement(X1),complement(X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', compatibility2)).
cnf(c_0_8,negated_conjecture,(join(meet(complement(a),join(a,b)),join(complement(b),meet(a,b)))!=n1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_weak_orthomodular_law)).
cnf(c_0_9,axiom,(join(complement(X1),X1)=n1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', invertability1)).
cnf(c_0_10,plain,(join(X1,meet(X2,X1))=X1), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_11,plain,(join(X1,meet(complement(X1),join(X2,X1)))=join(X2,X1)), inference(spm,[status(thm)],[c_0_2, c_0_3])).
cnf(c_0_12,plain,(join(X1,join(meet(X2,X1),X3))=join(X1,X3)), inference(spm,[status(thm)],[c_0_4, c_0_10])).
cnf(c_0_13,plain,(join(meet(X1,X2),meet(X1,complement(meet(X1,X2))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_0]), c_0_1])).
cnf(c_0_14,plain,(join(X1,meet(X2,complement(meet(X2,X1))))=join(X1,X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_15,plain,(join(complement(X1),meet(complement(X2),join(X2,X1)))=complement(meet(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_5]), c_0_6]), c_0_7])).
cnf(c_0_16,negated_conjecture,(join(meet(a,b),join(complement(b),meet(join(a,b),complement(a))))!=n1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_1]), c_0_3]), c_0_3]), c_0_4])).
cnf(c_0_17,plain,(join(complement(X1),meet(join(X2,X1),complement(X2)))=complement(meet(X1,X2))), inference(spm,[status(thm)],[c_0_15, c_0_1])).
cnf(c_0_18,plain,(join(X1,complement(X1))=n1), inference(rw,[status(thm)],[c_0_9, c_0_3])).
cnf(c_0_19,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17]), c_0_1]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation.
EOF
