% Problem    : ROB005-1 : TPTP v6.1.0. Released v1.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n140.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:22:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(negate(add(negate(add(X1,X2)),negate(add(X1,negate(X2)))))=X1), file('/export/starexec/sandbox/benchmark/Axioms/ROB001-0.ax', robbins_axiom)).
cnf(c_0_1,hypothesis,(add(c,c)=c), file('/export/starexec/sandbox/benchmark/theBenchmark.p', idempotence)).
cnf(c_0_2,axiom,(add(add(X1,X2),X3)=add(X1,add(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/ROB001-0.ax', associativity_of_add)).
cnf(c_0_3,axiom,(add(X1,X2)=add(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/ROB001-0.ax', commutativity_of_add)).
cnf(c_0_4,negated_conjecture,(add(negate(add(a,negate(b))),negate(add(negate(a),negate(b))))!=b), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_huntingtons_axiom)).
cnf(c_0_5,hypothesis,(negate(add(negate(c),negate(add(c,negate(c)))))=c), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_6,plain,(add(X1,add(X2,X3))=add(X2,add(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_3]), c_0_2])).
cnf(c_0_7,hypothesis,(negate(add(c,negate(add(c,add(negate(c),negate(c))))))=negate(c)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_5]), c_0_6]), c_0_3])).
cnf(c_0_8,hypothesis,(add(c,add(c,X1))=add(c,X1)), inference(spm,[status(thm)],[c_0_2, c_0_1])).
cnf(c_0_9,plain,(negate(add(negate(add(X1,X2)),negate(add(negate(X2),X1))))=X1), inference(spm,[status(thm)],[c_0_0, c_0_3])).
cnf(c_0_10,hypothesis,(negate(add(negate(c),negate(add(c,add(negate(c),negate(c))))))=c), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_7]), c_0_8]), c_0_3])).
cnf(c_0_11,plain,(negate(add(negate(add(X1,X2)),negate(add(X2,negate(X1)))))=X2), inference(spm,[status(thm)],[c_0_0, c_0_3])).
cnf(c_0_12,hypothesis,(negate(add(c,add(negate(c),negate(c))))=negate(add(c,negate(c)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_3]), c_0_7]), c_0_3])).
cnf(c_0_13,hypothesis,(negate(add(negate(add(X1,c)),negate(add(X1,add(negate(c),negate(add(c,negate(c))))))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_5]), c_0_3])).
cnf(c_0_14,hypothesis,(negate(add(negate(add(c,X1)),negate(add(c,add(X1,negate(c))))))=add(c,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_8]), c_0_2])).
cnf(c_0_15,hypothesis,(negate(add(c,negate(add(c,negate(c)))))=negate(c)), inference(rw,[status(thm)],[c_0_7, c_0_12])).
cnf(c_0_16,hypothesis,(negate(add(negate(c),negate(add(c,add(negate(c),negate(add(c,negate(c))))))))=c), inference(spm,[status(thm)],[c_0_13, c_0_1])).
cnf(c_0_17,hypothesis,(add(c,negate(add(c,negate(c))))=c), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_3]), c_0_16])).
cnf(c_0_18,hypothesis,(add(c,add(X1,negate(add(c,negate(c)))))=add(X1,c)), inference(spm,[status(thm)],[c_0_6, c_0_17])).
cnf(c_0_19,plain,(add(X1,add(X2,X3))=add(X3,add(X1,X2))), inference(spm,[status(thm)],[c_0_3, c_0_2])).
cnf(c_0_20,hypothesis,(add(X1,negate(add(c,negate(c))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_18]), c_0_2]), c_0_3]), c_0_13])).
cnf(c_0_21,plain,(negate(add(X1,negate(add(negate(X2),add(X1,negate(add(X1,X2)))))))=negate(add(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_0]), c_0_19]), c_0_3]), c_0_3])).
cnf(c_0_22,hypothesis,(add(negate(add(c,negate(c))),X1)=X1), inference(spm,[status(thm)],[c_0_3, c_0_20])).
cnf(c_0_23,plain,(negate(add(X1,negate(add(negate(add(X1,X2)),negate(negate(add(X1,negate(X2))))))))=negate(add(X1,X2))), inference(spm,[status(thm)],[c_0_0, c_0_0])).
cnf(c_0_24,hypothesis,(negate(negate(add(negate(X1),negate(X1))))=negate(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_22]), c_0_22]), c_0_22])).
cnf(c_0_25,plain,(negate(add(negate(add(X1,X2)),negate(add(X1,add(negate(add(X2,X3)),negate(add(X2,negate(X3))))))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_0]), c_0_3])).
cnf(c_0_26,hypothesis,(negate(negate(add(negate(X1),negate(negate(negate(X1))))))=negate(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_22]), c_0_22]), c_0_22]), c_0_22])).
cnf(c_0_27,hypothesis,(negate(negate(add(X1,X1)))=X1), inference(spm,[status(thm)],[c_0_24, c_0_0])).
cnf(c_0_28,plain,(negate(add(negate(add(X1,X2)),negate(add(negate(X1),X2))))=X2), inference(spm,[status(thm)],[c_0_9, c_0_3])).
cnf(c_0_29,hypothesis,(negate(add(X1,negate(X1)))=negate(add(c,negate(c)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_22]), c_0_22]), c_0_0]), c_0_3])).
cnf(c_0_30,hypothesis,(negate(negate(add(X1,negate(negate(X1)))))=X1), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31,hypothesis,(negate(negate(X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_20]), c_0_30])).
cnf(c_0_32,negated_conjecture,(add(negate(add(a,negate(b))),negate(add(negate(b),negate(a))))!=b), inference(rw,[status(thm)],[c_0_4, c_0_3])).
cnf(c_0_33,hypothesis,(add(negate(add(X1,X2)),negate(add(X2,negate(X1))))=negate(X2)), inference(spm,[status(thm)],[c_0_31, c_0_11])).
cnf(c_0_34,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_31])]), ['proof']).
# SZS output end CNFRefutation.
EOF
