% Problem    : GRP122-1 : TPTP v6.1.0. Released v1.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n134.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:04:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___209_C18_F1_AE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(multiply(X1,multiply(multiply(X1,multiply(multiply(X1,X1),multiply(X2,X3))),multiply(X3,multiply(X3,X3))))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', single_axiom)).
cnf(c_0_1,axiom,(multiply(identity,identity)=identity), file('/export/starexec/sandbox/benchmark/theBenchmark.p', single_axiom2)).
cnf(c_0_2,negated_conjecture,(multiply(multiply(a,b),c)!=multiply(a,multiply(b,c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_order3)).
cnf(c_0_3,plain,(multiply(X1,multiply(multiply(X1,multiply(multiply(X1,X1),multiply(X2,identity))),identity))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_1])).
cnf(c_0_4,plain,(multiply(identity,multiply(multiply(identity,multiply(identity,multiply(X1,identity))),identity))=X1), inference(spm,[status(thm)],[c_0_3, c_0_1])).
cnf(c_0_5,plain,(multiply(multiply(X1,X1),multiply(multiply(multiply(X1,X1),multiply(X1,X1)),multiply(X2,identity)))=multiply(X1,multiply(multiply(X1,X2),identity))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_3]), c_0_1]), c_0_1])).
cnf(c_0_6,plain,(multiply(identity,multiply(multiply(identity,X1),identity))=multiply(identity,multiply(identity,multiply(X1,identity)))), inference(spm,[status(thm)],[c_0_4, c_0_4])).
cnf(c_0_7,plain,(multiply(multiply(X1,X1),multiply(multiply(X1,multiply(multiply(X1,X2),identity)),identity))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_5]), c_0_1]), c_0_1])).
cnf(c_0_8,plain,(multiply(identity,multiply(multiply(identity,multiply(identity,multiply(X1,X2))),multiply(X2,multiply(X2,X2))))=X1), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_9,plain,(multiply(identity,multiply(identity,multiply(identity,multiply(multiply(X1,identity),identity))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_6]), c_0_6])).
cnf(c_0_10,plain,(multiply(X1,multiply(multiply(X1,multiply(multiply(X1,X1),identity)),identity))=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_1]), c_0_1])).
cnf(c_0_11,plain,(multiply(multiply(identity,multiply(identity,multiply(X1,X2))),multiply(X2,multiply(X2,X2)))=multiply(identity,multiply(identity,multiply(multiply(X1,identity),identity)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_1]), c_0_6])).
cnf(c_0_12,plain,(multiply(multiply(identity,X1),identity)=multiply(identity,multiply(X1,identity))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_6]), c_0_1]), c_0_6]), c_0_6]), c_0_9])).
cnf(c_0_13,plain,(multiply(multiply(X1,X1),multiply(multiply(multiply(X1,X1),multiply(X1,X1)),identity))=multiply(X1,multiply(multiply(X1,identity),identity))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_10]), c_0_1]), c_0_1])).
cnf(c_0_14,plain,(multiply(multiply(identity,X1),multiply(multiply(X2,multiply(X2,X2)),multiply(multiply(X2,multiply(X2,X2)),multiply(X2,multiply(X2,X2)))))=multiply(identity,multiply(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_11]), c_0_9]), c_0_12]), c_0_12]), c_0_12]), c_0_12]), c_0_9])).
cnf(c_0_15,plain,(multiply(X1,multiply(multiply(X1,multiply(X1,multiply(multiply(X1,identity),identity))),identity))=multiply(multiply(X1,X1),multiply(X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_13]), c_0_1]), c_0_1])).
cnf(c_0_16,plain,(multiply(identity,multiply(multiply(X1,multiply(X1,X1)),multiply(multiply(X1,multiply(X1,X1)),multiply(X1,multiply(X1,X1)))))=multiply(identity,multiply(identity,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_1]), c_0_1]), c_0_1]), c_0_1]), c_0_1]), c_0_1]), c_0_1]), c_0_1])).
cnf(c_0_17,plain,(multiply(multiply(X1,multiply(multiply(X1,X1),multiply(X2,X3))),multiply(X3,multiply(X3,X3)))=multiply(multiply(X1,X1),multiply(multiply(X1,multiply(X2,identity)),identity))), inference(spm,[status(thm)],[c_0_7, c_0_0])).
cnf(c_0_18,plain,(multiply(multiply(X1,multiply(X1,X1)),multiply(multiply(X1,multiply(X1,X1)),multiply(X1,multiply(X1,X1))))=multiply(identity,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_16]), c_0_1]), c_0_12]), c_0_12]), c_0_12]), c_0_6]), c_0_6]), c_0_9])).
cnf(c_0_19,plain,(multiply(identity,multiply(multiply(identity,multiply(X1,multiply(X2,multiply(X2,X2)))),X2))=multiply(identity,multiply(identity,multiply(multiply(X1,identity),identity)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_14]), c_0_1]), c_0_1]), c_0_12])).
cnf(c_0_20,plain,(multiply(multiply(identity,X1),multiply(identity,X2))=multiply(identity,multiply(X1,X2))), inference(rw,[status(thm)],[c_0_14, c_0_18])).
cnf(c_0_21,plain,(multiply(multiply(identity,multiply(X1,multiply(X2,multiply(X2,X2)))),X2)=multiply(identity,multiply(multiply(X1,identity),identity))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_19]), c_0_1]), c_0_6]), c_0_6]), c_0_9])).
cnf(c_0_22,plain,(multiply(identity,multiply(multiply(identity,multiply(identity,multiply(multiply(X1,identity),identity))),X2))=multiply(X1,multiply(identity,X2))), inference(spm,[status(thm)],[c_0_20, c_0_9])).
cnf(c_0_23,plain,(multiply(X1,multiply(multiply(X1,multiply(multiply(X1,X1),multiply(identity,multiply(X2,identity)))),identity))=multiply(identity,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_12]), c_0_1]), c_0_1])).
cnf(c_0_24,plain,(multiply(multiply(X1,multiply(identity,multiply(X2,multiply(X2,X2)))),X2)=multiply(multiply(X1,identity),identity)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_12]), c_0_6]), c_0_12]), c_0_12]), c_0_9])).
cnf(c_0_25,plain,(multiply(X1,multiply(multiply(X1,multiply(multiply(X1,X1),X2)),identity))=multiply(identity,multiply(identity,multiply(identity,multiply(X2,identity))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_3]), c_0_1])).
cnf(c_0_26,plain,(multiply(multiply(multiply(X1,multiply(identity,multiply(identity,multiply(identity,X2)))),identity),identity)=multiply(multiply(multiply(X1,identity),identity),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_24]), c_0_20]), c_0_20]), c_0_18])).
cnf(c_0_27,plain,(multiply(multiply(X1,multiply(multiply(X1,X1),X2)),identity)=multiply(multiply(X1,X1),multiply(multiply(X1,X2),identity))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_25]), c_0_12]), c_0_12]), c_0_12]), c_0_9])).
cnf(c_0_28,plain,(multiply(X1,multiply(multiply(multiply(X1,identity),identity),multiply(multiply(multiply(X1,identity),identity),multiply(multiply(X1,identity),identity))))=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_9]), c_0_1]), c_0_1]), c_0_1]), c_0_1])).
cnf(c_0_29,plain,(multiply(multiply(multiply(X1,identity),identity),multiply(multiply(X2,identity),identity))=multiply(multiply(multiply(X1,X2),identity),identity)), inference(spm,[status(thm)],[c_0_26, c_0_9])).
cnf(c_0_30,plain,(multiply(X1,multiply(multiply(X1,X1),multiply(multiply(X1,X2),identity)))=multiply(identity,multiply(identity,multiply(identity,multiply(X2,identity))))), inference(rw,[status(thm)],[c_0_25, c_0_27])).
cnf(c_0_31,plain,(multiply(X1,multiply(multiply(multiply(X1,multiply(X1,X1)),identity),identity))=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29]), c_0_29])).
cnf(c_0_32,plain,(multiply(multiply(X1,multiply(X1,X1)),identity)=multiply(X1,multiply(multiply(X1,X1),identity))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_1]), c_0_9])).
cnf(c_0_33,plain,(multiply(X1,multiply(multiply(X1,X1),multiply(X1,multiply(multiply(X1,X1),identity))))=multiply(identity,multiply(identity,multiply(identity,multiply(multiply(X1,X1),identity))))), inference(spm,[status(thm)],[c_0_30, c_0_32])).
cnf(c_0_34,plain,(multiply(multiply(X1,X1),multiply(multiply(X1,multiply(X1,identity)),identity))=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_33]), c_0_11]), c_0_1]), c_0_1]), c_0_1]), c_0_1])).
cnf(c_0_35,plain,(multiply(multiply(X1,multiply(multiply(X1,X1),identity)),identity)=multiply(multiply(X1,X1),multiply(multiply(X1,identity),identity))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_10]), c_0_1])).
cnf(c_0_36,plain,(multiply(X1,identity)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_34]), c_0_9])).
cnf(c_0_37,plain,(multiply(multiply(X1,X1),multiply(multiply(X1,X1),multiply(multiply(X1,identity),identity)))=X1), inference(spm,[status(thm)],[c_0_7, c_0_35])).
cnf(c_0_38,plain,(multiply(multiply(X1,X1),multiply(X1,multiply(X1,X2)))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_36]), c_0_36])).
cnf(c_0_39,plain,(multiply(multiply(X1,X1),multiply(multiply(X1,X1),X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_36]), c_0_36])).
cnf(c_0_40,plain,(multiply(multiply(X1,X1),multiply(X1,X1))=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_36]), c_0_36])).
cnf(c_0_41,plain,(multiply(identity,X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_42,plain,(multiply(multiply(X1,multiply(X2,multiply(X2,X2))),X2)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_36]), c_0_36]), c_0_41])).
cnf(c_0_43,plain,(multiply(multiply(X1,X1),X2)=multiply(X1,multiply(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_38]), c_0_40]), c_0_41])).
cnf(c_0_44,plain,(multiply(multiply(X1,multiply(X1,multiply(X2,multiply(X2,X2)))),X2)=multiply(X1,X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_45,plain,(multiply(multiply(multiply(X1,X2),X1),X2)=multiply(multiply(X1,X2),multiply(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_42]), c_0_18]), c_0_41]), c_0_18]), c_0_41]), c_0_18]), c_0_41])).
cnf(c_0_46,plain,(multiply(X1,multiply(X1,multiply(X1,multiply(X1,X2))))=X2), inference(rw,[status(thm)],[c_0_38, c_0_43])).
cnf(c_0_47,plain,(multiply(multiply(X1,X2),multiply(X2,multiply(X2,multiply(X2,X1))))=multiply(X1,X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_48,plain,(multiply(multiply(X1,multiply(X2,X2)),multiply(X2,multiply(X2,X1)))=multiply(X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_43]), c_0_43]), c_0_46]), c_0_43])).
cnf(c_0_49,plain,(multiply(multiply(X1,multiply(X1,X1)),X1)=identity), inference(spm,[status(thm)],[c_0_42, c_0_41])).
cnf(c_0_50,plain,(multiply(X1,multiply(X1,multiply(X1,X1)))=identity), inference(rw,[status(thm)],[c_0_40, c_0_43])).
cnf(c_0_51,plain,(multiply(multiply(X1,multiply(X1,multiply(X1,multiply(X2,X3)))),multiply(X3,multiply(X3,X3)))=multiply(X1,multiply(X1,multiply(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_36]), c_0_36]), c_0_43]), c_0_43])).
cnf(c_0_52,plain,(multiply(multiply(X1,multiply(X1,X1)),multiply(multiply(X1,multiply(X1,X1)),X2))=multiply(X1,multiply(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_48]), c_0_43])).
cnf(c_0_53,plain,(multiply(multiply(X1,multiply(X1,X1)),multiply(X1,multiply(X1,X1)))=multiply(X1,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_49]), c_0_41])).
cnf(c_0_54,plain,(multiply(multiply(X1,multiply(X2,X2)),multiply(X2,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_50]), c_0_43]), c_0_36])).
cnf(c_0_55,plain,(multiply(multiply(X1,multiply(X1,X1)),X2)=multiply(X1,multiply(X1,multiply(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_50]), c_0_36]), c_0_18]), c_0_41])).
cnf(c_0_56,plain,(multiply(multiply(X1,multiply(X1,X1)),multiply(X1,multiply(X1,X2)))=multiply(X1,multiply(X1,multiply(multiply(X1,multiply(X1,X1)),X2)))), inference(spm,[status(thm)],[c_0_52, c_0_52])).
cnf(c_0_57,plain,(multiply(multiply(X1,X2),multiply(X2,multiply(X2,X2)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_53]), c_0_54])).
cnf(c_0_58,plain,(multiply(multiply(X1,multiply(X2,X3)),multiply(X3,multiply(X3,X3)))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_55]), c_0_55]), c_0_46]), c_0_56]), c_0_55]), c_0_46]), c_0_55]), c_0_55]), c_0_46]), c_0_55]), c_0_46])).
cnf(c_0_59,plain,(multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_18]), c_0_41])).
cnf(c_0_60,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_59])]), ['proof']).
# SZS output end CNFRefutation.
EOF
