% Problem    : GRP169-2 : TPTP v6.1.0. Bugfixed v1.2.1.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n125.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:43:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C02_F1_PI_AE_Q12_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', associativity)).
cnf(c_0_1,axiom,(multiply(inverse(X1),X1)=identity), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_inverse)).
cnf(c_0_2,axiom,(multiply(identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_identity)).
cnf(c_0_3,axiom,(multiply(X1,greatest_lower_bound(X2,X3))=greatest_lower_bound(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_glb1)).
cnf(c_0_4,axiom,(multiply(greatest_lower_bound(X1,X2),X3)=greatest_lower_bound(multiply(X1,X3),multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_glb2)).
cnf(c_0_5,axiom,(greatest_lower_bound(X1,X2)=greatest_lower_bound(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', symmetry_of_glb)).
cnf(c_0_6,axiom,(least_upper_bound(X1,greatest_lower_bound(X1,X2))=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', lub_absorbtion)).
cnf(c_0_7,hypothesis,(greatest_lower_bound(inverse(a),inverse(b))=inverse(a)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', p02b_1)).
cnf(c_0_8,axiom,(greatest_lower_bound(X1,least_upper_bound(X1,X2))=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', glb_absorbtion)).
cnf(c_0_9,axiom,(least_upper_bound(X1,X2)=least_upper_bound(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', symmetry_of_lub)).
cnf(c_0_10,negated_conjecture,(greatest_lower_bound(a,b)!=b), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_p02b)).
cnf(c_0_11,plain,(multiply(inverse(X1),multiply(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_12,plain,(multiply(inverse(inverse(X1)),identity)=X1), inference(spm,[status(thm)],[c_0_11, c_0_1])).
cnf(c_0_13,plain,(multiply(inverse(inverse(X1)),X2)=multiply(X1,X2)), inference(spm,[status(thm)],[c_0_11, c_0_11])).
cnf(c_0_14,plain,(multiply(X1,identity)=X1), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_15,plain,(greatest_lower_bound(X1,multiply(X1,X2))=multiply(X1,greatest_lower_bound(identity,X2))), inference(spm,[status(thm)],[c_0_3, c_0_14])).
cnf(c_0_16,plain,(multiply(greatest_lower_bound(X1,inverse(X2)),X2)=greatest_lower_bound(identity,multiply(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_1]), c_0_5])).
cnf(c_0_17,plain,(least_upper_bound(X1,multiply(X1,greatest_lower_bound(identity,X2)))=X1), inference(spm,[status(thm)],[c_0_6, c_0_15])).
cnf(c_0_18,hypothesis,(greatest_lower_bound(identity,multiply(inverse(a),b))=multiply(inverse(a),b)), inference(spm,[status(thm)],[c_0_16, c_0_7])).
cnf(c_0_19,hypothesis,(least_upper_bound(X1,multiply(X1,multiply(inverse(a),b)))=X1), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_20,plain,(inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_14])).
cnf(c_0_21,plain,(greatest_lower_bound(X1,least_upper_bound(X2,X1))=X1), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_22,hypothesis,(least_upper_bound(a,b)=a), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_11]), c_0_20]), c_0_20])).
cnf(c_0_23,hypothesis,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_5]), c_0_10]), ['proof']).
# SZS output end CNFRefutation.
EOF
