% Problem    : ALG236-1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n102.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:19:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(mult(X1,mult(X2,mult(X3,X4)))=mult(X3,mult(X2,mult(X1,X4)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c02)).
cnf(c_0_1,axiom,(mult(X1,mult(X2,mult(X1,X2)))=mult(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c01)).
cnf(c_0_2,axiom,(mult(mult(X1,mult(X2,mult(X3,X2))),X4)=mult(X1,mult(X4,mult(mult(X3,X2),X4)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c03)).
cnf(c_0_3,negated_conjecture,(mult(mult(a,b),mult(c,mult(d,e)))!=mult(a,mult(c,mult(mult(d,b),e)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
cnf(c_0_4,plain,(mult(X1,mult(X2,mult(X3,mult(X2,X1))))=mult(X3,mult(X2,X1))), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_5,plain,(mult(X1,mult(X2,mult(mult(X1,X2),X2)))=mult(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_4]), c_0_1]), c_0_0])).
cnf(c_0_6,plain,(mult(X1,mult(X2,mult(mult(X1,X3),X2)))=mult(mult(X1,X3),X2)), inference(spm,[status(thm)],[c_0_2, c_0_1])).
cnf(c_0_7,plain,(mult(mult(X1,X2),X2)=mult(X1,X2)), inference(rw,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_8,plain,(mult(mult(X1,X2),mult(X3,mult(X4,X2)))=mult(X4,mult(X3,mult(X1,X2)))), inference(spm,[status(thm)],[c_0_0, c_0_7])).
cnf(c_0_9,plain,(mult(mult(X1,X2),mult(X3,mult(mult(X4,X2),X3)))=mult(mult(X4,mult(X2,mult(X1,X2))),X3)), inference(spm,[status(thm)],[c_0_2, c_0_8])).
cnf(c_0_10,plain,(mult(X1,mult(X2,mult(mult(X3,mult(X4,X3)),X2)))=mult(mult(X1,mult(X4,X3)),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_9]), c_0_7]), c_0_1]), c_0_7]), c_0_7])).
cnf(c_0_11,plain,(mult(mult(X1,mult(X2,X1)),mult(X3,mult(X4,X3)))=mult(mult(X4,mult(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_0, c_0_10])).
cnf(c_0_12,plain,(mult(mult(mult(X1,X2),mult(X3,X4)),X2)=mult(mult(X1,mult(X3,X4)),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_7]), c_0_11])).
cnf(c_0_13,plain,(mult(mult(X1,mult(X2,mult(X3,X2))),X4)=mult(mult(X3,X2),mult(X4,mult(X1,X4)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_8]), c_0_12])).
cnf(c_0_14,plain,(mult(mult(X1,mult(X2,mult(X3,X2))),mult(X4,X1))=mult(mult(X3,X2),mult(X4,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_9]), c_0_7]), c_0_7]), c_0_1]), c_0_7])).
cnf(c_0_15,plain,(mult(mult(X1,mult(X2,mult(X3,X2))),mult(X4,X3))=mult(mult(X1,X2),mult(X4,X3))), inference(spm,[status(thm)],[c_0_14, c_0_0])).
cnf(c_0_16,plain,(mult(mult(mult(X1,mult(X2,X3)),X4),mult(X2,X3))=mult(mult(X1,X4),mult(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1]), c_0_11])).
cnf(c_0_17,plain,(mult(X1,mult(X2,mult(mult(mult(X1,X3),X4),X2)))=mult(mult(mult(X1,X3),X4),X2)), inference(spm,[status(thm)],[c_0_6, c_0_6])).
cnf(c_0_18,plain,(mult(X1,mult(X2,mult(X1,mult(X2,X3))))=mult(X1,mult(X2,X3))), inference(spm,[status(thm)],[c_0_1, c_0_8])).
cnf(c_0_19,plain,(mult(X1,mult(X2,mult(mult(X3,X4),mult(X2,X3))))=mult(mult(X1,X4),mult(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_15]), c_0_8])).
cnf(c_0_20,plain,(mult(mult(mult(X1,X2),X3),X2)=mult(mult(X1,X3),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_1]), c_0_0]), c_0_17]), c_0_0]), c_0_6])).
cnf(c_0_21,plain,(mult(X1,mult(X2,mult(X1,mult(mult(X2,X3),X4))))=mult(X1,mult(mult(X2,X3),X4))), inference(spm,[status(thm)],[c_0_18, c_0_6])).
cnf(c_0_22,plain,(mult(mult(X1,X2),mult(X3,mult(X4,X5)))=mult(X1,mult(X3,mult(mult(X4,X2),X5)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_0]), c_0_20]), c_0_21])).
cnf(c_0_23,negated_conjecture,(mult(d,mult(c,mult(mult(a,b),e)))!=mult(a,mult(c,mult(mult(d,b),e)))), inference(rw,[status(thm)],[c_0_3, c_0_0])).
cnf(c_0_24,plain,(mult(X1,mult(X2,mult(mult(X3,X4),X5)))=mult(X3,mult(X2,mult(mult(X1,X4),X5)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_20]), c_0_22]), c_0_22]), c_0_7])).
cnf(c_0_25,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24])]), ['proof']).
# SZS output end CNFRefutation.
EOF
