% Problem    : COL042-8 : TPTP v6.1.0. Released v2.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n120.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:54:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___213_C18_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(apply(apply(w1,X1),X2)=apply(apply(X2,X1),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', w1_definition)).
cnf(c_0_1,axiom,(apply(apply(apply(b,X1),X2),X3)=apply(X1,apply(X2,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', b_definition)).
cnf(c_0_2,axiom,(strong_fixed_point=apply(apply(b,apply(apply(b,apply(w1,w1)),apply(apply(b,apply(b,w1)),b))),b)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', strong_fixed_point)).
cnf(c_0_3,negated_conjecture,(apply(strong_fixed_point,fixed_pt)!=apply(fixed_pt,apply(strong_fixed_point,fixed_pt))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_strong_fixed_point)).
cnf(c_0_4,plain,(apply(apply(w1,X1),apply(apply(b,X2),X3))=apply(apply(X2,apply(X3,X1)),X1)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_5,plain,(apply(apply(w1,X1),apply(apply(b,apply(b,X2)),X3))=apply(X2,apply(apply(X3,X1),X1))), inference(spm,[status(thm)],[c_0_1, c_0_4])).
cnf(c_0_6,plain,(apply(apply(w1,X1),apply(apply(b,apply(b,X2)),X3))=apply(X2,apply(apply(w1,X1),X3))), inference(spm,[status(thm)],[c_0_5, c_0_0])).
cnf(c_0_7,plain,(apply(X1,apply(apply(w1,apply(apply(b,apply(b,X1)),X2)),X2))=apply(apply(w1,apply(apply(b,apply(b,X1)),X2)),w1)), inference(spm,[status(thm)],[c_0_0, c_0_6])).
cnf(c_0_8,plain,(apply(apply(w1,apply(apply(b,apply(b,X1)),w1)),w1)=apply(strong_fixed_point,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1, c_0_2]), c_0_1]), c_0_1]), c_0_4])).
cnf(c_0_9,plain,(apply(X1,apply(strong_fixed_point,X1))=apply(strong_fixed_point,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_8])).
cnf(c_0_10,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_9])]), ['proof']).
# SZS output end CNFRefutation.
EOF
