% Problem    : GRP519-1 : TPTP v6.1.0. Released v2.6.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n141.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:29:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___209_C18_F1_AE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(multiply(X1,multiply(multiply(inverse(multiply(X1,X2)),X3),X2))=X3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', single_axiom)).
cnf(c_0_1,negated_conjecture,(multiply(multiply(a3,b3),c3)!=multiply(a3,multiply(b3,c3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_these_axioms_3)).
cnf(c_0_2,plain,(multiply(multiply(inverse(multiply(inverse(multiply(X1,X2)),X3)),X4),X3)=multiply(X1,multiply(X4,X2))), inference(spm,[status(thm)],[c_0_0, c_0_0])).
cnf(c_0_3,plain,(multiply(inverse(multiply(X1,X2)),multiply(X1,multiply(X3,X2)))=X3), inference(spm,[status(thm)],[c_0_0, c_0_2])).
cnf(c_0_4,plain,(multiply(inverse(multiply(inverse(multiply(X1,X2)),multiply(X3,X2))),X3)=X1), inference(spm,[status(thm)],[c_0_3, c_0_3])).
cnf(c_0_5,plain,(multiply(inverse(X1),multiply(inverse(multiply(inverse(multiply(X2,X3)),X3)),X1))=X2), inference(spm,[status(thm)],[c_0_4, c_0_0])).
cnf(c_0_6,plain,(inverse(multiply(inverse(multiply(X1,X2)),X2))=X1), inference(spm,[status(thm)],[c_0_3, c_0_5])).
cnf(c_0_7,plain,(multiply(inverse(X1),multiply(X2,X1))=X2), inference(rw,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_8,plain,(multiply(inverse(multiply(X1,X2)),X1)=inverse(X2)), inference(spm,[status(thm)],[c_0_7, c_0_7])).
cnf(c_0_9,plain,(multiply(X1,multiply(inverse(X2),X2))=X1), inference(spm,[status(thm)],[c_0_0, c_0_8])).
cnf(c_0_10,plain,(inverse(inverse(X1))=X1), inference(spm,[status(thm)],[c_0_6, c_0_8])).
cnf(c_0_11,plain,(multiply(multiply(inverse(X1),X2),multiply(X3,multiply(X1,X4)))=multiply(X3,multiply(X2,X4))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_0]), c_0_2])).
cnf(c_0_12,plain,(multiply(multiply(X1,X2),inverse(X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_9]), c_0_10])).
cnf(c_0_13,plain,(multiply(multiply(inverse(X1),X2),X1)=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_7]), c_0_7])).
cnf(c_0_14,plain,(multiply(X1,X2)=multiply(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_12]), c_0_10])).
cnf(c_0_15,plain,(multiply(X1,multiply(inverse(X1),X2))=X2), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_16,plain,(multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_15]), c_0_10])).
cnf(c_0_17,negated_conjecture,(multiply(c3,multiply(a3,b3))!=multiply(a3,multiply(b3,c3))), inference(rw,[status(thm)],[c_0_1, c_0_14])).
cnf(c_0_18,plain,(multiply(X1,multiply(X2,X3))=multiply(X3,multiply(X1,X2))), inference(spm,[status(thm)],[c_0_14, c_0_16])).
cnf(c_0_19,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18])]), ['proof']).
# SZS output end CNFRefutation.
EOF
