% Problem    : GRP706-1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n141.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:35:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___209_C18_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(mult(mult(X1,X2),X3)=mult(mult(X1,rd(X3,X3)),mult(X2,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c06)).
cnf(c_0_1,axiom,(mult(rd(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c03)).
cnf(c_0_2,axiom,(rd(mult(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c04)).
cnf(c_0_3,axiom,(mult(X1,mult(X2,X3))=mult(mult(X1,X2),mult(ld(X1,X1),X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c05)).
cnf(c_0_4,axiom,(mult(X1,ld(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c01)).
cnf(c_0_5,axiom,(ld(X1,mult(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c02)).
cnf(c_0_6,negated_conjecture,(f(a,f(b,f(a,c)))!=f(f(f(a,b),a),c)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
cnf(c_0_7,axiom,(f(X1,X2)=mult(rd(X1,op_c),ld(op_c,X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c07)).
cnf(c_0_8,plain,(mult(mult(rd(X1,rd(X2,X2)),X3),X2)=mult(X1,mult(X3,X2))), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_9,plain,(mult(rd(X1,rd(X2,X2)),X3)=rd(mult(X1,mult(X3,X2)),X2)), inference(spm,[status(thm)],[c_0_2, c_0_8])).
cnf(c_0_10,plain,(rd(mult(mult(X1,X2),X3),mult(X2,X3))=mult(X1,rd(X3,X3))), inference(spm,[status(thm)],[c_0_2, c_0_0])).
cnf(c_0_11,plain,(rd(rd(mult(X1,mult(X2,X3)),X3),X2)=rd(X1,rd(X3,X3))), inference(spm,[status(thm)],[c_0_2, c_0_9])).
cnf(c_0_12,plain,(mult(X1,mult(X2,ld(ld(X1,X1),X3)))=mult(mult(X1,X2),X3)), inference(spm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_13,plain,(ld(rd(X1,X2),X1)=X2), inference(spm,[status(thm)],[c_0_5, c_0_1])).
cnf(c_0_14,plain,(rd(mult(X1,X2),mult(ld(X3,X1),X2))=mult(X3,rd(X2,X2))), inference(spm,[status(thm)],[c_0_10, c_0_4])).
cnf(c_0_15,plain,(rd(rd(mult(X1,X2),X3),rd(X2,X3))=rd(X1,rd(X3,X3))), inference(spm,[status(thm)],[c_0_11, c_0_1])).
cnf(c_0_16,plain,(mult(X1,ld(ld(X2,X2),X3))=ld(X2,mult(mult(X2,X1),X3))), inference(spm,[status(thm)],[c_0_5, c_0_12])).
cnf(c_0_17,plain,(ld(mult(X1,rd(X2,X2)),mult(X3,X2))=mult(ld(X1,X3),X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18,plain,(rd(rd(X1,X2),rd(ld(X3,X1),X2))=rd(X3,rd(X2,X2))), inference(spm,[status(thm)],[c_0_15, c_0_4])).
cnf(c_0_19,plain,(ld(X1,ld(X2,mult(mult(X2,X1),X3)))=ld(ld(X2,X2),X3)), inference(spm,[status(thm)],[c_0_5, c_0_16])).
cnf(c_0_20,plain,(ld(mult(X1,X2),mult(X1,mult(X2,X3)))=mult(ld(X1,X1),X3)), inference(spm,[status(thm)],[c_0_5, c_0_3])).
cnf(c_0_21,plain,(ld(mult(X1,rd(X2,X2)),X3)=mult(ld(X1,rd(X3,X2)),X2)), inference(spm,[status(thm)],[c_0_17, c_0_1])).
cnf(c_0_22,plain,(ld(rd(X1,rd(X2,X2)),rd(X3,X2))=rd(ld(X1,X3),X2)), inference(spm,[status(thm)],[c_0_13, c_0_18])).
cnf(c_0_23,plain,(rd(X1,rd(ld(X2,mult(X1,X3)),X3))=rd(X2,rd(X3,X3))), inference(spm,[status(thm)],[c_0_18, c_0_2])).
cnf(c_0_24,plain,(ld(ld(X1,X2),ld(X1,mult(X2,X3)))=ld(ld(X1,X1),X3)), inference(spm,[status(thm)],[c_0_19, c_0_4])).
cnf(c_0_25,plain,(mult(ld(X1,rd(mult(X1,mult(rd(X2,X2),X3)),X2)),X2)=mult(ld(X1,X1),X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26,plain,(ld(rd(X1,rd(X2,X2)),X3)=rd(ld(X1,mult(X3,X2)),X2)), inference(spm,[status(thm)],[c_0_22, c_0_2])).
cnf(c_0_27,plain,(rd(rd(mult(X1,X2),X3),rd(X2,X2))=rd(X1,rd(X3,X2))), inference(spm,[status(thm)],[c_0_23, c_0_13])).
cnf(c_0_28,plain,(ld(ld(X1,rd(X2,X3)),ld(X1,X2))=ld(ld(X1,X1),X3)), inference(spm,[status(thm)],[c_0_24, c_0_1])).
cnf(c_0_29,plain,(ld(X1,rd(mult(X1,mult(rd(X2,X2),X3)),X2))=rd(mult(ld(X1,X1),X3),X2)), inference(spm,[status(thm)],[c_0_2, c_0_25])).
cnf(c_0_30,plain,(rd(ld(rd(mult(X1,X2),X3),mult(X4,X2)),X2)=ld(rd(X1,rd(X3,X2)),X4)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31,plain,(ld(rd(mult(ld(X1,X1),X2),X3),mult(rd(X3,X3),X2))=ld(ld(X1,X1),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_5])).
cnf(c_0_32,plain,(mult(ld(X1,X1),ld(X2,X3))=ld(mult(X1,X2),mult(X1,X3))), inference(spm,[status(thm)],[c_0_20, c_0_4])).
cnf(c_0_33,plain,(ld(rd(ld(X1,X1),rd(X2,X3)),rd(X2,X2))=rd(ld(ld(X1,X1),X2),X3)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34,plain,(rd(X1,ld(X2,X1))=X2), inference(spm,[status(thm)],[c_0_2, c_0_4])).
cnf(c_0_35,plain,(rd(ld(mult(X1,X2),mult(X1,X3)),ld(X2,X3))=ld(X1,X1)), inference(spm,[status(thm)],[c_0_2, c_0_32])).
cnf(c_0_36,plain,(rd(ld(ld(X1,X1),X2),ld(X3,X2))=ld(rd(ld(X1,X1),X3),rd(X2,X2))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_37,plain,(rd(ld(mult(rd(X1,X2),X3),X1),ld(X3,X2))=ld(rd(X1,X2),rd(X1,X2))), inference(spm,[status(thm)],[c_0_35, c_0_1])).
cnf(c_0_38,plain,(mult(rd(X1,X2),rd(X3,X3))=rd(mult(X1,X3),mult(X2,X3))), inference(spm,[status(thm)],[c_0_10, c_0_1])).
cnf(c_0_39,plain,(mult(ld(rd(ld(X1,X1),X2),rd(X3,X3)),ld(X2,X3))=ld(ld(X1,X1),X3)), inference(spm,[status(thm)],[c_0_1, c_0_36])).
cnf(c_0_40,plain,(rd(ld(X1,X1),ld(X2,X2))=ld(rd(X1,X2),rd(X1,X2))), inference(spm,[status(thm)],[c_0_37, c_0_1])).
cnf(c_0_41,plain,(ld(ld(rd(X1,X2),rd(X1,X2)),X3)=ld(X2,ld(rd(X1,X2),mult(X1,X3)))), inference(spm,[status(thm)],[c_0_19, c_0_1])).
cnf(c_0_42,plain,(mult(ld(rd(X1,X2),rd(X3,X4)),X4)=ld(rd(mult(X1,X4),mult(X2,X4)),X3)), inference(spm,[status(thm)],[c_0_21, c_0_38])).
cnf(c_0_43,plain,(ld(rd(X1,rd(X2,X3)),rd(mult(X1,X3),X2))=rd(X3,X3)), inference(spm,[status(thm)],[c_0_13, c_0_27])).
cnf(c_0_44,plain,(ld(X1,mult(ld(rd(X2,X1),mult(X2,rd(X3,X3))),X3))=ld(ld(X2,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_39]), c_0_40]), c_0_41]), c_0_4])).
cnf(c_0_45,plain,(ld(rd(mult(X1,X2),mult(rd(X2,X3),X2)),mult(X1,X3))=mult(rd(X3,X3),X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_46,plain,(mult(ld(rd(X1,X2),mult(X1,rd(X3,X3))),X3)=ld(X1,mult(mult(X1,X2),X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_44]), c_0_16])).
cnf(c_0_47,plain,(ld(mult(X1,rd(X2,X3)),mult(X1,X2))=mult(ld(X1,X1),X3)), inference(spm,[status(thm)],[c_0_20, c_0_1])).
cnf(c_0_48,plain,(ld(rd(mult(rd(X1,X2),X3),mult(rd(X3,X2),X3)),X1)=mult(rd(X2,X2),X3)), inference(spm,[status(thm)],[c_0_45, c_0_1])).
cnf(c_0_49,plain,(ld(rd(X1,X2),mult(X1,rd(X3,X3)))=rd(ld(X1,mult(mult(X1,X2),X3)),X3)), inference(spm,[status(thm)],[c_0_2, c_0_46])).
cnf(c_0_50,plain,(ld(X1,mult(mult(X1,rd(mult(ld(X1,X1),X2),X3)),X3))=mult(rd(X3,X3),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_31]), c_0_16])).
cnf(c_0_51,plain,(rd(mult(rd(mult(X1,X2),X3),mult(X4,X2)),X2)=mult(rd(X1,rd(X3,X2)),X4)), inference(spm,[status(thm)],[c_0_9, c_0_27])).
cnf(c_0_52,plain,(rd(mult(X1,X2),mult(ld(X1,X1),X3))=mult(X1,rd(X2,X3))), inference(spm,[status(thm)],[c_0_34, c_0_47])).
cnf(c_0_53,plain,(rd(mult(rd(X1,X2),X3),mult(rd(X3,X2),X3))=rd(X1,mult(rd(X2,X2),X3))), inference(spm,[status(thm)],[c_0_34, c_0_48])).
cnf(c_0_54,plain,(mult(rd(X1,X2),rd(ld(X1,mult(mult(X1,X2),X3)),X3))=mult(X1,rd(X3,X3))), inference(spm,[status(thm)],[c_0_4, c_0_49])).
cnf(c_0_55,plain,(mult(mult(X1,rd(mult(ld(X1,X1),X2),X3)),X3)=mult(X1,mult(rd(X3,X3),X2))), inference(spm,[status(thm)],[c_0_4, c_0_50])).
cnf(c_0_56,plain,(mult(rd(X1,rd(X2,X3)),rd(X4,X3))=rd(mult(rd(mult(X1,X3),X2),X4),X3)), inference(spm,[status(thm)],[c_0_51, c_0_1])).
cnf(c_0_57,plain,(rd(X1,mult(ld(X2,X2),X3))=mult(X2,rd(ld(X2,X1),X3))), inference(spm,[status(thm)],[c_0_52, c_0_4])).
cnf(c_0_58,plain,(rd(rd(X1,X2),rd(X3,X2))=rd(rd(X1,X3),rd(X2,X2))), inference(spm,[status(thm)],[c_0_15, c_0_1])).
cnf(c_0_59,plain,(mult(rd(X1,mult(rd(X2,X2),X3)),mult(rd(X3,X2),X3))=mult(rd(X1,X2),X3)), inference(spm,[status(thm)],[c_0_1, c_0_53])).
cnf(c_0_60,plain,(rd(mult(mult(X1,rd(X2,X3)),mult(rd(X2,X2),X3)),X2)=mult(X1,rd(X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_5]), c_0_56]), c_0_57]), c_0_5])).
cnf(c_0_61,plain,(mult(rd(rd(X1,X2),rd(X3,X2)),X4)=rd(mult(rd(X1,X3),mult(X4,X2)),X2)), inference(spm,[status(thm)],[c_0_9, c_0_58])).
cnf(c_0_62,plain,(mult(rd(mult(X1,mult(rd(X2,X2),X3)),X2),X3)=mult(X1,mult(rd(X3,X2),X3))), inference(spm,[status(thm)],[c_0_59, c_0_2])).
cnf(c_0_63,plain,(rd(mult(X1,mult(rd(X2,X2),X3)),X2)=rd(mult(X1,X2),mult(rd(X2,X3),X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_1]), c_0_38])).
cnf(c_0_64,plain,(rd(mult(rd(X1,mult(X2,X3)),mult(X4,X3)),X3)=mult(rd(rd(X1,X3),X2),X4)), inference(spm,[status(thm)],[c_0_61, c_0_2])).
cnf(c_0_65,plain,(mult(rd(mult(X1,X2),mult(rd(X2,X3),X2)),X3)=mult(X1,mult(rd(X3,X2),X3))), inference(rw,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_66,plain,(mult(rd(X1,rd(X2,mult(X3,X2))),X3)=rd(mult(X1,mult(X3,mult(X3,X2))),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_2]), c_0_2])).
cnf(c_0_67,plain,(mult(rd(X1,rd(ld(X2,X3),X3)),X2)=rd(mult(X1,mult(X2,X3)),ld(X2,X3))), inference(spm,[status(thm)],[c_0_66, c_0_4])).
cnf(c_0_68,plain,(ld(ld(X1,rd(mult(X1,X2),X3)),X2)=ld(ld(X1,X1),X3)), inference(spm,[status(thm)],[c_0_28, c_0_5])).
cnf(c_0_69,plain,(rd(ld(X1,X2),ld(ld(X1,X1),X3))=ld(X1,rd(X2,X3))), inference(spm,[status(thm)],[c_0_34, c_0_28])).
cnf(c_0_70,plain,(rd(rd(mult(X1,mult(X2,X3)),ld(X2,X3)),X2)=rd(X1,rd(ld(X2,X3),X3))), inference(spm,[status(thm)],[c_0_2, c_0_67])).
cnf(c_0_71,plain,(ld(ld(X1,X1),ld(X2,mult(X1,X3)))=ld(ld(X1,X2),X3)), inference(spm,[status(thm)],[c_0_68, c_0_34])).
cnf(c_0_72,plain,(rd(X1,ld(ld(X2,X2),X3))=ld(X2,rd(mult(X2,X1),X3))), inference(spm,[status(thm)],[c_0_69, c_0_5])).
cnf(c_0_73,plain,(rd(rd(X1,mult(rd(X2,X2),X3)),rd(X2,X3))=rd(rd(X1,X2),rd(X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_59]), c_0_13]), c_0_15]), c_0_13])).
cnf(c_0_74,plain,(mult(X1,rd(ld(X1,X2),ld(X3,X4)))=rd(X2,ld(mult(X1,X3),mult(X1,X4)))), inference(spm,[status(thm)],[c_0_57, c_0_32])).
cnf(c_0_75,plain,(rd(ld(X1,mult(X2,X3)),ld(ld(X2,X1),X3))=ld(X2,X2)), inference(spm,[status(thm)],[c_0_34, c_0_71])).
cnf(c_0_76,plain,(ld(X1,rd(mult(X1,X2),ld(X3,mult(X1,X4))))=rd(X2,ld(ld(X1,X3),X4))), inference(spm,[status(thm)],[c_0_72, c_0_71])).
cnf(c_0_77,plain,(rd(rd(X1,X2),rd(X3,ld(rd(X3,X3),X2)))=rd(rd(X1,X3),rd(X3,X3))), inference(spm,[status(thm)],[c_0_73, c_0_4])).
cnf(c_0_78,plain,(rd(mult(X1,X2),ld(mult(X3,ld(X1,X3)),mult(X3,X2)))=mult(X3,ld(X1,X1))), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_79,plain,(rd(ld(X1,X2),ld(ld(X1,X3),X4))=ld(X1,rd(X2,ld(X3,mult(X1,X4))))), inference(spm,[status(thm)],[c_0_76, c_0_4])).
cnf(c_0_80,plain,(rd(X1,rd(X2,ld(rd(X2,X2),X3)))=rd(rd(mult(X1,X3),X2),rd(X2,X2))), inference(spm,[status(thm)],[c_0_77, c_0_2])).
cnf(c_0_81,plain,(ld(rd(rd(X1,X2),rd(X3,X2)),X4)=rd(ld(rd(X1,X3),mult(X4,X2)),X2)), inference(spm,[status(thm)],[c_0_26, c_0_58])).
cnf(c_0_82,plain,(rd(X1,ld(mult(X2,ld(X3,X2)),mult(X2,ld(X3,X1))))=mult(X2,ld(X3,X3))), inference(spm,[status(thm)],[c_0_78, c_0_4])).
cnf(c_0_83,plain,(rd(mult(X1,ld(X2,X3)),ld(mult(X2,ld(X1,X2)),X3))=mult(X2,ld(X1,X1))), inference(spm,[status(thm)],[c_0_78, c_0_4])).
cnf(c_0_84,plain,(ld(rd(X1,ld(X1,X2)),rd(X1,ld(X1,X2)))=ld(X1,rd(X1,ld(X2,X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_40]), c_0_4])).
cnf(c_0_85,plain,(rd(ld(rd(mult(X1,X2),X3),mult(X1,X3)),X3)=rd(X3,ld(rd(X3,X3),X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_80]), c_0_81])).
cnf(c_0_86,plain,(ld(mult(X1,ld(X2,X1)),mult(X1,ld(X2,X3)))=ld(mult(X1,ld(X2,X2)),X3)), inference(spm,[status(thm)],[c_0_13, c_0_82])).
cnf(c_0_87,plain,(ld(mult(X1,ld(X2,X2)),mult(X2,ld(X1,X3)))=ld(mult(X1,ld(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_13, c_0_83])).
cnf(c_0_88,plain,(mult(rd(X1,ld(X1,X2)),ld(X1,rd(X1,ld(X2,X2))))=rd(X1,ld(X1,X2))), inference(spm,[status(thm)],[c_0_4, c_0_84])).
cnf(c_0_89,plain,(rd(X1,ld(rd(X1,X1),ld(X2,X3)))=rd(ld(rd(X3,X1),mult(X2,X1)),X1)), inference(spm,[status(thm)],[c_0_85, c_0_4])).
cnf(c_0_90,plain,(mult(mult(X1,ld(X2,X1)),ld(mult(X1,ld(X2,X2)),X3))=mult(X1,ld(X2,X3))), inference(spm,[status(thm)],[c_0_4, c_0_86])).
cnf(c_0_91,plain,(ld(rd(X1,ld(X2,X2)),rd(X1,ld(X1,X2)))=ld(X2,rd(X1,ld(X2,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_84]), c_0_4]), c_0_13]), c_0_4])).
cnf(c_0_92,plain,(ld(rd(ld(rd(X1,X2),mult(X3,X2)),X2),X2)=ld(rd(X2,X2),ld(X3,X1))), inference(spm,[status(thm)],[c_0_13, c_0_89])).
cnf(c_0_93,plain,(mult(X1,ld(X2,mult(mult(X1,ld(X2,X2)),X3)))=mult(mult(X1,ld(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_90, c_0_5])).
cnf(c_0_94,plain,(mult(rd(X1,ld(X2,X2)),ld(X2,rd(X1,ld(X2,X2))))=rd(X1,ld(X1,X2))), inference(spm,[status(thm)],[c_0_4, c_0_91])).
cnf(c_0_95,plain,(mult(rd(mult(X1,X2),mult(rd(X2,X3),X2)),X2)=mult(X1,mult(rd(X2,X2),X3))), inference(spm,[status(thm)],[c_0_1, c_0_63])).
cnf(c_0_96,plain,(ld(rd(ld(rd(X1,X2),X3),X2),X2)=ld(rd(X2,X2),ld(rd(X3,X2),X1))), inference(spm,[status(thm)],[c_0_92, c_0_1])).
cnf(c_0_97,plain,(rd(ld(X1,mult(mult(X1,ld(X2,X1)),X3)),X3)=ld(X2,mult(X1,rd(X3,X3)))), inference(spm,[status(thm)],[c_0_49, c_0_34])).
cnf(c_0_98,plain,(mult(rd(X1,ld(X2,X2)),ld(X2,mult(X1,X3)))=mult(rd(X1,ld(X1,X2)),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_1]), c_0_94])).
cnf(c_0_99,plain,(mult(X1,mult(rd(X2,X2),ld(X3,X2)))=mult(rd(mult(X1,X2),mult(X3,X2)),X2)), inference(spm,[status(thm)],[c_0_95, c_0_34])).
cnf(c_0_100,plain,(ld(rd(ld(X1,mult(X2,X3)),X3),X3)=ld(rd(X3,X3),ld(X2,mult(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_30]), c_0_26]), c_0_2])).
cnf(c_0_101,plain,(rd(mult(X1,ld(X2,X1)),ld(X2,mult(X1,rd(X3,X3))))=rd(X1,rd(X3,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_97]), c_0_2])).
cnf(c_0_102,plain,(ld(X1,rd(mult(X1,X2),mult(rd(X2,X3),X2)))=rd(mult(ld(X1,X1),X3),X2)), inference(rw,[status(thm)],[c_0_29, c_0_63])).
cnf(c_0_103,plain,(ld(rd(X1,ld(X2,X2)),mult(rd(X1,ld(X1,X2)),X3))=ld(X2,mult(X1,X3))), inference(spm,[status(thm)],[c_0_5, c_0_98])).
cnf(c_0_104,plain,(mult(rd(mult(X1,X2),ld(X3,mult(X4,X2))),X2)=mult(X1,ld(X4,mult(X3,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_4]), c_0_1])).
cnf(c_0_105,plain,(mult(rd(X1,mult(rd(X2,X3),X2)),X2)=mult(rd(X1,X2),mult(rd(X2,X2),X3))), inference(spm,[status(thm)],[c_0_95, c_0_1])).
cnf(c_0_106,plain,(rd(mult(X1,mult(ld(X2,mult(X1,rd(X3,X3))),X3)),X3)=mult(X1,ld(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1, c_0_101]), c_0_9])).
cnf(c_0_107,plain,(rd(ld(mult(X1,X2),mult(X1,X3)),X3)=ld(X1,rd(mult(X1,X3),mult(X2,X3)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_34]), c_0_32])).
cnf(c_0_108,plain,(ld(rd(X1,ld(X1,X2)),mult(rd(X1,ld(X2,X2)),X3))=ld(X1,mult(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_91]), c_0_84]), c_0_34]), c_0_34])).
cnf(c_0_109,plain,(mult(rd(mult(X1,X2),ld(X3,X4)),X2)=mult(X1,ld(rd(X4,X2),mult(X3,X2)))), inference(spm,[status(thm)],[c_0_104, c_0_1])).
cnf(c_0_110,plain,(mult(mult(X1,rd(X2,X3)),mult(rd(X2,X2),X3))=mult(mult(X1,rd(X2,X2)),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_10]), c_0_2])).
cnf(c_0_111,plain,(mult(X1,mult(ld(X2,mult(X1,rd(X3,X3))),X3))=mult(mult(X1,ld(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_1, c_0_106])).
cnf(c_0_112,plain,(ld(rd(X1,ld(X2,X2)),mult(X2,rd(X3,X3)))=rd(ld(X1,mult(X2,X3)),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_94]), c_0_108]), c_0_14])).
cnf(c_0_113,plain,(mult(X1,ld(rd(mult(X2,X3),X4),mult(X2,X4)))=mult(rd(mult(X1,X4),X3),X4)), inference(spm,[status(thm)],[c_0_109, c_0_5])).
cnf(c_0_114,plain,(mult(mult(X1,rd(X2,ld(rd(X2,X2),X3))),X3)=mult(mult(X1,rd(X2,X2)),X2)), inference(spm,[status(thm)],[c_0_110, c_0_4])).
cnf(c_0_115,plain,(mult(rd(X1,ld(rd(X2,X3),mult(X4,X3))),X3)=mult(rd(X1,X3),ld(X4,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_89]), c_0_1]), c_0_4])).
cnf(c_0_116,plain,(mult(mult(X1,ld(rd(X2,ld(X1,X1)),X1)),X3)=mult(X1,ld(X2,mult(X1,X3)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_1])).
cnf(c_0_117,plain,(mult(X1,ld(rd(mult(rd(X2,X3),X4),X3),X2))=mult(rd(mult(X1,X3),X4),X3)), inference(spm,[status(thm)],[c_0_113, c_0_1])).
cnf(c_0_118,plain,(mult(X1,rd(X2,ld(rd(X2,X2),X3)))=rd(mult(mult(X1,rd(X2,X2)),X2),X3)), inference(spm,[status(thm)],[c_0_2, c_0_114])).
cnf(c_0_119,plain,(mult(rd(X1,ld(rd(X2,X3),X4)),X3)=mult(rd(X1,X3),ld(rd(X4,X3),X2))), inference(spm,[status(thm)],[c_0_115, c_0_1])).
cnf(c_0_120,plain,(rd(mult(rd(X1,X2),ld(X3,X4)),X2)=rd(X1,ld(rd(X4,X2),mult(X3,X2)))), inference(spm,[status(thm)],[c_0_2, c_0_115])).
cnf(c_0_121,plain,(mult(X1,ld(mult(X2,ld(X1,X1)),mult(X1,X3)))=mult(mult(X1,ld(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_116, c_0_2])).
cnf(c_0_122,plain,(mult(mult(rd(X1,X2),ld(X1,X1)),X3)=mult(X1,ld(mult(X1,X2),mult(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_4]), c_0_34])).
cnf(c_0_123,plain,(rd(mult(X1,ld(X2,mult(X3,X4))),X4)=rd(mult(X1,X4),ld(X3,mult(X2,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_80]), c_0_80]), c_0_81]), c_0_61]), c_0_1]), c_0_120]), c_0_2]), c_0_1]), c_0_80]), c_0_81]), c_0_56]), c_0_119]), c_0_2]), c_0_2]), c_0_1])).
cnf(c_0_124,plain,(ld(mult(X1,ld(X2,X2)),mult(X2,X3))=ld(X2,mult(mult(X2,ld(X1,X2)),X3))), inference(spm,[status(thm)],[c_0_5, c_0_121])).
cnf(c_0_125,plain,(rd(mult(X1,X2),ld(X1,mult(mult(X1,X3),X2)))=mult(rd(X1,X3),ld(X1,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_122]), c_0_123])).
cnf(c_0_126,plain,(ld(X1,mult(mult(X1,ld(X2,X1)),ld(X1,X3)))=ld(mult(X2,ld(X1,X1)),X3)), inference(spm,[status(thm)],[c_0_124, c_0_4])).
cnf(c_0_127,plain,(ld(X1,rd(mult(mult(X1,rd(X2,X2)),X2),X3))=rd(X2,ld(rd(X2,X2),X3))), inference(spm,[status(thm)],[c_0_5, c_0_118])).
cnf(c_0_128,plain,(rd(X1,ld(X2,mult(mult(X2,X3),ld(X2,X1))))=mult(rd(X2,X3),ld(X2,X2))), inference(spm,[status(thm)],[c_0_125, c_0_4])).
cnf(c_0_129,plain,(ld(rd(mult(X1,X2),mult(X3,X2)),mult(X4,X2))=mult(ld(rd(X1,X3),X4),X2)), inference(spm,[status(thm)],[c_0_42, c_0_2])).
cnf(c_0_130,plain,(mult(mult(X1,ld(X2,X1)),ld(X1,X3))=mult(X1,ld(mult(X2,ld(X1,X1)),X3))), inference(spm,[status(thm)],[c_0_4, c_0_126])).
cnf(c_0_131,plain,(mult(X1,rd(ld(X1,mult(X2,mult(ld(X1,X1),X3))),X3))=X2), inference(spm,[status(thm)],[c_0_2, c_0_57])).
cnf(c_0_132,plain,(ld(rd(mult(X1,X2),ld(X3,X1)),X1)=ld(X3,mult(rd(X1,X2),ld(X1,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_89]), c_0_123]), c_0_21]), c_0_2]), c_0_129]), c_0_2])).
cnf(c_0_133,plain,(ld(ld(X1,X2),ld(mult(X1,ld(X2,X2)),X3))=ld(ld(X2,X2),ld(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_130]), c_0_5])).
cnf(c_0_134,plain,(rd(ld(X1,mult(X2,mult(ld(X1,X1),X3))),X3)=ld(X1,X2)), inference(spm,[status(thm)],[c_0_5, c_0_131])).
cnf(c_0_135,plain,(rd(X1,ld(X2,mult(rd(X1,X3),ld(X1,X1))))=rd(mult(X1,X3),ld(X2,X1))), inference(spm,[status(thm)],[c_0_34, c_0_132])).
cnf(c_0_136,plain,(ld(ld(rd(X1,ld(X2,X2)),X2),ld(X1,X3))=ld(ld(X2,X2),ld(X2,X3))), inference(spm,[status(thm)],[c_0_133, c_0_1])).
cnf(c_0_137,plain,(rd(mult(X1,X2),ld(rd(X3,X2),mult(X4,X2)))=rd(mult(X1,ld(X4,X3)),X2)), inference(spm,[status(thm)],[c_0_120, c_0_2])).
cnf(c_0_138,plain,(ld(X1,mult(X2,mult(ld(X1,X1),X3)))=mult(ld(X1,X2),X3)), inference(spm,[status(thm)],[c_0_1, c_0_134])).
cnf(c_0_139,plain,(ld(ld(X1,X1),ld(mult(X1,X2),X3))=ld(X2,ld(X1,X3))), inference(spm,[status(thm)],[c_0_19, c_0_4])).
cnf(c_0_140,plain,(rd(mult(X1,ld(X2,X1)),ld(X3,X1))=rd(X1,ld(X3,mult(X2,ld(X1,X1))))), inference(spm,[status(thm)],[c_0_135, c_0_34])).
cnf(c_0_141,plain,(ld(ld(rd(X1,ld(X2,X2)),X2),X3)=ld(ld(X2,X2),ld(X2,mult(X1,X3)))), inference(spm,[status(thm)],[c_0_136, c_0_5])).
cnf(c_0_142,plain,(ld(rd(mult(X1,ld(X2,X3)),X4),mult(X1,X4))=ld(rd(X3,X4),mult(X2,X4))), inference(spm,[status(thm)],[c_0_13, c_0_137])).
cnf(c_0_143,plain,(rd(mult(X1,ld(mult(X1,X2),X3)),ld(X1,X3))=mult(rd(X1,X2),ld(X1,X1))), inference(spm,[status(thm)],[c_0_125, c_0_4])).
cnf(c_0_144,plain,(ld(X1,mult(X2,ld(mult(X1,X3),mult(X1,X4))))=mult(ld(X1,X2),ld(X3,X4))), inference(spm,[status(thm)],[c_0_138, c_0_32])).
cnf(c_0_145,plain,(ld(X1,mult(rd(X2,ld(X2,X3)),ld(X2,X2)))=ld(rd(X3,ld(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_132, c_0_4])).
cnf(c_0_146,plain,(ld(X1,rd(mult(X1,X2),ld(mult(X1,X3),X4)))=rd(X2,ld(X3,ld(X1,X4)))), inference(spm,[status(thm)],[c_0_72, c_0_139])).
cnf(c_0_147,plain,(rd(X1,ld(rd(X1,X2),mult(X3,ld(X1,X1))))=rd(mult(X1,ld(X3,X1)),X2)), inference(spm,[status(thm)],[c_0_140, c_0_13])).
cnf(c_0_148,plain,(ld(X1,rd(mult(X1,X2),ld(X1,mult(X3,X2))))=ld(rd(X3,ld(X1,X1)),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_141]), c_0_72])).
cnf(c_0_149,plain,(ld(mult(rd(X1,X2),ld(X1,X1)),X3)=ld(X1,mult(mult(X1,X2),ld(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_4]), c_0_34])).
cnf(c_0_150,plain,(mult(rd(X1,ld(X1,ld(X2,X3))),ld(X1,X1))=ld(rd(X3,X1),mult(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_5]), c_0_146])).
cnf(c_0_151,plain,(rd(mult(X1,ld(rd(X2,ld(X1,X1)),X1)),X3)=rd(X1,ld(rd(X1,X3),X2))), inference(spm,[status(thm)],[c_0_147, c_0_1])).
cnf(c_0_152,plain,(ld(rd(rd(X1,X2),ld(X3,X3)),X3)=ld(X3,rd(mult(X3,X2),ld(X3,X1)))), inference(spm,[status(thm)],[c_0_148, c_0_1])).
cnf(c_0_153,negated_conjecture,(mult(rd(mult(rd(mult(rd(a,op_c),ld(op_c,b)),op_c),ld(op_c,a)),op_c),ld(op_c,c))!=mult(rd(a,op_c),ld(op_c,mult(rd(b,op_c),ld(op_c,mult(rd(a,op_c),ld(op_c,c))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_7]), c_0_7]), c_0_7]), c_0_7]), c_0_7]), c_0_7])).
cnf(c_0_154,plain,(ld(ld(rd(X1,X2),mult(X3,X2)),X4)=ld(X2,mult(ld(X3,X1),ld(X2,X4)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_4])).
cnf(c_0_155,plain,(ld(X1,mult(rd(mult(X1,X2),mult(X3,X2)),X2))=mult(rd(X2,X2),ld(X3,X2))), inference(spm,[status(thm)],[c_0_5, c_0_99])).
cnf(c_0_156,plain,(rd(X1,ld(rd(X1,X2),rd(X3,X4)))=rd(rd(mult(X1,X4),ld(X1,X3)),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_4])).
cnf(c_0_157,negated_conjecture,(mult(rd(mult(rd(a,op_c),ld(op_c,b)),ld(rd(a,op_c),mult(op_c,op_c))),ld(op_c,c))!=mult(rd(a,op_c),ld(op_c,mult(rd(b,op_c),ld(op_c,mult(rd(a,op_c),ld(op_c,c))))))), inference(rw,[status(thm)],[c_0_153, c_0_120])).
cnf(c_0_158,plain,(mult(rd(mult(X1,X2),ld(X1,X3)),X4)=mult(X1,ld(rd(X3,X2),mult(X1,X4)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_152]), c_0_4])).
cnf(c_0_159,plain,(ld(mult(rd(X1,X1),ld(X2,X1)),X3)=ld(X1,mult(mult(X2,X1),ld(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_1]), c_0_13])).
cnf(c_0_160,plain,(mult(rd(X1,X2),ld(rd(rd(X3,X4),X2),X1))=rd(mult(X1,X4),ld(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_156]), c_0_1])).
cnf(c_0_161,negated_conjecture,(mult(rd(a,op_c),ld(rd(mult(op_c,op_c),ld(op_c,b)),mult(rd(a,op_c),ld(op_c,c))))!=mult(rd(a,op_c),ld(op_c,mult(rd(b,op_c),ld(op_c,mult(rd(a,op_c),ld(op_c,c))))))), inference(rw,[status(thm)],[c_0_157, c_0_158])).
cnf(c_0_162,plain,(ld(rd(mult(X1,X2),ld(X1,X3)),X4)=ld(X1,mult(rd(X3,X2),ld(X1,X4)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_1])).
cnf(c_0_163,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_161, c_0_162])]), ['proof']).
# SZS output end CNFRefutation.
EOF
