% Problem    : REL029-1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n191.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:49:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4d
# and selection function SelectCQIPrecWNTNp.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(converse(composition(X1,X2))=composition(converse(X2),converse(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_multiplicativity_10)).
cnf(c_0_1,axiom,(composition(X1,one)=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_identity_6)).
cnf(c_0_2,axiom,(converse(converse(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_idempotence_8)).
cnf(c_0_3,axiom,(join(composition(converse(X1),complement(composition(X1,X2))),complement(X2))=complement(X2)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_cancellativity_11)).
cnf(c_0_4,axiom,(join(X1,X2)=join(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux1_join_commutativity_1)).
cnf(c_0_5,axiom,(zero=meet(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', def_zero_13)).
cnf(c_0_6,axiom,(meet(X1,X2)=complement(join(complement(X1),complement(X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux4_definiton_of_meet_4)).
cnf(c_0_7,axiom,(top=join(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', def_top_12)).
cnf(c_0_8,axiom,(X1=join(complement(join(complement(X1),complement(X2))),complement(join(complement(X1),X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux3_a_kind_of_de_Morgan_3)).
cnf(c_0_9,axiom,(join(X1,join(X2,X3))=join(join(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux2_join_associativity_2)).
cnf(c_0_10,axiom,(converse(join(X1,X2))=join(converse(X1),converse(X2))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_additivity_9)).
cnf(c_0_11,axiom,(composition(join(X1,X2),X3)=join(composition(X1,X3),composition(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_distributivity_7)).
cnf(c_0_12,negated_conjecture,(join(sk1,one)=one), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_14)).
cnf(c_0_13,negated_conjecture,(join(sk2,one)=one), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_15)).
cnf(c_0_14,axiom,(composition(X1,composition(X2,X3))=composition(composition(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_associativity_5)).
cnf(c_0_15,negated_conjecture,(meet(composition(sk1,sk3),composition(sk2,sk3))!=composition(meet(sk1,sk2),sk3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_16)).
cnf(c_0_16,plain,(composition(converse(one),converse(X1))=converse(X1)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_17,plain,(composition(converse(one),X1)=X1), inference(spm,[status(thm)],[c_0_16, c_0_2])).
cnf(c_0_18,plain,(converse(one)=one), inference(spm,[status(thm)],[c_0_1, c_0_17])).
cnf(c_0_19,plain,(join(complement(X1),composition(converse(X2),complement(composition(X2,X1))))=complement(X1)), inference(rw,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_20,plain,(composition(one,X1)=X1), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21,plain,(complement(join(complement(X1),complement(complement(X1))))=zero), inference(rw,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_22,plain,(join(complement(X1),complement(X1))=complement(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_18]), c_0_20])).
cnf(c_0_23,plain,(complement(top)=zero), inference(rw,[status(thm)],[c_0_21, c_0_7])).
cnf(c_0_24,plain,(join(zero,zero)=zero), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,plain,(join(complement(join(complement(X1),X2)),complement(join(complement(X1),complement(X2))))=X1), inference(rw,[status(thm)],[c_0_8, c_0_4])).
cnf(c_0_26,plain,(join(zero,join(zero,X1))=join(zero,X1)), inference(spm,[status(thm)],[c_0_9, c_0_24])).
cnf(c_0_27,plain,(join(zero,complement(complement(X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_22]), c_0_7]), c_0_23]), c_0_4])).
cnf(c_0_28,plain,(join(zero,X1)=X1), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_29,plain,(complement(complement(X1))=X1), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_30,plain,(join(X1,X1)=X1), inference(spm,[status(thm)],[c_0_22, c_0_29])).
cnf(c_0_31,plain,(join(X1,join(X1,X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_9, c_0_30])).
cnf(c_0_32,plain,(join(X1,complement(join(complement(X1),X2)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_25]), c_0_4])).
cnf(c_0_33,plain,(join(X1,join(X2,X1))=join(X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_4])).
cnf(c_0_34,plain,(join(converse(X1),converse(complement(X1)))=converse(top)), inference(spm,[status(thm)],[c_0_10, c_0_7])).
cnf(c_0_35,plain,(join(X1,complement(join(X2,complement(X1))))=X1), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36,plain,(join(X1,top)=top), inference(spm,[status(thm)],[c_0_31, c_0_7])).
cnf(c_0_37,plain,(join(converse(X1),converse(top))=converse(top)), inference(spm,[status(thm)],[c_0_31, c_0_34])).
cnf(c_0_38,plain,(join(complement(join(complement(X1),X2)),complement(join(complement(X2),complement(X1))))=X1), inference(spm,[status(thm)],[c_0_25, c_0_4])).
cnf(c_0_39,plain,(join(complement(join(X1,X2)),complement(join(X1,complement(X2))))=complement(X1)), inference(spm,[status(thm)],[c_0_25, c_0_29])).
cnf(c_0_40,plain,(join(complement(X1),complement(join(X2,X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_35, c_0_29])).
cnf(c_0_41,plain,(join(top,X1)=top), inference(spm,[status(thm)],[c_0_4, c_0_36])).
cnf(c_0_42,plain,(join(X1,converse(top))=converse(top)), inference(spm,[status(thm)],[c_0_37, c_0_2])).
cnf(c_0_43,plain,(join(X1,complement(join(X1,complement(X2))))=join(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_29]), c_0_29]), c_0_9]), c_0_40])).
cnf(c_0_44,plain,(converse(top)=top), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45,plain,(join(X1,complement(join(X1,X2)))=join(X1,complement(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_43]), c_0_31])).
cnf(c_0_46,plain,(join(converse(X1),converse(complement(X1)))=top), inference(rw,[status(thm)],[c_0_34, c_0_44])).
cnf(c_0_47,plain,(join(X1,zero)=X1), inference(spm,[status(thm)],[c_0_4, c_0_28])).
cnf(c_0_48,plain,(join(converse(X1),complement(converse(complement(X1))))=converse(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_23]), c_0_47])).
cnf(c_0_49,plain,(join(X1,converse(complement(converse(complement(X1)))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_48]), c_0_2]), c_0_2])).
cnf(c_0_50,plain,(join(X1,converse(complement(converse(X1))))=top), inference(spm,[status(thm)],[c_0_46, c_0_2])).
cnf(c_0_51,plain,(join(X1,join(complement(X1),X2))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_41])).
cnf(c_0_52,plain,(join(complement(X1),converse(complement(converse(X1))))=complement(X1)), inference(spm,[status(thm)],[c_0_49, c_0_29])).
cnf(c_0_53,plain,(join(X1,complement(converse(complement(converse(X1)))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_50]), c_0_23]), c_0_47])).
cnf(c_0_54,plain,(join(complement(X1),join(X1,X2))=top), inference(spm,[status(thm)],[c_0_51, c_0_29])).
cnf(c_0_55,negated_conjecture,(join(X1,composition(sk1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_20]), c_0_4]), c_0_20])).
cnf(c_0_56,plain,(complement(converse(complement(converse(X1))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_52]), c_0_29]), c_0_4]), c_0_53])).
cnf(c_0_57,negated_conjecture,(join(one,complement(sk2))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_13]), c_0_4])).
cnf(c_0_58,plain,(join(converse(zero),converse(X1))=converse(X1)), inference(spm,[status(thm)],[c_0_10, c_0_28])).
cnf(c_0_59,negated_conjecture,(join(X1,join(composition(sk1,X1),X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_9, c_0_55])).
cnf(c_0_60,plain,(converse(complement(converse(X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_29, c_0_56])).
cnf(c_0_61,negated_conjecture,(join(X1,composition(complement(sk2),X1))=composition(top,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_57]), c_0_20])).
cnf(c_0_62,plain,(join(converse(zero),X1)=X1), inference(spm,[status(thm)],[c_0_58, c_0_2])).
cnf(c_0_63,negated_conjecture,(join(X1,complement(composition(sk1,X1)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_7]), c_0_36])).
cnf(c_0_64,plain,(converse(complement(X1))=complement(converse(X1))), inference(spm,[status(thm)],[c_0_2, c_0_60])).
cnf(c_0_65,negated_conjecture,(join(X1,composition(top,complement(X1)))=top), inference(spm,[status(thm)],[c_0_51, c_0_61])).
cnf(c_0_66,plain,(join(converse(complement(X1)),composition(converse(complement(composition(X2,X1))),X2))=converse(complement(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_19]), c_0_0]), c_0_2])).
cnf(c_0_67,plain,(converse(zero)=zero), inference(spm,[status(thm)],[c_0_47, c_0_62])).
cnf(c_0_68,negated_conjecture,(join(converse(X1),complement(composition(converse(X1),converse(sk1))))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_63]), c_0_44]), c_0_64]), c_0_0])).
cnf(c_0_69,negated_conjecture,(join(converse(X1),composition(converse(complement(X1)),top))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_65]), c_0_44]), c_0_0]), c_0_44])).
cnf(c_0_70,plain,(composition(converse(complement(composition(X1,top))),X1)=zero), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_23]), c_0_67]), c_0_28]), c_0_67])).
cnf(c_0_71,negated_conjecture,(join(X1,complement(composition(X1,converse(sk1))))=top), inference(spm,[status(thm)],[c_0_68, c_0_2])).
cnf(c_0_72,negated_conjecture,(join(converse(X1),composition(complement(converse(X1)),top))=top), inference(rw,[status(thm)],[c_0_69, c_0_64])).
cnf(c_0_73,plain,(join(composition(X1,X2),composition(complement(X1),X2))=composition(top,X2)), inference(spm,[status(thm)],[c_0_11, c_0_7])).
cnf(c_0_74,plain,(composition(complement(composition(top,converse(X1))),X1)=zero), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_64]), c_0_0]), c_0_44])).
cnf(c_0_75,plain,(join(X1,join(X2,X3))=join(X2,join(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_4]), c_0_9])).
cnf(c_0_76,negated_conjecture,(join(X1,composition(X1,converse(sk1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_71]), c_0_23]), c_0_47]), c_0_29])).
cnf(c_0_77,negated_conjecture,(join(one,complement(sk1))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_12]), c_0_4])).
cnf(c_0_78,negated_conjecture,(join(X1,composition(complement(X1),top))=top), inference(spm,[status(thm)],[c_0_72, c_0_2])).
cnf(c_0_79,negated_conjecture,(join(one,converse(sk1))=one), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_12]), c_0_18]), c_0_4]), c_0_18])).
cnf(c_0_80,plain,(composition(top,composition(converse(X1),X1))=composition(top,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_29]), c_0_14]), c_0_28])).
cnf(c_0_81,negated_conjecture,(join(X1,join(X2,composition(X1,converse(sk1))))=join(X2,X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_82,negated_conjecture,(join(X1,composition(complement(sk1),X1))=composition(top,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_77]), c_0_20])).
cnf(c_0_83,negated_conjecture,(join(one,converse(sk2))=one), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_13]), c_0_18]), c_0_4]), c_0_18])).
cnf(c_0_84,negated_conjecture,(join(complement(X1),composition(X1,top))=top), inference(spm,[status(thm)],[c_0_78, c_0_29])).
cnf(c_0_85,negated_conjecture,(join(X1,composition(converse(sk1),X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_79]), c_0_20]), c_0_20])).
cnf(c_0_86,plain,(composition(converse(X1),composition(X1,top))=composition(converse(X1),top)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_80]), c_0_0]), c_0_44]), c_0_0]), c_0_2]), c_0_44]), c_0_14])).
cnf(c_0_87,negated_conjecture,(join(complement(sk1),composition(top,converse(sk1)))=join(complement(sk1),converse(sk1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_4])).
cnf(c_0_88,negated_conjecture,(join(X1,composition(converse(sk2),X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_83]), c_0_20]), c_0_20])).
cnf(c_0_89,negated_conjecture,(join(complement(X1),join(X2,composition(X1,top)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_84]), c_0_36])).
cnf(c_0_90,negated_conjecture,(join(composition(sk1,top),composition(converse(sk1),top))=composition(sk1,top)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_91,negated_conjecture,(join(composition(sk1,top),complement(converse(sk1)))=join(sk1,complement(converse(sk1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_87]), c_0_10]), c_0_64]), c_0_2]), c_0_64]), c_0_0]), c_0_2]), c_0_44]), c_0_4]), c_0_4])).
cnf(c_0_92,negated_conjecture,(join(converse(X1),composition(converse(X1),sk2))=converse(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_88]), c_0_0]), c_0_2])).
cnf(c_0_93,negated_conjecture,(join(sk1,complement(converse(sk1)))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_4]), c_0_91])).
cnf(c_0_94,negated_conjecture,(join(X1,composition(X1,sk2))=X1), inference(spm,[status(thm)],[c_0_92, c_0_2])).
cnf(c_0_95,negated_conjecture,(join(one,converse(complement(sk2)))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_57]), c_0_44]), c_0_18])).
cnf(c_0_96,negated_conjecture,(join(sk1,converse(sk1))=sk1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_93]), c_0_23]), c_0_47]), c_0_29])).
cnf(c_0_97,negated_conjecture,(join(X1,join(X2,composition(X1,sk2)))=join(X2,X1)), inference(spm,[status(thm)],[c_0_75, c_0_94])).
cnf(c_0_98,negated_conjecture,(join(X1,composition(converse(complement(sk2)),X1))=composition(top,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_95]), c_0_20])).
cnf(c_0_99,negated_conjecture,(join(converse(X1),composition(converse(X1),sk1))=converse(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_85]), c_0_0]), c_0_2])).
cnf(c_0_100,plain,(composition(complement(composition(top,X1)),converse(X1))=zero), inference(spm,[status(thm)],[c_0_74, c_0_2])).
cnf(c_0_101,negated_conjecture,(converse(sk1)=sk1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_96]), c_0_2]), c_0_4]), c_0_96])).
cnf(c_0_102,negated_conjecture,(join(composition(top,sk2),converse(complement(sk2)))=join(sk2,converse(complement(sk2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_4])).
cnf(c_0_103,negated_conjecture,(join(X1,composition(sk2,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_13]), c_0_20]), c_0_4]), c_0_20])).
cnf(c_0_104,plain,(composition(X1,composition(converse(X1),top))=composition(X1,top)), inference(spm,[status(thm)],[c_0_86, c_0_2])).
cnf(c_0_105,negated_conjecture,(join(composition(sk2,top),composition(converse(sk2),top))=composition(sk2,top)), inference(spm,[status(thm)],[c_0_88, c_0_86])).
cnf(c_0_106,negated_conjecture,(join(X1,composition(X1,sk1))=X1), inference(spm,[status(thm)],[c_0_99, c_0_2])).
cnf(c_0_107,negated_conjecture,(join(X1,join(composition(X1,sk2),X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_9, c_0_94])).
cnf(c_0_108,negated_conjecture,(join(composition(X1,X2),composition(X1,composition(sk2,X2)))=composition(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_94]), c_0_14])).
cnf(c_0_109,negated_conjecture,(composition(complement(composition(top,sk1)),sk1)=zero), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_110,negated_conjecture,(join(complement(sk2),composition(converse(sk2),top))=join(complement(sk2),converse(sk2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_102]), c_0_10]), c_0_2]), c_0_0]), c_0_44]), c_0_2]), c_0_4]), c_0_4])).
cnf(c_0_111,negated_conjecture,(composition(converse(sk2),top)=composition(sk2,top)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_4]), c_0_105])).
cnf(c_0_112,negated_conjecture,(join(X1,join(composition(X1,sk1),X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_9, c_0_106])).
cnf(c_0_113,negated_conjecture,(join(X1,composition(X1,composition(sk2,sk1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_106]), c_0_94]), c_0_14])).
cnf(c_0_114,negated_conjecture,(composition(complement(composition(top,sk1)),composition(sk2,sk1))=zero), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_28])).
cnf(c_0_115,negated_conjecture,(join(X1,join(X2,composition(complement(sk1),X1)))=join(X2,composition(top,X1))), inference(spm,[status(thm)],[c_0_75, c_0_82])).
cnf(c_0_116,plain,(join(X1,complement(join(X2,X1)))=join(X1,complement(X2))), inference(spm,[status(thm)],[c_0_45, c_0_4])).
cnf(c_0_117,negated_conjecture,(join(complement(sk2),converse(sk2))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_111]), c_0_84])).
cnf(c_0_118,negated_conjecture,(join(X1,composition(X1,composition(sk1,composition(sk2,sk1))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_106]), c_0_14])).
cnf(c_0_119,negated_conjecture,(composition(top,composition(sk1,composition(sk2,sk1)))=composition(top,composition(sk2,sk1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_114]), c_0_29]), c_0_14]), c_0_28])).
cnf(c_0_120,negated_conjecture,(join(complement(sk1),composition(top,composition(sk2,sk1)))=join(complement(sk1),composition(sk2,sk1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_113]), c_0_4])).
cnf(c_0_121,negated_conjecture,(join(sk2,converse(sk2))=converse(sk2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_23]), c_0_47]), c_0_29]), c_0_4])).
cnf(c_0_122,negated_conjecture,(join(complement(sk1),composition(sk1,composition(sk2,sk1)))=join(complement(sk1),composition(sk2,sk1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_118]), c_0_119]), c_0_4]), c_0_120])).
cnf(c_0_123,negated_conjecture,(converse(sk2)=sk2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_121]), c_0_2]), c_0_2]), c_0_4]), c_0_121])).
cnf(c_0_124,negated_conjecture,(join(complement(sk1),composition(sk2,sk1))=join(complement(sk1),composition(sk1,sk2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_122]), c_0_10]), c_0_64]), c_0_101]), c_0_0]), c_0_101]), c_0_123]), c_0_64]), c_0_101]), c_0_0]), c_0_0]), c_0_101]), c_0_123]), c_0_101]), c_0_14]), c_0_122])).
cnf(c_0_125,negated_conjecture,(join(complement(X1),complement(composition(X1,sk2)))=complement(composition(X1,sk2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_94]), c_0_4])).
cnf(c_0_126,negated_conjecture,(join(complement(X1),complement(composition(sk2,X1)))=complement(composition(sk2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_103]), c_0_4])).
cnf(c_0_127,negated_conjecture,(complement(composition(sk2,sk1))=complement(composition(sk1,sk2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_124]), c_0_45]), c_0_125]), c_0_126])).
cnf(c_0_128,negated_conjecture,(join(X1,join(X2,composition(X1,sk1)))=join(X2,X1)), inference(spm,[status(thm)],[c_0_75, c_0_106])).
cnf(c_0_129,negated_conjecture,(composition(sk2,sk1)=composition(sk1,sk2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_127]), c_0_38])).
cnf(c_0_130,negated_conjecture,(join(complement(sk2),composition(top,sk1))=join(sk1,complement(sk2))), inference(spm,[status(thm)],[c_0_128, c_0_61])).
cnf(c_0_131,negated_conjecture,(join(X1,composition(complement(X1),sk2))=join(X1,composition(top,sk2))), inference(spm,[status(thm)],[c_0_107, c_0_73])).
cnf(c_0_132,negated_conjecture,(join(X1,complement(join(X2,composition(X1,sk1))))=join(X1,complement(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_116]), c_0_112])).
cnf(c_0_133,negated_conjecture,(join(composition(sk1,sk2),composition(complement(sk2),sk1))=composition(top,sk1)), inference(spm,[status(thm)],[c_0_73, c_0_129])).
cnf(c_0_134,negated_conjecture,(join(complement(sk2),complement(composition(top,sk1)))=join(complement(sk1),complement(sk2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_130]), c_0_116]), c_0_4])).
cnf(c_0_135,negated_conjecture,(join(complement(X1),complement(composition(sk1,X1)))=complement(composition(sk1,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_55]), c_0_4])).
cnf(c_0_136,plain,(join(complement(join(X1,complement(X2))),complement(join(complement(X2),complement(X1))))=X2), inference(spm,[status(thm)],[c_0_25, c_0_4])).
cnf(c_0_137,negated_conjecture,(join(X1,complement(composition(complement(X1),sk2)))=join(X1,complement(composition(top,sk2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_131]), c_0_45])).
cnf(c_0_138,negated_conjecture,(join(X1,complement(composition(X1,sk2)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_7]), c_0_36])).
cnf(c_0_139,negated_conjecture,(join(X1,composition(complement(X1),sk1))=join(X1,composition(top,sk1))), inference(spm,[status(thm)],[c_0_112, c_0_73])).
cnf(c_0_140,negated_conjecture,(complement(composition(sk1,sk2))=join(complement(sk1),complement(sk2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_134]), c_0_135])).
cnf(c_0_141,negated_conjecture,(composition(complement(X1),sk2)=complement(join(X1,complement(composition(top,sk2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_4]), c_0_138]), c_0_23]), c_0_47])).
cnf(c_0_142,negated_conjecture,(join(X1,complement(composition(complement(X1),sk1)))=join(X1,complement(composition(top,sk1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_139]), c_0_45])).
cnf(c_0_143,negated_conjecture,(join(X1,complement(composition(X1,sk1)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_7]), c_0_36])).
cnf(c_0_144,plain,(join(X1,join(X2,X3))=join(X3,join(X1,X2))), inference(spm,[status(thm)],[c_0_4, c_0_9])).
cnf(c_0_145,negated_conjecture,(composition(sk1,sk2)=complement(join(complement(sk1),complement(sk2)))), inference(spm,[status(thm)],[c_0_29, c_0_140])).
cnf(c_0_146,negated_conjecture,(composition(sk2,complement(converse(X1)))=complement(join(complement(composition(sk2,top)),converse(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_141]), c_0_64]), c_0_10]), c_0_64]), c_0_0]), c_0_123]), c_0_44]), c_0_123]), c_0_64]), c_0_4])).
cnf(c_0_147,negated_conjecture,(composition(complement(X1),sk1)=complement(join(X1,complement(composition(top,sk1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_142]), c_0_4]), c_0_143]), c_0_23]), c_0_47])).
cnf(c_0_148,plain,(join(complement(X1),join(X2,join(X3,X1)))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_51]), c_0_36]), c_0_9])).
cnf(c_0_149,negated_conjecture,(complement(join(complement(composition(sk1,sk3)),complement(composition(sk2,sk3))))!=composition(complement(join(complement(sk1),complement(sk2))),sk3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_6]), c_0_6])).
cnf(c_0_150,negated_conjecture,(composition(complement(join(complement(sk1),complement(sk2))),X1)=composition(sk1,composition(sk2,X1))), inference(spm,[status(thm)],[c_0_14, c_0_145])).
cnf(c_0_151,negated_conjecture,(composition(sk2,complement(X1))=complement(join(complement(composition(sk2,top)),X1))), inference(spm,[status(thm)],[c_0_146, c_0_2])).
cnf(c_0_152,negated_conjecture,(composition(sk1,complement(converse(X1)))=complement(join(complement(composition(sk1,top)),converse(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_147]), c_0_64]), c_0_10]), c_0_64]), c_0_0]), c_0_101]), c_0_44]), c_0_101]), c_0_64]), c_0_4])).
cnf(c_0_153,negated_conjecture,(join(X1,join(complement(composition(sk1,X1)),X2))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_55]), c_0_144])).
cnf(c_0_154,negated_conjecture,(composition(sk1,composition(sk2,sk3))!=complement(join(complement(composition(sk1,sk3)),complement(composition(sk2,sk3))))), inference(rw,[status(thm)],[c_0_149, c_0_150])).
cnf(c_0_155,negated_conjecture,(composition(sk2,X1)=complement(join(complement(composition(sk2,top)),complement(X1)))), inference(spm,[status(thm)],[c_0_151, c_0_29])).
cnf(c_0_156,negated_conjecture,(composition(sk1,complement(X1))=complement(join(complement(composition(sk1,top)),X1))), inference(spm,[status(thm)],[c_0_152, c_0_2])).
cnf(c_0_157,negated_conjecture,(join(complement(X1),join(complement(composition(sk1,X1)),X2))=join(complement(composition(sk1,X1)),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_153]), c_0_23]), c_0_9]), c_0_47]), c_0_9]), c_0_144])).
cnf(c_0_158,negated_conjecture,(complement(join(complement(composition(sk1,sk3)),complement(composition(sk2,top))))!=complement(join(complement(sk3),join(complement(composition(sk1,top)),complement(composition(sk2,top)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_154, c_0_155]), c_0_4]), c_0_156]), c_0_75]), c_0_155]), c_0_4]), c_0_29]), c_0_75]), c_0_157])).
cnf(c_0_159,negated_conjecture,(composition(sk1,X1)=complement(join(complement(composition(sk1,top)),complement(X1)))), inference(spm,[status(thm)],[c_0_156, c_0_29])).
cnf(c_0_160,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_158, c_0_159]), c_0_4]), c_0_29]), c_0_9])]), ['proof']).
# SZS output end CNFRefutation.
EOF
