% Problem    : GRP729-1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n097.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:08:56 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___060_C18_F1_PI_AE_CS_SP_CO_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(mult(i(X1),mult(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c06)).
cnf(c_0_1,axiom,(mult(i(X1),X1)=unit), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c04)).
cnf(c_0_2,axiom,(mult(X1,unit)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c02)).
cnf(c_0_3,axiom,(mult(mult(X1,X2),X3)=mult(mult(X1,mult(X2,X3)),asoc(X1,X2,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c10)).
cnf(c_0_4,axiom,(asoc(asoc(X1,X2,X3),X4,X5)=unit), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c21)).
cnf(c_0_5,axiom,(i(mult(X1,X2))=mult(i(X1),i(X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c05)).
cnf(c_0_6,axiom,(mult(X1,X2)=mult(mult(X2,X1),op_k(X1,X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c11)).
cnf(c_0_7,axiom,(asoc(X1,X2,asoc(X3,X4,X5))=unit), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c22)).
cnf(c_0_8,axiom,(mult(rd(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c08)).
cnf(c_0_9,axiom,(rd(mult(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c07)).
cnf(c_0_10,axiom,(mult(mult(X1,mult(X2,X1)),X3)=mult(X1,mult(X2,mult(X1,X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c09)).
cnf(c_0_11,axiom,(mult(X1,i(X1))=unit), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c03)).
cnf(c_0_12,negated_conjecture,(asoc(a,b,op_k(c,d))!=unit), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
cnf(c_0_13,plain,(i(i(X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_14,plain,(mult(X1,asoc(X1,i(X2),X2))=mult(mult(X1,i(X2)),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_3, c_0_1]), c_0_2])).
cnf(c_0_15,plain,(mult(X1,mult(i(X1),X2))=X2), inference(spm,[status(thm)],[c_0_0, c_0_13])).
cnf(c_0_16,plain,(mult(mult(asoc(X1,X2,X3),i(X4)),X4)=asoc(X1,X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_4]), c_0_2])).
cnf(c_0_17,plain,(mult(X1,i(mult(X1,X2)))=i(X2)), inference(spm,[status(thm)],[c_0_15, c_0_5])).
cnf(c_0_18,plain,(mult(i(mult(X1,X2)),mult(X2,X1))=op_k(X2,X1)), inference(spm,[status(thm)],[c_0_0, c_0_6])).
cnf(c_0_19,plain,(mult(i(X1),mult(asoc(X2,X3,X4),X1))=asoc(X2,X3,X4)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,plain,(mult(mult(asoc(X1,X2,X3),X4),i(X4))=asoc(X1,X2,X3)), inference(spm,[status(thm)],[c_0_16, c_0_13])).
cnf(c_0_21,plain,(op_k(i(X1),mult(asoc(X2,X3,X4),X1))=unit), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_1])).
cnf(c_0_22,plain,(mult(X1,asoc(X2,X3,X4))=mult(asoc(X2,X3,X4),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_19]), c_0_13])).
cnf(c_0_23,plain,(op_k(i(X1),mult(X1,asoc(X2,X3,X4)))=unit), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_24,plain,(mult(mult(X1,X2),asoc(X3,X4,X5))=mult(X1,mult(X2,asoc(X3,X4,X5)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_3, c_0_7]), c_0_2])).
cnf(c_0_25,plain,(mult(rd(X1,X2),i(X1))=i(X2)), inference(spm,[status(thm)],[c_0_17, c_0_8])).
cnf(c_0_26,plain,(rd(X1,mult(X2,X1))=i(X2)), inference(spm,[status(thm)],[c_0_9, c_0_0])).
cnf(c_0_27,plain,(op_k(i(mult(X1,X2)),mult(X1,mult(X2,asoc(X3,X4,X5))))=unit), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28,plain,(mult(rd(i(X1),X2),X1)=i(X2)), inference(spm,[status(thm)],[c_0_25, c_0_13])).
cnf(c_0_29,plain,(i(rd(X1,X2))=rd(X2,X1)), inference(spm,[status(thm)],[c_0_26, c_0_8])).
cnf(c_0_30,plain,(op_k(X1,mult(X2,mult(i(mult(X2,X1)),asoc(X3,X4,X5))))=unit), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_17]), c_0_13])).
cnf(c_0_31,plain,(rd(mult(X1,mult(X2,mult(X1,X3))),X3)=mult(X1,mult(X2,X1))), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_32,plain,(rd(X1,i(X2))=rd(X2,i(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_28]), c_0_29])).
cnf(c_0_33,plain,(op_k(X1,mult(X2,mult(i(mult(mult(X2,X1),X3)),X3)))=unit), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_14]), c_0_5])).
cnf(c_0_34,plain,(mult(X1,mult(i(mult(X1,X2)),X1))=rd(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_1]), c_0_2])).
cnf(c_0_35,plain,(i(mult(i(X1),X2))=mult(X1,i(X2))), inference(spm,[status(thm)],[c_0_5, c_0_13])).
cnf(c_0_36,plain,(rd(X1,i(mult(X1,X2)))=mult(X1,mult(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_11]), c_0_2]), c_0_32])).
cnf(c_0_37,plain,(rd(i(X1),X2)=rd(i(X2),X1)), inference(spm,[status(thm)],[c_0_9, c_0_28])).
cnf(c_0_38,plain,(op_k(X1,mult(X2,mult(X3,i(mult(mult(X2,X1),X3)))))=unit), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_17]), c_0_13])).
cnf(c_0_39,plain,(mult(mult(X1,X2),i(X1))=mult(X1,rd(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_34]), c_0_29]), c_0_35])).
cnf(c_0_40,plain,(i(mult(X1,rd(X2,X3)))=mult(i(X1),rd(X3,X2))), inference(spm,[status(thm)],[c_0_5, c_0_29])).
cnf(c_0_41,plain,(mult(i(X1),rd(X1,X2))=mult(i(mult(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_0, c_0_34])).
cnf(c_0_42,plain,(rd(i(X1),mult(X1,X2))=i(mult(X1,mult(X2,X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_36]), c_0_37])).
cnf(c_0_43,plain,(mult(mult(X1,X2),i(mult(X2,X1)))=i(op_k(X2,X1))), inference(spm,[status(thm)],[c_0_17, c_0_6])).
cnf(c_0_44,plain,(op_k(X1,mult(i(mult(X2,X1)),X2))=unit), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_41]), c_0_15])).
cnf(c_0_45,plain,(mult(i(mult(X1,mult(X2,X1))),mult(X1,X2))=i(X1)), inference(spm,[status(thm)],[c_0_8, c_0_42])).
cnf(c_0_46,plain,(op_k(mult(asoc(X1,X2,X3),X4),i(X4))=unit), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_1])).
cnf(c_0_47,plain,(i(op_k(i(X1),i(X2)))=op_k(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_5]), c_0_5]), c_0_13]), c_0_18])).
cnf(c_0_48,plain,(rd(op_k(X1,X2),mult(X1,X2))=i(mult(X2,X1))), inference(spm,[status(thm)],[c_0_26, c_0_6])).
cnf(c_0_49,plain,(op_k(mult(X1,X2),i(mult(X2,X1)))=unit), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_13]), c_0_17])).
cnf(c_0_50,plain,(rd(unit,X1)=i(X1)), inference(spm,[status(thm)],[c_0_9, c_0_1])).
cnf(c_0_51,plain,(i(mult(X1,i(X2)))=mult(i(X1),X2)), inference(spm,[status(thm)],[c_0_5, c_0_13])).
cnf(c_0_52,plain,(op_k(X1,mult(asoc(X2,X3,X4),i(X1)))=unit), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_15]), c_0_35])).
cnf(c_0_53,plain,(op_k(i(X1),i(X2))=i(op_k(X1,X2))), inference(spm,[status(thm)],[c_0_13, c_0_47])).
cnf(c_0_54,plain,(i(op_k(X1,X2))=op_k(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_51]), c_0_18]), c_0_18])).
cnf(c_0_55,plain,(op_k(X1,mult(i(X1),asoc(X2,X3,X4)))=unit), inference(spm,[status(thm)],[c_0_52, c_0_22])).
cnf(c_0_56,plain,(op_k(i(X1),i(X2))=op_k(X2,X1)), inference(rw,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_57,plain,(op_k(X1,mult(i(mult(X1,X2)),X2))=unit), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_14]), c_0_5])).
cnf(c_0_58,plain,(mult(X1,mult(rd(X2,X1),mult(X1,X3)))=mult(mult(X1,X2),X3)), inference(spm,[status(thm)],[c_0_10, c_0_8])).
cnf(c_0_59,plain,(mult(X1,asoc(X1,X2,i(X2)))=mult(mult(X1,X2),i(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_3, c_0_11]), c_0_2])).
cnf(c_0_60,plain,(mult(i(mult(X1,X2)),op_k(X1,X2))=i(mult(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_56]), c_0_5]), c_0_5])).
cnf(c_0_61,plain,(op_k(rd(X1,X2),mult(i(X1),X2))=unit), inference(spm,[status(thm)],[c_0_57, c_0_8])).
cnf(c_0_62,plain,(i(mult(rd(X1,X2),X3))=mult(rd(X2,X1),i(X3))), inference(spm,[status(thm)],[c_0_5, c_0_29])).
cnf(c_0_63,plain,(mult(rd(X1,X2),mult(X2,X3))=mult(i(X2),mult(mult(X2,X1),X3))), inference(spm,[status(thm)],[c_0_0, c_0_58])).
cnf(c_0_64,plain,(mult(i(X1),mult(mult(X1,X2),i(X2)))=asoc(X1,X2,i(X2))), inference(spm,[status(thm)],[c_0_0, c_0_59])).
cnf(c_0_65,plain,(mult(i(mult(X1,mult(X2,X1))),X1)=i(mult(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_42]), c_0_13])).
cnf(c_0_66,plain,(i(op_k(X1,i(X2)))=op_k(i(X1),X2)), inference(spm,[status(thm)],[c_0_53, c_0_13])).
cnf(c_0_67,plain,(mult(mult(X1,i(X2)),rd(X2,X1))=asoc(X1,X2,i(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_35]), c_0_63]), c_0_64]), c_0_2]), c_0_40]), c_0_35])).
cnf(c_0_68,plain,(op_k(i(X1),mult(X1,mult(X2,X1)))=op_k(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_65]), c_0_17]), c_0_13]), c_0_18]), c_0_66])).
cnf(c_0_69,plain,(rd(mult(X1,X2),op_k(X1,X2))=mult(X2,X1)), inference(spm,[status(thm)],[c_0_9, c_0_6])).
cnf(c_0_70,plain,(rd(X1,unit)=X1), inference(spm,[status(thm)],[c_0_2, c_0_8])).
cnf(c_0_71,plain,(mult(i(X1),mult(X2,asoc(X1,X2,i(X2))))=rd(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_67]), c_0_51]), c_0_24])).
cnf(c_0_72,plain,(mult(mult(asoc(X1,X2,X3),X4),X5)=mult(asoc(X1,X2,X3),mult(X4,X5))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_3, c_0_4]), c_0_2])).
cnf(c_0_73,plain,(op_k(rd(X1,X2),X2)=op_k(i(X2),mult(X2,X1))), inference(spm,[status(thm)],[c_0_68, c_0_8])).
cnf(c_0_74,plain,(i(op_k(i(X1),X2))=op_k(X1,i(X2))), inference(spm,[status(thm)],[c_0_53, c_0_13])).
cnf(c_0_75,plain,(mult(mult(X1,X2),i(mult(X2,X1)))=op_k(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_49]), c_0_70]), c_0_18])).
cnf(c_0_76,plain,(mult(X1,rd(X2,X1))=mult(X2,asoc(X1,X2,i(X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_71]), c_0_13])).
cnf(c_0_77,plain,(mult(mult(X1,asoc(X2,X3,X4)),X5)=mult(asoc(X2,X3,X4),mult(X1,X5))), inference(spm,[status(thm)],[c_0_72, c_0_22])).
cnf(c_0_78,plain,(op_k(X1,rd(X2,X1))=op_k(X1,i(mult(X1,X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_73]), c_0_74])).
cnf(c_0_79,plain,(asoc(X1,X2,i(X2))=op_k(X1,i(mult(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_8]), c_0_77]), c_0_11]), c_0_2]), c_0_78])).
cnf(c_0_80,plain,(op_k(i(mult(X1,mult(X2,X1))),X1)=i(op_k(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_65]), c_0_17]), c_0_53])).
cnf(c_0_81,plain,(asoc(X1,X2,op_k(X3,i(mult(X3,X4))))=unit), inference(spm,[status(thm)],[c_0_7, c_0_79])).
cnf(c_0_82,plain,(op_k(i(mult(X1,mult(X2,X1))),X1)=op_k(X2,X1)), inference(rw,[status(thm)],[c_0_80, c_0_54])).
cnf(c_0_83,plain,(asoc(X1,X2,op_k(X3,X4))=unit), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_45]), c_0_13]), c_0_82])).
cnf(c_0_84,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_83])]), ['proof']).
# SZS output end CNFRefutation.
EOF
