% Problem    : GRP185-1 : TPTP v6.1.0. Bugfixed v1.2.1.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n135.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:12:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4d
# and selection function SelectCQIPrecWNTNp.
#
# No success with AutoSched1
# Trying AutoSched2 for 24 seconds
# AutoSched2-Mode selected heuristic G_E___107_B03_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(multiply(identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_identity)).
cnf(c_0_1,axiom,(multiply(inverse(X1),X1)=identity), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', left_inverse)).
cnf(c_0_2,axiom,(multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-0.ax', associativity)).
cnf(c_0_3,negated_conjecture,(least_upper_bound(least_upper_bound(multiply(a,b),identity),multiply(least_upper_bound(a,identity),least_upper_bound(b,identity)))!=multiply(least_upper_bound(a,identity),least_upper_bound(b,identity))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_p22a)).
cnf(c_0_4,axiom,(least_upper_bound(X1,X2)=least_upper_bound(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', symmetry_of_lub)).
cnf(c_0_5,axiom,(multiply(least_upper_bound(X1,X2),X3)=least_upper_bound(multiply(X1,X3),multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_lub2)).
cnf(c_0_6,axiom,(multiply(X1,least_upper_bound(X2,X3))=least_upper_bound(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', monotony_lub1)).
cnf(c_0_7,axiom,(least_upper_bound(X1,least_upper_bound(X2,X3))=least_upper_bound(least_upper_bound(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', associativity_of_lub)).
cnf(c_0_8,axiom,(least_upper_bound(X1,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/GRP004-2.ax', idempotence_of_lub)).
cnf(c_0_9,plain,(multiply(inverse(X1),multiply(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_10,plain,(multiply(inverse(inverse(X1)),identity)=X1), inference(spm,[status(thm)],[c_0_9, c_0_1])).
cnf(c_0_11,plain,(multiply(inverse(inverse(inverse(X1))),X1)=identity), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12,plain,(inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_11]), c_0_10])).
cnf(c_0_13,negated_conjecture,(least_upper_bound(identity,least_upper_bound(multiply(a,b),least_upper_bound(identity,least_upper_bound(b,least_upper_bound(multiply(a,identity),multiply(a,b))))))!=least_upper_bound(identity,least_upper_bound(b,least_upper_bound(multiply(a,identity),multiply(a,b))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_4]), c_0_4]), c_0_4]), c_0_5]), c_0_0]), c_0_6]), c_0_7]), c_0_7]), c_0_4]), c_0_4]), c_0_5]), c_0_0]), c_0_6]), c_0_7])).
cnf(c_0_14,plain,(multiply(X1,identity)=X1), inference(rw,[status(thm)],[c_0_10, c_0_12])).
cnf(c_0_15,negated_conjecture,(least_upper_bound(identity,least_upper_bound(multiply(a,b),least_upper_bound(identity,least_upper_bound(b,least_upper_bound(a,multiply(a,b))))))!=least_upper_bound(identity,least_upper_bound(b,least_upper_bound(a,multiply(a,b))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_14])).
cnf(c_0_16,negated_conjecture,(least_upper_bound(multiply(inverse(X1),X1),least_upper_bound(multiply(a,b),least_upper_bound(multiply(inverse(X1),X1),least_upper_bound(b,least_upper_bound(a,multiply(a,b))))))!=least_upper_bound(multiply(inverse(X1),X1),least_upper_bound(b,least_upper_bound(a,multiply(a,b))))), inference(spm,[status(thm)],[c_0_15, c_0_1])).
cnf(c_0_17,plain,(least_upper_bound(X1,least_upper_bound(X2,X3))=least_upper_bound(X3,least_upper_bound(X1,X2))), inference(spm,[status(thm)],[c_0_4, c_0_7])).
cnf(c_0_18,plain,(least_upper_bound(X1,least_upper_bound(X1,X2))=least_upper_bound(X1,X2)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_19,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_7]), c_0_7]), c_0_8]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation.
EOF
