% Problem    : LAT043-1 : TPTP v6.1.0. Released v2.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n110.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:42:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C02_F1_PI_AE_Q12_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(meet(complement(X1),X1)=n0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', invertability2)).
cnf(c_0_1,axiom,(meet(X1,X2)=meet(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/LAT001-0.ax', commutativity_of_meet)).
cnf(c_0_2,axiom,(join(complement(X1),X1)=n1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', invertability1)).
cnf(c_0_3,axiom,(join(X1,X2)=join(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/LAT001-0.ax', commutativity_of_join)).
cnf(c_0_4,axiom,(join(X1,meet(X1,X2))=X1), file('/export/starexec/sandbox/benchmark/Axioms/LAT001-0.ax', absorption2)).
cnf(c_0_5,axiom,(meet(X1,join(X1,X2))=X1), file('/export/starexec/sandbox/benchmark/Axioms/LAT001-0.ax', absorption1)).
cnf(c_0_6,axiom,(meet(X1,join(X2,X3))=join(meet(X1,X2),meet(X1,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', distributivity)).
cnf(c_0_7,axiom,(join(join(X1,X2),X3)=join(X1,join(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/LAT001-0.ax', associativity_of_join)).
cnf(c_0_8,axiom,(complement(complement(X1))=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', invertability3)).
cnf(c_0_9,negated_conjecture,(complement(join(c,d))!=meet(complement(c),complement(d))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_compatability_law)).
cnf(c_0_10,plain,(meet(X1,complement(X1))=n0), inference(rw,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_11,plain,(join(X1,complement(X1))=n1), inference(rw,[status(thm)],[c_0_2, c_0_3])).
cnf(c_0_12,plain,(join(X1,n0)=X1), inference(spm,[status(thm)],[c_0_4, c_0_10])).
cnf(c_0_13,plain,(join(X1,meet(X2,X1))=X1), inference(spm,[status(thm)],[c_0_4, c_0_1])).
cnf(c_0_14,plain,(meet(X1,n1)=X1), inference(spm,[status(thm)],[c_0_5, c_0_11])).
cnf(c_0_15,plain,(join(n0,X1)=X1), inference(spm,[status(thm)],[c_0_3, c_0_12])).
cnf(c_0_16,plain,(meet(X1,join(X2,complement(X1)))=meet(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_10]), c_0_12])).
cnf(c_0_17,plain,(join(n1,X1)=n1), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18,plain,(join(X1,join(meet(X2,X1),X3))=join(X1,X3)), inference(spm,[status(thm)],[c_0_7, c_0_13])).
cnf(c_0_19,plain,(meet(X1,join(complement(X1),X2))=meet(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_10]), c_0_15])).
cnf(c_0_20,plain,(join(X1,join(X2,complement(join(X1,X2))))=n1), inference(spm,[status(thm)],[c_0_7, c_0_11])).
cnf(c_0_21,plain,(meet(complement(X1),join(X2,X1))=meet(complement(X1),X2)), inference(spm,[status(thm)],[c_0_16, c_0_8])).
cnf(c_0_22,plain,(join(X1,join(complement(X1),X2))=n1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_11]), c_0_17])).
cnf(c_0_23,plain,(join(X1,join(X2,X1))=join(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_5]), c_0_7])).
cnf(c_0_24,plain,(join(X1,n1)=n1), inference(spm,[status(thm)],[c_0_3, c_0_17])).
cnf(c_0_25,plain,(join(X1,meet(X2,join(X1,X3)))=join(X1,meet(X2,X3))), inference(spm,[status(thm)],[c_0_18, c_0_6])).
cnf(c_0_26,plain,(meet(X1,join(X2,complement(join(complement(X1),X2))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_14])).
cnf(c_0_27,plain,(meet(X1,complement(join(complement(X1),X2)))=complement(join(complement(X1),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_14]), c_0_1])).
cnf(c_0_28,plain,(join(complement(X1),join(X2,X1))=n1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_22]), c_0_7]), c_0_24]), c_0_24]), c_0_7])).
cnf(c_0_29,plain,(meet(complement(X1),join(X1,X2))=meet(complement(X1),X2)), inference(spm,[status(thm)],[c_0_19, c_0_8])).
cnf(c_0_30,plain,(join(X1,complement(join(complement(X2),X1)))=join(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
cnf(c_0_31,plain,(meet(complement(X1),complement(join(X2,X1)))=complement(join(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_28]), c_0_14]), c_0_1])).
cnf(c_0_32,plain,(complement(join(complement(X1),X2))=meet(complement(X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_29]), c_0_31])).
cnf(c_0_33,plain,(meet(complement(X1),complement(X2))=complement(join(X2,X1))), inference(spm,[status(thm)],[c_0_32, c_0_8])).
cnf(c_0_34,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_33]), c_0_3])]), ['proof']).
# SZS output end CNFRefutation.
EOF
