% Problem    : REL042-1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n093.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:27:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(converse(composition(X1,X2))=composition(converse(X2),converse(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_multiplicativity_10)).
cnf(c_0_1,axiom,(converse(converse(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_idempotence_8)).
cnf(c_0_2,axiom,(composition(X1,one)=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_identity_6)).
cnf(c_0_3,axiom,(join(composition(converse(X1),complement(composition(X1,X2))),complement(X2))=complement(X2)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_cancellativity_11)).
cnf(c_0_4,axiom,(join(X1,X2)=join(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux1_join_commutativity_1)).
cnf(c_0_5,axiom,(zero=meet(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', def_zero_13)).
cnf(c_0_6,axiom,(meet(X1,X2)=complement(join(complement(X1),complement(X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux4_definiton_of_meet_4)).
cnf(c_0_7,axiom,(top=join(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', def_top_12)).
cnf(c_0_8,axiom,(X1=join(complement(join(complement(X1),complement(X2))),complement(join(complement(X1),X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux3_a_kind_of_de_Morgan_3)).
cnf(c_0_9,axiom,(join(X1,join(X2,X3))=join(join(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux2_join_associativity_2)).
cnf(c_0_10,axiom,(converse(join(X1,X2))=join(converse(X1),converse(X2))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_additivity_9)).
cnf(c_0_11,negated_conjecture,(meet(composition(sk1,X1),composition(sk1,complement(X1)))=zero), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_14)).
cnf(c_0_12,axiom,(composition(join(X1,X2),X3)=join(composition(X1,X3),composition(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_distributivity_7)).
cnf(c_0_13,axiom,(composition(X1,composition(X2,X3))=composition(composition(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_associativity_5)).
cnf(c_0_14,negated_conjecture,(join(composition(converse(sk1),sk1),one)!=one), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_15)).
cnf(c_0_15,plain,(converse(composition(converse(X1),X2))=composition(converse(X2),X1)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_16,plain,(composition(converse(one),X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_2]), c_0_1])).
cnf(c_0_17,plain,(converse(one)=one), inference(spm,[status(thm)],[c_0_2, c_0_16])).
cnf(c_0_18,plain,(join(complement(X1),composition(converse(X2),complement(composition(X2,X1))))=complement(X1)), inference(rw,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_19,plain,(composition(one,X1)=X1), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,plain,(complement(join(complement(X1),complement(complement(X1))))=zero), inference(rw,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_21,plain,(join(complement(X1),complement(X1))=complement(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_17]), c_0_19])).
cnf(c_0_22,plain,(complement(top)=zero), inference(rw,[status(thm)],[c_0_20, c_0_7])).
cnf(c_0_23,plain,(join(zero,zero)=zero), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_24,plain,(join(complement(join(complement(X1),X2)),complement(join(complement(X1),complement(X2))))=X1), inference(rw,[status(thm)],[c_0_8, c_0_4])).
cnf(c_0_25,plain,(join(zero,join(zero,X1))=join(zero,X1)), inference(spm,[status(thm)],[c_0_9, c_0_23])).
cnf(c_0_26,plain,(join(zero,complement(complement(X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_21]), c_0_7]), c_0_22]), c_0_4])).
cnf(c_0_27,plain,(join(zero,X1)=X1), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_28,plain,(complement(complement(X1))=X1), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_29,plain,(join(X1,X1)=X1), inference(spm,[status(thm)],[c_0_21, c_0_28])).
cnf(c_0_30,plain,(join(X1,join(X1,X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_9, c_0_29])).
cnf(c_0_31,plain,(converse(join(converse(X1),X2))=join(X1,converse(X2))), inference(spm,[status(thm)],[c_0_10, c_0_1])).
cnf(c_0_32,plain,(join(X1,top)=top), inference(spm,[status(thm)],[c_0_30, c_0_7])).
cnf(c_0_33,plain,(join(X1,converse(complement(converse(X1))))=converse(top)), inference(spm,[status(thm)],[c_0_31, c_0_7])).
cnf(c_0_34,plain,(join(top,X1)=top), inference(spm,[status(thm)],[c_0_4, c_0_32])).
cnf(c_0_35,plain,(converse(top)=top), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_36,plain,(join(X1,converse(complement(converse(X1))))=top), inference(rw,[status(thm)],[c_0_33, c_0_35])).
cnf(c_0_37,plain,(complement(join(complement(X1),complement(converse(complement(converse(complement(X1)))))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_36]), c_0_22]), c_0_27])).
cnf(c_0_38,plain,(join(X1,complement(join(complement(X1),X2)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_24]), c_0_4])).
cnf(c_0_39,plain,(join(X1,join(X2,X1))=join(X2,X1)), inference(spm,[status(thm)],[c_0_30, c_0_4])).
cnf(c_0_40,plain,(join(complement(one),composition(converse(X1),complement(X1)))=complement(one)), inference(spm,[status(thm)],[c_0_18, c_0_2])).
cnf(c_0_41,plain,(complement(join(X1,complement(converse(complement(converse(X1))))))=complement(X1)), inference(spm,[status(thm)],[c_0_37, c_0_28])).
cnf(c_0_42,plain,(join(X1,complement(join(X2,complement(X1))))=X1), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43,plain,(join(complement(one),composition(converse(complement(X1)),X1))=complement(one)), inference(spm,[status(thm)],[c_0_40, c_0_28])).
cnf(c_0_44,negated_conjecture,(complement(join(complement(composition(sk1,X1)),complement(composition(sk1,complement(X1)))))=zero), inference(rw,[status(thm)],[c_0_11, c_0_6])).
cnf(c_0_45,plain,(join(X1,complement(converse(complement(converse(X1)))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_41]), c_0_24])).
cnf(c_0_46,plain,(join(complement(join(complement(X1),X2)),complement(join(complement(X2),complement(X1))))=X1), inference(spm,[status(thm)],[c_0_24, c_0_4])).
cnf(c_0_47,plain,(join(complement(join(X1,X2)),complement(join(X1,complement(X2))))=complement(X1)), inference(spm,[status(thm)],[c_0_24, c_0_28])).
cnf(c_0_48,plain,(join(complement(X1),complement(join(X2,X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_42, c_0_28])).
cnf(c_0_49,plain,(converse(join(composition(converse(X1),X2),X3))=join(composition(converse(X2),X1),converse(X3))), inference(spm,[status(thm)],[c_0_31, c_0_15])).
cnf(c_0_50,plain,(join(composition(X1,converse(X2)),converse(composition(X2,X3)))=composition(join(X1,converse(X3)),converse(X2))), inference(spm,[status(thm)],[c_0_12, c_0_0])).
cnf(c_0_51,plain,(converse(join(X1,converse(X2)))=join(converse(X1),X2)), inference(spm,[status(thm)],[c_0_10, c_0_1])).
cnf(c_0_52,plain,(join(complement(one),converse(complement(one)))=complement(one)), inference(spm,[status(thm)],[c_0_43, c_0_2])).
cnf(c_0_53,plain,(join(X1,join(X2,X3))=join(X3,join(X1,X2))), inference(spm,[status(thm)],[c_0_4, c_0_9])).
cnf(c_0_54,negated_conjecture,(complement(join(complement(sk1),complement(composition(sk1,complement(one)))))=zero), inference(spm,[status(thm)],[c_0_44, c_0_2])).
cnf(c_0_55,plain,(join(X1,converse(complement(converse(complement(X1)))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_45]), c_0_1]), c_0_1])).
cnf(c_0_56,plain,(join(X1,complement(join(X1,complement(X2))))=join(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_28]), c_0_28]), c_0_9]), c_0_48])).
cnf(c_0_57,plain,(join(X1,zero)=X1), inference(spm,[status(thm)],[c_0_4, c_0_27])).
cnf(c_0_58,plain,(join(composition(X1,X2),composition(X1,X3))=composition(X1,join(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_10]), c_0_0]), c_0_1]), c_0_1]), c_0_1])).
cnf(c_0_59,plain,(join(complement(X1),composition(X2,complement(composition(converse(X2),X1))))=complement(X1)), inference(spm,[status(thm)],[c_0_18, c_0_1])).
cnf(c_0_60,plain,(join(X1,composition(X2,X1))=composition(join(one,X2),X1)), inference(spm,[status(thm)],[c_0_12, c_0_19])).
cnf(c_0_61,plain,(converse(complement(one))=complement(one)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_4]), c_0_52])).
cnf(c_0_62,plain,(join(X1,join(complement(X1),X2))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_7]), c_0_32])).
cnf(c_0_63,negated_conjecture,(complement(join(complement(sk1),composition(sk1,complement(one))))=sk1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_54]), c_0_28]), c_0_27])).
cnf(c_0_64,plain,(join(complement(X1),converse(complement(converse(X1))))=complement(X1)), inference(spm,[status(thm)],[c_0_55, c_0_28])).
cnf(c_0_65,plain,(join(complement(join(X1,complement(X2))),complement(join(complement(X2),complement(X1))))=X2), inference(spm,[status(thm)],[c_0_24, c_0_4])).
cnf(c_0_66,plain,(join(X1,complement(join(X1,X2)))=join(X1,complement(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_56]), c_0_30])).
cnf(c_0_67,negated_conjecture,(join(complement(composition(sk1,X1)),complement(composition(sk1,complement(X1))))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_44]), c_0_9]), c_0_57])).
cnf(c_0_68,plain,(join(X1,composition(X1,X2))=composition(X1,join(X2,one))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_2]), c_0_4])).
cnf(c_0_69,plain,(join(X1,composition(X2,complement(composition(converse(X2),complement(X1)))))=X1), inference(spm,[status(thm)],[c_0_59, c_0_28])).
cnf(c_0_70,plain,(join(X1,composition(complement(one),X1))=composition(top,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_36]), c_0_17]), c_0_61])).
cnf(c_0_71,plain,(converse(composition(X1,converse(X2)))=composition(X2,converse(X1))), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_72,plain,(join(X1,join(X2,complement(X1)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_28]), c_0_53])).
cnf(c_0_73,plain,(join(converse(X1),join(converse(X2),X3))=join(converse(join(X1,X2)),X3)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_74,negated_conjecture,(join(complement(sk1),composition(sk1,complement(one)))=complement(sk1)), inference(spm,[status(thm)],[c_0_28, c_0_63])).
cnf(c_0_75,plain,(complement(converse(complement(converse(X1))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_64]), c_0_28]), c_0_4]), c_0_45])).
cnf(c_0_76,plain,(join(complement(join(X1,X2)),complement(join(X2,complement(X1))))=complement(X2)), inference(spm,[status(thm)],[c_0_65, c_0_28])).
cnf(c_0_77,negated_conjecture,(join(complement(composition(sk1,X1)),composition(sk1,complement(X1)))=complement(composition(sk1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_22]), c_0_57]), c_0_28])).
cnf(c_0_78,plain,(composition(X1,join(one,complement(composition(converse(X1),complement(X1)))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_4])).
cnf(c_0_79,plain,(join(one,complement(composition(converse(X1),complement(X1))))=complement(composition(converse(X1),complement(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_40]), c_0_28]), c_0_4])).
cnf(c_0_80,plain,(join(complement(join(X1,X2)),complement(join(complement(X2),X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_46, c_0_28])).
cnf(c_0_81,plain,(join(X1,composition(X1,complement(one)))=composition(X1,top)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_70]), c_0_71]), c_0_35]), c_0_71]), c_0_61])).
cnf(c_0_82,plain,(join(complement(converse(X1)),converse(join(X1,X2)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_4])).
cnf(c_0_83,negated_conjecture,(join(sk1,complement(composition(sk1,complement(one))))=complement(composition(sk1,complement(one)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_74]), c_0_28]), c_0_4])).
cnf(c_0_84,plain,(converse(complement(converse(X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_28, c_0_75])).
cnf(c_0_85,negated_conjecture,(join(composition(sk1,X1),complement(composition(sk1,top)))=complement(composition(sk1,complement(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_28]), c_0_28]), c_0_58]), c_0_4]), c_0_7])).
cnf(c_0_86,plain,(composition(X1,complement(composition(converse(X1),complement(X1))))=X1), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_87,negated_conjecture,(complement(composition(sk1,complement(one)))=join(sk1,complement(composition(sk1,top)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_74]), c_0_4]), c_0_81]), c_0_28]), c_0_4])).
cnf(c_0_88,negated_conjecture,(join(complement(converse(sk1)),converse(complement(composition(sk1,complement(one)))))=top), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_89,plain,(converse(complement(X1))=complement(converse(X1))), inference(spm,[status(thm)],[c_0_1, c_0_84])).
cnf(c_0_90,plain,(converse(composition(X1,complement(one)))=composition(complement(one),converse(X1))), inference(spm,[status(thm)],[c_0_0, c_0_61])).
cnf(c_0_91,negated_conjecture,(complement(composition(sk1,composition(converse(sk1),complement(sk1))))=join(sk1,complement(composition(sk1,top)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_28])).
cnf(c_0_92,negated_conjecture,(complement(join(sk1,complement(composition(sk1,top))))=composition(sk1,complement(one))), inference(spm,[status(thm)],[c_0_28, c_0_87])).
cnf(c_0_93,negated_conjecture,(join(complement(converse(sk1)),complement(composition(complement(one),converse(sk1))))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_89]), c_0_90])).
cnf(c_0_94,plain,(converse(composition(X1,composition(converse(X2),X3)))=composition(converse(X3),composition(X2,converse(X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_15]), c_0_13])).
cnf(c_0_95,negated_conjecture,(composition(sk1,composition(converse(sk1),complement(sk1)))=composition(sk1,complement(one))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_91]), c_0_92])).
cnf(c_0_96,negated_conjecture,(complement(join(complement(converse(sk1)),composition(complement(one),converse(sk1))))=converse(sk1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_93]), c_0_22]), c_0_28]), c_0_27])).
cnf(c_0_97,negated_conjecture,(composition(complement(converse(sk1)),composition(sk1,converse(sk1)))=composition(complement(one),converse(sk1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_90]), c_0_89])).
cnf(c_0_98,negated_conjecture,(join(complement(converse(sk1)),composition(complement(one),converse(sk1)))=complement(converse(sk1))), inference(spm,[status(thm)],[c_0_28, c_0_96])).
cnf(c_0_99,plain,(join(complement(converse(X1)),composition(X2,complement(converse(composition(X1,X2)))))=complement(converse(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_0]), c_0_1])).
cnf(c_0_100,negated_conjecture,(composition(complement(converse(sk1)),join(one,composition(sk1,converse(sk1))))=complement(converse(sk1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_97]), c_0_98]), c_0_4])).
cnf(c_0_101,plain,(join(X1,composition(join(one,X2),X1))=composition(join(one,X2),X1)), inference(spm,[status(thm)],[c_0_30, c_0_60])).
cnf(c_0_102,negated_conjecture,(join(sk1,composition(sk1,composition(converse(sk1),sk1)))=sk1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_84]), c_0_28]), c_0_84]), c_0_28]), c_0_101]), c_0_84]), c_0_28]), c_0_60]), c_0_13])).
cnf(c_0_103,negated_conjecture,(join(sk1,join(composition(sk1,composition(converse(sk1),sk1)),X1))=join(sk1,X1)), inference(spm,[status(thm)],[c_0_9, c_0_102])).
cnf(c_0_104,plain,(join(X1,complement(composition(X2,complement(composition(converse(X2),X1)))))=complement(composition(X2,complement(composition(converse(X2),X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_59]), c_0_28]), c_0_4])).
cnf(c_0_105,negated_conjecture,(complement(composition(sk1,complement(composition(converse(sk1),sk1))))=join(sk1,complement(composition(sk1,top)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_85]), c_0_104])).
cnf(c_0_106,plain,(converse(join(complement(one),X1))=join(complement(one),converse(X1))), inference(spm,[status(thm)],[c_0_31, c_0_61])).
cnf(c_0_107,plain,(join(complement(one),composition(converse(X1),composition(X2,complement(composition(converse(X2),X1)))))=complement(one)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_15]), c_0_13])).
cnf(c_0_108,negated_conjecture,(composition(sk1,complement(composition(converse(sk1),sk1)))=composition(sk1,complement(one))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_105]), c_0_92])).
cnf(c_0_109,plain,(join(X1,composition(X1,X2))=composition(X1,join(one,X2))), inference(spm,[status(thm)],[c_0_58, c_0_2])).
cnf(c_0_110,plain,(join(complement(one),composition(join(complement(one),converse(X1)),complement(join(complement(one),X1))))=complement(one)), inference(spm,[status(thm)],[c_0_40, c_0_106])).
cnf(c_0_111,negated_conjecture,(join(complement(one),composition(converse(sk1),composition(sk1,complement(one))))=complement(one)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_112,plain,(composition(join(X1,composition(X1,X2)),X3)=composition(X1,composition(join(X2,one),X3))), inference(spm,[status(thm)],[c_0_13, c_0_68])).
cnf(c_0_113,plain,(join(X1,composition(X1,join(one,X2)))=composition(X1,join(one,X2))), inference(spm,[status(thm)],[c_0_30, c_0_109])).
cnf(c_0_114,negated_conjecture,(composition(complement(one),join(one,composition(converse(sk1),sk1)))=complement(one)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_15]), c_0_90]), c_0_13]), c_0_28]), c_0_112]), c_0_4]), c_0_2]), c_0_113])).
cnf(c_0_115,negated_conjecture,(join(one,composition(converse(sk1),sk1))!=one), inference(rw,[status(thm)],[c_0_14, c_0_4])).
cnf(c_0_116,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_114]), c_0_89]), c_0_17]), c_0_28]), c_0_89]), c_0_17]), c_0_28]), c_0_2]), c_0_30]), c_0_89]), c_0_17]), c_0_28]), c_0_115]), ['proof']).
# SZS output end CNFRefutation.
EOF
