% Problem    : GRP754-1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n121.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:41:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C48_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(mult(X1,mult(X2,X3))=mult(mult(rd(X1,X1),X2),mult(X1,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', f05)).
cnf(c_0_1,axiom,(mult(X1,ld(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', f01)).
cnf(c_0_2,axiom,(rd(mult(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', f04)).
cnf(c_0_3,axiom,(mult(mult(X1,X2),X3)=mult(mult(X1,X3),mult(X2,ld(X3,X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', f06)).
cnf(c_0_4,axiom,(ld(X1,mult(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', f02)).
cnf(c_0_5,axiom,(mult(rd(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', f03)).
cnf(c_0_6,negated_conjecture,(mult(mult(a,a),mult(b,c))!=mult(mult(a,b),mult(a,c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
cnf(c_0_7,plain,(mult(mult(rd(X1,X1),X2),X3)=mult(X1,mult(X2,ld(X1,X3)))), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_8,plain,(rd(mult(X1,mult(X2,ld(X1,X3))),X3)=mult(rd(X1,X1),X2)), inference(spm,[status(thm)],[c_0_2, c_0_7])).
cnf(c_0_9,plain,(rd(mult(mult(X1,X2),X3),mult(X2,ld(X3,X3)))=mult(X1,X3)), inference(spm,[status(thm)],[c_0_2, c_0_3])).
cnf(c_0_10,plain,(rd(mult(X1,mult(X2,X3)),mult(X1,X3))=mult(rd(X1,X1),X2)), inference(spm,[status(thm)],[c_0_8, c_0_4])).
cnf(c_0_11,plain,(mult(X1,mult(X2,ld(ld(X3,X1),ld(X3,X1))))=mult(mult(X3,X2),ld(X3,X1))), inference(spm,[status(thm)],[c_0_3, c_0_1])).
cnf(c_0_12,plain,(rd(mult(X1,X2),mult(X3,ld(X2,X2)))=mult(rd(X1,X3),X2)), inference(spm,[status(thm)],[c_0_9, c_0_5])).
cnf(c_0_13,plain,(mult(rd(mult(X1,X2),X3),ld(X1,X3))=mult(rd(X3,X3),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])).
cnf(c_0_14,plain,(mult(rd(X1,X1),rd(X2,X3))=rd(mult(X1,X2),mult(X1,X3))), inference(spm,[status(thm)],[c_0_10, c_0_5])).
cnf(c_0_15,plain,(mult(rd(mult(X1,X2),mult(X1,X2)),mult(X3,X2))=mult(X1,mult(X3,ld(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_10]), c_0_4]), c_0_7])).
cnf(c_0_16,plain,(mult(rd(mult(X1,X2),mult(X1,X3)),X4)=mult(X1,mult(rd(X2,X3),ld(X1,X4)))), inference(spm,[status(thm)],[c_0_7, c_0_14])).
cnf(c_0_17,plain,(ld(mult(X1,X2),mult(mult(X1,X3),X2))=mult(X3,ld(X2,X2))), inference(spm,[status(thm)],[c_0_4, c_0_3])).
cnf(c_0_18,plain,(mult(X1,mult(rd(X2,X2),ld(X1,mult(X3,X2))))=mult(X1,mult(X3,ld(X1,X2)))), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(ld(mult(X1,X2),mult(X3,X2))=mult(ld(X1,X3),ld(X2,X2))), inference(spm,[status(thm)],[c_0_17, c_0_1])).
cnf(c_0_20,plain,(mult(rd(X1,X1),ld(X2,mult(X3,X1)))=mult(X3,ld(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_18]), c_0_4])).
cnf(c_0_21,plain,(mult(ld(rd(X1,X2),X3),ld(X2,X2))=ld(X1,mult(X3,X2))), inference(spm,[status(thm)],[c_0_19, c_0_5])).
cnf(c_0_22,plain,(mult(X1,mult(ld(X2,mult(X3,X1)),ld(X1,X4)))=mult(mult(X3,ld(X2,X1)),X4)), inference(spm,[status(thm)],[c_0_7, c_0_20])).
cnf(c_0_23,plain,(mult(mult(X1,ld(rd(X2,X3),X4)),X3)=mult(mult(X1,X3),ld(X2,mult(X4,X3)))), inference(spm,[status(thm)],[c_0_3, c_0_21])).
cnf(c_0_24,plain,(mult(mult(X1,X2),ld(X3,mult(X2,X2)))=mult(X2,ld(X3,mult(mult(X1,X2),X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_21]), c_0_23])).
cnf(c_0_25,plain,(mult(X1,ld(X2,mult(X3,X3)))=mult(X3,ld(X2,mult(X1,X3)))), inference(spm,[status(thm)],[c_0_24, c_0_5])).
cnf(c_0_26,plain,(ld(rd(X1,X2),X1)=X2), inference(spm,[status(thm)],[c_0_4, c_0_5])).
cnf(c_0_27,plain,(mult(X1,ld(rd(mult(X1,X2),X3),mult(X2,X2)))=mult(X2,X3)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_28,plain,(ld(rd(mult(X1,X2),X3),mult(X2,X2))=ld(X1,mult(X2,X3))), inference(spm,[status(thm)],[c_0_4, c_0_27])).
cnf(c_0_29,plain,(mult(rd(mult(X1,X2),X3),ld(X1,mult(X2,X3)))=mult(X2,X2)), inference(spm,[status(thm)],[c_0_1, c_0_28])).
cnf(c_0_30,plain,(mult(X1,ld(rd(mult(X1,X2),X3),X2))=mult(rd(X2,X2),X3)), inference(spm,[status(thm)],[c_0_20, c_0_26])).
cnf(c_0_31,plain,(mult(rd(mult(rd(mult(X1,X2),X3),X1),X2),X3)=mult(X1,X1)), inference(spm,[status(thm)],[c_0_29, c_0_26])).
cnf(c_0_32,plain,(rd(X1,ld(X2,X1))=X2), inference(spm,[status(thm)],[c_0_2, c_0_1])).
cnf(c_0_33,plain,(ld(rd(mult(X1,X2),X3),X2)=ld(X1,mult(rd(X2,X2),X3))), inference(spm,[status(thm)],[c_0_4, c_0_30])).
cnf(c_0_34,plain,(rd(mult(rd(mult(X1,X2),X3),X1),X2)=rd(mult(X1,X1),X3)), inference(spm,[status(thm)],[c_0_2, c_0_31])).
cnf(c_0_35,plain,(rd(X1,ld(X2,mult(rd(X1,X1),X3)))=rd(mult(X2,X1),X3)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36,plain,(ld(X1,mult(X2,ld(X3,mult(X1,X1))))=ld(X3,mult(X2,X1))), inference(spm,[status(thm)],[c_0_4, c_0_25])).
cnf(c_0_37,plain,(mult(mult(rd(X1,X2),X3),X2)=mult(X1,mult(X3,ld(X2,X2)))), inference(spm,[status(thm)],[c_0_3, c_0_5])).
cnf(c_0_38,plain,(rd(mult(rd(X1,X1),X2),ld(X3,X1))=rd(mult(X3,X2),X1)), inference(spm,[status(thm)],[c_0_2, c_0_13])).
cnf(c_0_39,plain,(mult(rd(mult(X1,X1),X2),X3)=mult(rd(mult(X1,X3),X2),X1)), inference(spm,[status(thm)],[c_0_5, c_0_34])).
cnf(c_0_40,plain,(rd(mult(X1,X2),ld(X3,mult(X1,X1)))=rd(mult(X3,X2),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_35])).
cnf(c_0_41,plain,(rd(mult(X1,mult(X2,ld(X3,X3))),X3)=mult(rd(X1,X3),X2)), inference(spm,[status(thm)],[c_0_2, c_0_37])).
cnf(c_0_42,plain,(mult(rd(X1,X2),rd(X3,X2))=rd(mult(X1,X3),mult(X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_16]), c_0_40]), c_0_41])).
cnf(c_0_43,plain,(rd(mult(X1,mult(X2,X3)),mult(X3,X3))=mult(rd(X1,X3),X2)), inference(spm,[status(thm)],[c_0_42, c_0_2])).
cnf(c_0_44,plain,(mult(mult(rd(X1,X2),X3),mult(X2,X2))=mult(X1,mult(X3,X2))), inference(spm,[status(thm)],[c_0_5, c_0_43])).
cnf(c_0_45,plain,(mult(mult(X1,X1),mult(X2,X3))=mult(mult(X1,X2),mult(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_39]), c_0_44])).
cnf(c_0_46,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_45])]), ['proof']).
# SZS output end CNFRefutation.
EOF
