% Problem    : REL045-2 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n126.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:39:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4d
# and selection function SelectCQIPrecWNTNp.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(converse(composition(X1,X2))=composition(converse(X2),converse(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_multiplicativity_10)).
cnf(c_0_1,axiom,(composition(X1,one)=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_identity_6)).
cnf(c_0_2,axiom,(converse(converse(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_idempotence_8)).
cnf(c_0_3,axiom,(join(composition(converse(X1),complement(composition(X1,X2))),complement(X2))=complement(X2)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_cancellativity_11)).
cnf(c_0_4,axiom,(join(X1,X2)=join(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux1_join_commutativity_1)).
cnf(c_0_5,axiom,(zero=meet(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', def_zero_13)).
cnf(c_0_6,axiom,(meet(X1,X2)=complement(join(complement(X1),complement(X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux4_definiton_of_meet_4)).
cnf(c_0_7,axiom,(top=join(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', def_top_12)).
cnf(c_0_8,axiom,(X1=join(complement(join(complement(X1),complement(X2))),complement(join(complement(X1),X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux3_a_kind_of_de_Morgan_3)).
cnf(c_0_9,axiom,(join(X1,join(X2,X3))=join(join(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', maddux2_join_associativity_2)).
cnf(c_0_10,axiom,(converse(join(X1,X2))=join(converse(X1),converse(X2))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', converse_additivity_9)).
cnf(c_0_11,axiom,(join(meet(composition(X1,X2),X3),composition(meet(X1,composition(X3,converse(X2))),meet(X2,composition(converse(X1),X3))))=composition(meet(X1,composition(X3,converse(X2))),meet(X2,composition(converse(X1),X3)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-1.ax', dedekind_law_14)).
cnf(c_0_12,axiom,(composition(join(X1,X2),X3)=join(composition(X1,X3),composition(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_distributivity_7)).
cnf(c_0_13,negated_conjecture,(join(sk1,composition(composition(sk1,converse(sk1)),sk1))!=composition(composition(sk1,converse(sk1)),sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_17)).
cnf(c_0_14,axiom,(composition(X1,composition(X2,X3))=composition(composition(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001-0.ax', composition_associativity_5)).
cnf(c_0_15,plain,(composition(converse(one),converse(X1))=converse(X1)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_16,plain,(composition(converse(one),X1)=X1), inference(spm,[status(thm)],[c_0_15, c_0_2])).
cnf(c_0_17,plain,(converse(one)=one), inference(spm,[status(thm)],[c_0_1, c_0_16])).
cnf(c_0_18,plain,(join(complement(X1),composition(converse(X2),complement(composition(X2,X1))))=complement(X1)), inference(rw,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_19,plain,(composition(one,X1)=X1), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,plain,(complement(join(complement(X1),complement(complement(X1))))=zero), inference(rw,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_21,plain,(join(complement(X1),complement(X1))=complement(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_17]), c_0_19])).
cnf(c_0_22,plain,(complement(top)=zero), inference(rw,[status(thm)],[c_0_20, c_0_7])).
cnf(c_0_23,plain,(join(zero,zero)=zero), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_24,plain,(join(complement(join(complement(X1),X2)),complement(join(complement(X1),complement(X2))))=X1), inference(rw,[status(thm)],[c_0_8, c_0_4])).
cnf(c_0_25,plain,(join(zero,join(zero,X1))=join(zero,X1)), inference(spm,[status(thm)],[c_0_9, c_0_23])).
cnf(c_0_26,plain,(join(zero,complement(complement(X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_21]), c_0_7]), c_0_22]), c_0_4])).
cnf(c_0_27,plain,(join(zero,X1)=X1), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_28,plain,(complement(complement(X1))=X1), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_29,plain,(join(X1,X1)=X1), inference(spm,[status(thm)],[c_0_21, c_0_28])).
cnf(c_0_30,plain,(join(X1,join(X1,X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_9, c_0_29])).
cnf(c_0_31,plain,(join(converse(X1),converse(complement(X1)))=converse(top)), inference(spm,[status(thm)],[c_0_10, c_0_7])).
cnf(c_0_32,plain,(join(X1,top)=top), inference(spm,[status(thm)],[c_0_30, c_0_7])).
cnf(c_0_33,plain,(join(converse(X1),converse(top))=converse(top)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34,plain,(join(top,X1)=top), inference(spm,[status(thm)],[c_0_4, c_0_32])).
cnf(c_0_35,plain,(join(X1,converse(top))=converse(top)), inference(spm,[status(thm)],[c_0_33, c_0_2])).
cnf(c_0_36,plain,(join(one,converse(complement(one)))=converse(top)), inference(spm,[status(thm)],[c_0_31, c_0_17])).
cnf(c_0_37,plain,(converse(top)=top), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38,plain,(join(one,converse(complement(one)))=top), inference(rw,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_39,plain,(join(X1,complement(join(complement(X1),X2)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_24]), c_0_4])).
cnf(c_0_40,plain,(join(complement(join(complement(composition(X1,X2)),complement(X3))),composition(complement(join(complement(X1),complement(composition(X3,converse(X2))))),complement(join(complement(X2),complement(composition(converse(X1),X3))))))=composition(complement(join(complement(X1),complement(composition(X3,converse(X2))))),complement(join(complement(X2),complement(composition(converse(X1),X3)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_6]), c_0_6]), c_0_6]), c_0_6]), c_0_6])).
cnf(c_0_41,plain,(join(X1,composition(converse(complement(one)),X1))=composition(top,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_38]), c_0_19])).
cnf(c_0_42,plain,(join(complement(X1),complement(join(X1,X2)))=complement(X1)), inference(spm,[status(thm)],[c_0_39, c_0_28])).
cnf(c_0_43,plain,(join(complement(join(complement(composition(converse(X1),X2)),complement(X3))),composition(complement(join(complement(converse(X1)),complement(composition(X3,converse(X2))))),complement(join(complement(X2),complement(composition(X1,X3))))))=composition(complement(join(complement(converse(X1)),complement(composition(X3,converse(X2))))),complement(join(complement(X2),complement(composition(X1,X3)))))), inference(spm,[status(thm)],[c_0_40, c_0_2])).
cnf(c_0_44,plain,(join(X1,complement(composition(top,complement(X1))))=X1), inference(spm,[status(thm)],[c_0_39, c_0_41])).
cnf(c_0_45,plain,(join(complement(X1),complement(composition(top,X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_42, c_0_41])).
cnf(c_0_46,negated_conjecture,(join(sk1,composition(sk1,composition(converse(sk1),sk1)))!=composition(sk1,composition(converse(sk1),sk1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_14])).
cnf(c_0_47,plain,(join(X1,composition(X1,composition(converse(X1),X1)))=composition(X1,composition(converse(X1),X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_37]), c_0_28]), c_0_4]), c_0_45]), c_0_28]), c_0_37]), c_0_22]), c_0_28]), c_0_27]), c_0_28]), c_0_28]), c_0_14]), c_0_37]), c_0_22]), c_0_28]), c_0_27]), c_0_28]), c_0_28]), c_0_14])).
cnf(c_0_48,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47])]), ['proof']).
# SZS output end CNFRefutation.
EOF
