% Problem    : GRP471-1 : TPTP v6.1.0. Released v2.6.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n132.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:38:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___209_C18_F1_AE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(divide(inverse(divide(X1,divide(X2,divide(X3,X4)))),divide(divide(X4,X3),X1))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', single_axiom)).
cnf(c_0_1,negated_conjecture,(multiply(multiply(a3,b3),c3)!=multiply(a3,multiply(b3,c3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_these_axioms_3)).
cnf(c_0_2,axiom,(multiply(X1,X2)=divide(X1,inverse(X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', multiply)).
cnf(c_0_3,plain,(divide(inverse(divide(X1,X2)),divide(divide(X3,divide(X4,X5)),X1))=inverse(divide(X3,divide(X2,divide(X5,X4))))), inference(spm,[status(thm)],[c_0_0, c_0_0])).
cnf(c_0_4,plain,(inverse(divide(X1,divide(divide(X2,divide(divide(X3,X4),X1)),divide(X4,X3))))=X2), inference(spm,[status(thm)],[c_0_0, c_0_3])).
cnf(c_0_5,plain,(inverse(inverse(divide(X1,divide(X2,divide(inverse(divide(divide(X3,X4),X2)),divide(X4,X3))))))=X1), inference(spm,[status(thm)],[c_0_4, c_0_3])).
cnf(c_0_6,plain,(inverse(divide(divide(inverse(divide(divide(X1,X2),divide(X3,X4))),divide(X2,X1)),divide(X5,divide(X4,X3))))=inverse(inverse(X5))), inference(spm,[status(thm)],[c_0_5, c_0_0])).
cnf(c_0_7,plain,(divide(inverse(inverse(X1)),divide(divide(X2,X3),divide(inverse(divide(divide(X4,X5),divide(X2,X3))),divide(X5,X4))))=X1), inference(spm,[status(thm)],[c_0_0, c_0_6])).
cnf(c_0_8,plain,(divide(inverse(X1),divide(divide(X2,X3),inverse(divide(divide(X3,X2),divide(X1,divide(X4,X5))))))=divide(X5,X4)), inference(spm,[status(thm)],[c_0_0, c_0_0])).
cnf(c_0_9,plain,(divide(inverse(inverse(X1)),divide(X2,divide(inverse(divide(divide(X3,X4),X2)),divide(X4,X3))))=X1), inference(spm,[status(thm)],[c_0_7, c_0_7])).
cnf(c_0_10,plain,(divide(divide(divide(X1,X2),inverse(divide(divide(X2,X1),divide(X3,divide(X4,X5))))),inverse(X3))=divide(X4,X5)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_8]), c_0_8])).
cnf(c_0_11,plain,(divide(inverse(X1),divide(divide(divide(X2,X3),inverse(divide(divide(X3,X2),X4))),inverse(inverse(X1))))=X4), inference(spm,[status(thm)],[c_0_0, c_0_9])).
cnf(c_0_12,plain,(divide(divide(divide(X1,X2),inverse(divide(divide(X2,X1),divide(X3,X4)))),inverse(X3))=X4), inference(spm,[status(thm)],[c_0_10, c_0_0])).
cnf(c_0_13,plain,(divide(X1,divide(divide(divide(X2,X3),inverse(divide(divide(X3,X2),X4))),inverse(X1)))=X4), inference(spm,[status(thm)],[c_0_11, c_0_5])).
cnf(c_0_14,plain,(divide(divide(divide(X1,X2),X3),inverse(divide(X3,divide(divide(X4,X5),divide(X2,X1)))))=divide(X5,X4)), inference(spm,[status(thm)],[c_0_12, c_0_4])).
cnf(c_0_15,plain,(divide(divide(inverse(divide(divide(X1,X2),X3)),divide(divide(X4,X5),divide(X1,X2))),divide(X5,X4))=X3), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_16,plain,(divide(divide(inverse(divide(X1,X2)),divide(divide(X3,X4),X1)),divide(X4,X3))=X2), inference(spm,[status(thm)],[c_0_15, c_0_15])).
cnf(c_0_17,plain,(divide(divide(divide(divide(divide(X1,X2),X3),inverse(divide(X3,X4))),inverse(X4)),inverse(X2))=X1), inference(spm,[status(thm)],[c_0_12, c_0_16])).
cnf(c_0_18,plain,(divide(divide(X1,inverse(X2)),inverse(inverse(divide(X3,X4))))=divide(divide(X1,divide(inverse(X4),X2)),X3)), inference(spm,[status(thm)],[c_0_17, c_0_17])).
cnf(c_0_19,plain,(inverse(divide(divide(X1,X2),divide(X3,divide(divide(divide(X2,X1),X4),inverse(X3)))))=X4), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_18]), c_0_3])).
cnf(c_0_20,plain,(divide(X1,divide(divide(inverse(X2),divide(divide(X3,X4),X1)),divide(X4,X3)))=X2), inference(spm,[status(thm)],[c_0_0, c_0_19])).
cnf(c_0_21,plain,(divide(divide(X1,X2),divide(X3,divide(divide(divide(X2,X1),inverse(X4)),inverse(X3))))=X4), inference(spm,[status(thm)],[c_0_20, c_0_20])).
cnf(c_0_22,plain,(divide(divide(X1,X2),inverse(divide(divide(X2,X1),divide(X3,divide(X4,inverse(X3))))))=X4), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_18]), c_0_3])).
cnf(c_0_23,plain,(divide(X1,divide(X2,inverse(X1)))=divide(X3,divide(X2,inverse(X3)))), inference(spm,[status(thm)],[c_0_13, c_0_22])).
cnf(c_0_24,plain,(divide(divide(X1,inverse(X2)),X2)=divide(divide(X1,inverse(X3)),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_23]), c_0_8])).
cnf(c_0_25,plain,(divide(divide(X1,X2),divide(X3,divide(divide(divide(X2,X1),inverse(X4)),X4)))=inverse(X3)), inference(spm,[status(thm)],[c_0_21, c_0_24])).
cnf(c_0_26,plain,(divide(divide(inverse(divide(X1,inverse(X2))),divide(divide(X3,X2),X1)),divide(X4,X3))=inverse(X4)), inference(spm,[status(thm)],[c_0_25, c_0_17])).
cnf(c_0_27,plain,(divide(inverse(divide(X1,divide(inverse(divide(X2,X3)),divide(X3,X2)))),divide(X4,X1))=inverse(X4)), inference(spm,[status(thm)],[c_0_26, c_0_3])).
cnf(c_0_28,plain,(divide(inverse(divide(X1,X2)),divide(X2,X1))=divide(inverse(divide(X3,X4)),divide(X4,X3))), inference(spm,[status(thm)],[c_0_16, c_0_27])).
cnf(c_0_29,plain,(divide(inverse(inverse(divide(X1,X2))),divide(X3,inverse(X4)))=divide(X1,divide(X3,divide(inverse(X2),X4)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_18]), c_0_8])).
cnf(c_0_30,plain,(divide(divide(inverse(divide(X1,X2)),divide(X2,X1)),divide(X3,X4))=divide(X4,X3)), inference(spm,[status(thm)],[c_0_16, c_0_28])).
cnf(c_0_31,plain,(divide(divide(inverse(divide(X1,X2)),X3),X4)=divide(X2,divide(X4,divide(inverse(X1),X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_29]), c_0_30])).
cnf(c_0_32,plain,(divide(X1,divide(divide(X2,X3),divide(inverse(X4),divide(X1,X4))))=divide(X3,X2)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_33,plain,(divide(divide(X1,divide(inverse(X2),X3)),X4)=divide(X1,divide(inverse(divide(X2,X4)),X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_31]), c_0_8])).
cnf(c_0_34,plain,(inverse(divide(inverse(X1),divide(X2,divide(inverse(divide(divide(X3,X4),X2)),divide(X4,X3)))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_13]), c_0_3])).
cnf(c_0_35,plain,(divide(inverse(divide(divide(X1,X2),X3)),divide(X2,X1))=X3), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_32]), c_0_33]), c_0_0])).
cnf(c_0_36,plain,(inverse(divide(inverse(X1),divide(X2,X2)))=X1), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_37,plain,(divide(inverse(X1),divide(X2,X1))=inverse(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_12]), c_0_35])).
cnf(c_0_38,plain,(divide(divide(X1,X1),X2)=inverse(X2)), inference(spm,[status(thm)],[c_0_36, c_0_35])).
cnf(c_0_39,plain,(inverse(inverse(X1))=X1), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_40,plain,(divide(X1,divide(X2,X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_38]), c_0_39])).
cnf(c_0_41,plain,(inverse(divide(inverse(X1),divide(divide(X2,X3),divide(inverse(divide(divide(X4,X5),divide(X1,X3))),divide(X5,X4)))))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_12]), c_0_3])).
cnf(c_0_42,plain,(inverse(divide(X1,X2))=divide(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_40]), c_0_37])).
cnf(c_0_43,plain,(divide(divide(divide(X1,X2),divide(X3,X2)),inverse(X3))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_35]), c_0_42])).
cnf(c_0_44,negated_conjecture,(divide(divide(a3,inverse(b3)),inverse(c3))!=divide(a3,inverse(divide(b3,inverse(c3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_1, c_0_2]), c_0_2]), c_0_2]), c_0_2])).
cnf(c_0_45,plain,(divide(divide(X1,X2),divide(X3,X2))=divide(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_43]), c_0_39]), c_0_42])).
cnf(c_0_46,plain,(divide(divide(X1,X2),inverse(X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_37]), c_0_39])).
cnf(c_0_47,negated_conjecture,(divide(divide(a3,inverse(b3)),inverse(c3))!=divide(a3,divide(inverse(c3),b3))), inference(rw,[status(thm)],[c_0_44, c_0_42])).
cnf(c_0_48,plain,(divide(divide(X1,inverse(X2)),X3)=divide(X1,divide(X3,X2))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_49,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48])]), ['proof']).
# SZS output end CNFRefutation.
EOF
