% Problem    : GRP769-1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n091.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:27:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C48_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(product(quotient(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos06)).
cnf(c_0_1,axiom,(difference(X1,product(product(X1,X2),X3))=quotient(product(X2,product(X3,X1)),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos07)).
cnf(c_0_2,axiom,(product(X1,difference(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos03)).
cnf(c_0_3,axiom,(difference(X1,product(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos04)).
cnf(c_0_4,axiom,(product(X1,one)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos01)).
cnf(c_0_5,axiom,(product(j(product(X1,X2)),X1)=j(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos24)).
cnf(c_0_6,axiom,(j(X1)=quotient(one,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos10)).
cnf(c_0_7,axiom,(product(one,X1)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos02)).
cnf(c_0_8,axiom,(difference(product(X1,X2),product(X1,product(X2,X3)))=quotient(quotient(product(X3,product(X1,X2)),X2),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos08)).
cnf(c_0_9,axiom,(quotient(product(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos05)).
cnf(c_0_10,axiom,(product(X1,i(product(X2,X1)))=i(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos23)).
cnf(c_0_11,axiom,(i(X1)=difference(X1,one)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos09)).
cnf(c_0_12,axiom,(product(i(X1),X1)=product(X1,j(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos11)).
cnf(c_0_13,axiom,(eta(X1)=product(i(X1),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos12)).
cnf(c_0_14,axiom,(product(X1,product(eta(X1),X2))=product(j(j(X1)),X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos14)).
cnf(c_0_15,axiom,(product(eta(X1),product(X2,X3))=product(product(eta(X1),X2),X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos16)).
cnf(c_0_16,negated_conjecture,(product(product(x0,x1),x2)!=product(product(x0,x2),difference(x2,product(x1,x2)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
cnf(c_0_17,plain,(product(difference(X1,product(product(X1,X2),X3)),X1)=product(X2,product(X3,X1))), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_18,plain,(product(difference(X1,product(X2,X3)),X1)=product(difference(X1,X2),product(X3,X1))), inference(spm,[status(thm)],[c_0_17, c_0_2])).
cnf(c_0_19,plain,(product(difference(X1,X2),product(difference(X2,X3),X1))=product(difference(X1,X3),X1)), inference(spm,[status(thm)],[c_0_18, c_0_2])).
cnf(c_0_20,plain,(product(X1,product(difference(product(X2,X1),X3),X2))=product(difference(X2,X3),X2)), inference(spm,[status(thm)],[c_0_19, c_0_3])).
cnf(c_0_21,plain,(difference(X1,X1)=one), inference(spm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_22,plain,(difference(X1,product(difference(X2,X3),X2))=product(difference(product(X2,X1),X3),X2)), inference(spm,[status(thm)],[c_0_3, c_0_20])).
cnf(c_0_23,plain,(difference(difference(X1,X2),product(difference(X1,X3),X1))=product(difference(X2,X3),X1)), inference(spm,[status(thm)],[c_0_3, c_0_19])).
cnf(c_0_24,plain,(product(quotient(one,product(X1,X2)),X1)=quotient(one,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_6]), c_0_6])).
cnf(c_0_25,plain,(product(difference(product(X1,product(difference(X1,X2),X1)),X2),X1)=one), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26,plain,(quotient(X1,one)=X1), inference(spm,[status(thm)],[c_0_4, c_0_0])).
cnf(c_0_27,plain,(difference(difference(X1,X2),X1)=product(difference(X2,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_21]), c_0_7])).
cnf(c_0_28,plain,(difference(product(X1,product(difference(X1,X2),X1)),X2)=quotient(one,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_7])).
cnf(c_0_29,plain,(product(difference(product(X1,X2),X1),X1)=difference(X2,X1)), inference(spm,[status(thm)],[c_0_27, c_0_3])).
cnf(c_0_30,plain,(quotient(difference(X1,product(product(X1,X2),X3)),X3)=difference(product(X3,X1),product(X3,product(X1,X2)))), inference(rw,[status(thm)],[c_0_8, c_0_1])).
cnf(c_0_31,plain,(difference(product(X1,product(X2,X1)),product(X1,X2))=quotient(one,X1)), inference(spm,[status(thm)],[c_0_28, c_0_3])).
cnf(c_0_32,plain,(quotient(difference(X1,X2),X2)=difference(product(X2,X1),X2)), inference(spm,[status(thm)],[c_0_9, c_0_29])).
cnf(c_0_33,plain,(difference(quotient(X1,X2),X1)=X2), inference(spm,[status(thm)],[c_0_3, c_0_0])).
cnf(c_0_34,plain,(quotient(difference(X1,product(X2,X3)),X3)=difference(product(X3,X1),product(X3,X2))), inference(spm,[status(thm)],[c_0_30, c_0_2])).
cnf(c_0_35,plain,(difference(product(X1,X2),product(X1,quotient(X2,X1)))=quotient(one,X1)), inference(spm,[status(thm)],[c_0_31, c_0_0])).
cnf(c_0_36,plain,(quotient(quotient(one,X1),X2)=quotient(one,product(X2,X1))), inference(spm,[status(thm)],[c_0_9, c_0_24])).
cnf(c_0_37,plain,(quotient(X1,difference(X2,X1))=X2), inference(spm,[status(thm)],[c_0_9, c_0_2])).
cnf(c_0_38,plain,(difference(product(X1,quotient(X1,X2)),X1)=quotient(X2,X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39,plain,(product(X1,difference(product(X2,X1),one))=difference(X2,one)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11]), c_0_11])).
cnf(c_0_40,plain,(difference(product(quotient(X1,X2),product(X2,X1)),X1)=quotient(one,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_0]), c_0_0])).
cnf(c_0_41,plain,(quotient(X1,quotient(X2,X1))=product(X1,quotient(X1,X2))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_42,plain,(difference(product(X1,X2),one)=difference(X2,difference(X1,one))), inference(spm,[status(thm)],[c_0_3, c_0_39])).
cnf(c_0_43,plain,(product(quotient(X1,X2),product(X2,X1))=product(X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_40]), c_0_41]), c_0_26])).
cnf(c_0_44,plain,(difference(difference(X1,X2),difference(X1,one))=difference(X2,one)), inference(spm,[status(thm)],[c_0_42, c_0_2])).
cnf(c_0_45,plain,(product(difference(X1,one),X1)=product(X1,quotient(one,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_11]), c_0_6])).
cnf(c_0_46,plain,(product(difference(X1,one),product(X2,X1))=product(difference(X1,X2),X1)), inference(spm,[status(thm)],[c_0_18, c_0_7])).
cnf(c_0_47,plain,(product(product(X1,quotient(X1,X2)),X2)=product(X1,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_0]), c_0_41])).
cnf(c_0_48,plain,(product(difference(X1,X2),difference(X2,one))=difference(X1,one)), inference(spm,[status(thm)],[c_0_2, c_0_44])).
cnf(c_0_49,plain,(difference(X1,product(difference(X1,one),X1))=quotient(one,X1)), inference(spm,[status(thm)],[c_0_3, c_0_45])).
cnf(c_0_50,plain,(product(difference(X1,one),product(X2,X2))=product(difference(X1,X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_18]), c_0_0])).
cnf(c_0_51,plain,(difference(difference(X1,X2),product(X3,X1))=product(difference(X2,product(X1,X3)),X1)), inference(spm,[status(thm)],[c_0_23, c_0_3])).
cnf(c_0_52,plain,(product(difference(X1,quotient(one,X2)),X2)=difference(X1,one)), inference(spm,[status(thm)],[c_0_48, c_0_33])).
cnf(c_0_53,plain,(difference(difference(quotient(one,X1),X2),X1)=difference(X2,one)), inference(spm,[status(thm)],[c_0_44, c_0_33])).
cnf(c_0_54,plain,(product(difference(X1,difference(X1,one)),product(X1,X1))=one), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_49]), c_0_0])).
cnf(c_0_55,plain,(difference(difference(X1,one),product(difference(X1,X2),X2))=product(X2,X2)), inference(spm,[status(thm)],[c_0_3, c_0_50])).
cnf(c_0_56,plain,(difference(quotient(X1,X2),one)=product(X2,difference(X1,one))), inference(spm,[status(thm)],[c_0_48, c_0_33])).
cnf(c_0_57,plain,(product(difference(X1,product(X2,product(difference(X2,X1),X3))),X2)=quotient(product(X3,X1),difference(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1, c_0_2]), c_0_51])).
cnf(c_0_58,plain,(product(difference(X1,one),difference(difference(X2,one),one))=difference(difference(X2,X1),one)), inference(spm,[status(thm)],[c_0_48, c_0_44])).
cnf(c_0_59,plain,(difference(one,X1)=X1), inference(spm,[status(thm)],[c_0_7, c_0_2])).
cnf(c_0_60,plain,(quotient(difference(X1,one),X2)=difference(X1,quotient(one,X2))), inference(spm,[status(thm)],[c_0_9, c_0_52])).
cnf(c_0_61,plain,(product(X1,product(difference(X1,X2),X2))=product(X2,X2)), inference(spm,[status(thm)],[c_0_43, c_0_37])).
cnf(c_0_62,plain,(quotient(X1,difference(X2,one))=difference(quotient(one,X1),X2)), inference(spm,[status(thm)],[c_0_37, c_0_53])).
cnf(c_0_63,plain,(quotient(one,product(X1,X1))=difference(X1,difference(X1,one))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_54]), c_0_26]), c_0_7])).
cnf(c_0_64,plain,(product(product(X1,X2),difference(X2,difference(X1,one)))=one), inference(spm,[status(thm)],[c_0_2, c_0_42])).
cnf(c_0_65,plain,(difference(product(X1,difference(X2,one)),product(X1,X2))=product(X2,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_33]), c_0_56])).
cnf(c_0_66,plain,(product(product(X1,difference(difference(X2,X1),one)),X1)=difference(difference(X2,one),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_4]), c_0_60]), c_0_37])).
cnf(c_0_67,plain,(difference(difference(X1,difference(X1,one)),X2)=product(product(X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_50]), c_0_61]), c_0_59]), c_0_4]), c_0_62]), c_0_63])).
cnf(c_0_68,plain,(difference(difference(X1,difference(X1,one)),one)=product(X1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_54])).
cnf(c_0_69,plain,(difference(difference(product(X1,X1),X2),one)=product(difference(X2,one),product(X1,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_54]), c_0_64])).
cnf(c_0_70,plain,(product(product(X1,difference(X2,one)),product(X2,X2))=product(X1,X2)), inference(spm,[status(thm)],[c_0_2, c_0_65])).
cnf(c_0_71,plain,(product(difference(X1,difference(X1,one)),X2)=difference(product(X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_42]), c_0_42]), c_0_2]), c_0_68])).
cnf(c_0_72,plain,(difference(difference(product(X1,X1),X2),one)=product(difference(X2,X1),X1)), inference(rw,[status(thm)],[c_0_69, c_0_50])).
cnf(c_0_73,plain,(difference(product(X1,X1),product(X2,X1))=product(difference(product(X1,X1),X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_70]), c_0_42]), c_0_71]), c_0_18]), c_0_46]), c_0_21]), c_0_7])).
cnf(c_0_74,plain,(product(difference(product(X1,X1),X2),product(difference(X2,X1),X1))=one), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_50]), c_0_42]), c_0_54])).
cnf(c_0_75,plain,(product(difference(X1,one),product(product(X1,X1),X2))=product(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_65]), c_0_3])).
cnf(c_0_76,plain,(difference(X1,product(product(X1,difference(X2,one)),X1))=difference(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1, c_0_50]), c_0_9])).
cnf(c_0_77,plain,(difference(X1,product(difference(X2,X1),X1))=product(difference(product(X2,X1),X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_42]), c_0_69]), c_0_50])).
cnf(c_0_78,plain,(product(difference(product(X1,X1),product(X1,X2)),difference(X2,X1))=one), inference(spm,[status(thm)],[c_0_74, c_0_29])).
cnf(c_0_79,plain,(quotient(one,difference(X1,X2))=product(quotient(one,X2),X1)), inference(spm,[status(thm)],[c_0_24, c_0_2])).
cnf(c_0_80,plain,(product(difference(X1,product(X2,quotient(X3,X2))),X2)=difference(difference(X2,X1),X3)), inference(spm,[status(thm)],[c_0_51, c_0_0])).
cnf(c_0_81,plain,(product(difference(X1,one),X1)=eta(X1)), inference(rw,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_82,plain,(product(X1,difference(product(X1,X1),X2))=product(difference(X1,one),X2)), inference(spm,[status(thm)],[c_0_75, c_0_2])).
cnf(c_0_83,plain,(product(product(X1,difference(X2,one)),X1)=product(X1,difference(X2,X1))), inference(spm,[status(thm)],[c_0_2, c_0_76])).
cnf(c_0_84,plain,(difference(product(X1,X1),product(X1,difference(X2,X1)))=difference(product(X2,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_77]), c_0_9])).
cnf(c_0_85,plain,(difference(product(X1,X1),product(X1,X2))=product(quotient(one,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_78]), c_0_26]), c_0_7]), c_0_79])).
cnf(c_0_86,plain,(product(product(difference(X1,difference(X2,one)),X2),X2)=difference(difference(X2,X1),one)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_45]), c_0_22]), c_0_42])).
cnf(c_0_87,plain,(product(X1,product(product(difference(X1,one),X1),X2))=product(quotient(one,quotient(one,X1)),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_6]), c_0_6]), c_0_81])).
cnf(c_0_88,plain,(difference(X1,product(difference(X1,one),X2))=difference(product(X1,X1),X2)), inference(spm,[status(thm)],[c_0_3, c_0_82])).
cnf(c_0_89,plain,(product(difference(X1,one),product(X1,difference(X2,X1)))=product(difference(X2,one),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_83]), c_0_3])).
cnf(c_0_90,plain,(product(quotient(one,X1),difference(X2,X1))=difference(product(X2,X1),X1)), inference(rw,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_91,plain,(difference(difference(X1,difference(X1,X2)),one)=product(product(difference(X2,one),X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_53]), c_0_37])).
cnf(c_0_92,plain,(difference(X1,product(quotient(one,quotient(one,X1)),X2))=product(product(difference(X1,one),X1),X2)), inference(spm,[status(thm)],[c_0_3, c_0_87])).
cnf(c_0_93,plain,(difference(X1,product(difference(X2,one),X1))=difference(product(X2,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_85]), c_0_90])).
cnf(c_0_94,plain,(difference(product(X1,X2),difference(difference(X2,X3),one))=difference(X2,difference(difference(X2,product(X3,X1)),one))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_18]), c_0_42])).
cnf(c_0_95,plain,(difference(difference(X1,product(product(difference(X1,one),X1),X2)),one)=product(difference(X2,one),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_42]), c_0_33]), c_0_52])).
cnf(c_0_96,plain,(product(product(product(difference(X1,one),X1),X2),X3)=product(product(difference(X1,one),X1),product(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_81]), c_0_81])).
cnf(c_0_97,plain,(difference(X1,product(difference(X2,difference(X3,one)),X1))=difference(product(product(X3,X2),X1),X1)), inference(spm,[status(thm)],[c_0_93, c_0_42])).
cnf(c_0_98,plain,(product(difference(quotient(one,X1),X2),X2)=product(X1,product(X2,X2))), inference(spm,[status(thm)],[c_0_50, c_0_33])).
cnf(c_0_99,plain,(difference(product(X1,X2),product(difference(X3,one),X2))=difference(product(product(X3,X1),X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96]), c_0_95]), c_0_42]), c_0_97])).
cnf(c_0_100,plain,(difference(X1,product(product(X1,X2),X1))=difference(quotient(one,X2),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_98]), c_0_1])).
cnf(c_0_101,plain,(product(product(X1,X2),difference(product(product(X3,X1),X2),X2))=product(difference(X3,one),X2)), inference(spm,[status(thm)],[c_0_2, c_0_99])).
cnf(c_0_102,plain,(difference(product(quotient(one,X1),X2),X2)=difference(X2,product(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_2]), c_0_79])).
cnf(c_0_103,plain,(product(product(X1,X2),difference(X2,product(X3,X2)))=product(product(X1,X3),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_24]), c_0_102]), c_0_56]), c_0_59]), c_0_4])).
cnf(c_0_104,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_103])]), ['proof']).
# SZS output end CNFRefutation.
EOF
