% Problem    : RNG016-6 : TPTP v6.1.0. Released v1.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n179.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:30:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(add(X1,add(X2,X3))=add(add(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', associativity_for_addition)).
cnf(c_0_1,axiom,(add(X1,additive_inverse(X1))=additive_identity), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', right_additive_inverse)).
cnf(c_0_2,axiom,(add(additive_identity,X1)=X1), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', left_additive_identity)).
cnf(c_0_3,axiom,(additive_inverse(additive_inverse(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', additive_inverse_additive_inverse)).
cnf(c_0_4,axiom,(add(X1,X2)=add(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', commutativity_for_addition)).
cnf(c_0_5,axiom,(multiply(X1,add(X2,X3))=add(multiply(X1,X2),multiply(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', distribute1)).
cnf(c_0_6,axiom,(multiply(X1,additive_identity)=additive_identity), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', right_multiplicative_zero)).
cnf(c_0_7,axiom,(add(X1,additive_identity)=X1), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', right_additive_identity)).
cnf(c_0_8,axiom,(multiply(add(X1,X2),X3)=add(multiply(X1,X3),multiply(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', distribute2)).
cnf(c_0_9,axiom,(multiply(additive_identity,X1)=additive_identity), file('/export/starexec/sandbox/benchmark/Axioms/RNG003-0.ax', left_multiplicative_zero)).
cnf(c_0_10,negated_conjecture,(multiply(add(x,additive_inverse(y)),z)!=add(multiply(x,z),additive_inverse(multiply(y,z)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prove_distributivity)).
cnf(c_0_11,plain,(add(X1,add(additive_inverse(X1),X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_12,plain,(add(additive_inverse(X1),add(X1,X2))=X2), inference(spm,[status(thm)],[c_0_11, c_0_3])).
cnf(c_0_13,plain,(add(additive_inverse(X1),add(X2,X1))=X2), inference(spm,[status(thm)],[c_0_12, c_0_4])).
cnf(c_0_14,plain,(add(X1,additive_inverse(add(X1,X2)))=additive_inverse(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_13]), c_0_4])).
cnf(c_0_15,plain,(add(multiply(X1,X2),additive_inverse(multiply(X1,add(X2,X3))))=additive_inverse(multiply(X1,X3))), inference(spm,[status(thm)],[c_0_14, c_0_5])).
cnf(c_0_16,plain,(additive_inverse(additive_identity)=additive_identity), inference(spm,[status(thm)],[c_0_2, c_0_1])).
cnf(c_0_17,plain,(additive_inverse(multiply(X1,additive_inverse(X2)))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_1]), c_0_6]), c_0_16]), c_0_7])).
cnf(c_0_18,plain,(additive_inverse(multiply(X1,X2))=multiply(X1,additive_inverse(X2))), inference(spm,[status(thm)],[c_0_3, c_0_17])).
cnf(c_0_19,plain,(add(multiply(X1,X2),multiply(add(X1,X3),additive_inverse(X2)))=multiply(X3,additive_inverse(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_8]), c_0_18]), c_0_18])).
cnf(c_0_20,plain,(multiply(additive_inverse(X1),additive_inverse(X2))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_1]), c_0_9]), c_0_7])).
cnf(c_0_21,plain,(multiply(additive_inverse(X1),X2)=multiply(X1,additive_inverse(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_20]), c_0_18]), c_0_3])).
cnf(c_0_22,negated_conjecture,(add(multiply(x,z),multiply(y,additive_inverse(z)))!=multiply(add(x,additive_inverse(y)),z)), inference(rw,[status(thm)],[c_0_10, c_0_18])).
cnf(c_0_23,plain,(add(multiply(X1,X2),multiply(X3,additive_inverse(X2)))=multiply(add(X1,additive_inverse(X3)),X2)), inference(spm,[status(thm)],[c_0_8, c_0_21])).
cnf(c_0_24,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23])]), ['proof']).
# SZS output end CNFRefutation.
EOF
