% Problem    : GRP694-1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n091.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:33:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___209_C18_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(ld(X1,mult(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c02)).
cnf(c_0_1,axiom,(mult(X1,i(X1))=unit), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c10)).
cnf(c_0_2,axiom,(mult(mult(X1,X2),X3)=mult(mult(X1,X3),ld(X3,mult(X2,X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c08)).
cnf(c_0_3,axiom,(mult(i(X1),X1)=unit), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c09)).
cnf(c_0_4,axiom,(rd(mult(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c04)).
cnf(c_0_5,axiom,(mult(X1,ld(X1,X2))=X2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c01)).
cnf(c_0_6,axiom,(mult(unit,X1)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c06)).
cnf(c_0_7,axiom,(mult(X1,unit)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c05)).
cnf(c_0_8,axiom,(mult(rd(X1,X2),X2)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c03)).
cnf(c_0_9,axiom,(mult(X1,mult(X2,X3))=mult(rd(mult(X1,X2),X1),mult(X1,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c07)).
cnf(c_0_10,axiom,(mult(op_c,mult(op_c,op_c))=op_d), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c11)).
cnf(c_0_11,axiom,(mult(op_d,op_d)=op_e), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c12)).
cnf(c_0_12,axiom,(mult(op_e,op_e)=op_f), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c13)).
cnf(c_0_13,negated_conjecture,(mult(a,mult(op_f,b))!=mult(mult(a,op_f),b)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
cnf(c_0_14,plain,(i(X1)=ld(X1,unit)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_15,plain,(ld(mult(X1,X2),mult(mult(X1,X3),X2))=ld(X2,mult(X3,X2))), inference(spm,[status(thm)],[c_0_0, c_0_2])).
cnf(c_0_16,plain,(mult(ld(X1,unit),X1)=unit), inference(rw,[status(thm)],[c_0_3, c_0_14])).
cnf(c_0_17,plain,(rd(X1,ld(X2,X1))=X2), inference(spm,[status(thm)],[c_0_4, c_0_5])).
cnf(c_0_18,plain,(ld(mult(ld(X1,unit),X2),X2)=ld(X2,mult(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_6])).
cnf(c_0_19,plain,(ld(unit,X1)=X1), inference(spm,[status(thm)],[c_0_0, c_0_6])).
cnf(c_0_20,plain,(mult(mult(ld(X1,unit),X2),X1)=ld(X1,mult(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_16]), c_0_6])).
cnf(c_0_21,plain,(rd(X1,ld(X1,mult(X2,X1)))=mult(ld(X2,unit),X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22,plain,(ld(X1,mult(ld(X2,unit),X1))=ld(mult(X2,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_18]), c_0_7]), c_0_19]), c_0_7])).
cnf(c_0_23,plain,(rd(ld(X1,mult(X2,X1)),X1)=mult(ld(X1,unit),X2)), inference(spm,[status(thm)],[c_0_4, c_0_20])).
cnf(c_0_24,plain,(mult(ld(rd(X1,X2),unit),X2)=rd(X2,ld(X2,X1))), inference(spm,[status(thm)],[c_0_21, c_0_8])).
cnf(c_0_25,plain,(mult(X1,ld(mult(X2,X1),X1))=mult(ld(X2,unit),X1)), inference(spm,[status(thm)],[c_0_5, c_0_22])).
cnf(c_0_26,plain,(mult(rd(mult(X1,X2),X1),X3)=mult(X1,mult(X2,ld(X1,X3)))), inference(spm,[status(thm)],[c_0_9, c_0_5])).
cnf(c_0_27,plain,(mult(ld(X1,unit),rd(X2,X1))=rd(ld(X1,X2),X1)), inference(spm,[status(thm)],[c_0_23, c_0_8])).
cnf(c_0_28,plain,(ld(X1,rd(X1,ld(X1,X2)))=ld(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_24]), c_0_8])).
cnf(c_0_29,plain,(mult(ld(rd(X1,X2),unit),X2)=mult(X2,ld(X1,X2))), inference(spm,[status(thm)],[c_0_25, c_0_8])).
cnf(c_0_30,plain,(rd(mult(X1,mult(X2,ld(X1,X3))),X3)=rd(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_4, c_0_26])).
cnf(c_0_31,plain,(rd(rd(ld(X1,X2),X1),rd(X2,X1))=ld(X1,unit)), inference(spm,[status(thm)],[c_0_4, c_0_27])).
cnf(c_0_32,plain,(ld(ld(X1,unit),unit)=X1), inference(spm,[status(thm)],[c_0_0, c_0_16])).
cnf(c_0_33,plain,(rd(X1,ld(X1,X2))=mult(X1,ld(X2,X1))), inference(spm,[status(thm)],[c_0_5, c_0_28])).
cnf(c_0_34,plain,(rd(mult(X1,ld(X2,X1)),X1)=ld(rd(X2,X1),unit)), inference(spm,[status(thm)],[c_0_4, c_0_29])).
cnf(c_0_35,plain,(ld(rd(X1,X2),X1)=X2), inference(spm,[status(thm)],[c_0_0, c_0_8])).
cnf(c_0_36,plain,(rd(mult(X1,ld(ld(X1,X2),unit)),X1)=rd(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_16]), c_0_7])).
cnf(c_0_37,plain,(rd(rd(X1,ld(X1,unit)),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_17]), c_0_32])).
cnf(c_0_38,plain,(rd(rd(X1,X2),rd(mult(X2,X1),X2))=ld(X2,unit)), inference(spm,[status(thm)],[c_0_31, c_0_0])).
cnf(c_0_39,plain,(mult(mult(X1,ld(X2,X1)),ld(X1,X2))=X1), inference(spm,[status(thm)],[c_0_8, c_0_33])).
cnf(c_0_40,plain,(ld(rd(rd(X1,X2),X1),unit)=rd(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41,plain,(mult(X1,ld(ld(X1,X2),unit))=mult(rd(X1,X2),X1)), inference(spm,[status(thm)],[c_0_8, c_0_36])).
cnf(c_0_42,plain,(rd(X1,ld(X1,unit))=mult(X1,X1)), inference(spm,[status(thm)],[c_0_8, c_0_37])).
cnf(c_0_43,plain,(ld(ld(X1,unit),rd(X2,X1))=rd(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_35, c_0_38])).
cnf(c_0_44,plain,(mult(X1,mult(X2,ld(X1,rd(rd(X1,X2),X1))))=unit), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_6]), c_0_19]), c_0_26])).
cnf(c_0_45,plain,(mult(mult(rd(X1,X2),X3),X2)=mult(X1,ld(X2,mult(X3,X2)))), inference(spm,[status(thm)],[c_0_2, c_0_8])).
cnf(c_0_46,plain,(ld(X1,mult(rd(X1,X2),X1))=ld(ld(X1,X2),unit)), inference(spm,[status(thm)],[c_0_0, c_0_41])).
cnf(c_0_47,plain,(mult(X1,mult(X2,ld(X1,unit)))=rd(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_7, c_0_26])).
cnf(c_0_48,plain,(mult(mult(X1,X1),ld(X1,unit))=X1), inference(spm,[status(thm)],[c_0_8, c_0_42])).
cnf(c_0_49,plain,(rd(mult(X1,mult(X2,X1)),X1)=ld(ld(X1,unit),X2)), inference(spm,[status(thm)],[c_0_43, c_0_4])).
cnf(c_0_50,plain,(mult(X1,ld(X2,rd(rd(X2,X1),X2)))=ld(X2,unit)), inference(spm,[status(thm)],[c_0_0, c_0_44])).
cnf(c_0_51,plain,(rd(mult(X1,ld(X2,mult(X3,X2))),X2)=mult(rd(X1,X2),X3)), inference(spm,[status(thm)],[c_0_4, c_0_45])).
cnf(c_0_52,plain,(ld(ld(X1,ld(X2,X1)),unit)=ld(X1,mult(X2,X1))), inference(spm,[status(thm)],[c_0_46, c_0_17])).
cnf(c_0_53,plain,(ld(ld(X1,unit),X1)=mult(X1,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_54,plain,(mult(ld(ld(X1,unit),X2),X1)=mult(X1,mult(X2,X1))), inference(spm,[status(thm)],[c_0_8, c_0_49])).
cnf(c_0_55,plain,(ld(X1,rd(rd(X1,X2),X1))=ld(X2,ld(X1,unit))), inference(spm,[status(thm)],[c_0_0, c_0_50])).
cnf(c_0_56,plain,(rd(mult(X1,ld(X2,X3)),X2)=mult(rd(X1,X2),rd(X3,X2))), inference(spm,[status(thm)],[c_0_51, c_0_8])).
cnf(c_0_57,plain,(mult(ld(X1,ld(X2,X1)),ld(X1,mult(X2,X1)))=unit), inference(spm,[status(thm)],[c_0_5, c_0_52])).
cnf(c_0_58,plain,(rd(unit,X1)=ld(X1,unit)), inference(spm,[status(thm)],[c_0_4, c_0_16])).
cnf(c_0_59,plain,(ld(X1,rd(mult(X1,X2),X1))=mult(X2,ld(X1,unit))), inference(spm,[status(thm)],[c_0_0, c_0_47])).
cnf(c_0_60,plain,(mult(ld(X1,unit),mult(X1,X1))=X1), inference(spm,[status(thm)],[c_0_5, c_0_53])).
cnf(c_0_61,plain,(mult(mult(X1,X1),X1)=mult(X1,mult(X1,X1))), inference(spm,[status(thm)],[c_0_54, c_0_53])).
cnf(c_0_62,plain,(ld(ld(X1,X2),ld(X2,unit))=ld(X2,rd(X1,X2))), inference(spm,[status(thm)],[c_0_55, c_0_17])).
cnf(c_0_63,plain,(mult(rd(ld(X1,ld(X2,X1)),X1),X2)=ld(X1,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), c_0_4])).
cnf(c_0_64,plain,(ld(ld(X1,unit),mult(X1,X1))=mult(X1,mult(X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_42]), c_0_32]), c_0_61])).
cnf(c_0_65,plain,(mult(ld(X1,X2),ld(X2,rd(X1,X2)))=ld(X2,unit)), inference(spm,[status(thm)],[c_0_5, c_0_62])).
cnf(c_0_66,plain,(rd(ld(X1,ld(X2,X1)),X1)=rd(ld(X1,unit),X2)), inference(spm,[status(thm)],[c_0_4, c_0_63])).
cnf(c_0_67,plain,(rd(mult(X1,X1),mult(X1,mult(X1,X1)))=ld(X1,unit)), inference(spm,[status(thm)],[c_0_17, c_0_64])).
cnf(c_0_68,plain,(mult(ld(mult(X1,X2),X2),ld(X2,X1))=ld(X2,unit)), inference(spm,[status(thm)],[c_0_65, c_0_4])).
cnf(c_0_69,plain,(rd(ld(X1,unit),rd(X1,X2))=rd(ld(X1,X2),X1)), inference(spm,[status(thm)],[c_0_66, c_0_35])).
cnf(c_0_70,plain,(rd(mult(X1,mult(X2,X3)),mult(X1,X3))=rd(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_30, c_0_0])).
cnf(c_0_71,plain,(rd(mult(op_c,op_c),op_d)=ld(op_c,unit)), inference(spm,[status(thm)],[c_0_67, c_0_10])).
cnf(c_0_72,plain,(mult(rd(ld(X1,unit),X2),X1)=ld(X1,ld(X2,X1))), inference(spm,[status(thm)],[c_0_8, c_0_66])).
cnf(c_0_73,plain,(rd(ld(X1,unit),ld(X1,X2))=ld(mult(X2,X1),X1)), inference(spm,[status(thm)],[c_0_4, c_0_68])).
cnf(c_0_74,plain,(mult(rd(ld(X1,X2),X1),rd(X1,X2))=ld(X1,unit)), inference(spm,[status(thm)],[c_0_8, c_0_69])).
cnf(c_0_75,plain,(rd(mult(X1,rd(X2,X3)),X1)=rd(mult(X1,X2),mult(X1,X3))), inference(spm,[status(thm)],[c_0_70, c_0_8])).
cnf(c_0_76,plain,(ld(mult(rd(X1,X2),X3),mult(X1,X3))=ld(X3,mult(X2,X3))), inference(spm,[status(thm)],[c_0_15, c_0_8])).
cnf(c_0_77,plain,(mult(ld(op_c,unit),op_d)=mult(op_c,op_c)), inference(spm,[status(thm)],[c_0_8, c_0_71])).
cnf(c_0_78,plain,(ld(X1,ld(ld(X1,X2),X1))=mult(ld(mult(X2,X1),X1),X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_79,plain,(mult(rd(X1,X2),rd(X2,mult(X2,X1)))=ld(X2,unit)), inference(spm,[status(thm)],[c_0_74, c_0_0])).
cnf(c_0_80,plain,(mult(rd(mult(X1,X2),mult(X1,X3)),X1)=mult(X1,rd(X2,X3))), inference(spm,[status(thm)],[c_0_8, c_0_75])).
cnf(c_0_81,plain,(rd(mult(X1,X2),ld(X2,mult(X3,X2)))=mult(rd(X1,X3),X2)), inference(spm,[status(thm)],[c_0_17, c_0_76])).
cnf(c_0_82,plain,(ld(mult(X1,X1),X1)=ld(X1,unit)), inference(spm,[status(thm)],[c_0_35, c_0_42])).
cnf(c_0_83,plain,(rd(X1,mult(X1,X1))=ld(X1,unit)), inference(spm,[status(thm)],[c_0_17, c_0_53])).
cnf(c_0_84,plain,(ld(op_c,mult(op_d,op_c))=op_d), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_77]), c_0_61]), c_0_10])).
cnf(c_0_85,plain,(mult(X1,mult(ld(mult(X2,X1),X1),X1))=ld(ld(X1,X2),X1)), inference(spm,[status(thm)],[c_0_5, c_0_78])).
cnf(c_0_86,plain,(ld(rd(X1,X2),ld(X2,unit))=rd(X2,mult(X2,X1))), inference(spm,[status(thm)],[c_0_0, c_0_79])).
cnf(c_0_87,plain,(ld(mult(X1,ld(X2,X1)),X1)=ld(X1,X2)), inference(spm,[status(thm)],[c_0_35, c_0_33])).
cnf(c_0_88,plain,(mult(rd(X1,mult(rd(X1,X2),X1)),X1)=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_46]), c_0_32]), c_0_5])).
cnf(c_0_89,plain,(mult(X1,rd(X2,ld(X1,X3)))=mult(rd(mult(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_80, c_0_5])).
cnf(c_0_90,plain,(rd(mult(X1,X2),ld(X2,X3))=mult(rd(X1,rd(X3,X2)),X2)), inference(spm,[status(thm)],[c_0_81, c_0_8])).
cnf(c_0_91,plain,(ld(X1,mult(ld(X2,X3),X1))=ld(mult(X2,X1),mult(X3,X1))), inference(spm,[status(thm)],[c_0_15, c_0_5])).
cnf(c_0_92,plain,(mult(X1,ld(mult(X1,X1),unit))=ld(X1,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_61]), c_0_4]), c_0_82])).
cnf(c_0_93,plain,(rd(ld(X1,unit),X1)=ld(mult(X1,X1),unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_61]), c_0_83]), c_0_4])).
cnf(c_0_94,plain,(mult(mult(X1,op_c),op_d)=mult(mult(X1,op_d),op_c)), inference(spm,[status(thm)],[c_0_2, c_0_84])).
cnf(c_0_95,plain,(rd(ld(ld(X1,X2),X1),mult(X1,X1))=ld(X2,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_85]), c_0_25]), c_0_4])).
cnf(c_0_96,plain,(rd(X1,mult(X1,mult(X2,X1)))=ld(X2,ld(X1,unit))), inference(spm,[status(thm)],[c_0_86, c_0_4])).
cnf(c_0_97,plain,(ld(X1,rd(X1,X2))=ld(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_87, c_0_35])).
cnf(c_0_98,plain,(rd(X1,mult(rd(X1,X2),X1))=rd(X2,X1)), inference(spm,[status(thm)],[c_0_4, c_0_88])).
cnf(c_0_99,plain,(ld(X1,mult(rd(mult(X1,X2),X3),X1))=rd(X2,ld(X1,X3))), inference(spm,[status(thm)],[c_0_0, c_0_89])).
cnf(c_0_100,plain,(mult(rd(X1,rd(mult(X2,X3),X2)),X2)=rd(mult(X1,X2),X3)), inference(spm,[status(thm)],[c_0_90, c_0_0])).
cnf(c_0_101,plain,(mult(X1,ld(mult(X2,X1),mult(X3,X1)))=mult(ld(X2,X3),X1)), inference(spm,[status(thm)],[c_0_5, c_0_91])).
cnf(c_0_102,plain,(rd(ld(X1,unit),ld(mult(X1,X1),unit))=X1), inference(spm,[status(thm)],[c_0_4, c_0_92])).
cnf(c_0_103,plain,(ld(ld(mult(X1,X1),unit),ld(X1,unit))=X1), inference(spm,[status(thm)],[c_0_35, c_0_93])).
cnf(c_0_104,plain,(rd(mult(mult(X1,op_d),op_c),op_d)=mult(X1,op_c)), inference(spm,[status(thm)],[c_0_4, c_0_94])).
cnf(c_0_105,plain,(ld(ld(X1,X2),X1)=mult(ld(X2,unit),mult(X1,X1))), inference(spm,[status(thm)],[c_0_8, c_0_95])).
cnf(c_0_106,plain,(ld(ld(X1,ld(X2,unit)),X2)=mult(X2,mult(X1,X2))), inference(spm,[status(thm)],[c_0_35, c_0_96])).
cnf(c_0_107,plain,(ld(op_e,op_d)=ld(op_d,unit)), inference(spm,[status(thm)],[c_0_82, c_0_11])).
cnf(c_0_108,plain,(ld(mult(X1,mult(rd(X1,X2),X1)),X1)=ld(X1,rd(X2,X1))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_109,plain,(ld(X1,rd(mult(mult(X1,X2),X1),X3))=rd(X2,mult(X3,ld(X1,unit)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_59])).
cnf(c_0_110,plain,(mult(ld(X1,rd(X2,X3)),X3)=mult(X3,ld(mult(X1,X3),X2))), inference(spm,[status(thm)],[c_0_101, c_0_8])).
cnf(c_0_111,plain,(rd(X1,ld(mult(X1,X1),unit))=mult(X1,mult(X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_102]), c_0_32]), c_0_61]), c_0_103])).
cnf(c_0_112,plain,(rd(mult(X1,op_c),op_d)=mult(rd(X1,op_d),op_c)), inference(spm,[status(thm)],[c_0_104, c_0_8])).
cnf(c_0_113,plain,(rd(ld(X1,X2),X2)=rd(X2,mult(X1,X2))), inference(spm,[status(thm)],[c_0_98, c_0_17])).
cnf(c_0_114,plain,(mult(rd(X1,mult(X1,X2)),X1)=mult(X1,ld(X2,unit))), inference(spm,[status(thm)],[c_0_41, c_0_0])).
cnf(c_0_115,plain,(rd(X1,mult(ld(X2,unit),mult(X1,X1)))=ld(X1,X2)), inference(spm,[status(thm)],[c_0_17, c_0_105])).
cnf(c_0_116,plain,(ld(ld(X1,rd(X2,X1)),X1)=mult(X1,mult(ld(X2,X1),X1))), inference(spm,[status(thm)],[c_0_106, c_0_62])).
cnf(c_0_117,plain,(mult(op_e,ld(op_d,unit))=op_d), inference(spm,[status(thm)],[c_0_5, c_0_107])).
cnf(c_0_118,plain,(ld(ld(ld(X1,unit),rd(X1,X2)),unit)=rd(rd(X2,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_108]), c_0_5]), c_0_49])).
cnf(c_0_119,plain,(ld(ld(op_c,unit),rd(op_c,X1))=rd(op_d,mult(X1,op_c))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_77]), c_0_48]), c_0_32])).
cnf(c_0_120,plain,(ld(rd(X1,X2),X2)=mult(rd(X2,X1),X2)), inference(spm,[status(thm)],[c_0_35, c_0_98])).
cnf(c_0_121,plain,(rd(mult(X1,ld(mult(X2,X1),X3)),X1)=ld(X2,rd(X3,X1))), inference(spm,[status(thm)],[c_0_4, c_0_110])).
cnf(c_0_122,plain,(mult(X1,ld(X2,ld(X1,unit)))=rd(rd(X1,X2),X1)), inference(spm,[status(thm)],[c_0_5, c_0_55])).
cnf(c_0_123,plain,(ld(mult(X1,mult(X1,X1)),X1)=ld(mult(X1,X1),unit)), inference(spm,[status(thm)],[c_0_35, c_0_111])).
cnf(c_0_124,plain,(mult(rd(rd(X1,op_c),op_d),op_c)=rd(X1,op_d)), inference(spm,[status(thm)],[c_0_112, c_0_8])).
cnf(c_0_125,plain,(mult(ld(X1,unit),ld(rd(X2,X1),unit))=rd(X1,mult(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_34]), c_0_0]), c_0_113])).
cnf(c_0_126,plain,(rd(X1,unit)=X1), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_127,plain,(rd(mult(X1,ld(X2,unit)),X1)=rd(X1,mult(X1,X2))), inference(spm,[status(thm)],[c_0_4, c_0_114])).
cnf(c_0_128,plain,(rd(X1,mult(X2,mult(X1,X1)))=ld(X1,ld(X2,unit))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_35]), c_0_58])).
cnf(c_0_129,plain,(ld(mult(ld(X1,X2),X2),unit)=ld(mult(X2,X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_116]), c_0_0]), c_0_91]), c_0_8])).
cnf(c_0_130,plain,(rd(op_d,op_e)=ld(op_d,unit)), inference(spm,[status(thm)],[c_0_83, c_0_11])).
cnf(c_0_131,plain,(rd(mult(op_d,op_e),op_d)=op_e), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_117]), c_0_11])).
cnf(c_0_132,plain,(ld(rd(op_d,mult(X1,op_c)),unit)=rd(rd(X1,op_c),op_c)), inference(spm,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_133,plain,(mult(rd(X1,mult(X2,X1)),X1)=ld(X2,X1)), inference(spm,[status(thm)],[c_0_120, c_0_4])).
cnf(c_0_134,plain,(rd(rd(rd(X1,mult(X2,X1)),X1),X1)=ld(X2,ld(mult(X1,X1),unit))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_93])).
cnf(c_0_135,plain,(ld(mult(op_c,op_c),unit)=ld(op_d,op_c)), inference(spm,[status(thm)],[c_0_123, c_0_10])).
cnf(c_0_136,plain,(rd(rd(X1,op_c),op_d)=rd(rd(X1,op_d),op_c)), inference(spm,[status(thm)],[c_0_4, c_0_124])).
cnf(c_0_137,plain,(mult(rd(X1,X2),rd(X2,mult(X1,X2)))=rd(rd(X1,X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_125]), c_0_126]), c_0_127]), c_0_8])).
cnf(c_0_138,plain,(rd(rd(mult(X1,X2),X1),X1)=rd(X1,rd(mult(X1,X1),X2))), inference(spm,[status(thm)],[c_0_98, c_0_100])).
cnf(c_0_139,plain,(ld(op_d,ld(X1,unit))=rd(op_d,mult(X1,op_e))), inference(spm,[status(thm)],[c_0_128, c_0_11])).
cnf(c_0_140,plain,(ld(mult(X1,X1),rd(X1,X2))=ld(mult(X2,X1),unit)), inference(spm,[status(thm)],[c_0_129, c_0_35])).
cnf(c_0_141,plain,(mult(ld(op_d,unit),op_e)=op_d), inference(spm,[status(thm)],[c_0_8, c_0_130])).
cnf(c_0_142,plain,(mult(op_e,op_d)=mult(op_d,op_e)), inference(spm,[status(thm)],[c_0_8, c_0_131])).
cnf(c_0_143,plain,(mult(mult(X1,rd(X2,X3)),X3)=mult(mult(X1,X3),ld(X3,X2))), inference(spm,[status(thm)],[c_0_2, c_0_8])).
cnf(c_0_144,plain,(rd(mult(X1,rd(X2,ld(X1,unit))),X1)=mult(X1,X2)), inference(spm,[status(thm)],[c_0_47, c_0_8])).
cnf(c_0_145,plain,(ld(rd(op_d,ld(X1,op_c)),unit)=ld(X1,ld(op_d,op_c))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_134]), c_0_135])).
cnf(c_0_146,plain,(mult(rd(rd(X1,op_d),op_c),op_d)=rd(X1,op_c)), inference(spm,[status(thm)],[c_0_8, c_0_136])).
cnf(c_0_147,plain,(mult(rd(rd(X1,X2),X1),X2)=mult(X1,ld(X2,ld(X1,X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_137]), c_0_133])).
cnf(c_0_148,plain,(mult(rd(X1,rd(mult(X1,X1),X2)),X1)=rd(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_8, c_0_138])).
cnf(c_0_149,plain,(rd(rd(mult(X1,X2),X3),X2)=rd(X1,rd(mult(X2,X3),X2))), inference(spm,[status(thm)],[c_0_4, c_0_100])).
cnf(c_0_150,plain,(mult(op_d,rd(op_d,mult(X1,op_e)))=ld(X1,unit)), inference(spm,[status(thm)],[c_0_5, c_0_139])).
cnf(c_0_151,plain,(ld(ld(mult(X1,X1),X2),unit)=mult(ld(X2,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_129]), c_0_19]), c_0_126]), c_0_19])).
cnf(c_0_152,plain,(mult(mult(X1,X1),ld(mult(X2,X1),unit))=rd(X1,X2)), inference(spm,[status(thm)],[c_0_5, c_0_140])).
cnf(c_0_153,plain,(ld(op_f,op_e)=ld(op_e,unit)), inference(spm,[status(thm)],[c_0_82, c_0_12])).
cnf(c_0_154,plain,(rd(rd(op_e,ld(op_d,unit)),op_e)=op_d), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_141]), c_0_42]), c_0_11]), c_0_32])).
cnf(c_0_155,plain,(ld(ld(op_d,unit),op_e)=mult(op_d,op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_141]), c_0_42]), c_0_11]), c_0_32]), c_0_142])).
cnf(c_0_156,plain,(mult(mult(X1,op_e),op_d)=mult(mult(X1,op_d),op_e)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_131]), c_0_0])).
cnf(c_0_157,plain,(mult(X1,rd(X2,ld(X1,unit)))=mult(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_8, c_0_144])).
cnf(c_0_158,plain,(ld(ld(X1,ld(op_d,op_c)),unit)=rd(op_d,ld(X1,op_c))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_145]), c_0_19]), c_0_126]), c_0_19])).
cnf(c_0_159,plain,(rd(mult(X1,op_d),op_c)=mult(rd(X1,op_c),op_d)), inference(spm,[status(thm)],[c_0_146, c_0_4])).
cnf(c_0_160,plain,(rd(X1,rd(mult(X2,X1),X2))=mult(X1,rd(rd(X2,X1),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_147]), c_0_120]), c_0_114]), c_0_0]), c_0_26]), c_0_122]), c_0_138]), c_0_148]), c_0_149])).
cnf(c_0_161,plain,(ld(op_e,rd(op_d,X1))=ld(mult(X1,op_d),unit)), inference(spm,[status(thm)],[c_0_140, c_0_11])).
cnf(c_0_162,plain,(ld(ld(X1,mult(X2,X1)),unit)=ld(X1,ld(X2,X1))), inference(spm,[status(thm)],[c_0_46, c_0_133])).
cnf(c_0_163,plain,(ld(rd(X1,op_e),unit)=mult(op_d,rd(op_d,X1))), inference(spm,[status(thm)],[c_0_150, c_0_8])).
cnf(c_0_164,plain,(ld(ld(op_e,X1),unit)=mult(ld(X1,op_d),op_d)), inference(spm,[status(thm)],[c_0_151, c_0_11])).
cnf(c_0_165,plain,(rd(X1,rd(X2,X1))=mult(mult(X1,X1),ld(X2,unit))), inference(spm,[status(thm)],[c_0_152, c_0_8])).
cnf(c_0_166,plain,(mult(op_f,ld(op_e,unit))=op_e), inference(spm,[status(thm)],[c_0_5, c_0_153])).
cnf(c_0_167,plain,(rd(op_e,ld(op_d,unit))=mult(op_d,op_e)), inference(spm,[status(thm)],[c_0_8, c_0_154])).
cnf(c_0_168,plain,(mult(op_d,mult(op_d,op_e))=op_f), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_155]), c_0_156]), c_0_11]), c_0_12]), c_0_142])).
cnf(c_0_169,plain,(rd(mult(mult(X1,X2),X1),rd(X2,ld(X1,unit)))=X1), inference(spm,[status(thm)],[c_0_4, c_0_157])).
cnf(c_0_170,plain,(mult(ld(rd(X1,X2),X3),X2)=mult(X2,ld(X1,mult(X3,X2)))), inference(spm,[status(thm)],[c_0_101, c_0_8])).
cnf(c_0_171,plain,(mult(rd(X1,ld(op_d,op_c)),X1)=mult(mult(rd(X1,op_c),op_d),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_158]), c_0_89]), c_0_159])).
cnf(c_0_172,plain,(mult(mult(X1,X2),ld(mult(X2,X1),X2))=rd(mult(X1,X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_160]), c_0_143]), c_0_97])).
cnf(c_0_173,plain,(mult(op_e,ld(mult(X1,op_d),unit))=rd(op_d,X1)), inference(spm,[status(thm)],[c_0_5, c_0_161])).
cnf(c_0_174,plain,(rd(mult(X1,X2),mult(X1,ld(X3,X2)))=rd(mult(X1,X3),X1)), inference(spm,[status(thm)],[c_0_70, c_0_5])).
cnf(c_0_175,plain,(mult(ld(mult(mult(X1,X2),X1),X1),X1)=ld(X1,ld(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_15]), c_0_162])).
cnf(c_0_176,plain,(mult(op_d,rd(op_d,rd(op_e,X1)))=rd(mult(op_e,X1),op_e)), inference(spm,[status(thm)],[c_0_40, c_0_163])).
cnf(c_0_177,plain,(mult(ld(mult(op_e,X1),op_d),op_d)=ld(X1,unit)), inference(spm,[status(thm)],[c_0_164, c_0_0])).
cnf(c_0_178,plain,(rd(op_e,ld(op_d,X1))=mult(mult(op_e,ld(X1,unit)),op_d)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_11]), c_0_165]), c_0_11])).
cnf(c_0_179,plain,(mult(X1,mult(ld(X1,X2),ld(X1,X3)))=mult(rd(X2,X1),X3)), inference(spm,[status(thm)],[c_0_26, c_0_5])).
cnf(c_0_180,plain,(rd(mult(op_e,op_f),op_e)=op_f), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_166]), c_0_12])).
cnf(c_0_181,plain,(rd(mult(op_d,op_e),ld(op_d,unit))=op_f), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_167]), c_0_32]), c_0_168]), c_0_155])).
cnf(c_0_182,plain,(ld(X1,mult(mult(X1,X2),X1))=rd(X2,ld(X1,unit))), inference(spm,[status(thm)],[c_0_35, c_0_169])).
cnf(c_0_183,plain,(mult(mult(rd(X1,X2),X1),X1)=mult(X1,ld(X2,mult(X1,X1)))), inference(spm,[status(thm)],[c_0_170, c_0_120])).
cnf(c_0_184,plain,(rd(X1,ld(op_d,op_c))=mult(rd(X1,op_c),op_d)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_171]), c_0_4])).
cnf(c_0_185,plain,(mult(mult(X1,mult(X2,X2)),X2)=mult(mult(X1,X2),mult(X2,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_61]), c_0_0])).
cnf(c_0_186,plain,(ld(mult(rd(X1,X2),X1),rd(X1,X2))=rd(X2,mult(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_137]), c_0_97])).
cnf(c_0_187,plain,(ld(mult(mult(X1,X2),X1),mult(X1,X2))=ld(mult(X2,X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_172]), c_0_97])).
cnf(c_0_188,plain,(ld(mult(X1,X1),mult(X1,mult(X1,X1)))=X1), inference(spm,[status(thm)],[c_0_0, c_0_61])).
cnf(c_0_189,plain,(rd(op_d,rd(X1,op_d))=mult(op_e,ld(X1,unit))), inference(spm,[status(thm)],[c_0_173, c_0_8])).
cnf(c_0_190,plain,(ld(rd(mult(X1,X2),X1),mult(X1,X3))=mult(X1,ld(X2,X3))), inference(spm,[status(thm)],[c_0_35, c_0_174])).
cnf(c_0_191,plain,(ld(mult(X1,ld(op_d,unit)),unit)=ld(op_d,rd(op_e,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176]), c_0_26]), c_0_107]), c_0_177]), c_0_35])).
cnf(c_0_192,plain,(rd(op_e,ld(X1,ld(op_d,unit)))=mult(mult(op_d,op_e),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_55]), c_0_40]), c_0_143]), c_0_142]), c_0_0])).
cnf(c_0_193,plain,(mult(ld(X1,X2),ld(X1,unit))=ld(X1,rd(X2,X1))), inference(spm,[status(thm)],[c_0_59, c_0_5])).
cnf(c_0_194,plain,(ld(X1,ld(X1,unit))=ld(mult(X1,X1),unit)), inference(spm,[status(thm)],[c_0_0, c_0_92])).
cnf(c_0_195,plain,(ld(X1,mult(rd(X2,X1),X3))=mult(ld(X1,X2),ld(X1,X3))), inference(spm,[status(thm)],[c_0_0, c_0_179])).
cnf(c_0_196,plain,(mult(mult(X1,op_f),op_e)=mult(mult(X1,op_e),op_f)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_180]), c_0_0])).
cnf(c_0_197,plain,(mult(op_f,ld(op_d,unit))=mult(op_d,op_e)), inference(spm,[status(thm)],[c_0_8, c_0_181])).
cnf(c_0_198,plain,(rd(op_e,rd(op_f,op_d))=ld(op_d,unit)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_168]), c_0_131])).
cnf(c_0_199,plain,(rd(ld(ld(X1,X2),unit),ld(X1,unit))=ld(X2,mult(X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_41]), c_0_183]), c_0_0])).
cnf(c_0_200,plain,(ld(mult(rd(X1,op_c),op_d),X1)=ld(op_d,op_c)), inference(spm,[status(thm)],[c_0_35, c_0_184])).
cnf(c_0_201,plain,(mult(mult(X1,X1),mult(X1,X1))=mult(X1,mult(X1,mult(X1,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_64]), c_0_61]), c_0_185])).
cnf(c_0_202,plain,(mult(op_c,op_d)=mult(op_d,op_c)), inference(spm,[status(thm)],[c_0_5, c_0_84])).
cnf(c_0_203,plain,(rd(ld(X1,X2),mult(X2,ld(X1,X2)))=ld(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_186, c_0_17])).
cnf(c_0_204,plain,(mult(X1,ld(mult(X1,X2),X1))=rd(X1,X2)), inference(spm,[status(thm)],[c_0_33, c_0_0])).
cnf(c_0_205,plain,(ld(mult(ld(X1,X2),X1),ld(X1,X2))=ld(mult(X2,X1),X2)), inference(spm,[status(thm)],[c_0_187, c_0_5])).
cnf(c_0_206,plain,(mult(ld(mult(X1,X2),unit),mult(X1,X1))=ld(X2,X1)), inference(spm,[status(thm)],[c_0_105, c_0_0])).
cnf(c_0_207,plain,(mult(op_d,ld(mult(X1,op_d),op_e))=mult(ld(X1,op_d),op_d)), inference(spm,[status(thm)],[c_0_101, c_0_11])).
cnf(c_0_208,plain,(rd(mult(X1,X2),mult(X2,X2))=mult(rd(X1,mult(X2,X2)),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_188]), c_0_4])).
cnf(c_0_209,plain,(ld(mult(op_e,ld(X1,unit)),op_d)=rd(X1,op_d)), inference(spm,[status(thm)],[c_0_35, c_0_189])).
cnf(c_0_210,plain,(ld(rd(mult(X1,X2),X1),X3)=mult(X1,ld(X2,ld(X1,X3)))), inference(spm,[status(thm)],[c_0_190, c_0_5])).
cnf(c_0_211,plain,(rd(mult(X1,mult(X1,ld(X2,X1))),X1)=rd(mult(X1,X1),X2)), inference(spm,[status(thm)],[c_0_30, c_0_39])).
cnf(c_0_212,plain,(ld(op_e,ld(X1,unit))=rd(op_e,mult(X1,op_f))), inference(spm,[status(thm)],[c_0_128, c_0_12])).
cnf(c_0_213,plain,(ld(op_f,rd(op_e,X1))=ld(mult(X1,op_e),unit)), inference(spm,[status(thm)],[c_0_140, c_0_12])).
cnf(c_0_214,plain,(rd(op_f,mult(op_e,ld(X1,op_e)))=rd(mult(op_e,X1),op_e)), inference(spm,[status(thm)],[c_0_174, c_0_12])).
cnf(c_0_215,plain,(ld(ld(op_e,unit),X1)=ld(op_d,mult(mult(op_d,op_e),X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_191]), c_0_192]), c_0_193]), c_0_58]), c_0_194]), c_0_11])).
cnf(c_0_216,plain,(mult(ld(X1,mult(X2,X1)),ld(X1,X3))=ld(X1,mult(X2,X3))), inference(spm,[status(thm)],[c_0_195, c_0_4])).
cnf(c_0_217,plain,(rd(mult(mult(X1,op_e),op_f),op_e)=mult(X1,op_f)), inference(spm,[status(thm)],[c_0_4, c_0_196])).
cnf(c_0_218,plain,(rd(mult(op_d,op_f),op_d)=op_f), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_197]), c_0_168])).
cnf(c_0_219,plain,(rd(op_f,op_d)=mult(op_d,op_e)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_198]), c_0_155])).
cnf(c_0_220,plain,(ld(X1,mult(op_d,op_c))=mult(ld(X1,op_c),op_d)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200]), c_0_58]), c_0_77]), c_0_151]), c_0_184]), c_0_4]), c_0_58]), c_0_77]), c_0_58]), c_0_77]), c_0_201]), c_0_10]), c_0_202])).
cnf(c_0_221,plain,(mult(ld(mult(X1,X2),X1),mult(X2,ld(X1,X2)))=ld(X1,X2)), inference(spm,[status(thm)],[c_0_8, c_0_203])).
cnf(c_0_222,plain,(mult(ld(X1,X2),ld(mult(X2,X1),X2))=rd(ld(X1,X2),X1)), inference(spm,[status(thm)],[c_0_204, c_0_205])).
cnf(c_0_223,plain,(ld(ld(mult(X1,op_d),op_e),op_d)=mult(ld(op_e,X1),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207]), c_0_129]), c_0_11]), c_0_11])).
cnf(c_0_224,plain,(mult(rd(rd(X1,X2),mult(X2,X2)),X2)=rd(X1,mult(X2,X2))), inference(spm,[status(thm)],[c_0_208, c_0_8])).
cnf(c_0_225,plain,(ld(mult(rd(op_e,X1),op_e),op_d)=rd(ld(op_e,X1),op_d)), inference(spm,[status(thm)],[c_0_209, c_0_41])).
cnf(c_0_226,plain,(mult(X1,ld(ld(X1,X2),ld(X1,X3)))=ld(rd(X2,X1),X3)), inference(spm,[status(thm)],[c_0_210, c_0_5])).
cnf(c_0_227,plain,(rd(mult(X1,X1),rd(X1,X2))=rd(mult(X1,mult(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_211, c_0_35])).
cnf(c_0_228,plain,(rd(op_e,rd(op_d,mult(X1,op_e)))=mult(mult(op_e,X1),op_d)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_139]), c_0_32])).
cnf(c_0_229,plain,(ld(rd(X1,X2),rd(X3,X2))=rd(mult(X2,ld(X1,X3)),X2)), inference(spm,[status(thm)],[c_0_121, c_0_8])).
cnf(c_0_230,plain,(rd(op_f,op_e)=op_e), inference(spm,[status(thm)],[c_0_4, c_0_12])).
cnf(c_0_231,plain,(mult(X1,rd(ld(X1,X2),X3))=mult(rd(X2,mult(X1,X3)),X1)), inference(spm,[status(thm)],[c_0_80, c_0_5])).
cnf(c_0_232,plain,(rd(op_e,mult(ld(X1,unit),op_f))=ld(op_e,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_35]), c_0_58])).
cnf(c_0_233,plain,(mult(ld(mult(mult(X1,X2),X1),X1),X2)=ld(X1,unit)), inference(spm,[status(thm)],[c_0_68, c_0_0])).
cnf(c_0_234,plain,(mult(op_f,ld(mult(X1,op_e),unit))=rd(op_e,X1)), inference(spm,[status(thm)],[c_0_5, c_0_213])).
cnf(c_0_235,plain,(ld(op_d,mult(mult(op_d,op_e),ld(X1,unit)))=rd(op_f,mult(X1,op_e))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_18]), c_0_5]), c_0_49]), c_0_215])).
cnf(c_0_236,plain,(ld(X1,mult(X2,mult(X1,X3)))=mult(ld(X1,mult(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_216, c_0_0])).
cnf(c_0_237,plain,(ld(op_d,op_f)=mult(op_d,op_e)), inference(spm,[status(thm)],[c_0_0, c_0_168])).
cnf(c_0_238,plain,(rd(mult(X1,op_f),op_e)=mult(rd(X1,op_e),op_f)), inference(spm,[status(thm)],[c_0_217, c_0_8])).
cnf(c_0_239,plain,(mult(ld(op_d,unit),op_f)=mult(op_d,op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_218]), c_0_0]), c_0_219])).
cnf(c_0_240,plain,(mult(X1,mult(ld(X1,op_c),op_d))=mult(op_d,op_c)), inference(spm,[status(thm)],[c_0_5, c_0_220])).
cnf(c_0_241,plain,(mult(ld(X1,X2),rd(ld(X2,X1),X2))=ld(mult(X1,X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_205]), c_0_187]), c_0_5]), c_0_222])).
cnf(c_0_242,plain,(ld(mult(mult(X1,X2),X2),mult(X1,X2))=ld(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_87, c_0_0])).
cnf(c_0_243,plain,(mult(op_e,ld(op_f,rd(X1,op_d)))=rd(ld(op_e,X1),op_d)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224]), c_0_11]), c_0_120]), c_0_225]), c_0_11]), c_0_110]), c_0_12])).
cnf(c_0_244,plain,(rd(mult(mult(X1,op_d),op_e),op_d)=mult(X1,op_e)), inference(spm,[status(thm)],[c_0_4, c_0_156])).
cnf(c_0_245,plain,(rd(mult(op_d,ld(op_e,X1)),op_d)=ld(op_d,rd(X1,op_d))), inference(spm,[status(thm)],[c_0_121, c_0_11])).
cnf(c_0_246,plain,(ld(ld(X1,X2),ld(X1,X3))=ld(X1,ld(rd(X2,X1),X3))), inference(spm,[status(thm)],[c_0_0, c_0_226])).
cnf(c_0_247,plain,(rd(mult(op_d,mult(op_d,X1)),op_d)=rd(op_e,rd(op_d,X1))), inference(spm,[status(thm)],[c_0_227, c_0_11])).
cnf(c_0_248,plain,(rd(op_e,rd(op_d,X1))=mult(mult(op_e,rd(X1,op_e)),op_d)), inference(spm,[status(thm)],[c_0_228, c_0_8])).
cnf(c_0_249,plain,(rd(ld(X1,mult(X2,X3)),ld(X1,X3))=ld(X1,mult(X2,X1))), inference(spm,[status(thm)],[c_0_4, c_0_216])).
cnf(c_0_250,plain,(rd(mult(op_e,ld(X1,op_f)),op_e)=mult(rd(op_e,X1),op_e)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230]), c_0_120])).
cnf(c_0_251,plain,(ld(X1,mult(rd(X2,mult(X1,X3)),X1))=rd(ld(X1,X2),X3)), inference(spm,[status(thm)],[c_0_0, c_0_231])).
cnf(c_0_252,plain,(rd(op_e,mult(op_f,ld(X1,op_f)))=ld(op_e,rd(X1,op_f))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_170]), c_0_6])).
cnf(c_0_253,plain,(rd(ld(X1,unit),X2)=ld(mult(mult(X1,X2),X1),X1)), inference(spm,[status(thm)],[c_0_4, c_0_233])).
cnf(c_0_254,plain,(rd(op_e,rd(X1,op_e))=mult(op_f,ld(X1,unit))), inference(spm,[status(thm)],[c_0_234, c_0_8])).
cnf(c_0_255,plain,(mult(mult(op_d,op_e),rd(op_d,X1))=rd(op_f,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_163]), c_0_236]), c_0_156]), c_0_11]), c_0_12]), c_0_237]), c_0_8])).
cnf(c_0_256,plain,(mult(op_d,ld(op_e,unit))=ld(op_d,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_142]), c_0_4]), c_0_107])).
cnf(c_0_257,plain,(mult(rd(ld(op_d,unit),op_e),op_f)=op_d), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239]), c_0_4])).
cnf(c_0_258,plain,(rd(X1,mult(X2,ld(X3,ld(X2,X1))))=rd(mult(X2,X3),X2)), inference(spm,[status(thm)],[c_0_174, c_0_5])).
cnf(c_0_259,plain,(rd(mult(op_d,op_c),mult(X1,op_d))=rd(op_c,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_240]), c_0_5])).
cnf(c_0_260,plain,(mult(X1,rd(X2,mult(mult(X2,X1),X2)))=rd(X1,mult(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_97]), c_0_35]), c_0_113]), c_0_186])).
cnf(c_0_261,plain,(mult(ld(mult(X1,mult(X2,X1)),X1),X2)=ld(mult(X2,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_242]), c_0_187]), c_0_5])).
cnf(c_0_262,plain,(ld(op_e,rd(ld(op_e,X1),op_d))=ld(op_f,rd(X1,op_d))), inference(spm,[status(thm)],[c_0_0, c_0_243])).
cnf(c_0_263,plain,(rd(mult(X1,op_e),op_d)=mult(rd(X1,op_d),op_e)), inference(spm,[status(thm)],[c_0_244, c_0_8])).
cnf(c_0_264,plain,(mult(mult(rd(X1,op_f),op_d),op_e)=mult(X1,ld(op_d,unit))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_245]), c_0_246]), c_0_58]), c_0_43]), c_0_59]), c_0_247]), c_0_248]), c_0_231]), c_0_12]), c_0_156])).
cnf(c_0_265,plain,(rd(ld(X1,X2),ld(X1,X3))=ld(X1,mult(rd(X2,X3),X1))), inference(spm,[status(thm)],[c_0_249, c_0_8])).
cnf(c_0_266,plain,(ld(op_e,op_f)=op_e), inference(spm,[status(thm)],[c_0_0, c_0_12])).
cnf(c_0_267,plain,(mult(ld(X1,op_f),ld(op_e,unit))=mult(ld(X1,op_d),op_d)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_250]), c_0_46]), c_0_164])).
cnf(c_0_268,plain,(ld(mult(mult(op_e,ld(X1,op_f)),op_e),op_e)=ld(mult(op_e,op_f),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_252]), c_0_110]), c_0_0]), c_0_153]), c_0_253])).
cnf(c_0_269,plain,(ld(mult(op_f,ld(X1,unit)),op_e)=rd(X1,op_e)), inference(spm,[status(thm)],[c_0_35, c_0_254])).
cnf(c_0_270,plain,(rd(rd(op_f,X1),rd(op_d,X1))=mult(op_d,op_e)), inference(spm,[status(thm)],[c_0_4, c_0_255])).
cnf(c_0_271,plain,(rd(ld(op_d,unit),ld(op_e,unit))=op_d), inference(spm,[status(thm)],[c_0_4, c_0_256])).
cnf(c_0_272,plain,(ld(ld(op_e,unit),ld(op_d,unit))=op_d), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_141]), c_0_142]), c_0_4])).
cnf(c_0_273,plain,(rd(ld(op_d,unit),op_e)=rd(op_d,op_f)), inference(spm,[status(thm)],[c_0_4, c_0_257])).
cnf(c_0_274,plain,(mult(op_d,mult(op_f,ld(op_d,X1)))=mult(op_f,X1)), inference(spm,[status(thm)],[c_0_26, c_0_218])).
cnf(c_0_275,plain,(rd(mult(rd(X1,X2),X1),X2)=rd(X1,mult(rd(X2,X1),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_82]), c_0_41]), c_0_179])).
cnf(c_0_276,plain,(rd(op_c,rd(X1,op_d))=rd(mult(op_d,op_c),X1)), inference(spm,[status(thm)],[c_0_259, c_0_8])).
cnf(c_0_277,plain,(ld(mult(X1,op_d),mult(X1,op_c))=ld(op_d,op_c)), inference(spm,[status(thm)],[c_0_200, c_0_4])).
cnf(c_0_278,plain,(mult(ld(X1,X2),rd(X1,mult(X2,X1)))=rd(X2,mult(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_5]), c_0_113])).
cnf(c_0_279,plain,(mult(ld(X1,rd(X2,X1)),rd(X1,X2))=rd(X2,mult(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_108]), c_0_186])).
cnf(c_0_280,plain,(ld(op_f,rd(mult(op_e,X1),op_d))=ld(op_e,rd(X1,op_d))), inference(spm,[status(thm)],[c_0_262, c_0_0])).
cnf(c_0_281,plain,(mult(rd(X1,op_d),ld(op_d,unit))=mult(rd(X1,op_f),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_264]), c_0_56]), c_0_58]), c_0_4])).
cnf(c_0_282,plain,(ld(op_e,mult(rd(X1,op_f),op_e))=rd(ld(op_e,X1),op_e)), inference(spm,[status(thm)],[c_0_265, c_0_266])).
cnf(c_0_283,plain,(mult(ld(op_f,unit),ld(op_e,X1))=ld(mult(op_e,op_f),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_267]), c_0_193]), c_0_58]), c_0_194]), c_0_12]), c_0_129]), c_0_11]), c_0_253]), c_0_268])).
cnf(c_0_284,plain,(ld(op_e,mult(rd(op_f,X1),op_e))=mult(op_e,ld(X1,op_e))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_266]), c_0_33])).
cnf(c_0_285,plain,(rd(ld(X1,X2),mult(X2,X2))=ld(mult(X2,X1),unit)), inference(spm,[status(thm)],[c_0_4, c_0_206])).
cnf(c_0_286,plain,(ld(mult(rd(op_f,X1),op_f),op_e)=rd(ld(op_f,X1),op_e)), inference(spm,[status(thm)],[c_0_269, c_0_41])).
cnf(c_0_287,plain,(mult(op_f,op_e)=mult(op_e,op_f)), inference(spm,[status(thm)],[c_0_8, c_0_180])).
cnf(c_0_288,plain,(rd(op_e,op_f)=ld(op_e,unit)), inference(spm,[status(thm)],[c_0_83, c_0_12])).
cnf(c_0_289,plain,(mult(rd(rd(X1,op_f),op_e),op_f)=rd(X1,op_e)), inference(spm,[status(thm)],[c_0_238, c_0_8])).
cnf(c_0_290,plain,(rd(mult(mult(X1,X2),ld(X2,X3)),X2)=mult(X1,rd(X3,X2))), inference(spm,[status(thm)],[c_0_4, c_0_143])).
cnf(c_0_291,plain,(rd(X1,rd(op_d,ld(X1,op_f)))=mult(op_d,op_e)), inference(spm,[status(thm)],[c_0_270, c_0_17])).
cnf(c_0_292,plain,(rd(op_d,ld(op_e,unit))=mult(op_d,op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_271]), c_0_32]), c_0_142]), c_0_272])).
cnf(c_0_293,plain,(mult(rd(op_d,op_f),op_e)=ld(op_d,unit)), inference(spm,[status(thm)],[c_0_8, c_0_273])).
cnf(c_0_294,plain,(ld(op_d,mult(op_f,X1))=mult(op_f,ld(op_d,X1))), inference(spm,[status(thm)],[c_0_0, c_0_274])).
cnf(c_0_295,plain,(ld(op_d,op_e)=op_d), inference(spm,[status(thm)],[c_0_0, c_0_11])).
cnf(c_0_296,plain,(mult(rd(X1,mult(rd(X2,X1),X2)),X2)=mult(rd(X1,X2),X1)), inference(spm,[status(thm)],[c_0_8, c_0_275])).
cnf(c_0_297,plain,(rd(mult(X1,X2),ld(mult(X3,X2),X2))=mult(rd(X1,ld(X3,unit)),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_25]), c_0_4])).
cnf(c_0_298,plain,(rd(mult(op_d,op_c),X1)=mult(op_d,rd(op_c,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_276]), c_0_80]), c_0_202])).
cnf(c_0_299,plain,(mult(mult(X1,op_d),ld(op_d,op_c))=mult(X1,op_c)), inference(spm,[status(thm)],[c_0_5, c_0_277])).
cnf(c_0_300,plain,(rd(rd(X1,X2),mult(X2,rd(X1,X2)))=rd(X1,mult(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_8]), c_0_279])).
cnf(c_0_301,plain,(ld(op_f,rd(mult(op_e,X1),op_e))=rd(ld(op_e,X1),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_30]), c_0_107]), c_0_56]), c_0_58]), c_0_281]), c_0_282])).
cnf(c_0_302,plain,(ld(mult(op_e,op_f),mult(op_e,X1))=mult(ld(op_f,unit),X1)), inference(spm,[status(thm)],[c_0_283, c_0_0])).
cnf(c_0_303,plain,(mult(ld(op_e,unit),rd(op_f,X1))=mult(op_d,rd(op_d,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_284]), c_0_34]), c_0_163])).
cnf(c_0_304,plain,(ld(mult(X1,rd(X1,X2)),unit)=rd(X2,mult(X1,X1))), inference(spm,[status(thm)],[c_0_285, c_0_35])).
cnf(c_0_305,plain,(ld(rd(op_f,X1),ld(op_e,unit))=rd(X1,mult(op_e,op_f))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_286]), c_0_231]), c_0_287]), c_0_4]), c_0_288])).
cnf(c_0_306,plain,(rd(rd(X1,op_f),op_e)=rd(rd(X1,op_e),op_f)), inference(spm,[status(thm)],[c_0_4, c_0_289])).
cnf(c_0_307,plain,(rd(mult(mult(X1,X2),X3),X2)=mult(X1,rd(mult(X2,X3),X2))), inference(spm,[status(thm)],[c_0_290, c_0_0])).
cnf(c_0_308,plain,(ld(op_e,mult(op_d,op_e))=op_d), inference(spm,[status(thm)],[c_0_35, c_0_131])).
cnf(c_0_309,plain,(ld(ld(op_f,X1),unit)=mult(ld(X1,op_e),op_e)), inference(spm,[status(thm)],[c_0_151, c_0_12])).
cnf(c_0_310,plain,(rd(op_e,op_d)=op_d), inference(spm,[status(thm)],[c_0_4, c_0_11])).
cnf(c_0_311,plain,(rd(op_d,ld(X1,op_f))=ld(mult(op_d,op_e),X1)), inference(spm,[status(thm)],[c_0_35, c_0_291])).
cnf(c_0_312,plain,(mult(mult(op_d,op_e),ld(op_e,unit))=op_d), inference(spm,[status(thm)],[c_0_8, c_0_292])).
cnf(c_0_313,plain,(mult(mult(rd(op_d,op_f),op_d),op_e)=unit), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_293]), c_0_16])).
cnf(c_0_314,plain,(mult(op_f,ld(op_d,ld(op_e,unit)))=op_d), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_166]), c_0_295])).
cnf(c_0_315,plain,(rd(X1,mult(X1,ld(X2,unit)))=rd(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_98, c_0_114])).
cnf(c_0_316,plain,(rd(X1,ld(X2,ld(X3,mult(X2,X1))))=ld(X2,mult(X3,X2))), inference(spm,[status(thm)],[c_0_99, c_0_17])).
cnf(c_0_317,plain,(ld(mult(rd(X1,X2),X1),mult(X1,X2))=ld(X1,mult(X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_296]), c_0_183]), c_0_0])).
cnf(c_0_318,plain,(mult(rd(op_d,ld(X1,unit)),op_c)=mult(op_d,mult(X1,op_c))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_298]), c_0_17])).
cnf(c_0_319,plain,(rd(mult(op_d,mult(X1,op_c)),op_c)=ld(ld(op_d,unit),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_299]), c_0_49])).
cnf(c_0_320,plain,(ld(rd(X1,mult(X2,X1)),rd(X1,X2))=mult(X2,rd(X1,X2))), inference(spm,[status(thm)],[c_0_35, c_0_300])).
cnf(c_0_321,plain,(rd(op_e,mult(op_d,ld(X1,op_d)))=rd(mult(op_d,X1),op_d)), inference(spm,[status(thm)],[c_0_174, c_0_11])).
cnf(c_0_322,plain,(mult(op_e,mult(ld(op_f,unit),X1))=ld(op_e,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_301]), c_0_8]), c_0_287]), c_0_302])).
cnf(c_0_323,plain,(mult(rd(X1,op_e),ld(op_f,unit))=rd(X1,mult(op_e,op_f))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_303]), c_0_304]), c_0_11]), c_0_193]), c_0_58]), c_0_194]), c_0_12]), c_0_305])).
cnf(c_0_324,plain,(mult(rd(rd(X1,op_e),op_f),op_e)=rd(X1,op_f)), inference(spm,[status(thm)],[c_0_8, c_0_306])).
cnf(c_0_325,plain,(ld(mult(rd(rd(X1,X2),X3),X2),X1)=ld(X2,mult(X3,X2))), inference(spm,[status(thm)],[c_0_76, c_0_8])).
cnf(c_0_326,plain,(mult(rd(X1,X2),rd(mult(X2,X3),X2))=rd(mult(X1,X3),X2)), inference(spm,[status(thm)],[c_0_307, c_0_8])).
cnf(c_0_327,plain,(rd(mult(X1,ld(op_d,unit)),op_e)=mult(rd(X1,op_f),op_d)), inference(spm,[status(thm)],[c_0_4, c_0_264])).
cnf(c_0_328,plain,(mult(ld(op_e,unit),op_d)=ld(op_d,unit)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_308]), c_0_130])).
cnf(c_0_329,plain,(mult(ld(ld(X1,op_f),op_e),op_e)=ld(op_f,mult(X1,op_f))), inference(spm,[status(thm)],[c_0_52, c_0_309])).
cnf(c_0_330,plain,(rd(mult(op_d,ld(X1,op_e)),op_d)=mult(rd(op_d,X1),op_d)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_310]), c_0_120])).
cnf(c_0_331,plain,(ld(ld(op_d,ld(X1,op_f)),unit)=ld(op_f,mult(X1,op_d))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_311]), c_0_91]), c_0_156]), c_0_11]), c_0_12])).
cnf(c_0_332,plain,(ld(ld(op_e,unit),op_d)=mult(op_d,op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_312]), c_0_142]), c_0_49])).
cnf(c_0_333,plain,(mult(rd(op_d,op_f),op_d)=ld(op_e,unit)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_313]), c_0_58])).
cnf(c_0_334,plain,(ld(ld(X1,X2),ld(X1,rd(X2,X1)))=ld(X1,unit)), inference(spm,[status(thm)],[c_0_0, c_0_193])).
cnf(c_0_335,plain,(rd(op_f,mult(op_d,op_e))=op_d), inference(spm,[status(thm)],[c_0_4, c_0_168])).
cnf(c_0_336,plain,(ld(mult(op_d,op_e),op_f)=op_d), inference(spm,[status(thm)],[c_0_35, c_0_219])).
cnf(c_0_337,plain,(ld(mult(op_d,op_e),op_d)=ld(op_e,unit)), inference(spm,[status(thm)],[c_0_35, c_0_292])).
cnf(c_0_338,plain,(ld(op_d,ld(op_e,unit))=ld(op_f,op_d)), inference(spm,[status(thm)],[c_0_0, c_0_314])).
cnf(c_0_339,plain,(ld(mult(X1,mult(X1,ld(X2,unit))),X1)=mult(X2,ld(X1,unit))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_315]), c_0_59])).
cnf(c_0_340,plain,(ld(ld(X1,mult(X2,X1)),X3)=ld(X1,ld(X2,mult(X1,X3)))), inference(spm,[status(thm)],[c_0_35, c_0_316])).
cnf(c_0_341,plain,(mult(op_e,rd(op_e,mult(X1,op_f)))=ld(X1,unit)), inference(spm,[status(thm)],[c_0_5, c_0_212])).
cnf(c_0_342,plain,(rd(mult(X1,X2),ld(X1,mult(X2,X2)))=mult(rd(X1,X2),X1)), inference(spm,[status(thm)],[c_0_17, c_0_317])).
cnf(c_0_343,plain,(rd(op_d,ld(X1,unit))=ld(ld(op_d,unit),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_318]), c_0_319])).
cnf(c_0_344,plain,(ld(rd(X1,mult(ld(X2,X1),X1)),X2)=mult(ld(X2,X1),X2)), inference(spm,[status(thm)],[c_0_320, c_0_17])).
cnf(c_0_345,plain,(ld(ld(op_d,unit),ld(X1,unit))=rd(op_e,mult(X1,op_d))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_18]), c_0_5]), c_0_49])).
cnf(c_0_346,plain,(mult(op_e,mult(op_f,ld(op_e,X1)))=mult(op_f,X1)), inference(spm,[status(thm)],[c_0_26, c_0_180])).
cnf(c_0_347,plain,(mult(op_e,rd(ld(op_f,X1),op_f))=ld(op_e,rd(X1,op_f))), inference(spm,[status(thm)],[c_0_322, c_0_27])).
cnf(c_0_348,plain,(ld(op_f,mult(mult(op_e,X1),op_e))=ld(op_e,mult(X1,op_e))), inference(spm,[status(thm)],[c_0_15, c_0_12])).
cnf(c_0_349,plain,(rd(mult(X1,op_e),mult(op_e,op_f))=mult(X1,ld(op_f,unit))), inference(spm,[status(thm)],[c_0_323, c_0_4])).
cnf(c_0_350,plain,(rd(mult(X1,op_e),op_f)=mult(rd(X1,op_f),op_e)), inference(spm,[status(thm)],[c_0_324, c_0_4])).
cnf(c_0_351,plain,(mult(rd(mult(op_e,X1),op_f),op_e)=mult(op_e,rd(X1,op_e))), inference(spm,[status(thm)],[c_0_80, c_0_12])).
cnf(c_0_352,plain,(ld(mult(X1,ld(X2,X1)),X2)=ld(X1,mult(X2,ld(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_83]), c_0_170]), c_0_6]), c_0_45]), c_0_8])).
cnf(c_0_353,plain,(ld(op_f,mult(op_e,op_f))=op_e), inference(spm,[status(thm)],[c_0_35, c_0_180])).
cnf(c_0_354,plain,(mult(rd(X1,op_e),ld(op_d,unit))=mult(rd(X1,op_f),op_d)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_327]), c_0_288]), c_0_328]), c_0_327])).
cnf(c_0_355,plain,(rd(ld(op_f,mult(X1,op_f)),op_e)=ld(ld(X1,op_f),op_e)), inference(spm,[status(thm)],[c_0_4, c_0_329])).
cnf(c_0_356,plain,(mult(ld(X1,op_e),ld(op_d,unit))=ld(ld(op_d,X1),unit)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_330]), c_0_46])).
cnf(c_0_357,plain,(ld(op_d,ld(X1,op_f))=mult(ld(mult(X1,op_d),op_e),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_331]), c_0_19]), c_0_309]), c_0_126]), c_0_19])).
cnf(c_0_358,plain,(ld(op_f,mult(op_d,op_f))=op_d), inference(spm,[status(thm)],[c_0_35, c_0_218])).
cnf(c_0_359,plain,(rd(op_d,mult(op_d,op_e))=ld(op_e,unit)), inference(spm,[status(thm)],[c_0_17, c_0_332])).
cnf(c_0_360,plain,(rd(ld(op_e,unit),op_d)=rd(op_d,op_f)), inference(spm,[status(thm)],[c_0_4, c_0_333])).
cnf(c_0_361,plain,(ld(mult(op_d,op_e),unit)=ld(op_f,op_d)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_335]), c_0_336]), c_0_337]), c_0_338])).
cnf(c_0_362,plain,(mult(op_f,mult(op_d,X1))=mult(op_d,mult(op_f,X1))), inference(spm,[status(thm)],[c_0_274, c_0_0])).
cnf(c_0_363,plain,(ld(X1,ld(rd(mult(X2,X2),X1),mult(X2,X1)))=ld(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_216]), c_0_7]), c_0_216]), c_0_246]), c_0_340]), c_0_7]), c_0_5])).
cnf(c_0_364,plain,(ld(rd(X1,op_f),unit)=mult(op_e,rd(op_e,X1))), inference(spm,[status(thm)],[c_0_341, c_0_8])).
cnf(c_0_365,plain,(mult(rd(X1,mult(X2,X1)),ld(X2,X1))=ld(X2,ld(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_17]), c_0_113])).
cnf(c_0_366,plain,(rd(mult(X1,op_d),ld(X1,op_e))=mult(rd(X1,op_d),X1)), inference(spm,[status(thm)],[c_0_342, c_0_11])).
cnf(c_0_367,plain,(rd(op_e,mult(X1,op_d))=rd(op_d,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_344]), c_0_19]), c_0_7]), c_0_19]), c_0_83]), c_0_345])).
cnf(c_0_368,plain,(rd(rd(X1,rd(ld(X1,unit),X2)),X1)=mult(X1,X2)), inference(spm,[status(thm)],[c_0_122, c_0_35])).
cnf(c_0_369,plain,(ld(op_e,mult(op_f,X1))=mult(op_f,ld(op_e,X1))), inference(spm,[status(thm)],[c_0_0, c_0_346])).
cnf(c_0_370,plain,(mult(mult(X1,ld(op_f,unit)),op_e)=rd(X1,op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_348]), c_0_231]), c_0_349]), c_0_350]), c_0_351]), c_0_0])).
cnf(c_0_371,plain,(ld(mult(X1,X2),rd(X1,X2))=ld(mult(mult(X1,X2),X2),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_0]), c_0_204])).
cnf(c_0_372,plain,(mult(ld(op_f,unit),op_e)=ld(op_e,unit)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_353]), c_0_288])).
cnf(c_0_373,plain,(mult(mult(ld(op_f,unit),X1),op_d)=ld(op_f,mult(X1,op_d))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_355]), c_0_356]), c_0_357]), c_0_129]), c_0_12]), c_0_23])).
cnf(c_0_374,plain,(mult(ld(op_f,unit),op_d)=rd(op_d,op_f)), inference(spm,[status(thm)],[c_0_23, c_0_358])).
cnf(c_0_375,plain,(rd(op_d,op_f)=ld(op_f,op_d)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_359]), c_0_156]), c_0_11]), c_0_12]), c_0_335]), c_0_360]), c_0_361])).
cnf(c_0_376,plain,(ld(op_f,mult(op_d,mult(op_f,X1)))=mult(op_d,X1)), inference(spm,[status(thm)],[c_0_0, c_0_362])).
cnf(c_0_377,plain,(mult(mult(X1,op_f),op_d)=mult(mult(X1,op_d),op_f)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_218]), c_0_0])).
cnf(c_0_378,plain,(ld(ld(rd(X1,X2),X3),X2)=mult(X2,ld(ld(X1,mult(X3,X2)),unit))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_170]), c_0_114])).
cnf(c_0_379,plain,(ld(rd(mult(X1,X1),X2),mult(X1,X2))=mult(X2,ld(X1,X2))), inference(spm,[status(thm)],[c_0_5, c_0_363])).
cnf(c_0_380,plain,(mult(op_e,rd(op_e,rd(op_f,X1)))=rd(mult(op_f,X1),op_f)), inference(spm,[status(thm)],[c_0_40, c_0_364])).
cnf(c_0_381,plain,(mult(ld(mult(op_f,X1),op_e),op_e)=ld(X1,unit)), inference(spm,[status(thm)],[c_0_309, c_0_0])).
cnf(c_0_382,plain,(rd(op_f,ld(op_e,X1))=mult(mult(op_f,ld(X1,unit)),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_12]), c_0_165]), c_0_12])).
cnf(c_0_383,plain,(rd(ld(X1,ld(X1,X2)),X2)=rd(ld(X1,X2),mult(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_113]), c_0_365]), c_0_17])).
cnf(c_0_384,plain,(rd(ld(X1,X2),mult(X1,X2))=rd(X2,mult(X1,mult(X1,X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_133]), c_0_4])).
cnf(c_0_385,plain,(rd(op_d,ld(X1,ld(X1,op_e)))=ld(op_d,mult(X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_366]), c_0_183]), c_0_0])).
cnf(c_0_386,plain,(mult(op_e,ld(X1,unit))=rd(op_e,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_8]), c_0_165]), c_0_11])).
cnf(c_0_387,plain,(rd(X1,ld(X2,mult(X3,X2)))=mult(rd(rd(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_17, c_0_325])).
cnf(c_0_388,plain,(rd(X1,rd(ld(X1,unit),X2))=mult(mult(X1,X2),X1)), inference(spm,[status(thm)],[c_0_8, c_0_368])).
cnf(c_0_389,plain,(rd(ld(X1,unit),op_d)=ld(mult(op_e,X1),op_d)), inference(spm,[status(thm)],[c_0_4, c_0_177])).
cnf(c_0_390,plain,(mult(op_f,ld(op_e,ld(op_f,X1)))=ld(op_e,X1)), inference(spm,[status(thm)],[c_0_369, c_0_5])).
cnf(c_0_391,plain,(ld(mult(mult(X1,op_e),op_e),X1)=ld(op_f,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_370]), c_0_371]), c_0_372]), c_0_194]), c_0_12])).
cnf(c_0_392,plain,(ld(mult(mult(X1,rd(X1,X2)),X1),X1)=rd(ld(X1,X2),X1)), inference(spm,[status(thm)],[c_0_69, c_0_253])).
cnf(c_0_393,plain,(mult(ld(op_f,op_d),ld(op_d,X1))=ld(op_f,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_373]), c_0_8]), c_0_374]), c_0_375])).
cnf(c_0_394,plain,(ld(op_f,mult(op_d,X1))=mult(op_d,ld(op_f,X1))), inference(spm,[status(thm)],[c_0_376, c_0_5])).
cnf(c_0_395,plain,(rd(mult(mult(X1,op_d),op_f),op_d)=mult(X1,op_f)), inference(spm,[status(thm)],[c_0_4, c_0_377])).
cnf(c_0_396,plain,(mult(X1,mult(ld(mult(mult(X2,X1),X1),X2),X2))=ld(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_379]), c_0_87]), c_0_151])).
cnf(c_0_397,plain,(mult(op_e,ld(op_f,unit))=ld(op_e,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_287]), c_0_4]), c_0_153])).
cnf(c_0_398,plain,(ld(mult(X1,ld(op_e,unit)),unit)=ld(op_e,rd(op_f,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_380]), c_0_26]), c_0_153]), c_0_381]), c_0_35])).
cnf(c_0_399,plain,(rd(op_f,ld(X1,ld(op_e,unit)))=mult(mult(op_e,op_f),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_55]), c_0_40]), c_0_143]), c_0_287]), c_0_0])).
cnf(c_0_400,plain,(rd(ld(X1,ld(X1,X2)),X2)=rd(X2,mult(X1,mult(X1,X2)))), inference(rw,[status(thm)],[c_0_383, c_0_384])).
cnf(c_0_401,plain,(ld(X1,ld(X1,op_e))=mult(ld(mult(X1,X1),unit),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_385]), c_0_105]), c_0_11])).
cnf(c_0_402,plain,(mult(ld(mult(X1,X1),unit),X1)=ld(X1,unit)), inference(spm,[status(thm)],[c_0_8, c_0_93])).
cnf(c_0_403,plain,(mult(rd(rd(op_e,X1),X2),X1)=mult(op_e,ld(X1,ld(X2,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_340]), c_0_7]), c_0_387])).
cnf(c_0_404,plain,(mult(mult(op_e,ld(mult(op_d,X1),unit)),op_d)=rd(op_e,X1)), inference(spm,[status(thm)],[c_0_178, c_0_0])).
cnf(c_0_405,plain,(mult(rd(op_e,mult(op_d,X1)),op_d)=mult(op_d,rd(op_d,X1))), inference(spm,[status(thm)],[c_0_80, c_0_11])).
cnf(c_0_406,plain,(rd(X1,ld(mult(op_e,X1),op_d))=mult(mult(X1,op_d),X1)), inference(spm,[status(thm)],[c_0_388, c_0_389])).
cnf(c_0_407,plain,(mult(op_e,rd(ld(op_f,X1),op_e))=rd(ld(op_e,X1),op_e)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_390]), c_0_230])).
cnf(c_0_408,plain,(ld(op_f,mult(mult(op_f,X1),op_f))=mult(mult(X1,op_e),op_e)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_391]), c_0_182])).
cnf(c_0_409,plain,(rd(ld(op_e,mult(X1,op_f)),op_e)=ld(op_e,mult(X1,op_e))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_341]), c_0_18])).
cnf(c_0_410,plain,(mult(ld(op_f,op_d),X1)=mult(op_d,ld(op_f,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_0]), c_0_394])).
cnf(c_0_411,plain,(rd(mult(X1,op_f),op_d)=mult(rd(X1,op_d),op_f)), inference(spm,[status(thm)],[c_0_395, c_0_8])).
cnf(c_0_412,plain,(rd(X1,mult(X2,mult(X2,X1)))=ld(mult(mult(X1,X2),X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_396]), c_0_384]), c_0_121]), c_0_371])).
cnf(c_0_413,plain,(rd(X1,mult(op_e,mult(op_e,X1)))=ld(op_f,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_322]), c_0_383]), c_0_384]), c_0_397]), c_0_93]), c_0_12])).
cnf(c_0_414,plain,(ld(ld(op_f,unit),X1)=ld(op_e,mult(mult(op_e,op_f),X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_398]), c_0_399]), c_0_193]), c_0_58]), c_0_194]), c_0_12])).
cnf(c_0_415,plain,(rd(op_e,mult(X1,mult(X1,op_e)))=ld(mult(X1,X1),unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_401]), c_0_265]), c_0_4]), c_0_402]), c_0_194]), c_0_5]), c_0_384])).
cnf(c_0_416,plain,(rd(X1,rd(X2,ld(X1,X3)))=mult(ld(rd(mult(X1,X2),X3),unit),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_89]), c_0_25])).
cnf(c_0_417,plain,(mult(rd(op_e,X1),rd(X1,mult(X2,X1)))=rd(rd(op_e,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_403]), c_0_56]), c_0_113])).
cnf(c_0_418,plain,(mult(op_d,rd(op_d,X1))=rd(op_e,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_404, c_0_386]), c_0_405])).
cnf(c_0_419,plain,(ld(rd(op_d,X1),op_e)=mult(X1,op_d)), inference(spm,[status(thm)],[c_0_35, c_0_367])).
cnf(c_0_420,plain,(rd(X1,ld(mult(mult(X2,X1),X2),X2))=mult(mult(X1,X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_233]), c_0_32])).
cnf(c_0_421,plain,(ld(mult(mult(X1,op_d),X1),X1)=ld(mult(op_e,X1),op_d)), inference(spm,[status(thm)],[c_0_35, c_0_406])).
cnf(c_0_422,plain,(ld(op_e,mult(mult(op_f,X1),op_e))=mult(op_e,mult(X1,op_e))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_407, c_0_408]), c_0_4]), c_0_409])).
cnf(c_0_423,plain,(mult(ld(ld(X1,op_e),op_d),op_d)=ld(op_e,mult(X1,op_e))), inference(spm,[status(thm)],[c_0_52, c_0_164])).
cnf(c_0_424,plain,(rd(mult(op_d,ld(op_f,X1)),X1)=ld(op_f,op_d)), inference(spm,[status(thm)],[c_0_4, c_0_410])).
cnf(c_0_425,plain,(mult(rd(rd(X1,op_f),op_d),op_f)=rd(X1,op_d)), inference(spm,[status(thm)],[c_0_411, c_0_8])).
cnf(c_0_426,plain,(rd(ld(X1,ld(X1,X2)),X2)=ld(mult(mult(X2,X1),X1),X2)), inference(rw,[status(thm)],[c_0_400, c_0_412])).
cnf(c_0_427,plain,(ld(op_e,mult(mult(op_e,op_f),X1))=mult(op_e,mult(op_e,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_413]), c_0_414])).
cnf(c_0_428,plain,(ld(op_e,ld(op_e,X1))=mult(ld(op_f,unit),X1)), inference(spm,[status(thm)],[c_0_0, c_0_322])).
cnf(c_0_429,plain,(ld(ld(mult(X1,X1),unit),op_e)=mult(X1,mult(X1,op_e))), inference(spm,[status(thm)],[c_0_35, c_0_415])).
cnf(c_0_430,plain,(rd(mult(X1,ld(X2,mult(X3,X1))),X1)=ld(rd(X2,X1),X3)), inference(spm,[status(thm)],[c_0_4, c_0_170])).
cnf(c_0_431,plain,(rd(mult(X1,X1),rd(X2,ld(X1,unit)))=mult(X1,ld(X2,unit))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_82]), c_0_307]), c_0_206]), c_0_120]), c_0_114])).
cnf(c_0_432,plain,(ld(mult(rd(op_e,X1),X2),rd(op_e,X1))=rd(X1,mult(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_417]), c_0_97])).
cnf(c_0_433,plain,(ld(rd(op_e,X1),unit)=rd(X1,op_e)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_418]), c_0_11])).
cnf(c_0_434,plain,(ld(mult(op_e,X1),op_e)=ld(X1,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_420]), c_0_11]), c_0_421]), c_0_177])).
cnf(c_0_435,plain,(ld(X1,ld(op_e,unit))=mult(ld(mult(op_f,X1),unit),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_422]), c_0_96]), c_0_25])).
cnf(c_0_436,plain,(mult(ld(X1,unit),op_e)=ld(X1,op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_423]), c_0_33]), c_0_25]), c_0_17])).
cnf(c_0_437,plain,(ld(rd(X1,op_e),unit)=rd(op_e,X1)), inference(rw,[status(thm)],[c_0_163, c_0_418])).
cnf(c_0_438,plain,(mult(rd(X1,X2),mult(X2,X3))=mult(X2,mult(ld(X2,X1),X3))), inference(spm,[status(thm)],[c_0_179, c_0_0])).
cnf(c_0_439,plain,(rd(mult(op_d,X1),mult(op_f,X1))=ld(op_f,op_d)), inference(spm,[status(thm)],[c_0_424, c_0_0])).
cnf(c_0_440,plain,(mult(mult(mult(op_d,op_e),X1),ld(X1,op_d))=op_f), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_255]), c_0_8])).
cnf(c_0_441,plain,(rd(rd(X1,op_f),op_d)=rd(rd(X1,op_d),op_f)), inference(spm,[status(thm)],[c_0_4, c_0_425])).
cnf(c_0_442,plain,(rd(mult(op_e,X1),mult(mult(op_e,op_f),X1))=ld(op_f,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_427]), c_0_0]), c_0_391])).
cnf(c_0_443,plain,(mult(ld(op_f,unit),mult(op_e,X1))=ld(op_e,X1)), inference(spm,[status(thm)],[c_0_428, c_0_0])).
cnf(c_0_444,plain,(ld(mult(ld(X1,unit),X2),X1)=mult(mult(ld(X2,unit),X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_23]), c_0_33]), c_0_25])).
cnf(c_0_445,plain,(mult(ld(mult(X1,X1),unit),mult(X1,mult(X1,op_e)))=op_e), inference(spm,[status(thm)],[c_0_5, c_0_429])).
cnf(c_0_446,plain,(ld(X1,ld(rd(X2,X1),X3))=mult(ld(X2,mult(X3,X1)),ld(X1,unit))), inference(spm,[status(thm)],[c_0_59, c_0_430])).
cnf(c_0_447,plain,(mult(X1,ld(mult(X2,ld(X1,unit)),unit))=rd(mult(X1,X1),X2)), inference(spm,[status(thm)],[c_0_431, c_0_4])).
cnf(c_0_448,plain,(rd(X1,mult(X2,ld(X1,unit)))=mult(X1,mult(X1,ld(X2,unit)))), inference(spm,[status(thm)],[c_0_17, c_0_339])).
cnf(c_0_449,plain,(rd(X1,mult(op_e,ld(mult(X2,X1),X1)))=mult(X2,rd(X1,op_e))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_432]), c_0_45]), c_0_91]), c_0_6]), c_0_433])).
cnf(c_0_450,plain,(mult(X1,ld(op_e,unit))=rd(X1,op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_434]), c_0_435]), c_0_436]), c_0_269]), c_0_59])).
cnf(c_0_451,plain,(rd(X1,mult(X1,rd(X2,op_e)))=rd(mult(X1,op_e),mult(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_437]), c_0_75])).
cnf(c_0_452,plain,(mult(X1,mult(ld(X1,mult(X2,X1)),X3))=mult(X2,mult(X1,X3))), inference(spm,[status(thm)],[c_0_438, c_0_4])).
cnf(c_0_453,plain,(ld(mult(op_d,op_e),ld(X1,unit))=rd(ld(op_f,op_d),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_439]), c_0_49]), c_0_311])).
cnf(c_0_454,plain,(ld(ld(X1,unit),mult(op_d,op_e))=mult(rd(X1,op_d),op_f)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_440]), c_0_411]), c_0_49])).
cnf(c_0_455,plain,(mult(ld(mult(mult(X1,X1),X2),X1),X1)=ld(X2,unit)), inference(spm,[status(thm)],[c_0_151, c_0_0])).
cnf(c_0_456,plain,(mult(rd(rd(X1,op_d),op_f),op_d)=rd(X1,op_f)), inference(spm,[status(thm)],[c_0_8, c_0_441])).
cnf(c_0_457,plain,(ld(mult(op_e,X1),ld(op_e,X1))=ld(X1,ld(op_f,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_442]), c_0_443]), c_0_15]), c_0_340]), c_0_7])).
cnf(c_0_458,plain,(ld(rd(X1,mult(X2,X1)),ld(X2,X1))=X1), inference(spm,[status(thm)],[c_0_35, c_0_113])).
cnf(c_0_459,plain,(rd(mult(X1,X2),mult(X2,mult(X1,X2)))=rd(X1,mult(X2,X1))), inference(spm,[status(thm)],[c_0_300, c_0_4])).
cnf(c_0_460,plain,(mult(ld(mult(X1,op_e),X1),mult(X1,X1))=ld(op_e,mult(X1,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_445]), c_0_206])).
cnf(c_0_461,plain,(ld(mult(mult(X1,X2),ld(X2,X1)),X1)=mult(ld(X1,X2),ld(X2,unit))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_68]), c_0_5]), c_0_246]), c_0_446]), c_0_6])).
cnf(c_0_462,plain,(mult(op_e,rd(op_e,X1))=rd(op_f,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_447, c_0_386]), c_0_448]), c_0_386]), c_0_12])).
cnf(c_0_463,plain,(ld(mult(X1,rd(X2,op_e)),X2)=mult(op_e,ld(mult(X1,X2),X2))), inference(spm,[status(thm)],[c_0_35, c_0_449])).
cnf(c_0_464,plain,(mult(rd(X1,X2),rd(X2,op_e))=mult(rd(X1,mult(X2,op_e)),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_438, c_0_450]), c_0_450]), c_0_231])).
cnf(c_0_465,plain,(rd(X1,mult(rd(X2,mult(X1,op_e)),X1))=rd(mult(X1,op_e),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_231]), c_0_5])).
cnf(c_0_466,plain,(mult(rd(mult(X1,X2),mult(X2,op_e)),X2)=mult(X1,rd(X2,op_e))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_452, c_0_450]), c_0_231]), c_0_450])).
cnf(c_0_467,plain,(ld(X1,ld(X2,mult(X1,op_e)))=mult(ld(X1,ld(X2,X1)),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_436, c_0_340]), c_0_7]), c_0_340])).
cnf(c_0_468,plain,(rd(X1,ld(op_f,op_d))=mult(rd(X1,op_d),op_f)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_453]), c_0_27]), c_0_86]), c_0_5]), c_0_454])).
cnf(c_0_469,plain,(rd(ld(X1,unit),X2)=ld(mult(mult(X2,X2),X1),X2)), inference(spm,[status(thm)],[c_0_4, c_0_455])).
cnf(c_0_470,plain,(mult(ld(mult(op_d,op_e),X1),ld(X1,op_f))=op_d), inference(spm,[status(thm)],[c_0_8, c_0_311])).
cnf(c_0_471,plain,(rd(mult(X1,op_d),op_f)=mult(rd(X1,op_f),op_d)), inference(spm,[status(thm)],[c_0_456, c_0_4])).
cnf(c_0_472,plain,(mult(op_d,mult(op_e,ld(op_d,X1)))=mult(op_e,X1)), inference(spm,[status(thm)],[c_0_26, c_0_131])).
cnf(c_0_473,plain,(ld(op_e,rd(X1,mult(op_e,X1)))=rd(X1,mult(op_f,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_457]), c_0_5]), c_0_113]), c_0_113])).
cnf(c_0_474,plain,(mult(ld(mult(X1,X2),X1),rd(X2,X1))=ld(mult(X2,X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_204]), c_0_187])).
cnf(c_0_475,plain,(rd(rd(X1,X2),ld(mult(X1,X2),X1))=X1), inference(spm,[status(thm)],[c_0_17, c_0_97])).
cnf(c_0_476,plain,(ld(rd(X1,X2),rd(mult(X1,X2),X1))=rd(mult(X2,X2),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_317]), c_0_5])).
cnf(c_0_477,plain,(mult(ld(X1,op_d),op_d)=ld(X1,op_e)), inference(spm,[status(thm)],[c_0_419, c_0_17])).
cnf(c_0_478,plain,(ld(rd(X1,mult(X2,X1)),ld(X2,mult(X1,X2)))=mult(X1,X2)), inference(spm,[status(thm)],[c_0_458, c_0_459])).
cnf(c_0_479,plain,(rd(mult(X1,X1),mult(op_e,mult(X1,X1)))=ld(mult(X1,op_e),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_460]), c_0_113])).
cnf(c_0_480,plain,(rd(X1,mult(op_e,X1))=ld(op_e,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_462]), c_0_35]), c_0_8]), c_0_153]), c_0_97]), c_0_434]), c_0_125])).
cnf(c_0_481,plain,(ld(mult(rd(X1,mult(X2,op_e)),X2),X2)=mult(op_e,ld(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_463, c_0_464]), c_0_8])).
cnf(c_0_482,plain,(rd(X1,mult(X2,rd(X1,op_e)))=rd(mult(X1,op_e),mult(X2,X1))), inference(spm,[status(thm)],[c_0_465, c_0_466])).
cnf(c_0_483,plain,(mult(X1,mult(ld(X1,ld(X2,X1)),op_e))=ld(X2,mult(X1,op_e))), inference(spm,[status(thm)],[c_0_5, c_0_467])).
cnf(c_0_484,plain,(ld(mult(rd(X1,op_d),op_f),X1)=ld(op_f,op_d)), inference(spm,[status(thm)],[c_0_35, c_0_468])).
cnf(c_0_485,plain,(mult(ld(mult(op_e,X1),op_d),op_f)=ld(X1,mult(op_d,op_e))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_454, c_0_35]), c_0_58]), c_0_469]), c_0_11])).
cnf(c_0_486,plain,(ld(rd(mult(op_d,op_e),X1),unit)=mult(rd(X1,op_f),op_d)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_470]), c_0_471]), c_0_34])).
cnf(c_0_487,plain,(ld(op_d,mult(op_e,X1))=mult(op_e,ld(op_d,X1))), inference(spm,[status(thm)],[c_0_0, c_0_472])).
cnf(c_0_488,plain,(ld(op_e,mult(mult(op_d,X1),op_d))=ld(op_d,mult(X1,op_d))), inference(spm,[status(thm)],[c_0_15, c_0_11])).
cnf(c_0_489,plain,(mult(op_e,rd(X1,mult(op_f,X1)))=rd(X1,mult(op_e,X1))), inference(spm,[status(thm)],[c_0_5, c_0_473])).
cnf(c_0_490,plain,(ld(ld(mult(X1,X2),X1),rd(X2,X1))=mult(X2,rd(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_474]), c_0_475])).
cnf(c_0_491,plain,(rd(op_e,ld(X1,op_d))=mult(op_d,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_477]), c_0_113]), c_0_265]), c_0_310]), c_0_478]), c_0_11])).
cnf(c_0_492,plain,(ld(mult(X1,op_e),X1)=ld(op_e,unit)), inference(rw,[status(thm)],[c_0_479, c_0_480])).
cnf(c_0_493,plain,(mult(op_e,mult(op_d,X1))=mult(op_d,mult(op_e,X1))), inference(spm,[status(thm)],[c_0_472, c_0_0])).
cnf(c_0_494,plain,(ld(X1,mult(mult(X1,X1),X2))=mult(mult(X1,X1),ld(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_61]), c_0_0])).
cnf(c_0_495,plain,(mult(ld(op_e,unit),X1)=ld(op_e,X1)), inference(spm,[status(thm)],[c_0_133, c_0_480])).
cnf(c_0_496,plain,(ld(mult(X1,rd(X2,op_e)),X1)=mult(op_e,ld(mult(X1,X2),X1))), inference(spm,[status(thm)],[c_0_481, c_0_80])).
cnf(c_0_497,plain,(rd(ld(X1,X2),rd(X2,op_e))=rd(ld(X1,mult(X2,op_e)),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_482]), c_0_251])).
cnf(c_0_498,plain,(ld(rd(X1,X2),mult(X1,op_e))=mult(X1,mult(ld(X1,X2),op_e))), inference(spm,[status(thm)],[c_0_483, c_0_35])).
cnf(c_0_499,plain,(mult(rd(X1,op_f),op_d)=mult(X1,ld(op_f,op_d))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_484]), c_0_469]), c_0_11]), c_0_485]), c_0_33]), c_0_34]), c_0_486])).
cnf(c_0_500,plain,(mult(op_e,ld(op_d,ld(op_e,X1)))=ld(op_d,X1)), inference(spm,[status(thm)],[c_0_487, c_0_5])).
cnf(c_0_501,plain,(ld(op_e,mult(ld(op_d,unit),X1))=mult(op_d,ld(op_f,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_488]), c_0_23]), c_0_4]), c_0_394])).
cnf(c_0_502,plain,(ld(X1,mult(ld(X1,unit),X2))=mult(ld(X1,unit),ld(X1,X2))), inference(spm,[status(thm)],[c_0_195, c_0_58])).
cnf(c_0_503,plain,(mult(rd(X1,mult(op_f,X1)),op_e)=rd(X1,mult(op_e,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_489]), c_0_5]), c_0_113]), c_0_346])).
cnf(c_0_504,plain,(ld(ld(X1,op_e),unit)=ld(op_e,X1)), inference(spm,[status(thm)],[c_0_87, c_0_434])).
cnf(c_0_505,plain,(mult(mult(op_d,op_e),X1)=mult(op_d,mult(op_e,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_490, c_0_491]), c_0_492]), c_0_215]), c_0_236]), c_0_156]), c_0_11]), c_0_12]), c_0_237]), c_0_493])).
cnf(c_0_506,plain,(mult(X1,mult(mult(X1,X1),ld(X1,X2)))=mult(mult(X1,X1),X2)), inference(spm,[status(thm)],[c_0_5, c_0_494])).
cnf(c_0_507,plain,(rd(mult(X1,ld(op_e,X2)),mult(X1,X2))=rd(X1,mult(X1,op_e))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_495]), c_0_127])).
cnf(c_0_508,plain,(mult(X1,mult(op_e,ld(mult(X1,X2),X1)))=rd(X1,rd(X2,op_e))), inference(spm,[status(thm)],[c_0_204, c_0_496])).
cnf(c_0_509,plain,(rd(mult(X1,mult(ld(X1,X2),op_e)),X1)=rd(X2,rd(X1,op_e))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_497, c_0_35]), c_0_498])).
cnf(c_0_510,plain,(ld(mult(op_e,rd(op_e,X1)),unit)=rd(X1,op_f)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_364]), c_0_19]), c_0_126]), c_0_19])).
cnf(c_0_511,plain,(rd(mult(X1,ld(op_d,unit)),op_e)=mult(X1,ld(op_f,op_d))), inference(rw,[status(thm)],[c_0_327, c_0_499])).
cnf(c_0_512,plain,(mult(ld(X1,ld(X2,unit)),X2)=mult(X2,ld(mult(X1,X2),unit))), inference(spm,[status(thm)],[c_0_114, c_0_96])).
cnf(c_0_513,plain,(mult(ld(op_d,unit),ld(op_d,X1))=mult(op_e,ld(op_f,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_501]), c_0_0]), c_0_502])).
cnf(c_0_514,plain,(rd(X1,mult(X1,rd(op_e,X2)))=rd(mult(X1,X2),mult(X1,op_e))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_437]), c_0_7]), c_0_75])).
cnf(c_0_515,plain,(rd(rd(X1,mult(op_e,X1)),op_e)=rd(X1,mult(op_f,X1))), inference(spm,[status(thm)],[c_0_4, c_0_503])).
cnf(c_0_516,plain,(mult(ld(X1,unit),mult(X1,X2))=mult(X1,mult(ld(X1,unit),X2))), inference(spm,[status(thm)],[c_0_438, c_0_58])).
cnf(c_0_517,plain,(mult(rd(X1,op_e),X1)=mult(X1,ld(op_e,X1))), inference(spm,[status(thm)],[c_0_41, c_0_504])).
cnf(c_0_518,plain,(ld(ld(op_e,unit),X1)=mult(op_e,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_215, c_0_505]), c_0_0])).
cnf(c_0_519,plain,(mult(op_d,mult(ld(op_d,X1),op_d))=mult(rd(X1,op_d),op_e)), inference(spm,[status(thm)],[c_0_179, c_0_295])).
cnf(c_0_520,plain,(mult(mult(X1,ld(X2,unit)),X2)=mult(mult(X1,X2),ld(X2,unit))), inference(spm,[status(thm)],[c_0_2, c_0_16])).
cnf(c_0_521,plain,(mult(mult(X1,X1),mult(X2,X1))=mult(X1,mult(mult(X1,X2),X1))), inference(spm,[status(thm)],[c_0_506, c_0_2])).
cnf(c_0_522,plain,(rd(rd(X1,X2),rd(X1,rd(X2,op_e)))=rd(X1,mult(X1,op_e))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_507, c_0_508]), c_0_0]), c_0_204])).
cnf(c_0_523,plain,(mult(rd(X1,rd(X2,op_e)),X2)=mult(X2,mult(ld(X2,X1),op_e))), inference(spm,[status(thm)],[c_0_8, c_0_509])).
cnf(c_0_524,plain,(rd(mult(X1,op_d),op_f)=mult(X1,ld(op_f,op_d))), inference(rw,[status(thm)],[c_0_471, c_0_499])).
cnf(c_0_525,plain,(rd(ld(X1,op_e),op_f)=ld(mult(op_e,X1),unit)), inference(spm,[status(thm)],[c_0_510, c_0_17])).
cnf(c_0_526,plain,(mult(ld(X1,op_d),ld(op_f,op_d))=mult(ld(mult(op_f,X1),unit),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_511, c_0_512]), c_0_191]), c_0_513]), c_0_213]), c_0_121]), c_0_58]), c_0_32]), c_0_435])).
cnf(c_0_527,plain,(rd(X1,mult(rd(mult(X1,op_e),X2),X1))=rd(X2,mult(X1,op_e))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_514, c_0_89]), c_0_5])).
cnf(c_0_528,plain,(mult(rd(X1,mult(X2,X1)),mult(X2,mult(X1,X2)))=mult(X1,X2)), inference(spm,[status(thm)],[c_0_8, c_0_459])).
cnf(c_0_529,plain,(rd(X1,mult(op_f,X1))=ld(op_f,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_515, c_0_480]), c_0_93]), c_0_12])).
cnf(c_0_530,plain,(mult(X1,mult(ld(X1,unit),mult(X2,X1)))=rd(X2,ld(X1,unit))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_32]), c_0_54]), c_0_7]), c_0_516])).
cnf(c_0_531,plain,(mult(ld(op_e,X1),ld(X1,X2))=mult(X1,ld(mult(op_e,X1),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_495]), c_0_110]), c_0_495])).
cnf(c_0_532,plain,(ld(X1,ld(rd(X1,op_e),mult(X1,X2)))=ld(ld(op_e,X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_517]), c_0_0])).
cnf(c_0_533,plain,(mult(X1,ld(ld(op_e,X1),X2))=mult(mult(op_e,X1),ld(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_518]), c_0_143]), c_0_495])).
cnf(c_0_534,plain,(mult(rd(op_d,rd(op_e,X1)),op_e)=ld(ld(op_d,unit),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_519, c_0_59]), c_0_520]), c_0_47]), c_0_49]), c_0_138]), c_0_11])).
cnf(c_0_535,plain,(rd(op_d,rd(op_e,mult(op_d,X1)))=mult(op_e,rd(X1,op_e))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_405]), c_0_11]), c_0_304]), c_0_11])).
cnf(c_0_536,plain,(ld(ld(X1,unit),mult(X1,X2))=mult(X1,ld(ld(X1,unit),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_521]), c_0_49]), c_0_49])).
cnf(c_0_537,plain,(ld(mult(rd(X1,mult(X1,op_e)),X2),X1)=mult(ld(X2,X1),op_e)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_522]), c_0_523]), c_0_0])).
cnf(c_0_538,plain,(mult(rd(X1,X2),rd(X2,mult(X2,X3)))=rd(mult(X1,ld(X3,unit)),X2)), inference(spm,[status(thm)],[c_0_326, c_0_127])).
cnf(c_0_539,plain,(ld(rd(mult(X1,X2),mult(X1,X3)),X1)=mult(X1,ld(rd(X2,X3),unit))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_75]), c_0_114])).
cnf(c_0_540,plain,(rd(ld(X1,unit),op_e)=ld(mult(op_f,X1),op_e)), inference(spm,[status(thm)],[c_0_4, c_0_381])).
cnf(c_0_541,plain,(ld(mult(op_f,X1),op_e)=ld(mult(op_e,X1),unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_524, c_0_477]), c_0_525]), c_0_526]), c_0_436])).
cnf(c_0_542,plain,(ld(rd(mult(X1,op_e),X2),X1)=mult(rd(X2,mult(X1,op_e)),X1)), inference(spm,[status(thm)],[c_0_133, c_0_527])).
cnf(c_0_543,plain,(mult(mult(X1,op_e),op_e)=mult(X1,op_f)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_528, c_0_529]), c_0_516]), c_0_530]), c_0_182]), c_0_408])).
cnf(c_0_544,plain,(ld(X1,mult(rd(X1,op_e),X2))=mult(X1,ld(mult(op_e,X1),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_517]), c_0_0]), c_0_531])).
cnf(c_0_545,plain,(ld(rd(X1,op_e),mult(X1,X2))=mult(mult(op_e,X1),ld(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_532]), c_0_533])).
cnf(c_0_546,plain,(mult(op_d,ld(ld(op_d,unit),X1))=mult(op_f,ld(op_e,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_535]), c_0_143]), c_0_12]), c_0_536])).
cnf(c_0_547,plain,(rd(op_e,ld(X1,unit))=mult(op_e,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_233]), c_0_421]), c_0_17])).
cnf(c_0_548,plain,(mult(mult(X1,op_f),ld(op_e,X2))=mult(X1,mult(op_e,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_537, c_0_538]), c_0_539]), c_0_540]), c_0_541]), c_0_32]), c_0_542]), c_0_26]), c_0_492]), c_0_450]), c_0_143]), c_0_543])).
cnf(c_0_549,plain,(ld(X1,mult(mult(op_e,X1),ld(X1,X2)))=ld(ld(op_e,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_544]), c_0_25]), c_0_495]), c_0_545])).
cnf(c_0_550,plain,(ld(X1,X1)=unit), inference(spm,[status(thm)],[c_0_0, c_0_7])).
cnf(c_0_551,plain,(mult(op_f,ld(op_e,X1))=mult(op_e,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_343]), c_0_546]), c_0_547])).
cnf(c_0_552,plain,(mult(op_e,mult(op_e,X1))=mult(op_f,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_529]), c_0_414]), c_0_427])).
cnf(c_0_553,plain,(mult(mult(X1,op_f),X2)=mult(X1,mult(op_f,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_549]), c_0_550]), c_0_19]), c_0_12]), c_0_551]), c_0_552])).
cnf(c_0_554,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_553])]), ['proof']).
# SZS output end CNFRefutation.
EOF
