% Problem    : HWV069-1 : TPTP v6.1.0. Bugfixed v5.5.1.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n110.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:44:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____011_C18_F1_PI_AE_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(esk8318_4(X3,X1,X4,X2)|~p(X1)|~p(X2)|~esk2_0|~p(X4)|~p(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_23159)).
cnf(c_0_1,axiom,(esk2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_25550)).
cnf(c_0_2,axiom,(esk6073_4(X1,X2,X3,X4)|~esk6291_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15196)).
cnf(c_0_3,axiom,(~esk6073_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_19547)).
cnf(c_0_4,axiom,(esk8284_4(X1,X2,X3,X4)|~esk8318_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_18921)).
cnf(c_0_5,axiom,(esk6291_4(X1,X2,X3,X4)|~esk6767_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_11826)).
cnf(c_0_6,axiom,(esk6184_4(X1,X2,X3,X4)|~esk8284_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_19690)).
cnf(c_0_7,axiom,(esk6766_4(X1,X2,X3,X4)|esk6767_4(X1,X2,X3,X4)|~esk6768_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_8688)).
cnf(c_0_8,axiom,(esk6768_4(X1,X2,X3,X4)|~esk6184_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1764)).
cnf(c_0_9,axiom,(esk6735_4(X1,X2,X3,X4)|~esk6766_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_22868)).
cnf(c_0_10,axiom,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|~esk6735_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_16916)).
cnf(c_0_11,axiom,(esk6646_4(X1,X2,X3,X4)|~esk6734_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_14829)).
cnf(c_0_12,axiom,(esk6073_4(X1,X2,X3,X4)|~esk6639_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15864)).
cnf(c_0_13,axiom,(esk6073_4(X1,X2,X3,X4)|~esk6640_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_12215)).
cnf(c_0_14,axiom,(esk6641_4(X1,X2,X3,X4)|~esk6646_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_9395)).
cnf(c_0_15,axiom,(esk6638_4(X1,X2,X3,X4)|esk6639_4(X1,X2,X3,X4)|esk6640_4(X1,X2,X3,X4)|~esk6641_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_18547)).
cnf(c_0_16,axiom,(esk6185_4(X1,X2,X3,X4)|~esk6733_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_23466)).
cnf(c_0_17,axiom,(esk3974_4(X1,X2,X3,X4)|~esk4192_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_21421)).
cnf(c_0_18,axiom,(~esk3974_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_27224)).
cnf(c_0_19,axiom,(esk6185_4(X1,X2,X3,X4)|~esk6638_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_10711)).
cnf(c_0_20,axiom,(esk4192_4(X1,X2,X3,X4)|~esk4668_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_17214)).
cnf(c_0_21,axiom,(esk4085_4(X1,X2,X3,X4)|~esk6185_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_6045)).
cnf(c_0_22,axiom,(esk4667_4(X1,X2,X3,X4)|esk4668_4(X1,X2,X3,X4)|~esk4669_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15178)).
cnf(c_0_23,axiom,(esk4669_4(X1,X2,X3,X4)|~esk4085_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_20864)).
cnf(c_0_24,axiom,(esk4636_4(X1,X2,X3,X4)|~esk4667_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_18477)).
cnf(c_0_25,axiom,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|~esk4636_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_23515)).
cnf(c_0_26,axiom,(esk4547_4(X1,X2,X3,X4)|~esk4635_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_16619)).
cnf(c_0_27,axiom,(esk3974_4(X1,X2,X3,X4)|~esk4540_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_26066)).
cnf(c_0_28,axiom,(esk3974_4(X1,X2,X3,X4)|~esk4541_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_11583)).
cnf(c_0_29,axiom,(esk4542_4(X1,X2,X3,X4)|~esk4547_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_9637)).
cnf(c_0_30,axiom,(esk4540_4(X1,X2,X3,X4)|esk4541_4(X1,X2,X3,X4)|esk4539_4(X1,X2,X3,X4)|~esk4542_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_11296)).
cnf(c_0_31,axiom,(esk4086_4(X1,X2,X3,X4)|~esk4634_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_26197)).
cnf(c_0_32,axiom,(esk1875_4(X1,X2,X3,X4)|~esk2093_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15739)).
cnf(c_0_33,axiom,(~esk1875_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_20170)).
cnf(c_0_34,axiom,(esk4086_4(X1,X2,X3,X4)|~esk4539_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_12501)).
cnf(c_0_35,axiom,(esk2093_4(X1,X2,X3,X4)|~esk2569_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_24728)).
cnf(c_0_36,axiom,(esk1986_4(X1,X2,X3,X4)|~esk4086_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_23528)).
cnf(c_0_37,axiom,(esk2568_4(X1,X2,X3,X4)|esk2569_4(X1,X2,X3,X4)|~esk2570_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_21252)).
cnf(c_0_38,axiom,(esk2570_4(X1,X2,X3,X4)|~esk1986_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_8638)).
cnf(c_0_39,axiom,(esk2537_4(X1,X2,X3,X4)|~esk2568_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_23309)).
cnf(c_0_40,axiom,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|~esk2537_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15727)).
cnf(c_0_41,axiom,(esk2448_4(X1,X2,X3,X4)|~esk2536_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15876)).
cnf(c_0_42,axiom,(esk1875_4(X1,X2,X3,X4)|~esk2442_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_273)).
cnf(c_0_43,axiom,(esk1875_4(X1,X2,X3,X4)|~esk2441_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_24050)).
cnf(c_0_44,axiom,(esk2443_4(X1,X2,X3,X4)|~esk2448_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_546)).
cnf(c_0_45,axiom,(esk2440_4(X1,X2,X3,X4)|esk2441_4(X1,X2,X3,X4)|esk2442_4(X1,X2,X3,X4)|~esk2443_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_26462)).
cnf(c_0_46,axiom,(esk1987_4(X1,X2,X3,X4)|~esk2535_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_9693)).
cnf(c_0_47,axiom,(esk1987_4(X1,X2,X3,X4)|~esk2440_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_24283)).
cnf(c_0_48,axiom,(esk377_0|~p(X1)|~esk1987_4(X2,X1,X3,X4)|~p(X2)|~p(X3)|~p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_5425)).
cnf(c_0_49,axiom,(p(true)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule_true)).
cnf(c_0_50,axiom,(esk8172_4(X1,X2,X3,X4)|~esk8390_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_26573)).
cnf(c_0_51,axiom,(~esk8172_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_21066)).
cnf(c_0_52,axiom,(esk8390_4(X1,X2,X3,X4)|~esk8866_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_19272)).
cnf(c_0_53,axiom,(esk8283_4(X1,X2,X3,X4)|p(X4)|~esk377_0|~p(X2)|~p(X1)|~p(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_24756)).
cnf(c_0_54,axiom,(esk8865_4(X1,X2,X3,X4)|esk8866_4(X1,X2,X3,X4)|~esk8867_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_24196)).
cnf(c_0_55,axiom,(esk8867_4(X1,X2,X3,X4)|~esk8283_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_11106)).
cnf(c_0_56,axiom,(esk8834_4(X1,X2,X3,X4)|~esk8865_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_3087)).
cnf(c_0_57,axiom,(esk8832_4(X1,X2,X3,X4)|esk8833_4(X1,X2,X3,X4)|~esk8834_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_27521)).
cnf(c_0_58,axiom,(esk8284_4(X1,X2,X3,X4)|~esk8832_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_23096)).
cnf(c_0_59,axiom,(esk8172_4(X1,X2,X3,X4)|~esk8738_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_19519)).
cnf(c_0_60,axiom,(esk8172_4(X1,X2,X3,X4)|~esk8739_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_20836)).
cnf(c_0_61,axiom,(esk8745_4(X1,X2,X3,X4)|~esk8833_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_9016)).
cnf(c_0_62,axiom,(esk8739_4(X1,X2,X3,X4)|esk8738_4(X1,X2,X3,X4)|esk8737_4(X1,X2,X3,X4)|~esk8740_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15537)).
cnf(c_0_63,axiom,(esk8740_4(X1,X2,X3,X4)|~esk8745_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_11834)).
cnf(c_0_64,axiom,(esk8284_4(X1,X2,X3,X4)|~esk8737_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_22524)).
cnf(c_0_65,axiom,(esk376_0|p(X4)|~esk1987_4(X1,X2,X3,X4)|~p(X1)|~p(X3)|~p(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_16563)).
cnf(c_0_66,axiom,(~p(false)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule_false)).
cnf(c_0_67,axiom,(p(X4)|esk8283_4(X2,X1,X4,X3)|~p(X1)|~p(X2)|~p(X3)|~esk376_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_2605)).
cnf(c_0_68,axiom,(esk375_0|p(X3)|~p(X1)|~p(X2)|~esk1987_4(X2,X1,X3,X4)|~p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_2786)).
cnf(c_0_69,axiom,(p(X3)|p(X4)|esk8283_4(X1,X2,X3,X4)|~p(X1)|~esk375_0|~p(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_19195)).
cnf(c_0_70,axiom,(p(X2)|p(X3)|esk374_0|~p(X1)|~p(X4)|~esk1987_4(X4,X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1335)).
cnf(c_0_71,axiom,(p(X1)|esk8283_4(X2,X1,X3,X4)|~esk374_0|~p(X2)|~p(X3)|~p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_12562)).
cnf(c_0_72,axiom,(esk373_0|p(X4)|~p(X1)|~p(X2)|~esk1987_4(X3,X4,X1,X2)|~p(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_24350)).
cnf(c_0_73,axiom,(p(X2)|p(X3)|esk8283_4(X1,X2,X4,X3)|~p(X1)|~esk373_0|~p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_14240)).
cnf(c_0_74,axiom,(esk372_0|p(X1)|p(X4)|~p(X2)|~p(X3)|~esk1987_4(X3,X1,X2,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_3798)).
cnf(c_0_75,axiom,(p(X3)|esk8283_4(X2,X4,X3,X1)|p(X4)|~p(X1)|~p(X2)|~esk372_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_6515)).
cnf(c_0_76,axiom,(p(X1)|esk371_0|p(X4)|~p(X2)|~p(X3)|~esk1987_4(X3,X4,X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15771)).
cnf(c_0_77,axiom,(p(X1)|p(X2)|p(X4)|esk8283_4(X3,X1,X4,X2)|~p(X3)|~esk371_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_13911)).
cnf(c_0_78,axiom,(p(X1)|p(X3)|p(X4)|esk370_0|~p(X2)|~esk1987_4(X2,X4,X1,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_6614)).
cnf(c_0_79,axiom,(p(X1)|esk8283_4(X1,X2,X3,X4)|~esk370_0|~p(X3)|~p(X2)|~p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_14554)).
cnf(c_0_80,axiom,(p(X2)|esk369_0|~p(X1)|~p(X3)|~p(X4)|~esk1987_4(X2,X4,X1,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1861)).
cnf(c_0_81,axiom,(esk8283_4(X1,X2,X3,X4)|p(X1)|p(X4)|~esk369_0|~p(X2)|~p(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_23911)).
cnf(c_0_82,axiom,(p(X1)|esk368_0|p(X2)|~esk1987_4(X2,X3,X4,X1)|~p(X3)|~p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_13042)).
cnf(c_0_83,axiom,(esk8283_4(X2,X1,X3,X4)|p(X3)|p(X2)|~p(X1)|~esk368_0|~p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_20186)).
cnf(c_0_84,axiom,(p(X1)|p(X3)|esk367_0|~esk1987_4(X1,X2,X3,X4)|~p(X4)|~p(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_8043)).
cnf(c_0_85,axiom,(p(X1)|p(X2)|esk8283_4(X3,X4,X2,X1)|p(X3)|~esk367_0|~p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_8445)).
cnf(c_0_86,axiom,(p(X3)|p(X1)|p(X4)|esk366_0|~esk1987_4(X1,X2,X3,X4)|~p(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_19128)).
cnf(c_0_87,axiom,(p(X1)|esk8283_4(X3,X1,X4,X2)|p(X3)|~esk366_0|~p(X2)|~p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_17302)).
cnf(c_0_88,axiom,(p(X3)|esk365_0|p(X4)|~p(X1)|~p(X2)|~esk1987_4(X3,X4,X2,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_9318)).
cnf(c_0_89,axiom,(p(X2)|p(X3)|p(X4)|esk8283_4(X4,X3,X1,X2)|~p(X1)|~esk365_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_2940)).
cnf(c_0_90,axiom,(p(X1)|esk364_0|p(X2)|p(X4)|~esk1987_4(X1,X2,X3,X4)|~p(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_3839)).
cnf(c_0_91,axiom,(p(X1)|p(X2)|p(X4)|esk8283_4(X1,X2,X4,X3)|~p(X3)|~esk364_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_17378)).
cnf(c_0_92,axiom,(p(X2)|p(X3)|esk363_0|p(X4)|~p(X1)|~esk1987_4(X3,X4,X2,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_17980)).
cnf(c_0_93,axiom,(esk8283_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|p(X3)|~esk363_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1996)).
cnf(c_0_94,axiom,(p(X1)|p(X2)|p(X3)|esk1_0|p(X4)|~esk2874_4(X2,X1,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15456)).
cnf(c_0_95,axiom,(~esk1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_15881)).
cnf(c_0_96,axiom,(esk2874_4(X1,X2,X3,X4)|~esk1682_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_7219)).
cnf(c_0_97,axiom,(esk2446_4(X1,X2,X3,X4)|esk1682_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_10200)).
cnf(c_0_98,axiom,(esk2567_4(X1,X2,X3,X4)|~esk2568_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_14154)).
cnf(c_0_99,axiom,(~esk2446_4(X1,X2,X3,X4)|~esk1682_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_2106)).
cnf(c_0_100,axiom,(esk1878_4(X1,X2,X3,X4)|esk1682_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_27437)).
cnf(c_0_101,axiom,(~esk1878_4(X1,X2,X3,X4)|~esk2567_4(X1,X2,X3,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_4600)).
cnf(c_0_102,plain,(esk8318_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1])])).
cnf(c_0_103,plain,(~esk6291_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_2, c_0_3])).
cnf(c_0_104,plain,(esk8284_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_4, c_0_102])).
cnf(c_0_105,plain,(~esk6767_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_5, c_0_103])).
cnf(c_0_106,plain,(esk6184_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_6, c_0_104])).
cnf(c_0_107,plain,(esk6766_4(X1,X2,X3,X4)|~esk6768_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_7, c_0_105])).
cnf(c_0_108,plain,(esk6768_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_8, c_0_106])).
cnf(c_0_109,plain,(esk6766_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_110,plain,(esk6735_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_9, c_0_109])).
cnf(c_0_111,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_10, c_0_110])).
cnf(c_0_112,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_11, c_0_111])).
cnf(c_0_113,plain,(~esk6639_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_12, c_0_3])).
cnf(c_0_114,plain,(~esk6640_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_13, c_0_3])).
cnf(c_0_115,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_14, c_0_112])).
cnf(c_0_116,plain,(esk6638_4(X1,X2,X3,X4)|~esk6641_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_15, c_0_113]), c_0_114])).
cnf(c_0_117,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_16, c_0_115])).
cnf(c_0_118,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_119,plain,(~esk4192_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_120,plain,(esk6185_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_19, c_0_118])).
cnf(c_0_121,plain,(~esk4668_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_20, c_0_119])).
cnf(c_0_122,plain,(esk4085_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_21, c_0_120])).
cnf(c_0_123,plain,(esk4667_4(X1,X2,X3,X4)|~esk4669_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_22, c_0_121])).
cnf(c_0_124,plain,(esk4669_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_23, c_0_122])).
cnf(c_0_125,plain,(esk4667_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_126,plain,(esk4636_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_24, c_0_125])).
cnf(c_0_127,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_25, c_0_126])).
cnf(c_0_128,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_26, c_0_127])).
cnf(c_0_129,plain,(~esk4540_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_27, c_0_18])).
cnf(c_0_130,plain,(~esk4541_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_28, c_0_18])).
cnf(c_0_131,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_29, c_0_128])).
cnf(c_0_132,plain,(esk4539_4(X1,X2,X3,X4)|~esk4542_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_30, c_0_129]), c_0_130])).
cnf(c_0_133,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_31, c_0_131])).
cnf(c_0_134,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_135,plain,(~esk2093_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_136,plain,(esk4086_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_34, c_0_134])).
cnf(c_0_137,plain,(~esk2569_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_35, c_0_135])).
cnf(c_0_138,plain,(esk1986_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_36, c_0_136])).
cnf(c_0_139,plain,(esk2568_4(X1,X2,X3,X4)|~esk2570_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_37, c_0_137])).
cnf(c_0_140,plain,(esk2570_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_38, c_0_138])).
cnf(c_0_141,plain,(esk2568_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_142,plain,(esk2537_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_39, c_0_141])).
cnf(c_0_143,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_40, c_0_142])).
cnf(c_0_144,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_41, c_0_143])).
cnf(c_0_145,plain,(~esk2442_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_42, c_0_33])).
cnf(c_0_146,plain,(~esk2441_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_43, c_0_33])).
cnf(c_0_147,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_44, c_0_144])).
cnf(c_0_148,plain,(esk2440_4(X1,X2,X3,X4)|~esk2443_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_45, c_0_145]), c_0_146])).
cnf(c_0_149,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_46, c_0_147])).
cnf(c_0_150,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_148, c_0_149])).
cnf(c_0_151,plain,(esk1987_4(X1,X2,X3,X4)|~p(X3)|~p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_47, c_0_150])).
cnf(c_0_152,plain,(esk377_0|~p(X1)|~p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_48, c_0_151])).
cnf(c_0_153,plain,(esk377_0|~p(X1)|~p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_152, c_0_49])).
cnf(c_0_154,plain,(esk377_0|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_153, c_0_49])).
cnf(c_0_155,plain,(esk377_0|~p(X1)), inference(spm,[status(thm)],[c_0_154, c_0_49])).
cnf(c_0_156,plain,(~esk8390_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_157,plain,(esk377_0), inference(spm,[status(thm)],[c_0_155, c_0_49])).
cnf(c_0_158,plain,(~esk8866_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_52, c_0_156])).
cnf(c_0_159,plain,(esk8283_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_157])])).
cnf(c_0_160,plain,(esk8865_4(X1,X2,X3,X4)|~esk8867_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_54, c_0_158])).
cnf(c_0_161,plain,(esk8867_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_55, c_0_159])).
cnf(c_0_162,plain,(esk8865_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_163,plain,(esk8834_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_56, c_0_162])).
cnf(c_0_164,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_57, c_0_163])).
cnf(c_0_165,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_58, c_0_164])).
cnf(c_0_166,plain,(~esk8738_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_59, c_0_51])).
cnf(c_0_167,plain,(~esk8739_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[c_0_60, c_0_51])).
cnf(c_0_168,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_61, c_0_165])).
cnf(c_0_169,plain,(esk8737_4(X1,X2,X3,X4)|~esk8740_4(X1,X2,X3,X4)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_62, c_0_166]), c_0_167])).
cnf(c_0_170,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_63, c_0_168])).
cnf(c_0_171,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_169, c_0_170])).
cnf(c_0_172,plain,(esk8284_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_64, c_0_171])).
cnf(c_0_173,plain,(esk6184_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_6, c_0_172])).
cnf(c_0_174,plain,(esk6768_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_8, c_0_173])).
cnf(c_0_175,plain,(esk6766_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_107, c_0_174])).
cnf(c_0_176,plain,(esk6735_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_9, c_0_175])).
cnf(c_0_177,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_10, c_0_176])).
cnf(c_0_178,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_11, c_0_177])).
cnf(c_0_179,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_14, c_0_178])).
cnf(c_0_180,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_16, c_0_179])).
cnf(c_0_181,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_116, c_0_180])).
cnf(c_0_182,plain,(esk6185_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_19, c_0_181])).
cnf(c_0_183,plain,(esk4085_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_21, c_0_182])).
cnf(c_0_184,plain,(esk4669_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_23, c_0_183])).
cnf(c_0_185,plain,(esk4667_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_123, c_0_184])).
cnf(c_0_186,plain,(esk4636_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_24, c_0_185])).
cnf(c_0_187,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_25, c_0_186])).
cnf(c_0_188,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_26, c_0_187])).
cnf(c_0_189,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_29, c_0_188])).
cnf(c_0_190,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_31, c_0_189])).
cnf(c_0_191,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_132, c_0_190])).
cnf(c_0_192,plain,(esk4086_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_34, c_0_191])).
cnf(c_0_193,plain,(esk1986_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_36, c_0_192])).
cnf(c_0_194,plain,(esk2570_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_38, c_0_193])).
cnf(c_0_195,plain,(esk2568_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_139, c_0_194])).
cnf(c_0_196,plain,(esk2537_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_39, c_0_195])).
cnf(c_0_197,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_40, c_0_196])).
cnf(c_0_198,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_41, c_0_197])).
cnf(c_0_199,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_44, c_0_198])).
cnf(c_0_200,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_46, c_0_199])).
cnf(c_0_201,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_148, c_0_200])).
cnf(c_0_202,plain,(esk1987_4(X1,X2,X3,X4)|p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_47, c_0_201])).
cnf(c_0_203,plain,(esk376_0|p(X1)|~p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_65, c_0_202])).
cnf(c_0_204,plain,(esk376_0|p(X1)|~p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_203, c_0_49])).
cnf(c_0_205,plain,(esk376_0|p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_204, c_0_49])).
cnf(c_0_206,plain,(esk376_0|p(X1)), inference(spm,[status(thm)],[c_0_205, c_0_49])).
cnf(c_0_207,plain,(esk376_0), inference(spm,[status(thm)],[c_0_66, c_0_206])).
cnf(c_0_208,plain,(esk8283_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_207])])).
cnf(c_0_209,plain,(esk8867_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_55, c_0_208])).
cnf(c_0_210,plain,(esk8865_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_160, c_0_209])).
cnf(c_0_211,plain,(esk8834_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_56, c_0_210])).
cnf(c_0_212,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_57, c_0_211])).
cnf(c_0_213,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_58, c_0_212])).
cnf(c_0_214,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_61, c_0_213])).
cnf(c_0_215,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_63, c_0_214])).
cnf(c_0_216,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_169, c_0_215])).
cnf(c_0_217,plain,(esk8284_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_64, c_0_216])).
cnf(c_0_218,plain,(esk6184_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_6, c_0_217])).
cnf(c_0_219,plain,(esk6768_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_8, c_0_218])).
cnf(c_0_220,plain,(esk6766_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_107, c_0_219])).
cnf(c_0_221,plain,(esk6735_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_9, c_0_220])).
cnf(c_0_222,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_10, c_0_221])).
cnf(c_0_223,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_11, c_0_222])).
cnf(c_0_224,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_14, c_0_223])).
cnf(c_0_225,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_16, c_0_224])).
cnf(c_0_226,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_116, c_0_225])).
cnf(c_0_227,plain,(esk6185_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_19, c_0_226])).
cnf(c_0_228,plain,(esk4085_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_21, c_0_227])).
cnf(c_0_229,plain,(esk4669_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_23, c_0_228])).
cnf(c_0_230,plain,(esk4667_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_123, c_0_229])).
cnf(c_0_231,plain,(esk4636_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_24, c_0_230])).
cnf(c_0_232,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_25, c_0_231])).
cnf(c_0_233,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_26, c_0_232])).
cnf(c_0_234,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_29, c_0_233])).
cnf(c_0_235,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_31, c_0_234])).
cnf(c_0_236,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_132, c_0_235])).
cnf(c_0_237,plain,(esk4086_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_34, c_0_236])).
cnf(c_0_238,plain,(esk1986_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_36, c_0_237])).
cnf(c_0_239,plain,(esk2570_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_38, c_0_238])).
cnf(c_0_240,plain,(esk2568_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_139, c_0_239])).
cnf(c_0_241,plain,(esk2537_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_39, c_0_240])).
cnf(c_0_242,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_40, c_0_241])).
cnf(c_0_243,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_41, c_0_242])).
cnf(c_0_244,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_44, c_0_243])).
cnf(c_0_245,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_46, c_0_244])).
cnf(c_0_246,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_148, c_0_245])).
cnf(c_0_247,plain,(esk1987_4(X1,X2,X3,X4)|p(X3)|~p(X4)|~p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_47, c_0_246])).
cnf(c_0_248,plain,(esk375_0|p(X1)|~p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_68, c_0_247])).
cnf(c_0_249,plain,(esk375_0|p(X1)|~p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_248, c_0_49])).
cnf(c_0_250,plain,(esk375_0|p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_249, c_0_49])).
cnf(c_0_251,plain,(esk375_0|p(X1)), inference(spm,[status(thm)],[c_0_250, c_0_49])).
cnf(c_0_252,plain,(esk375_0), inference(spm,[status(thm)],[c_0_66, c_0_251])).
cnf(c_0_253,plain,(esk8283_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_252])])).
cnf(c_0_254,plain,(esk8867_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_55, c_0_253])).
cnf(c_0_255,plain,(esk8865_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_160, c_0_254])).
cnf(c_0_256,plain,(esk8834_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_56, c_0_255])).
cnf(c_0_257,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_57, c_0_256])).
cnf(c_0_258,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_58, c_0_257])).
cnf(c_0_259,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_61, c_0_258])).
cnf(c_0_260,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_63, c_0_259])).
cnf(c_0_261,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_169, c_0_260])).
cnf(c_0_262,plain,(esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_64, c_0_261])).
cnf(c_0_263,plain,(esk6184_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_6, c_0_262])).
cnf(c_0_264,plain,(esk6768_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_8, c_0_263])).
cnf(c_0_265,plain,(esk6766_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_107, c_0_264])).
cnf(c_0_266,plain,(esk6735_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_9, c_0_265])).
cnf(c_0_267,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_10, c_0_266])).
cnf(c_0_268,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_11, c_0_267])).
cnf(c_0_269,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_14, c_0_268])).
cnf(c_0_270,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_16, c_0_269])).
cnf(c_0_271,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_116, c_0_270])).
cnf(c_0_272,plain,(esk6185_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_19, c_0_271])).
cnf(c_0_273,plain,(esk4085_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_21, c_0_272])).
cnf(c_0_274,plain,(esk4669_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_23, c_0_273])).
cnf(c_0_275,plain,(esk4667_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_123, c_0_274])).
cnf(c_0_276,plain,(esk4636_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_24, c_0_275])).
cnf(c_0_277,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_25, c_0_276])).
cnf(c_0_278,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_26, c_0_277])).
cnf(c_0_279,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_29, c_0_278])).
cnf(c_0_280,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_31, c_0_279])).
cnf(c_0_281,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_132, c_0_280])).
cnf(c_0_282,plain,(esk4086_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_34, c_0_281])).
cnf(c_0_283,plain,(esk1986_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_36, c_0_282])).
cnf(c_0_284,plain,(esk2570_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_38, c_0_283])).
cnf(c_0_285,plain,(esk2568_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_139, c_0_284])).
cnf(c_0_286,plain,(esk2537_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_39, c_0_285])).
cnf(c_0_287,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_40, c_0_286])).
cnf(c_0_288,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_41, c_0_287])).
cnf(c_0_289,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_44, c_0_288])).
cnf(c_0_290,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_46, c_0_289])).
cnf(c_0_291,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_148, c_0_290])).
cnf(c_0_292,plain,(esk1987_4(X1,X2,X3,X4)|p(X4)|p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_47, c_0_291])).
cnf(c_0_293,plain,(esk374_0|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_70, c_0_292])).
cnf(c_0_294,plain,(esk374_0|p(X1)|p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_293, c_0_49])).
cnf(c_0_295,plain,(esk374_0|p(X1)|p(X2)), inference(spm,[status(thm)],[c_0_294, c_0_49])).
cnf(c_0_296,plain,(esk374_0|p(X1)), inference(ef,[status(thm)],[c_0_295])).
cnf(c_0_297,plain,(esk374_0), inference(spm,[status(thm)],[c_0_66, c_0_296])).
cnf(c_0_298,plain,(esk8283_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_297])])).
cnf(c_0_299,plain,(esk8867_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_55, c_0_298])).
cnf(c_0_300,plain,(esk8865_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_160, c_0_299])).
cnf(c_0_301,plain,(esk8834_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_56, c_0_300])).
cnf(c_0_302,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_57, c_0_301])).
cnf(c_0_303,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_58, c_0_302])).
cnf(c_0_304,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_61, c_0_303])).
cnf(c_0_305,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_63, c_0_304])).
cnf(c_0_306,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_169, c_0_305])).
cnf(c_0_307,plain,(esk8284_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_64, c_0_306])).
cnf(c_0_308,plain,(esk6184_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_6, c_0_307])).
cnf(c_0_309,plain,(esk6768_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_8, c_0_308])).
cnf(c_0_310,plain,(esk6766_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_107, c_0_309])).
cnf(c_0_311,plain,(esk6735_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_9, c_0_310])).
cnf(c_0_312,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_10, c_0_311])).
cnf(c_0_313,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_11, c_0_312])).
cnf(c_0_314,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_14, c_0_313])).
cnf(c_0_315,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_16, c_0_314])).
cnf(c_0_316,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_116, c_0_315])).
cnf(c_0_317,plain,(esk6185_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_19, c_0_316])).
cnf(c_0_318,plain,(esk4085_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_21, c_0_317])).
cnf(c_0_319,plain,(esk4669_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_23, c_0_318])).
cnf(c_0_320,plain,(esk4667_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_123, c_0_319])).
cnf(c_0_321,plain,(esk4636_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_24, c_0_320])).
cnf(c_0_322,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_25, c_0_321])).
cnf(c_0_323,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_26, c_0_322])).
cnf(c_0_324,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_29, c_0_323])).
cnf(c_0_325,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_31, c_0_324])).
cnf(c_0_326,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_132, c_0_325])).
cnf(c_0_327,plain,(esk4086_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_34, c_0_326])).
cnf(c_0_328,plain,(esk1986_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_36, c_0_327])).
cnf(c_0_329,plain,(esk2570_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_38, c_0_328])).
cnf(c_0_330,plain,(esk2568_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_139, c_0_329])).
cnf(c_0_331,plain,(esk2537_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_39, c_0_330])).
cnf(c_0_332,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_40, c_0_331])).
cnf(c_0_333,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_41, c_0_332])).
cnf(c_0_334,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_44, c_0_333])).
cnf(c_0_335,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_46, c_0_334])).
cnf(c_0_336,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_148, c_0_335])).
cnf(c_0_337,plain,(esk1987_4(X1,X2,X3,X4)|p(X2)|~p(X4)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_47, c_0_336])).
cnf(c_0_338,plain,(esk373_0|p(X1)|~p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_72, c_0_337])).
cnf(c_0_339,plain,(esk373_0|p(X1)|~p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_338, c_0_49])).
cnf(c_0_340,plain,(esk373_0|p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_339, c_0_49])).
cnf(c_0_341,plain,(esk373_0|p(X1)), inference(spm,[status(thm)],[c_0_340, c_0_49])).
cnf(c_0_342,plain,(esk373_0), inference(spm,[status(thm)],[c_0_66, c_0_341])).
cnf(c_0_343,plain,(esk8283_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_342])])).
cnf(c_0_344,plain,(esk8867_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_55, c_0_343])).
cnf(c_0_345,plain,(esk8865_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_160, c_0_344])).
cnf(c_0_346,plain,(esk8834_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_56, c_0_345])).
cnf(c_0_347,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_57, c_0_346])).
cnf(c_0_348,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_58, c_0_347])).
cnf(c_0_349,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_61, c_0_348])).
cnf(c_0_350,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_63, c_0_349])).
cnf(c_0_351,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_169, c_0_350])).
cnf(c_0_352,plain,(esk8284_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_64, c_0_351])).
cnf(c_0_353,plain,(esk6184_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_6, c_0_352])).
cnf(c_0_354,plain,(esk6768_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_8, c_0_353])).
cnf(c_0_355,plain,(esk6766_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_107, c_0_354])).
cnf(c_0_356,plain,(esk6735_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_9, c_0_355])).
cnf(c_0_357,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_10, c_0_356])).
cnf(c_0_358,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_11, c_0_357])).
cnf(c_0_359,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_14, c_0_358])).
cnf(c_0_360,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_16, c_0_359])).
cnf(c_0_361,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_116, c_0_360])).
cnf(c_0_362,plain,(esk6185_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_19, c_0_361])).
cnf(c_0_363,plain,(esk4085_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_21, c_0_362])).
cnf(c_0_364,plain,(esk4669_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_23, c_0_363])).
cnf(c_0_365,plain,(esk4667_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_123, c_0_364])).
cnf(c_0_366,plain,(esk4636_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_24, c_0_365])).
cnf(c_0_367,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_25, c_0_366])).
cnf(c_0_368,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_26, c_0_367])).
cnf(c_0_369,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_29, c_0_368])).
cnf(c_0_370,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_31, c_0_369])).
cnf(c_0_371,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_132, c_0_370])).
cnf(c_0_372,plain,(esk4086_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_34, c_0_371])).
cnf(c_0_373,plain,(esk1986_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_36, c_0_372])).
cnf(c_0_374,plain,(esk2570_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_38, c_0_373])).
cnf(c_0_375,plain,(esk2568_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_139, c_0_374])).
cnf(c_0_376,plain,(esk2537_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_39, c_0_375])).
cnf(c_0_377,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_40, c_0_376])).
cnf(c_0_378,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_41, c_0_377])).
cnf(c_0_379,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_44, c_0_378])).
cnf(c_0_380,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_46, c_0_379])).
cnf(c_0_381,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_148, c_0_380])).
cnf(c_0_382,plain,(esk1987_4(X1,X2,X3,X4)|p(X4)|p(X2)|~p(X3)|~p(X1)), inference(spm,[status(thm)],[c_0_47, c_0_381])).
cnf(c_0_383,plain,(esk372_0|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_74, c_0_382])).
cnf(c_0_384,plain,(esk372_0|p(X1)|p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_383, c_0_49])).
cnf(c_0_385,plain,(esk372_0|p(X1)|p(X2)), inference(spm,[status(thm)],[c_0_384, c_0_49])).
cnf(c_0_386,plain,(esk372_0|p(X1)), inference(ef,[status(thm)],[c_0_385])).
cnf(c_0_387,plain,(esk372_0), inference(spm,[status(thm)],[c_0_66, c_0_386])).
cnf(c_0_388,plain,(esk8283_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_387])])).
cnf(c_0_389,plain,(esk8867_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_55, c_0_388])).
cnf(c_0_390,plain,(esk8865_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_160, c_0_389])).
cnf(c_0_391,plain,(esk8834_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_56, c_0_390])).
cnf(c_0_392,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_57, c_0_391])).
cnf(c_0_393,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_58, c_0_392])).
cnf(c_0_394,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_61, c_0_393])).
cnf(c_0_395,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_63, c_0_394])).
cnf(c_0_396,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_169, c_0_395])).
cnf(c_0_397,plain,(esk8284_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_64, c_0_396])).
cnf(c_0_398,plain,(esk6184_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_6, c_0_397])).
cnf(c_0_399,plain,(esk6768_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_8, c_0_398])).
cnf(c_0_400,plain,(esk6766_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_107, c_0_399])).
cnf(c_0_401,plain,(esk6735_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_9, c_0_400])).
cnf(c_0_402,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_10, c_0_401])).
cnf(c_0_403,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_11, c_0_402])).
cnf(c_0_404,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_14, c_0_403])).
cnf(c_0_405,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_16, c_0_404])).
cnf(c_0_406,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_116, c_0_405])).
cnf(c_0_407,plain,(esk6185_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_19, c_0_406])).
cnf(c_0_408,plain,(esk4085_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_21, c_0_407])).
cnf(c_0_409,plain,(esk4669_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_23, c_0_408])).
cnf(c_0_410,plain,(esk4667_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_123, c_0_409])).
cnf(c_0_411,plain,(esk4636_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_24, c_0_410])).
cnf(c_0_412,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_25, c_0_411])).
cnf(c_0_413,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_26, c_0_412])).
cnf(c_0_414,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_29, c_0_413])).
cnf(c_0_415,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_31, c_0_414])).
cnf(c_0_416,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_132, c_0_415])).
cnf(c_0_417,plain,(esk4086_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_34, c_0_416])).
cnf(c_0_418,plain,(esk1986_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_36, c_0_417])).
cnf(c_0_419,plain,(esk2570_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_38, c_0_418])).
cnf(c_0_420,plain,(esk2568_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_139, c_0_419])).
cnf(c_0_421,plain,(esk2537_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_39, c_0_420])).
cnf(c_0_422,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_40, c_0_421])).
cnf(c_0_423,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_41, c_0_422])).
cnf(c_0_424,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_44, c_0_423])).
cnf(c_0_425,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_46, c_0_424])).
cnf(c_0_426,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_148, c_0_425])).
cnf(c_0_427,plain,(esk1987_4(X1,X2,X3,X4)|p(X2)|p(X3)|~p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_47, c_0_426])).
cnf(c_0_428,plain,(esk371_0|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_76, c_0_427])).
cnf(c_0_429,plain,(esk371_0|p(X1)|p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_428, c_0_49])).
cnf(c_0_430,plain,(esk371_0|p(X1)|p(X2)), inference(spm,[status(thm)],[c_0_429, c_0_49])).
cnf(c_0_431,plain,(esk371_0|p(X1)), inference(ef,[status(thm)],[c_0_430])).
cnf(c_0_432,plain,(esk371_0), inference(spm,[status(thm)],[c_0_66, c_0_431])).
cnf(c_0_433,plain,(esk8283_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_432])])).
cnf(c_0_434,plain,(esk8867_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_55, c_0_433])).
cnf(c_0_435,plain,(esk8865_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_160, c_0_434])).
cnf(c_0_436,plain,(esk8834_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_56, c_0_435])).
cnf(c_0_437,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_57, c_0_436])).
cnf(c_0_438,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_58, c_0_437])).
cnf(c_0_439,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_61, c_0_438])).
cnf(c_0_440,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_63, c_0_439])).
cnf(c_0_441,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_169, c_0_440])).
cnf(c_0_442,plain,(esk8284_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_64, c_0_441])).
cnf(c_0_443,plain,(esk6184_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_6, c_0_442])).
cnf(c_0_444,plain,(esk6768_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_8, c_0_443])).
cnf(c_0_445,plain,(esk6766_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_107, c_0_444])).
cnf(c_0_446,plain,(esk6735_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_9, c_0_445])).
cnf(c_0_447,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_10, c_0_446])).
cnf(c_0_448,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_11, c_0_447])).
cnf(c_0_449,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_14, c_0_448])).
cnf(c_0_450,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_16, c_0_449])).
cnf(c_0_451,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_116, c_0_450])).
cnf(c_0_452,plain,(esk6185_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_19, c_0_451])).
cnf(c_0_453,plain,(esk4085_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_21, c_0_452])).
cnf(c_0_454,plain,(esk4669_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_23, c_0_453])).
cnf(c_0_455,plain,(esk4667_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_123, c_0_454])).
cnf(c_0_456,plain,(esk4636_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_24, c_0_455])).
cnf(c_0_457,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_25, c_0_456])).
cnf(c_0_458,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_26, c_0_457])).
cnf(c_0_459,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_29, c_0_458])).
cnf(c_0_460,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_31, c_0_459])).
cnf(c_0_461,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_132, c_0_460])).
cnf(c_0_462,plain,(esk4086_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_34, c_0_461])).
cnf(c_0_463,plain,(esk1986_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_36, c_0_462])).
cnf(c_0_464,plain,(esk2570_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_38, c_0_463])).
cnf(c_0_465,plain,(esk2568_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_139, c_0_464])).
cnf(c_0_466,plain,(esk2537_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_39, c_0_465])).
cnf(c_0_467,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_40, c_0_466])).
cnf(c_0_468,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_41, c_0_467])).
cnf(c_0_469,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_44, c_0_468])).
cnf(c_0_470,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_46, c_0_469])).
cnf(c_0_471,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_148, c_0_470])).
cnf(c_0_472,plain,(esk1987_4(X1,X2,X3,X4)|p(X3)|p(X4)|p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_47, c_0_471])).
cnf(c_0_473,plain,(esk370_0|p(X1)|p(X2)|p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_78, c_0_472])).
cnf(c_0_474,plain,(esk370_0|p(X1)|p(X2)|p(X3)), inference(spm,[status(thm)],[c_0_473, c_0_49])).
cnf(c_0_475,plain,(esk370_0|p(X1)|p(X2)), inference(ef,[status(thm)],[c_0_474])).
cnf(c_0_476,plain,(esk370_0|p(X1)), inference(ef,[status(thm)],[c_0_475])).
cnf(c_0_477,plain,(esk370_0), inference(spm,[status(thm)],[c_0_66, c_0_476])).
cnf(c_0_478,plain,(esk8283_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_477])])).
cnf(c_0_479,plain,(esk8867_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_55, c_0_478])).
cnf(c_0_480,plain,(esk8865_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_160, c_0_479])).
cnf(c_0_481,plain,(esk8834_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_56, c_0_480])).
cnf(c_0_482,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_57, c_0_481])).
cnf(c_0_483,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_58, c_0_482])).
cnf(c_0_484,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_61, c_0_483])).
cnf(c_0_485,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_63, c_0_484])).
cnf(c_0_486,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_169, c_0_485])).
cnf(c_0_487,plain,(esk8284_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_64, c_0_486])).
cnf(c_0_488,plain,(esk6184_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_6, c_0_487])).
cnf(c_0_489,plain,(esk6768_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_8, c_0_488])).
cnf(c_0_490,plain,(esk6766_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_107, c_0_489])).
cnf(c_0_491,plain,(esk6735_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_9, c_0_490])).
cnf(c_0_492,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_10, c_0_491])).
cnf(c_0_493,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_11, c_0_492])).
cnf(c_0_494,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_14, c_0_493])).
cnf(c_0_495,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_16, c_0_494])).
cnf(c_0_496,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_116, c_0_495])).
cnf(c_0_497,plain,(esk6185_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_19, c_0_496])).
cnf(c_0_498,plain,(esk4085_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_21, c_0_497])).
cnf(c_0_499,plain,(esk4669_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_23, c_0_498])).
cnf(c_0_500,plain,(esk4667_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_123, c_0_499])).
cnf(c_0_501,plain,(esk4636_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_24, c_0_500])).
cnf(c_0_502,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_25, c_0_501])).
cnf(c_0_503,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_26, c_0_502])).
cnf(c_0_504,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_29, c_0_503])).
cnf(c_0_505,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_31, c_0_504])).
cnf(c_0_506,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_132, c_0_505])).
cnf(c_0_507,plain,(esk4086_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_34, c_0_506])).
cnf(c_0_508,plain,(esk1986_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_36, c_0_507])).
cnf(c_0_509,plain,(esk2570_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_38, c_0_508])).
cnf(c_0_510,plain,(esk2568_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_139, c_0_509])).
cnf(c_0_511,plain,(esk2537_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_39, c_0_510])).
cnf(c_0_512,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_40, c_0_511])).
cnf(c_0_513,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_41, c_0_512])).
cnf(c_0_514,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_44, c_0_513])).
cnf(c_0_515,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_46, c_0_514])).
cnf(c_0_516,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_148, c_0_515])).
cnf(c_0_517,plain,(esk1987_4(X1,X2,X3,X4)|p(X1)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_47, c_0_516])).
cnf(c_0_518,plain,(esk369_0|p(X1)|~p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_80, c_0_517])).
cnf(c_0_519,plain,(esk369_0|p(X1)|~p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_518, c_0_49])).
cnf(c_0_520,plain,(esk369_0|p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_519, c_0_49])).
cnf(c_0_521,plain,(esk369_0|p(X1)), inference(spm,[status(thm)],[c_0_520, c_0_49])).
cnf(c_0_522,plain,(esk369_0), inference(spm,[status(thm)],[c_0_66, c_0_521])).
cnf(c_0_523,plain,(esk8283_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_522])])).
cnf(c_0_524,plain,(esk8867_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_55, c_0_523])).
cnf(c_0_525,plain,(esk8865_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_160, c_0_524])).
cnf(c_0_526,plain,(esk8834_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_56, c_0_525])).
cnf(c_0_527,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_57, c_0_526])).
cnf(c_0_528,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_58, c_0_527])).
cnf(c_0_529,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_61, c_0_528])).
cnf(c_0_530,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_63, c_0_529])).
cnf(c_0_531,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_169, c_0_530])).
cnf(c_0_532,plain,(esk8284_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_64, c_0_531])).
cnf(c_0_533,plain,(esk6184_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_6, c_0_532])).
cnf(c_0_534,plain,(esk6768_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_8, c_0_533])).
cnf(c_0_535,plain,(esk6766_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_107, c_0_534])).
cnf(c_0_536,plain,(esk6735_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_9, c_0_535])).
cnf(c_0_537,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_10, c_0_536])).
cnf(c_0_538,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_11, c_0_537])).
cnf(c_0_539,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_14, c_0_538])).
cnf(c_0_540,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_16, c_0_539])).
cnf(c_0_541,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_116, c_0_540])).
cnf(c_0_542,plain,(esk6185_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_19, c_0_541])).
cnf(c_0_543,plain,(esk4085_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_21, c_0_542])).
cnf(c_0_544,plain,(esk4669_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_23, c_0_543])).
cnf(c_0_545,plain,(esk4667_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_123, c_0_544])).
cnf(c_0_546,plain,(esk4636_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_24, c_0_545])).
cnf(c_0_547,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_25, c_0_546])).
cnf(c_0_548,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_26, c_0_547])).
cnf(c_0_549,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_29, c_0_548])).
cnf(c_0_550,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_31, c_0_549])).
cnf(c_0_551,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_132, c_0_550])).
cnf(c_0_552,plain,(esk4086_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_34, c_0_551])).
cnf(c_0_553,plain,(esk1986_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_36, c_0_552])).
cnf(c_0_554,plain,(esk2570_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_38, c_0_553])).
cnf(c_0_555,plain,(esk2568_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_139, c_0_554])).
cnf(c_0_556,plain,(esk2537_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_39, c_0_555])).
cnf(c_0_557,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_40, c_0_556])).
cnf(c_0_558,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_41, c_0_557])).
cnf(c_0_559,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_44, c_0_558])).
cnf(c_0_560,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_46, c_0_559])).
cnf(c_0_561,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_148, c_0_560])).
cnf(c_0_562,plain,(esk1987_4(X1,X2,X3,X4)|p(X4)|p(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_47, c_0_561])).
cnf(c_0_563,plain,(esk368_0|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_82, c_0_562])).
cnf(c_0_564,plain,(esk368_0|p(X1)|p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_563, c_0_49])).
cnf(c_0_565,plain,(esk368_0|p(X1)|p(X2)), inference(spm,[status(thm)],[c_0_564, c_0_49])).
cnf(c_0_566,plain,(esk368_0|p(X1)), inference(ef,[status(thm)],[c_0_565])).
cnf(c_0_567,plain,(esk368_0), inference(spm,[status(thm)],[c_0_66, c_0_566])).
cnf(c_0_568,plain,(esk8283_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_567])])).
cnf(c_0_569,plain,(esk8867_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_55, c_0_568])).
cnf(c_0_570,plain,(esk8865_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_160, c_0_569])).
cnf(c_0_571,plain,(esk8834_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_56, c_0_570])).
cnf(c_0_572,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_57, c_0_571])).
cnf(c_0_573,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_58, c_0_572])).
cnf(c_0_574,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_61, c_0_573])).
cnf(c_0_575,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_63, c_0_574])).
cnf(c_0_576,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_169, c_0_575])).
cnf(c_0_577,plain,(esk8284_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_64, c_0_576])).
cnf(c_0_578,plain,(esk6184_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_6, c_0_577])).
cnf(c_0_579,plain,(esk6768_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_8, c_0_578])).
cnf(c_0_580,plain,(esk6766_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_107, c_0_579])).
cnf(c_0_581,plain,(esk6735_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_9, c_0_580])).
cnf(c_0_582,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_10, c_0_581])).
cnf(c_0_583,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_11, c_0_582])).
cnf(c_0_584,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_14, c_0_583])).
cnf(c_0_585,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_16, c_0_584])).
cnf(c_0_586,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_116, c_0_585])).
cnf(c_0_587,plain,(esk6185_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_19, c_0_586])).
cnf(c_0_588,plain,(esk4085_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_21, c_0_587])).
cnf(c_0_589,plain,(esk4669_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_23, c_0_588])).
cnf(c_0_590,plain,(esk4667_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_123, c_0_589])).
cnf(c_0_591,plain,(esk4636_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_24, c_0_590])).
cnf(c_0_592,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_25, c_0_591])).
cnf(c_0_593,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_26, c_0_592])).
cnf(c_0_594,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_29, c_0_593])).
cnf(c_0_595,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_31, c_0_594])).
cnf(c_0_596,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_132, c_0_595])).
cnf(c_0_597,plain,(esk4086_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_34, c_0_596])).
cnf(c_0_598,plain,(esk1986_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_36, c_0_597])).
cnf(c_0_599,plain,(esk2570_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_38, c_0_598])).
cnf(c_0_600,plain,(esk2568_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_139, c_0_599])).
cnf(c_0_601,plain,(esk2537_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_39, c_0_600])).
cnf(c_0_602,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_40, c_0_601])).
cnf(c_0_603,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_41, c_0_602])).
cnf(c_0_604,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_44, c_0_603])).
cnf(c_0_605,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_46, c_0_604])).
cnf(c_0_606,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_148, c_0_605])).
cnf(c_0_607,plain,(esk1987_4(X1,X2,X3,X4)|p(X3)|p(X1)|~p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_47, c_0_606])).
cnf(c_0_608,plain,(esk367_0|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_84, c_0_607])).
cnf(c_0_609,plain,(esk367_0|p(X1)|p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_608, c_0_49])).
cnf(c_0_610,plain,(esk367_0|p(X1)|p(X2)), inference(spm,[status(thm)],[c_0_609, c_0_49])).
cnf(c_0_611,plain,(esk367_0|p(X1)), inference(ef,[status(thm)],[c_0_610])).
cnf(c_0_612,plain,(esk367_0), inference(spm,[status(thm)],[c_0_66, c_0_611])).
cnf(c_0_613,plain,(esk8283_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_612])])).
cnf(c_0_614,plain,(esk8867_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_55, c_0_613])).
cnf(c_0_615,plain,(esk8865_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_160, c_0_614])).
cnf(c_0_616,plain,(esk8834_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_56, c_0_615])).
cnf(c_0_617,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_57, c_0_616])).
cnf(c_0_618,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_58, c_0_617])).
cnf(c_0_619,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_61, c_0_618])).
cnf(c_0_620,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_63, c_0_619])).
cnf(c_0_621,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_169, c_0_620])).
cnf(c_0_622,plain,(esk8284_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_64, c_0_621])).
cnf(c_0_623,plain,(esk6184_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_6, c_0_622])).
cnf(c_0_624,plain,(esk6768_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_8, c_0_623])).
cnf(c_0_625,plain,(esk6766_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_107, c_0_624])).
cnf(c_0_626,plain,(esk6735_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_9, c_0_625])).
cnf(c_0_627,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_10, c_0_626])).
cnf(c_0_628,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_11, c_0_627])).
cnf(c_0_629,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_14, c_0_628])).
cnf(c_0_630,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_16, c_0_629])).
cnf(c_0_631,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_116, c_0_630])).
cnf(c_0_632,plain,(esk6185_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_19, c_0_631])).
cnf(c_0_633,plain,(esk4085_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_21, c_0_632])).
cnf(c_0_634,plain,(esk4669_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_23, c_0_633])).
cnf(c_0_635,plain,(esk4667_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_123, c_0_634])).
cnf(c_0_636,plain,(esk4636_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_24, c_0_635])).
cnf(c_0_637,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_25, c_0_636])).
cnf(c_0_638,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_26, c_0_637])).
cnf(c_0_639,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_29, c_0_638])).
cnf(c_0_640,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_31, c_0_639])).
cnf(c_0_641,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_132, c_0_640])).
cnf(c_0_642,plain,(esk4086_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_34, c_0_641])).
cnf(c_0_643,plain,(esk1986_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_36, c_0_642])).
cnf(c_0_644,plain,(esk2570_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_38, c_0_643])).
cnf(c_0_645,plain,(esk2568_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_139, c_0_644])).
cnf(c_0_646,plain,(esk2537_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_39, c_0_645])).
cnf(c_0_647,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_40, c_0_646])).
cnf(c_0_648,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_41, c_0_647])).
cnf(c_0_649,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_44, c_0_648])).
cnf(c_0_650,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_46, c_0_649])).
cnf(c_0_651,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_148, c_0_650])).
cnf(c_0_652,plain,(esk1987_4(X1,X2,X3,X4)|p(X1)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_47, c_0_651])).
cnf(c_0_653,plain,(esk366_0|p(X1)|p(X2)|p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_86, c_0_652])).
cnf(c_0_654,plain,(esk366_0|p(X1)|p(X2)|p(X3)), inference(spm,[status(thm)],[c_0_653, c_0_49])).
cnf(c_0_655,plain,(esk366_0|p(X1)|p(X2)), inference(ef,[status(thm)],[c_0_654])).
cnf(c_0_656,plain,(esk366_0|p(X1)), inference(ef,[status(thm)],[c_0_655])).
cnf(c_0_657,plain,(esk366_0), inference(spm,[status(thm)],[c_0_66, c_0_656])).
cnf(c_0_658,plain,(esk8283_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_657])])).
cnf(c_0_659,plain,(esk8867_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_55, c_0_658])).
cnf(c_0_660,plain,(esk8865_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_160, c_0_659])).
cnf(c_0_661,plain,(esk8834_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_56, c_0_660])).
cnf(c_0_662,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_57, c_0_661])).
cnf(c_0_663,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_58, c_0_662])).
cnf(c_0_664,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_61, c_0_663])).
cnf(c_0_665,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_63, c_0_664])).
cnf(c_0_666,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_169, c_0_665])).
cnf(c_0_667,plain,(esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_64, c_0_666])).
cnf(c_0_668,plain,(esk6184_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_6, c_0_667])).
cnf(c_0_669,plain,(esk6768_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_8, c_0_668])).
cnf(c_0_670,plain,(esk6766_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_107, c_0_669])).
cnf(c_0_671,plain,(esk6735_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_9, c_0_670])).
cnf(c_0_672,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_10, c_0_671])).
cnf(c_0_673,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_11, c_0_672])).
cnf(c_0_674,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_14, c_0_673])).
cnf(c_0_675,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_16, c_0_674])).
cnf(c_0_676,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_116, c_0_675])).
cnf(c_0_677,plain,(esk6185_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_19, c_0_676])).
cnf(c_0_678,plain,(esk4085_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_21, c_0_677])).
cnf(c_0_679,plain,(esk4669_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_23, c_0_678])).
cnf(c_0_680,plain,(esk4667_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_123, c_0_679])).
cnf(c_0_681,plain,(esk4636_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_24, c_0_680])).
cnf(c_0_682,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_25, c_0_681])).
cnf(c_0_683,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_26, c_0_682])).
cnf(c_0_684,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_29, c_0_683])).
cnf(c_0_685,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_31, c_0_684])).
cnf(c_0_686,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_132, c_0_685])).
cnf(c_0_687,plain,(esk4086_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_34, c_0_686])).
cnf(c_0_688,plain,(esk1986_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_36, c_0_687])).
cnf(c_0_689,plain,(esk2570_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_38, c_0_688])).
cnf(c_0_690,plain,(esk2568_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_139, c_0_689])).
cnf(c_0_691,plain,(esk2537_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_39, c_0_690])).
cnf(c_0_692,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_40, c_0_691])).
cnf(c_0_693,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_41, c_0_692])).
cnf(c_0_694,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_44, c_0_693])).
cnf(c_0_695,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_46, c_0_694])).
cnf(c_0_696,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_148, c_0_695])).
cnf(c_0_697,plain,(esk1987_4(X1,X2,X3,X4)|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_47, c_0_696])).
cnf(c_0_698,plain,(esk365_0|p(X1)|p(X2)|~p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_88, c_0_697])).
cnf(c_0_699,plain,(esk365_0|p(X1)|p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_698, c_0_49])).
cnf(c_0_700,plain,(esk365_0|p(X1)|p(X2)), inference(spm,[status(thm)],[c_0_699, c_0_49])).
cnf(c_0_701,plain,(esk365_0|p(X1)), inference(ef,[status(thm)],[c_0_700])).
cnf(c_0_702,plain,(esk365_0), inference(spm,[status(thm)],[c_0_66, c_0_701])).
cnf(c_0_703,plain,(esk8283_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_702])])).
cnf(c_0_704,plain,(esk8867_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_55, c_0_703])).
cnf(c_0_705,plain,(esk8865_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_160, c_0_704])).
cnf(c_0_706,plain,(esk8834_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_56, c_0_705])).
cnf(c_0_707,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_57, c_0_706])).
cnf(c_0_708,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_58, c_0_707])).
cnf(c_0_709,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_61, c_0_708])).
cnf(c_0_710,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_63, c_0_709])).
cnf(c_0_711,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_169, c_0_710])).
cnf(c_0_712,plain,(esk8284_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_64, c_0_711])).
cnf(c_0_713,plain,(esk6184_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_6, c_0_712])).
cnf(c_0_714,plain,(esk6768_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_8, c_0_713])).
cnf(c_0_715,plain,(esk6766_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_107, c_0_714])).
cnf(c_0_716,plain,(esk6735_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_9, c_0_715])).
cnf(c_0_717,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_10, c_0_716])).
cnf(c_0_718,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_11, c_0_717])).
cnf(c_0_719,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_14, c_0_718])).
cnf(c_0_720,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_16, c_0_719])).
cnf(c_0_721,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_116, c_0_720])).
cnf(c_0_722,plain,(esk6185_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_19, c_0_721])).
cnf(c_0_723,plain,(esk4085_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_21, c_0_722])).
cnf(c_0_724,plain,(esk4669_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_23, c_0_723])).
cnf(c_0_725,plain,(esk4667_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_123, c_0_724])).
cnf(c_0_726,plain,(esk4636_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_24, c_0_725])).
cnf(c_0_727,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_25, c_0_726])).
cnf(c_0_728,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_26, c_0_727])).
cnf(c_0_729,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_29, c_0_728])).
cnf(c_0_730,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_31, c_0_729])).
cnf(c_0_731,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_132, c_0_730])).
cnf(c_0_732,plain,(esk4086_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_34, c_0_731])).
cnf(c_0_733,plain,(esk1986_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_36, c_0_732])).
cnf(c_0_734,plain,(esk2570_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_38, c_0_733])).
cnf(c_0_735,plain,(esk2568_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_139, c_0_734])).
cnf(c_0_736,plain,(esk2537_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_39, c_0_735])).
cnf(c_0_737,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_40, c_0_736])).
cnf(c_0_738,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_41, c_0_737])).
cnf(c_0_739,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_44, c_0_738])).
cnf(c_0_740,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_46, c_0_739])).
cnf(c_0_741,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_148, c_0_740])).
cnf(c_0_742,plain,(esk1987_4(X1,X2,X3,X4)|p(X1)|p(X2)|p(X4)|~p(X3)), inference(spm,[status(thm)],[c_0_47, c_0_741])).
cnf(c_0_743,plain,(esk364_0|p(X1)|p(X2)|p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_90, c_0_742])).
cnf(c_0_744,plain,(esk364_0|p(X1)|p(X2)|p(X3)), inference(spm,[status(thm)],[c_0_743, c_0_49])).
cnf(c_0_745,plain,(esk364_0|p(X1)|p(X2)), inference(ef,[status(thm)],[c_0_744])).
cnf(c_0_746,plain,(esk364_0|p(X1)), inference(ef,[status(thm)],[c_0_745])).
cnf(c_0_747,plain,(esk364_0), inference(spm,[status(thm)],[c_0_66, c_0_746])).
cnf(c_0_748,plain,(esk8283_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_747])])).
cnf(c_0_749,plain,(esk8867_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_55, c_0_748])).
cnf(c_0_750,plain,(esk8865_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_160, c_0_749])).
cnf(c_0_751,plain,(esk8834_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_56, c_0_750])).
cnf(c_0_752,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_57, c_0_751])).
cnf(c_0_753,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_58, c_0_752])).
cnf(c_0_754,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_61, c_0_753])).
cnf(c_0_755,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_63, c_0_754])).
cnf(c_0_756,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_169, c_0_755])).
cnf(c_0_757,plain,(esk8284_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_64, c_0_756])).
cnf(c_0_758,plain,(esk6184_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_6, c_0_757])).
cnf(c_0_759,plain,(esk6768_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_8, c_0_758])).
cnf(c_0_760,plain,(esk6766_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_107, c_0_759])).
cnf(c_0_761,plain,(esk6735_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_9, c_0_760])).
cnf(c_0_762,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_10, c_0_761])).
cnf(c_0_763,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_11, c_0_762])).
cnf(c_0_764,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_14, c_0_763])).
cnf(c_0_765,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_16, c_0_764])).
cnf(c_0_766,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_116, c_0_765])).
cnf(c_0_767,plain,(esk6185_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_19, c_0_766])).
cnf(c_0_768,plain,(esk4085_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_21, c_0_767])).
cnf(c_0_769,plain,(esk4669_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_23, c_0_768])).
cnf(c_0_770,plain,(esk4667_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_123, c_0_769])).
cnf(c_0_771,plain,(esk4636_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_24, c_0_770])).
cnf(c_0_772,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_25, c_0_771])).
cnf(c_0_773,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_26, c_0_772])).
cnf(c_0_774,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_29, c_0_773])).
cnf(c_0_775,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_31, c_0_774])).
cnf(c_0_776,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_132, c_0_775])).
cnf(c_0_777,plain,(esk4086_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_34, c_0_776])).
cnf(c_0_778,plain,(esk1986_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_36, c_0_777])).
cnf(c_0_779,plain,(esk2570_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_38, c_0_778])).
cnf(c_0_780,plain,(esk2568_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_139, c_0_779])).
cnf(c_0_781,plain,(esk2537_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_39, c_0_780])).
cnf(c_0_782,plain,(esk2535_4(X1,X2,X3,X4)|esk2536_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_40, c_0_781])).
cnf(c_0_783,plain,(esk2535_4(X1,X2,X3,X4)|esk2448_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_41, c_0_782])).
cnf(c_0_784,plain,(esk2535_4(X1,X2,X3,X4)|esk2443_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_44, c_0_783])).
cnf(c_0_785,plain,(esk2443_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_46, c_0_784])).
cnf(c_0_786,plain,(esk2440_4(X1,X2,X3,X4)|esk1987_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_148, c_0_785])).
cnf(c_0_787,plain,(esk1987_4(X1,X2,X3,X4)|p(X3)|p(X2)|p(X1)|~p(X4)), inference(spm,[status(thm)],[c_0_47, c_0_786])).
cnf(c_0_788,plain,(esk363_0|p(X1)|p(X2)|p(X3)|~p(X4)), inference(spm,[status(thm)],[c_0_92, c_0_787])).
cnf(c_0_789,plain,(esk363_0|p(X1)|p(X2)|p(X3)), inference(spm,[status(thm)],[c_0_788, c_0_49])).
cnf(c_0_790,plain,(esk363_0|p(X1)|p(X2)), inference(ef,[status(thm)],[c_0_789])).
cnf(c_0_791,plain,(esk363_0|p(X1)), inference(ef,[status(thm)],[c_0_790])).
cnf(c_0_792,plain,(esk363_0), inference(spm,[status(thm)],[c_0_66, c_0_791])).
cnf(c_0_793,plain,(esk8283_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_792])])).
cnf(c_0_794,plain,(esk8867_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_55, c_0_793])).
cnf(c_0_795,plain,(esk8865_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_160, c_0_794])).
cnf(c_0_796,plain,(esk8834_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_56, c_0_795])).
cnf(c_0_797,plain,(esk8833_4(X1,X2,X3,X4)|esk8832_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_57, c_0_796])).
cnf(c_0_798,plain,(esk8833_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_58, c_0_797])).
cnf(c_0_799,plain,(esk8745_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_61, c_0_798])).
cnf(c_0_800,plain,(esk8740_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_63, c_0_799])).
cnf(c_0_801,plain,(esk8737_4(X1,X2,X3,X4)|esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_169, c_0_800])).
cnf(c_0_802,plain,(esk8284_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_64, c_0_801])).
cnf(c_0_803,plain,(esk6184_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_6, c_0_802])).
cnf(c_0_804,plain,(esk6768_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_8, c_0_803])).
cnf(c_0_805,plain,(esk6766_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_107, c_0_804])).
cnf(c_0_806,plain,(esk6735_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_9, c_0_805])).
cnf(c_0_807,plain,(esk6733_4(X1,X2,X3,X4)|esk6734_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_10, c_0_806])).
cnf(c_0_808,plain,(esk6733_4(X1,X2,X3,X4)|esk6646_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_11, c_0_807])).
cnf(c_0_809,plain,(esk6733_4(X1,X2,X3,X4)|esk6641_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_14, c_0_808])).
cnf(c_0_810,plain,(esk6641_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_16, c_0_809])).
cnf(c_0_811,plain,(esk6638_4(X1,X2,X3,X4)|esk6185_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_116, c_0_810])).
cnf(c_0_812,plain,(esk6185_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_19, c_0_811])).
cnf(c_0_813,plain,(esk4085_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_21, c_0_812])).
cnf(c_0_814,plain,(esk4669_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_23, c_0_813])).
cnf(c_0_815,plain,(esk4667_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_123, c_0_814])).
cnf(c_0_816,plain,(esk4636_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_24, c_0_815])).
cnf(c_0_817,plain,(esk4635_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_25, c_0_816])).
cnf(c_0_818,plain,(esk4547_4(X1,X2,X3,X4)|esk4634_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_26, c_0_817])).
cnf(c_0_819,plain,(esk4634_4(X1,X2,X3,X4)|esk4542_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_29, c_0_818])).
cnf(c_0_820,plain,(esk4542_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_31, c_0_819])).
cnf(c_0_821,plain,(esk4539_4(X1,X2,X3,X4)|esk4086_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_132, c_0_820])).
cnf(c_0_822,plain,(esk4086_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_34, c_0_821])).
cnf(c_0_823,plain,(esk1986_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_36, c_0_822])).
cnf(c_0_824,plain,(esk2570_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_38, c_0_823])).
cnf(c_0_825,plain,(esk2568_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_139, c_0_824])).
cnf(c_0_826,plain,(p(X1)|p(X2)|p(X3)|p(X4)|~esk2874_4(X3,X4,X2,X1)), inference(sr,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_827,plain,(esk2874_4(X1,X2,X3,X4)|esk2446_4(X1,X2,X3,X4)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_828,plain,(esk2567_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X2)|p(X1)), inference(spm,[status(thm)],[c_0_98, c_0_825])).
cnf(c_0_829,plain,(esk1878_4(X1,X2,X3,X4)|~esk2446_4(X1,X2,X3,X4)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_830,plain,(esk2446_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X1)|p(X2)), inference(spm,[status(thm)],[c_0_826, c_0_827])).
cnf(c_0_831,plain,(p(X1)|p(X2)|p(X3)|p(X4)|~esk1878_4(X4,X3,X2,X1)), inference(spm,[status(thm)],[c_0_101, c_0_828])).
cnf(c_0_832,plain,(esk1878_4(X1,X2,X3,X4)|p(X4)|p(X3)|p(X1)|p(X2)), inference(spm,[status(thm)],[c_0_829, c_0_830])).
cnf(c_0_833,plain,(p(X1)|p(X2)|p(X3)|p(X4)), inference(spm,[status(thm)],[c_0_831, c_0_832])).
cnf(c_0_834,plain,(p(X1)|p(X2)|p(X3)), inference(ef,[status(thm)],[c_0_833])).
cnf(c_0_835,plain,(p(X1)|p(X2)), inference(ef,[status(thm)],[c_0_834])).
cnf(c_0_836,plain,(p(X1)), inference(ef,[status(thm)],[c_0_835])).
cnf(c_0_837,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_836])]), ['proof']).
# SZS output end CNFRefutation.
EOF
