% Problem    : SYN445-1 : TPTP v6.1.0. Released v2.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n145.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:24:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_N___023_B07_F1_SP_PI_Q7_CS_SP_CO_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,negated_conjecture,(ndr1_0|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause5)).
cnf(c_0_1,negated_conjecture,(hskp26|hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1)).
cnf(c_0_2,negated_conjecture,(ndr1_0|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause15)).
cnf(c_0_3,negated_conjecture,(c1_1(X1)|c2_1(X1)|hskp16|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause128)).
cnf(c_0_4,negated_conjecture,(c3_1(a336)|~hskp21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause51)).
cnf(c_0_5,negated_conjecture,(~c1_1(a336)|~hskp21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause93)).
cnf(c_0_6,negated_conjecture,(c3_1(a321)|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause59)).
cnf(c_0_7,negated_conjecture,(c2_1(a321)|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause58)).
cnf(c_0_8,negated_conjecture,(~c1_1(a321)|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause100)).
cnf(c_0_9,negated_conjecture,(c2_1(X1)|c3_1(X1)|hskp21|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause129)).
cnf(c_0_10,negated_conjecture,(c0_1(a349)|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause47)).
cnf(c_0_11,negated_conjecture,(~c2_1(a349)|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause81)).
cnf(c_0_12,negated_conjecture,(~c3_1(a349)|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause82)).
cnf(c_0_13,negated_conjecture,(~c0_1(a336)|~hskp21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause92)).
cnf(c_0_14,negated_conjecture,(c0_1(a295)|~hskp1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause79)).
cnf(c_0_15,negated_conjecture,(c2_1(a295)|~hskp1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause80)).
cnf(c_0_16,negated_conjecture,(~c1_1(a295)|~hskp1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause124)).
cnf(c_0_17,negated_conjecture,(c3_1(a310)|~hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause65)).
cnf(c_0_18,negated_conjecture,(~c2_1(a310)|~hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause106)).
cnf(c_0_19,negated_conjecture,(~c0_1(a310)|~hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause105)).
cnf(c_0_20,negated_conjecture,(c3_1(a309)|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause67)).
cnf(c_0_21,negated_conjecture,(c1_1(a309)|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause66)).
cnf(c_0_22,negated_conjecture,(~c2_1(a309)|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause107)).
cnf(c_0_23,negated_conjecture,(hskp26|hskp11|hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause36)).
cnf(c_0_24,negated_conjecture,(c1_1(X1)|c1_1(X2)|hskp4|~c0_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause167)).
cnf(c_0_25,negated_conjecture,(c2_1(X2)|~c0_1(X1)|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c1_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause169)).
cnf(c_0_26,negated_conjecture,(c1_1(X1)|c0_1(X2)|c1_1(X2)|hskp1|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause154)).
cnf(c_0_27,negated_conjecture,(c2_1(X1)|c0_1(X2)|c1_1(X2)|hskp0|~c1_1(X1)|~c3_1(X1)|~ndr1_0|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause155)).
cnf(c_0_28,negated_conjecture,(c1_1(a333)|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause52)).
cnf(c_0_29,negated_conjecture,(~c2_1(a333)|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause94)).
cnf(c_0_30,negated_conjecture,(~c3_1(a333)|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause95)).
cnf(c_0_31,negated_conjecture,(~c0_1(a294)|~hskp0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause125)).
cnf(c_0_32,negated_conjecture,(~c2_1(a294)|~hskp0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause126)).
cnf(c_0_33,negated_conjecture,(~c3_1(a294)|~hskp0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause127)).
cnf(c_0_34,negated_conjecture,(c3_1(a299)|~hskp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause77)).
cnf(c_0_35,negated_conjecture,(~c2_1(a299)|~hskp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause121)).
cnf(c_0_36,negated_conjecture,(c0_1(a299)|~hskp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause76)).
cnf(c_0_37,negated_conjecture,(~c3_1(a346)|~hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause85)).
cnf(c_0_38,negated_conjecture,(~c0_1(a346)|~hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause83)).
cnf(c_0_39,negated_conjecture,(~c1_1(a346)|~hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause84)).
cnf(c_0_40,negated_conjecture,(hskp20|hskp12|hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause37)).
cnf(c_0_41,negated_conjecture,(c0_1(X1)|hskp14|hskp0|~c2_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause137)).
cnf(c_0_42,negated_conjecture,(~c0_1(a300)|~hskp4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause118)).
cnf(c_0_43,negated_conjecture,(~c1_1(a300)|~hskp4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause119)).
cnf(c_0_44,negated_conjecture,(~c2_1(a300)|~hskp4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause120)).
cnf(c_0_45,negated_conjecture,(c1_1(X1)|c0_1(X2)|c2_1(X2)|c0_1(X3)|c2_1(X3)|c3_1(X3)|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c1_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause176)).
cnf(c_0_46,negated_conjecture,(c0_1(X1)|c0_1(X2)|c2_1(X2)|c3_1(X2)|c0_1(X3)|c1_1(X3)|c3_1(X3)|~c2_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause174)).
cnf(c_0_47,negated_conjecture,(~c0_1(a315)|~hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause103)).
cnf(c_0_48,negated_conjecture,(c1_1(a315)|~hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause61)).
cnf(c_0_49,negated_conjecture,(c2_1(a315)|~hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause62)).
cnf(c_0_50,negated_conjecture,(c2_1(X1)|c0_1(X2)|c1_1(X2)|c2_1(X2)|hskp0|~c0_1(X1)|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause149)).
cnf(c_0_51,negated_conjecture,(c1_1(X1)|hskp27|hskp3|~c2_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause141)).
cnf(c_0_52,negated_conjecture,(c1_1(X2)|c0_1(X3)|c1_1(X3)|c2_1(X3)|~c0_1(X1)|~c1_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause179)).
cnf(c_0_53,negated_conjecture,(c2_1(X1)|c3_1(X1)|c0_1(X2)|hskp12|~c0_1(X1)|~ndr1_0|~c1_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause162)).
cnf(c_0_54,negated_conjecture,(c2_1(X1)|c3_1(X1)|c0_1(X2)|hskp3|~c1_1(X1)|~ndr1_0|~c2_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause163)).
cnf(c_0_55,negated_conjecture,(c3_1(a334)|~hskp27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause46)).
cnf(c_0_56,negated_conjecture,(c0_1(a334)|~hskp27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause44)).
cnf(c_0_57,negated_conjecture,(c0_1(X1)|c2_1(X1)|hskp10|hskp11|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause134)).
cnf(c_0_58,negated_conjecture,(c3_1(a308)|~hskp10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause69)).
cnf(c_0_59,negated_conjecture,(~c0_1(a308)|~hskp10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause108)).
cnf(c_0_60,negated_conjecture,(c1_1(a308)|~hskp10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause68)).
cnf(c_0_61,negated_conjecture,(c0_1(X1)|c1_1(X1)|hskp2|hskp0|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause131)).
cnf(c_0_62,negated_conjecture,(~c0_1(a297)|~hskp2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause122)).
cnf(c_0_63,negated_conjecture,(c1_1(a297)|~hskp2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause78)).
cnf(c_0_64,negated_conjecture,(~c3_1(a297)|~hskp2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause123)).
cnf(c_0_65,negated_conjecture,(c0_1(X1)|c0_1(X2)|c2_1(X2)|c3_1(X2)|c0_1(X3)|c1_1(X3)|c2_1(X3)|~c1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause172)).
cnf(c_0_66,negated_conjecture,(c1_1(X1)|hskp17|hskp18|~c0_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause138)).
cnf(c_0_67,negated_conjecture,(c2_1(a329)|~hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause55)).
cnf(c_0_68,negated_conjecture,(~c3_1(a329)|~hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause98)).
cnf(c_0_69,negated_conjecture,(~c0_1(a329)|~hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause97)).
cnf(c_0_70,negated_conjecture,(~c0_1(a305)|~hskp8), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause111)).
cnf(c_0_71,negated_conjecture,(c2_1(a305)|~hskp8), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause71)).
cnf(c_0_72,negated_conjecture,(~c1_1(a305)|~hskp8), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause112)).
cnf(c_0_73,negated_conjecture,(c3_1(X1)|c0_1(X2)|c3_1(X2)|hskp8|~c1_1(X1)|~c2_1(X1)|~ndr1_0|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause161)).
cnf(c_0_74,negated_conjecture,(c1_1(X1)|hskp11|hskp20|~c0_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause140)).
cnf(c_0_75,negated_conjecture,(c2_1(a323)|~hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause57)).
cnf(c_0_76,negated_conjecture,(c0_1(a323)|~hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause56)).
cnf(c_0_77,negated_conjecture,(~c3_1(a323)|~hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause99)).
cnf(c_0_78,negated_conjecture,(c3_1(X2)|hskp11|~c0_1(X1)|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause170)).
cnf(c_0_79,negated_conjecture,(c0_1(X1)|c0_1(X2)|c3_1(X2)|c0_1(X3)|c2_1(X3)|~c1_1(X1)|~c3_1(X1)|~ndr1_0|~c1_1(X2)|~c1_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause178)).
fof(c_0_80, plain, ((~epred31_0<=>![X1]:((((hskp1|c1_1(X1))|~ndr1_0)|~c2_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_81,negated_conjecture,(ndr1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
cnf(c_0_82,negated_conjecture,(epred31_0|c1_1(X1)|hskp1|~c3_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_80]), c_0_81])])).
cnf(c_0_83,negated_conjecture,(c2_1(X1)|c1_1(X1)|hskp16|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_81])])).
fof(c_0_84, plain, ((~epred32_0<=>![X2]:((c0_1(X2)|c1_1(X2))|~c2_1(X2)))), introduced(definition)).
fof(c_0_85, plain, ((~epred34_0<=>![X2]:((c0_1(X2)|c1_1(X2))|~c3_1(X2)))), introduced(definition)).
fof(c_0_86, plain, ((~epred28_0<=>![X1]:((c1_1(X1)|~c0_1(X1))|~c2_1(X1)))), introduced(definition)).
cnf(c_0_87,negated_conjecture,(epred31_0|hskp1|~c2_1(a336)|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_4]), c_0_5])).
cnf(c_0_88,negated_conjecture,(c2_1(a336)|hskp16|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_4]), c_0_5])).
cnf(c_0_89,negated_conjecture,(epred31_0|hskp1|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_6]), c_0_7]), c_0_8])).
cnf(c_0_90,negated_conjecture,(c3_1(X1)|c2_1(X1)|hskp21|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_81])])).
cnf(c_0_91,negated_conjecture,(epred32_0|c1_1(X1)|c0_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_84])).
cnf(c_0_92,negated_conjecture,(epred34_0|c1_1(X1)|c0_1(X1)|~c3_1(X1)), inference(split_equiv,[status(thm)],[c_0_85])).
fof(c_0_93, plain, ((~epred1_0<=>![X2]:(((c2_1(X2)|~ndr1_0)|~c1_1(X2))|~c3_1(X2)))), introduced(definition)).
fof(c_0_94, plain, ((~epred2_0<=>![X1]:((~c0_1(X1)|~c2_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_95,negated_conjecture,(epred28_0|c1_1(X1)|~c2_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_86])).
cnf(c_0_96,negated_conjecture,(epred31_0|hskp1|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89])).
cnf(c_0_97,negated_conjecture,(hskp21|~hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_10]), c_0_11]), c_0_12])).
cnf(c_0_98,negated_conjecture,(epred32_0|~c2_1(a336)|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_91]), c_0_5])).
cnf(c_0_99,negated_conjecture,(epred34_0|~hskp21), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_4]), c_0_13]), c_0_5])).
cnf(c_0_100,negated_conjecture,(epred1_0|c2_1(X1)|~c3_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_93]), c_0_81])])).
cnf(c_0_101,negated_conjecture,(epred2_0|~c3_1(X1)|~c2_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_94])).
cnf(c_0_102,negated_conjecture,(epred28_0|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_14]), c_0_15]), c_0_16])).
cnf(c_0_103,negated_conjecture,(epred31_0|hskp1|~hskp26), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_104,negated_conjecture,(epred32_0|epred28_0|c1_1(X1)|~c2_1(X1)), inference(spm,[status(thm)],[c_0_95, c_0_91])).
cnf(c_0_105,negated_conjecture,(epred32_0|hskp16|~hskp21), inference(spm,[status(thm)],[c_0_98, c_0_88])).
cnf(c_0_106,negated_conjecture,(epred34_0|~hskp26), inference(spm,[status(thm)],[c_0_99, c_0_97])).
cnf(c_0_107,negated_conjecture,(epred1_0|~c1_1(a310)|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_17]), c_0_18])).
cnf(c_0_108,negated_conjecture,(epred34_0|c1_1(a310)|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_17]), c_0_19])).
cnf(c_0_109,negated_conjecture,(epred2_0|~c0_1(a321)|~hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_6]), c_0_7])).
cnf(c_0_110,negated_conjecture,(epred34_0|c0_1(a321)|~hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_6]), c_0_8])).
fof(c_0_111, plain, ((~epred37_0<=>![X1]:((((hskp4|c1_1(X1))|~ndr1_0)|~c0_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_112,negated_conjecture,(epred31_0|epred28_0|~hskp26), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_113,negated_conjecture,(epred31_0|epred28_0|~hskp16), inference(spm,[status(thm)],[c_0_102, c_0_89])).
cnf(c_0_114,negated_conjecture,(epred32_0|epred28_0|~hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_7]), c_0_8])).
cnf(c_0_115,negated_conjecture,(epred32_0|hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_97]), c_0_1])).
fof(c_0_116, plain, ((~epred30_0<=>![X2]:((c0_1(X2)|~c2_1(X2))|~c3_1(X2)))), introduced(definition)).
cnf(c_0_117,negated_conjecture,(epred1_0|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_20]), c_0_21]), c_0_22])).
cnf(c_0_118,negated_conjecture,(epred34_0|hskp11|hskp12), inference(spm,[status(thm)],[c_0_106, c_0_23])).
cnf(c_0_119,negated_conjecture,(epred34_0|epred1_0|~hskp12), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_120,negated_conjecture,(epred34_0|epred2_0|~hskp16), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_121,negated_conjecture,(epred34_0|hskp16), inference(spm,[status(thm)],[c_0_106, c_0_1])).
fof(c_0_122, plain, ((~epred45_0<=>![X1]:((c1_1(X1)|~c2_1(X1))|~c3_1(X1)))), introduced(definition)).
fof(c_0_123, plain, ((~epred10_0<=>![X1]:((((hskp0|c2_1(X1))|~ndr1_0)|~c0_1(X1))|~c1_1(X1)))), introduced(definition)).
fof(c_0_124, plain, ((~epred9_0<=>![X2]:((c0_1(X2)|c2_1(X2))|~c1_1(X2)))), introduced(definition)).
cnf(c_0_125,negated_conjecture,(~epred37_0|~epred28_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_24, c_0_111]), c_0_86])).
cnf(c_0_126,negated_conjecture,(epred31_0|epred28_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_1]), c_0_113])).
cnf(c_0_127,negated_conjecture,(epred32_0|epred28_0), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_128,negated_conjecture,(epred30_0|c0_1(X1)|~c3_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_116])).
fof(c_0_129, plain, ((~epred33_0<=>![X1]:((((hskp0|c2_1(X1))|~ndr1_0)|~c1_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_130,negated_conjecture,(~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_25, c_0_93]), c_0_94])).
cnf(c_0_131,negated_conjecture,(epred34_0|epred1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])).
cnf(c_0_132,negated_conjecture,(epred34_0|epred2_0), inference(spm,[status(thm)],[c_0_120, c_0_121])).
fof(c_0_133, plain, ((~epred29_0<=>![X1]:((((hskp3|c2_1(X1))|c3_1(X1))|~ndr1_0)|~c1_1(X1)))), introduced(definition)).
fof(c_0_134, plain, ((~epred49_0<=>![X3]:(((c0_1(X3)|c1_1(X3))|c2_1(X3))|~ndr1_0))), introduced(definition)).
cnf(c_0_135,negated_conjecture,(epred45_0|c1_1(X1)|~c3_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_122])).
cnf(c_0_136,negated_conjecture,(epred10_0|c2_1(X1)|hskp0|~c1_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_123]), c_0_81])])).
cnf(c_0_137,negated_conjecture,(epred9_0|c2_1(X1)|c0_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_124])).
fof(c_0_138, plain, ((~epred42_0<=>![X3]:((c0_1(X3)|c1_1(X3))|c3_1(X3)))), introduced(definition)).
cnf(c_0_139,negated_conjecture,(~epred32_0|~epred31_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_26, c_0_80]), c_0_84])).
cnf(c_0_140,negated_conjecture,(epred31_0|~epred37_0), inference(spm,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_141,negated_conjecture,(epred32_0|~epred37_0), inference(spm,[status(thm)],[c_0_125, c_0_127])).
cnf(c_0_142,negated_conjecture,(epred30_0|~c2_1(a336)|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_4]), c_0_13])).
cnf(c_0_143,negated_conjecture,(~epred34_0|~epred33_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_27, c_0_129]), c_0_85])).
cnf(c_0_144,negated_conjecture,(epred34_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_132])).
cnf(c_0_145,negated_conjecture,(epred29_0|c3_1(X1)|c2_1(X1)|hskp3|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_133]), c_0_81])])).
cnf(c_0_146,negated_conjecture,(epred49_0|c2_1(X1)|c1_1(X1)|c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_134]), c_0_81])])).
fof(c_0_147, plain, ((~epred16_0<=>![X1]:((((hskp12|c2_1(X1))|c3_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
cnf(c_0_148,negated_conjecture,(epred45_0|~c2_1(a336)|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_4]), c_0_5])).
cnf(c_0_149,negated_conjecture,(epred45_0|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_6]), c_0_7]), c_0_8])).
cnf(c_0_150,negated_conjecture,(epred10_0|epred9_0|c2_1(X1)|hskp0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_136, c_0_137])).
cnf(c_0_151,negated_conjecture,(epred42_0|c3_1(X1)|c1_1(X1)|c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_138])).
cnf(c_0_152,negated_conjecture,(epred37_0|c1_1(X1)|hskp4|~c3_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_111]), c_0_81])])).
cnf(c_0_153,negated_conjecture,(~epred37_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_141])).
cnf(c_0_154,negated_conjecture,(epred30_0|hskp16|~hskp21), inference(spm,[status(thm)],[c_0_142, c_0_88])).
cnf(c_0_155,negated_conjecture,(epred33_0|c2_1(X1)|hskp0|~c3_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_129]), c_0_81])])).
cnf(c_0_156,negated_conjecture,(~epred33_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_143, c_0_144])])).
cnf(c_0_157,negated_conjecture,(epred29_0|hskp3|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_28]), c_0_29]), c_0_30])).
fof(c_0_158, plain, ((~epred11_0<=>![X2]:((c0_1(X2)|c1_1(X2))|c2_1(X2)))), introduced(definition)).
cnf(c_0_159,negated_conjecture,(epred9_0|~c1_1(a294)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_137]), c_0_32])).
cnf(c_0_160,negated_conjecture,(epred49_0|c1_1(a294)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_146]), c_0_32])).
cnf(c_0_161,negated_conjecture,(epred16_0|c3_1(X1)|c2_1(X1)|hskp12|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_147]), c_0_81])])).
cnf(c_0_162,negated_conjecture,(epred45_0|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_88]), c_0_149])).
cnf(c_0_163,negated_conjecture,(epred10_0|epred9_0|hskp0|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_28]), c_0_29])).
cnf(c_0_164,negated_conjecture,(epred42_0|c1_1(a294)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_151]), c_0_31])).
cnf(c_0_165,negated_conjecture,(hskp4|~c0_1(a321)|~hskp16), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_6]), c_0_153]), c_0_8])).
cnf(c_0_166,negated_conjecture,(epred30_0|c0_1(a321)|~hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_6]), c_0_7])).
cnf(c_0_167,negated_conjecture,(epred30_0|hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_97]), c_0_1])).
cnf(c_0_168,negated_conjecture,(hskp0|~c1_1(a299)|~hskp3), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_34]), c_0_156]), c_0_35])).
cnf(c_0_169,negated_conjecture,(c1_1(a299)|hskp4|~hskp3), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_34]), c_0_153]), c_0_36])).
cnf(c_0_170,negated_conjecture,(epred42_0|~hskp25), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_151]), c_0_38]), c_0_39])).
cnf(c_0_171,negated_conjecture,(epred29_0|hskp3|hskp12|hskp25), inference(spm,[status(thm)],[c_0_157, c_0_40])).
fof(c_0_172, plain, ((~epred41_0<=>![X2]:(((c0_1(X2)|c2_1(X2))|c3_1(X2))|~ndr1_0))), introduced(definition)).
cnf(c_0_173,negated_conjecture,(epred11_0|c2_1(X1)|c1_1(X1)|c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_158])).
fof(c_0_174, plain, ((~epred17_0<=>![X2]:((c0_1(X2)|~c1_1(X2))|~c2_1(X2)))), introduced(definition)).
cnf(c_0_175,negated_conjecture,(c0_1(X1)|hskp0|hskp14|~c3_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_81])])).
cnf(c_0_176,negated_conjecture,(epred49_0|epred9_0|~hskp0), inference(spm,[status(thm)],[c_0_159, c_0_160])).
cnf(c_0_177,negated_conjecture,(hskp0|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_20]), c_0_156]), c_0_21]), c_0_22])).
cnf(c_0_178,negated_conjecture,(epred16_0|hskp12|~hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_10]), c_0_11]), c_0_12])).
cnf(c_0_179,negated_conjecture,(epred45_0|~hskp26), inference(spm,[status(thm)],[c_0_162, c_0_97])).
cnf(c_0_180,negated_conjecture,(epred9_0|~c1_1(a310)|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_137]), c_0_18])).
cnf(c_0_181,negated_conjecture,(epred49_0|c1_1(a310)|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_146]), c_0_18])).
cnf(c_0_182,negated_conjecture,(epred10_0|epred9_0|hskp0|hskp12|hskp25), inference(spm,[status(thm)],[c_0_163, c_0_40])).
cnf(c_0_183,negated_conjecture,(epred42_0|epred9_0|~hskp0), inference(spm,[status(thm)],[c_0_159, c_0_164])).
cnf(c_0_184,negated_conjecture,(epred49_0|~hskp4), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_146]), c_0_43]), c_0_44])).
cnf(c_0_185,negated_conjecture,(epred30_0|hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_167])).
cnf(c_0_186,negated_conjecture,(hskp0|hskp4|~hskp3), inference(spm,[status(thm)],[c_0_168, c_0_169])).
cnf(c_0_187,negated_conjecture,(epred42_0|epred29_0|hskp3|hskp12), inference(spm,[status(thm)],[c_0_170, c_0_171])).
cnf(c_0_188,negated_conjecture,(epred41_0|c3_1(X1)|c2_1(X1)|c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_172]), c_0_81])])).
cnf(c_0_189,negated_conjecture,(hskp0|~c1_1(a310)|~hskp12), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_17]), c_0_156]), c_0_18])).
cnf(c_0_190,negated_conjecture,(epred11_0|c1_1(a310)|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_173]), c_0_18])).
cnf(c_0_191,negated_conjecture,(epred17_0|c0_1(X1)|~c2_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_174])).
cnf(c_0_192,negated_conjecture,(hskp0|hskp14|~c2_1(a336)|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_4]), c_0_13])).
cnf(c_0_193,negated_conjecture,(epred49_0|epred9_0|~hskp11), inference(spm,[status(thm)],[c_0_176, c_0_177])).
cnf(c_0_194,negated_conjecture,(epred16_0|hskp11|hskp12), inference(spm,[status(thm)],[c_0_178, c_0_23])).
cnf(c_0_195,negated_conjecture,(~epred45_0|~epred41_0|~epred9_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_45, c_0_172]), c_0_124]), c_0_122])).
cnf(c_0_196,negated_conjecture,(epred45_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_1]), c_0_149])).
cnf(c_0_197,negated_conjecture,(epred49_0|epred9_0|~hskp12), inference(spm,[status(thm)],[c_0_180, c_0_181])).
cnf(c_0_198,negated_conjecture,(epred42_0|epred10_0|epred9_0|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_182]), c_0_183])).
cnf(c_0_199,negated_conjecture,(~epred42_0|~epred41_0|~epred30_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_46, c_0_172]), c_0_138]), c_0_116])).
cnf(c_0_200,negated_conjecture,(epred49_0|epred30_0), inference(spm,[status(thm)],[c_0_184, c_0_185])).
cnf(c_0_201,negated_conjecture,(epred42_0|epred29_0|hskp0|hskp4|hskp12), inference(spm,[status(thm)],[c_0_186, c_0_187])).
cnf(c_0_202,negated_conjecture,(epred41_0|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_188]), c_0_31]), c_0_32])).
cnf(c_0_203,negated_conjecture,(epred11_0|hskp0|~hskp12), inference(spm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_204,negated_conjecture,(c0_1(a321)|hskp0|hskp14|~hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_6]), c_0_7])).
cnf(c_0_205,negated_conjecture,(epred17_0|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_191]), c_0_48]), c_0_49])).
cnf(c_0_206,negated_conjecture,(hskp0|hskp14|hskp16|~hskp21), inference(spm,[status(thm)],[c_0_192, c_0_88])).
cnf(c_0_207,negated_conjecture,(epred49_0|epred16_0|epred9_0|hskp12), inference(spm,[status(thm)],[c_0_193, c_0_194])).
cnf(c_0_208,negated_conjecture,(~epred41_0|~epred9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_195, c_0_196])])).
cnf(c_0_209,negated_conjecture,(epred49_0|epred42_0|epred10_0|epred9_0), inference(spm,[status(thm)],[c_0_197, c_0_198])).
cnf(c_0_210,negated_conjecture,(epred49_0|~epred42_0|~epred41_0), inference(spm,[status(thm)],[c_0_199, c_0_200])).
cnf(c_0_211,negated_conjecture,(epred11_0|~hskp4), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_173]), c_0_43]), c_0_44])).
cnf(c_0_212,negated_conjecture,(epred42_0|epred29_0|epred9_0|hskp4|hskp12), inference(spm,[status(thm)],[c_0_183, c_0_201])).
cnf(c_0_213,negated_conjecture,(epred11_0|epred9_0|~hskp12), inference(spm,[status(thm)],[c_0_180, c_0_190])).
cnf(c_0_214,negated_conjecture,(epred42_0|epred41_0|epred29_0|hskp4|hskp12), inference(spm,[status(thm)],[c_0_202, c_0_201])).
cnf(c_0_215,negated_conjecture,(epred41_0|epred11_0|~hskp12), inference(spm,[status(thm)],[c_0_202, c_0_203])).
cnf(c_0_216,negated_conjecture,(hskp0|hskp4|hskp14|~hskp16), inference(spm,[status(thm)],[c_0_165, c_0_204])).
cnf(c_0_217,negated_conjecture,(epred17_0|hskp0|hskp16|~hskp21), inference(spm,[status(thm)],[c_0_205, c_0_206])).
cnf(c_0_218,negated_conjecture,(epred49_0|hskp0|~hskp12), inference(spm,[status(thm)],[c_0_189, c_0_181])).
cnf(c_0_219,negated_conjecture,(epred41_0|~hskp11), inference(spm,[status(thm)],[c_0_202, c_0_177])).
cnf(c_0_220,negated_conjecture,(epred49_0|epred16_0|epred9_0), inference(spm,[status(thm)],[c_0_197, c_0_207])).
cnf(c_0_221,negated_conjecture,(~epred11_0|~epred10_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_50, c_0_123]), c_0_158])).
cnf(c_0_222,negated_conjecture,(epred49_0|epred10_0|~epred41_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209]), c_0_210])).
cnf(c_0_223,negated_conjecture,(epred42_0|epred29_0|epred11_0|epred9_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212]), c_0_213])).
cnf(c_0_224,negated_conjecture,(epred42_0|epred41_0|epred29_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_214]), c_0_215])).
fof(c_0_225, plain, ((~epred35_0<=>![X1]:((~c0_1(X1)|~c1_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_226,negated_conjecture,(epred17_0|hskp0|hskp4|~hskp16), inference(spm,[status(thm)],[c_0_205, c_0_216])).
cnf(c_0_227,negated_conjecture,(epred17_0|hskp0|hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_97]), c_0_1])).
cnf(c_0_228,negated_conjecture,(epred49_0|epred41_0|~hskp12), inference(spm,[status(thm)],[c_0_202, c_0_218])).
cnf(c_0_229,negated_conjecture,(epred41_0|epred16_0|hskp12), inference(spm,[status(thm)],[c_0_219, c_0_194])).
cnf(c_0_230,negated_conjecture,(epred49_0|epred16_0|~epred41_0), inference(spm,[status(thm)],[c_0_208, c_0_220])).
cnf(c_0_231,negated_conjecture,(epred49_0|~epred41_0|~epred11_0), inference(spm,[status(thm)],[c_0_221, c_0_222])).
cnf(c_0_232,negated_conjecture,(epred42_0|epred29_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_223]), c_0_224])).
cnf(c_0_233,negated_conjecture,(c1_1(X1)|hskp3|hskp27|~c3_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_81])])).
cnf(c_0_234,negated_conjecture,(~epred49_0|~epred35_0|~epred28_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_52, c_0_134]), c_0_86]), c_0_225])).
cnf(c_0_235,negated_conjecture,(epred17_0|hskp0|hskp4), inference(spm,[status(thm)],[c_0_226, c_0_227])).
cnf(c_0_236,negated_conjecture,(~epred17_0|~epred16_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_53, c_0_147]), c_0_174])).
cnf(c_0_237,negated_conjecture,(epred49_0|epred16_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_230])).
cnf(c_0_238,negated_conjecture,(~epred30_0|~epred29_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_54, c_0_133]), c_0_116])).
cnf(c_0_239,negated_conjecture,(epred49_0|epred29_0|~epred41_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_232]), c_0_210])).
cnf(c_0_240,negated_conjecture,(epred35_0|~c3_1(X1)|~c1_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_225])).
cnf(c_0_241,negated_conjecture,(hskp3|hskp27|~c2_1(a336)|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_4]), c_0_5])).
cnf(c_0_242,negated_conjecture,(hskp3|hskp27|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_6]), c_0_7]), c_0_8])).
cnf(c_0_243,negated_conjecture,(epred31_0|~epred49_0|~epred35_0), inference(spm,[status(thm)],[c_0_234, c_0_126])).
cnf(c_0_244,negated_conjecture,(epred32_0|~epred49_0|~epred35_0), inference(spm,[status(thm)],[c_0_234, c_0_127])).
cnf(c_0_245,negated_conjecture,(epred41_0|epred17_0|hskp4), inference(spm,[status(thm)],[c_0_202, c_0_235])).
cnf(c_0_246,negated_conjecture,(epred49_0|~epred17_0), inference(spm,[status(thm)],[c_0_236, c_0_237])).
cnf(c_0_247,negated_conjecture,(epred49_0|~epred41_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239]), c_0_200])).
cnf(c_0_248,negated_conjecture,(epred35_0|~c1_1(a334)|~hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_55]), c_0_56])).
cnf(c_0_249,negated_conjecture,(c1_1(a334)|hskp4|~hskp27), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_55]), c_0_153]), c_0_56])).
cnf(c_0_250,negated_conjecture,(hskp3|hskp27|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_88]), c_0_242])).
cnf(c_0_251,negated_conjecture,(~epred49_0|~epred35_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_243]), c_0_244])).
cnf(c_0_252,negated_conjecture,(epred49_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_245]), c_0_246]), c_0_247])).
cnf(c_0_253,negated_conjecture,(c2_1(X1)|c0_1(X1)|hskp10|hskp11|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_81])])).
cnf(c_0_254,negated_conjecture,(epred35_0|~c1_1(a299)|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_34]), c_0_36])).
cnf(c_0_255,negated_conjecture,(epred35_0|hskp4|~hskp27), inference(spm,[status(thm)],[c_0_248, c_0_249])).
cnf(c_0_256,negated_conjecture,(hskp3|hskp27|~hskp26), inference(spm,[status(thm)],[c_0_250, c_0_97])).
cnf(c_0_257,negated_conjecture,(~epred35_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_251, c_0_252])])).
cnf(c_0_258,negated_conjecture,(c2_1(a336)|hskp10|hskp11|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_4]), c_0_13])).
cnf(c_0_259,negated_conjecture,(epred35_0|hskp4|~hskp3), inference(spm,[status(thm)],[c_0_254, c_0_169])).
cnf(c_0_260,negated_conjecture,(hskp3|hskp4|~hskp26), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_256]), c_0_257])).
cnf(c_0_261,negated_conjecture,(hskp3|hskp4|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_242]), c_0_257])).
cnf(c_0_262,negated_conjecture,(epred30_0|~c2_1(a308)|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_58]), c_0_59])).
cnf(c_0_263,negated_conjecture,(epred9_0|c2_1(a308)|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_137]), c_0_60])).
cnf(c_0_264,negated_conjecture,(epred30_0|hskp10|hskp11|~hskp21), inference(spm,[status(thm)],[c_0_142, c_0_258])).
cnf(c_0_265,negated_conjecture,(epred11_0|c1_1(a294)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_173]), c_0_32])).
cnf(c_0_266,negated_conjecture,(hskp4|~hskp26), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_260]), c_0_257])).
cnf(c_0_267,negated_conjecture,(hskp4|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_261]), c_0_257])).
cnf(c_0_268,negated_conjecture,(epred30_0|epred9_0|~hskp10), inference(spm,[status(thm)],[c_0_262, c_0_263])).
cnf(c_0_269,negated_conjecture,(epred30_0|hskp10|hskp11|~hskp26), inference(spm,[status(thm)],[c_0_264, c_0_97])).
cnf(c_0_270,negated_conjecture,(hskp10|hskp11|~hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_17]), c_0_19]), c_0_18])).
cnf(c_0_271,negated_conjecture,(epred11_0|epred9_0|~hskp0), inference(spm,[status(thm)],[c_0_159, c_0_265])).
cnf(c_0_272,negated_conjecture,(hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_1]), c_0_267])).
cnf(c_0_273,negated_conjecture,(c1_1(X1)|c0_1(X1)|hskp0|hskp2|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_81])])).
cnf(c_0_274,negated_conjecture,(epred41_0|c2_1(X1)|c0_1(X1)|~c1_1(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_188]), c_0_156]), c_0_202])).
cnf(c_0_275,negated_conjecture,(epred30_0|epred9_0|hskp11|~hskp26), inference(spm,[status(thm)],[c_0_268, c_0_269])).
cnf(c_0_276,negated_conjecture,(epred30_0|epred9_0|hskp11|~hskp12), inference(spm,[status(thm)],[c_0_268, c_0_270])).
cnf(c_0_277,negated_conjecture,(epred11_0|epred9_0|~hskp11), inference(spm,[status(thm)],[c_0_271, c_0_177])).
cnf(c_0_278,negated_conjecture,(epred17_0|~c2_1(a297)|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_191]), c_0_63])).
cnf(c_0_279,negated_conjecture,(epred9_0|c2_1(a297)|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_137]), c_0_63])).
cnf(c_0_280,negated_conjecture,(~c0_1(a300)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_272])])).
cnf(c_0_281,negated_conjecture,(epred42_0|c1_1(X1)|c0_1(X1)|hskp0|hskp2), inference(spm,[status(thm)],[c_0_273, c_0_151])).
cnf(c_0_282,negated_conjecture,(~c1_1(a300)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_272])])).
cnf(c_0_283,negated_conjecture,(epred41_0|c2_1(a297)|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_188]), c_0_62])).
cnf(c_0_284,negated_conjecture,(epred41_0|c2_1(X1)|c0_1(X1)|hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_188]), c_0_202]), c_0_274])).
cnf(c_0_285,negated_conjecture,(~c2_1(a300)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_272])])).
cnf(c_0_286,negated_conjecture,(epred30_0|epred9_0|hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_23]), c_0_276])).
cnf(c_0_287,negated_conjecture,(epred30_0|epred2_0|~hskp16), inference(spm,[status(thm)],[c_0_109, c_0_166])).
cnf(c_0_288,negated_conjecture,(c2_1(a308)|hskp0|~hskp10), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_58]), c_0_156]), c_0_60])).
cnf(c_0_289,negated_conjecture,(epred41_0|c2_1(X1)|c0_1(X1)|hskp10|hskp11), inference(spm,[status(thm)],[c_0_253, c_0_188])).
cnf(c_0_290,negated_conjecture,(epred16_0|epred11_0|epred9_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_194]), c_0_213])).
cnf(c_0_291,negated_conjecture,(epred41_0|epred16_0|epred11_0), inference(spm,[status(thm)],[c_0_215, c_0_229])).
cnf(c_0_292,negated_conjecture,(epred17_0|epred9_0|~hskp2), inference(spm,[status(thm)],[c_0_278, c_0_279])).
cnf(c_0_293,negated_conjecture,(epred42_0|hskp0|hskp2), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281]), c_0_282])).
cnf(c_0_294,negated_conjecture,(epred41_0|epred17_0|~hskp2), inference(spm,[status(thm)],[c_0_278, c_0_283])).
cnf(c_0_295,negated_conjecture,(epred41_0|hskp2), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_284]), c_0_285])).
cnf(c_0_296,negated_conjecture,(epred30_0|epred9_0|epred1_0), inference(spm,[status(thm)],[c_0_117, c_0_286])).
cnf(c_0_297,negated_conjecture,(epred30_0|epred2_0), inference(spm,[status(thm)],[c_0_287, c_0_167])).
cnf(c_0_298,negated_conjecture,(epred30_0|hskp0|~hskp10), inference(spm,[status(thm)],[c_0_262, c_0_288])).
cnf(c_0_299,negated_conjecture,(epred41_0|hskp10), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_289]), c_0_285]), c_0_219])).
cnf(c_0_300,negated_conjecture,(~epred41_0|~epred17_0|~epred11_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_65, c_0_172]), c_0_158]), c_0_174])).
cnf(c_0_301,negated_conjecture,(epred16_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_290]), c_0_291])).
cnf(c_0_302,negated_conjecture,(epred42_0|epred17_0|epred9_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293]), c_0_183])).
cnf(c_0_303,negated_conjecture,(epred41_0|epred17_0), inference(spm,[status(thm)],[c_0_294, c_0_295])).
fof(c_0_304, plain, ((~epred21_0<=>![X2]:((c0_1(X2)|c3_1(X2))|~c2_1(X2)))), introduced(definition)).
fof(c_0_305, plain, ((~epred47_0<=>![X1]:((c0_1(X1)|~c1_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_306,negated_conjecture,(epred30_0|epred9_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_296]), c_0_297])).
cnf(c_0_307,negated_conjecture,(epred41_0|epred30_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299]), c_0_202])).
cnf(c_0_308,negated_conjecture,(~epred41_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_301]), c_0_236])).
cnf(c_0_309,negated_conjecture,(epred42_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_302]), c_0_303])).
cnf(c_0_310,negated_conjecture,(epred21_0|c3_1(X1)|c0_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_304])).
cnf(c_0_311,negated_conjecture,(epred47_0|c0_1(X1)|~c3_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_305])).
cnf(c_0_312,negated_conjecture,(epred41_0|c2_1(X1)|c1_1(X1)|c0_1(X1)|hskp16), inference(spm,[status(thm)],[c_0_83, c_0_188])).
cnf(c_0_313,negated_conjecture,(epred30_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_306]), c_0_307])).
cnf(c_0_314,negated_conjecture,(epred42_0|~epred41_0), inference(spm,[status(thm)],[c_0_308, c_0_309])).
cnf(c_0_315,negated_conjecture,(epred21_0|~c2_1(a297)|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_310]), c_0_62])).
cnf(c_0_316,negated_conjecture,(epred47_0|~hskp10), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_58]), c_0_60]), c_0_59])).
fof(c_0_317, plain, ((~epred20_0<=>![X1]:((((hskp8|c3_1(X1))|~ndr1_0)|~c1_1(X1))|~c2_1(X1)))), introduced(definition)).
cnf(c_0_318,negated_conjecture,(epred41_0|hskp16), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_312]), c_0_285]), c_0_282])).
cnf(c_0_319,negated_conjecture,(~epred41_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_199, c_0_313])]), c_0_314])).
cnf(c_0_320,negated_conjecture,(epred41_0|epred21_0|~hskp2), inference(spm,[status(thm)],[c_0_315, c_0_283])).
cnf(c_0_321,negated_conjecture,(epred47_0|epred41_0|c2_1(X1)|c0_1(X1)|hskp11), inference(spm,[status(thm)],[c_0_316, c_0_289])).
fof(c_0_322, plain, ((~epred38_0<=>![X2]:((((hskp11|c3_1(X2))|~ndr1_0)|~c0_1(X2))|~c2_1(X2)))), introduced(definition)).
cnf(c_0_323,negated_conjecture,(epred20_0|c3_1(X1)|hskp8|~c2_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_317]), c_0_81])])).
cnf(c_0_324,negated_conjecture,(c1_1(X1)|hskp17|hskp18|~c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_81])])).
cnf(c_0_325,negated_conjecture,(hskp16), inference(sr,[status(thm)],[c_0_318, c_0_319])).
cnf(c_0_326,negated_conjecture,(epred41_0|epred21_0), inference(spm,[status(thm)],[c_0_320, c_0_295])).
cnf(c_0_327,negated_conjecture,(epred47_0|epred41_0|c2_1(X1)|c0_1(X1)), inference(spm,[status(thm)],[c_0_219, c_0_321])).
cnf(c_0_328,negated_conjecture,(epred38_0|c3_1(X1)|hskp11|~c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_322]), c_0_81])])).
cnf(c_0_329,negated_conjecture,(epred20_0|hskp8|~c1_1(a329)|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_67]), c_0_68])).
cnf(c_0_330,negated_conjecture,(epred32_0|c1_1(a329)|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_91]), c_0_67])).
cnf(c_0_331,negated_conjecture,(epred32_0|~hskp8), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_91]), c_0_71]), c_0_72])).
cnf(c_0_332,negated_conjecture,(epred32_0|c1_1(X1)|hskp17|hskp18|~c2_1(X1)), inference(spm,[status(thm)],[c_0_324, c_0_91])).
cnf(c_0_333,negated_conjecture,(c2_1(a321)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_325])])).
cnf(c_0_334,negated_conjecture,(~c1_1(a321)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_325])])).
cnf(c_0_335,negated_conjecture,(~epred21_0|~epred20_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_73, c_0_317]), c_0_304])).
cnf(c_0_336,negated_conjecture,(epred21_0), inference(sr,[status(thm)],[c_0_326, c_0_319])).
cnf(c_0_337,negated_conjecture,(epred2_0|~c0_1(a321)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_325])])).
cnf(c_0_338,negated_conjecture,(epred47_0|epred41_0|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_327]), c_0_18])).
fof(c_0_339, plain, ((~epred46_0<=>![X2]:(((c0_1(X2)|c3_1(X2))|~ndr1_0)|~c1_1(X2)))), introduced(definition)).
cnf(c_0_340,negated_conjecture,(c1_1(X1)|hskp11|hskp20|~c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_81])])).
cnf(c_0_341,negated_conjecture,(epred38_0|hskp11|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_75]), c_0_76]), c_0_77])).
cnf(c_0_342,negated_conjecture,(epred32_0|epred20_0|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_330]), c_0_331])).
cnf(c_0_343,negated_conjecture,(epred32_0|hskp17|hskp18), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_333]), c_0_334])).
cnf(c_0_344,negated_conjecture,(~epred20_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_335, c_0_336])])).
cnf(c_0_345,negated_conjecture,(~epred38_0|~epred2_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_78, c_0_322]), c_0_94])).
cnf(c_0_346,negated_conjecture,(epred32_0|epred2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_91]), c_0_333])]), c_0_334])).
cnf(c_0_347,negated_conjecture,(epred47_0|epred41_0|epred16_0), inference(spm,[status(thm)],[c_0_338, c_0_229])).
cnf(c_0_348,negated_conjecture,(epred46_0|c3_1(X1)|c0_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_339]), c_0_81])])).
cnf(c_0_349,negated_conjecture,(hskp11|hskp20|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_14]), c_0_15]), c_0_16])).
cnf(c_0_350,negated_conjecture,(epred41_0|epred38_0|~hskp17), inference(spm,[status(thm)],[c_0_219, c_0_341])).
cnf(c_0_351,negated_conjecture,(epred32_0|hskp17), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_343]), c_0_344])).
cnf(c_0_352,negated_conjecture,(epred32_0|~epred38_0), inference(spm,[status(thm)],[c_0_345, c_0_346])).
cnf(c_0_353,negated_conjecture,(epred47_0|epred41_0|~hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_327]), c_0_44])).
cnf(c_0_354,negated_conjecture,(epred47_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_347]), c_0_308])).
cnf(c_0_355,negated_conjecture,(epred46_0|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_348]), c_0_63]), c_0_62])).
cnf(c_0_356,negated_conjecture,(epred31_0|hskp11|hskp20|~hskp16), inference(spm,[status(thm)],[c_0_349, c_0_89])).
cnf(c_0_357,negated_conjecture,(epred32_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_351]), c_0_319]), c_0_352])).
cnf(c_0_358,negated_conjecture,(epred47_0|epred41_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_245]), c_0_354])).
cnf(c_0_359,negated_conjecture,(epred46_0|epred41_0), inference(spm,[status(thm)],[c_0_355, c_0_295])).
cnf(c_0_360,negated_conjecture,(epred31_0|hskp11|hskp20), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_356, c_0_325])])).
cnf(c_0_361,negated_conjecture,(~epred31_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_139, c_0_357])])).
cnf(c_0_362,negated_conjecture,(epred11_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_211, c_0_272])])).
cnf(c_0_363,negated_conjecture,(~epred47_0|~epred46_0|~epred9_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_79, c_0_339]), c_0_124]), c_0_305])).
cnf(c_0_364,negated_conjecture,(epred47_0), inference(sr,[status(thm)],[c_0_358, c_0_319])).
cnf(c_0_365,negated_conjecture,(epred46_0), inference(sr,[status(thm)],[c_0_359, c_0_319])).
cnf(c_0_366,negated_conjecture,(hskp11|hskp20), inference(sr,[status(thm)],[c_0_360, c_0_361])).
cnf(c_0_367,negated_conjecture,(~epred10_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_221, c_0_362])])).
cnf(c_0_368,negated_conjecture,(~epred9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_363, c_0_364])]), c_0_365])])).
cnf(c_0_369,negated_conjecture,(hskp0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_366]), c_0_367]), c_0_368]), c_0_177])).
cnf(c_0_370,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_202, c_0_369])]), c_0_319]), ['proof']).
# SZS output end CNFRefutation.
EOF
