% Problem    : SYO594-1 : TPTP v6.1.0. Bugfixed v5.5.1.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n131.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:02:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___092_C01_F1_AE_CS_SP_PS_CO_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___107_C37_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(esk297_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk296_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_838)).
cnf(c_0_1,axiom,(esk296_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk184_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_293)).
cnf(c_0_2,axiom,(~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk297_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1196)).
cnf(c_0_3,axiom,(esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk282_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_995)).
cnf(c_0_4,axiom,(esk282_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk281_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk190_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_223)).
cnf(c_0_5,axiom,(~p(X1)|~p(X2)|~esk280_23(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X1,X2,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1000)).
cnf(c_0_6,axiom,(esk280_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk191_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_507)).
cnf(c_0_7,axiom,(~esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk192_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1203)).
cnf(c_0_8,axiom,(esk192_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_771)).
cnf(c_0_9,axiom,(esk184_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_982)).
cnf(c_0_10,axiom,(p(X17)|~esk281_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_407)).
cnf(c_0_11,axiom,(esk191_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_649)).
cnf(c_0_12,axiom,(esk190_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_743)).
cnf(c_0_13,axiom,(~p(X1)|~esk300_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X1,X19,X20,X21,X22,X23)|~esk160_16(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~p(X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_964)).
cnf(c_0_14,axiom,(esk300_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk173_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1099)).
cnf(c_0_15,axiom,(esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk303_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1058)).
cnf(c_0_16,axiom,(esk303_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk302_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk169_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_842)).
cnf(c_0_17,axiom,(~p(X1)|~p(X2)|~esk301_23(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X2,X1,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_821)).
cnf(c_0_18,axiom,(esk301_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk170_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_133)).
cnf(c_0_19,axiom,(esk173_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk172_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_707)).
cnf(c_0_20,axiom,(p(X16)|~p(X1)|~esk302_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X1,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1065)).
cnf(c_0_21,axiom,(esk170_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_168)).
cnf(c_0_22,axiom,(esk323_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk322_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1162)).
cnf(c_0_23,axiom,(esk322_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk161_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_799)).
cnf(c_0_24,axiom,(p(true)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule_true)).
cnf(c_0_25,axiom,(esk313_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk312_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_158)).
cnf(c_0_26,axiom,(esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk312_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk164_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1112)).
cnf(c_0_27,axiom,(~esk171_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1052)).
cnf(c_0_28,axiom,(esk171_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_935)).
cnf(c_0_29,axiom,(esk169_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_144)).
cnf(c_0_30,axiom,(~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk323_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_132)).
cnf(c_0_31,axiom,(p(X1)|~esk328_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X13)|~p(X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_276)).
cnf(c_0_32,axiom,(esk328_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk145_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_459)).
cnf(c_0_33,axiom,(esk351_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk350_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_896)).
cnf(c_0_34,axiom,(esk350_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk135_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1092)).
cnf(c_0_35,axiom,(~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk313_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_357)).
cnf(c_0_36,axiom,(esk161_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_810)).
cnf(c_0_37,axiom,(~esk326_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X16)|~p(X15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1236)).
cnf(c_0_38,axiom,(esk326_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk149_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_296)).
cnf(c_0_39,axiom,(~esk147_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_207)).
cnf(c_0_40,axiom,(esk147_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_695)).
cnf(c_0_41,axiom,(esk145_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1182)).
cnf(c_0_42,axiom,(~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk351_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_997)).
cnf(c_0_43,axiom,(esk309_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk308_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk130_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_522)).
cnf(c_0_44,axiom,(esk308_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk166_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_986)).
cnf(c_0_45,axiom,(esk149_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk324_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk326_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk325_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_441)).
cnf(c_0_46,axiom,(p(X1)|esk324_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X1,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_53)).
cnf(c_0_47,axiom,(esk326_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_679)).
cnf(c_0_48,axiom,(esk172_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_870)).
cnf(c_0_49,axiom,(esk164_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_240)).
cnf(c_0_50,axiom,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk305_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1010)).
cnf(c_0_51,axiom,(esk305_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk304_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk168_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_380)).
cnf(c_0_52,axiom,(esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk305_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_174)).
cnf(c_0_53,axiom,(esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk149_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1151)).
cnf(c_0_54,axiom,(esk135_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_413)).
cnf(c_0_55,axiom,(p(X14)|~p(X1)|~p(X2)|~esk109_11(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~esk354_23(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X2,X1,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_683)).
cnf(c_0_56,axiom,(esk354_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_509)).
cnf(c_0_57,axiom,(~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk309_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_315)).
cnf(c_0_58,axiom,(esk325_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_667)).
cnf(c_0_59,axiom,(p(X1)|~esk304_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X1,X17,X18,X19,X20,X21,X22,X23)|~p(X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_535)).
cnf(c_0_60,axiom,(esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk352_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk353_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk355_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk354_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_594)).
cnf(c_0_61,axiom,(p(X1)|esk353_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1155)).
cnf(c_0_62,axiom,(p(X1)|esk355_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1008)).
cnf(c_0_63,axiom,(esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1056)).
cnf(c_0_64,axiom,(esk364_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk363_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk80_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1045)).
cnf(c_0_65,axiom,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk363_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk114_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_245)).
cnf(c_0_66,axiom,(p(X2)|~p(X1)|~esk357_23(X3,X4,X5,X6,X7,X8,X9,X10,X11,X2,X1,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_33)).
cnf(c_0_67,axiom,(esk358_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk357_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk117_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1125)).
cnf(c_0_68,axiom,(esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk378_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk377_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk376_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_194)).
cnf(c_0_69,axiom,(p(X1)|esk378_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X1,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_637)).
cnf(c_0_70,axiom,(esk376_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X10)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_309)).
cnf(c_0_71,axiom,(esk166_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_461)).
cnf(c_0_72,axiom,(p(X16)|~p(X1)|~esk324_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X1,X16,X17,X18,X19,X20,X21,X22,X23)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_430)).
cnf(c_0_73,axiom,(esk324_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk149_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_166)).
cnf(c_0_74,axiom,(esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk149_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1135)).
cnf(c_0_75,axiom,(esk168_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_675)).
cnf(c_0_76,axiom,(esk352_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X1,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_723)).
cnf(c_0_77,axiom,(~esk364_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_630)).
cnf(c_0_78,axiom,(esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk358_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_632)).
cnf(c_0_79,axiom,(p(X1)|~esk83_9(X2,X3,X4,X5,X6,X7,X8,X9,X10)|~esk377_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X1,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_806)).
cnf(c_0_80,axiom,(esk377_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_76)).
cnf(c_0_81,axiom,(esk378_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1085)).
cnf(c_0_82,axiom,(esk377_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_790)).
cnf(c_0_83,axiom,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk130_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk347_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_577)).
cnf(c_0_84,axiom,(esk347_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk346_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1044)).
cnf(c_0_85,axiom,(esk349_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk100_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk345_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_581)).
cnf(c_0_86,axiom,(esk346_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk345_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk136_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1223)).
cnf(c_0_87,axiom,(esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk330_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_678)).
cnf(c_0_88,axiom,(esk330_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk329_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk144_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_359)).
cnf(c_0_89,axiom,(~p(X1)|~p(X2)|~p(X3)|~esk327_23(X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X3,X1,X2,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_814)).
cnf(c_0_90,axiom,(esk327_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk146_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_906)).
cnf(c_0_91,axiom,(esk354_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_250)).
cnf(c_0_92,axiom,(esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1047)).
cnf(c_0_93,axiom,(esk114_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_616)).
cnf(c_0_94,axiom,(esk117_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_573)).
cnf(c_0_95,axiom,(esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_893)).
cnf(c_0_96,axiom,(esk80_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk391_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_563)).
cnf(c_0_97,axiom,(esk391_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk390_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_725)).
cnf(c_0_98,axiom,(esk393_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk389_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_164)).
cnf(c_0_99,axiom,(esk389_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk390_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_260)).
cnf(c_0_100,axiom,(esk376_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_688)).
cnf(c_0_101,axiom,(esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_205)).
cnf(c_0_102,axiom,(p(X2)|~p(X1)|~esk329_23(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X2,X1,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_67)).
cnf(c_0_103,axiom,(esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_958)).
cnf(c_0_104,axiom,(esk146_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1061)).
cnf(c_0_105,axiom,(esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1060)).
cnf(c_0_106,axiom,(esk377_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X1,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_900)).
cnf(c_0_107,axiom,(esk73_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk396_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk397_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk398_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_511)).
cnf(c_0_108,axiom,(esk396_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1030)).
cnf(c_0_109,axiom,(p(X1)|esk398_23(X2,X3,X4,X5,X6,X7,X8,X1,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_264)).
cnf(c_0_110,axiom,(~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk349_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1177)).
cnf(c_0_111,axiom,(~esk355_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X14)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X13)|~p(X12)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_836)).
cnf(c_0_112,axiom,(esk355_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_760)).
cnf(c_0_113,axiom,(esk336_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk335_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk105_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_143)).
cnf(c_0_114,axiom,(esk335_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk141_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_658)).
cnf(c_0_115,axiom,(~esk393_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_756)).
cnf(c_0_116,axiom,(esk397_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1081)).
cnf(c_0_117,axiom,(esk144_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_496)).
cnf(c_0_118,axiom,(~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk336_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_379)).
cnf(c_0_119,axiom,(~esk398_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X8)|~p(X9)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_883)).
cnf(c_0_120,axiom,(esk398_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk73_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_602)).
cnf(c_0_121,axiom,(esk72_7(X1,X2,X3,X4,X5,X6,X7)|~esk73_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_500)).
cnf(c_0_122,axiom,(~esk379_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X9)|~p(X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_543)).
cnf(c_0_123,axiom,(esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk379_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk92_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_23)).
cnf(c_0_124,axiom,(esk136_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_21)).
cnf(c_0_125,axiom,(esk141_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_344)).
cnf(c_0_126,axiom,(esk373_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk367_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_472)).
cnf(c_0_127,axiom,(esk367_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_529)).
cnf(c_0_128,axiom,(esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_613)).
cnf(c_0_129,axiom,(esk73_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk72_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1064)).
cnf(c_0_130,axiom,(esk92_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_931)).
cnf(c_0_131,axiom,(~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk373_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_238)).
cnf(c_0_132,axiom,(esk100_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1027)).
cnf(c_0_133,axiom,(esk105_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk370_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_571)).
cnf(c_0_134,axiom,(esk370_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_425)).
cnf(c_0_135,axiom,(esk413_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk409_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk35_5(X1,X2,X3,X4,X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_733)).
cnf(c_0_136,axiom,(esk409_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk410_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk64_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_409)).
cnf(c_0_137,axiom,(esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk410_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_332)).
cnf(c_0_138,axiom,(esk51_7(X1,X2,X3,X4,X5,X6,X7)|~esk418_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk417_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk416_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_717)).
cnf(c_0_139,axiom,(p(X1)|esk417_23(X2,X3,X4,X5,X6,X7,X1,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_984)).
cnf(c_0_140,axiom,(esk418_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_233)).
cnf(c_0_141,axiom,(esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_578)).
cnf(c_0_142,axiom,(~esk41_5(X1,X2,X3,X4,X5)|~esk413_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1200)).
cnf(c_0_143,axiom,(p(X1)|esk416_23(X2,X3,X4,X5,X6,X1,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1102)).
cnf(c_0_144,axiom,(~esk41_5(X1,X2,X3,X4,X5)|~p(X6)|~p(X7)|~esk418_23(X1,X2,X3,X4,X5,X7,X6,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_553)).
cnf(c_0_145,axiom,(esk418_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk51_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_335)).
cnf(c_0_146,axiom,(esk50_5(X1,X2,X3,X4,X5)|~esk51_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_226)).
cnf(c_0_147,axiom,(~p(X1)|~esk399_23(X2,X3,X4,X5,X6,X1,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_278)).
cnf(c_0_148,axiom,(esk399_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk70_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_868)).
cnf(c_0_149,axiom,(esk30_5(X1,X2,X3,X4,X5)|~esk436_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk435_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk434_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_272)).
cnf(c_0_150,axiom,(p(X1)|esk434_23(X2,X3,X4,X1,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1021)).
cnf(c_0_151,axiom,(esk436_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_764)).
cnf(c_0_152,axiom,(esk64_7(X1,X2,X3,X4,X5,X6,X7)|~esk52_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_600)).
cnf(c_0_153,axiom,(esk51_7(X1,X2,X3,X4,X5,X6,X7)|esk52_7(X1,X2,X3,X4,X5,X6,X7)|~esk50_5(X1,X2,X3,X4,X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_652)).
cnf(c_0_154,axiom,(esk70_7(X1,X2,X3,X4,X5,X6,X7)|~esk52_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_501)).
cnf(c_0_155,axiom,(~p(X1)|~esk419_23(X2,X3,X4,X1,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_808)).
cnf(c_0_156,axiom,(esk41_5(X1,X2,X3,X4,X5)|esk419_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk48_5(X1,X2,X3,X4,X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_373)).
cnf(c_0_157,axiom,(~esk23_3(X1,X2,X3)|~p(X4)|~p(X5)|~esk436_23(X1,X2,X3,X5,X4,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_292)).
cnf(c_0_158,axiom,(esk436_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk30_5(X1,X2,X3,X4,X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_759)).
cnf(c_0_159,axiom,(esk435_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_47)).
cnf(c_0_160,axiom,(esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk358_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_597)).
cnf(c_0_161,axiom,(esk48_5(X1,X2,X3,X4,X5)|~esk31_5(X1,X2,X3,X4,X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_615)).
cnf(c_0_162,axiom,(esk31_5(X1,X2,X3,X4,X5)|esk30_5(X1,X2,X3,X4,X5)|~esk29_3(X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1040)).
cnf(c_0_163,axiom,(esk29_3(X1,X2,X3)|~esk30_5(X1,X2,X3,X4,X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_598)).
cnf(c_0_164,axiom,(esk433_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk429_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk17_3(X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_584)).
cnf(c_0_165,axiom,(esk429_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk430_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk42_5(X1,X2,X3,X4,X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_187)).
cnf(c_0_166,axiom,(esk35_5(X1,X2,X3,X4,X5)|~esk430_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_819)).
cnf(c_0_167,axiom,(esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk390_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_625)).
cnf(c_0_168,axiom,(~esk433_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk23_3(X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_567)).
cnf(c_0_169,axiom,(~esk10_1(X1)|~esk446_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_921)).
cnf(c_0_170,axiom,(esk446_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk7_1(X1)|~esk442_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_112)).
cnf(c_0_171,axiom,(esk42_5(X1,X2,X3,X4,X5)|~esk31_5(X1,X2,X3,X4,X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_708)).
cnf(c_0_172,axiom,(esk443_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk442_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk24_3(X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_218)).
cnf(c_0_173,axiom,(esk13_3(X1,X2,X3)|~esk447_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk448_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1166)).
cnf(c_0_174,axiom,(esk447_23(X2,X1,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_826)).
cnf(c_0_175,axiom,(esk17_3(X1,X2,X3)|~esk443_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_403)).
cnf(c_0_176,axiom,(~esk437_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X2)|~p(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_382)).
cnf(c_0_177,axiom,(esk437_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk23_3(X1,X2,X3)|~esk27_3(X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_224)).
cnf(c_0_178,axiom,(~esk10_1(X1)|~esk448_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X3)|~p(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1229)).
cnf(c_0_179,axiom,(esk448_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk13_3(X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_172)).
cnf(c_0_180,axiom,(p(X1)|esk448_23(X2,X3,X1,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_565)).
cnf(c_0_181,axiom,(esk3_0|~esk1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_14)).
cnf(c_0_182,axiom,(esk1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_873)).
cnf(c_0_183,axiom,(esk27_3(X1,X2,X3)|~esk14_3(X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1234)).
cnf(c_0_184,axiom,(esk14_3(X1,X2,X3)|esk13_3(X1,X2,X3)|~esk12_1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_904)).
cnf(c_0_185,axiom,(esk12_1(X1)|~esk13_3(X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_71)).
cnf(c_0_186,axiom,(esk4_1(X1)|esk5_1(X1)|~esk3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1163)).
cnf(c_0_187,axiom,(esk24_3(X1,X2,X3)|~esk14_3(X1,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_815)).
cnf(c_0_188,axiom,(esk2_0|~esk1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_635)).
cnf(c_0_189,axiom,(esk10_1(X1)|~esk449_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_361)).
cnf(c_0_190,axiom,(esk449_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X1)|~esk11_1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1048)).
cnf(c_0_191,axiom,(esk11_1(X1)|~esk5_1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_855)).
cnf(c_0_192,axiom,(~p(X1)|~esk4_1(X1)|~esk2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_674)).
cnf(c_0_193,axiom,(p(X1)|esk4_1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_718)).
cnf(c_0_194,axiom,(esk375_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk374_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_462)).
cnf(c_0_195,axiom,(esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk374_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk111_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1220)).
cnf(c_0_196,axiom,(esk7_1(X1)|~esk449_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1041)).
cnf(c_0_197,axiom,(~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk375_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_6)).
cnf(c_0_198,axiom,(esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk360_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_681)).
cnf(c_0_199,axiom,(esk359_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk360_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk116_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_898)).
cnf(c_0_200,axiom,(~p(X1)|~esk356_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X1,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_178)).
cnf(c_0_201,axiom,(esk356_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk118_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_700)).
cnf(c_0_202,axiom,(~esk119_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_633)).
cnf(c_0_203,axiom,(esk119_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_142)).
cnf(c_0_204,axiom,(esk111_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_592)).
cnf(c_0_205,axiom,(p(X11)|~esk359_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X10)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_436)).
cnf(c_0_206,axiom,(esk118_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_668)).
cnf(c_0_207,axiom,(esk116_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_923)).
cnf(c_0_208,axiom,(~p(X1)|~esk84_9(X2,X3,X4,X5,X6,X7,X8,X9,X10)|~esk378_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X1,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_789)).
cnf(c_0_209,axiom,(esk378_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_353)).
cnf(c_0_210,axiom,(esk386_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk385_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk58_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_338)).
cnf(c_0_211,axiom,(esk385_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk88_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_734)).
cnf(c_0_212,axiom,(~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk386_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_549)).
cnf(c_0_213,axiom,(esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_965)).
cnf(c_0_214,axiom,(esk88_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_416)).
cnf(c_0_215,axiom,(esk61_7(X1,X2,X3,X4,X5,X6,X7)|esk58_7(X1,X2,X3,X4,X5,X6,X7)|~esk411_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_493)).
cnf(c_0_216,axiom,(esk411_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk410_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_390)).
cnf(c_0_217,axiom,(esk344_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk101_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk343_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_37)).
cnf(c_0_218,axiom,(esk343_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk137_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_962)).
cnf(c_0_219,axiom,(esk334_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk333_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_788)).
cnf(c_0_220,axiom,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk333_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk142_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_827)).
cnf(c_0_221,axiom,(~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk344_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_754)).
cnf(c_0_222,axiom,(~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk334_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_70)).
cnf(c_0_223,axiom,(esk137_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1152)).
cnf(c_0_224,axiom,(esk142_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_644)).
cnf(c_0_225,axiom,(~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk334_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_3)).
cnf(c_0_226,axiom,(esk395_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk394_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk55_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_183)).
cnf(c_0_227,axiom,(esk394_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk75_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk85_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_417)).
cnf(c_0_228,axiom,(~esk395_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk61_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1114)).
cnf(c_0_229,axiom,(~esk93_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk75_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk76_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_777)).
cnf(c_0_230,axiom,(esk93_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_849)).
cnf(c_0_231,axiom,(esk85_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_301)).
cnf(c_0_232,axiom,(p(X9)|~p(X1)|~esk380_23(X2,X3,X4,X5,X6,X7,X8,X9,X1,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_914)).
cnf(c_0_233,axiom,(esk76_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk380_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk91_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_169)).
cnf(c_0_234,axiom,(esk91_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_105)).
cnf(c_0_235,axiom,(p(X8)|~esk397_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk61_7(X1,X2,X3,X4,X5,X6,X7)|~p(X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_371)).
cnf(c_0_236,axiom,(esk397_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk73_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_60)).
cnf(c_0_237,axiom,(esk406_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk37_5(X1,X2,X3,X4,X5)|~esk405_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_190)).
cnf(c_0_238,axiom,(esk55_7(X1,X2,X3,X4,X5,X6,X7)|esk405_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk66_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_605)).
cnf(c_0_239,axiom,(~esk41_5(X1,X2,X3,X4,X5)|~esk406_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_189)).
cnf(c_0_240,axiom,(esk72_7(X1,X2,X3,X4,X5,X6,X7)|~esk52_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_519)).
cnf(c_0_241,axiom,(esk66_7(X1,X2,X3,X4,X5,X6,X7)|~esk52_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_865)).
cnf(c_0_242,axiom,(esk40_5(X1,X2,X3,X4,X5)|esk37_5(X1,X2,X3,X4,X5)|~esk431_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_162)).
cnf(c_0_243,axiom,(esk431_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk430_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1149)).
cnf(c_0_244,axiom,(esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk371_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_975)).
cnf(c_0_245,axiom,(esk371_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_439)).
cnf(c_0_246,axiom,(~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk277_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_157)).
cnf(c_0_247,axiom,(esk277_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk276_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_328)).
cnf(c_0_248,axiom,(esk196_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|esk276_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk205_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_875)).
cnf(c_0_249,axiom,(~esk196_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk197_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk212_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_714)).
cnf(c_0_250,axiom,(esk212_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1035)).
cnf(c_0_251,axiom,(esk205_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_495)).
cnf(c_0_252,axiom,(~esk262_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X20)|~p(X19)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_656)).
cnf(c_0_253,axiom,(esk262_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk197_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk211_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_120)).
cnf(c_0_254,axiom,(esk211_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1082)).
cnf(c_0_255,axiom,(~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X19)|~p(X20)|~esk279_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X20,X19,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_983)).
cnf(c_0_256,axiom,(esk279_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk194_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_43)).
cnf(c_0_257,axiom,(esk194_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk193_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_940)).
cnf(c_0_258,axiom,(esk173_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk300_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk298_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk299_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_396)).
cnf(c_0_259,axiom,(p(X1)|esk300_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X1,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_317)).
cnf(c_0_260,axiom,(esk298_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X17)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1132)).
cnf(c_0_261,axiom,(p(X1)|esk299_23(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X1,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_201)).
cnf(c_0_262,axiom,(~esk393_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk61_7(X1,X2,X3,X4,X5,X6,X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_13)).
cnf(c_0_263,axiom,(esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk284_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1074)).
cnf(c_0_264,axiom,(esk284_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk283_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk189_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_457)).
cnf(c_0_265,axiom,(esk290_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk289_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1055)).
cnf(c_0_266,axiom,(esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|esk289_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk186_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_147)).
cnf(c_0_267,axiom,(esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk284_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_639)).
cnf(c_0_268,axiom,(esk172_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk173_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_395)).
cnf(c_0_269,axiom,(p(X18)|~esk283_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X17)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1156)).
cnf(c_0_270,axiom,(~esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk290_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_234)).
cnf(c_0_271,axiom,(p(X1)|~p(X2)|~esk158_16(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk298_23(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X2,X1,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_778)).
cnf(c_0_272,axiom,(esk298_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk173_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_360)).
cnf(c_0_273,axiom,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk315_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1066)).
cnf(c_0_274,axiom,(esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk315_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk314_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk163_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_817)).
cnf(c_0_275,axiom,(esk189_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_32)).
cnf(c_0_276,axiom,(esk193_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_928)).
cnf(c_0_277,axiom,(esk186_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_456)).
cnf(c_0_278,axiom,(esk316_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk314_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_138)).
cnf(c_0_279,axiom,(~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk316_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_454)).
cnf(c_0_280,axiom,(~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk313_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_22)).
cnf(c_0_281,axiom,(~esk40_5(X1,X2,X3,X4,X5)|~esk413_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_843)).
cnf(c_0_282,axiom,(esk163_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_985)).
cnf(c_0_283,axiom,(esk311_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk310_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_954)).
cnf(c_0_284,axiom,(esk310_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk165_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_209)).
cnf(c_0_285,axiom,(~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk311_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_110)).
cnf(c_0_286,axiom,(esk165_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_141)).
cnf(c_0_287,axiom,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk330_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1068)).
cnf(c_0_288,axiom,(esk342_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk341_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk102_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_237)).
cnf(c_0_289,axiom,(esk341_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk138_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_944)).
cnf(c_0_290,axiom,(esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk332_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1139)).
cnf(c_0_291,axiom,(esk331_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk332_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk143_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_159)).
cnf(c_0_292,axiom,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk332_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_858)).
cnf(c_0_293,axiom,(~esk342_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_314)).
cnf(c_0_294,axiom,(p(X1)|p(X2)|~p(X3)|~esk331_23(X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X1,X2,X3,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_7)).
cnf(c_0_295,axiom,(p(X14)|~p(X1)|~p(X2)|~esk108_11(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~esk353_23(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X2,X14,X1,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_599)).
cnf(c_0_296,axiom,(esk353_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1051)).
cnf(c_0_297,axiom,(esk138_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_978)).
cnf(c_0_298,axiom,(esk338_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk337_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk104_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_636)).
cnf(c_0_299,axiom,(esk337_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk140_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_795)).
cnf(c_0_300,axiom,(esk143_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_303)).
cnf(c_0_301,axiom,(~esk338_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_266)).
cnf(c_0_302,axiom,(p(X12)|p(X13)|~esk352_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1167)).
cnf(c_0_303,axiom,(esk352_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_951)).
cnf(c_0_304,axiom,(esk140_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_228)).
cnf(c_0_305,axiom,(esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk104_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk372_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1171)).
cnf(c_0_306,axiom,(esk372_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_1201)).
cnf(c_0_307,axiom,(esk102_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk358_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', i_0_671)).
cnf(c_0_308,axiom,(~p(false)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule_false)).
cnf(c_0_309,plain,(esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|esk297_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk184_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_310,plain,(esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk184_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_2, c_0_309])).
cnf(c_0_311,plain,(esk281_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk190_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_312,plain,(esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk191_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X18)|~p(X17)), inference(spm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_313,plain,(~esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_314,plain,(esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_310, c_0_9])).
cnf(c_0_315,plain,(esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|p(X17)|~esk190_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X18)), inference(spm,[status(thm)],[c_0_10, c_0_311])).
cnf(c_0_316,plain,(esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X18)|~p(X17)), inference(spm,[status(thm)],[c_0_312, c_0_11])).
cnf(c_0_317,plain,(~esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_313, c_0_314])).
cnf(c_0_318,plain,(esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X18)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_12]), c_0_316])).
cnf(c_0_319,plain,(~esk173_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X17)|~p(X18)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_320, plain, ((~epred29_0<=>![X18]:~p(X18))), introduced(definition)).
fof(c_0_321, plain, ((~epred30_0<=>![X17]:~p(X17))), introduced(definition)).
cnf(c_0_322,plain,(esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk302_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk169_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_323,plain,(esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk170_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X15)|~p(X16)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_324, plain, ((~epred28_0<=>![X16]:![X15]:![X14]:![X13]:![X12]:![X11]:![X10]:![X9]:![X8]:![X7]:![X6]:![X5]:![X4]:![X3]:![X2]:![X1]:((~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16))|~esk172_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)))), introduced(definition)).
cnf(c_0_325,plain,(~esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X18)), inference(spm,[status(thm)],[c_0_317, c_0_318])).
cnf(c_0_326,plain,(esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk172_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X17)|~p(X18)), inference(spm,[status(thm)],[c_0_319, c_0_19])).
cnf(c_0_327,plain,(epred29_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_320])).
cnf(c_0_328,plain,(epred30_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_321])).
cnf(c_0_329,plain,(esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X15)|~esk169_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X16)), inference(spm,[status(thm)],[c_0_20, c_0_322])).
cnf(c_0_330,plain,(esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X15)|~p(X16)), inference(spm,[status(thm)],[c_0_323, c_0_21])).
cnf(c_0_331,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk323_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk161_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_332,plain,(~epred30_0|~epred29_0|~epred28_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_324]), c_0_320]), c_0_321])).
cnf(c_0_333,plain,(epred29_0), inference(spm,[status(thm)],[c_0_327, c_0_24])).
cnf(c_0_334,plain,(epred30_0), inference(spm,[status(thm)],[c_0_328, c_0_24])).
cnf(c_0_335,plain,(esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk313_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk164_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_336,plain,(~esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_337,plain,(esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X16)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_29]), c_0_330])).
cnf(c_0_338,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk161_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_30, c_0_331])).
cnf(c_0_339,plain,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|~esk145_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X14)|~p(X13)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_340,plain,(esk351_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk135_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_341,plain,(epred28_0|~esk172_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(split_equiv,[status(thm)],[c_0_324])).
cnf(c_0_342,plain,(~epred28_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_332, c_0_333])]), c_0_334])])).
cnf(c_0_343,plain,(esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk164_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_35, c_0_335])).
cnf(c_0_344,plain,(~esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X16)), inference(spm,[status(thm)],[c_0_336, c_0_337])).
cnf(c_0_345,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_338, c_0_36])).
cnf(c_0_346,plain,(~esk149_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X16)|~p(X15)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_347, plain, ((~epred22_0<=>![X16]:~p(X16))), introduced(definition)).
fof(c_0_348, plain, ((~epred23_0<=>![X15]:~p(X15))), introduced(definition)).
cnf(c_0_349,plain,(~esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_350,plain,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X14)|~p(X13)), inference(spm,[status(thm)],[c_0_339, c_0_41])).
cnf(c_0_351,plain,(esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk135_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_42, c_0_340])).
cnf(c_0_352,plain,(esk309_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk130_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk166_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_353,plain,(esk149_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X15)|~esk325_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])).
cnf(c_0_354,plain,(~esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_48]), c_0_342])).
cnf(c_0_355,plain,(esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_343, c_0_49])).
cnf(c_0_356,plain,(esk304_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk168_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_357,plain,(esk304_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk168_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_52, c_0_51])).
fof(c_0_358, plain, ((~epred21_0<=>![X14]:![X13]:![X12]:![X11]:![X10]:![X9]:![X8]:![X7]:![X6]:![X5]:![X4]:![X3]:![X2]:![X1]:((~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14))|~esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)))), introduced(definition)).
cnf(c_0_359,plain,(~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X16)), inference(spm,[status(thm)],[c_0_344, c_0_345])).
cnf(c_0_360,plain,(esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X16)|~p(X15)), inference(spm,[status(thm)],[c_0_346, c_0_53])).
cnf(c_0_361,plain,(epred22_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_347])).
cnf(c_0_362,plain,(epred23_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_348])).
cnf(c_0_363,plain,(p(X1)|~esk123_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13,X14)|~esk122_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13,X14)|~p(X14)|~p(X13)), inference(spm,[status(thm)],[c_0_349, c_0_350])).
cnf(c_0_364,plain,(esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_351, c_0_54])).
cnf(c_0_365,plain,(p(X1)|~esk121_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13,X14)|~esk109_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_366,plain,(esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk130_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk166_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_57, c_0_352])).
cnf(c_0_367,plain,(esk149_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X16)|p(X15)), inference(spm,[status(thm)],[c_0_353, c_0_58])).
cnf(c_0_368,plain,(esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_354, c_0_355])).
cnf(c_0_369,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X16)|~esk168_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X15)), inference(spm,[status(thm)],[c_0_59, c_0_356])).
cnf(c_0_370,plain,(esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X16)|~esk168_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X15)), inference(spm,[status(thm)],[c_0_59, c_0_357])).
cnf(c_0_371,plain,(~epred23_0|~epred22_0|~epred21_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_360]), c_0_358]), c_0_347]), c_0_348])).
cnf(c_0_372,plain,(epred22_0), inference(spm,[status(thm)],[c_0_361, c_0_24])).
cnf(c_0_373,plain,(epred23_0), inference(spm,[status(thm)],[c_0_362, c_0_24])).
cnf(c_0_374,plain,(esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|~esk352_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk354_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])).
cnf(c_0_375,plain,(p(X1)|~esk99_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk122_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13,X14)|~esk109_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X14)|~p(X13)), inference(spm,[status(thm)],[c_0_363, c_0_364])).
cnf(c_0_376,plain,(esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_365, c_0_63])).
cnf(c_0_377,plain,(esk364_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk80_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk114_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_378,plain,(esk358_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|p(X10)|~esk117_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X11)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_379,plain,(esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|~esk377_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])).
cnf(c_0_380,plain,(esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk130_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_366, c_0_71])).
cnf(c_0_381,plain,(p(X1)|~esk149_16(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X1)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_382,plain,(esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X15)|p(X16)), inference(spm,[status(thm)],[c_0_74, c_0_367])).
cnf(c_0_383,plain,(~esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_336, c_0_368])).
cnf(c_0_384,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X15)), inference(spm,[status(thm)],[c_0_369, c_0_75])).
cnf(c_0_385,plain,(esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X15)), inference(spm,[status(thm)],[c_0_370, c_0_75])).
cnf(c_0_386,plain,(~epred21_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_371, c_0_372])]), c_0_373])])).
cnf(c_0_387,plain,(esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|~esk354_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X13)), inference(spm,[status(thm)],[c_0_374, c_0_76])).
cnf(c_0_388,plain,(p(X1)|~esk120_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk99_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk109_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_375, c_0_376])).
cnf(c_0_389,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk80_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk114_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_77, c_0_377])).
cnf(c_0_390,plain,(esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|~esk117_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X11)), inference(spm,[status(thm)],[c_0_78, c_0_378])).
cnf(c_0_391,plain,(p(X1)|~esk95_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X1,X11)|~esk83_9(X2,X3,X4,X5,X6,X7,X8,X9,X10)|~p(X11)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_392,plain,(esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk376_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk377_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_68, c_0_81])).
cnf(c_0_393,plain,(esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X11)|p(X10)), inference(spm,[status(thm)],[c_0_379, c_0_82])).
cnf(c_0_394,plain,(~esk130_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_354, c_0_380])).
cnf(c_0_395,plain,(esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X16)|~esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X15)), inference(spm,[status(thm)],[c_0_381, c_0_53])).
cnf(c_0_396,plain,(esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X15)), inference(ef,[status(thm)],[c_0_382])).
cnf(c_0_397,plain,(esk130_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk346_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_398,plain,(esk346_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk349_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk100_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk136_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_399,plain,(p(X1)|~esk150_16(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X1)|~esk128_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk134_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_384]), c_0_385])).
cnf(c_0_400,plain,(esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk329_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk144_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_401,plain,(~esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_358]), c_0_386])).
cnf(c_0_402,plain,(esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk146_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X12)|~p(X14)|~p(X13)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_403,plain,(esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X14)|p(X12)|~p(X13)), inference(spm,[status(thm)],[c_0_387, c_0_91])).
cnf(c_0_404,plain,(p(X1)|~esk96_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk99_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk109_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_388, c_0_92])).
cnf(c_0_405,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk80_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_389, c_0_93])).
cnf(c_0_406,plain,(esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X11)), inference(spm,[status(thm)],[c_0_390, c_0_94])).
cnf(c_0_407,plain,(esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|~esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)), inference(spm,[status(thm)],[c_0_391, c_0_95])).
cnf(c_0_408,plain,(esk80_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk390_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_409,plain,(esk390_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk393_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_410,plain,(esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk377_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X11)), inference(spm,[status(thm)],[c_0_392, c_0_100])).
cnf(c_0_411,plain,(esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)), inference(spm,[status(thm)],[c_0_101, c_0_393])).
cnf(c_0_412,plain,(p(X1)|~esk130_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk160_16(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X1)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_394, c_0_395]), c_0_396])).
cnf(c_0_413,plain,(esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X16)|~esk148_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X15)), inference(spm,[status(thm)],[c_0_385, c_0_395])).
cnf(c_0_414,plain,(esk130_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk349_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk100_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk136_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_397, c_0_398])).
cnf(c_0_415,plain,(p(X1)|~esk128_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk134_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_395]), c_0_396])).
cnf(c_0_416,plain,(esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X13)|~esk144_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X12)|~p(X14)), inference(spm,[status(thm)],[c_0_102, c_0_400])).
cnf(c_0_417,plain,(~esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_401, c_0_103])).
cnf(c_0_418,plain,(esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X12)|~p(X14)|~p(X13)), inference(spm,[status(thm)],[c_0_402, c_0_104])).
cnf(c_0_419,plain,(esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_105, c_0_403])).
cnf(c_0_420,plain,(p(X1)|~esk80_9(X2,X3,X4,X5,X6,X7,X8,X9,X10)|~esk96_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk83_9(X2,X3,X4,X5,X6,X7,X8,X9,X10)|~esk109_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_404, c_0_405])).
cnf(c_0_421,plain,(esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|~esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)), inference(spm,[status(thm)],[c_0_406, c_0_407])).
cnf(c_0_422,plain,(esk80_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk393_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_408, c_0_409])).
cnf(c_0_423,plain,(esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)|~p(X10)), inference(spm,[status(thm)],[c_0_410, c_0_106])).
cnf(c_0_424,plain,(esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)), inference(ef,[status(thm)],[c_0_411])).
cnf(c_0_425,plain,(esk73_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X8)|~esk397_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109])).
cnf(c_0_426,plain,(p(X1)|~esk130_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_413]), c_0_396])).
cnf(c_0_427,plain,(esk130_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk100_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk136_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_110, c_0_414])).
cnf(c_0_428,plain,(p(X1)|p(X2)|~esk144_14(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X1,X15)|~esk132_14(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X1,X15)|~esk134_14(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X1,X15)|~p(X16)|~p(X14)|~p(X15)), inference(spm,[status(thm)],[c_0_415, c_0_416])).
cnf(c_0_429,plain,(~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X12)|~p(X14)|~p(X13)), inference(spm,[status(thm)],[c_0_417, c_0_418])).
cnf(c_0_430,plain,(~esk121_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X14)|~p(X13)|~p(X12)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_431,plain,(esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X13)), inference(spm,[status(thm)],[c_0_419, c_0_24])).
cnf(c_0_432,plain,(esk336_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk141_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk105_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_433,plain,(p(X1)|p(X2)|~esk80_9(X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk94_9(X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X12)|~p(X13)|~p(X14)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_407]), c_0_421])).
cnf(c_0_434,plain,(esk80_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_115, c_0_422])).
cnf(c_0_435,plain,(esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_423]), c_0_424]), c_0_424])).
cnf(c_0_436,plain,(esk73_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X9)|p(X8)), inference(spm,[status(thm)],[c_0_425, c_0_116])).
cnf(c_0_437,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X15)|~esk100_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk136_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X16)), inference(spm,[status(thm)],[c_0_426, c_0_427])).
cnf(c_0_438,plain,(p(X1)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk122_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk134_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)|~p(X13)|~p(X15)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_117]), c_0_429])).
cnf(c_0_439,plain,(esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X14)|~p(X13)|~p(X12)), inference(spm,[status(thm)],[c_0_430, c_0_63])).
cnf(c_0_440,plain,(esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)), inference(ef,[status(thm)],[c_0_431])).
cnf(c_0_441,plain,(esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk141_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk105_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_118, c_0_432])).
cnf(c_0_442,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X12)|~p(X13)|~p(X14)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_434]), c_0_435])).
cnf(c_0_443,plain,(~esk73_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_444,plain,(esk72_7(X1,X2,X3,X4,X5,X6,X7)|p(X8)|p(X9)), inference(spm,[status(thm)],[c_0_121, c_0_436])).
cnf(c_0_445,plain,(esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk92_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_446,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X15)|~esk100_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X16)), inference(spm,[status(thm)],[c_0_437, c_0_124])).
cnf(c_0_447,plain,(p(X1)|~esk110_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk134_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)|~p(X13)|~p(X15)|~p(X14)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_438, c_0_439]), c_0_440])).
cnf(c_0_448,plain,(esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk105_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_441, c_0_125])).
cnf(c_0_449,plain,(esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk373_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_450,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X12)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_442, c_0_128])).
cnf(c_0_451,plain,(esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk72_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_443, c_0_129])).
cnf(c_0_452,plain,(esk72_7(X1,X2,X3,X4,X5,X6,X7)|p(X8)), inference(ef,[status(thm)],[c_0_444])).
cnf(c_0_453,plain,(esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_445, c_0_130])).
cnf(c_0_454,plain,(p(X1)|~esk100_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk110_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)|~p(X15)|~p(X14)|~p(X13)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_446, c_0_439]), c_0_440]), c_0_447])).
cnf(c_0_455,plain,(esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk105_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X14)|~p(X13)|~p(X12)), inference(spm,[status(thm)],[c_0_448, c_0_439])).
cnf(c_0_456,plain,(esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_131, c_0_449])).
cnf(c_0_457,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X12)|~p(X13)|~p(X14)|~p(X9)|~p(X8)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_451]), c_0_452])).
cnf(c_0_458,plain,(esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk72_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_453, c_0_451])).
cnf(c_0_459,plain,(p(X1)|~esk100_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk105_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk110_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X13)|~p(X14)|~p(X15)|~p(X16)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_454, c_0_455]), c_0_440])).
cnf(c_0_460,plain,(esk100_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_132, c_0_456])).
cnf(c_0_461,plain,(esk105_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_462,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X12)|~p(X13)|~p(X14)|~p(X9)|~p(X8)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_457, c_0_458]), c_0_452])).
cnf(c_0_463,plain,(p(X1)|~esk112_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk78_9(X2,X3,X4,X5,X6,X7,X8,X9,X10)|~esk105_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk110_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk83_9(X2,X3,X4,X5,X6,X7,X8,X9,X10)|~p(X13)|~p(X14)|~p(X15)|~p(X16)), inference(spm,[status(thm)],[c_0_459, c_0_460])).
cnf(c_0_464,plain,(esk105_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_461, c_0_456])).
cnf(c_0_465,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X12)|~p(X13)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_462, c_0_24])).
cnf(c_0_466,plain,(esk413_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk410_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk64_7(X1,X2,X3,X4,X5,X6,X7)|~esk35_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_467,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X13)|~p(X14)|~p(X15)|~p(X16)), inference(spm,[status(thm)],[c_0_463, c_0_464])).
cnf(c_0_468,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X12)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_465, c_0_24])).
cnf(c_0_469,plain,(esk413_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk64_7(X1,X2,X3,X4,X5,X6,X7)|~esk35_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_137, c_0_466])).
cnf(c_0_470,plain,(esk51_7(X1,X2,X3,X4,X5,X6,X7)|p(X7)|~esk416_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_140])).
cnf(c_0_471,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X13)|~p(X14)|~p(X15)|~p(X16)), inference(spm,[status(thm)],[c_0_467, c_0_141])).
cnf(c_0_472,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_468, c_0_24])).
cnf(c_0_473,plain,(esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk41_5(X1,X2,X3,X4,X5)|~esk64_7(X1,X2,X3,X4,X5,X6,X7)|~esk35_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_142, c_0_469])).
cnf(c_0_474,plain,(esk51_7(X1,X2,X3,X4,X5,X6,X7)|p(X6)|p(X7)), inference(spm,[status(thm)],[c_0_470, c_0_143])).
cnf(c_0_475,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X13)|~p(X14)|~p(X15)), inference(spm,[status(thm)],[c_0_471, c_0_24])).
cnf(c_0_476,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk41_5(X1,X2,X3,X4,X5)|~esk64_7(X1,X2,X3,X4,X5,X6,X7)|~esk35_5(X1,X2,X3,X4,X5)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_472, c_0_473])).
cnf(c_0_477,plain,(~esk41_5(X1,X2,X3,X4,X5)|~esk51_7(X1,X2,X3,X4,X5,X6,X7)|~p(X6)|~p(X7)), inference(spm,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_478,plain,(esk50_5(X1,X2,X3,X4,X5)|p(X6)|p(X7)), inference(spm,[status(thm)],[c_0_146, c_0_474])).
cnf(c_0_479,plain,(esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk70_7(X1,X2,X3,X4,X5,X6,X7)|~p(X7)|~p(X6)), inference(spm,[status(thm)],[c_0_147, c_0_148])).
cnf(c_0_480,plain,(esk30_5(X1,X2,X3,X4,X5)|p(X4)|~esk435_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_151])).
cnf(c_0_481,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_475, c_0_24])).
cnf(c_0_482,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk52_7(X1,X2,X3,X4,X5,X6,X7)|~esk41_5(X1,X2,X3,X4,X5)|~esk35_5(X1,X2,X3,X4,X5)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_476, c_0_152])).
cnf(c_0_483,plain,(esk52_7(X1,X2,X3,X4,X5,X6,X7)|~esk50_5(X1,X2,X3,X4,X5)|~esk41_5(X1,X2,X3,X4,X5)|~p(X6)|~p(X7)), inference(spm,[status(thm)],[c_0_477, c_0_153])).
cnf(c_0_484,plain,(esk50_5(X1,X2,X3,X4,X5)|p(X6)), inference(ef,[status(thm)],[c_0_478])).
cnf(c_0_485,plain,(esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk52_7(X1,X2,X3,X4,X5,X6,X7)|~p(X7)|~p(X6)), inference(spm,[status(thm)],[c_0_479, c_0_154])).
cnf(c_0_486,plain,(esk41_5(X1,X2,X3,X4,X5)|~esk48_5(X1,X2,X3,X4,X5)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_487,plain,(~esk30_5(X1,X2,X3,X4,X5)|~esk23_3(X1,X2,X3)|~p(X4)|~p(X5)), inference(spm,[status(thm)],[c_0_157, c_0_158])).
cnf(c_0_488,plain,(esk30_5(X1,X2,X3,X4,X5)|p(X5)|p(X4)), inference(spm,[status(thm)],[c_0_480, c_0_159])).
cnf(c_0_489,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X13)), inference(spm,[status(thm)],[c_0_481, c_0_24])).
cnf(c_0_490,plain,(esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|~esk117_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X11)), inference(spm,[status(thm)],[c_0_160, c_0_378])).
cnf(c_0_491,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk41_5(X1,X2,X3,X4,X5)|~esk35_5(X1,X2,X3,X4,X5)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X9)|~p(X8)|~p(X6)|~p(X7)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_483]), c_0_484])).
cnf(c_0_492,plain,(esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk50_5(X1,X2,X3,X4,X5)|~esk41_5(X1,X2,X3,X4,X5)|~p(X7)|~p(X6)), inference(spm,[status(thm)],[c_0_485, c_0_483])).
cnf(c_0_493,plain,(esk41_5(X1,X2,X3,X4,X5)|~esk31_5(X1,X2,X3,X4,X5)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_486, c_0_161])).
cnf(c_0_494,plain,(esk31_5(X1,X2,X3,X4,X5)|~esk23_3(X1,X2,X3)|~esk29_3(X1,X2,X3)|~p(X4)|~p(X5)), inference(spm,[status(thm)],[c_0_487, c_0_162])).
cnf(c_0_495,plain,(esk29_3(X1,X2,X3)|p(X4)|p(X5)), inference(spm,[status(thm)],[c_0_163, c_0_488])).
cnf(c_0_496,plain,(esk433_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk430_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk17_3(X1,X2,X3)|~esk42_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_497,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_489, c_0_24])).
cnf(c_0_498,plain,(esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X11)), inference(spm,[status(thm)],[c_0_490, c_0_94])).
cnf(c_0_499,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk41_5(X1,X2,X3,X4,X5)|~esk35_5(X1,X2,X3,X4,X5)|~p(X9)|~p(X8)|~p(X6)|~p(X7)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_492]), c_0_484])).
cnf(c_0_500,plain,(esk41_5(X1,X2,X3,X4,X5)|~esk23_3(X1,X2,X3)|~esk29_3(X1,X2,X3)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_493, c_0_494])).
cnf(c_0_501,plain,(esk29_3(X1,X2,X3)|p(X4)), inference(ef,[status(thm)],[c_0_495])).
cnf(c_0_502,plain,(esk433_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk35_5(X1,X2,X3,X4,X5)|~esk17_3(X1,X2,X3)|~esk42_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_166, c_0_496])).
cnf(c_0_503,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_497, c_0_407]), c_0_424])).
cnf(c_0_504,plain,(esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|~esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)), inference(spm,[status(thm)],[c_0_498, c_0_407])).
cnf(c_0_505,plain,(esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk393_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_167, c_0_409])).
cnf(c_0_506,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk23_3(X1,X2,X3)|~esk35_5(X1,X2,X3,X4,X5)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_500]), c_0_501])).
cnf(c_0_507,plain,(esk35_5(X1,X2,X3,X4,X5)|~esk17_3(X1,X2,X3)|~esk42_5(X1,X2,X3,X4,X5)|~esk23_3(X1,X2,X3)), inference(spm,[status(thm)],[c_0_168, c_0_502])).
cnf(c_0_508,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_503, c_0_504]), c_0_424])).
cnf(c_0_509,plain,(esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_115, c_0_505])).
cnf(c_0_510,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk17_3(X1,X2,X3)|~esk42_5(X1,X2,X3,X4,X5)|~esk23_3(X1,X2,X3)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_506, c_0_507])).
cnf(c_0_511,plain,(~esk10_1(X1)|~esk7_1(X1)|~esk442_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_169, c_0_170])).
cnf(c_0_512,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)), inference(spm,[status(thm)],[c_0_508, c_0_509])).
cnf(c_0_513,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk17_3(X1,X2,X3)|~esk23_3(X1,X2,X3)|~esk31_5(X1,X2,X3,X4,X5)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_510, c_0_171])).
cnf(c_0_514,plain,(esk443_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk10_1(X1)|~esk24_3(X1,X2,X3)|~esk7_1(X1)), inference(spm,[status(thm)],[c_0_511, c_0_172])).
cnf(c_0_515,plain,(esk13_3(X1,X2,X3)|~esk448_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~p(X2)), inference(spm,[status(thm)],[c_0_173, c_0_174])).
cnf(c_0_516,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)), inference(spm,[status(thm)],[c_0_512, c_0_128])).
cnf(c_0_517,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk17_3(X1,X2,X3)|~esk23_3(X1,X2,X3)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_494]), c_0_501])).
cnf(c_0_518,plain,(esk17_3(X1,X2,X3)|~esk10_1(X1)|~esk24_3(X1,X2,X3)|~esk7_1(X1)), inference(spm,[status(thm)],[c_0_175, c_0_514])).
cnf(c_0_519,plain,(esk23_3(X1,X2,X3)|~esk27_3(X1,X2,X3)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_176, c_0_177])).
cnf(c_0_520,plain,(~esk10_1(X1)|~esk13_3(X1,X2,X3)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_178, c_0_179])).
cnf(c_0_521,plain,(esk13_3(X1,X2,X3)|p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_515, c_0_180])).
cnf(c_0_522,plain,(esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_181, c_0_182])])).
cnf(c_0_523,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)|~p(X9)|~p(X8)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_516, c_0_451]), c_0_452])).
cnf(c_0_524,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk10_1(X1)|~esk24_3(X1,X2,X3)|~esk23_3(X1,X2,X3)|~esk7_1(X1)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_517, c_0_518])).
cnf(c_0_525,plain,(esk23_3(X1,X2,X3)|~esk14_3(X1,X2,X3)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_519, c_0_183])).
cnf(c_0_526,plain,(esk14_3(X1,X2,X3)|~esk10_1(X1)|~esk12_1(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_520, c_0_184])).
cnf(c_0_527,plain,(esk12_1(X1)|p(X2)|~p(X3)), inference(spm,[status(thm)],[c_0_185, c_0_521])).
cnf(c_0_528,plain,(esk5_1(X1)|esk4_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_186, c_0_522])])).
cnf(c_0_529,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X11)|~p(X9)|~p(X8)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_523, c_0_458]), c_0_452])).
cnf(c_0_530,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk10_1(X1)|~esk23_3(X1,X2,X3)|~esk7_1(X1)|~esk14_3(X1,X2,X3)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_524, c_0_187])).
cnf(c_0_531,plain,(esk23_3(X1,X2,X3)|~esk10_1(X1)|~esk12_1(X1)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_525, c_0_526])).
cnf(c_0_532,plain,(esk12_1(X1)|p(X2)), inference(spm,[status(thm)],[c_0_527, c_0_24])).
cnf(c_0_533,plain,(esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_188, c_0_182])])).
cnf(c_0_534,plain,(esk10_1(X1)|~esk11_1(X1)|~p(X1)), inference(spm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_535,plain,(esk11_1(X1)|esk4_1(X1)), inference(spm,[status(thm)],[c_0_191, c_0_528])).
cnf(c_0_536,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk41_5(X1,X2,X3,X4,X5)|~esk64_7(X1,X2,X3,X4,X5,X6,X7)|~esk35_5(X1,X2,X3,X4,X5)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X11)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_529, c_0_473])).
cnf(c_0_537,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk10_1(X1)|~esk7_1(X1)|~esk14_3(X1,X2,X3)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_530, c_0_531]), c_0_532])).
cnf(c_0_538,plain,(~esk4_1(X1)|~p(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_192, c_0_533])])).
cnf(c_0_539,plain,(esk10_1(X1)|esk4_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_535]), c_0_193])).
cnf(c_0_540,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk52_7(X1,X2,X3,X4,X5,X6,X7)|~esk41_5(X1,X2,X3,X4,X5)|~esk35_5(X1,X2,X3,X4,X5)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X11)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_536, c_0_152])).
cnf(c_0_541,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk10_1(X1)|~esk7_1(X1)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_537, c_0_526]), c_0_532])).
cnf(c_0_542,plain,(esk10_1(X1)|~p(X1)), inference(spm,[status(thm)],[c_0_538, c_0_539])).
cnf(c_0_543,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk41_5(X1,X2,X3,X4,X5)|~esk35_5(X1,X2,X3,X4,X5)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_540, c_0_483]), c_0_484])).
cnf(c_0_544,plain,(esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk375_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk111_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_194, c_0_195])).
cnf(c_0_545,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk7_1(X1)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_541, c_0_542])).
cnf(c_0_546,plain,(esk7_1(X1)|~esk11_1(X1)|~p(X1)), inference(spm,[status(thm)],[c_0_196, c_0_190])).
cnf(c_0_547,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk41_5(X1,X2,X3,X4,X5)|~esk35_5(X1,X2,X3,X4,X5)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_543, c_0_492]), c_0_484])).
cnf(c_0_548,plain,(esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk111_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_197, c_0_544])).
cnf(c_0_549,plain,(esk359_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk116_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_198, c_0_199])).
cnf(c_0_550,plain,(esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk118_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X10)|~p(X11)), inference(spm,[status(thm)],[c_0_200, c_0_201])).
cnf(c_0_551,plain,(esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~esk11_1(X1)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_545, c_0_546])).
cnf(c_0_552,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk23_3(X1,X2,X3)|~esk35_5(X1,X2,X3,X4,X5)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_547, c_0_500]), c_0_501])).
cnf(c_0_553,plain,(~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_202, c_0_203])).
cnf(c_0_554,plain,(esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_548, c_0_204])).
cnf(c_0_555,plain,(esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X11)|~esk116_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X10)), inference(spm,[status(thm)],[c_0_205, c_0_549])).
cnf(c_0_556,plain,(esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X10)|~p(X11)), inference(spm,[status(thm)],[c_0_550, c_0_206])).
cnf(c_0_557,plain,(esk4_1(X1)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X10)|p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_551, c_0_535]), c_0_193])).
cnf(c_0_558,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk17_3(X1,X2,X3)|~esk42_5(X1,X2,X3,X4,X5)|~esk23_3(X1,X2,X3)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_552, c_0_507])).
cnf(c_0_559,plain,(~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_553, c_0_554])).
cnf(c_0_560,plain,(esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X10)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_207]), c_0_556])).
cnf(c_0_561,plain,(~esk95_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)|~p(X10)), inference(spm,[status(thm)],[c_0_208, c_0_209])).
fof(c_0_562, plain, ((~epred13_0<=>![X10]:~p(X10))), introduced(definition)).
cnf(c_0_563,plain,(esk4_1(X1)|esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,true)|p(X9)|p(X10)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_557, c_0_24])).
cnf(c_0_564,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk17_3(X1,X2,X3)|~esk23_3(X1,X2,X3)|~esk31_5(X1,X2,X3,X4,X5)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_558, c_0_171])).
fof(c_0_565, plain, ((~epred12_0<=>![X9]:![X8]:![X7]:![X6]:![X5]:![X4]:![X3]:![X2]:![X1]:((~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9))|~esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)))), introduced(definition)).
fof(c_0_566, plain, ((~epred14_0<=>![X11]:~p(X11))), introduced(definition)).
cnf(c_0_567,plain,(~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X10)), inference(spm,[status(thm)],[c_0_559, c_0_560])).
cnf(c_0_568,plain,(esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X11)|~p(X10)), inference(spm,[status(thm)],[c_0_561, c_0_95])).
cnf(c_0_569,plain,(epred13_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_562])).
cnf(c_0_570,plain,(esk4_1(X1)|esk84_9(X1,X2,X3,X4,X5,X6,X7,true,true)|p(X8)|p(X9)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_563, c_0_24])).
cnf(c_0_571,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk17_3(X1,X2,X3)|~esk23_3(X1,X2,X3)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_564, c_0_494]), c_0_501])).
cnf(c_0_572,plain,(~epred14_0|~epred13_0|~epred12_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(spm,[status(thm)],[c_0_567, c_0_568]), c_0_565]), c_0_562]), c_0_566])).
cnf(c_0_573,plain,(epred13_0), inference(spm,[status(thm)],[c_0_569, c_0_24])).
cnf(c_0_574,plain,(epred14_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_566])).
cnf(c_0_575,plain,(esk4_1(X1)|esk84_9(X1,X2,X3,X4,X5,true,X6,true,true)|p(X7)|p(X8)|~p(X6)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_570, c_0_24])).
cnf(c_0_576,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk10_1(X1)|~esk24_3(X1,X2,X3)|~esk23_3(X1,X2,X3)|~esk7_1(X1)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_571, c_0_518])).
cnf(c_0_577,plain,(~epred14_0|~epred12_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_572, c_0_573])])).
cnf(c_0_578,plain,(epred14_0), inference(spm,[status(thm)],[c_0_574, c_0_24])).
cnf(c_0_579,plain,(esk4_1(X1)|esk84_9(X1,X2,X3,X4,X5,true,true,true,true)|p(X6)|p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_575, c_0_24])).
cnf(c_0_580,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk10_1(X1)|~esk23_3(X1,X2,X3)|~esk7_1(X1)|~esk14_3(X1,X2,X3)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)), inference(spm,[status(thm)],[c_0_576, c_0_187])).
cnf(c_0_581,plain,(~epred12_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_577, c_0_578])])).
cnf(c_0_582,plain,(esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk386_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk58_7(X1,X2,X3,X4,X5,X6,X7)|~esk88_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_210, c_0_211])).
cnf(c_0_583,plain,(esk4_1(X1)|esk84_9(X1,X2,X3,X4,true,true,true,true,true)|p(X5)|p(X6)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_579, c_0_24])).
cnf(c_0_584,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk10_1(X1)|~esk7_1(X1)|~esk14_3(X1,X2,X3)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_580, c_0_531]), c_0_532])).
cnf(c_0_585,plain,(~esk94_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_565]), c_0_581])).
cnf(c_0_586,plain,(esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk58_7(X1,X2,X3,X4,X5,X6,X7)|~esk88_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_212, c_0_582])).
cnf(c_0_587,plain,(esk4_1(X1)|esk84_9(X1,X2,X3,true,true,true,true,true,true)|p(X4)|p(X5)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_583, c_0_24])).
cnf(c_0_588,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk10_1(X1)|~esk7_1(X1)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_584, c_0_526]), c_0_532])).
cnf(c_0_589,plain,(~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_585, c_0_213])).
cnf(c_0_590,plain,(esk77_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk58_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_586, c_0_214])).
cnf(c_0_591,plain,(esk4_1(X1)|esk84_9(X1,X2,true,true,true,true,true,true,true)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_587, c_0_24])).
cnf(c_0_592,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~esk7_1(X1)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_588, c_0_542])).
cnf(c_0_593,plain,(~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk58_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_589, c_0_590])).
cnf(c_0_594,plain,(esk4_1(X1)|esk84_9(X1,true,true,true,true,true,true,true,true)|p(X2)|p(X3)), inference(spm,[status(thm)],[c_0_591, c_0_24])).
cnf(c_0_595,plain,(esk58_7(X1,X2,X3,X4,X5,X6,X7)|esk61_7(X1,X2,X3,X4,X5,X6,X7)|~esk410_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_215, c_0_216])).
cnf(c_0_596,plain,(esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X12)|p(X10)|~esk11_1(X1)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)|~p(X1)), inference(spm,[status(thm)],[c_0_592, c_0_546])).
cnf(c_0_597,plain,(~esk72_7(X1,X2,X3,X4,X5,X6,X7)|~esk58_7(X1,X2,X3,X4,X5,X6,X7)|~esk62_7(X1,X2,X3,X4,X5,X6,X7)|~esk84_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X9)|~p(X8)), inference(spm,[status(thm)],[c_0_593, c_0_451])).
cnf(c_0_598,plain,(esk4_1(X1)|esk84_9(X1,true,true,true,true,true,true,true,true)|p(X2)), inference(ef,[status(thm)],[c_0_594])).
cnf(c_0_599,plain,(esk413_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk58_7(X1,X2,X3,X4,X5,X6,X7)|esk61_7(X1,X2,X3,X4,X5,X6,X7)|~esk64_7(X1,X2,X3,X4,X5,X6,X7)|~esk35_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_595, c_0_466])).
cnf(c_0_600,plain,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk344_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk137_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk101_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_217, c_0_218])).
cnf(c_0_601,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk334_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk142_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_219, c_0_220])).
cnf(c_0_602,plain,(esk4_1(X1)|esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|p(X12)|~p(X11)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_596, c_0_535]), c_0_193])).
cnf(c_0_603,plain,(esk4_1(X1)|p(X2)|~esk58_7(X1,true,true,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_597, c_0_598]), c_0_24])]), c_0_452])).
cnf(c_0_604,plain,(esk58_7(X1,X2,X3,X4,X5,X6,X7)|esk61_7(X1,X2,X3,X4,X5,X6,X7)|~esk41_5(X1,X2,X3,X4,X5)|~esk64_7(X1,X2,X3,X4,X5,X6,X7)|~esk35_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_142, c_0_599])).
cnf(c_0_605,plain,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk137_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk101_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_221, c_0_600])).
cnf(c_0_606,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk142_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_222, c_0_601])).
fof(c_0_607, plain, ((~epred25_0<=>![X14]:~p(X14))), introduced(definition)).
cnf(c_0_608,plain,(esk4_1(X1)|esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,true)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,true)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,true)|p(X11)|p(X10)|~p(X9)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_602, c_0_24])).
cnf(c_0_609,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk41_5(X1,true,true,true,true)|~esk64_7(X1,true,true,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)), inference(spm,[status(thm)],[c_0_603, c_0_604])).
cnf(c_0_610,plain,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk125_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk101_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_605, c_0_223])).
fof(c_0_611, plain, ((~epred24_0<=>![X11]:![X10]:![X9]:![X8]:![X7]:![X6]:![X5]:![X4]:![X3]:![X2]:![X1]:((~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11))|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)))), introduced(definition)).
fof(c_0_612, plain, ((~epred26_0<=>![X13]:~p(X13))), introduced(definition)).
fof(c_0_613, plain, ((~epred27_0<=>![X12]:~p(X12))), introduced(definition)).
cnf(c_0_614,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_606, c_0_224])).
cnf(c_0_615,plain,(~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X12)|~p(X14)|~p(X13)), inference(spm,[status(thm)],[c_0_429, c_0_439])).
cnf(c_0_616,plain,(epred25_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_607])).
cnf(c_0_617,plain,(esk4_1(X1)|esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,true,X9,true)|esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,true,X9,true)|esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,true,X9,true)|p(X9)|p(X10)|~p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_608, c_0_24])).
cnf(c_0_618,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk52_7(X1,true,true,true,true,true,true)|~esk41_5(X1,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)), inference(spm,[status(thm)],[c_0_609, c_0_152])).
cnf(c_0_619,plain,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk101_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_417, c_0_610])).
cnf(c_0_620,plain,(~epred27_0|~epred26_0|~epred25_0|~epred24_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_614, c_0_439]), c_0_615]), c_0_611]), c_0_607]), c_0_612]), c_0_613])).
cnf(c_0_621,plain,(epred25_0), inference(spm,[status(thm)],[c_0_616, c_0_24])).
cnf(c_0_622,plain,(epred26_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_612])).
cnf(c_0_623,plain,(esk4_1(X1)|esk99_11(X1,X2,X3,X4,X5,X6,X7,true,true,X8,true)|esk98_11(X1,X2,X3,X4,X5,X6,X7,true,true,X8,true)|esk107_11(X1,X2,X3,X4,X5,X6,X7,true,true,X8,true)|p(X9)|p(X8)|~p(X6)|~p(X7)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_617, c_0_24])).
cnf(c_0_624,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk41_5(X1,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_618, c_0_483]), c_0_24])]), c_0_484])).
cnf(c_0_625,plain,(~esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk101_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_349, c_0_619])).
cnf(c_0_626,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk142_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_225, c_0_601])).
cnf(c_0_627,plain,(~epred27_0|~epred26_0|~epred24_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_620, c_0_621])])).
cnf(c_0_628,plain,(epred26_0), inference(spm,[status(thm)],[c_0_622, c_0_24])).
cnf(c_0_629,plain,(epred27_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_613])).
cnf(c_0_630,plain,(esk4_1(X1)|esk99_11(X1,X2,X3,X4,X5,true,X6,true,true,X7,true)|esk98_11(X1,X2,X3,X4,X5,true,X6,true,true,X7,true)|esk107_11(X1,X2,X3,X4,X5,true,X6,true,true,X7,true)|p(X7)|p(X8)|~p(X6)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_623, c_0_24])).
cnf(c_0_631,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk41_5(X1,true,true,true,true)|~esk35_5(X1,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_492]), c_0_24])]), c_0_484])).
cnf(c_0_632,plain,(~esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk101_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_625, c_0_364])).
cnf(c_0_633,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_626, c_0_224])).
cnf(c_0_634,plain,(~epred27_0|~epred24_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_627, c_0_628])])).
cnf(c_0_635,plain,(epred27_0), inference(spm,[status(thm)],[c_0_629, c_0_24])).
cnf(c_0_636,plain,(esk4_1(X1)|esk99_11(X1,X2,X3,X4,X5,true,true,true,true,X6,true)|esk98_11(X1,X2,X3,X4,X5,true,true,true,true,X6,true)|esk107_11(X1,X2,X3,X4,X5,true,true,true,true,X6,true)|p(X7)|p(X6)|~p(X5)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_630, c_0_24])).
cnf(c_0_637,plain,(esk75_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|esk395_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk85_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk55_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_226, c_0_227])).
cnf(c_0_638,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk23_3(X1,true,true)|~esk35_5(X1,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_631, c_0_500]), c_0_24])]), c_0_501])).
cnf(c_0_639,plain,(p(X1)|~esk99_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk101_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk134_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13,X14)|~esk109_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X13)|~p(X14)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_632, c_0_376]), c_0_440])).
cnf(c_0_640,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_633, c_0_376])).
cnf(c_0_641,plain,(~epred24_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_634, c_0_635])])).
cnf(c_0_642,plain,(esk4_1(X1)|esk99_11(X1,X2,X3,X4,true,true,true,true,true,X5,true)|esk98_11(X1,X2,X3,X4,true,true,true,true,true,X5,true)|esk107_11(X1,X2,X3,X4,true,true,true,true,true,X5,true)|p(X5)|p(X6)|~p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_636, c_0_24])).
cnf(c_0_643,plain,(esk75_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk85_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk55_7(X1,X2,X3,X4,X5,X6,X7)|~esk61_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_228, c_0_637])).
cnf(c_0_644,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk17_3(X1,true,true)|~esk42_5(X1,true,true,true,true)|~esk23_3(X1,true,true)), inference(spm,[status(thm)],[c_0_638, c_0_507])).
cnf(c_0_645,plain,(p(X1)|~esk99_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk101_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk109_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X13)|~p(X14)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_639, c_0_640]), c_0_388]), c_0_440])).
cnf(c_0_646,plain,(~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_611]), c_0_641])).
cnf(c_0_647,plain,(esk4_1(X1)|esk99_11(X1,X2,X3,true,true,true,true,true,true,X4,true)|esk98_11(X1,X2,X3,true,true,true,true,true,true,X4,true)|esk107_11(X1,X2,X3,true,true,true,true,true,true,X4,true)|p(X5)|p(X4)|~p(X3)|~p(X2)), inference(spm,[status(thm)],[c_0_642, c_0_24])).
cnf(c_0_648,plain,(~esk75_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk76_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_229, c_0_230])).
cnf(c_0_649,plain,(esk75_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk55_7(X1,X2,X3,X4,X5,X6,X7)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk61_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_643, c_0_231])).
cnf(c_0_650,plain,(esk76_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X8)|~esk91_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X9)), inference(spm,[status(thm)],[c_0_232, c_0_233])).
cnf(c_0_651,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)|~esk31_5(X1,true,true,true,true)), inference(spm,[status(thm)],[c_0_644, c_0_171])).
cnf(c_0_652,plain,(~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk99_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk101_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk109_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X14)|~p(X13)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_632, c_0_439]), c_0_645])).
cnf(c_0_653,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|~esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_614, c_0_376]), c_0_646]), c_0_646]), c_0_646]), c_0_440])).
cnf(c_0_654,plain,(esk4_1(X1)|esk99_11(X1,X2,true,true,true,true,true,true,true,X3,true)|esk98_11(X1,X2,true,true,true,true,true,true,true,X3,true)|esk107_11(X1,X2,true,true,true,true,true,true,true,X3,true)|p(X3)|p(X4)|~p(X2)), inference(spm,[status(thm)],[c_0_647, c_0_24])).
cnf(c_0_655,plain,(~esk55_7(X1,X2,X3,X4,X5,X6,X7)|~esk76_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk61_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_648, c_0_649])).
cnf(c_0_656,plain,(esk76_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X8)|~esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X9)), inference(spm,[status(thm)],[c_0_650, c_0_234])).
cnf(c_0_657,plain,(p(X1)|~esk73_9(X2,X3,X4,X5,X6,X7,X8,X1,X9)|~esk61_7(X2,X3,X4,X5,X6,X7,X8)|~p(X9)), inference(spm,[status(thm)],[c_0_235, c_0_236])).
cnf(c_0_658,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_651, c_0_494]), c_0_24])]), c_0_501])).
cnf(c_0_659,plain,(p(X1)|~esk99_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~esk109_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X13)|~p(X14)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_653]), c_0_645]), c_0_388]), c_0_388]), c_0_440])).
cnf(c_0_660,plain,(esk4_1(X1)|esk99_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)), inference(spm,[status(thm)],[c_0_654, c_0_24])).
cnf(c_0_661,plain,(p(X1)|~esk55_7(X2,X3,X4,X5,X6,X7,X8)|~esk74_9(X2,X3,X4,X5,X6,X7,X8,X1,X9)|~esk61_7(X2,X3,X4,X5,X6,X7,X8)|~p(X9)), inference(spm,[status(thm)],[c_0_655, c_0_656])).
cnf(c_0_662,plain,(esk74_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|p(X8)|~esk72_7(X1,X2,X3,X4,X5,X6,X7)|~esk61_7(X1,X2,X3,X4,X5,X6,X7)|~p(X9)), inference(spm,[status(thm)],[c_0_657, c_0_129])).
cnf(c_0_663,plain,(esk55_7(X1,X2,X3,X4,X5,X6,X7)|esk406_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk37_5(X1,X2,X3,X4,X5)|~esk66_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_237, c_0_238])).
cnf(c_0_664,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk24_3(X1,true,true)|~esk23_3(X1,true,true)|~esk7_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_658, c_0_518]), c_0_539])).
cnf(c_0_665,plain,(esk4_1(X1)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|p(X4)|~esk109_11(X1,true,true,true,true,true,true,true,true,X2,true)|~p(X5)|~p(X6)), inference(spm,[status(thm)],[c_0_659, c_0_660])).
cnf(c_0_666,plain,(p(X1)|~esk72_7(X2,X3,X4,X5,X6,X7,X8)|~esk55_7(X2,X3,X4,X5,X6,X7,X8)|~esk61_7(X2,X3,X4,X5,X6,X7,X8)|~p(X9)), inference(spm,[status(thm)],[c_0_661, c_0_662])).
cnf(c_0_667,plain,(esk55_7(X1,X2,X3,X4,X5,X6,X7)|~esk37_5(X1,X2,X3,X4,X5)|~esk41_5(X1,X2,X3,X4,X5)|~esk66_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_239, c_0_663])).
cnf(c_0_668,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk23_3(X1,true,true)|~esk7_1(X1)|~esk14_3(X1,true,true)), inference(spm,[status(thm)],[c_0_664, c_0_187])).
cnf(c_0_669,plain,(esk4_1(X1)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X4)|p(X2)|~esk109_11(X1,true,true,true,true,true,true,true,true,X2,true)|~p(X5)), inference(spm,[status(thm)],[c_0_665, c_0_24])).
cnf(c_0_670,plain,(p(X1)|~esk52_7(X2,X3,X4,X5,X6,X7,X8)|~esk55_7(X2,X3,X4,X5,X6,X7,X8)|~esk61_7(X2,X3,X4,X5,X6,X7,X8)|~p(X9)), inference(spm,[status(thm)],[c_0_666, c_0_240])).
cnf(c_0_671,plain,(esk55_7(X1,X2,X3,X4,X5,X6,X7)|~esk52_7(X1,X2,X3,X4,X5,X6,X7)|~esk37_5(X1,X2,X3,X4,X5)|~esk41_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_667, c_0_241])).
cnf(c_0_672,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk7_1(X1)|~esk14_3(X1,true,true)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_531]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_673,plain,(esk4_1(X1)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|p(X4)|~esk109_11(X1,true,true,true,true,true,true,true,true,X2,true)), inference(spm,[status(thm)],[c_0_669, c_0_24])).
cnf(c_0_674,plain,(p(X1)|~esk52_7(X2,X3,X4,X5,X6,X7,X8)|~esk37_5(X2,X3,X4,X5,X6)|~esk41_5(X2,X3,X4,X5,X6)|~esk61_7(X2,X3,X4,X5,X6,X7,X8)|~p(X9)), inference(spm,[status(thm)],[c_0_670, c_0_671])).
cnf(c_0_675,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)|~esk7_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_672, c_0_526]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_676,plain,(esk37_5(X1,X2,X3,X4,X5)|esk40_5(X1,X2,X3,X4,X5)|~esk430_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_242, c_0_243])).
cnf(c_0_677,plain,(esk4_1(X1)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X4)|p(X2)|~esk83_9(X1,true,true,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_673, c_0_421]), c_0_24])]), c_0_424])).
cnf(c_0_678,plain,(esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_244, c_0_245])).
cnf(c_0_679,plain,(p(X1)|~esk50_5(X2,X3,X4,X5,X6)|~esk37_5(X2,X3,X4,X5,X6)|~esk41_5(X2,X3,X4,X5,X6)|~esk61_7(X2,X3,X4,X5,X6,X7,X8)|~p(X9)|~p(X7)|~p(X8)), inference(spm,[status(thm)],[c_0_674, c_0_483])).
cnf(c_0_680,plain,(esk4_1(X1)|esk61_7(X1,true,true,true,true,true,true)|p(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_675, c_0_546]), c_0_193]), c_0_535])).
cnf(c_0_681,plain,(esk433_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk37_5(X1,X2,X3,X4,X5)|esk40_5(X1,X2,X3,X4,X5)|~esk17_3(X1,X2,X3)|~esk42_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_676, c_0_496])).
cnf(c_0_682,plain,(esk4_1(X1)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|p(X4)|~esk62_7(X1,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_677, c_0_458]), c_0_24])]), c_0_452])).
cnf(c_0_683,plain,(esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk106_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_678, c_0_456])).
cnf(c_0_684,plain,(esk4_1(X1)|p(X2)|p(X3)|~esk37_5(X1,true,true,true,true)|~esk41_5(X1,true,true,true,true)|~p(X4)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_679, c_0_680]), c_0_24])]), c_0_484])).
cnf(c_0_685,plain,(esk37_5(X1,X2,X3,X4,X5)|esk40_5(X1,X2,X3,X4,X5)|~esk17_3(X1,X2,X3)|~esk42_5(X1,X2,X3,X4,X5)|~esk23_3(X1,X2,X3)), inference(spm,[status(thm)],[c_0_168, c_0_681])).
cnf(c_0_686,plain,(esk4_1(X1)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X4)|p(X2)|~esk41_5(X1,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_682, c_0_492]), c_0_24])]), c_0_484])).
cnf(c_0_687,plain,(esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_646, c_0_683])).
cnf(c_0_688,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)|~esk17_3(X1,true,true)|~esk42_5(X1,true,true,true,true)|~esk23_3(X1,true,true)|~esk41_5(X1,true,true,true,true)|~p(X4)), inference(spm,[status(thm)],[c_0_684, c_0_685])).
cnf(c_0_689,plain,(~esk276_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_246, c_0_247])).
cnf(c_0_690,plain,(esk4_1(X1)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|p(X4)|~esk23_3(X1,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_686, c_0_500]), c_0_24])]), c_0_501])).
cnf(c_0_691,plain,(esk97_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_687, c_0_141]), c_0_92])).
cnf(c_0_692,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)|~esk41_5(X1,true,true,true,true)|~esk31_5(X1,true,true,true,true)|~p(X4)), inference(spm,[status(thm)],[c_0_688, c_0_171])).
cnf(c_0_693,plain,(esk196_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk205_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_689, c_0_248])).
cnf(c_0_694,plain,(esk4_1(X1)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X4)|p(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_531]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_695,plain,(esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk98_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_553, c_0_691])).
cnf(c_0_696,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)|~esk41_5(X1,true,true,true,true)|~p(X4)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_494]), c_0_24])]), c_0_501])).
cnf(c_0_697,plain,(~esk196_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk197_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_249, c_0_250])).
cnf(c_0_698,plain,(esk196_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_693, c_0_251])).
cnf(c_0_699,plain,(esk197_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk211_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~p(X20)|~p(X19)), inference(spm,[status(thm)],[c_0_252, c_0_253])).
cnf(c_0_700,plain,(esk4_1(X1)|esk98_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)), inference(ef,[status(thm)],[c_0_694])).
cnf(c_0_701,plain,(esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk110_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~p(X10)), inference(spm,[status(thm)],[c_0_695, c_0_560])).
cnf(c_0_702,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)|~p(X4)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_500]), c_0_24])]), c_0_501])).
cnf(c_0_703,plain,(~esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk197_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_697, c_0_698])).
cnf(c_0_704,plain,(esk197_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~p(X20)|~p(X19)), inference(spm,[status(thm)],[c_0_699, c_0_254])).
cnf(c_0_705,plain,(~esk194_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X19)|~p(X20)), inference(spm,[status(thm)],[c_0_255, c_0_256])).
fof(c_0_706, plain, ((~epred2_0<=>![X20]:~p(X20))), introduced(definition)).
fof(c_0_707, plain, ((~epred3_0<=>![X19]:~p(X19))), introduced(definition)).
cnf(c_0_708,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|~esk78_9(X1,true,true,true,true,true,true,true,true)|~esk96_11(X1,true,true,true,true,true,true,true,true,X2,true)|~esk110_11(X1,true,true,true,true,true,true,true,true,X2,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_700]), c_0_701])).
cnf(c_0_709,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)), inference(spm,[status(thm)],[c_0_702, c_0_24])).
fof(c_0_710, plain, ((~epred1_0<=>![X18]:![X17]:![X16]:![X15]:![X14]:![X13]:![X12]:![X11]:![X10]:![X9]:![X8]:![X7]:![X6]:![X5]:![X4]:![X3]:![X2]:![X1]:((~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18))|~esk193_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)))), introduced(definition)).
cnf(c_0_711,plain,(~esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~p(X20)|~p(X19)), inference(spm,[status(thm)],[c_0_703, c_0_704])).
cnf(c_0_712,plain,(esk195_20(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~esk193_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X19)|~p(X20)), inference(spm,[status(thm)],[c_0_705, c_0_257])).
cnf(c_0_713,plain,(epred2_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_706])).
cnf(c_0_714,plain,(epred3_0|~p(X1)), inference(split_equiv,[status(thm)],[c_0_707])).
cnf(c_0_715,plain,(esk173_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|p(X17)|~esk299_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_259]), c_0_260])).
cnf(c_0_716,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|~esk78_9(X1,true,true,true,true,true,true,true,true)|~esk110_11(X1,true,true,true,true,true,true,true,true,X2,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_708, c_0_407]), c_0_24])]), c_0_424])).
cnf(c_0_717,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)|~esk24_3(X1,true,true)|~esk23_3(X1,true,true)|~esk7_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_709, c_0_518]), c_0_539])).
cnf(c_0_718,plain,(~epred3_0|~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(spm,[status(thm)],[c_0_711, c_0_712]), c_0_710]), c_0_706]), c_0_707])).
cnf(c_0_719,plain,(epred2_0), inference(spm,[status(thm)],[c_0_713, c_0_24])).
cnf(c_0_720,plain,(epred3_0), inference(spm,[status(thm)],[c_0_714, c_0_24])).
cnf(c_0_721,plain,(esk173_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|p(X18)|p(X17)), inference(spm,[status(thm)],[c_0_715, c_0_261])).
cnf(c_0_722,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)|~esk78_9(X1,true,true,true,true,true,true,true,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_716, c_0_504]), c_0_24])]), c_0_424])).
cnf(c_0_723,plain,(esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk86_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk61_7(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_262, c_0_505])).
cnf(c_0_724,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)|~esk23_3(X1,true,true)|~esk7_1(X1)|~esk14_3(X1,true,true)), inference(spm,[status(thm)],[c_0_717, c_0_187])).
cnf(c_0_725,plain,(esk283_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk189_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_263, c_0_264])).
cnf(c_0_726,plain,(~epred1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_718, c_0_719])]), c_0_720])])).
cnf(c_0_727,plain,(esk290_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk186_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_265, c_0_266])).
cnf(c_0_728,plain,(esk283_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk189_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_267, c_0_264])).
cnf(c_0_729,plain,(esk172_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X17)|p(X18)), inference(spm,[status(thm)],[c_0_268, c_0_721])).
cnf(c_0_730,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|~esk86_9(X1,true,true,true,true,true,true,true,true)|~esk56_7(X1,true,true,true,true,true,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_722, c_0_723]), c_0_680])).
cnf(c_0_731,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)|~esk7_1(X1)|~esk14_3(X1,true,true)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_724, c_0_531]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_732,plain,(esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|p(X18)|~esk189_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X17)), inference(spm,[status(thm)],[c_0_269, c_0_725])).
cnf(c_0_733,plain,(~esk193_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_710]), c_0_726])).
cnf(c_0_734,plain,(esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk186_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_270, c_0_727])).
cnf(c_0_735,plain,(esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|p(X18)|~esk189_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X17)), inference(spm,[status(thm)],[c_0_269, c_0_728])).
cnf(c_0_736,plain,(p(X1)|~esk173_18(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X1)|~esk158_16(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~p(X18)), inference(spm,[status(thm)],[c_0_271, c_0_272])).
cnf(c_0_737,plain,(esk172_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X17)), inference(ef,[status(thm)],[c_0_729])).
cnf(c_0_738,plain,(esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk314_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk163_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_273, c_0_274])).
cnf(c_0_739,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)|~esk56_7(X1,true,true,true,true,true,true)|~esk74_9(X1,true,true,true,true,true,true,true,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)), inference(spm,[status(thm)],[c_0_730, c_0_128])).
cnf(c_0_740,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)|~esk7_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_526]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_741,plain,(esk175_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|p(X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X17)), inference(spm,[status(thm)],[c_0_732, c_0_275])).
cnf(c_0_742,plain,(~esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_733, c_0_276])).
cnf(c_0_743,plain,(esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|esk177_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_734, c_0_277])).
cnf(c_0_744,plain,(esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|p(X18)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~p(X17)), inference(spm,[status(thm)],[c_0_735, c_0_275])).
cnf(c_0_745,plain,(esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|p(X18)|~esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X17)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_736, c_0_19]), c_0_737])).
cnf(c_0_746,plain,(esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk316_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk163_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_278, c_0_738])).
cnf(c_0_747,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|~esk56_7(X1,true,true,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_739, c_0_451]), c_0_24])]), c_0_452])).
cnf(c_0_748,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)|p(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_740, c_0_546]), c_0_193]), c_0_535])).
cnf(c_0_749,plain,(p(X1)|~esk176_18(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X1)|~esk174_18(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X1)|~p(X18)), inference(spm,[status(thm)],[c_0_313, c_0_741])).
cnf(c_0_750,plain,(esk176_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk174_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_742, c_0_743])).
cnf(c_0_751,plain,(esk183_18(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|p(X18)|~esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~p(X17)), inference(spm,[status(thm)],[c_0_744, c_0_745])).
cnf(c_0_752,plain,(esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk163_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_279, c_0_746])).
cnf(c_0_753,plain,(esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk164_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_280, c_0_335])).
cnf(c_0_754,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)|~esk56_7(X1,true,true,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_747, c_0_458]), c_0_24])]), c_0_452])).
cnf(c_0_755,plain,(esk56_7(X1,X2,X3,X4,X5,X6,X7)|~esk64_7(X1,X2,X3,X4,X5,X6,X7)|~esk35_5(X1,X2,X3,X4,X5)|~esk40_5(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_281, c_0_469])).
cnf(c_0_756,plain,(esk4_1(X1)|esk40_5(X1,true,true,true,true)|p(X2)), inference(ef,[status(thm)],[c_0_748])).
cnf(c_0_757,plain,(p(X1)|~esk155_16(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~esk158_16(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~p(X18)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_749, c_0_750]), c_0_751]), c_0_745])).
cnf(c_0_758,plain,(esk155_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_752, c_0_282])).
cnf(c_0_759,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk311_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk165_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_283, c_0_284])).
cnf(c_0_760,plain,(esk153_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_753, c_0_49])).
cnf(c_0_761,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|~esk64_7(X1,true,true,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_754, c_0_755]), c_0_756])).
cnf(c_0_762,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X17)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X18)), inference(spm,[status(thm)],[c_0_757, c_0_758])).
cnf(c_0_763,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk165_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_285, c_0_759])).
cnf(c_0_764,plain,(esk152_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_354, c_0_760])).
cnf(c_0_765,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)|~esk52_7(X1,true,true,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)), inference(spm,[status(thm)],[c_0_761, c_0_152])).
cnf(c_0_766,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X17)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_762, c_0_24])).
cnf(c_0_767,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_763, c_0_286])).
cnf(c_0_768,plain,(~esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk128_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk160_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_336, c_0_764])).
cnf(c_0_769,plain,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk329_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk144_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_287, c_0_88])).
cnf(c_0_770,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|~esk41_5(X1,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_765, c_0_483]), c_0_24])]), c_0_484])).
cnf(c_0_771,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X17)|~esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X16)|~p(X15)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_766, c_0_360]), c_0_396])).
cnf(c_0_772,plain,(esk158_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X16)), inference(spm,[status(thm)],[c_0_344, c_0_767])).
cnf(c_0_773,plain,(p(X1)|~esk150_16(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X1)|~esk128_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_768, c_0_384]), c_0_385])).
cnf(c_0_774,plain,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X13)|~esk144_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X12)|~p(X14)), inference(spm,[status(thm)],[c_0_102, c_0_769])).
cnf(c_0_775,plain,(esk342_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk138_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk102_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_288, c_0_289])).
cnf(c_0_776,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)|~esk41_5(X1,true,true,true,true)|~esk35_5(X1,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_770, c_0_492]), c_0_24])]), c_0_484])).
cnf(c_0_777,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X17)|~esk150_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X15)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_771, c_0_772]), c_0_384])).
cnf(c_0_778,plain,(esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk331_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk143_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_290, c_0_291])).
cnf(c_0_779,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk331_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~esk143_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_292, c_0_291])).
cnf(c_0_780,plain,(p(X1)|~esk128_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~esk132_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~p(X16)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_773, c_0_395]), c_0_396])).
cnf(c_0_781,plain,(esk124_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X13)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X12)|~p(X14)), inference(spm,[status(thm)],[c_0_774, c_0_117])).
cnf(c_0_782,plain,(esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk138_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk102_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_293, c_0_775])).
cnf(c_0_783,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|~esk23_3(X1,true,true)|~esk35_5(X1,true,true,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_776, c_0_500]), c_0_24])]), c_0_501])).
cnf(c_0_784,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X17)|~esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X15)|~p(X16)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_777, c_0_360]), c_0_396])).
cnf(c_0_785,plain,(esk127_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|p(X13)|~esk143_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X14)), inference(spm,[status(thm)],[c_0_294, c_0_778])).
cnf(c_0_786,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|p(X13)|~esk143_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X14)), inference(spm,[status(thm)],[c_0_294, c_0_779])).
cnf(c_0_787,plain,(p(X1)|p(X2)|~esk144_14(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X1,X15)|~esk132_14(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X1,X15)|~p(X16)|~p(X14)|~p(X15)), inference(spm,[status(thm)],[c_0_780, c_0_416])).
cnf(c_0_788,plain,(p(X1)|~esk121_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X1,X14)|~esk108_11(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_295, c_0_296])).
cnf(c_0_789,plain,(p(X1)|~esk123_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X1,X14)|~esk122_14(X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X1,X14)|~p(X13)|~p(X14)), inference(spm,[status(thm)],[c_0_349, c_0_781])).
cnf(c_0_790,plain,(esk123_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk102_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_782, c_0_297])).
cnf(c_0_791,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)|~esk17_3(X1,true,true)|~esk42_5(X1,true,true,true,true)|~esk23_3(X1,true,true)), inference(spm,[status(thm)],[c_0_783, c_0_507])).
cnf(c_0_792,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X13)|p(X12)|p(X17)|~esk143_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X15)|~p(X16)|~p(X14)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_784, c_0_785]), c_0_786])).
cnf(c_0_793,plain,(esk338_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk104_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk140_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_298, c_0_299])).
cnf(c_0_794,plain,(p(X1)|p(X2)|~esk132_14(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X2,X15)|~esk122_14(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X2,X15)|~p(X16)|~p(X14)|~p(X15)), inference(spm,[status(thm)],[c_0_787, c_0_117])).
cnf(c_0_795,plain,(esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X13)|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X12)|~p(X14)), inference(spm,[status(thm)],[c_0_788, c_0_63])).
cnf(c_0_796,plain,(esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X13)|~esk102_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X12)|~p(X14)), inference(spm,[status(thm)],[c_0_789, c_0_790])).
cnf(c_0_797,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)|~esk31_5(X1,true,true,true,true)), inference(spm,[status(thm)],[c_0_791, c_0_171])).
cnf(c_0_798,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X17)|p(X12)|p(X13)|~esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X15)|~p(X16)|~p(X14)), inference(spm,[status(thm)],[c_0_792, c_0_300])).
cnf(c_0_799,plain,(esk129_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk104_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk140_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(spm,[status(thm)],[c_0_301, c_0_793])).
cnf(c_0_800,plain,(p(X1)|p(X2)|~esk121_14(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X1,X2,X14)|~esk107_11(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~p(X14)), inference(spm,[status(thm)],[c_0_302, c_0_303])).
cnf(c_0_801,plain,(p(X1)|p(X2)|~esk108_11(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~esk132_14(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X1,X15)|~p(X16)|~p(X14)|~p(X15)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_795]), c_0_440])).
cnf(c_0_802,plain,(esk132_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X13)|~esk102_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X12)|~p(X14)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_796, c_0_795]), c_0_440])).
cnf(c_0_803,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_797, c_0_494]), c_0_24])]), c_0_501])).
cnf(c_0_804,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X13)|p(X12)|p(X17)|~esk104_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X15)|~p(X16)|~p(X14)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_798, c_0_799]), c_0_304])).
cnf(c_0_805,plain,(esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X13)|p(X12)|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X14)), inference(spm,[status(thm)],[c_0_800, c_0_63])).
cnf(c_0_806,plain,(esk104_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk368_23(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_305, c_0_306])).
cnf(c_0_807,plain,(p(X1)|p(X2)|~esk102_11(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~esk108_11(X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~p(X14)|~p(X15)|~p(X16)), inference(spm,[status(thm)],[c_0_801, c_0_802])).
cnf(c_0_808,plain,(esk102_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X10)|~esk117_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X11)), inference(spm,[status(thm)],[c_0_307, c_0_378])).
cnf(c_0_809,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|~esk24_3(X1,true,true)|~esk23_3(X1,true,true)|~esk7_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_803, c_0_518]), c_0_539])).
cnf(c_0_810,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X17)|p(X12)|p(X13)|~esk104_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X15)|~p(X16)|~p(X14)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_804, c_0_805]), c_0_440])).
cnf(c_0_811,plain,(esk104_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_806, c_0_456])).
cnf(c_0_812,plain,(p(X1)|p(X2)|p(X3)|~esk108_11(X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13)|~esk117_11(X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13)|~p(X14)|~p(X15)|~p(X16)|~p(X13)), inference(spm,[status(thm)],[c_0_807, c_0_808])).
cnf(c_0_813,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)|~esk23_3(X1,true,true)|~esk7_1(X1)|~esk14_3(X1,true,true)), inference(spm,[status(thm)],[c_0_809, c_0_187])).
cnf(c_0_814,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X13)|p(X12)|p(X17)|~esk112_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X15)|~p(X16)|~p(X14)), inference(spm,[status(thm)],[c_0_810, c_0_811])).
cnf(c_0_815,plain,(p(X1)|p(X2)|p(X3)|~esk96_11(X4,X5,X6,X7,X8,X9,X10,X11,X12,X3,X13)|~esk108_11(X4,X5,X6,X7,X8,X9,X10,X11,X12,X3,X13)|~p(X14)|~p(X15)|~p(X16)|~p(X13)), inference(spm,[status(thm)],[c_0_812, c_0_94])).
cnf(c_0_816,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|~esk7_1(X1)|~esk14_3(X1,true,true)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_813, c_0_531]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_817,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|esk108_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|p(X17)|p(X12)|p(X13)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk96_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X15)|~p(X16)|~p(X14)), inference(spm,[status(thm)],[c_0_814, c_0_141])).
cnf(c_0_818,plain,(p(X1)|p(X2)|p(X3)|~esk108_11(X4,X5,X6,X7,X8,X9,X10,X11,X12,X1,X13)|~esk83_9(X4,X5,X6,X7,X8,X9,X10,X11,X12)|~p(X14)|~p(X15)|~p(X16)|~p(X13)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_815, c_0_407]), c_0_424])).
cnf(c_0_819,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)|~esk7_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_816, c_0_526]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_820,plain,(esk151_16(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|p(X10)|p(X13)|p(X12)|p(X17)|~esk78_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk83_9(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X15)|~p(X16)|~p(X14)|~p(X11)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_817, c_0_407]), c_0_424]), c_0_818])).
cnf(c_0_821,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|esk107_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_819, c_0_546]), c_0_193]), c_0_535])).
cnf(c_0_822,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X8)|p(X9)|p(X3)|p(X4)|p(X2)|~esk78_9(X1,true,true,true,true,true,true,true,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_820, c_0_821]), c_0_24])])).
cnf(c_0_823,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X8)|p(X9)|~esk86_9(X1,true,true,true,true,true,true,true,true)|~esk56_7(X1,true,true,true,true,true,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_822, c_0_723]), c_0_680])).
cnf(c_0_824,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X8)|p(X9)|p(X3)|p(X4)|p(X2)|~esk56_7(X1,true,true,true,true,true,true)|~esk74_9(X1,true,true,true,true,true,true,true,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(spm,[status(thm)],[c_0_823, c_0_128])).
cnf(c_0_825,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X8)|p(X9)|~esk56_7(X1,true,true,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)|~esk83_9(X1,true,true,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_824, c_0_451]), c_0_24])]), c_0_452])).
cnf(c_0_826,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X8)|p(X9)|p(X3)|p(X4)|p(X2)|~esk56_7(X1,true,true,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_825, c_0_458]), c_0_24])]), c_0_452])).
cnf(c_0_827,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X8)|p(X9)|~esk64_7(X1,true,true,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_826, c_0_755]), c_0_756])).
cnf(c_0_828,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X8)|p(X9)|p(X3)|p(X4)|p(X2)|~esk52_7(X1,true,true,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(spm,[status(thm)],[c_0_827, c_0_152])).
cnf(c_0_829,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X8)|p(X9)|~esk41_5(X1,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~esk62_7(X1,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_828, c_0_483]), c_0_24])]), c_0_484])).
cnf(c_0_830,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X8)|p(X9)|p(X3)|p(X4)|p(X2)|~esk41_5(X1,true,true,true,true)|~esk35_5(X1,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_829, c_0_492]), c_0_24])]), c_0_484])).
cnf(c_0_831,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X8)|p(X9)|~esk23_3(X1,true,true)|~esk35_5(X1,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_830, c_0_500]), c_0_24])]), c_0_501])).
cnf(c_0_832,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X8)|p(X9)|p(X3)|p(X4)|p(X2)|~esk17_3(X1,true,true)|~esk42_5(X1,true,true,true,true)|~esk23_3(X1,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(spm,[status(thm)],[c_0_831, c_0_507])).
cnf(c_0_833,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X8)|p(X9)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)|~esk31_5(X1,true,true,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(spm,[status(thm)],[c_0_832, c_0_171])).
cnf(c_0_834,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X8)|p(X9)|p(X3)|p(X4)|p(X2)|~esk17_3(X1,true,true)|~esk23_3(X1,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_833, c_0_494]), c_0_24])]), c_0_501])).
cnf(c_0_835,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X8)|p(X9)|~esk24_3(X1,true,true)|~esk23_3(X1,true,true)|~esk7_1(X1)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_834, c_0_518]), c_0_539])).
cnf(c_0_836,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X8)|p(X9)|p(X3)|p(X4)|p(X2)|~esk23_3(X1,true,true)|~esk7_1(X1)|~esk14_3(X1,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(spm,[status(thm)],[c_0_835, c_0_187])).
cnf(c_0_837,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X8)|p(X9)|~esk7_1(X1)|~esk14_3(X1,true,true)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_836, c_0_531]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_838,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X8)|p(X9)|p(X3)|p(X4)|p(X2)|~esk7_1(X1)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_837, c_0_526]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_839,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,X6,X7)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X8)|p(X9)|~p(X6)|~p(X7)|~p(X5)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_838, c_0_546]), c_0_193]), c_0_535])).
cnf(c_0_840,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,true,X6)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X7)|p(X8)|p(X3)|p(X4)|p(X2)|~p(X6)|~p(X5)), inference(spm,[status(thm)],[c_0_839, c_0_24])).
cnf(c_0_841,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,X5,true,true)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X6)|p(X7)|~p(X5)), inference(spm,[status(thm)],[c_0_840, c_0_24])).
cnf(c_0_842,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,true,true,true)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X5)|p(X6)|p(X3)|p(X4)|p(X2)), inference(spm,[status(thm)],[c_0_841, c_0_24])).
cnf(c_0_843,plain,(esk151_16(X1,true,true,true,true,true,true,true,true,X2,true,X3,X4,true,true,true)|esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X4)|p(X3)|p(X5)), inference(ef,[status(thm)],[c_0_842])).
cnf(c_0_844,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X4)|p(X5)|p(X2)|~esk150_16(X1,true,true,true,true,true,true,true,true,X2,true,X4,X5,true,true,true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_843]), c_0_24])])).
cnf(c_0_845,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X13)|p(X12)|~esk122_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~p(X14)), inference(spm,[status(thm)],[c_0_786, c_0_300])).
cnf(c_0_846,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|p(X4)|p(X5)|~esk134_14(X1,true,true,true,true,true,true,true,true,X2,true,X4,X3,true)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_844, c_0_360]), c_0_24])]), c_0_396])).
cnf(c_0_847,plain,(esk134_14(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|p(X12)|p(X13)|~esk120_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~esk107_11(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~p(X14)), inference(spm,[status(thm)],[c_0_845, c_0_805])).
cnf(c_0_848,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X4)|p(X5)|p(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_846, c_0_847]), c_0_24])]), c_0_821]), c_0_440])).
cnf(c_0_849,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X2)|p(X3)|p(X4)), inference(ef,[status(thm)],[c_0_848])).
cnf(c_0_850,plain,(esk4_1(X1)|esk108_11(X1,true,true,true,true,true,true,true,true,X2,true)|p(X3)|p(X2)), inference(ef,[status(thm)],[c_0_849])).
cnf(c_0_851,plain,(esk4_1(X1)|p(X2)|p(X3)|p(X4)|p(X5)|~esk83_9(X1,true,true,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_818, c_0_850]), c_0_24])])).
cnf(c_0_852,plain,(esk4_1(X1)|p(X2)|p(X3)|p(X4)|p(X5)|~esk62_7(X1,true,true,true,true,true,true)|~p(X6)|~p(X7)|~p(X8)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_851, c_0_458]), c_0_24])]), c_0_452])).
cnf(c_0_853,plain,(esk4_1(X1)|p(X2)|p(X3)|p(X4)|p(X5)|~esk41_5(X1,true,true,true,true)|~p(X6)|~p(X7)|~p(X8)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_852, c_0_492]), c_0_24])]), c_0_484])).
cnf(c_0_854,plain,(esk4_1(X1)|p(X2)|p(X3)|p(X4)|p(X5)|~esk23_3(X1,true,true)|~p(X6)|~p(X7)|~p(X8)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_853, c_0_500]), c_0_24])]), c_0_501])).
cnf(c_0_855,plain,(esk4_1(X1)|p(X2)|p(X3)|p(X4)|p(X5)|~p(X6)|~p(X7)|~p(X8)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_854, c_0_531]), c_0_24])]), c_0_532]), c_0_539])).
cnf(c_0_856,plain,(esk4_1(X1)|p(X2)|p(X3)|p(X4)|p(X5)|~p(X6)|~p(X7)), inference(spm,[status(thm)],[c_0_855, c_0_24])).
cnf(c_0_857,plain,(esk4_1(X1)|p(X2)|p(X3)|p(X4)|p(X5)|~p(X6)), inference(spm,[status(thm)],[c_0_856, c_0_24])).
cnf(c_0_858,plain,(esk4_1(X1)|p(X2)|p(X3)|p(X4)|p(X5)), inference(spm,[status(thm)],[c_0_857, c_0_24])).
cnf(c_0_859,plain,(esk4_1(X1)|p(X2)|p(X3)|p(X4)), inference(ef,[status(thm)],[c_0_858])).
cnf(c_0_860,plain,(esk4_1(X1)|p(X2)|p(X3)), inference(ef,[status(thm)],[c_0_859])).
cnf(c_0_861,plain,(esk4_1(X1)|p(X2)), inference(ef,[status(thm)],[c_0_860])).
cnf(c_0_862,plain,(p(X1)|~p(X2)), inference(spm,[status(thm)],[c_0_538, c_0_861])).
cnf(c_0_863,plain,(p(X1)), inference(spm,[status(thm)],[c_0_862, c_0_24])).
cnf(c_0_864,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_308, c_0_863])]), ['proof']).
# SZS output end CNFRefutation.
EOF
