% Problem    : PLA031-1.004 : TPTP v6.1.0. Released v3.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n145.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:32:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Presaturation interreduction done

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(s(X1,X2,X3,truck1,X5,X6,X7,X4,X8)|~s(X1,X2,X3,X4,X5,X6,X7,X4,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', load13)).
cnf(c_0_1,axiom,(s(s1,s0,s2,s0,s1,s2,s2,s0,s1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', init)).
cnf(c_0_2,axiom,(s(X1,X2,X3,X4,X5,X6,X10,X8,X9)|~s(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~path(X7,X10)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', walk3)).
cnf(c_0_3,axiom,(s(X1,X2,X3,X4,X5,X6,truck1,X7,X8)|~s(X1,X2,X3,X4,X5,X6,X7,X7,X8)|~neq(X5,truck1)|~neq(X6,truck1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', board9)).
cnf(c_0_4,axiom,(neq(s2,truck1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', neq37)).
cnf(c_0_5,axiom,(neq(s1,truck1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', neq28)).
cnf(c_0_6,axiom,(path(s2,p2_0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', map9)).
cnf(c_0_7,axiom,(path(p2_0,s0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', map12)).
cnf(c_0_8,axiom,(s(X1,X2,X3,X4,X5,X6,truck1,X9,X8)|~s(X1,X2,X3,X4,X5,X6,truck1,X7,X8)|~link(X7,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', drive5)).
cnf(c_0_9,axiom,(link(s0,s2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', map13)).
cnf(c_0_10,axiom,(s(X1,X2,X3,X4,X5,X6,X7,X4,X8)|~s(X1,X2,X3,truck1,X5,X6,X7,X4,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', load14)).
cnf(c_0_11,axiom,(s(X1,X2,truck1,X4,X5,X6,X7,X3,X8)|~s(X1,X2,X3,X4,X5,X6,X7,X3,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', load9)).
cnf(c_0_12,axiom,(link(s2,s1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', map17)).
cnf(c_0_13,axiom,(s(X1,X2,X3,X4,X5,X6,X7,X3,X8)|~s(X1,X2,truck1,X4,X5,X6,X7,X3,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', load10)).
cnf(c_0_14,axiom,(s(truck1,X2,X3,X4,X5,X6,X7,X1,X8)|~s(X1,X2,X3,X4,X5,X6,X7,X1,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', load1)).
cnf(c_0_15,axiom,(link(s1,s0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', map15)).
cnf(c_0_16,axiom,(s(X1,X2,X3,X4,X5,X6,X7,X1,X8)|~s(truck1,X2,X3,X4,X5,X6,X7,X1,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', load2)).
cnf(c_0_17,axiom,(s(X1,X2,X3,X4,X5,truck1,X7,X6,X8)|~s(X1,X2,X3,X4,X5,X6,X7,X6,X8)|~neq(X5,truck1)|~neq(X7,truck1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', board5)).
cnf(c_0_18,axiom,(s(X1,X2,X3,X4,X5,X6,X7,X7,X8)|~s(X1,X2,X3,X4,X5,X6,truck1,X7,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', board10)).
cnf(c_0_19,axiom,(s(X1,X2,X3,X4,X5,truck1,X6,X9,X8)|~s(X1,X2,X3,X4,X5,truck1,X6,X7,X8)|~link(X7,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', drive3)).
cnf(c_0_20,axiom,(s(X1,X2,X3,X4,truck1,X6,X7,X5,X8)|~s(X1,X2,X3,X4,X5,X6,X7,X5,X8)|~neq(X6,truck1)|~neq(X7,truck1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', board1)).
cnf(c_0_21,axiom,(s(X1,X2,X3,X4,X5,X6,X7,X6,X8)|~s(X1,X2,X3,X4,X5,truck1,X7,X6,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', board6)).
cnf(c_0_22,axiom,(s(X1,X2,X3,X4,truck1,X5,X6,X9,X8)|~s(X1,X2,X3,X4,truck1,X5,X6,X7,X8)|~link(X7,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', drive1)).
cnf(c_0_23,axiom,(link(s1,s2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', map18)).
cnf(c_0_24,axiom,(s(X1,X2,X3,X4,X5,X6,X7,X5,X8)|~s(X1,X2,X3,X4,truck1,X6,X7,X5,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', board2)).
cnf(c_0_25,axiom,(link(s2,s0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', map14)).
cnf(c_0_26,negated_conjecture,(~s(s0,s0,s1,s2,s2,X1,X2,s0,s1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goal)).
cnf(c_0_27,plain,(s(s1,s0,s2,truck1,s1,s2,s2,s0,s1)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_28,plain,(s(s1,s0,s2,truck1,s1,s2,X1,s0,s1)|~path(s2,X1)), inference(spm,[status(thm)],[c_0_2, c_0_27])).
cnf(c_0_29,plain,(s(X1,X2,X3,X4,X5,s2,truck1,X6,X7)|~neq(X5,truck1)|~s(X1,X2,X3,X4,X5,s2,X6,X6,X7)), inference(spm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_30,plain,(s(s1,s0,s2,truck1,s1,s2,X1,s0,s1)|~path(s2,X2)|~path(X2,X1)), inference(spm,[status(thm)],[c_0_2, c_0_28])).
cnf(c_0_31,plain,(s(X1,X2,X3,X4,s1,s2,truck1,X5,X6)|~s(X1,X2,X3,X4,s1,s2,X5,X5,X6)), inference(spm,[status(thm)],[c_0_29, c_0_5])).
cnf(c_0_32,plain,(s(s1,s0,s2,truck1,s1,s2,X1,s0,s1)|~path(p2_0,X1)), inference(spm,[status(thm)],[c_0_30, c_0_6])).
cnf(c_0_33,plain,(s(s1,s0,s2,truck1,s1,s2,truck1,s0,s1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_7])])).
cnf(c_0_34,plain,(s(s1,s0,s2,truck1,s1,s2,truck1,X1,s1)|~link(s0,X1)), inference(spm,[status(thm)],[c_0_8, c_0_33])).
cnf(c_0_35,plain,(s(s1,s0,s2,truck1,s1,s2,truck1,s2,s1)), inference(spm,[status(thm)],[c_0_34, c_0_9])).
cnf(c_0_36,plain,(s(s1,s0,s2,s2,s1,s2,truck1,s2,s1)), inference(spm,[status(thm)],[c_0_10, c_0_35])).
cnf(c_0_37,plain,(s(s1,s0,truck1,s2,s1,s2,truck1,s2,s1)), inference(spm,[status(thm)],[c_0_11, c_0_36])).
cnf(c_0_38,plain,(s(s1,s0,truck1,s2,s1,s2,truck1,X1,s1)|~link(s2,X1)), inference(spm,[status(thm)],[c_0_8, c_0_37])).
cnf(c_0_39,plain,(s(s1,s0,truck1,s2,s1,s2,truck1,s1,s1)), inference(spm,[status(thm)],[c_0_38, c_0_12])).
cnf(c_0_40,plain,(s(s1,s0,s1,s2,s1,s2,truck1,s1,s1)), inference(spm,[status(thm)],[c_0_13, c_0_39])).
cnf(c_0_41,plain,(s(truck1,s0,s1,s2,s1,s2,truck1,s1,s1)), inference(spm,[status(thm)],[c_0_14, c_0_40])).
cnf(c_0_42,plain,(s(truck1,s0,s1,s2,s1,s2,truck1,X1,s1)|~link(s1,X1)), inference(spm,[status(thm)],[c_0_8, c_0_41])).
cnf(c_0_43,plain,(s(truck1,s0,s1,s2,s1,s2,truck1,s0,s1)), inference(spm,[status(thm)],[c_0_42, c_0_15])).
cnf(c_0_44,plain,(s(s0,s0,s1,s2,s1,s2,truck1,s0,s1)), inference(spm,[status(thm)],[c_0_16, c_0_43])).
cnf(c_0_45,plain,(s(s0,s0,s1,s2,s1,s2,truck1,X1,s1)|~link(s0,X1)), inference(spm,[status(thm)],[c_0_8, c_0_44])).
cnf(c_0_46,plain,(s(X1,X2,X3,X4,X5,truck1,s2,X6,X7)|~neq(X5,truck1)|~s(X1,X2,X3,X4,X5,X6,s2,X6,X7)), inference(spm,[status(thm)],[c_0_17, c_0_4])).
cnf(c_0_47,plain,(s(s0,s0,s1,s2,s1,s2,truck1,s2,s1)), inference(spm,[status(thm)],[c_0_45, c_0_9])).
cnf(c_0_48,plain,(s(X1,X2,X3,X4,s1,truck1,s2,X5,X6)|~s(X1,X2,X3,X4,s1,X5,s2,X5,X6)), inference(spm,[status(thm)],[c_0_46, c_0_5])).
cnf(c_0_49,plain,(s(s0,s0,s1,s2,s1,s2,s2,s2,s1)), inference(spm,[status(thm)],[c_0_18, c_0_47])).
cnf(c_0_50,plain,(s(s0,s0,s1,s2,s1,truck1,s2,s2,s1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_51,plain,(s(s0,s0,s1,s2,s1,truck1,s2,X1,s1)|~link(s2,X1)), inference(spm,[status(thm)],[c_0_19, c_0_50])).
cnf(c_0_52,plain,(s(X1,X2,X3,X4,truck1,X5,s2,X6,X7)|~neq(X5,truck1)|~s(X1,X2,X3,X4,X6,X5,s2,X6,X7)), inference(spm,[status(thm)],[c_0_20, c_0_4])).
cnf(c_0_53,plain,(s(s0,s0,s1,s2,s1,truck1,s2,s1,s1)), inference(spm,[status(thm)],[c_0_51, c_0_12])).
cnf(c_0_54,plain,(s(X1,X2,X3,X4,truck1,s1,s2,X5,X6)|~s(X1,X2,X3,X4,X5,s1,s2,X5,X6)), inference(spm,[status(thm)],[c_0_52, c_0_5])).
cnf(c_0_55,plain,(s(s0,s0,s1,s2,s1,s1,s2,s1,s1)), inference(spm,[status(thm)],[c_0_21, c_0_53])).
cnf(c_0_56,plain,(s(s0,s0,s1,s2,truck1,s1,s2,s1,s1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_57,plain,(s(s0,s0,s1,s2,truck1,s1,s2,X1,s1)|~link(s1,X1)), inference(spm,[status(thm)],[c_0_22, c_0_56])).
cnf(c_0_58,plain,(s(X1,X2,X3,X4,X5,s1,truck1,X6,X7)|~neq(X5,truck1)|~s(X1,X2,X3,X4,X5,s1,X6,X6,X7)), inference(spm,[status(thm)],[c_0_3, c_0_5])).
cnf(c_0_59,plain,(s(s0,s0,s1,s2,truck1,s1,s2,s2,s1)), inference(spm,[status(thm)],[c_0_57, c_0_23])).
cnf(c_0_60,plain,(s(X1,X2,X3,X4,s2,s1,truck1,X5,X6)|~s(X1,X2,X3,X4,s2,s1,X5,X5,X6)), inference(spm,[status(thm)],[c_0_58, c_0_4])).
cnf(c_0_61,plain,(s(s0,s0,s1,s2,s2,s1,s2,s2,s1)), inference(spm,[status(thm)],[c_0_24, c_0_59])).
cnf(c_0_62,plain,(s(s0,s0,s1,s2,s2,s1,truck1,s2,s1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_63,plain,(s(s0,s0,s1,s2,s2,s1,truck1,X1,s1)|~link(s2,X1)), inference(spm,[status(thm)],[c_0_8, c_0_62])).
cnf(c_0_64,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_25]), c_0_26]), ['proof']).
# SZS output end CNFRefutation.
EOF
