% Problem    : SWV422-1.065 : TPTP v6.1.0. Released v3.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n150.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:15:57 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____011_C07_F1_PI_AE_OS_S1U
# and selection function SelectComplexAHPExceptRRHorn.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(~m_main_v_state2(X1,c_n2)|~m_main_v_state2(X1,c_t2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_6)).
cnf(c_0_1,axiom,(m_main_v_state2(s0,c_n2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_12)).
cnf(c_0_2,axiom,(~m_main_v_state1(X1,c_n1)|~m_main_v_state1(X1,c_t1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_2)).
cnf(c_0_3,axiom,(m_main_v_state1(s0,c_n1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_11)).
cnf(c_0_4,axiom,(~m_main_v_state2(X1,c_n2)|~m_main_v_state2(X1,c_c2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_7)).
cnf(c_0_5,axiom,(m_main_v_state2(X1,c_t2)|~node1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_15)).
cnf(c_0_6,axiom,(node1(X1)|node2(X1)|node3(X1)|node4(X1)|node5(X1)|m_main_v_state1(X1,c_c1)|node6(X1,X2)|~node7(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_33)).
cnf(c_0_7,axiom,(node7(X1,X2)|~trans(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_34)).
cnf(c_0_8,axiom,(trans(X1,X2)|~succ(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound66)).
cnf(c_0_9,axiom,(succ(s0,s1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound1)).
cnf(c_0_10,axiom,(m_main_v_state1(X1,c_t1)|~node5(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_22)).
cnf(c_0_11,axiom,(m_main_v_state1(X1,c_t1)|~node4(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_20)).
cnf(c_0_12,axiom,(m_main_v_state2(X1,c_c2)|~node3(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_19)).
cnf(c_0_13,axiom,(~m_main_v_state1(X1,c_n1)|~m_main_v_state1(X1,c_c1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_3)).
cnf(c_0_14,axiom,(node15(X1)|node16(X1)|node17(X1,X2)|~node18(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_64)).
cnf(c_0_15,axiom,(node18(X1,X2)|~trans(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_65)).
cnf(c_0_16,axiom,(m_main_v_state1(X1,c_t1)|~node16(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_59)).
cnf(c_0_17,axiom,(m_main_v_state2(X1,c_t2)|~node15(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_57)).
cnf(c_0_18,axiom,(node1(X1)|m_main_v_state1(X2,c_t1)|~m_main_v_state1(X1,c_n1)|~m_main_v_state2(X1,c_n2)|~node7(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_28)).
cnf(c_0_19,axiom,(m_main_v_state1(X1,c_t1)|~node8(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_36)).
cnf(c_0_20,axiom,(node8(X1)|m_main_v_state2(X2,c_t2)|~m_main_v_state2(X1,c_n2)|~m_main_v_state1(X1,c_n1)|~node14(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_49)).
cnf(c_0_21,axiom,(m_main_v_state1(X1,c_n1)|~node2(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_16)).
cnf(c_0_22,axiom,(m_main_v_state2(X1,c_n2)|~node2(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_17)).
cnf(c_0_23,axiom,(m_main_v_state1(X1,X2)|~m_main_v_state1(X3,X2)|~node6(X3,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_25)).
cnf(c_0_24,axiom,(m_main_v_turn(X1,X2)|~m_main_v_turn(X3,X2)|~node17(X3,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_60)).
cnf(c_0_25,axiom,(m_main_v_state1(X1,c_n1)|~node11(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_42)).
cnf(c_0_26,axiom,(~m_main_v_state1(X1,c_t1)|~m_main_v_state1(X1,c_c1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_4)).
cnf(c_0_27,axiom,(m_main_v_state1(X1,c_n1)|~node9(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_38)).
cnf(c_0_28,axiom,(m_main_v_state2(X1,c_n2)|~node8(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_35)).
cnf(c_0_29,axiom,(node14(X1,X2)|~trans(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_55)).
cnf(c_0_30,axiom,(~m_main_v_turn(X1,c_1)|~m_main_v_turn(X1,c_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_10)).
cnf(c_0_31,axiom,(m_main_v_turn(s0,c_1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_13)).
cnf(c_0_32,axiom,(m_main_v_state1(X1,c_c1)|~node10(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_40)).
cnf(c_0_33,axiom,(node8(X1)|node9(X1)|node10(X1)|node11(X1)|node12(X1)|m_main_v_state2(X1,c_c2)|node13(X1,X2)|~node14(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_54)).
cnf(c_0_34,axiom,(succ(s1,s2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound2)).
cnf(c_0_35,axiom,(m_main_v_turn(X1,c_2)|~node12(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_45)).
cnf(c_0_36,axiom,(~m_main_v_state2(X1,c_t2)|~m_main_v_state2(X1,c_c2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_8)).
cnf(c_0_37,axiom,(m_main_v_state2(X1,c_n2)|~node16(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_58)).
cnf(c_0_38,axiom,(m_main_v_state1(X1,c_n1)|~node15(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_56)).
cnf(c_0_39,axiom,(m_main_v_state2(X1,c_n2)|~node4(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_21)).
cnf(c_0_40,axiom,(m_main_v_state1(X1,c_n1)|~node3(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_18)).
cnf(c_0_41,axiom,(m_main_v_state1(X1,c_n1)|~node1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_14)).
cnf(c_0_42,axiom,(succ(s2,s3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound3)).
cnf(c_0_43,axiom,(m_main_v_state2(X1,X2)|~m_main_v_state2(X3,X2)|~node13(X3,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_46)).
cnf(c_0_44,axiom,(m_main_v_turn(X1,c_1)|m_main_v_turn(X1,c_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_9)).
cnf(c_0_45,axiom,(m_main_v_turn(X2,c_2)|~m_main_v_state1(X1,c_n1)|~m_main_v_state2(X1,c_t2)|~node18(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_62)).
cnf(c_0_46,axiom,(node1(X1)|node2(X1)|node3(X1)|node4(X1)|m_main_v_state1(X2,c_c1)|~m_main_v_state1(X1,c_t1)|~m_main_v_state2(X1,c_t2)|~m_main_v_turn(X1,c_1)|~node7(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_31)).
cnf(c_0_47,axiom,(m_main_v_state2(X3,X2)|~m_main_v_state2(X1,X2)|~node13(X3,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_47)).
cnf(c_0_48,axiom,(m_main_v_state2(X1,c_t2)|~node11(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_41)).
cnf(c_0_49,axiom,(m_main_v_state1(X3,X2)|~m_main_v_state1(X1,X2)|~node6(X3,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_26)).
cnf(c_0_50,axiom,(m_main_v_state2(X1,c_n2)|~node10(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_39)).
cnf(c_0_51,axiom,(m_main_v_state2(X1,c_n2)|~node9(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_37)).
cnf(c_0_52,axiom,(node1(X1)|node2(X1)|node3(X1)|node4(X1)|node5(X1)|m_main_v_state1(X2,c_n1)|~m_main_v_state1(X1,c_c1)|~node7(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_32)).
cnf(c_0_53,axiom,(node8(X1)|node9(X1)|m_main_v_state2(X2,c_t2)|~m_main_v_state2(X1,c_n2)|~m_main_v_state1(X1,c_c1)|~node14(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_50)).
cnf(c_0_54,axiom,(succ(s3,s4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound4)).
cnf(c_0_55,axiom,(m_main_v_state2(X1,c_n2)|m_main_v_state2(X1,c_t2)|m_main_v_state2(X1,c_c2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_5)).
cnf(c_0_56,axiom,(m_main_v_state1(X2,c_t1)|~m_main_v_state1(X1,c_n1)|~m_main_v_state2(X1,c_t2)|~node7(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_27)).
cnf(c_0_57,axiom,(node8(X1)|node9(X1)|node10(X1)|m_main_v_state2(X2,c_c2)|~m_main_v_state2(X1,c_t2)|~m_main_v_state1(X1,c_n1)|~node14(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_51)).
cnf(c_0_58,axiom,(m_main_v_state2(X1,c_t2)|~node12(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_43)).
cnf(c_0_59,axiom,(node8(X1)|node9(X1)|node10(X1)|node11(X1)|node12(X1)|m_main_v_state2(X2,c_n2)|~m_main_v_state2(X1,c_c2)|~node14(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_53)).
cnf(c_0_60,axiom,(m_main_v_turn(X1,c_1)|~node5(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_24)).
cnf(c_0_61,axiom,(succ(s4,s5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound5)).
cnf(c_0_62,axiom,(succ(s5,s6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound6)).
cnf(c_0_63,axiom,(node15(X1)|m_main_v_turn(X2,c_1)|~m_main_v_state2(X1,c_n2)|~m_main_v_state1(X1,c_t1)|~node18(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_63)).
cnf(c_0_64,axiom,(m_main_v_state2(X2,c_t2)|~m_main_v_state2(X1,c_n2)|~m_main_v_state1(X1,c_t1)|~node14(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_48)).
cnf(c_0_65,axiom,(node1(X1)|node2(X1)|node3(X1)|m_main_v_state1(X2,c_c1)|~m_main_v_state1(X1,c_t1)|~m_main_v_state2(X1,c_n2)|~node7(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_30)).
cnf(c_0_66,axiom,(m_main_v_turn(X3,X2)|~m_main_v_turn(X1,X2)|~node17(X3,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_61)).
cnf(c_0_67,axiom,(succ(s6,s7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound7)).
cnf(c_0_68,axiom,(m_main_v_state1(X1,c_t1)|~node12(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_44)).
cnf(c_0_69,axiom,(succ(s7,s8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound8)).
cnf(c_0_70,axiom,(succ(s8,s9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound9)).
cnf(c_0_71,axiom,(succ(s9,s10)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound10)).
cnf(c_0_72,axiom,(succ(s10,s11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound11)).
cnf(c_0_73,axiom,(succ(s11,s12)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound12)).
cnf(c_0_74,axiom,(succ(s12,s13)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound13)).
cnf(c_0_75,axiom,(succ(s13,s14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound14)).
cnf(c_0_76,axiom,(succ(s14,s15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound15)).
cnf(c_0_77,axiom,(succ(s15,s16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound16)).
cnf(c_0_78,axiom,(succ(s16,s17)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound17)).
cnf(c_0_79,axiom,(succ(s17,s18)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound18)).
cnf(c_0_80,axiom,(succ(s18,s19)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound19)).
cnf(c_0_81,axiom,(succ(s19,s20)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound20)).
cnf(c_0_82,axiom,(succ(s20,s21)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound21)).
cnf(c_0_83,axiom,(succ(s21,s22)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound22)).
cnf(c_0_84,axiom,(succ(s22,s23)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound23)).
cnf(c_0_85,axiom,(succ(s23,s24)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound24)).
cnf(c_0_86,axiom,(succ(s24,s25)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound25)).
cnf(c_0_87,axiom,(succ(s25,s26)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound26)).
cnf(c_0_88,axiom,(succ(s26,s27)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound27)).
cnf(c_0_89,axiom,(succ(s27,s28)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound28)).
cnf(c_0_90,axiom,(succ(s28,s29)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound29)).
cnf(c_0_91,axiom,(succ(s29,s30)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound30)).
cnf(c_0_92,axiom,(succ(s30,s31)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound31)).
cnf(c_0_93,axiom,(succ(s31,s32)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound32)).
cnf(c_0_94,axiom,(succ(s32,s33)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound33)).
cnf(c_0_95,axiom,(succ(s33,s34)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound34)).
cnf(c_0_96,axiom,(succ(s34,s35)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound35)).
cnf(c_0_97,axiom,(succ(s35,s36)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound36)).
cnf(c_0_98,axiom,(succ(s36,s37)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound37)).
cnf(c_0_99,axiom,(succ(s37,s38)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound38)).
cnf(c_0_100,axiom,(succ(s38,s39)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound39)).
cnf(c_0_101,axiom,(succ(s39,s40)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound40)).
cnf(c_0_102,axiom,(succ(s40,s41)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound41)).
cnf(c_0_103,axiom,(succ(s41,s42)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound42)).
cnf(c_0_104,axiom,(succ(s42,s43)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound43)).
cnf(c_0_105,axiom,(succ(s43,s44)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound44)).
cnf(c_0_106,axiom,(succ(s44,s45)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound45)).
cnf(c_0_107,axiom,(succ(s45,s46)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound46)).
cnf(c_0_108,axiom,(succ(s46,s47)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound47)).
cnf(c_0_109,axiom,(succ(s47,s48)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound48)).
cnf(c_0_110,axiom,(succ(s48,s49)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound49)).
cnf(c_0_111,negated_conjecture,(node26(X1)|xuntil28(X1)|~until27(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty19)).
cnf(c_0_112,negated_conjecture,(until27(s0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty26)).
cnf(c_0_113,negated_conjecture,(m_main_v_state2(X1,c_t2)|~node26(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty17)).
cnf(c_0_114,axiom,(succ(s49,s50)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound50)).
cnf(c_0_115,negated_conjecture,(until27(X1)|~succ(X2,X1)|~xuntil28(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty20)).
cnf(c_0_116,axiom,(succ(s50,s51)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound51)).
cnf(c_0_117,axiom,(succ(s51,s52)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound52)).
cnf(c_0_118,axiom,(succ(s52,s53)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound53)).
cnf(c_0_119,axiom,(succ(s53,s54)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound54)).
cnf(c_0_120,axiom,(succ(s54,s55)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound55)).
cnf(c_0_121,axiom,(succ(s55,s56)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound56)).
cnf(c_0_122,axiom,(succ(s56,s57)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound57)).
cnf(c_0_123,axiom,(succ(s57,s58)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound58)).
cnf(c_0_124,axiom,(succ(s58,s59)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound59)).
cnf(c_0_125,axiom,(succ(s59,s60)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound60)).
cnf(c_0_126,axiom,(succ(s60,s61)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound61)).
cnf(c_0_127,axiom,(succ(s61,s62)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound62)).
cnf(c_0_128,axiom,(succ(s62,s63)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound63)).
cnf(c_0_129,axiom,(succ(s63,s64)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound64)).
cnf(c_0_130,negated_conjecture,(loop|~last(X1)|~xuntil28(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty21)).
cnf(c_0_131,axiom,(last(s64)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound65)).
cnf(c_0_132,negated_conjecture,(always25(X1)|~node26(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty18)).
cnf(c_0_133,negated_conjecture,(~m_main_v_state2(X1,c_c2)|~always25(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty14)).
cnf(c_0_134,negated_conjecture,(always25(X1)|~trans(X2,X1)|~always25(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty15)).
cnf(c_0_135,axiom,(m_main_v_state2(X1,c_t2)|~node5(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_23)).
cnf(c_0_136,axiom,(node8(X1)|node9(X1)|node10(X1)|node11(X1)|m_main_v_state2(X2,c_c2)|~m_main_v_state2(X1,c_t2)|~m_main_v_state1(X1,c_t1)|~m_main_v_turn(X1,c_2)|~node14(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', m_main_52)).
cnf(c_0_137,axiom,(trans(s64,s0)|trans(s64,s1)|trans(s64,s2)|trans(s64,s3)|trans(s64,s4)|trans(s64,s5)|trans(s64,s6)|trans(s64,s7)|trans(s64,s8)|trans(s64,s9)|trans(s64,s10)|trans(s64,s11)|trans(s64,s12)|trans(s64,s13)|trans(s64,s14)|trans(s64,s15)|trans(s64,s16)|trans(s64,s17)|trans(s64,s18)|trans(s64,s19)|trans(s64,s20)|trans(s64,s21)|trans(s64,s22)|trans(s64,s23)|trans(s64,s24)|trans(s64,s25)|trans(s64,s26)|trans(s64,s27)|trans(s64,s28)|trans(s64,s29)|trans(s64,s30)|trans(s64,s31)|trans(s64,s32)|trans(s64,s33)|trans(s64,s34)|trans(s64,s35)|trans(s64,s36)|trans(s64,s37)|trans(s64,s38)|trans(s64,s39)|trans(s64,s40)|trans(s64,s41)|trans(s64,s42)|trans(s64,s43)|trans(s64,s44)|trans(s64,s45)|trans(s64,s46)|trans(s64,s47)|trans(s64,s48)|trans(s64,s49)|trans(s64,s50)|trans(s64,s51)|trans(s64,s52)|trans(s64,s53)|trans(s64,s54)|trans(s64,s55)|trans(s64,s56)|trans(s64,s57)|trans(s64,s58)|trans(s64,s59)|trans(s64,s60)|trans(s64,s61)|trans(s64,s62)|trans(s64,s63)|trans(s64,s64)|~loop), file('/export/starexec/sandbox/benchmark/theBenchmark.p', bound67)).
cnf(c_0_138,negated_conjecture,(until2p29(X1)|~trans(X2,X1)|~last(X2)|~xuntil28(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty22)).
cnf(c_0_139,negated_conjecture,(node26(X1)|xuntil2p30(X1)|~until2p29(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty23)).
cnf(c_0_140,negated_conjecture,(until2p29(X1)|~succ(X2,X1)|~xuntil2p30(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty24)).
cnf(c_0_141,negated_conjecture,(~last(X1)|~xuntil2p30(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', prpty25)).
cnf(c_0_142,plain,(~m_main_v_state2(s0,c_t2)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_143,plain,(~m_main_v_state1(s0,c_t1)), inference(spm,[status(thm)],[c_0_2, c_0_3])).
cnf(c_0_144,plain,(~m_main_v_state2(s0,c_c2)), inference(spm,[status(thm)],[c_0_4, c_0_1])).
cnf(c_0_145,plain,(~node1(s0)), inference(spm,[status(thm)],[c_0_142, c_0_5])).
cnf(c_0_146,plain,(node6(X1,X2)|node5(X1)|node4(X1)|node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X1,c_c1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_147,plain,(trans(s0,s1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_148,plain,(~node5(s0)), inference(spm,[status(thm)],[c_0_143, c_0_10])).
cnf(c_0_149,plain,(~node4(s0)), inference(spm,[status(thm)],[c_0_143, c_0_11])).
cnf(c_0_150,plain,(~node3(s0)), inference(spm,[status(thm)],[c_0_144, c_0_12])).
cnf(c_0_151,plain,(~m_main_v_state1(s0,c_c1)), inference(spm,[status(thm)],[c_0_13, c_0_3])).
cnf(c_0_152,plain,(node17(X1,X2)|node16(X1)|node15(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_153,plain,(~node16(s0)), inference(spm,[status(thm)],[c_0_143, c_0_16])).
cnf(c_0_154,plain,(~node15(s0)), inference(spm,[status(thm)],[c_0_142, c_0_17])).
cnf(c_0_155,plain,(m_main_v_state1(X1,c_t1)|~node7(s0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_1]), c_0_3])]), c_0_145])).
cnf(c_0_156,plain,(~node8(s0)), inference(spm,[status(thm)],[c_0_143, c_0_19])).
cnf(c_0_157,plain,(node8(X1)|m_main_v_state2(X2,c_t2)|~node14(X1,X2)|~node2(X1)|~m_main_v_state2(X1,c_n2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_158,plain,(node6(s0,s1)|node2(s0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_148]), c_0_149]), c_0_150]), c_0_145]), c_0_151])).
cnf(c_0_159,plain,(node17(s0,s1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_147]), c_0_153]), c_0_154])).
cnf(c_0_160,plain,(m_main_v_state1(X1,c_t1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_155, c_0_7])).
cnf(c_0_161,plain,(m_main_v_state2(X1,c_t2)|~node14(s0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_3]), c_0_1])]), c_0_156])).
cnf(c_0_162,plain,(node8(X1)|m_main_v_state2(X2,c_t2)|~node14(X1,X2)|~node2(X1)), inference(spm,[status(thm)],[c_0_157, c_0_22])).
cnf(c_0_163,plain,(node2(s0)|m_main_v_state1(s1,X1)|~m_main_v_state1(s0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_158])).
cnf(c_0_164,plain,(m_main_v_turn(s1,X1)|~m_main_v_turn(s0,X1)), inference(spm,[status(thm)],[c_0_24, c_0_159])).
cnf(c_0_165,plain,(~node11(X1)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_2, c_0_25])).
cnf(c_0_166,plain,(~m_main_v_state1(X1,c_c1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_26, c_0_160])).
cnf(c_0_167,plain,(~node9(X1)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_2, c_0_27])).
cnf(c_0_168,plain,(~node8(X1)|~m_main_v_state2(X1,c_t2)), inference(spm,[status(thm)],[c_0_0, c_0_28])).
cnf(c_0_169,plain,(m_main_v_state2(X1,c_t2)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_161, c_0_29])).
cnf(c_0_170,plain,(node8(X1)|m_main_v_state2(X2,c_t2)|~node2(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_162, c_0_29])).
cnf(c_0_171,plain,(node2(s0)|m_main_v_state1(s1,c_n1)), inference(spm,[status(thm)],[c_0_163, c_0_3])).
cnf(c_0_172,plain,(~m_main_v_turn(s1,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_164]), c_0_31])])).
cnf(c_0_173,plain,(~node11(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_165, c_0_160])).
cnf(c_0_174,plain,(~node10(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_166, c_0_32])).
cnf(c_0_175,plain,(~node9(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_167, c_0_160])).
cnf(c_0_176,plain,(~node8(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_168, c_0_169])).
cnf(c_0_177,plain,(m_main_v_state2(s1,c_t2)|~node2(s0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_147]), c_0_156])).
cnf(c_0_178,plain,(node2(s0)|~m_main_v_state1(s1,c_t1)), inference(spm,[status(thm)],[c_0_2, c_0_171])).
cnf(c_0_179,plain,(node13(X1,X2)|node12(X1)|node11(X1)|node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X1,c_c2)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_33, c_0_29])).
cnf(c_0_180,plain,(trans(s1,s2)), inference(spm,[status(thm)],[c_0_8, c_0_34])).
cnf(c_0_181,plain,(~node12(s1)), inference(spm,[status(thm)],[c_0_172, c_0_35])).
cnf(c_0_182,plain,(~node11(s1)), inference(spm,[status(thm)],[c_0_173, c_0_147])).
cnf(c_0_183,plain,(~node10(s1)), inference(spm,[status(thm)],[c_0_174, c_0_147])).
cnf(c_0_184,plain,(~node9(s1)), inference(spm,[status(thm)],[c_0_175, c_0_147])).
cnf(c_0_185,plain,(~node8(s1)), inference(spm,[status(thm)],[c_0_176, c_0_147])).
cnf(c_0_186,plain,(~node2(s0)|~m_main_v_state2(s1,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_177])).
cnf(c_0_187,plain,(node2(s0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_160]), c_0_147])])).
cnf(c_0_188,plain,(~node16(X1)|~m_main_v_state2(X1,c_t2)), inference(spm,[status(thm)],[c_0_0, c_0_37])).
cnf(c_0_189,plain,(~node15(X1)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_2, c_0_38])).
cnf(c_0_190,plain,(node1(X1)|m_main_v_state1(X2,c_t1)|~node7(X1,X2)|~node2(X1)|~m_main_v_state1(X1,c_n1)), inference(spm,[status(thm)],[c_0_18, c_0_22])).
cnf(c_0_191,plain,(node13(s1,s2)|m_main_v_state2(s1,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_181]), c_0_182]), c_0_183]), c_0_184]), c_0_185])).
cnf(c_0_192,plain,(~m_main_v_state2(s1,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_186, c_0_187])])).
cnf(c_0_193,plain,(~node16(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_188, c_0_169])).
cnf(c_0_194,plain,(~node15(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_189, c_0_160])).
cnf(c_0_195,plain,(node1(X1)|m_main_v_state1(X2,c_t1)|~node7(X1,X2)|~node2(X1)), inference(spm,[status(thm)],[c_0_190, c_0_21])).
cnf(c_0_196,plain,(node13(s1,s2)), inference(sr,[status(thm)],[c_0_191, c_0_192])).
cnf(c_0_197,plain,(~node16(s1)), inference(spm,[status(thm)],[c_0_193, c_0_147])).
cnf(c_0_198,plain,(~node15(s1)), inference(spm,[status(thm)],[c_0_194, c_0_147])).
cnf(c_0_199,plain,(node1(X1)|m_main_v_state1(X2,c_t1)|~node2(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_195, c_0_7])).
cnf(c_0_200,plain,(~node4(X1)|~m_main_v_state2(X1,c_t2)), inference(spm,[status(thm)],[c_0_0, c_0_39])).
cnf(c_0_201,plain,(~node3(X1)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_2, c_0_40])).
cnf(c_0_202,plain,(~node2(X1)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_2, c_0_21])).
cnf(c_0_203,plain,(~node1(X1)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_2, c_0_41])).
cnf(c_0_204,plain,(trans(s2,s3)), inference(spm,[status(thm)],[c_0_8, c_0_42])).
cnf(c_0_205,plain,(m_main_v_state2(s2,X1)|~m_main_v_state2(s1,X1)), inference(spm,[status(thm)],[c_0_43, c_0_196])).
cnf(c_0_206,plain,(m_main_v_state2(s1,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_177, c_0_187])])).
cnf(c_0_207,plain,(node17(s1,s2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_180]), c_0_197]), c_0_198])).
cnf(c_0_208,plain,(m_main_v_state1(s1,c_t1)|~node2(s0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_147]), c_0_145])).
cnf(c_0_209,plain,(~node4(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_200, c_0_169])).
cnf(c_0_210,plain,(~node3(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_201, c_0_160])).
cnf(c_0_211,plain,(~node2(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_202, c_0_160])).
cnf(c_0_212,plain,(~node1(X1)|~trans(s0,X1)), inference(spm,[status(thm)],[c_0_203, c_0_160])).
cnf(c_0_213,plain,(node17(s2,s3)|node16(s2)|node15(s2)), inference(spm,[status(thm)],[c_0_152, c_0_204])).
cnf(c_0_214,plain,(~node16(s2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_205]), c_0_206])])).
cnf(c_0_215,plain,(m_main_v_turn(s2,X1)|~m_main_v_turn(s1,X1)), inference(spm,[status(thm)],[c_0_24, c_0_207])).
cnf(c_0_216,plain,(m_main_v_turn(s1,c_1)), inference(spm,[status(thm)],[c_0_172, c_0_44])).
cnf(c_0_217,plain,(m_main_v_state1(s1,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_208, c_0_187])])).
cnf(c_0_218,plain,(~node4(s1)), inference(spm,[status(thm)],[c_0_209, c_0_147])).
cnf(c_0_219,plain,(~node3(s1)), inference(spm,[status(thm)],[c_0_210, c_0_147])).
cnf(c_0_220,plain,(~node2(s1)), inference(spm,[status(thm)],[c_0_211, c_0_147])).
cnf(c_0_221,plain,(~node1(s1)), inference(spm,[status(thm)],[c_0_212, c_0_147])).
cnf(c_0_222,plain,(node17(s2,s3)|node15(s2)), inference(sr,[status(thm)],[c_0_213, c_0_214])).
cnf(c_0_223,plain,(~m_main_v_turn(s2,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_215]), c_0_216])])).
cnf(c_0_224,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node1(X2)|~m_main_v_state1(X2,c_n1)), inference(spm,[status(thm)],[c_0_45, c_0_5])).
cnf(c_0_225,plain,(m_main_v_state1(X1,c_c1)|~node7(s1,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_164]), c_0_31])]), c_0_206]), c_0_217])]), c_0_218]), c_0_219]), c_0_220]), c_0_221])).
cnf(c_0_226,plain,(node15(s2)|m_main_v_turn(s3,X1)|~m_main_v_turn(s2,X1)), inference(spm,[status(thm)],[c_0_24, c_0_222])).
cnf(c_0_227,plain,(m_main_v_turn(s2,c_1)), inference(spm,[status(thm)],[c_0_223, c_0_44])).
cnf(c_0_228,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node1(X2)), inference(spm,[status(thm)],[c_0_224, c_0_41])).
cnf(c_0_229,plain,(~node15(X1)|~m_main_v_state1(X1,c_c1)), inference(spm,[status(thm)],[c_0_13, c_0_38])).
cnf(c_0_230,plain,(m_main_v_state1(X1,c_c1)|~trans(s1,X1)), inference(spm,[status(thm)],[c_0_225, c_0_7])).
cnf(c_0_231,plain,(node15(s2)|m_main_v_turn(s3,c_1)), inference(spm,[status(thm)],[c_0_226, c_0_227])).
cnf(c_0_232,plain,(m_main_v_turn(X1,c_2)|~node1(X2)|~trans(X2,X1)), inference(spm,[status(thm)],[c_0_228, c_0_15])).
cnf(c_0_233,plain,(~node5(X1)|~m_main_v_state1(X1,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_10])).
cnf(c_0_234,plain,(~node15(X1)|~trans(s1,X1)), inference(spm,[status(thm)],[c_0_229, c_0_230])).
cnf(c_0_235,plain,(node15(s2)|~m_main_v_turn(s3,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_231])).
cnf(c_0_236,plain,(m_main_v_turn(s3,c_2)|~node1(s2)), inference(spm,[status(thm)],[c_0_232, c_0_204])).
cnf(c_0_237,plain,(~node5(X1)|~trans(s1,X1)), inference(spm,[status(thm)],[c_0_233, c_0_230])).
cnf(c_0_238,plain,(m_main_v_state2(s1,X1)|~m_main_v_state2(s2,X1)), inference(spm,[status(thm)],[c_0_47, c_0_196])).
cnf(c_0_239,plain,(~node2(X1)|~m_main_v_state2(X1,c_t2)), inference(spm,[status(thm)],[c_0_0, c_0_22])).
cnf(c_0_240,plain,(~node15(s2)), inference(spm,[status(thm)],[c_0_234, c_0_180])).
cnf(c_0_241,plain,(node15(s2)|~node1(s2)), inference(spm,[status(thm)],[c_0_235, c_0_236])).
cnf(c_0_242,plain,(~node5(s2)), inference(spm,[status(thm)],[c_0_237, c_0_180])).
cnf(c_0_243,plain,(~node4(s2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_205]), c_0_206])])).
cnf(c_0_244,plain,(~node3(s2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_12]), c_0_192])).
cnf(c_0_245,plain,(~node2(s2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_205]), c_0_206])])).
cnf(c_0_246,plain,(~node1(s2)), inference(spm,[status(thm)],[c_0_240, c_0_241])).
cnf(c_0_247,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node11(X2)|~m_main_v_state1(X2,c_n1)), inference(spm,[status(thm)],[c_0_45, c_0_48])).
cnf(c_0_248,plain,(node6(s2,s3)|m_main_v_state1(s2,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_204]), c_0_242]), c_0_243]), c_0_244]), c_0_245]), c_0_246])).
cnf(c_0_249,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node11(X2)), inference(spm,[status(thm)],[c_0_247, c_0_25])).
cnf(c_0_250,plain,(m_main_v_state1(s2,c_c1)|m_main_v_state1(s3,X1)|~m_main_v_state1(s2,X1)), inference(spm,[status(thm)],[c_0_23, c_0_248])).
cnf(c_0_251,plain,(m_main_v_turn(X1,c_2)|~node11(X2)|~trans(X2,X1)), inference(spm,[status(thm)],[c_0_249, c_0_15])).
cnf(c_0_252,plain,(m_main_v_state1(s2,c_c1)|m_main_v_state1(s2,X1)|~m_main_v_state1(s3,X1)), inference(spm,[status(thm)],[c_0_49, c_0_248])).
cnf(c_0_253,plain,(m_main_v_state1(s2,c_c1)|m_main_v_state1(s3,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_230]), c_0_180])])).
cnf(c_0_254,plain,(m_main_v_turn(s3,c_2)|~node11(s2)), inference(spm,[status(thm)],[c_0_251, c_0_204])).
cnf(c_0_255,plain,(m_main_v_state1(s2,c_c1)), inference(spm,[status(thm)],[c_0_252, c_0_253])).
cnf(c_0_256,plain,(node15(s2)|~node11(s2)), inference(spm,[status(thm)],[c_0_235, c_0_254])).
cnf(c_0_257,plain,(~node10(X1)|~m_main_v_state2(X1,c_t2)), inference(spm,[status(thm)],[c_0_0, c_0_50])).
cnf(c_0_258,plain,(~node9(X1)|~m_main_v_state2(X1,c_t2)), inference(spm,[status(thm)],[c_0_0, c_0_51])).
cnf(c_0_259,plain,(m_main_v_state1(X1,c_n1)|~node7(s2,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_255]), c_0_242]), c_0_243]), c_0_244]), c_0_245]), c_0_246])).
cnf(c_0_260,plain,(~node12(s2)), inference(spm,[status(thm)],[c_0_223, c_0_35])).
cnf(c_0_261,plain,(~node11(s2)), inference(spm,[status(thm)],[c_0_240, c_0_256])).
cnf(c_0_262,plain,(~node10(s2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_205]), c_0_206])])).
cnf(c_0_263,plain,(~node9(s2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_205]), c_0_206])])).
cnf(c_0_264,plain,(~node8(s2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_205]), c_0_206])])).
cnf(c_0_265,plain,(~m_main_v_state2(s2,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_205]), c_0_206])])).
cnf(c_0_266,plain,(m_main_v_state1(X1,c_n1)|~trans(s2,X1)), inference(spm,[status(thm)],[c_0_259, c_0_7])).
cnf(c_0_267,plain,(node13(s2,s3)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_204]), c_0_260]), c_0_261]), c_0_262]), c_0_263]), c_0_264]), c_0_265])).
cnf(c_0_268,plain,(~m_main_v_state1(X1,c_c1)|~trans(s2,X1)), inference(spm,[status(thm)],[c_0_13, c_0_266])).
cnf(c_0_269,plain,(~m_main_v_state1(X1,c_t1)|~trans(s2,X1)), inference(spm,[status(thm)],[c_0_2, c_0_266])).
cnf(c_0_270,plain,(m_main_v_state2(s3,X1)|~m_main_v_state2(s2,X1)), inference(spm,[status(thm)],[c_0_43, c_0_267])).
cnf(c_0_271,plain,(m_main_v_state2(X1,c_t2)|~node14(s2,X1)|~m_main_v_state2(s2,c_n2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_255]), c_0_263]), c_0_264])).
cnf(c_0_272,plain,(node15(s2)|~node12(s3)), inference(spm,[status(thm)],[c_0_235, c_0_35])).
cnf(c_0_273,plain,(~node10(X1)|~trans(s2,X1)), inference(spm,[status(thm)],[c_0_268, c_0_32])).
cnf(c_0_274,plain,(~node8(X1)|~trans(s2,X1)), inference(spm,[status(thm)],[c_0_269, c_0_19])).
cnf(c_0_275,plain,(m_main_v_state2(s2,X1)|~m_main_v_state2(s3,X1)), inference(spm,[status(thm)],[c_0_47, c_0_267])).
cnf(c_0_276,plain,(~m_main_v_state2(s2,c_t2)|~m_main_v_state2(s1,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_205])).
cnf(c_0_277,plain,(~m_main_v_state2(s3,c_t2)|~m_main_v_state2(s2,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_270])).
cnf(c_0_278,plain,(m_main_v_state2(X1,c_t2)|~m_main_v_state2(s2,c_n2)|~trans(s2,X1)), inference(spm,[status(thm)],[c_0_271, c_0_29])).
cnf(c_0_279,plain,(trans(s3,s4)), inference(spm,[status(thm)],[c_0_8, c_0_54])).
cnf(c_0_280,plain,(~node12(s3)), inference(spm,[status(thm)],[c_0_240, c_0_272])).
cnf(c_0_281,plain,(~node10(s3)), inference(spm,[status(thm)],[c_0_273, c_0_204])).
cnf(c_0_282,plain,(~node8(s3)), inference(spm,[status(thm)],[c_0_274, c_0_204])).
cnf(c_0_283,plain,(m_main_v_state2(s2,c_n2)|~node9(s3)), inference(spm,[status(thm)],[c_0_275, c_0_51])).
cnf(c_0_284,plain,(~m_main_v_state2(s1,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_205]), c_0_206])])).
cnf(c_0_285,plain,(m_main_v_state2(s2,c_n2)|m_main_v_state2(s2,c_t2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_55]), c_0_192])).
cnf(c_0_286,plain,(~m_main_v_state2(s2,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_278]), c_0_204])])).
cnf(c_0_287,plain,(node13(s3,s4)|node11(s3)|node9(s3)|m_main_v_state2(s3,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_279]), c_0_280]), c_0_281]), c_0_282])).
cnf(c_0_288,plain,(~node9(s3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_283]), c_0_284])).
cnf(c_0_289,plain,(~m_main_v_state2(s3,c_c2)|~m_main_v_state2(s2,c_t2)), inference(spm,[status(thm)],[c_0_36, c_0_270])).
cnf(c_0_290,plain,(m_main_v_state2(s2,c_t2)), inference(sr,[status(thm)],[c_0_285, c_0_286])).
cnf(c_0_291,plain,(node13(s3,s4)|node11(s3)|m_main_v_state2(s3,c_c2)), inference(sr,[status(thm)],[c_0_287, c_0_288])).
cnf(c_0_292,plain,(~m_main_v_state2(s3,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_289, c_0_290])])).
cnf(c_0_293,plain,(node13(s3,s4)|node11(s3)), inference(sr,[status(thm)],[c_0_291, c_0_292])).
cnf(c_0_294,plain,(node11(s3)|m_main_v_state2(s4,X1)|~m_main_v_state2(s3,X1)), inference(spm,[status(thm)],[c_0_43, c_0_293])).
cnf(c_0_295,plain,(~node4(X1)|~trans(s2,X1)), inference(spm,[status(thm)],[c_0_269, c_0_11])).
cnf(c_0_296,plain,(~node5(X1)|~trans(s2,X1)), inference(spm,[status(thm)],[c_0_269, c_0_10])).
cnf(c_0_297,plain,(node11(s3)|m_main_v_state2(s4,X1)|~m_main_v_state2(s2,X1)), inference(spm,[status(thm)],[c_0_294, c_0_270])).
cnf(c_0_298,plain,(node6(s3,s4)|node5(s3)|node4(s3)|node3(s3)|node2(s3)|node1(s3)|m_main_v_state1(s3,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_279])).
cnf(c_0_299,plain,(~node4(s3)), inference(spm,[status(thm)],[c_0_295, c_0_204])).
cnf(c_0_300,plain,(~node5(s3)), inference(spm,[status(thm)],[c_0_296, c_0_204])).
cnf(c_0_301,plain,(m_main_v_state2(s2,c_n2)|~node2(s3)), inference(spm,[status(thm)],[c_0_275, c_0_22])).
cnf(c_0_302,plain,(node11(s3)|m_main_v_state2(s3,X1)|~m_main_v_state2(s4,X1)), inference(spm,[status(thm)],[c_0_47, c_0_293])).
cnf(c_0_303,plain,(node11(s3)|m_main_v_state2(s4,c_t2)), inference(spm,[status(thm)],[c_0_297, c_0_290])).
cnf(c_0_304,plain,(node6(s3,s4)|node3(s3)|node2(s3)|node1(s3)|m_main_v_state1(s3,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_298, c_0_299]), c_0_300])).
cnf(c_0_305,plain,(~node3(s3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_12]), c_0_265])).
cnf(c_0_306,plain,(~node2(s3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_301]), c_0_284])).
cnf(c_0_307,plain,(node11(s3)|m_main_v_state2(s3,c_t2)), inference(spm,[status(thm)],[c_0_302, c_0_303])).
cnf(c_0_308,plain,(node6(s3,s4)|node1(s3)|m_main_v_state1(s3,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_304, c_0_305]), c_0_306])).
cnf(c_0_309,plain,(node11(s3)|m_main_v_state1(X1,c_t1)|~node7(s3,X1)|~m_main_v_state1(s3,c_n1)), inference(spm,[status(thm)],[c_0_56, c_0_307])).
cnf(c_0_310,plain,(m_main_v_state1(X1,c_t1)|~node11(X2)|~node7(X2,X1)|~m_main_v_state1(X2,c_n1)), inference(spm,[status(thm)],[c_0_56, c_0_48])).
cnf(c_0_311,plain,(node1(s3)|m_main_v_state1(s3,c_c1)|m_main_v_state1(s4,X1)|~m_main_v_state1(s3,X1)), inference(spm,[status(thm)],[c_0_23, c_0_308])).
cnf(c_0_312,plain,(node11(s3)|m_main_v_state1(X1,c_t1)|~node7(s3,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_266]), c_0_204])])).
cnf(c_0_313,plain,(m_main_v_state1(X1,c_t1)|~node11(X2)|~node7(X2,X1)), inference(spm,[status(thm)],[c_0_310, c_0_25])).
cnf(c_0_314,plain,(node1(s3)|m_main_v_state1(s3,c_c1)|m_main_v_state1(s4,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_266]), c_0_204])])).
cnf(c_0_315,plain,(node11(s3)|m_main_v_state1(X1,c_t1)|~trans(s3,X1)), inference(spm,[status(thm)],[c_0_312, c_0_7])).
cnf(c_0_316,plain,(m_main_v_state1(X1,c_t1)|~node11(X2)|~trans(X2,X1)), inference(spm,[status(thm)],[c_0_313, c_0_7])).
cnf(c_0_317,plain,(node1(s3)|m_main_v_state1(s3,c_c1)|~m_main_v_state1(s4,c_t1)), inference(spm,[status(thm)],[c_0_2, c_0_314])).
cnf(c_0_318,plain,(node11(s3)|m_main_v_state1(s4,c_t1)), inference(spm,[status(thm)],[c_0_315, c_0_279])).
cnf(c_0_319,plain,(node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X2,c_c2)|~node14(X1,X2)|~node11(X1)|~m_main_v_state2(X1,c_t2)), inference(spm,[status(thm)],[c_0_57, c_0_25])).
cnf(c_0_320,plain,(node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X2,c_c2)|~node14(X1,X2)|~node1(X1)|~m_main_v_state2(X1,c_t2)), inference(spm,[status(thm)],[c_0_57, c_0_41])).
cnf(c_0_321,plain,(m_main_v_state1(s4,c_t1)|~node11(s3)), inference(spm,[status(thm)],[c_0_316, c_0_279])).
cnf(c_0_322,plain,(node11(s3)|node1(s3)|m_main_v_state1(s3,c_c1)), inference(spm,[status(thm)],[c_0_317, c_0_318])).
cnf(c_0_323,plain,(m_main_v_state1(X1,c_t1)|~node7(X2,X1)|~node1(X2)|~m_main_v_state1(X2,c_n1)), inference(spm,[status(thm)],[c_0_56, c_0_5])).
cnf(c_0_324,plain,(node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X2,c_c2)|~node14(X1,X2)|~node11(X1)), inference(spm,[status(thm)],[c_0_319, c_0_48])).
cnf(c_0_325,plain,(node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X2,c_c2)|~node14(X1,X2)|~node1(X1)), inference(spm,[status(thm)],[c_0_320, c_0_5])).
cnf(c_0_326,plain,(node1(s3)|m_main_v_state1(s3,c_c1)|~node11(s3)), inference(spm,[status(thm)],[c_0_317, c_0_321])).
cnf(c_0_327,plain,(node11(s3)|node1(s3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_322]), c_0_204])])).
cnf(c_0_328,plain,(m_main_v_state1(X1,c_t1)|~node7(X2,X1)|~node1(X2)), inference(spm,[status(thm)],[c_0_323, c_0_41])).
cnf(c_0_329,plain,(node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X2,c_c2)|~node11(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_324, c_0_29])).
cnf(c_0_330,plain,(node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X2,c_c2)|~node1(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_325, c_0_29])).
cnf(c_0_331,plain,(node1(s3)|m_main_v_state1(s3,c_c1)), inference(spm,[status(thm)],[c_0_326, c_0_327])).
cnf(c_0_332,plain,(m_main_v_state1(X1,c_t1)|~node1(X2)|~trans(X2,X1)), inference(spm,[status(thm)],[c_0_328, c_0_7])).
cnf(c_0_333,plain,(node9(s3)|m_main_v_state2(s4,c_c2)|~node11(s3)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_279]), c_0_281]), c_0_282])).
cnf(c_0_334,plain,(node11(s3)|~node10(s4)), inference(spm,[status(thm)],[c_0_257, c_0_303])).
cnf(c_0_335,plain,(node11(s3)|~node8(s4)), inference(spm,[status(thm)],[c_0_168, c_0_303])).
cnf(c_0_336,plain,(node9(s3)|m_main_v_state2(s4,c_c2)|~node1(s3)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_279]), c_0_281]), c_0_282])).
cnf(c_0_337,plain,(node1(s3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_331]), c_0_204])])).
cnf(c_0_338,plain,(m_main_v_state1(s4,c_t1)|~node1(s3)), inference(spm,[status(thm)],[c_0_332, c_0_279])).
cnf(c_0_339,plain,(~node10(X1)|~m_main_v_state2(X1,c_c2)), inference(spm,[status(thm)],[c_0_4, c_0_50])).
cnf(c_0_340,plain,(m_main_v_state2(s4,c_c2)|~node10(s4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_334]), c_0_288])).
cnf(c_0_341,plain,(~node11(s3)|~node9(s4)), inference(spm,[status(thm)],[c_0_167, c_0_321])).
cnf(c_0_342,plain,(node11(s3)|~node9(s4)), inference(spm,[status(thm)],[c_0_258, c_0_303])).
cnf(c_0_343,plain,(~node8(X1)|~m_main_v_state2(X1,c_c2)), inference(spm,[status(thm)],[c_0_4, c_0_28])).
cnf(c_0_344,plain,(m_main_v_state2(s4,c_c2)|~node8(s4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_335]), c_0_288])).
cnf(c_0_345,plain,(~node12(X1)|~m_main_v_state2(X1,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_58])).
cnf(c_0_346,plain,(m_main_v_state2(s4,c_c2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_336, c_0_337])]), c_0_288])).
cnf(c_0_347,plain,(~node11(s4)|~node1(s3)), inference(spm,[status(thm)],[c_0_165, c_0_338])).
cnf(c_0_348,plain,(~node10(s4)), inference(spm,[status(thm)],[c_0_339, c_0_340])).
cnf(c_0_349,plain,(~node9(s4)), inference(spm,[status(thm)],[c_0_341, c_0_342])).
cnf(c_0_350,plain,(~node8(s4)), inference(spm,[status(thm)],[c_0_343, c_0_344])).
cnf(c_0_351,plain,(~node12(s4)), inference(spm,[status(thm)],[c_0_345, c_0_346])).
cnf(c_0_352,plain,(~node11(s4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_347, c_0_337])])).
cnf(c_0_353,plain,(m_main_v_state2(X1,c_n2)|~node14(s4,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_346]), c_0_348]), c_0_349]), c_0_350]), c_0_351]), c_0_352])).
cnf(c_0_354,plain,(~node5(X1)|~m_main_v_turn(X1,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_60])).
cnf(c_0_355,plain,(m_main_v_turn(s4,c_2)|~node1(s3)), inference(spm,[status(thm)],[c_0_232, c_0_279])).
cnf(c_0_356,plain,(m_main_v_state2(X1,c_n2)|~trans(s4,X1)), inference(spm,[status(thm)],[c_0_353, c_0_29])).
cnf(c_0_357,plain,(trans(s4,s5)), inference(spm,[status(thm)],[c_0_8, c_0_61])).
cnf(c_0_358,plain,(~node11(s3)|~node3(s4)), inference(spm,[status(thm)],[c_0_201, c_0_321])).
cnf(c_0_359,plain,(node11(s3)|~node3(s4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_12]), c_0_292])).
cnf(c_0_360,plain,(~node5(s4)|~node1(s3)), inference(spm,[status(thm)],[c_0_354, c_0_355])).
cnf(c_0_361,plain,(~node4(X1)|~m_main_v_state2(X1,c_c2)), inference(spm,[status(thm)],[c_0_4, c_0_39])).
cnf(c_0_362,plain,(~node11(s3)|~node2(s4)), inference(spm,[status(thm)],[c_0_202, c_0_321])).
cnf(c_0_363,plain,(node11(s3)|~node2(s4)), inference(spm,[status(thm)],[c_0_239, c_0_303])).
cnf(c_0_364,plain,(~node11(s3)|~node1(s4)), inference(spm,[status(thm)],[c_0_203, c_0_321])).
cnf(c_0_365,plain,(node11(s3)|~node1(s4)), inference(spm,[status(thm)],[c_0_203, c_0_318])).
cnf(c_0_366,plain,(~node1(s3)|~m_main_v_state1(s4,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_338])).
cnf(c_0_367,plain,(~m_main_v_state2(X1,c_t2)|~trans(s4,X1)), inference(spm,[status(thm)],[c_0_0, c_0_356])).
cnf(c_0_368,plain,(node6(s4,s5)|node5(s4)|node4(s4)|node3(s4)|node2(s4)|node1(s4)|m_main_v_state1(s4,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_357])).
cnf(c_0_369,plain,(~node3(s4)), inference(spm,[status(thm)],[c_0_358, c_0_359])).
cnf(c_0_370,plain,(~node5(s4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_360, c_0_337])])).
cnf(c_0_371,plain,(~node4(s4)), inference(spm,[status(thm)],[c_0_361, c_0_346])).
cnf(c_0_372,plain,(~node2(s4)), inference(spm,[status(thm)],[c_0_362, c_0_363])).
cnf(c_0_373,plain,(~node1(s4)), inference(spm,[status(thm)],[c_0_364, c_0_365])).
cnf(c_0_374,plain,(~m_main_v_state1(s4,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_366, c_0_337])])).
cnf(c_0_375,plain,(trans(s5,s6)), inference(spm,[status(thm)],[c_0_8, c_0_62])).
cnf(c_0_376,plain,(~node15(X1)|~trans(s4,X1)), inference(spm,[status(thm)],[c_0_367, c_0_17])).
cnf(c_0_377,plain,(node6(s4,s5)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_368, c_0_369]), c_0_370]), c_0_371]), c_0_372]), c_0_373]), c_0_374])).
cnf(c_0_378,plain,(node17(s4,s5)|node16(s4)|node15(s4)), inference(spm,[status(thm)],[c_0_152, c_0_357])).
cnf(c_0_379,plain,(node17(s5,s6)|node16(s5)|node15(s5)), inference(spm,[status(thm)],[c_0_152, c_0_375])).
cnf(c_0_380,plain,(~node15(s5)), inference(spm,[status(thm)],[c_0_376, c_0_357])).
cnf(c_0_381,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s4,X1)), inference(spm,[status(thm)],[c_0_63, c_0_356])).
cnf(c_0_382,plain,(m_main_v_state1(s5,X1)|~m_main_v_state1(s4,X1)), inference(spm,[status(thm)],[c_0_23, c_0_377])).
cnf(c_0_383,plain,(m_main_v_state1(s4,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_338, c_0_337])])).
cnf(c_0_384,plain,(node16(s4)|node15(s4)|m_main_v_turn(s5,X1)|~m_main_v_turn(s4,X1)), inference(spm,[status(thm)],[c_0_24, c_0_378])).
cnf(c_0_385,plain,(~node15(s4)|~node1(s3)), inference(spm,[status(thm)],[c_0_189, c_0_338])).
cnf(c_0_386,plain,(m_main_v_state2(X1,c_t2)|~node16(X2)|~node14(X2,X1)|~m_main_v_state2(X2,c_n2)), inference(spm,[status(thm)],[c_0_64, c_0_16])).
cnf(c_0_387,plain,(node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X2,c_c1)|~node16(X1)|~node7(X1,X2)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_65, c_0_37])).
cnf(c_0_388,plain,(node17(s5,s6)|node16(s5)), inference(sr,[status(thm)],[c_0_379, c_0_380])).
cnf(c_0_389,plain,(m_main_v_turn(X1,c_1)|~node18(s5,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_382]), c_0_357]), c_0_383])]), c_0_380])).
cnf(c_0_390,plain,(node16(s4)|node15(s4)|m_main_v_turn(s5,c_2)|~node1(s3)), inference(spm,[status(thm)],[c_0_384, c_0_355])).
cnf(c_0_391,plain,(~node15(s4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_385, c_0_337])])).
cnf(c_0_392,plain,(~node16(X1)|~m_main_v_state2(X1,c_c2)), inference(spm,[status(thm)],[c_0_4, c_0_37])).
cnf(c_0_393,plain,(m_main_v_state2(X1,c_t2)|~node16(X2)|~node14(X2,X1)), inference(spm,[status(thm)],[c_0_386, c_0_37])).
cnf(c_0_394,plain,(node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X2,c_c1)|~node16(X1)|~node7(X1,X2)), inference(spm,[status(thm)],[c_0_387, c_0_16])).
cnf(c_0_395,plain,(node16(s5)|m_main_v_turn(s5,X1)|~m_main_v_turn(s6,X1)), inference(spm,[status(thm)],[c_0_66, c_0_388])).
cnf(c_0_396,plain,(m_main_v_turn(X1,c_1)|~trans(s5,X1)), inference(spm,[status(thm)],[c_0_389, c_0_15])).
cnf(c_0_397,plain,(node16(s4)|m_main_v_turn(s5,c_2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_390, c_0_337])]), c_0_391])).
cnf(c_0_398,plain,(~node16(s4)), inference(spm,[status(thm)],[c_0_392, c_0_346])).
cnf(c_0_399,plain,(~m_main_v_state2(X1,c_c2)|~trans(s4,X1)), inference(spm,[status(thm)],[c_0_4, c_0_356])).
cnf(c_0_400,plain,(m_main_v_state2(X1,c_t2)|~node16(X2)|~trans(X2,X1)), inference(spm,[status(thm)],[c_0_393, c_0_29])).
cnf(c_0_401,plain,(node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X2,c_c1)|~node16(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_394, c_0_7])).
cnf(c_0_402,plain,(node16(s5)|m_main_v_turn(s5,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_396]), c_0_375])])).
cnf(c_0_403,plain,(m_main_v_turn(s5,c_2)), inference(sr,[status(thm)],[c_0_397, c_0_398])).
cnf(c_0_404,plain,(~node3(X1)|~trans(s4,X1)), inference(spm,[status(thm)],[c_0_399, c_0_12])).
cnf(c_0_405,plain,(~node1(X1)|~trans(s4,X1)), inference(spm,[status(thm)],[c_0_367, c_0_5])).
cnf(c_0_406,plain,(m_main_v_state2(s6,c_t2)|~node16(s5)), inference(spm,[status(thm)],[c_0_400, c_0_375])).
cnf(c_0_407,plain,(node3(s5)|node2(s5)|node1(s5)|m_main_v_state1(s6,c_c1)|~node16(s5)), inference(spm,[status(thm)],[c_0_401, c_0_375])).
cnf(c_0_408,plain,(node16(s5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_402]), c_0_403])])).
cnf(c_0_409,plain,(~node3(s5)), inference(spm,[status(thm)],[c_0_404, c_0_357])).
cnf(c_0_410,plain,(~node2(s5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_382]), c_0_383])])).
cnf(c_0_411,plain,(~node1(s5)), inference(spm,[status(thm)],[c_0_405, c_0_357])).
cnf(c_0_412,plain,(~node16(s5)|~node4(s6)), inference(spm,[status(thm)],[c_0_200, c_0_406])).
cnf(c_0_413,plain,(~node16(s5)|~node2(s6)), inference(spm,[status(thm)],[c_0_239, c_0_406])).
cnf(c_0_414,plain,(m_main_v_state1(s6,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_407, c_0_408])]), c_0_409]), c_0_410]), c_0_411])).
cnf(c_0_415,plain,(~node3(X1)|~m_main_v_state1(X1,c_c1)), inference(spm,[status(thm)],[c_0_13, c_0_40])).
cnf(c_0_416,plain,(~node1(X1)|~m_main_v_state1(X1,c_c1)), inference(spm,[status(thm)],[c_0_13, c_0_41])).
cnf(c_0_417,plain,(~node4(s6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_412, c_0_408])])).
cnf(c_0_418,plain,(~node2(s6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_413, c_0_408])])).
cnf(c_0_419,plain,(~node5(s6)), inference(spm,[status(thm)],[c_0_233, c_0_414])).
cnf(c_0_420,plain,(~node3(s6)), inference(spm,[status(thm)],[c_0_415, c_0_414])).
cnf(c_0_421,plain,(~node1(s6)), inference(spm,[status(thm)],[c_0_416, c_0_414])).
cnf(c_0_422,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~node16(X1)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_63, c_0_37])).
cnf(c_0_423,plain,(m_main_v_state1(X1,c_n1)|~node7(s6,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_414]), c_0_417]), c_0_418]), c_0_419]), c_0_420]), c_0_421])).
cnf(c_0_424,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~node16(X1)), inference(spm,[status(thm)],[c_0_422, c_0_16])).
cnf(c_0_425,plain,(m_main_v_state1(X1,c_n1)|~trans(s6,X1)), inference(spm,[status(thm)],[c_0_423, c_0_7])).
cnf(c_0_426,plain,(trans(s6,s7)), inference(spm,[status(thm)],[c_0_8, c_0_67])).
cnf(c_0_427,plain,(~node16(s5)|~node8(s6)), inference(spm,[status(thm)],[c_0_168, c_0_406])).
cnf(c_0_428,plain,(~node12(X1)|~m_main_v_state1(X1,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_68])).
cnf(c_0_429,plain,(~node11(X1)|~m_main_v_state1(X1,c_c1)), inference(spm,[status(thm)],[c_0_13, c_0_25])).
cnf(c_0_430,plain,(~node16(s5)|~node10(s6)), inference(spm,[status(thm)],[c_0_257, c_0_406])).
cnf(c_0_431,plain,(~node16(s5)|~node9(s6)), inference(spm,[status(thm)],[c_0_258, c_0_406])).
cnf(c_0_432,plain,(~node16(s5)|~m_main_v_state2(s6,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_406])).
cnf(c_0_433,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node16(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_424, c_0_15])).
cnf(c_0_434,plain,(~m_main_v_state1(X1,c_t1)|~trans(s6,X1)), inference(spm,[status(thm)],[c_0_2, c_0_425])).
cnf(c_0_435,plain,(node13(s6,s7)|node12(s6)|node11(s6)|node10(s6)|node9(s6)|node8(s6)|m_main_v_state2(s6,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_426])).
cnf(c_0_436,plain,(~node8(s6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_427, c_0_408])])).
cnf(c_0_437,plain,(~node12(s6)), inference(spm,[status(thm)],[c_0_428, c_0_414])).
cnf(c_0_438,plain,(~node11(s6)), inference(spm,[status(thm)],[c_0_429, c_0_414])).
cnf(c_0_439,plain,(~node10(s6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_430, c_0_408])])).
cnf(c_0_440,plain,(~node9(s6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_431, c_0_408])])).
cnf(c_0_441,plain,(~m_main_v_state2(s6,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_432, c_0_408])])).
cnf(c_0_442,plain,(node17(s6,s7)|node16(s6)|node15(s6)), inference(spm,[status(thm)],[c_0_152, c_0_426])).
cnf(c_0_443,plain,(node15(s5)|m_main_v_turn(s6,c_1)|~node16(s5)), inference(spm,[status(thm)],[c_0_433, c_0_375])).
cnf(c_0_444,plain,(~node16(s6)|~node16(s5)), inference(spm,[status(thm)],[c_0_188, c_0_406])).
cnf(c_0_445,plain,(trans(s7,s8)), inference(spm,[status(thm)],[c_0_8, c_0_69])).
cnf(c_0_446,plain,(~node16(X1)|~trans(s6,X1)), inference(spm,[status(thm)],[c_0_434, c_0_16])).
cnf(c_0_447,plain,(node13(s6,s7)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_435, c_0_436]), c_0_437]), c_0_438]), c_0_439]), c_0_440]), c_0_441])).
cnf(c_0_448,plain,(node16(s6)|node15(s6)|m_main_v_turn(s7,X1)|~m_main_v_turn(s6,X1)), inference(spm,[status(thm)],[c_0_24, c_0_442])).
cnf(c_0_449,plain,(m_main_v_turn(s6,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_443, c_0_408])]), c_0_380])).
cnf(c_0_450,plain,(~node16(s6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_444, c_0_408])])).
cnf(c_0_451,plain,(node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X2,c_c2)|~node15(X1)|~node14(X1,X2)|~m_main_v_state2(X1,c_t2)), inference(spm,[status(thm)],[c_0_57, c_0_38])).
cnf(c_0_452,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node15(X2)|~m_main_v_state1(X2,c_n1)), inference(spm,[status(thm)],[c_0_45, c_0_17])).
cnf(c_0_453,plain,(node17(s7,s8)|node16(s7)|node15(s7)), inference(spm,[status(thm)],[c_0_152, c_0_445])).
cnf(c_0_454,plain,(~node16(s7)), inference(spm,[status(thm)],[c_0_446, c_0_426])).
cnf(c_0_455,plain,(m_main_v_state2(s7,X1)|~m_main_v_state2(s6,X1)), inference(spm,[status(thm)],[c_0_43, c_0_447])).
cnf(c_0_456,plain,(m_main_v_state2(s6,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_406, c_0_408])])).
cnf(c_0_457,plain,(node15(s6)|m_main_v_turn(s7,c_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_448, c_0_449]), c_0_450])).
cnf(c_0_458,plain,(~node15(s6)), inference(spm,[status(thm)],[c_0_229, c_0_414])).
cnf(c_0_459,plain,(m_main_v_state1(X1,c_t1)|~node15(X2)|~node7(X2,X1)|~m_main_v_state1(X2,c_n1)), inference(spm,[status(thm)],[c_0_56, c_0_17])).
cnf(c_0_460,plain,(node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X2,c_c2)|~node15(X1)|~node14(X1,X2)), inference(spm,[status(thm)],[c_0_451, c_0_17])).
cnf(c_0_461,plain,(trans(s8,s9)), inference(spm,[status(thm)],[c_0_8, c_0_70])).
cnf(c_0_462,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node15(X2)), inference(spm,[status(thm)],[c_0_452, c_0_38])).
cnf(c_0_463,plain,(node17(s7,s8)|node15(s7)), inference(sr,[status(thm)],[c_0_453, c_0_454])).
cnf(c_0_464,plain,(m_main_v_turn(X1,c_2)|~node18(s7,X1)|~m_main_v_state1(s7,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_455]), c_0_456])])).
cnf(c_0_465,plain,(m_main_v_turn(s7,c_1)), inference(sr,[status(thm)],[c_0_457, c_0_458])).
cnf(c_0_466,plain,(m_main_v_state1(X1,c_t1)|~node15(X2)|~node7(X2,X1)), inference(spm,[status(thm)],[c_0_459, c_0_38])).
cnf(c_0_467,plain,(node10(X1)|node9(X1)|node8(X1)|m_main_v_state2(X2,c_c2)|~node15(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_460, c_0_29])).
cnf(c_0_468,plain,(node15(s6)|~m_main_v_turn(s7,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_457])).
cnf(c_0_469,plain,(node17(s8,s9)|node16(s8)|node15(s8)), inference(spm,[status(thm)],[c_0_152, c_0_461])).
cnf(c_0_470,plain,(m_main_v_turn(X1,c_2)|~node15(X2)|~trans(X2,X1)), inference(spm,[status(thm)],[c_0_462, c_0_15])).
cnf(c_0_471,plain,(node15(s7)|m_main_v_turn(s7,X1)|~m_main_v_turn(s8,X1)), inference(spm,[status(thm)],[c_0_66, c_0_463])).
cnf(c_0_472,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s7,c_n1)|~trans(s7,X1)), inference(spm,[status(thm)],[c_0_464, c_0_15])).
cnf(c_0_473,plain,(~m_main_v_turn(s7,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_465])).
cnf(c_0_474,plain,(m_main_v_state1(X1,c_t1)|~node15(X2)|~trans(X2,X1)), inference(spm,[status(thm)],[c_0_466, c_0_7])).
cnf(c_0_475,plain,(~m_main_v_state1(X1,c_c1)|~trans(s6,X1)), inference(spm,[status(thm)],[c_0_13, c_0_425])).
cnf(c_0_476,plain,(node10(s6)|node9(s6)|node8(s6)|m_main_v_state2(s7,c_c2)|~node15(s6)), inference(spm,[status(thm)],[c_0_467, c_0_426])).
cnf(c_0_477,plain,(node15(s6)|~node12(s7)), inference(spm,[status(thm)],[c_0_468, c_0_35])).
cnf(c_0_478,plain,(node16(s8)|node15(s8)|m_main_v_turn(s9,X1)|~m_main_v_turn(s8,X1)), inference(spm,[status(thm)],[c_0_24, c_0_469])).
cnf(c_0_479,plain,(m_main_v_turn(s8,c_2)|~node15(s7)), inference(spm,[status(thm)],[c_0_470, c_0_445])).
cnf(c_0_480,plain,(node15(s7)|~m_main_v_state1(s7,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_471, c_0_472]), c_0_445])]), c_0_473])).
cnf(c_0_481,plain,(m_main_v_state1(s8,c_t1)|~node15(s7)), inference(spm,[status(thm)],[c_0_474, c_0_445])).
cnf(c_0_482,plain,(~node10(X1)|~trans(s6,X1)), inference(spm,[status(thm)],[c_0_475, c_0_32])).
cnf(c_0_483,plain,(~node8(X1)|~trans(s6,X1)), inference(spm,[status(thm)],[c_0_434, c_0_19])).
cnf(c_0_484,plain,(m_main_v_state2(s7,c_c2)|~node12(s7)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_477]), c_0_439]), c_0_440]), c_0_436])).
cnf(c_0_485,plain,(node16(s8)|node15(s8)|m_main_v_turn(s9,c_2)|~node15(s7)), inference(spm,[status(thm)],[c_0_478, c_0_479])).
cnf(c_0_486,plain,(node15(s7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_425]), c_0_426])])).
cnf(c_0_487,plain,(~node15(s8)|~node15(s7)), inference(spm,[status(thm)],[c_0_189, c_0_481])).
cnf(c_0_488,plain,(node10(s7)|node9(s7)|node8(s7)|m_main_v_state2(s8,c_c2)|~node15(s7)), inference(spm,[status(thm)],[c_0_467, c_0_445])).
cnf(c_0_489,plain,(~node10(s7)), inference(spm,[status(thm)],[c_0_482, c_0_426])).
cnf(c_0_490,plain,(~node9(s7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_455]), c_0_456])])).
cnf(c_0_491,plain,(~node8(s7)), inference(spm,[status(thm)],[c_0_483, c_0_426])).
cnf(c_0_492,plain,(node13(s7,s8)|node12(s7)|node11(s7)|node10(s7)|node9(s7)|node8(s7)|m_main_v_state2(s7,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_445])).
cnf(c_0_493,plain,(~node12(s7)), inference(spm,[status(thm)],[c_0_345, c_0_484])).
cnf(c_0_494,plain,(~m_main_v_state2(s7,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_455]), c_0_456])])).
cnf(c_0_495,plain,(node16(s8)|node15(s8)|m_main_v_turn(s9,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_485, c_0_486])])).
cnf(c_0_496,plain,(~node15(s8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_487, c_0_486])])).
cnf(c_0_497,plain,(~node15(s7)|~node1(s8)), inference(spm,[status(thm)],[c_0_203, c_0_481])).
cnf(c_0_498,plain,(~node15(s7)|~node5(s8)), inference(spm,[status(thm)],[c_0_354, c_0_479])).
cnf(c_0_499,plain,(m_main_v_state2(s8,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_488, c_0_486])]), c_0_489]), c_0_490]), c_0_491])).
cnf(c_0_500,plain,(~node15(s7)|~node3(s8)), inference(spm,[status(thm)],[c_0_201, c_0_481])).
cnf(c_0_501,plain,(~node15(s7)|~node2(s8)), inference(spm,[status(thm)],[c_0_202, c_0_481])).
cnf(c_0_502,plain,(~node15(s7)|~m_main_v_state1(s8,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_481])).
cnf(c_0_503,plain,(~node4(X1)|~trans(s6,X1)), inference(spm,[status(thm)],[c_0_434, c_0_11])).
cnf(c_0_504,plain,(~node5(X1)|~trans(s6,X1)), inference(spm,[status(thm)],[c_0_434, c_0_10])).
cnf(c_0_505,plain,(m_main_v_state2(s6,X1)|~m_main_v_state2(s7,X1)), inference(spm,[status(thm)],[c_0_47, c_0_447])).
cnf(c_0_506,plain,(node13(s7,s8)|node11(s7)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_492, c_0_493]), c_0_489]), c_0_490]), c_0_491]), c_0_494])).
cnf(c_0_507,plain,(node16(s8)|m_main_v_turn(s9,c_2)), inference(sr,[status(thm)],[c_0_495, c_0_496])).
cnf(c_0_508,plain,(node6(s8,s9)|node5(s8)|node4(s8)|node3(s8)|node2(s8)|node1(s8)|m_main_v_state1(s8,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_461])).
cnf(c_0_509,plain,(~node1(s8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_497, c_0_486])])).
cnf(c_0_510,plain,(~node5(s8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_498, c_0_486])])).
cnf(c_0_511,plain,(~node4(s8)), inference(spm,[status(thm)],[c_0_361, c_0_499])).
cnf(c_0_512,plain,(~node3(s8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_500, c_0_486])])).
cnf(c_0_513,plain,(~node2(s8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_501, c_0_486])])).
cnf(c_0_514,plain,(~m_main_v_state1(s8,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_502, c_0_486])])).
cnf(c_0_515,plain,(node6(s7,s8)|node5(s7)|node4(s7)|node3(s7)|node2(s7)|node1(s7)|m_main_v_state1(s7,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_445])).
cnf(c_0_516,plain,(~node4(s7)), inference(spm,[status(thm)],[c_0_503, c_0_426])).
cnf(c_0_517,plain,(~node5(s7)), inference(spm,[status(thm)],[c_0_504, c_0_426])).
cnf(c_0_518,plain,(~node3(s7)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_505, c_0_12]), c_0_441])).
cnf(c_0_519,plain,(~node2(s7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_455]), c_0_456])])).
cnf(c_0_520,plain,(node11(s7)|m_main_v_state2(s8,X1)|~m_main_v_state2(s7,X1)), inference(spm,[status(thm)],[c_0_43, c_0_506])).
cnf(c_0_521,plain,(node3(s8)|node2(s8)|node1(s8)|m_main_v_state1(s9,c_c1)|~node16(s8)), inference(spm,[status(thm)],[c_0_401, c_0_461])).
cnf(c_0_522,plain,(node16(s8)|~node5(s9)), inference(spm,[status(thm)],[c_0_354, c_0_507])).
cnf(c_0_523,plain,(node6(s8,s9)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_508, c_0_509]), c_0_510]), c_0_511]), c_0_512]), c_0_513]), c_0_514])).
cnf(c_0_524,plain,(node6(s7,s8)|node1(s7)|m_main_v_state1(s7,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_515, c_0_516]), c_0_517]), c_0_518]), c_0_519])).
cnf(c_0_525,plain,(node11(s7)|m_main_v_state2(s8,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_520, c_0_17]), c_0_486])])).
cnf(c_0_526,plain,(~node15(s7)|~node11(s8)), inference(spm,[status(thm)],[c_0_165, c_0_481])).
cnf(c_0_527,plain,(~node15(s7)|~node9(s8)), inference(spm,[status(thm)],[c_0_167, c_0_481])).
cnf(c_0_528,plain,(trans(s9,s10)), inference(spm,[status(thm)],[c_0_8, c_0_71])).
cnf(c_0_529,plain,(m_main_v_state1(s9,c_c1)|~node5(s9)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_521, c_0_522]), c_0_512]), c_0_513]), c_0_509])).
cnf(c_0_530,plain,(m_main_v_state1(s9,X1)|~m_main_v_state1(s8,X1)), inference(spm,[status(thm)],[c_0_23, c_0_523])).
cnf(c_0_531,plain,(m_main_v_state1(s8,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_481, c_0_486])])).
cnf(c_0_532,plain,(node1(s7)|m_main_v_state1(s7,c_c1)|m_main_v_state1(s8,X1)|~m_main_v_state1(s7,X1)), inference(spm,[status(thm)],[c_0_23, c_0_524])).
cnf(c_0_533,plain,(node11(s7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_525]), c_0_499])])).
cnf(c_0_534,plain,(m_main_v_state1(s8,X1)|~m_main_v_state1(s9,X1)), inference(spm,[status(thm)],[c_0_49, c_0_523])).
cnf(c_0_535,plain,(~node12(s8)), inference(spm,[status(thm)],[c_0_345, c_0_499])).
cnf(c_0_536,plain,(~node11(s8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_526, c_0_486])])).
cnf(c_0_537,plain,(~node10(s8)), inference(spm,[status(thm)],[c_0_339, c_0_499])).
cnf(c_0_538,plain,(~node9(s8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_527, c_0_486])])).
cnf(c_0_539,plain,(~node8(s8)), inference(spm,[status(thm)],[c_0_343, c_0_499])).
cnf(c_0_540,plain,(node6(s9,s10)|node5(s9)|node4(s9)|node3(s9)|node2(s9)|node1(s9)|m_main_v_state1(s9,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_528])).
cnf(c_0_541,plain,(~node5(s9)), inference(spm,[status(thm)],[c_0_233, c_0_529])).
cnf(c_0_542,plain,(~node3(s9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_530]), c_0_531])])).
cnf(c_0_543,plain,(~node2(s9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_530]), c_0_531])])).
cnf(c_0_544,plain,(~node1(s9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_530]), c_0_531])])).
cnf(c_0_545,plain,(~m_main_v_state1(s9,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_530]), c_0_531])])).
cnf(c_0_546,plain,(node1(s7)|m_main_v_state1(s7,c_c1)|m_main_v_state1(s8,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_25]), c_0_533])])).
cnf(c_0_547,plain,(node13(s9,s10)|node12(s9)|node11(s9)|node10(s9)|node9(s9)|node8(s9)|m_main_v_state2(s9,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_528])).
cnf(c_0_548,plain,(~node11(s9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_530]), c_0_531])])).
cnf(c_0_549,plain,(~node10(s9)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_32]), c_0_514])).
cnf(c_0_550,plain,(~node9(s9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_530]), c_0_531])])).
cnf(c_0_551,plain,(m_main_v_state2(X1,c_n2)|~node14(s8,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_499]), c_0_535]), c_0_536]), c_0_537]), c_0_538]), c_0_539])).
cnf(c_0_552,plain,(node6(s9,s10)|node4(s9)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_540, c_0_541]), c_0_542]), c_0_543]), c_0_544]), c_0_545])).
cnf(c_0_553,plain,(m_main_v_state1(X1,c_t1)|~node7(s7,X1)|~m_main_v_state1(s7,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_455]), c_0_456])])).
cnf(c_0_554,plain,(node1(s7)|m_main_v_state1(s7,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_546]), c_0_531])])).
cnf(c_0_555,plain,(node13(s9,s10)|node12(s9)|node8(s9)|m_main_v_state2(s9,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_547, c_0_548]), c_0_549]), c_0_550])).
cnf(c_0_556,plain,(m_main_v_state2(X1,c_t2)|~node14(s9,X1)|~m_main_v_state2(s9,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_530]), c_0_531])])).
cnf(c_0_557,plain,(m_main_v_state2(X1,c_n2)|~trans(s8,X1)), inference(spm,[status(thm)],[c_0_551, c_0_29])).
cnf(c_0_558,plain,(node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X2,c_c1)|~node7(X1,X2)|~node4(X1)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_65, c_0_39])).
cnf(c_0_559,plain,(node4(s9)|m_main_v_state1(s10,X1)|~m_main_v_state1(s9,X1)), inference(spm,[status(thm)],[c_0_23, c_0_552])).
cnf(c_0_560,plain,(m_main_v_state1(X1,c_t1)|~m_main_v_state1(s7,c_n1)|~trans(s7,X1)), inference(spm,[status(thm)],[c_0_553, c_0_7])).
cnf(c_0_561,plain,(node1(s7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_554]), c_0_486])])).
cnf(c_0_562,plain,(node12(s9)|node8(s9)|m_main_v_state2(s9,c_c2)|m_main_v_state2(s9,X1)|~m_main_v_state2(s10,X1)), inference(spm,[status(thm)],[c_0_47, c_0_555])).
cnf(c_0_563,plain,(m_main_v_state2(X1,c_t2)|~m_main_v_state2(s9,c_n2)|~trans(s9,X1)), inference(spm,[status(thm)],[c_0_556, c_0_29])).
cnf(c_0_564,plain,(~m_main_v_state2(X1,c_t2)|~trans(s8,X1)), inference(spm,[status(thm)],[c_0_0, c_0_557])).
cnf(c_0_565,plain,(node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X2,c_c1)|~node7(X1,X2)|~node4(X1)), inference(spm,[status(thm)],[c_0_558, c_0_11])).
cnf(c_0_566,plain,(node4(s9)|m_main_v_state1(s10,X1)|~m_main_v_state1(s8,X1)), inference(spm,[status(thm)],[c_0_559, c_0_530])).
cnf(c_0_567,plain,(m_main_v_state1(X1,c_t1)|~trans(s7,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_560, c_0_41]), c_0_561])])).
cnf(c_0_568,plain,(node12(s9)|node8(s9)|m_main_v_state2(s9,c_c2)|m_main_v_state2(s9,c_t2)|~m_main_v_state2(s9,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_562, c_0_563]), c_0_528])])).
cnf(c_0_569,plain,(~node12(X1)|~trans(s8,X1)), inference(spm,[status(thm)],[c_0_564, c_0_58])).
cnf(c_0_570,plain,(node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X2,c_c1)|~node4(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_565, c_0_7])).
cnf(c_0_571,plain,(node4(s9)|m_main_v_state1(s10,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_567]), c_0_445])])).
cnf(c_0_572,plain,(m_main_v_state2(X1,c_t2)|~node14(X2,X1)|~node8(X2)|~m_main_v_state2(X2,c_n2)), inference(spm,[status(thm)],[c_0_64, c_0_19])).
cnf(c_0_573,plain,(node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X2,c_c1)|~node8(X1)|~node7(X1,X2)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_65, c_0_28])).
cnf(c_0_574,plain,(node12(s9)|node8(s9)|m_main_v_state2(s9,c_c2)|m_main_v_state2(s9,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_568, c_0_557]), c_0_461])])).
cnf(c_0_575,plain,(~node12(s9)), inference(spm,[status(thm)],[c_0_569, c_0_461])).
cnf(c_0_576,plain,(node3(s9)|node2(s9)|node1(s9)|m_main_v_state1(s10,c_c1)|~node4(s9)), inference(spm,[status(thm)],[c_0_570, c_0_528])).
cnf(c_0_577,plain,(node4(s9)|~node11(s10)), inference(spm,[status(thm)],[c_0_165, c_0_571])).
cnf(c_0_578,plain,(node4(s9)|~node9(s10)), inference(spm,[status(thm)],[c_0_167, c_0_571])).
cnf(c_0_579,plain,(m_main_v_state2(X1,c_t2)|~node14(X2,X1)|~node8(X2)), inference(spm,[status(thm)],[c_0_572, c_0_28])).
cnf(c_0_580,plain,(node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X2,c_c1)|~node8(X1)|~node7(X1,X2)), inference(spm,[status(thm)],[c_0_573, c_0_19])).
cnf(c_0_581,plain,(~m_main_v_state2(X1,c_c2)|~trans(s8,X1)), inference(spm,[status(thm)],[c_0_4, c_0_557])).
cnf(c_0_582,plain,(node8(s9)|m_main_v_state2(s9,c_c2)|m_main_v_state2(s9,c_t2)), inference(sr,[status(thm)],[c_0_574, c_0_575])).
cnf(c_0_583,plain,(~node16(X1)|~m_main_v_state2(s9,c_n2)|~trans(s9,X1)), inference(spm,[status(thm)],[c_0_188, c_0_563])).
cnf(c_0_584,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~node8(X1)|~m_main_v_state1(X1,c_t1)), inference(spm,[status(thm)],[c_0_63, c_0_28])).
cnf(c_0_585,plain,(trans(s10,s11)), inference(spm,[status(thm)],[c_0_8, c_0_72])).
cnf(c_0_586,plain,(m_main_v_state1(s10,c_c1)|~node11(s10)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_576, c_0_577]), c_0_542]), c_0_543]), c_0_544])).
cnf(c_0_587,plain,(~node9(X1)|~m_main_v_state1(X1,c_c1)), inference(spm,[status(thm)],[c_0_13, c_0_27])).
cnf(c_0_588,plain,(m_main_v_state1(s10,c_c1)|~node9(s10)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_576, c_0_578]), c_0_542]), c_0_543]), c_0_544])).
cnf(c_0_589,plain,(~node8(X1)|~m_main_v_state2(s9,c_n2)|~trans(s9,X1)), inference(spm,[status(thm)],[c_0_168, c_0_563])).
cnf(c_0_590,plain,(~node10(X1)|~m_main_v_state2(s9,c_n2)|~trans(s9,X1)), inference(spm,[status(thm)],[c_0_257, c_0_563])).
cnf(c_0_591,plain,(m_main_v_state2(X1,c_t2)|~node8(X2)|~trans(X2,X1)), inference(spm,[status(thm)],[c_0_579, c_0_29])).
cnf(c_0_592,plain,(node3(X1)|node2(X1)|node1(X1)|m_main_v_state1(X2,c_c1)|~node8(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_580, c_0_7])).
cnf(c_0_593,plain,(node8(s9)|m_main_v_state2(s9,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_581, c_0_582]), c_0_461])])).
cnf(c_0_594,plain,(node17(s9,s10)|node16(s9)|node15(s9)), inference(spm,[status(thm)],[c_0_152, c_0_528])).
cnf(c_0_595,plain,(~node16(X1)|~trans(s9,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_557]), c_0_461])])).
cnf(c_0_596,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~node8(X1)), inference(spm,[status(thm)],[c_0_584, c_0_19])).
cnf(c_0_597,plain,(node13(s10,s11)|node12(s10)|node11(s10)|node10(s10)|node9(s10)|node8(s10)|m_main_v_state2(s10,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_585])).
cnf(c_0_598,plain,(~node11(s10)), inference(spm,[status(thm)],[c_0_429, c_0_586])).
cnf(c_0_599,plain,(~node9(s10)), inference(spm,[status(thm)],[c_0_587, c_0_588])).
cnf(c_0_600,plain,(~node8(X1)|~trans(s9,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_589, c_0_557]), c_0_461])])).
cnf(c_0_601,plain,(~node10(X1)|~trans(s9,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_590, c_0_557]), c_0_461])])).
cnf(c_0_602,plain,(m_main_v_state2(s10,c_t2)|~node8(s9)), inference(spm,[status(thm)],[c_0_591, c_0_528])).
cnf(c_0_603,plain,(node3(s9)|node2(s9)|node1(s9)|m_main_v_state1(s10,c_c1)|~node8(s9)), inference(spm,[status(thm)],[c_0_592, c_0_528])).
cnf(c_0_604,plain,(node8(s9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_564, c_0_593]), c_0_461])])).
cnf(c_0_605,plain,(node16(s9)|node15(s9)|m_main_v_turn(s10,X1)|~m_main_v_turn(s9,X1)), inference(spm,[status(thm)],[c_0_24, c_0_594])).
cnf(c_0_606,plain,(node17(s10,s11)|node16(s10)|node15(s10)), inference(spm,[status(thm)],[c_0_152, c_0_585])).
cnf(c_0_607,plain,(~node16(s10)), inference(spm,[status(thm)],[c_0_595, c_0_528])).
cnf(c_0_608,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node8(X1)|~trans(X1,X2)), inference(spm,[status(thm)],[c_0_596, c_0_15])).
cnf(c_0_609,plain,(node13(s10,s11)|node12(s10)|node10(s10)|node8(s10)|m_main_v_state2(s10,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_597, c_0_598]), c_0_599])).
cnf(c_0_610,plain,(~node8(s10)), inference(spm,[status(thm)],[c_0_600, c_0_528])).
cnf(c_0_611,plain,(~node10(s10)), inference(spm,[status(thm)],[c_0_601, c_0_528])).
cnf(c_0_612,plain,(~node8(s9)|~m_main_v_state2(s10,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_602])).
cnf(c_0_613,plain,(m_main_v_state1(s10,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_603, c_0_604])]), c_0_542]), c_0_543]), c_0_544])).
cnf(c_0_614,plain,(node16(s9)|node16(s8)|node15(s9)|m_main_v_turn(s10,c_2)), inference(spm,[status(thm)],[c_0_605, c_0_507])).
cnf(c_0_615,plain,(~node16(s8)), inference(spm,[status(thm)],[c_0_392, c_0_499])).
cnf(c_0_616,plain,(m_main_v_state2(X1,c_t2)|~node14(X2,X1)|~node4(X2)|~m_main_v_state2(X2,c_n2)), inference(spm,[status(thm)],[c_0_64, c_0_11])).
cnf(c_0_617,plain,(node17(s10,s11)|node15(s10)), inference(sr,[status(thm)],[c_0_606, c_0_607])).
cnf(c_0_618,plain,(node15(s9)|m_main_v_turn(s10,c_1)|~node8(s9)), inference(spm,[status(thm)],[c_0_608, c_0_528])).
cnf(c_0_619,plain,(~node15(s9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_530]), c_0_531])])).
cnf(c_0_620,plain,(node13(s10,s11)|node12(s10)|m_main_v_state2(s10,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_609, c_0_610]), c_0_611])).
cnf(c_0_621,plain,(~m_main_v_state2(s10,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_612, c_0_604])])).
cnf(c_0_622,plain,(~node12(s10)), inference(spm,[status(thm)],[c_0_428, c_0_613])).
cnf(c_0_623,plain,(node16(s9)|node15(s9)|m_main_v_turn(s10,c_2)), inference(sr,[status(thm)],[c_0_614, c_0_615])).
cnf(c_0_624,plain,(m_main_v_state2(X1,c_t2)|~node14(X2,X1)|~node4(X2)), inference(spm,[status(thm)],[c_0_616, c_0_39])).
cnf(c_0_625,plain,(trans(s11,s12)), inference(spm,[status(thm)],[c_0_8, c_0_73])).
cnf(c_0_626,plain,(node15(s10)|m_main_v_turn(s11,X1)|~m_main_v_turn(s10,X1)), inference(spm,[status(thm)],[c_0_24, c_0_617])).
cnf(c_0_627,plain,(m_main_v_turn(s10,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_618, c_0_604])]), c_0_619])).
cnf(c_0_628,plain,(node13(s10,s11)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_620, c_0_621]), c_0_622])).
cnf(c_0_629,plain,(node16(s9)|m_main_v_turn(s10,c_2)), inference(sr,[status(thm)],[c_0_623, c_0_619])).
cnf(c_0_630,plain,(m_main_v_state2(X1,c_t2)|~node4(X2)|~trans(X2,X1)), inference(spm,[status(thm)],[c_0_624, c_0_29])).
cnf(c_0_631,plain,(node17(s11,s12)|node16(s11)|node15(s11)), inference(spm,[status(thm)],[c_0_152, c_0_625])).
cnf(c_0_632,plain,(node15(s10)|m_main_v_turn(s11,c_1)), inference(spm,[status(thm)],[c_0_626, c_0_627])).
cnf(c_0_633,plain,(~node15(s10)), inference(spm,[status(thm)],[c_0_229, c_0_613])).
cnf(c_0_634,plain,(m_main_v_state2(s11,X1)|~m_main_v_state2(s10,X1)), inference(spm,[status(thm)],[c_0_43, c_0_628])).
cnf(c_0_635,plain,(m_main_v_state2(s10,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_602, c_0_604])])).
cnf(c_0_636,plain,(node3(s9)|node2(s9)|node1(s9)|m_main_v_state1(s10,c_c1)|~node16(s9)), inference(spm,[status(thm)],[c_0_401, c_0_528])).
cnf(c_0_637,plain,(node16(s9)|~node5(s10)), inference(spm,[status(thm)],[c_0_354, c_0_629])).
cnf(c_0_638,plain,(~node4(X1)|~m_main_v_state2(s9,c_n2)|~trans(s9,X1)), inference(spm,[status(thm)],[c_0_200, c_0_563])).
cnf(c_0_639,plain,(node4(s9)|~node3(s10)), inference(spm,[status(thm)],[c_0_201, c_0_571])).
cnf(c_0_640,plain,(m_main_v_state2(s10,c_t2)|~node4(s9)), inference(spm,[status(thm)],[c_0_630, c_0_528])).
cnf(c_0_641,plain,(node4(s9)|~node1(s10)), inference(spm,[status(thm)],[c_0_203, c_0_571])).
cnf(c_0_642,plain,(node16(s11)|node15(s11)|m_main_v_turn(s12,X1)|~m_main_v_turn(s11,X1)), inference(spm,[status(thm)],[c_0_24, c_0_631])).
cnf(c_0_643,plain,(m_main_v_turn(s11,c_1)), inference(sr,[status(thm)],[c_0_632, c_0_633])).
cnf(c_0_644,plain,(~node16(s11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_634]), c_0_635])])).
cnf(c_0_645,plain,(m_main_v_state1(s10,c_c1)|~node5(s10)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_636, c_0_637]), c_0_542]), c_0_543]), c_0_544])).
cnf(c_0_646,plain,(~node4(X1)|~trans(s9,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_638, c_0_557]), c_0_461])])).
cnf(c_0_647,plain,(m_main_v_state1(s10,c_c1)|~node3(s10)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_576, c_0_639]), c_0_542]), c_0_543]), c_0_544])).
cnf(c_0_648,plain,(~node4(s9)|~node2(s10)), inference(spm,[status(thm)],[c_0_239, c_0_640])).
cnf(c_0_649,plain,(node4(s9)|~node2(s10)), inference(spm,[status(thm)],[c_0_202, c_0_571])).
cnf(c_0_650,plain,(m_main_v_state1(s10,c_c1)|~node1(s10)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_576, c_0_641]), c_0_542]), c_0_543]), c_0_544])).
cnf(c_0_651,plain,(node16(s11)|node15(s11)|m_main_v_turn(s12,c_1)), inference(spm,[status(thm)],[c_0_642, c_0_643])).
cnf(c_0_652,plain,(node17(s11,s12)|node15(s11)), inference(sr,[status(thm)],[c_0_631, c_0_644])).
cnf(c_0_653,plain,(m_main_v_turn(X1,c_2)|~node18(s11,X1)|~m_main_v_state1(s11,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_634]), c_0_635])])).
cnf(c_0_654,plain,(~node5(s10)), inference(spm,[status(thm)],[c_0_233, c_0_645])).
cnf(c_0_655,plain,(~node4(s10)), inference(spm,[status(thm)],[c_0_646, c_0_528])).
cnf(c_0_656,plain,(~node3(s10)), inference(spm,[status(thm)],[c_0_415, c_0_647])).
cnf(c_0_657,plain,(~node2(s10)), inference(spm,[status(thm)],[c_0_648, c_0_649])).
cnf(c_0_658,plain,(~node1(s10)), inference(spm,[status(thm)],[c_0_416, c_0_650])).
cnf(c_0_659,plain,(node15(s11)|m_main_v_turn(s12,c_1)), inference(sr,[status(thm)],[c_0_651, c_0_644])).
cnf(c_0_660,plain,(node15(s11)|m_main_v_turn(s11,X1)|~m_main_v_turn(s12,X1)), inference(spm,[status(thm)],[c_0_66, c_0_652])).
cnf(c_0_661,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s11,c_n1)|~trans(s11,X1)), inference(spm,[status(thm)],[c_0_653, c_0_15])).
cnf(c_0_662,plain,(~m_main_v_turn(s11,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_643])).
cnf(c_0_663,plain,(m_main_v_state1(X1,c_n1)|~node7(s10,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_613]), c_0_654]), c_0_655]), c_0_656]), c_0_657]), c_0_658])).
cnf(c_0_664,plain,(node15(s11)|~m_main_v_turn(s12,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_659])).
cnf(c_0_665,plain,(node15(s11)|~m_main_v_state1(s11,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_660, c_0_661]), c_0_625])]), c_0_662])).
cnf(c_0_666,plain,(m_main_v_state1(X1,c_n1)|~trans(s10,X1)), inference(spm,[status(thm)],[c_0_663, c_0_7])).
cnf(c_0_667,plain,(node10(s11)|node9(s11)|node8(s11)|m_main_v_state2(s12,c_c2)|~node15(s11)), inference(spm,[status(thm)],[c_0_467, c_0_625])).
cnf(c_0_668,plain,(node15(s11)|~node12(s12)), inference(spm,[status(thm)],[c_0_664, c_0_35])).
cnf(c_0_669,plain,(~node10(s11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_634]), c_0_635])])).
cnf(c_0_670,plain,(~node9(s11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_634]), c_0_635])])).
cnf(c_0_671,plain,(~node8(s11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_634]), c_0_635])])).
cnf(c_0_672,plain,(m_main_v_state1(s12,c_t1)|~node15(s11)), inference(spm,[status(thm)],[c_0_474, c_0_625])).
cnf(c_0_673,plain,(node15(s11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_665, c_0_666]), c_0_585])])).
cnf(c_0_674,plain,(m_main_v_state2(s12,c_c2)|~node12(s12)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_667, c_0_668]), c_0_669]), c_0_670]), c_0_671])).
cnf(c_0_675,plain,(~node15(s11)|~node11(s12)), inference(spm,[status(thm)],[c_0_165, c_0_672])).
cnf(c_0_676,plain,(m_main_v_state2(s12,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_667, c_0_673])]), c_0_669]), c_0_670]), c_0_671])).
cnf(c_0_677,plain,(~node15(s11)|~node9(s12)), inference(spm,[status(thm)],[c_0_167, c_0_672])).
cnf(c_0_678,plain,(~node12(s12)), inference(spm,[status(thm)],[c_0_345, c_0_674])).
cnf(c_0_679,plain,(~node11(s12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_675, c_0_673])])).
cnf(c_0_680,plain,(~node10(s12)), inference(spm,[status(thm)],[c_0_339, c_0_676])).
cnf(c_0_681,plain,(~node9(s12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_677, c_0_673])])).
cnf(c_0_682,plain,(~node8(s12)), inference(spm,[status(thm)],[c_0_343, c_0_676])).
cnf(c_0_683,plain,(m_main_v_state2(X1,c_n2)|~node14(s12,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_676]), c_0_678]), c_0_679]), c_0_680]), c_0_681]), c_0_682])).
cnf(c_0_684,plain,(m_main_v_turn(s12,c_2)|~node15(s11)), inference(spm,[status(thm)],[c_0_470, c_0_625])).
cnf(c_0_685,plain,(m_main_v_state2(X1,c_n2)|~trans(s12,X1)), inference(spm,[status(thm)],[c_0_683, c_0_29])).
cnf(c_0_686,plain,(trans(s12,s13)), inference(spm,[status(thm)],[c_0_8, c_0_74])).
cnf(c_0_687,plain,(~node15(s11)|~node5(s12)), inference(spm,[status(thm)],[c_0_354, c_0_684])).
cnf(c_0_688,plain,(~node15(s11)|~node3(s12)), inference(spm,[status(thm)],[c_0_201, c_0_672])).
cnf(c_0_689,plain,(~node15(s11)|~node2(s12)), inference(spm,[status(thm)],[c_0_202, c_0_672])).
cnf(c_0_690,plain,(~node15(s11)|~node1(s12)), inference(spm,[status(thm)],[c_0_203, c_0_672])).
cnf(c_0_691,plain,(~node15(s11)|~m_main_v_state1(s12,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_672])).
cnf(c_0_692,plain,(~m_main_v_state2(X1,c_t2)|~trans(s12,X1)), inference(spm,[status(thm)],[c_0_0, c_0_685])).
cnf(c_0_693,plain,(node6(s12,s13)|node5(s12)|node4(s12)|node3(s12)|node2(s12)|node1(s12)|m_main_v_state1(s12,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_686])).
cnf(c_0_694,plain,(~node5(s12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_687, c_0_673])])).
cnf(c_0_695,plain,(~node4(s12)), inference(spm,[status(thm)],[c_0_361, c_0_676])).
cnf(c_0_696,plain,(~node3(s12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_688, c_0_673])])).
cnf(c_0_697,plain,(~node2(s12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_689, c_0_673])])).
cnf(c_0_698,plain,(~node1(s12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_690, c_0_673])])).
cnf(c_0_699,plain,(~m_main_v_state1(s12,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_691, c_0_673])])).
cnf(c_0_700,plain,(node17(s12,s13)|node16(s12)|node15(s12)), inference(spm,[status(thm)],[c_0_152, c_0_686])).
cnf(c_0_701,plain,(trans(s13,s14)), inference(spm,[status(thm)],[c_0_8, c_0_75])).
cnf(c_0_702,plain,(~node15(X1)|~trans(s12,X1)), inference(spm,[status(thm)],[c_0_692, c_0_17])).
cnf(c_0_703,plain,(node6(s12,s13)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_693, c_0_694]), c_0_695]), c_0_696]), c_0_697]), c_0_698]), c_0_699])).
cnf(c_0_704,plain,(node16(s12)|node15(s12)|m_main_v_turn(s13,X1)|~m_main_v_turn(s12,X1)), inference(spm,[status(thm)],[c_0_24, c_0_700])).
cnf(c_0_705,plain,(node17(s13,s14)|node16(s13)|node15(s13)), inference(spm,[status(thm)],[c_0_152, c_0_701])).
cnf(c_0_706,plain,(~node15(s13)), inference(spm,[status(thm)],[c_0_702, c_0_686])).
cnf(c_0_707,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s12,X1)), inference(spm,[status(thm)],[c_0_63, c_0_685])).
cnf(c_0_708,plain,(m_main_v_state1(s13,X1)|~m_main_v_state1(s12,X1)), inference(spm,[status(thm)],[c_0_23, c_0_703])).
cnf(c_0_709,plain,(m_main_v_state1(s12,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_672, c_0_673])])).
cnf(c_0_710,plain,(node16(s12)|node15(s12)|m_main_v_turn(s13,c_2)|~node15(s11)), inference(spm,[status(thm)],[c_0_704, c_0_684])).
cnf(c_0_711,plain,(~node15(s12)|~node15(s11)), inference(spm,[status(thm)],[c_0_189, c_0_672])).
cnf(c_0_712,plain,(node17(s13,s14)|node16(s13)), inference(sr,[status(thm)],[c_0_705, c_0_706])).
cnf(c_0_713,plain,(m_main_v_turn(X1,c_1)|~node18(s13,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_708]), c_0_686]), c_0_709])]), c_0_706])).
cnf(c_0_714,plain,(node16(s12)|node15(s12)|m_main_v_turn(s13,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_710, c_0_673])])).
cnf(c_0_715,plain,(~node15(s12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_711, c_0_673])])).
cnf(c_0_716,plain,(node16(s13)|m_main_v_turn(s13,X1)|~m_main_v_turn(s14,X1)), inference(spm,[status(thm)],[c_0_66, c_0_712])).
cnf(c_0_717,plain,(m_main_v_turn(X1,c_1)|~trans(s13,X1)), inference(spm,[status(thm)],[c_0_713, c_0_15])).
cnf(c_0_718,plain,(node16(s12)|m_main_v_turn(s13,c_2)), inference(sr,[status(thm)],[c_0_714, c_0_715])).
cnf(c_0_719,plain,(~node16(s12)), inference(spm,[status(thm)],[c_0_392, c_0_676])).
cnf(c_0_720,plain,(trans(s14,s15)), inference(spm,[status(thm)],[c_0_8, c_0_76])).
cnf(c_0_721,plain,(m_main_v_state2(s14,c_t2)|~node16(s13)), inference(spm,[status(thm)],[c_0_400, c_0_701])).
cnf(c_0_722,plain,(node16(s13)|m_main_v_turn(s13,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_716, c_0_717]), c_0_701])])).
cnf(c_0_723,plain,(m_main_v_turn(s13,c_2)), inference(sr,[status(thm)],[c_0_718, c_0_719])).
cnf(c_0_724,plain,(~m_main_v_state2(X1,c_c2)|~trans(s12,X1)), inference(spm,[status(thm)],[c_0_4, c_0_685])).
cnf(c_0_725,plain,(node17(s14,s15)|node16(s14)|node15(s14)), inference(spm,[status(thm)],[c_0_152, c_0_720])).
cnf(c_0_726,plain,(~node16(s14)|~node16(s13)), inference(spm,[status(thm)],[c_0_188, c_0_721])).
cnf(c_0_727,plain,(node16(s13)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_722]), c_0_723])])).
cnf(c_0_728,plain,(~node3(X1)|~trans(s12,X1)), inference(spm,[status(thm)],[c_0_724, c_0_12])).
cnf(c_0_729,plain,(~node1(X1)|~trans(s12,X1)), inference(spm,[status(thm)],[c_0_692, c_0_5])).
cnf(c_0_730,plain,(node16(s14)|node15(s14)|m_main_v_turn(s15,X1)|~m_main_v_turn(s14,X1)), inference(spm,[status(thm)],[c_0_24, c_0_725])).
cnf(c_0_731,plain,(~node16(s14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_726, c_0_727])])).
cnf(c_0_732,plain,(node3(s13)|node2(s13)|node1(s13)|m_main_v_state1(s14,c_c1)|~node16(s13)), inference(spm,[status(thm)],[c_0_401, c_0_701])).
cnf(c_0_733,plain,(~node3(s13)), inference(spm,[status(thm)],[c_0_728, c_0_686])).
cnf(c_0_734,plain,(~node2(s13)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_708]), c_0_709])])).
cnf(c_0_735,plain,(~node1(s13)), inference(spm,[status(thm)],[c_0_729, c_0_686])).
cnf(c_0_736,plain,(node15(s14)|m_main_v_turn(s15,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_717]), c_0_701])]), c_0_731])).
cnf(c_0_737,plain,(m_main_v_state1(s14,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_732, c_0_727])]), c_0_733]), c_0_734]), c_0_735])).
cnf(c_0_738,plain,(~node16(s13)|~node8(s14)), inference(spm,[status(thm)],[c_0_168, c_0_721])).
cnf(c_0_739,plain,(node15(s14)|~m_main_v_turn(s15,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_736])).
cnf(c_0_740,plain,(m_main_v_turn(s15,c_2)|~node11(s14)), inference(spm,[status(thm)],[c_0_251, c_0_720])).
cnf(c_0_741,plain,(~node15(s14)), inference(spm,[status(thm)],[c_0_229, c_0_737])).
cnf(c_0_742,plain,(~node16(s13)|~node10(s14)), inference(spm,[status(thm)],[c_0_257, c_0_721])).
cnf(c_0_743,plain,(~node16(s13)|~node9(s14)), inference(spm,[status(thm)],[c_0_258, c_0_721])).
cnf(c_0_744,plain,(~node16(s13)|~m_main_v_state2(s14,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_721])).
cnf(c_0_745,plain,(node13(s14,s15)|node12(s14)|node11(s14)|node10(s14)|node9(s14)|node8(s14)|m_main_v_state2(s14,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_720])).
cnf(c_0_746,plain,(~node8(s14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_738, c_0_727])])).
cnf(c_0_747,plain,(~node12(s14)), inference(spm,[status(thm)],[c_0_428, c_0_737])).
cnf(c_0_748,plain,(~node11(s14)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_739, c_0_740]), c_0_741])).
cnf(c_0_749,plain,(~node10(s14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_742, c_0_727])])).
cnf(c_0_750,plain,(~node9(s14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_743, c_0_727])])).
cnf(c_0_751,plain,(~m_main_v_state2(s14,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_744, c_0_727])])).
cnf(c_0_752,plain,(node13(s14,s15)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_745, c_0_746]), c_0_747]), c_0_748]), c_0_749]), c_0_750]), c_0_751])).
cnf(c_0_753,plain,(trans(s15,s16)), inference(spm,[status(thm)],[c_0_8, c_0_77])).
cnf(c_0_754,plain,(node15(s14)|~node12(s15)), inference(spm,[status(thm)],[c_0_739, c_0_35])).
cnf(c_0_755,plain,(m_main_v_state2(s15,X1)|~m_main_v_state2(s14,X1)), inference(spm,[status(thm)],[c_0_43, c_0_752])).
cnf(c_0_756,plain,(m_main_v_state2(s14,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_721, c_0_727])])).
cnf(c_0_757,plain,(node13(s15,s16)|node12(s15)|node11(s15)|node10(s15)|node9(s15)|node8(s15)|m_main_v_state2(s15,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_753])).
cnf(c_0_758,plain,(~node12(s15)), inference(spm,[status(thm)],[c_0_741, c_0_754])).
cnf(c_0_759,plain,(~node10(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_755]), c_0_756])])).
cnf(c_0_760,plain,(~node9(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_755]), c_0_756])])).
cnf(c_0_761,plain,(~node8(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_755]), c_0_756])])).
cnf(c_0_762,plain,(~m_main_v_state2(s15,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_755]), c_0_756])])).
cnf(c_0_763,plain,(node13(s15,s16)|node11(s15)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_757, c_0_758]), c_0_759]), c_0_760]), c_0_761]), c_0_762])).
cnf(c_0_764,plain,(m_main_v_state2(X1,c_t2)|~node14(s13,X1)|~m_main_v_state2(s13,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_708]), c_0_709])])).
cnf(c_0_765,plain,(node17(s15,s16)|node16(s15)|node15(s15)), inference(spm,[status(thm)],[c_0_152, c_0_753])).
cnf(c_0_766,plain,(~node16(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_755]), c_0_756])])).
cnf(c_0_767,plain,(~node16(s13)|~node4(s14)), inference(spm,[status(thm)],[c_0_200, c_0_721])).
cnf(c_0_768,plain,(~node16(s13)|~node2(s14)), inference(spm,[status(thm)],[c_0_239, c_0_721])).
cnf(c_0_769,plain,(m_main_v_turn(s15,c_2)|~node1(s14)), inference(spm,[status(thm)],[c_0_232, c_0_720])).
cnf(c_0_770,plain,(node11(s15)|m_main_v_state2(s16,X1)|~m_main_v_state2(s15,X1)), inference(spm,[status(thm)],[c_0_43, c_0_763])).
cnf(c_0_771,plain,(m_main_v_state2(X1,c_t2)|~m_main_v_state2(s13,c_n2)|~trans(s13,X1)), inference(spm,[status(thm)],[c_0_764, c_0_29])).
cnf(c_0_772,plain,(node17(s15,s16)|node15(s15)), inference(sr,[status(thm)],[c_0_765, c_0_766])).
cnf(c_0_773,plain,(m_main_v_turn(X1,c_2)|~node18(s15,X1)|~m_main_v_state1(s15,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_755]), c_0_756])])).
cnf(c_0_774,plain,(m_main_v_turn(s15,c_1)), inference(sr,[status(thm)],[c_0_736, c_0_741])).
cnf(c_0_775,plain,(~node5(s14)), inference(spm,[status(thm)],[c_0_233, c_0_737])).
cnf(c_0_776,plain,(~node4(s14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_767, c_0_727])])).
cnf(c_0_777,plain,(~node3(s14)), inference(spm,[status(thm)],[c_0_415, c_0_737])).
cnf(c_0_778,plain,(~node2(s14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_768, c_0_727])])).
cnf(c_0_779,plain,(~node1(s14)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_739, c_0_769]), c_0_741])).
cnf(c_0_780,plain,(node16(s15)|node15(s15)|m_main_v_turn(s16,X1)|~m_main_v_turn(s15,X1)), inference(spm,[status(thm)],[c_0_24, c_0_765])).
cnf(c_0_781,plain,(node11(s15)|m_main_v_state2(s16,X1)|~m_main_v_state2(s14,X1)), inference(spm,[status(thm)],[c_0_770, c_0_755])).
cnf(c_0_782,plain,(m_main_v_state2(X1,c_t2)|~trans(s13,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_771, c_0_37]), c_0_727])])).
cnf(c_0_783,plain,(m_main_v_state2(s14,X1)|~m_main_v_state2(s15,X1)), inference(spm,[status(thm)],[c_0_47, c_0_752])).
cnf(c_0_784,plain,(node15(s15)|m_main_v_turn(s15,X1)|~m_main_v_turn(s16,X1)), inference(spm,[status(thm)],[c_0_66, c_0_772])).
cnf(c_0_785,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s15,c_n1)|~trans(s15,X1)), inference(spm,[status(thm)],[c_0_773, c_0_15])).
cnf(c_0_786,plain,(~m_main_v_turn(s15,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_774])).
cnf(c_0_787,plain,(m_main_v_state1(X1,c_n1)|~node7(s14,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_737]), c_0_775]), c_0_776]), c_0_777]), c_0_778]), c_0_779])).
cnf(c_0_788,plain,(node16(s15)|node15(s15)|m_main_v_turn(s16,c_1)), inference(spm,[status(thm)],[c_0_780, c_0_774])).
cnf(c_0_789,plain,(m_main_v_state1(s16,c_t1)|~node11(s15)), inference(spm,[status(thm)],[c_0_316, c_0_753])).
cnf(c_0_790,plain,(node11(s15)|m_main_v_state2(s15,X1)|~m_main_v_state2(s16,X1)), inference(spm,[status(thm)],[c_0_47, c_0_763])).
cnf(c_0_791,plain,(node11(s15)|m_main_v_state2(s16,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_781, c_0_782]), c_0_701])])).
cnf(c_0_792,plain,(node6(s15,s16)|node5(s15)|node4(s15)|node3(s15)|node2(s15)|node1(s15)|m_main_v_state1(s15,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_753])).
cnf(c_0_793,plain,(~node4(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_755]), c_0_756])])).
cnf(c_0_794,plain,(~node3(s15)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_783, c_0_12]), c_0_751])).
cnf(c_0_795,plain,(~node2(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_755]), c_0_756])])).
cnf(c_0_796,plain,(node15(s15)|~m_main_v_state1(s15,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_784, c_0_785]), c_0_753])]), c_0_786])).
cnf(c_0_797,plain,(m_main_v_state1(X1,c_n1)|~trans(s14,X1)), inference(spm,[status(thm)],[c_0_787, c_0_7])).
cnf(c_0_798,plain,(node15(s15)|m_main_v_turn(s16,c_1)), inference(sr,[status(thm)],[c_0_788, c_0_766])).
cnf(c_0_799,plain,(trans(s16,s17)), inference(spm,[status(thm)],[c_0_8, c_0_78])).
cnf(c_0_800,plain,(~node11(s15)|~node3(s16)), inference(spm,[status(thm)],[c_0_201, c_0_789])).
cnf(c_0_801,plain,(node11(s15)|~node3(s16)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_790, c_0_12]), c_0_762])).
cnf(c_0_802,plain,(~node11(s15)|~node2(s16)), inference(spm,[status(thm)],[c_0_202, c_0_789])).
cnf(c_0_803,plain,(node11(s15)|~node2(s16)), inference(spm,[status(thm)],[c_0_239, c_0_791])).
cnf(c_0_804,plain,(node6(s15,s16)|node5(s15)|node1(s15)|m_main_v_state1(s15,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_792, c_0_793]), c_0_794]), c_0_795])).
cnf(c_0_805,plain,(~node15(X1)|~node5(X1)), inference(spm,[status(thm)],[c_0_189, c_0_10])).
cnf(c_0_806,plain,(node15(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_796, c_0_797]), c_0_720])])).
cnf(c_0_807,plain,(node15(s15)|~m_main_v_turn(s16,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_798])).
cnf(c_0_808,plain,(m_main_v_turn(s16,c_2)|~node11(s15)), inference(spm,[status(thm)],[c_0_251, c_0_753])).
cnf(c_0_809,plain,(node6(s16,s17)|node5(s16)|node4(s16)|node3(s16)|node2(s16)|node1(s16)|m_main_v_state1(s16,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_799])).
cnf(c_0_810,plain,(~node3(s16)), inference(spm,[status(thm)],[c_0_800, c_0_801])).
cnf(c_0_811,plain,(~node2(s16)), inference(spm,[status(thm)],[c_0_802, c_0_803])).
cnf(c_0_812,plain,(node5(s15)|node1(s15)|m_main_v_state1(s15,c_c1)|m_main_v_state1(s16,X1)|~m_main_v_state1(s15,X1)), inference(spm,[status(thm)],[c_0_23, c_0_804])).
cnf(c_0_813,plain,(~node5(s15)), inference(spm,[status(thm)],[c_0_805, c_0_806])).
cnf(c_0_814,plain,(m_main_v_state1(s16,c_t1)|~node15(s15)), inference(spm,[status(thm)],[c_0_474, c_0_753])).
cnf(c_0_815,plain,(node10(s15)|node9(s15)|node8(s15)|m_main_v_state2(s16,c_c2)|~node15(s15)), inference(spm,[status(thm)],[c_0_467, c_0_753])).
cnf(c_0_816,plain,(node15(s15)|~node11(s15)), inference(spm,[status(thm)],[c_0_807, c_0_808])).
cnf(c_0_817,plain,(node6(s16,s17)|node5(s16)|node4(s16)|node1(s16)|m_main_v_state1(s16,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_809, c_0_810]), c_0_811])).
cnf(c_0_818,plain,(node1(s15)|m_main_v_state1(s15,c_c1)|m_main_v_state1(s16,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_812, c_0_797]), c_0_720])]), c_0_813])).
cnf(c_0_819,plain,(m_main_v_state1(s16,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_814, c_0_806])])).
cnf(c_0_820,plain,(m_main_v_turn(s16,c_2)|~node15(s15)), inference(spm,[status(thm)],[c_0_470, c_0_753])).
cnf(c_0_821,plain,(node15(s15)|~node12(s16)), inference(spm,[status(thm)],[c_0_807, c_0_35])).
cnf(c_0_822,plain,(m_main_v_state2(s16,c_c2)|~node11(s15)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_815, c_0_816]), c_0_759]), c_0_760]), c_0_761])).
cnf(c_0_823,plain,(node5(s16)|node4(s16)|node1(s16)|m_main_v_state1(s16,c_c1)|m_main_v_state1(s17,X1)|~m_main_v_state1(s16,X1)), inference(spm,[status(thm)],[c_0_23, c_0_817])).
cnf(c_0_824,plain,(m_main_v_state1(s16,c_t1)|~node1(s15)), inference(spm,[status(thm)],[c_0_332, c_0_753])).
cnf(c_0_825,plain,(node1(s15)|m_main_v_state1(s15,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_818]), c_0_819])])).
cnf(c_0_826,plain,(~node15(s15)|~node5(s16)), inference(spm,[status(thm)],[c_0_354, c_0_820])).
cnf(c_0_827,plain,(m_main_v_state2(s16,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_815, c_0_806])]), c_0_759]), c_0_760]), c_0_761])).
cnf(c_0_828,plain,(~node15(s15)|~node1(s16)), inference(spm,[status(thm)],[c_0_203, c_0_814])).
cnf(c_0_829,plain,(m_main_v_state2(s16,c_c2)|~node12(s16)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_815, c_0_821]), c_0_759]), c_0_760]), c_0_761])).
cnf(c_0_830,plain,(~node11(s15)|~node10(s16)), inference(spm,[status(thm)],[c_0_339, c_0_822])).
cnf(c_0_831,plain,(node11(s15)|~node10(s16)), inference(spm,[status(thm)],[c_0_257, c_0_791])).
cnf(c_0_832,plain,(~node11(s15)|~node9(s16)), inference(spm,[status(thm)],[c_0_167, c_0_789])).
cnf(c_0_833,plain,(node11(s15)|~node9(s16)), inference(spm,[status(thm)],[c_0_258, c_0_791])).
cnf(c_0_834,plain,(~node11(s15)|~node8(s16)), inference(spm,[status(thm)],[c_0_343, c_0_822])).
cnf(c_0_835,plain,(node11(s15)|~node8(s16)), inference(spm,[status(thm)],[c_0_168, c_0_791])).
cnf(c_0_836,plain,(~node15(s15)|~node11(s16)), inference(spm,[status(thm)],[c_0_165, c_0_814])).
cnf(c_0_837,plain,(node17(s16,s17)|node16(s16)|node15(s16)), inference(spm,[status(thm)],[c_0_152, c_0_799])).
cnf(c_0_838,plain,(node5(s16)|node4(s16)|node1(s16)|m_main_v_state1(s16,c_c1)|m_main_v_state1(s17,c_t1)|~node1(s15)), inference(spm,[status(thm)],[c_0_823, c_0_824])).
cnf(c_0_839,plain,(node1(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_825]), c_0_806])])).
cnf(c_0_840,plain,(~node5(s16)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_826, c_0_806])])).
cnf(c_0_841,plain,(~node4(s16)), inference(spm,[status(thm)],[c_0_361, c_0_827])).
cnf(c_0_842,plain,(~node1(s16)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_828, c_0_806])])).
cnf(c_0_843,plain,(~node12(s16)), inference(spm,[status(thm)],[c_0_345, c_0_829])).
cnf(c_0_844,plain,(~node10(s16)), inference(spm,[status(thm)],[c_0_830, c_0_831])).
cnf(c_0_845,plain,(~node9(s16)), inference(spm,[status(thm)],[c_0_832, c_0_833])).
cnf(c_0_846,plain,(~node8(s16)), inference(spm,[status(thm)],[c_0_834, c_0_835])).
cnf(c_0_847,plain,(~node11(s16)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_836, c_0_806])])).
cnf(c_0_848,plain,(~node15(s15)|~m_main_v_state1(s16,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_814])).
cnf(c_0_849,plain,(node16(s16)|node15(s16)|m_main_v_turn(s17,X1)|~m_main_v_turn(s16,X1)), inference(spm,[status(thm)],[c_0_24, c_0_837])).
cnf(c_0_850,plain,(~node15(s16)|~node15(s15)), inference(spm,[status(thm)],[c_0_189, c_0_814])).
cnf(c_0_851,plain,(m_main_v_state1(s16,c_c1)|m_main_v_state1(s17,c_t1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_838, c_0_839])]), c_0_840]), c_0_841]), c_0_842])).
cnf(c_0_852,plain,(m_main_v_state2(X1,c_n2)|~node14(s16,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_827]), c_0_843]), c_0_844]), c_0_845]), c_0_846]), c_0_847])).
cnf(c_0_853,plain,(~m_main_v_state1(s16,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_848, c_0_806])])).
cnf(c_0_854,plain,(trans(s17,s18)), inference(spm,[status(thm)],[c_0_8, c_0_79])).
cnf(c_0_855,plain,(node16(s16)|node15(s16)|m_main_v_turn(s17,c_2)|~m_main_v_state1(s15,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_849, c_0_785]), c_0_753])])).
cnf(c_0_856,plain,(~node15(s16)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_850, c_0_806])])).
cnf(c_0_857,plain,(m_main_v_state1(s16,c_c1)|~node15(s17)), inference(spm,[status(thm)],[c_0_189, c_0_851])).
cnf(c_0_858,plain,(m_main_v_state2(X1,c_n2)|~trans(s16,X1)), inference(spm,[status(thm)],[c_0_852, c_0_29])).
cnf(c_0_859,plain,(node6(s16,s17)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_817, c_0_842]), c_0_840]), c_0_841]), c_0_853])).
cnf(c_0_860,plain,(node17(s17,s18)|node16(s17)|node15(s17)), inference(spm,[status(thm)],[c_0_152, c_0_854])).
cnf(c_0_861,plain,(node16(s16)|m_main_v_turn(s17,c_2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_855, c_0_797]), c_0_720])]), c_0_856])).
cnf(c_0_862,plain,(~node16(s16)), inference(spm,[status(thm)],[c_0_392, c_0_827])).
cnf(c_0_863,plain,(~node15(s17)), inference(spm,[status(thm)],[c_0_853, c_0_857])).
cnf(c_0_864,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s16,X1)), inference(spm,[status(thm)],[c_0_63, c_0_858])).
cnf(c_0_865,plain,(m_main_v_state1(s17,X1)|~m_main_v_state1(s16,X1)), inference(spm,[status(thm)],[c_0_23, c_0_859])).
cnf(c_0_866,plain,(node16(s17)|node15(s17)|m_main_v_turn(s18,X1)|~m_main_v_turn(s17,X1)), inference(spm,[status(thm)],[c_0_24, c_0_860])).
cnf(c_0_867,plain,(m_main_v_turn(s17,c_2)), inference(sr,[status(thm)],[c_0_861, c_0_862])).
cnf(c_0_868,plain,(node17(s17,s18)|node16(s17)), inference(sr,[status(thm)],[c_0_860, c_0_863])).
cnf(c_0_869,plain,(m_main_v_turn(X1,c_1)|~node18(s17,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_864, c_0_865]), c_0_799]), c_0_819])]), c_0_863])).
cnf(c_0_870,plain,(node16(s17)|node15(s17)|m_main_v_turn(s18,c_2)), inference(spm,[status(thm)],[c_0_866, c_0_867])).
cnf(c_0_871,plain,(node16(s17)|m_main_v_turn(s17,X1)|~m_main_v_turn(s18,X1)), inference(spm,[status(thm)],[c_0_66, c_0_868])).
cnf(c_0_872,plain,(m_main_v_turn(X1,c_1)|~trans(s17,X1)), inference(spm,[status(thm)],[c_0_869, c_0_15])).
cnf(c_0_873,plain,(node16(s17)|m_main_v_turn(s18,c_2)), inference(sr,[status(thm)],[c_0_870, c_0_863])).
cnf(c_0_874,plain,(m_main_v_state1(s16,c_c1)|~node3(s17)), inference(spm,[status(thm)],[c_0_201, c_0_851])).
cnf(c_0_875,plain,(m_main_v_state1(s16,c_c1)|~node2(s17)), inference(spm,[status(thm)],[c_0_202, c_0_851])).
cnf(c_0_876,plain,(m_main_v_state1(s16,c_c1)|~node1(s17)), inference(spm,[status(thm)],[c_0_203, c_0_851])).
cnf(c_0_877,plain,(node16(s17)|m_main_v_turn(s17,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_871, c_0_872]), c_0_854])])).
cnf(c_0_878,plain,(node3(s17)|node2(s17)|node1(s17)|m_main_v_state1(s18,c_c1)|~node16(s17)), inference(spm,[status(thm)],[c_0_401, c_0_854])).
cnf(c_0_879,plain,(node16(s17)|~node5(s18)), inference(spm,[status(thm)],[c_0_354, c_0_873])).
cnf(c_0_880,plain,(~node3(s17)), inference(spm,[status(thm)],[c_0_853, c_0_874])).
cnf(c_0_881,plain,(~node2(s17)), inference(spm,[status(thm)],[c_0_853, c_0_875])).
cnf(c_0_882,plain,(~node1(s17)), inference(spm,[status(thm)],[c_0_853, c_0_876])).
cnf(c_0_883,plain,(m_main_v_state2(s18,c_t2)|~node16(s17)), inference(spm,[status(thm)],[c_0_400, c_0_854])).
cnf(c_0_884,plain,(node16(s17)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_877]), c_0_867])])).
cnf(c_0_885,plain,(m_main_v_state1(s18,c_c1)|~node5(s18)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_878, c_0_879]), c_0_880]), c_0_881]), c_0_882])).
cnf(c_0_886,plain,(~node16(s17)|~node4(s18)), inference(spm,[status(thm)],[c_0_200, c_0_883])).
cnf(c_0_887,plain,(m_main_v_state1(s18,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_878, c_0_884])]), c_0_880]), c_0_881]), c_0_882])).
cnf(c_0_888,plain,(~node16(s17)|~node2(s18)), inference(spm,[status(thm)],[c_0_239, c_0_883])).
cnf(c_0_889,plain,(~node5(s18)), inference(spm,[status(thm)],[c_0_233, c_0_885])).
cnf(c_0_890,plain,(~node4(s18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_886, c_0_884])])).
cnf(c_0_891,plain,(~node3(s18)), inference(spm,[status(thm)],[c_0_415, c_0_887])).
cnf(c_0_892,plain,(~node2(s18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_888, c_0_884])])).
cnf(c_0_893,plain,(~node1(s18)), inference(spm,[status(thm)],[c_0_416, c_0_887])).
cnf(c_0_894,plain,(m_main_v_state1(X1,c_n1)|~node7(s18,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_887]), c_0_889]), c_0_890]), c_0_891]), c_0_892]), c_0_893])).
cnf(c_0_895,plain,(trans(s18,s19)), inference(spm,[status(thm)],[c_0_8, c_0_80])).
cnf(c_0_896,plain,(m_main_v_state1(X1,c_n1)|~trans(s18,X1)), inference(spm,[status(thm)],[c_0_894, c_0_7])).
cnf(c_0_897,plain,(~node16(s17)|~node8(s18)), inference(spm,[status(thm)],[c_0_168, c_0_883])).
cnf(c_0_898,plain,(~node16(s17)|~node10(s18)), inference(spm,[status(thm)],[c_0_257, c_0_883])).
cnf(c_0_899,plain,(~node16(s17)|~node9(s18)), inference(spm,[status(thm)],[c_0_258, c_0_883])).
cnf(c_0_900,plain,(~node16(s17)|~m_main_v_state2(s18,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_883])).
cnf(c_0_901,plain,(node17(s18,s19)|node16(s18)|node15(s18)), inference(spm,[status(thm)],[c_0_152, c_0_895])).
cnf(c_0_902,plain,(~m_main_v_state1(X1,c_t1)|~trans(s18,X1)), inference(spm,[status(thm)],[c_0_2, c_0_896])).
cnf(c_0_903,plain,(node13(s18,s19)|node12(s18)|node11(s18)|node10(s18)|node9(s18)|node8(s18)|m_main_v_state2(s18,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_895])).
cnf(c_0_904,plain,(~node8(s18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_897, c_0_884])])).
cnf(c_0_905,plain,(~node12(s18)), inference(spm,[status(thm)],[c_0_428, c_0_887])).
cnf(c_0_906,plain,(~node11(s18)), inference(spm,[status(thm)],[c_0_429, c_0_887])).
cnf(c_0_907,plain,(~node10(s18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_898, c_0_884])])).
cnf(c_0_908,plain,(~node9(s18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_899, c_0_884])])).
cnf(c_0_909,plain,(~m_main_v_state2(s18,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_900, c_0_884])])).
cnf(c_0_910,plain,(node16(s18)|node15(s18)|m_main_v_turn(s19,X1)|~m_main_v_turn(s18,X1)), inference(spm,[status(thm)],[c_0_24, c_0_901])).
cnf(c_0_911,plain,(~node16(s18)|~node16(s17)), inference(spm,[status(thm)],[c_0_188, c_0_883])).
cnf(c_0_912,plain,(trans(s19,s20)), inference(spm,[status(thm)],[c_0_8, c_0_81])).
cnf(c_0_913,plain,(~node16(X1)|~trans(s18,X1)), inference(spm,[status(thm)],[c_0_902, c_0_16])).
cnf(c_0_914,plain,(node13(s18,s19)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_903, c_0_904]), c_0_905]), c_0_906]), c_0_907]), c_0_908]), c_0_909])).
cnf(c_0_915,plain,(node16(s18)|node15(s18)|m_main_v_turn(s19,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_910, c_0_872]), c_0_854])])).
cnf(c_0_916,plain,(~node16(s18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_911, c_0_884])])).
cnf(c_0_917,plain,(node17(s19,s20)|node16(s19)|node15(s19)), inference(spm,[status(thm)],[c_0_152, c_0_912])).
cnf(c_0_918,plain,(~node16(s19)), inference(spm,[status(thm)],[c_0_913, c_0_895])).
cnf(c_0_919,plain,(m_main_v_state2(s19,X1)|~m_main_v_state2(s18,X1)), inference(spm,[status(thm)],[c_0_43, c_0_914])).
cnf(c_0_920,plain,(m_main_v_state2(s18,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_883, c_0_884])])).
cnf(c_0_921,plain,(node15(s18)|m_main_v_turn(s19,c_1)), inference(sr,[status(thm)],[c_0_915, c_0_916])).
cnf(c_0_922,plain,(~node15(s18)), inference(spm,[status(thm)],[c_0_229, c_0_887])).
cnf(c_0_923,plain,(trans(s20,s21)), inference(spm,[status(thm)],[c_0_8, c_0_82])).
cnf(c_0_924,plain,(node17(s19,s20)|node15(s19)), inference(sr,[status(thm)],[c_0_917, c_0_918])).
cnf(c_0_925,plain,(m_main_v_turn(X1,c_2)|~node18(s19,X1)|~m_main_v_state1(s19,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_919]), c_0_920])])).
cnf(c_0_926,plain,(m_main_v_turn(s19,c_1)), inference(sr,[status(thm)],[c_0_921, c_0_922])).
cnf(c_0_927,plain,(node17(s20,s21)|node16(s20)|node15(s20)), inference(spm,[status(thm)],[c_0_152, c_0_923])).
cnf(c_0_928,plain,(node15(s19)|m_main_v_turn(s19,X1)|~m_main_v_turn(s20,X1)), inference(spm,[status(thm)],[c_0_66, c_0_924])).
cnf(c_0_929,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s19,c_n1)|~trans(s19,X1)), inference(spm,[status(thm)],[c_0_925, c_0_15])).
cnf(c_0_930,plain,(~m_main_v_turn(s19,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_926])).
cnf(c_0_931,plain,(~m_main_v_state1(X1,c_c1)|~trans(s18,X1)), inference(spm,[status(thm)],[c_0_13, c_0_896])).
cnf(c_0_932,plain,(node16(s20)|node15(s20)|m_main_v_turn(s21,X1)|~m_main_v_turn(s20,X1)), inference(spm,[status(thm)],[c_0_24, c_0_927])).
cnf(c_0_933,plain,(m_main_v_turn(s20,c_2)|~node15(s19)), inference(spm,[status(thm)],[c_0_470, c_0_912])).
cnf(c_0_934,plain,(node15(s19)|~m_main_v_state1(s19,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_928, c_0_929]), c_0_912])]), c_0_930])).
cnf(c_0_935,plain,(m_main_v_state1(s20,c_t1)|~node15(s19)), inference(spm,[status(thm)],[c_0_474, c_0_912])).
cnf(c_0_936,plain,(~node10(X1)|~trans(s18,X1)), inference(spm,[status(thm)],[c_0_931, c_0_32])).
cnf(c_0_937,plain,(~node8(X1)|~trans(s18,X1)), inference(spm,[status(thm)],[c_0_902, c_0_19])).
cnf(c_0_938,plain,(node16(s20)|node15(s20)|m_main_v_turn(s21,c_2)|~node15(s19)), inference(spm,[status(thm)],[c_0_932, c_0_933])).
cnf(c_0_939,plain,(node15(s19)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_934, c_0_896]), c_0_895])])).
cnf(c_0_940,plain,(~node15(s20)|~node15(s19)), inference(spm,[status(thm)],[c_0_189, c_0_935])).
cnf(c_0_941,plain,(node10(s19)|node9(s19)|node8(s19)|m_main_v_state2(s20,c_c2)|~node15(s19)), inference(spm,[status(thm)],[c_0_467, c_0_912])).
cnf(c_0_942,plain,(~node10(s19)), inference(spm,[status(thm)],[c_0_936, c_0_895])).
cnf(c_0_943,plain,(~node9(s19)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_919]), c_0_920])])).
cnf(c_0_944,plain,(~node8(s19)), inference(spm,[status(thm)],[c_0_937, c_0_895])).
cnf(c_0_945,plain,(node16(s20)|node15(s20)|m_main_v_turn(s21,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_938, c_0_939])])).
cnf(c_0_946,plain,(~node15(s20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_940, c_0_939])])).
cnf(c_0_947,plain,(~node15(s19)|~node5(s20)), inference(spm,[status(thm)],[c_0_354, c_0_933])).
cnf(c_0_948,plain,(m_main_v_state2(s20,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_941, c_0_939])]), c_0_942]), c_0_943]), c_0_944])).
cnf(c_0_949,plain,(~node15(s19)|~node3(s20)), inference(spm,[status(thm)],[c_0_201, c_0_935])).
cnf(c_0_950,plain,(~node15(s19)|~node2(s20)), inference(spm,[status(thm)],[c_0_202, c_0_935])).
cnf(c_0_951,plain,(~node15(s19)|~node1(s20)), inference(spm,[status(thm)],[c_0_203, c_0_935])).
cnf(c_0_952,plain,(~node15(s19)|~m_main_v_state1(s20,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_935])).
cnf(c_0_953,plain,(node16(s20)|m_main_v_turn(s21,c_2)), inference(sr,[status(thm)],[c_0_945, c_0_946])).
cnf(c_0_954,plain,(node6(s20,s21)|node5(s20)|node4(s20)|node3(s20)|node2(s20)|node1(s20)|m_main_v_state1(s20,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_923])).
cnf(c_0_955,plain,(~node5(s20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_947, c_0_939])])).
cnf(c_0_956,plain,(~node4(s20)), inference(spm,[status(thm)],[c_0_361, c_0_948])).
cnf(c_0_957,plain,(~node3(s20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_949, c_0_939])])).
cnf(c_0_958,plain,(~node2(s20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_950, c_0_939])])).
cnf(c_0_959,plain,(~node1(s20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_951, c_0_939])])).
cnf(c_0_960,plain,(~m_main_v_state1(s20,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_952, c_0_939])])).
cnf(c_0_961,plain,(node3(s20)|node2(s20)|node1(s20)|m_main_v_state1(s21,c_c1)|~node16(s20)), inference(spm,[status(thm)],[c_0_401, c_0_923])).
cnf(c_0_962,plain,(node16(s20)|~node5(s21)), inference(spm,[status(thm)],[c_0_354, c_0_953])).
cnf(c_0_963,plain,(node6(s20,s21)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_954, c_0_955]), c_0_956]), c_0_957]), c_0_958]), c_0_959]), c_0_960])).
cnf(c_0_964,plain,(~node15(s19)|~node11(s20)), inference(spm,[status(thm)],[c_0_165, c_0_935])).
cnf(c_0_965,plain,(~node15(s19)|~node9(s20)), inference(spm,[status(thm)],[c_0_167, c_0_935])).
cnf(c_0_966,plain,(trans(s21,s22)), inference(spm,[status(thm)],[c_0_8, c_0_83])).
cnf(c_0_967,plain,(m_main_v_state1(s21,c_c1)|~node5(s21)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_961, c_0_962]), c_0_957]), c_0_958]), c_0_959])).
cnf(c_0_968,plain,(m_main_v_state1(s21,X1)|~m_main_v_state1(s20,X1)), inference(spm,[status(thm)],[c_0_23, c_0_963])).
cnf(c_0_969,plain,(m_main_v_state1(s20,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_935, c_0_939])])).
cnf(c_0_970,plain,(~node12(s20)), inference(spm,[status(thm)],[c_0_345, c_0_948])).
cnf(c_0_971,plain,(~node11(s20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_964, c_0_939])])).
cnf(c_0_972,plain,(~node10(s20)), inference(spm,[status(thm)],[c_0_339, c_0_948])).
cnf(c_0_973,plain,(~node9(s20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_965, c_0_939])])).
cnf(c_0_974,plain,(~node8(s20)), inference(spm,[status(thm)],[c_0_343, c_0_948])).
cnf(c_0_975,plain,(node6(s21,s22)|node5(s21)|node4(s21)|node3(s21)|node2(s21)|node1(s21)|m_main_v_state1(s21,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_966])).
cnf(c_0_976,plain,(~node5(s21)), inference(spm,[status(thm)],[c_0_233, c_0_967])).
cnf(c_0_977,plain,(~node3(s21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_968]), c_0_969])])).
cnf(c_0_978,plain,(~node2(s21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_968]), c_0_969])])).
cnf(c_0_979,plain,(~node1(s21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_968]), c_0_969])])).
cnf(c_0_980,plain,(~m_main_v_state1(s21,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_968]), c_0_969])])).
cnf(c_0_981,plain,(m_main_v_state1(s20,X1)|~m_main_v_state1(s21,X1)), inference(spm,[status(thm)],[c_0_49, c_0_963])).
cnf(c_0_982,plain,(m_main_v_state2(X1,c_n2)|~node14(s20,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_948]), c_0_970]), c_0_971]), c_0_972]), c_0_973]), c_0_974])).
cnf(c_0_983,plain,(node6(s21,s22)|node4(s21)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_975, c_0_976]), c_0_977]), c_0_978]), c_0_979]), c_0_980])).
cnf(c_0_984,plain,(node13(s21,s22)|node12(s21)|node11(s21)|node10(s21)|node9(s21)|node8(s21)|m_main_v_state2(s21,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_966])).
cnf(c_0_985,plain,(~node11(s21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_968]), c_0_969])])).
cnf(c_0_986,plain,(~node10(s21)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_981, c_0_32]), c_0_960])).
cnf(c_0_987,plain,(~node9(s21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_968]), c_0_969])])).
cnf(c_0_988,plain,(m_main_v_state2(X1,c_n2)|~trans(s20,X1)), inference(spm,[status(thm)],[c_0_982, c_0_29])).
cnf(c_0_989,plain,(node4(s21)|m_main_v_state1(s22,X1)|~m_main_v_state1(s21,X1)), inference(spm,[status(thm)],[c_0_23, c_0_983])).
cnf(c_0_990,plain,(node13(s21,s22)|node12(s21)|node8(s21)|m_main_v_state2(s21,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_984, c_0_985]), c_0_986]), c_0_987])).
cnf(c_0_991,plain,(~m_main_v_state2(X1,c_t2)|~trans(s20,X1)), inference(spm,[status(thm)],[c_0_0, c_0_988])).
cnf(c_0_992,plain,(node4(s21)|m_main_v_state1(s22,X1)|~m_main_v_state1(s20,X1)), inference(spm,[status(thm)],[c_0_989, c_0_968])).
cnf(c_0_993,plain,(node12(s21)|node8(s21)|m_main_v_state2(s21,c_c2)|m_main_v_state2(s22,X1)|~m_main_v_state2(s21,X1)), inference(spm,[status(thm)],[c_0_43, c_0_990])).
cnf(c_0_994,plain,(~node12(X1)|~trans(s20,X1)), inference(spm,[status(thm)],[c_0_991, c_0_58])).
cnf(c_0_995,plain,(node4(s21)|m_main_v_state1(s22,c_t1)), inference(spm,[status(thm)],[c_0_992, c_0_969])).
cnf(c_0_996,plain,(node12(s21)|node8(s21)|m_main_v_state2(s21,c_c2)|m_main_v_state2(s22,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_993, c_0_988]), c_0_923])])).
cnf(c_0_997,plain,(~node12(s21)), inference(spm,[status(thm)],[c_0_994, c_0_923])).
cnf(c_0_998,plain,(m_main_v_state2(X1,c_t2)|~node14(s21,X1)|~m_main_v_state2(s21,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_968]), c_0_969])])).
cnf(c_0_999,plain,(node3(s21)|node2(s21)|node1(s21)|m_main_v_state1(s22,c_c1)|~node4(s21)), inference(spm,[status(thm)],[c_0_570, c_0_966])).
cnf(c_0_1000,plain,(node4(s21)|~node11(s22)), inference(spm,[status(thm)],[c_0_165, c_0_995])).
cnf(c_0_1001,plain,(node4(s21)|~node9(s22)), inference(spm,[status(thm)],[c_0_167, c_0_995])).
cnf(c_0_1002,plain,(node8(s21)|m_main_v_state2(s21,c_c2)|m_main_v_state2(s22,c_n2)), inference(sr,[status(thm)],[c_0_996, c_0_997])).
cnf(c_0_1003,plain,(m_main_v_state2(X1,c_t2)|~m_main_v_state2(s21,c_n2)|~trans(s21,X1)), inference(spm,[status(thm)],[c_0_998, c_0_29])).
cnf(c_0_1004,plain,(trans(s22,s23)), inference(spm,[status(thm)],[c_0_8, c_0_84])).
cnf(c_0_1005,plain,(m_main_v_state1(s22,c_c1)|~node11(s22)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1000]), c_0_977]), c_0_978]), c_0_979])).
cnf(c_0_1006,plain,(m_main_v_state1(s22,c_c1)|~node9(s22)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1001]), c_0_977]), c_0_978]), c_0_979])).
cnf(c_0_1007,plain,(node8(s21)|m_main_v_state2(s21,c_c2)|~m_main_v_state2(s22,c_t2)), inference(spm,[status(thm)],[c_0_0, c_0_1002])).
cnf(c_0_1008,plain,(m_main_v_state2(X1,c_t2)|~trans(s21,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_988]), c_0_923])])).
cnf(c_0_1009,plain,(~node16(X1)|~m_main_v_state2(s21,c_n2)|~trans(s21,X1)), inference(spm,[status(thm)],[c_0_188, c_0_1003])).
cnf(c_0_1010,plain,(node13(s22,s23)|node12(s22)|node11(s22)|node10(s22)|node9(s22)|node8(s22)|m_main_v_state2(s22,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1004])).
cnf(c_0_1011,plain,(~node11(s22)), inference(spm,[status(thm)],[c_0_429, c_0_1005])).
cnf(c_0_1012,plain,(~node9(s22)), inference(spm,[status(thm)],[c_0_587, c_0_1006])).
cnf(c_0_1013,plain,(~m_main_v_state2(X1,c_c2)|~trans(s20,X1)), inference(spm,[status(thm)],[c_0_4, c_0_988])).
cnf(c_0_1014,plain,(node8(s21)|m_main_v_state2(s21,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1007, c_0_1008]), c_0_966])])).
cnf(c_0_1015,plain,(node17(s21,s22)|node16(s21)|node15(s21)), inference(spm,[status(thm)],[c_0_152, c_0_966])).
cnf(c_0_1016,plain,(~node16(X1)|~trans(s21,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1009, c_0_988]), c_0_923])])).
cnf(c_0_1017,plain,(node13(s22,s23)|node12(s22)|node10(s22)|node8(s22)|m_main_v_state2(s22,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1010, c_0_1011]), c_0_1012])).
cnf(c_0_1018,plain,(~node10(X1)|~m_main_v_state2(s21,c_n2)|~trans(s21,X1)), inference(spm,[status(thm)],[c_0_257, c_0_1003])).
cnf(c_0_1019,plain,(~node8(X1)|~m_main_v_state2(s21,c_n2)|~trans(s21,X1)), inference(spm,[status(thm)],[c_0_168, c_0_1003])).
cnf(c_0_1020,plain,(node3(s21)|node2(s21)|node1(s21)|m_main_v_state1(s22,c_c1)|~node8(s21)), inference(spm,[status(thm)],[c_0_592, c_0_966])).
cnf(c_0_1021,plain,(node8(s21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1013, c_0_1014]), c_0_923])])).
cnf(c_0_1022,plain,(node16(s21)|node15(s21)|m_main_v_turn(s22,X1)|~m_main_v_turn(s21,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1015])).
cnf(c_0_1023,plain,(node17(s22,s23)|node16(s22)|node15(s22)), inference(spm,[status(thm)],[c_0_152, c_0_1004])).
cnf(c_0_1024,plain,(~node16(s22)), inference(spm,[status(thm)],[c_0_1016, c_0_966])).
cnf(c_0_1025,plain,(node12(s22)|node10(s22)|node8(s22)|m_main_v_state2(s22,c_c2)|m_main_v_state2(s23,X1)|~m_main_v_state2(s22,X1)), inference(spm,[status(thm)],[c_0_43, c_0_1017])).
cnf(c_0_1026,plain,(m_main_v_state2(s22,c_t2)|~node8(s21)), inference(spm,[status(thm)],[c_0_591, c_0_966])).
cnf(c_0_1027,plain,(~node10(X1)|~trans(s21,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1018, c_0_988]), c_0_923])])).
cnf(c_0_1028,plain,(~node8(X1)|~trans(s21,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1019, c_0_988]), c_0_923])])).
cnf(c_0_1029,plain,(m_main_v_state1(s22,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1020, c_0_1021])]), c_0_977]), c_0_978]), c_0_979])).
cnf(c_0_1030,plain,(node16(s21)|node16(s20)|node15(s21)|m_main_v_turn(s22,c_2)), inference(spm,[status(thm)],[c_0_1022, c_0_953])).
cnf(c_0_1031,plain,(~node16(s20)), inference(spm,[status(thm)],[c_0_392, c_0_948])).
cnf(c_0_1032,plain,(node17(s22,s23)|node15(s22)), inference(sr,[status(thm)],[c_0_1023, c_0_1024])).
cnf(c_0_1033,plain,(node15(s21)|m_main_v_turn(s22,c_1)|~node8(s21)), inference(spm,[status(thm)],[c_0_608, c_0_966])).
cnf(c_0_1034,plain,(~node15(s21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_968]), c_0_969])])).
cnf(c_0_1035,plain,(node12(s22)|node10(s22)|node8(s22)|m_main_v_state2(s22,c_c2)|m_main_v_state2(s23,c_t2)|~node8(s21)), inference(spm,[status(thm)],[c_0_1025, c_0_1026])).
cnf(c_0_1036,plain,(~node10(s22)), inference(spm,[status(thm)],[c_0_1027, c_0_966])).
cnf(c_0_1037,plain,(~node8(s22)), inference(spm,[status(thm)],[c_0_1028, c_0_966])).
cnf(c_0_1038,plain,(~node12(s22)), inference(spm,[status(thm)],[c_0_428, c_0_1029])).
cnf(c_0_1039,plain,(node16(s21)|node15(s21)|m_main_v_turn(s22,c_2)), inference(sr,[status(thm)],[c_0_1030, c_0_1031])).
cnf(c_0_1040,plain,(trans(s23,s24)), inference(spm,[status(thm)],[c_0_8, c_0_85])).
cnf(c_0_1041,plain,(node15(s22)|m_main_v_turn(s23,X1)|~m_main_v_turn(s22,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1032])).
cnf(c_0_1042,plain,(m_main_v_turn(s22,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1033, c_0_1021])]), c_0_1034])).
cnf(c_0_1043,plain,(~m_main_v_state2(s21,c_n2)|~m_main_v_state2(X1,c_c2)|~trans(s21,X1)), inference(spm,[status(thm)],[c_0_36, c_0_1003])).
cnf(c_0_1044,plain,(m_main_v_state2(s22,c_c2)|m_main_v_state2(s23,c_t2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1035, c_0_1021])]), c_0_1036]), c_0_1037]), c_0_1038])).
cnf(c_0_1045,plain,(~node8(s21)|~m_main_v_state2(s22,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_1026])).
cnf(c_0_1046,plain,(node16(s21)|m_main_v_turn(s22,c_2)), inference(sr,[status(thm)],[c_0_1039, c_0_1034])).
cnf(c_0_1047,plain,(node17(s23,s24)|node16(s23)|node15(s23)), inference(spm,[status(thm)],[c_0_152, c_0_1040])).
cnf(c_0_1048,plain,(node15(s22)|m_main_v_turn(s23,c_1)), inference(spm,[status(thm)],[c_0_1041, c_0_1042])).
cnf(c_0_1049,plain,(~node15(s22)), inference(spm,[status(thm)],[c_0_229, c_0_1029])).
cnf(c_0_1050,plain,(~m_main_v_state2(X1,c_c2)|~trans(s21,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1043, c_0_988]), c_0_923])])).
cnf(c_0_1051,plain,(m_main_v_state2(s22,c_c2)|~node16(s23)), inference(spm,[status(thm)],[c_0_188, c_0_1044])).
cnf(c_0_1052,plain,(~m_main_v_state2(s22,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1045, c_0_1021])])).
cnf(c_0_1053,plain,(node3(s21)|node2(s21)|node1(s21)|m_main_v_state1(s22,c_c1)|~node16(s21)), inference(spm,[status(thm)],[c_0_401, c_0_966])).
cnf(c_0_1054,plain,(node16(s21)|~node5(s22)), inference(spm,[status(thm)],[c_0_354, c_0_1046])).
cnf(c_0_1055,plain,(~node4(X1)|~m_main_v_state2(s21,c_n2)|~trans(s21,X1)), inference(spm,[status(thm)],[c_0_200, c_0_1003])).
cnf(c_0_1056,plain,(node4(s21)|~node3(s22)), inference(spm,[status(thm)],[c_0_201, c_0_995])).
cnf(c_0_1057,plain,(m_main_v_state2(s22,c_t2)|~node4(s21)), inference(spm,[status(thm)],[c_0_630, c_0_966])).
cnf(c_0_1058,plain,(node4(s21)|~node1(s22)), inference(spm,[status(thm)],[c_0_203, c_0_995])).
cnf(c_0_1059,plain,(node16(s23)|node15(s23)|m_main_v_turn(s24,X1)|~m_main_v_turn(s23,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1047])).
cnf(c_0_1060,plain,(m_main_v_turn(s23,c_1)), inference(sr,[status(thm)],[c_0_1048, c_0_1049])).
cnf(c_0_1061,plain,(~node16(s23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1050, c_0_1051]), c_0_966])])).
cnf(c_0_1062,plain,(m_main_v_state2(s23,c_t2)), inference(sr,[status(thm)],[c_0_1044, c_0_1052])).
cnf(c_0_1063,plain,(m_main_v_state1(s22,c_c1)|~node5(s22)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1053, c_0_1054]), c_0_977]), c_0_978]), c_0_979])).
cnf(c_0_1064,plain,(~node4(X1)|~trans(s21,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1055, c_0_988]), c_0_923])])).
cnf(c_0_1065,plain,(m_main_v_state1(s22,c_c1)|~node3(s22)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1056]), c_0_977]), c_0_978]), c_0_979])).
cnf(c_0_1066,plain,(~node4(s21)|~node2(s22)), inference(spm,[status(thm)],[c_0_239, c_0_1057])).
cnf(c_0_1067,plain,(node4(s21)|~node2(s22)), inference(spm,[status(thm)],[c_0_202, c_0_995])).
cnf(c_0_1068,plain,(m_main_v_state1(s22,c_c1)|~node1(s22)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1058]), c_0_977]), c_0_978]), c_0_979])).
cnf(c_0_1069,plain,(node16(s23)|node15(s23)|m_main_v_turn(s24,c_1)), inference(spm,[status(thm)],[c_0_1059, c_0_1060])).
cnf(c_0_1070,plain,(node17(s23,s24)|node15(s23)), inference(sr,[status(thm)],[c_0_1047, c_0_1061])).
cnf(c_0_1071,plain,(m_main_v_turn(X1,c_2)|~node18(s23,X1)|~m_main_v_state1(s23,c_n1)), inference(spm,[status(thm)],[c_0_45, c_0_1062])).
cnf(c_0_1072,plain,(~node5(s22)), inference(spm,[status(thm)],[c_0_233, c_0_1063])).
cnf(c_0_1073,plain,(~node4(s22)), inference(spm,[status(thm)],[c_0_1064, c_0_966])).
cnf(c_0_1074,plain,(~node3(s22)), inference(spm,[status(thm)],[c_0_415, c_0_1065])).
cnf(c_0_1075,plain,(~node2(s22)), inference(spm,[status(thm)],[c_0_1066, c_0_1067])).
cnf(c_0_1076,plain,(~node1(s22)), inference(spm,[status(thm)],[c_0_416, c_0_1068])).
cnf(c_0_1077,plain,(node15(s23)|m_main_v_turn(s24,c_1)), inference(sr,[status(thm)],[c_0_1069, c_0_1061])).
cnf(c_0_1078,plain,(node15(s23)|m_main_v_turn(s23,X1)|~m_main_v_turn(s24,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1070])).
cnf(c_0_1079,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s23,c_n1)|~trans(s23,X1)), inference(spm,[status(thm)],[c_0_1071, c_0_15])).
cnf(c_0_1080,plain,(~m_main_v_turn(s23,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1060])).
cnf(c_0_1081,plain,(m_main_v_state1(X1,c_n1)|~node7(s22,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1029]), c_0_1072]), c_0_1073]), c_0_1074]), c_0_1075]), c_0_1076])).
cnf(c_0_1082,plain,(node15(s23)|~m_main_v_turn(s24,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1077])).
cnf(c_0_1083,plain,(m_main_v_state2(s22,c_c2)|~node10(s23)), inference(spm,[status(thm)],[c_0_257, c_0_1044])).
cnf(c_0_1084,plain,(m_main_v_state2(s22,c_c2)|~node9(s23)), inference(spm,[status(thm)],[c_0_258, c_0_1044])).
cnf(c_0_1085,plain,(m_main_v_state2(s22,c_c2)|~node8(s23)), inference(spm,[status(thm)],[c_0_168, c_0_1044])).
cnf(c_0_1086,plain,(node15(s23)|~m_main_v_state1(s23,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1078, c_0_1079]), c_0_1040])]), c_0_1080])).
cnf(c_0_1087,plain,(m_main_v_state1(X1,c_n1)|~trans(s22,X1)), inference(spm,[status(thm)],[c_0_1081, c_0_7])).
cnf(c_0_1088,plain,(node10(s23)|node9(s23)|node8(s23)|m_main_v_state2(s24,c_c2)|~node15(s23)), inference(spm,[status(thm)],[c_0_467, c_0_1040])).
cnf(c_0_1089,plain,(node15(s23)|~node12(s24)), inference(spm,[status(thm)],[c_0_1082, c_0_35])).
cnf(c_0_1090,plain,(~node10(s23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1050, c_0_1083]), c_0_966])])).
cnf(c_0_1091,plain,(~node9(s23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1050, c_0_1084]), c_0_966])])).
cnf(c_0_1092,plain,(~node8(s23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1050, c_0_1085]), c_0_966])])).
cnf(c_0_1093,plain,(m_main_v_state1(s24,c_t1)|~node15(s23)), inference(spm,[status(thm)],[c_0_474, c_0_1040])).
cnf(c_0_1094,plain,(node15(s23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1086, c_0_1087]), c_0_1004])])).
cnf(c_0_1095,plain,(m_main_v_state2(s24,c_c2)|~node12(s24)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1088, c_0_1089]), c_0_1090]), c_0_1091]), c_0_1092])).
cnf(c_0_1096,plain,(~node15(s23)|~node11(s24)), inference(spm,[status(thm)],[c_0_165, c_0_1093])).
cnf(c_0_1097,plain,(~node15(s23)|~node9(s24)), inference(spm,[status(thm)],[c_0_167, c_0_1093])).
cnf(c_0_1098,plain,(m_main_v_state2(s24,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1088, c_0_1094])]), c_0_1090]), c_0_1091]), c_0_1092])).
cnf(c_0_1099,plain,(~node12(s24)), inference(spm,[status(thm)],[c_0_345, c_0_1095])).
cnf(c_0_1100,plain,(~node11(s24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1096, c_0_1094])])).
cnf(c_0_1101,plain,(~node9(s24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1097, c_0_1094])])).
cnf(c_0_1102,plain,(~node10(s24)), inference(spm,[status(thm)],[c_0_339, c_0_1098])).
cnf(c_0_1103,plain,(~node8(s24)), inference(spm,[status(thm)],[c_0_343, c_0_1098])).
cnf(c_0_1104,plain,(m_main_v_state2(X1,c_n2)|~node14(s24,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1098]), c_0_1099]), c_0_1100]), c_0_1101]), c_0_1102]), c_0_1103])).
cnf(c_0_1105,plain,(m_main_v_turn(s24,c_2)|~node15(s23)), inference(spm,[status(thm)],[c_0_470, c_0_1040])).
cnf(c_0_1106,plain,(m_main_v_state2(X1,c_n2)|~trans(s24,X1)), inference(spm,[status(thm)],[c_0_1104, c_0_29])).
cnf(c_0_1107,plain,(trans(s24,s25)), inference(spm,[status(thm)],[c_0_8, c_0_86])).
cnf(c_0_1108,plain,(~node15(s23)|~node1(s24)), inference(spm,[status(thm)],[c_0_203, c_0_1093])).
cnf(c_0_1109,plain,(~node15(s23)|~node5(s24)), inference(spm,[status(thm)],[c_0_354, c_0_1105])).
cnf(c_0_1110,plain,(~node15(s23)|~node3(s24)), inference(spm,[status(thm)],[c_0_201, c_0_1093])).
cnf(c_0_1111,plain,(~node15(s23)|~node2(s24)), inference(spm,[status(thm)],[c_0_202, c_0_1093])).
cnf(c_0_1112,plain,(~node15(s23)|~m_main_v_state1(s24,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_1093])).
cnf(c_0_1113,plain,(~m_main_v_state2(X1,c_t2)|~trans(s24,X1)), inference(spm,[status(thm)],[c_0_0, c_0_1106])).
cnf(c_0_1114,plain,(node6(s24,s25)|node5(s24)|node4(s24)|node3(s24)|node2(s24)|node1(s24)|m_main_v_state1(s24,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1107])).
cnf(c_0_1115,plain,(~node1(s24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1108, c_0_1094])])).
cnf(c_0_1116,plain,(~node5(s24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1109, c_0_1094])])).
cnf(c_0_1117,plain,(~node4(s24)), inference(spm,[status(thm)],[c_0_361, c_0_1098])).
cnf(c_0_1118,plain,(~node3(s24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1110, c_0_1094])])).
cnf(c_0_1119,plain,(~node2(s24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1111, c_0_1094])])).
cnf(c_0_1120,plain,(~m_main_v_state1(s24,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1112, c_0_1094])])).
cnf(c_0_1121,plain,(trans(s25,s26)), inference(spm,[status(thm)],[c_0_8, c_0_87])).
cnf(c_0_1122,plain,(~node15(X1)|~trans(s24,X1)), inference(spm,[status(thm)],[c_0_1113, c_0_17])).
cnf(c_0_1123,plain,(node6(s24,s25)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1114, c_0_1115]), c_0_1116]), c_0_1117]), c_0_1118]), c_0_1119]), c_0_1120])).
cnf(c_0_1124,plain,(node17(s24,s25)|node16(s24)|node15(s24)), inference(spm,[status(thm)],[c_0_152, c_0_1107])).
cnf(c_0_1125,plain,(node17(s25,s26)|node16(s25)|node15(s25)), inference(spm,[status(thm)],[c_0_152, c_0_1121])).
cnf(c_0_1126,plain,(~node15(s25)), inference(spm,[status(thm)],[c_0_1122, c_0_1107])).
cnf(c_0_1127,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s24,X1)), inference(spm,[status(thm)],[c_0_63, c_0_1106])).
cnf(c_0_1128,plain,(m_main_v_state1(s25,X1)|~m_main_v_state1(s24,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1123])).
cnf(c_0_1129,plain,(m_main_v_state1(s24,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1093, c_0_1094])])).
cnf(c_0_1130,plain,(node16(s24)|node15(s24)|m_main_v_turn(s25,X1)|~m_main_v_turn(s24,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1124])).
cnf(c_0_1131,plain,(~node15(s24)|~node15(s23)), inference(spm,[status(thm)],[c_0_189, c_0_1093])).
cnf(c_0_1132,plain,(node17(s25,s26)|node16(s25)), inference(sr,[status(thm)],[c_0_1125, c_0_1126])).
cnf(c_0_1133,plain,(m_main_v_turn(X1,c_1)|~node18(s25,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1127, c_0_1128]), c_0_1107]), c_0_1129])]), c_0_1126])).
cnf(c_0_1134,plain,(node16(s24)|node15(s24)|m_main_v_turn(s25,c_2)|~node15(s23)), inference(spm,[status(thm)],[c_0_1130, c_0_1105])).
cnf(c_0_1135,plain,(~node15(s24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1131, c_0_1094])])).
cnf(c_0_1136,plain,(node16(s25)|m_main_v_turn(s25,X1)|~m_main_v_turn(s26,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1132])).
cnf(c_0_1137,plain,(m_main_v_turn(X1,c_1)|~trans(s25,X1)), inference(spm,[status(thm)],[c_0_1133, c_0_15])).
cnf(c_0_1138,plain,(node16(s24)|m_main_v_turn(s25,c_2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1134, c_0_1094])]), c_0_1135])).
cnf(c_0_1139,plain,(~node16(s24)), inference(spm,[status(thm)],[c_0_392, c_0_1098])).
cnf(c_0_1140,plain,(~m_main_v_state2(X1,c_c2)|~trans(s24,X1)), inference(spm,[status(thm)],[c_0_4, c_0_1106])).
cnf(c_0_1141,plain,(trans(s26,s27)), inference(spm,[status(thm)],[c_0_8, c_0_88])).
cnf(c_0_1142,plain,(m_main_v_state2(s26,c_t2)|~node16(s25)), inference(spm,[status(thm)],[c_0_400, c_0_1121])).
cnf(c_0_1143,plain,(node16(s25)|m_main_v_turn(s25,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1136, c_0_1137]), c_0_1121])])).
cnf(c_0_1144,plain,(m_main_v_turn(s25,c_2)), inference(sr,[status(thm)],[c_0_1138, c_0_1139])).
cnf(c_0_1145,plain,(~node3(X1)|~trans(s24,X1)), inference(spm,[status(thm)],[c_0_1140, c_0_12])).
cnf(c_0_1146,plain,(~node1(X1)|~trans(s24,X1)), inference(spm,[status(thm)],[c_0_1113, c_0_5])).
cnf(c_0_1147,plain,(node17(s26,s27)|node16(s26)|node15(s26)), inference(spm,[status(thm)],[c_0_152, c_0_1141])).
cnf(c_0_1148,plain,(~node16(s26)|~node16(s25)), inference(spm,[status(thm)],[c_0_188, c_0_1142])).
cnf(c_0_1149,plain,(node16(s25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_1143]), c_0_1144])])).
cnf(c_0_1150,plain,(node3(s25)|node2(s25)|node1(s25)|m_main_v_state1(s26,c_c1)|~node16(s25)), inference(spm,[status(thm)],[c_0_401, c_0_1121])).
cnf(c_0_1151,plain,(~node3(s25)), inference(spm,[status(thm)],[c_0_1145, c_0_1107])).
cnf(c_0_1152,plain,(~node2(s25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_1128]), c_0_1129])])).
cnf(c_0_1153,plain,(~node1(s25)), inference(spm,[status(thm)],[c_0_1146, c_0_1107])).
cnf(c_0_1154,plain,(node16(s26)|node15(s26)|m_main_v_turn(s27,X1)|~m_main_v_turn(s26,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1147])).
cnf(c_0_1155,plain,(~node16(s26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1148, c_0_1149])])).
cnf(c_0_1156,plain,(~node16(s25)|~node8(s26)), inference(spm,[status(thm)],[c_0_168, c_0_1142])).
cnf(c_0_1157,plain,(m_main_v_state1(s26,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1150, c_0_1149])]), c_0_1151]), c_0_1152]), c_0_1153])).
cnf(c_0_1158,plain,(~node16(s25)|~node10(s26)), inference(spm,[status(thm)],[c_0_257, c_0_1142])).
cnf(c_0_1159,plain,(~node16(s25)|~node9(s26)), inference(spm,[status(thm)],[c_0_258, c_0_1142])).
cnf(c_0_1160,plain,(~node16(s25)|~m_main_v_state2(s26,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_1142])).
cnf(c_0_1161,plain,(node3(s24)|node2(s24)|node1(s24)|m_main_v_state1(s25,c_c1)|~node16(s24)), inference(spm,[status(thm)],[c_0_401, c_0_1107])).
cnf(c_0_1162,plain,(node16(s24)|~node5(s25)), inference(spm,[status(thm)],[c_0_354, c_0_1138])).
cnf(c_0_1163,plain,(node15(s26)|m_main_v_turn(s27,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1154, c_0_1137]), c_0_1121])]), c_0_1155])).
cnf(c_0_1164,plain,(node13(s26,s27)|node12(s26)|node11(s26)|node10(s26)|node9(s26)|node8(s26)|m_main_v_state2(s26,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1141])).
cnf(c_0_1165,plain,(~node8(s26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1156, c_0_1149])])).
cnf(c_0_1166,plain,(~node12(s26)), inference(spm,[status(thm)],[c_0_428, c_0_1157])).
cnf(c_0_1167,plain,(~node11(s26)), inference(spm,[status(thm)],[c_0_429, c_0_1157])).
cnf(c_0_1168,plain,(~node10(s26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1158, c_0_1149])])).
cnf(c_0_1169,plain,(~node9(s26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1159, c_0_1149])])).
cnf(c_0_1170,plain,(~m_main_v_state2(s26,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1160, c_0_1149])])).
cnf(c_0_1171,plain,(m_main_v_state1(s25,c_c1)|~node5(s25)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1161, c_0_1162]), c_0_1118]), c_0_1119]), c_0_1115])).
cnf(c_0_1172,plain,(node15(s26)|~m_main_v_turn(s27,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1163])).
cnf(c_0_1173,plain,(node13(s26,s27)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1164, c_0_1165]), c_0_1166]), c_0_1167]), c_0_1168]), c_0_1169]), c_0_1170])).
cnf(c_0_1174,plain,(node6(s25,s26)|node5(s25)|node4(s25)|node3(s25)|node2(s25)|node1(s25)|m_main_v_state1(s25,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1121])).
cnf(c_0_1175,plain,(~node5(s25)), inference(spm,[status(thm)],[c_0_233, c_0_1171])).
cnf(c_0_1176,plain,(~m_main_v_state1(s25,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_1128]), c_0_1129])])).
cnf(c_0_1177,plain,(trans(s27,s28)), inference(spm,[status(thm)],[c_0_8, c_0_89])).
cnf(c_0_1178,plain,(~node15(s26)), inference(spm,[status(thm)],[c_0_229, c_0_1157])).
cnf(c_0_1179,plain,(node15(s26)|~node12(s27)), inference(spm,[status(thm)],[c_0_1172, c_0_35])).
cnf(c_0_1180,plain,(m_main_v_state2(s27,X1)|~m_main_v_state2(s26,X1)), inference(spm,[status(thm)],[c_0_43, c_0_1173])).
cnf(c_0_1181,plain,(m_main_v_state2(s26,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1142, c_0_1149])])).
cnf(c_0_1182,plain,(node6(s25,s26)|node4(s25)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1174, c_0_1175]), c_0_1151]), c_0_1152]), c_0_1153]), c_0_1176])).
cnf(c_0_1183,plain,(node13(s27,s28)|node12(s27)|node11(s27)|node10(s27)|node9(s27)|node8(s27)|m_main_v_state2(s27,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1177])).
cnf(c_0_1184,plain,(~node12(s27)), inference(spm,[status(thm)],[c_0_1178, c_0_1179])).
cnf(c_0_1185,plain,(~node10(s27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_1180]), c_0_1181])])).
cnf(c_0_1186,plain,(~node9(s27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_1180]), c_0_1181])])).
cnf(c_0_1187,plain,(~node8(s27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_1180]), c_0_1181])])).
cnf(c_0_1188,plain,(~m_main_v_state2(s27,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_1180]), c_0_1181])])).
cnf(c_0_1189,plain,(node4(s25)|m_main_v_state1(s26,X1)|~m_main_v_state1(s25,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1182])).
cnf(c_0_1190,plain,(node17(s27,s28)|node16(s27)|node15(s27)), inference(spm,[status(thm)],[c_0_152, c_0_1177])).
cnf(c_0_1191,plain,(node13(s27,s28)|node11(s27)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1183, c_0_1184]), c_0_1185]), c_0_1186]), c_0_1187]), c_0_1188])).
cnf(c_0_1192,plain,(m_main_v_state2(X1,c_t2)|~node14(s25,X1)|~m_main_v_state2(s25,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_1128]), c_0_1129])])).
cnf(c_0_1193,plain,(node4(s25)|m_main_v_state1(s26,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1189, c_0_16]), c_0_1149])])).
cnf(c_0_1194,plain,(node16(s27)|node15(s27)|m_main_v_turn(s28,X1)|~m_main_v_turn(s27,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1190])).
cnf(c_0_1195,plain,(m_main_v_turn(s27,c_1)), inference(sr,[status(thm)],[c_0_1163, c_0_1178])).
cnf(c_0_1196,plain,(node11(s27)|m_main_v_state2(s28,X1)|~m_main_v_state2(s27,X1)), inference(spm,[status(thm)],[c_0_43, c_0_1191])).
cnf(c_0_1197,plain,(m_main_v_state2(X1,c_t2)|~m_main_v_state2(s25,c_n2)|~trans(s25,X1)), inference(spm,[status(thm)],[c_0_1192, c_0_29])).
cnf(c_0_1198,plain,(node4(s25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_1193]), c_0_1157])])).
cnf(c_0_1199,plain,(~node16(s27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_1180]), c_0_1181])])).
cnf(c_0_1200,plain,(~node16(s25)|~node4(s26)), inference(spm,[status(thm)],[c_0_200, c_0_1142])).
cnf(c_0_1201,plain,(~node16(s25)|~node2(s26)), inference(spm,[status(thm)],[c_0_239, c_0_1142])).
cnf(c_0_1202,plain,(node16(s27)|node15(s27)|m_main_v_turn(s28,c_1)), inference(spm,[status(thm)],[c_0_1194, c_0_1195])).
cnf(c_0_1203,plain,(node11(s27)|m_main_v_state2(s28,X1)|~m_main_v_state2(s26,X1)), inference(spm,[status(thm)],[c_0_1196, c_0_1180])).
cnf(c_0_1204,plain,(m_main_v_state2(X1,c_t2)|~trans(s25,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1197, c_0_39]), c_0_1198])])).
cnf(c_0_1205,plain,(node17(s27,s28)|node15(s27)), inference(sr,[status(thm)],[c_0_1190, c_0_1199])).
cnf(c_0_1206,plain,(m_main_v_turn(X1,c_2)|~node18(s27,X1)|~m_main_v_state1(s27,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1180]), c_0_1181])])).
cnf(c_0_1207,plain,(~node5(s26)), inference(spm,[status(thm)],[c_0_233, c_0_1157])).
cnf(c_0_1208,plain,(~node4(s26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1200, c_0_1149])])).
cnf(c_0_1209,plain,(~node3(s26)), inference(spm,[status(thm)],[c_0_415, c_0_1157])).
cnf(c_0_1210,plain,(~node2(s26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1201, c_0_1149])])).
cnf(c_0_1211,plain,(~node1(s26)), inference(spm,[status(thm)],[c_0_416, c_0_1157])).
cnf(c_0_1212,plain,(node15(s27)|m_main_v_turn(s28,c_1)), inference(sr,[status(thm)],[c_0_1202, c_0_1199])).
cnf(c_0_1213,plain,(m_main_v_state1(s28,c_t1)|~node11(s27)), inference(spm,[status(thm)],[c_0_316, c_0_1177])).
cnf(c_0_1214,plain,(node11(s27)|m_main_v_state2(s27,X1)|~m_main_v_state2(s28,X1)), inference(spm,[status(thm)],[c_0_47, c_0_1191])).
cnf(c_0_1215,plain,(node11(s27)|m_main_v_state2(s28,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1203, c_0_1204]), c_0_1121])])).
cnf(c_0_1216,plain,(node15(s27)|m_main_v_turn(s27,X1)|~m_main_v_turn(s28,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1205])).
cnf(c_0_1217,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s27,c_n1)|~trans(s27,X1)), inference(spm,[status(thm)],[c_0_1206, c_0_15])).
cnf(c_0_1218,plain,(~m_main_v_turn(s27,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1195])).
cnf(c_0_1219,plain,(m_main_v_state1(X1,c_n1)|~node7(s26,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1157]), c_0_1207]), c_0_1208]), c_0_1209]), c_0_1210]), c_0_1211])).
cnf(c_0_1220,plain,(node15(s27)|~m_main_v_turn(s28,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1212])).
cnf(c_0_1221,plain,(m_main_v_turn(s28,c_2)|~node11(s27)), inference(spm,[status(thm)],[c_0_251, c_0_1177])).
cnf(c_0_1222,plain,(trans(s28,s29)), inference(spm,[status(thm)],[c_0_8, c_0_90])).
cnf(c_0_1223,plain,(~node11(s27)|~node3(s28)), inference(spm,[status(thm)],[c_0_201, c_0_1213])).
cnf(c_0_1224,plain,(node11(s27)|~node3(s28)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1214, c_0_12]), c_0_1188])).
cnf(c_0_1225,plain,(~node11(s27)|~node2(s28)), inference(spm,[status(thm)],[c_0_202, c_0_1213])).
cnf(c_0_1226,plain,(node11(s27)|~node2(s28)), inference(spm,[status(thm)],[c_0_239, c_0_1215])).
cnf(c_0_1227,plain,(node15(s27)|~m_main_v_state1(s27,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1216, c_0_1217]), c_0_1177])]), c_0_1218])).
cnf(c_0_1228,plain,(m_main_v_state1(X1,c_n1)|~trans(s26,X1)), inference(spm,[status(thm)],[c_0_1219, c_0_7])).
cnf(c_0_1229,plain,(node10(s27)|node9(s27)|node8(s27)|m_main_v_state2(s28,c_c2)|~node15(s27)), inference(spm,[status(thm)],[c_0_467, c_0_1177])).
cnf(c_0_1230,plain,(node15(s27)|~node11(s27)), inference(spm,[status(thm)],[c_0_1220, c_0_1221])).
cnf(c_0_1231,plain,(node6(s28,s29)|node5(s28)|node4(s28)|node3(s28)|node2(s28)|node1(s28)|m_main_v_state1(s28,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1222])).
cnf(c_0_1232,plain,(~node3(s28)), inference(spm,[status(thm)],[c_0_1223, c_0_1224])).
cnf(c_0_1233,plain,(~node2(s28)), inference(spm,[status(thm)],[c_0_1225, c_0_1226])).
cnf(c_0_1234,plain,(m_main_v_turn(s28,c_2)|~node15(s27)), inference(spm,[status(thm)],[c_0_470, c_0_1177])).
cnf(c_0_1235,plain,(m_main_v_state1(s28,c_t1)|~node15(s27)), inference(spm,[status(thm)],[c_0_474, c_0_1177])).
cnf(c_0_1236,plain,(node15(s27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1227, c_0_1228]), c_0_1141])])).
cnf(c_0_1237,plain,(node15(s27)|~node12(s28)), inference(spm,[status(thm)],[c_0_1220, c_0_35])).
cnf(c_0_1238,plain,(m_main_v_state2(s28,c_c2)|~node11(s27)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1229, c_0_1230]), c_0_1185]), c_0_1186]), c_0_1187])).
cnf(c_0_1239,plain,(node6(s28,s29)|node5(s28)|node4(s28)|node1(s28)|m_main_v_state1(s28,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1231, c_0_1232]), c_0_1233])).
cnf(c_0_1240,plain,(~node15(s27)|~node5(s28)), inference(spm,[status(thm)],[c_0_354, c_0_1234])).
cnf(c_0_1241,plain,(~node15(s27)|~node1(s28)), inference(spm,[status(thm)],[c_0_203, c_0_1235])).
cnf(c_0_1242,plain,(m_main_v_state2(s28,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1229, c_0_1236])]), c_0_1185]), c_0_1186]), c_0_1187])).
cnf(c_0_1243,plain,(m_main_v_state2(s28,c_c2)|~node12(s28)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1229, c_0_1237]), c_0_1185]), c_0_1186]), c_0_1187])).
cnf(c_0_1244,plain,(~node15(s27)|~node11(s28)), inference(spm,[status(thm)],[c_0_165, c_0_1235])).
cnf(c_0_1245,plain,(~node11(s27)|~node10(s28)), inference(spm,[status(thm)],[c_0_339, c_0_1238])).
cnf(c_0_1246,plain,(node11(s27)|~node10(s28)), inference(spm,[status(thm)],[c_0_257, c_0_1215])).
cnf(c_0_1247,plain,(~node11(s27)|~node9(s28)), inference(spm,[status(thm)],[c_0_167, c_0_1213])).
cnf(c_0_1248,plain,(node11(s27)|~node9(s28)), inference(spm,[status(thm)],[c_0_258, c_0_1215])).
cnf(c_0_1249,plain,(~node11(s27)|~node8(s28)), inference(spm,[status(thm)],[c_0_343, c_0_1238])).
cnf(c_0_1250,plain,(node11(s27)|~node8(s28)), inference(spm,[status(thm)],[c_0_168, c_0_1215])).
cnf(c_0_1251,plain,(node5(s28)|node4(s28)|node1(s28)|m_main_v_state1(s28,c_c1)|m_main_v_state1(s29,X1)|~m_main_v_state1(s28,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1239])).
cnf(c_0_1252,plain,(m_main_v_state1(s28,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1235, c_0_1236])])).
cnf(c_0_1253,plain,(~node5(s28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1240, c_0_1236])])).
cnf(c_0_1254,plain,(~node1(s28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1241, c_0_1236])])).
cnf(c_0_1255,plain,(~node4(s28)), inference(spm,[status(thm)],[c_0_361, c_0_1242])).
cnf(c_0_1256,plain,(~node12(s28)), inference(spm,[status(thm)],[c_0_345, c_0_1243])).
cnf(c_0_1257,plain,(~node11(s28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1244, c_0_1236])])).
cnf(c_0_1258,plain,(~node10(s28)), inference(spm,[status(thm)],[c_0_1245, c_0_1246])).
cnf(c_0_1259,plain,(~node9(s28)), inference(spm,[status(thm)],[c_0_1247, c_0_1248])).
cnf(c_0_1260,plain,(~node8(s28)), inference(spm,[status(thm)],[c_0_1249, c_0_1250])).
cnf(c_0_1261,plain,(~node15(s27)|~m_main_v_state1(s28,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_1235])).
cnf(c_0_1262,plain,(m_main_v_state1(s28,c_c1)|m_main_v_state1(s29,c_t1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1251, c_0_1252]), c_0_1253]), c_0_1254]), c_0_1255])).
cnf(c_0_1263,plain,(m_main_v_state2(X1,c_n2)|~node14(s28,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1242]), c_0_1256]), c_0_1257]), c_0_1258]), c_0_1259]), c_0_1260])).
cnf(c_0_1264,plain,(~m_main_v_state1(s28,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1261, c_0_1236])])).
cnf(c_0_1265,plain,(trans(s29,s30)), inference(spm,[status(thm)],[c_0_8, c_0_91])).
cnf(c_0_1266,plain,(m_main_v_state1(s28,c_c1)|~node15(s29)), inference(spm,[status(thm)],[c_0_189, c_0_1262])).
cnf(c_0_1267,plain,(m_main_v_state2(X1,c_n2)|~trans(s28,X1)), inference(spm,[status(thm)],[c_0_1263, c_0_29])).
cnf(c_0_1268,plain,(node6(s28,s29)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1239, c_0_1253]), c_0_1255]), c_0_1254]), c_0_1264])).
cnf(c_0_1269,plain,(node17(s29,s30)|node16(s29)|node15(s29)), inference(spm,[status(thm)],[c_0_152, c_0_1265])).
cnf(c_0_1270,plain,(~node15(s29)), inference(spm,[status(thm)],[c_0_1264, c_0_1266])).
cnf(c_0_1271,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s28,X1)), inference(spm,[status(thm)],[c_0_63, c_0_1267])).
cnf(c_0_1272,plain,(m_main_v_state1(s29,X1)|~m_main_v_state1(s28,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1268])).
cnf(c_0_1273,plain,(node17(s28,s29)|node16(s28)|node15(s28)), inference(spm,[status(thm)],[c_0_152, c_0_1222])).
cnf(c_0_1274,plain,(~node15(s28)|~node15(s27)), inference(spm,[status(thm)],[c_0_189, c_0_1235])).
cnf(c_0_1275,plain,(node17(s29,s30)|node16(s29)), inference(sr,[status(thm)],[c_0_1269, c_0_1270])).
cnf(c_0_1276,plain,(m_main_v_turn(X1,c_1)|~node18(s29,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1271, c_0_1272]), c_0_1222]), c_0_1252])]), c_0_1270])).
cnf(c_0_1277,plain,(node16(s28)|node15(s28)|m_main_v_turn(s29,X1)|~m_main_v_turn(s28,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1273])).
cnf(c_0_1278,plain,(m_main_v_turn(s28,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1234, c_0_1236])])).
cnf(c_0_1279,plain,(~node15(s28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1274, c_0_1236])])).
cnf(c_0_1280,plain,(node16(s29)|m_main_v_turn(s29,X1)|~m_main_v_turn(s30,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1275])).
cnf(c_0_1281,plain,(m_main_v_turn(X1,c_1)|~trans(s29,X1)), inference(spm,[status(thm)],[c_0_1276, c_0_15])).
cnf(c_0_1282,plain,(node16(s28)|m_main_v_turn(s29,c_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1277, c_0_1278]), c_0_1279])).
cnf(c_0_1283,plain,(~node16(s28)), inference(spm,[status(thm)],[c_0_392, c_0_1242])).
cnf(c_0_1284,plain,(node16(s29)|m_main_v_turn(s29,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1280, c_0_1281]), c_0_1265])])).
cnf(c_0_1285,plain,(m_main_v_turn(s29,c_2)), inference(sr,[status(thm)],[c_0_1282, c_0_1283])).
cnf(c_0_1286,plain,(m_main_v_state1(s28,c_c1)|~node3(s29)), inference(spm,[status(thm)],[c_0_201, c_0_1262])).
cnf(c_0_1287,plain,(m_main_v_state1(s28,c_c1)|~node2(s29)), inference(spm,[status(thm)],[c_0_202, c_0_1262])).
cnf(c_0_1288,plain,(m_main_v_state1(s28,c_c1)|~node1(s29)), inference(spm,[status(thm)],[c_0_203, c_0_1262])).
cnf(c_0_1289,plain,(trans(s30,s31)), inference(spm,[status(thm)],[c_0_8, c_0_92])).
cnf(c_0_1290,plain,(m_main_v_state2(s30,c_t2)|~node16(s29)), inference(spm,[status(thm)],[c_0_400, c_0_1265])).
cnf(c_0_1291,plain,(node3(s29)|node2(s29)|node1(s29)|m_main_v_state1(s30,c_c1)|~node16(s29)), inference(spm,[status(thm)],[c_0_401, c_0_1265])).
cnf(c_0_1292,plain,(node16(s29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_1284]), c_0_1285])])).
cnf(c_0_1293,plain,(~node3(s29)), inference(spm,[status(thm)],[c_0_1264, c_0_1286])).
cnf(c_0_1294,plain,(~node2(s29)), inference(spm,[status(thm)],[c_0_1264, c_0_1287])).
cnf(c_0_1295,plain,(~node1(s29)), inference(spm,[status(thm)],[c_0_1264, c_0_1288])).
cnf(c_0_1296,plain,(node17(s30,s31)|node16(s30)|node15(s30)), inference(spm,[status(thm)],[c_0_152, c_0_1289])).
cnf(c_0_1297,plain,(~node16(s29)|~node8(s30)), inference(spm,[status(thm)],[c_0_168, c_0_1290])).
cnf(c_0_1298,plain,(m_main_v_state1(s30,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1291, c_0_1292])]), c_0_1293]), c_0_1294]), c_0_1295])).
cnf(c_0_1299,plain,(~node16(s29)|~node10(s30)), inference(spm,[status(thm)],[c_0_257, c_0_1290])).
cnf(c_0_1300,plain,(~node16(s29)|~node9(s30)), inference(spm,[status(thm)],[c_0_258, c_0_1290])).
cnf(c_0_1301,plain,(~node16(s29)|~m_main_v_state2(s30,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_1290])).
cnf(c_0_1302,plain,(node16(s29)|node15(s29)|m_main_v_turn(s30,X1)|~m_main_v_turn(s29,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1269])).
cnf(c_0_1303,plain,(node16(s30)|node15(s30)|m_main_v_turn(s31,X1)|~m_main_v_turn(s30,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1296])).
cnf(c_0_1304,plain,(~node16(s30)|~node16(s29)), inference(spm,[status(thm)],[c_0_188, c_0_1290])).
cnf(c_0_1305,plain,(node13(s30,s31)|node12(s30)|node11(s30)|node10(s30)|node9(s30)|node8(s30)|m_main_v_state2(s30,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1289])).
cnf(c_0_1306,plain,(~node8(s30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1297, c_0_1292])])).
cnf(c_0_1307,plain,(~node12(s30)), inference(spm,[status(thm)],[c_0_428, c_0_1298])).
cnf(c_0_1308,plain,(~node11(s30)), inference(spm,[status(thm)],[c_0_429, c_0_1298])).
cnf(c_0_1309,plain,(~node10(s30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1299, c_0_1292])])).
cnf(c_0_1310,plain,(~node9(s30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1300, c_0_1292])])).
cnf(c_0_1311,plain,(~m_main_v_state2(s30,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1301, c_0_1292])])).
cnf(c_0_1312,plain,(node16(s29)|node15(s29)|m_main_v_turn(s30,c_2)), inference(spm,[status(thm)],[c_0_1302, c_0_1285])).
cnf(c_0_1313,plain,(trans(s31,s32)), inference(spm,[status(thm)],[c_0_8, c_0_93])).
cnf(c_0_1314,plain,(node16(s30)|node15(s30)|m_main_v_turn(s31,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1303, c_0_1281]), c_0_1265])])).
cnf(c_0_1315,plain,(~node16(s30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1304, c_0_1292])])).
cnf(c_0_1316,plain,(node13(s30,s31)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1305, c_0_1306]), c_0_1307]), c_0_1308]), c_0_1309]), c_0_1310]), c_0_1311])).
cnf(c_0_1317,plain,(node16(s29)|m_main_v_turn(s30,c_2)), inference(sr,[status(thm)],[c_0_1312, c_0_1270])).
cnf(c_0_1318,plain,(node17(s31,s32)|node16(s31)|node15(s31)), inference(spm,[status(thm)],[c_0_152, c_0_1313])).
cnf(c_0_1319,plain,(node15(s30)|m_main_v_turn(s31,c_1)), inference(sr,[status(thm)],[c_0_1314, c_0_1315])).
cnf(c_0_1320,plain,(~node15(s30)), inference(spm,[status(thm)],[c_0_229, c_0_1298])).
cnf(c_0_1321,plain,(m_main_v_state2(s31,X1)|~m_main_v_state2(s30,X1)), inference(spm,[status(thm)],[c_0_43, c_0_1316])).
cnf(c_0_1322,plain,(m_main_v_state2(s30,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1290, c_0_1292])])).
cnf(c_0_1323,plain,(node16(s29)|~node5(s30)), inference(spm,[status(thm)],[c_0_354, c_0_1317])).
cnf(c_0_1324,plain,(node16(s31)|node15(s31)|m_main_v_turn(s32,X1)|~m_main_v_turn(s31,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1318])).
cnf(c_0_1325,plain,(m_main_v_turn(s31,c_1)), inference(sr,[status(thm)],[c_0_1319, c_0_1320])).
cnf(c_0_1326,plain,(~node16(s31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_1321]), c_0_1322])])).
cnf(c_0_1327,plain,(m_main_v_state1(s30,c_c1)|~node5(s30)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1291, c_0_1323]), c_0_1293]), c_0_1294]), c_0_1295])).
cnf(c_0_1328,plain,(~node16(s29)|~node4(s30)), inference(spm,[status(thm)],[c_0_200, c_0_1290])).
cnf(c_0_1329,plain,(~node16(s29)|~node2(s30)), inference(spm,[status(thm)],[c_0_239, c_0_1290])).
cnf(c_0_1330,plain,(node16(s31)|node15(s31)|m_main_v_turn(s32,c_1)), inference(spm,[status(thm)],[c_0_1324, c_0_1325])).
cnf(c_0_1331,plain,(node17(s31,s32)|node15(s31)), inference(sr,[status(thm)],[c_0_1318, c_0_1326])).
cnf(c_0_1332,plain,(m_main_v_turn(X1,c_2)|~node18(s31,X1)|~m_main_v_state1(s31,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1321]), c_0_1322])])).
cnf(c_0_1333,plain,(~node5(s30)), inference(spm,[status(thm)],[c_0_233, c_0_1327])).
cnf(c_0_1334,plain,(~node4(s30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1328, c_0_1292])])).
cnf(c_0_1335,plain,(~node3(s30)), inference(spm,[status(thm)],[c_0_415, c_0_1298])).
cnf(c_0_1336,plain,(~node2(s30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1329, c_0_1292])])).
cnf(c_0_1337,plain,(~node1(s30)), inference(spm,[status(thm)],[c_0_416, c_0_1298])).
cnf(c_0_1338,plain,(node15(s31)|m_main_v_turn(s32,c_1)), inference(sr,[status(thm)],[c_0_1330, c_0_1326])).
cnf(c_0_1339,plain,(node15(s31)|m_main_v_turn(s31,X1)|~m_main_v_turn(s32,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1331])).
cnf(c_0_1340,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s31,c_n1)|~trans(s31,X1)), inference(spm,[status(thm)],[c_0_1332, c_0_15])).
cnf(c_0_1341,plain,(~m_main_v_turn(s31,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1325])).
cnf(c_0_1342,plain,(m_main_v_state1(X1,c_n1)|~node7(s30,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1298]), c_0_1333]), c_0_1334]), c_0_1335]), c_0_1336]), c_0_1337])).
cnf(c_0_1343,plain,(node15(s31)|~m_main_v_turn(s32,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1338])).
cnf(c_0_1344,plain,(node15(s31)|~m_main_v_state1(s31,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1339, c_0_1340]), c_0_1313])]), c_0_1341])).
cnf(c_0_1345,plain,(m_main_v_state1(X1,c_n1)|~trans(s30,X1)), inference(spm,[status(thm)],[c_0_1342, c_0_7])).
cnf(c_0_1346,plain,(node10(s31)|node9(s31)|node8(s31)|m_main_v_state2(s32,c_c2)|~node15(s31)), inference(spm,[status(thm)],[c_0_467, c_0_1313])).
cnf(c_0_1347,plain,(node15(s31)|~node12(s32)), inference(spm,[status(thm)],[c_0_1343, c_0_35])).
cnf(c_0_1348,plain,(~node10(s31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_1321]), c_0_1322])])).
cnf(c_0_1349,plain,(~node9(s31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_1321]), c_0_1322])])).
cnf(c_0_1350,plain,(~node8(s31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_1321]), c_0_1322])])).
cnf(c_0_1351,plain,(m_main_v_state1(s32,c_t1)|~node15(s31)), inference(spm,[status(thm)],[c_0_474, c_0_1313])).
cnf(c_0_1352,plain,(node15(s31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1344, c_0_1345]), c_0_1289])])).
cnf(c_0_1353,plain,(m_main_v_turn(s32,c_2)|~node15(s31)), inference(spm,[status(thm)],[c_0_470, c_0_1313])).
cnf(c_0_1354,plain,(m_main_v_state2(s32,c_c2)|~node12(s32)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1346, c_0_1347]), c_0_1348]), c_0_1349]), c_0_1350])).
cnf(c_0_1355,plain,(~node15(s31)|~node11(s32)), inference(spm,[status(thm)],[c_0_165, c_0_1351])).
cnf(c_0_1356,plain,(m_main_v_state2(s32,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1346, c_0_1352])]), c_0_1348]), c_0_1349]), c_0_1350])).
cnf(c_0_1357,plain,(~node15(s31)|~node9(s32)), inference(spm,[status(thm)],[c_0_167, c_0_1351])).
cnf(c_0_1358,plain,(trans(s32,s33)), inference(spm,[status(thm)],[c_0_8, c_0_94])).
cnf(c_0_1359,plain,(~node15(s31)|~node1(s32)), inference(spm,[status(thm)],[c_0_203, c_0_1351])).
cnf(c_0_1360,plain,(~node15(s31)|~node5(s32)), inference(spm,[status(thm)],[c_0_354, c_0_1353])).
cnf(c_0_1361,plain,(~node15(s31)|~node3(s32)), inference(spm,[status(thm)],[c_0_201, c_0_1351])).
cnf(c_0_1362,plain,(~node15(s31)|~node2(s32)), inference(spm,[status(thm)],[c_0_202, c_0_1351])).
cnf(c_0_1363,plain,(~node15(s31)|~m_main_v_state1(s32,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_1351])).
cnf(c_0_1364,plain,(~node12(s32)), inference(spm,[status(thm)],[c_0_345, c_0_1354])).
cnf(c_0_1365,plain,(~node11(s32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1355, c_0_1352])])).
cnf(c_0_1366,plain,(~node10(s32)), inference(spm,[status(thm)],[c_0_339, c_0_1356])).
cnf(c_0_1367,plain,(~node9(s32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1357, c_0_1352])])).
cnf(c_0_1368,plain,(~node8(s32)), inference(spm,[status(thm)],[c_0_343, c_0_1356])).
cnf(c_0_1369,plain,(node6(s32,s33)|node5(s32)|node4(s32)|node3(s32)|node2(s32)|node1(s32)|m_main_v_state1(s32,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1358])).
cnf(c_0_1370,plain,(~node1(s32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1359, c_0_1352])])).
cnf(c_0_1371,plain,(~node5(s32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1360, c_0_1352])])).
cnf(c_0_1372,plain,(~node4(s32)), inference(spm,[status(thm)],[c_0_361, c_0_1356])).
cnf(c_0_1373,plain,(~node3(s32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1361, c_0_1352])])).
cnf(c_0_1374,plain,(~node2(s32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1362, c_0_1352])])).
cnf(c_0_1375,plain,(~m_main_v_state1(s32,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1363, c_0_1352])])).
cnf(c_0_1376,plain,(m_main_v_state2(X1,c_n2)|~node14(s32,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1356]), c_0_1364]), c_0_1365]), c_0_1366]), c_0_1367]), c_0_1368])).
cnf(c_0_1377,plain,(node6(s32,s33)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1369, c_0_1370]), c_0_1371]), c_0_1372]), c_0_1373]), c_0_1374]), c_0_1375])).
cnf(c_0_1378,plain,(node17(s32,s33)|node16(s32)|node15(s32)), inference(spm,[status(thm)],[c_0_152, c_0_1358])).
cnf(c_0_1379,plain,(m_main_v_state2(X1,c_n2)|~trans(s32,X1)), inference(spm,[status(thm)],[c_0_1376, c_0_29])).
cnf(c_0_1380,plain,(m_main_v_state1(s33,X1)|~m_main_v_state1(s32,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1377])).
cnf(c_0_1381,plain,(m_main_v_state1(s32,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1351, c_0_1352])])).
cnf(c_0_1382,plain,(node16(s32)|node15(s32)|m_main_v_turn(s33,X1)|~m_main_v_turn(s32,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1378])).
cnf(c_0_1383,plain,(trans(s33,s34)), inference(spm,[status(thm)],[c_0_8, c_0_95])).
cnf(c_0_1384,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s32,X1)), inference(spm,[status(thm)],[c_0_63, c_0_1379])).
cnf(c_0_1385,plain,(~node15(s33)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_1380]), c_0_1381])])).
cnf(c_0_1386,plain,(node16(s32)|node15(s32)|m_main_v_turn(s33,c_2)|~node15(s31)), inference(spm,[status(thm)],[c_0_1382, c_0_1353])).
cnf(c_0_1387,plain,(~node15(s32)|~node15(s31)), inference(spm,[status(thm)],[c_0_189, c_0_1351])).
cnf(c_0_1388,plain,(node17(s33,s34)|node16(s33)|node15(s33)), inference(spm,[status(thm)],[c_0_152, c_0_1383])).
cnf(c_0_1389,plain,(m_main_v_turn(X1,c_1)|~node18(s33,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1384, c_0_1380]), c_0_1358]), c_0_1381])]), c_0_1385])).
cnf(c_0_1390,plain,(node16(s32)|node15(s32)|m_main_v_turn(s33,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1386, c_0_1352])])).
cnf(c_0_1391,plain,(~node15(s32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1387, c_0_1352])])).
cnf(c_0_1392,plain,(node16(s33)|node15(s33)|m_main_v_turn(s33,X1)|~m_main_v_turn(s34,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1388])).
cnf(c_0_1393,plain,(m_main_v_turn(X1,c_1)|~trans(s33,X1)), inference(spm,[status(thm)],[c_0_1389, c_0_15])).
cnf(c_0_1394,plain,(node16(s32)|m_main_v_turn(s33,c_2)), inference(sr,[status(thm)],[c_0_1390, c_0_1391])).
cnf(c_0_1395,plain,(~node16(s32)), inference(spm,[status(thm)],[c_0_392, c_0_1356])).
cnf(c_0_1396,plain,(node16(s33)|m_main_v_turn(s33,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1392, c_0_1393]), c_0_1383])]), c_0_1385])).
cnf(c_0_1397,plain,(m_main_v_turn(s33,c_2)), inference(sr,[status(thm)],[c_0_1394, c_0_1395])).
cnf(c_0_1398,plain,(m_main_v_state2(s34,c_t2)|~node16(s33)), inference(spm,[status(thm)],[c_0_400, c_0_1383])).
cnf(c_0_1399,plain,(node3(s33)|node2(s33)|node1(s33)|m_main_v_state1(s34,c_c1)|~node16(s33)), inference(spm,[status(thm)],[c_0_401, c_0_1383])).
cnf(c_0_1400,plain,(node16(s33)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_1396]), c_0_1397])])).
cnf(c_0_1401,plain,(~node3(s33)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_1380]), c_0_1381])])).
cnf(c_0_1402,plain,(~node2(s33)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_1380]), c_0_1381])])).
cnf(c_0_1403,plain,(~node1(s33)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_1380]), c_0_1381])])).
cnf(c_0_1404,plain,(node16(s33)|node15(s33)|m_main_v_turn(s34,X1)|~m_main_v_turn(s33,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1388])).
cnf(c_0_1405,plain,(trans(s34,s35)), inference(spm,[status(thm)],[c_0_8, c_0_96])).
cnf(c_0_1406,plain,(~node16(s33)|~node8(s34)), inference(spm,[status(thm)],[c_0_168, c_0_1398])).
cnf(c_0_1407,plain,(m_main_v_state1(s34,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1399, c_0_1400])]), c_0_1401]), c_0_1402]), c_0_1403])).
cnf(c_0_1408,plain,(~node16(s33)|~node10(s34)), inference(spm,[status(thm)],[c_0_257, c_0_1398])).
cnf(c_0_1409,plain,(~node16(s33)|~node9(s34)), inference(spm,[status(thm)],[c_0_258, c_0_1398])).
cnf(c_0_1410,plain,(~node16(s33)|~m_main_v_state2(s34,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_1398])).
cnf(c_0_1411,plain,(node16(s33)|node16(s32)|node15(s33)|m_main_v_turn(s34,c_2)), inference(spm,[status(thm)],[c_0_1404, c_0_1394])).
cnf(c_0_1412,plain,(node17(s34,s35)|node16(s34)|node15(s34)), inference(spm,[status(thm)],[c_0_152, c_0_1405])).
cnf(c_0_1413,plain,(node15(s33)|m_main_v_turn(s34,c_1)|~node16(s33)), inference(spm,[status(thm)],[c_0_433, c_0_1383])).
cnf(c_0_1414,plain,(~node16(s34)|~node16(s33)), inference(spm,[status(thm)],[c_0_188, c_0_1398])).
cnf(c_0_1415,plain,(node13(s34,s35)|node12(s34)|node11(s34)|node10(s34)|node9(s34)|node8(s34)|m_main_v_state2(s34,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1405])).
cnf(c_0_1416,plain,(~node8(s34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1406, c_0_1400])])).
cnf(c_0_1417,plain,(~node12(s34)), inference(spm,[status(thm)],[c_0_428, c_0_1407])).
cnf(c_0_1418,plain,(~node11(s34)), inference(spm,[status(thm)],[c_0_429, c_0_1407])).
cnf(c_0_1419,plain,(~node10(s34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1408, c_0_1400])])).
cnf(c_0_1420,plain,(~node9(s34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1409, c_0_1400])])).
cnf(c_0_1421,plain,(~m_main_v_state2(s34,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1410, c_0_1400])])).
cnf(c_0_1422,plain,(node16(s33)|node15(s33)|m_main_v_turn(s34,c_2)), inference(sr,[status(thm)],[c_0_1411, c_0_1395])).
cnf(c_0_1423,plain,(trans(s35,s36)), inference(spm,[status(thm)],[c_0_8, c_0_97])).
cnf(c_0_1424,plain,(node16(s34)|node15(s34)|m_main_v_turn(s35,X1)|~m_main_v_turn(s34,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1412])).
cnf(c_0_1425,plain,(m_main_v_turn(s34,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1413, c_0_1400])]), c_0_1385])).
cnf(c_0_1426,plain,(~node16(s34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1414, c_0_1400])])).
cnf(c_0_1427,plain,(node13(s34,s35)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1415, c_0_1416]), c_0_1417]), c_0_1418]), c_0_1419]), c_0_1420]), c_0_1421])).
cnf(c_0_1428,plain,(node16(s33)|m_main_v_turn(s34,c_2)), inference(sr,[status(thm)],[c_0_1422, c_0_1385])).
cnf(c_0_1429,plain,(node17(s35,s36)|node16(s35)|node15(s35)), inference(spm,[status(thm)],[c_0_152, c_0_1423])).
cnf(c_0_1430,plain,(node15(s34)|m_main_v_turn(s35,c_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1424, c_0_1425]), c_0_1426])).
cnf(c_0_1431,plain,(~node15(s34)), inference(spm,[status(thm)],[c_0_229, c_0_1407])).
cnf(c_0_1432,plain,(m_main_v_state2(s35,X1)|~m_main_v_state2(s34,X1)), inference(spm,[status(thm)],[c_0_43, c_0_1427])).
cnf(c_0_1433,plain,(m_main_v_state2(s34,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1398, c_0_1400])])).
cnf(c_0_1434,plain,(node16(s33)|~node5(s34)), inference(spm,[status(thm)],[c_0_354, c_0_1428])).
cnf(c_0_1435,plain,(node16(s35)|node15(s35)|m_main_v_turn(s36,X1)|~m_main_v_turn(s35,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1429])).
cnf(c_0_1436,plain,(m_main_v_turn(s35,c_1)), inference(sr,[status(thm)],[c_0_1430, c_0_1431])).
cnf(c_0_1437,plain,(~node16(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_1432]), c_0_1433])])).
cnf(c_0_1438,plain,(m_main_v_state1(s34,c_c1)|~node5(s34)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1399, c_0_1434]), c_0_1401]), c_0_1402]), c_0_1403])).
cnf(c_0_1439,plain,(~node16(s33)|~node4(s34)), inference(spm,[status(thm)],[c_0_200, c_0_1398])).
cnf(c_0_1440,plain,(~node16(s33)|~node2(s34)), inference(spm,[status(thm)],[c_0_239, c_0_1398])).
cnf(c_0_1441,plain,(node16(s35)|node15(s35)|m_main_v_turn(s36,c_1)), inference(spm,[status(thm)],[c_0_1435, c_0_1436])).
cnf(c_0_1442,plain,(node17(s35,s36)|node15(s35)), inference(sr,[status(thm)],[c_0_1429, c_0_1437])).
cnf(c_0_1443,plain,(m_main_v_turn(X1,c_2)|~node18(s35,X1)|~m_main_v_state1(s35,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1432]), c_0_1433])])).
cnf(c_0_1444,plain,(~node5(s34)), inference(spm,[status(thm)],[c_0_233, c_0_1438])).
cnf(c_0_1445,plain,(~node4(s34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1439, c_0_1400])])).
cnf(c_0_1446,plain,(~node2(s34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1440, c_0_1400])])).
cnf(c_0_1447,plain,(~node3(s34)), inference(spm,[status(thm)],[c_0_415, c_0_1407])).
cnf(c_0_1448,plain,(~node1(s34)), inference(spm,[status(thm)],[c_0_416, c_0_1407])).
cnf(c_0_1449,plain,(node15(s35)|m_main_v_turn(s36,c_1)), inference(sr,[status(thm)],[c_0_1441, c_0_1437])).
cnf(c_0_1450,plain,(node15(s35)|m_main_v_turn(s35,X1)|~m_main_v_turn(s36,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1442])).
cnf(c_0_1451,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s35,c_n1)|~trans(s35,X1)), inference(spm,[status(thm)],[c_0_1443, c_0_15])).
cnf(c_0_1452,plain,(~m_main_v_turn(s35,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1436])).
cnf(c_0_1453,plain,(m_main_v_state1(X1,c_n1)|~node7(s34,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1407]), c_0_1444]), c_0_1445]), c_0_1446]), c_0_1447]), c_0_1448])).
cnf(c_0_1454,plain,(node15(s35)|~m_main_v_turn(s36,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1449])).
cnf(c_0_1455,plain,(node15(s35)|~m_main_v_state1(s35,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1450, c_0_1451]), c_0_1423])]), c_0_1452])).
cnf(c_0_1456,plain,(m_main_v_state1(X1,c_n1)|~trans(s34,X1)), inference(spm,[status(thm)],[c_0_1453, c_0_7])).
cnf(c_0_1457,plain,(node10(s35)|node9(s35)|node8(s35)|m_main_v_state2(s36,c_c2)|~node15(s35)), inference(spm,[status(thm)],[c_0_467, c_0_1423])).
cnf(c_0_1458,plain,(node15(s35)|~node12(s36)), inference(spm,[status(thm)],[c_0_1454, c_0_35])).
cnf(c_0_1459,plain,(~node10(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_1432]), c_0_1433])])).
cnf(c_0_1460,plain,(~node9(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_1432]), c_0_1433])])).
cnf(c_0_1461,plain,(~node8(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_1432]), c_0_1433])])).
cnf(c_0_1462,plain,(m_main_v_state1(s36,c_t1)|~node15(s35)), inference(spm,[status(thm)],[c_0_474, c_0_1423])).
cnf(c_0_1463,plain,(node15(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1455, c_0_1456]), c_0_1405])])).
cnf(c_0_1464,plain,(m_main_v_turn(s36,c_2)|~node15(s35)), inference(spm,[status(thm)],[c_0_470, c_0_1423])).
cnf(c_0_1465,plain,(m_main_v_state2(s36,c_c2)|~node12(s36)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1457, c_0_1458]), c_0_1459]), c_0_1460]), c_0_1461])).
cnf(c_0_1466,plain,(~node15(s35)|~node11(s36)), inference(spm,[status(thm)],[c_0_165, c_0_1462])).
cnf(c_0_1467,plain,(~node15(s35)|~node9(s36)), inference(spm,[status(thm)],[c_0_167, c_0_1462])).
cnf(c_0_1468,plain,(m_main_v_state2(s36,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1457, c_0_1463])]), c_0_1459]), c_0_1460]), c_0_1461])).
cnf(c_0_1469,plain,(trans(s36,s37)), inference(spm,[status(thm)],[c_0_8, c_0_98])).
cnf(c_0_1470,plain,(~node15(s35)|~node5(s36)), inference(spm,[status(thm)],[c_0_354, c_0_1464])).
cnf(c_0_1471,plain,(~node15(s35)|~node3(s36)), inference(spm,[status(thm)],[c_0_201, c_0_1462])).
cnf(c_0_1472,plain,(~node15(s35)|~node2(s36)), inference(spm,[status(thm)],[c_0_202, c_0_1462])).
cnf(c_0_1473,plain,(~node15(s35)|~node1(s36)), inference(spm,[status(thm)],[c_0_203, c_0_1462])).
cnf(c_0_1474,plain,(~node15(s35)|~m_main_v_state1(s36,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_1462])).
cnf(c_0_1475,plain,(~node12(s36)), inference(spm,[status(thm)],[c_0_345, c_0_1465])).
cnf(c_0_1476,plain,(~node11(s36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1466, c_0_1463])])).
cnf(c_0_1477,plain,(~node9(s36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1467, c_0_1463])])).
cnf(c_0_1478,plain,(~node10(s36)), inference(spm,[status(thm)],[c_0_339, c_0_1468])).
cnf(c_0_1479,plain,(~node8(s36)), inference(spm,[status(thm)],[c_0_343, c_0_1468])).
cnf(c_0_1480,plain,(node6(s36,s37)|node5(s36)|node4(s36)|node3(s36)|node2(s36)|node1(s36)|m_main_v_state1(s36,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1469])).
cnf(c_0_1481,plain,(~node5(s36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1470, c_0_1463])])).
cnf(c_0_1482,plain,(~node4(s36)), inference(spm,[status(thm)],[c_0_361, c_0_1468])).
cnf(c_0_1483,plain,(~node3(s36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1471, c_0_1463])])).
cnf(c_0_1484,plain,(~node2(s36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1472, c_0_1463])])).
cnf(c_0_1485,plain,(~node1(s36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1473, c_0_1463])])).
cnf(c_0_1486,plain,(~m_main_v_state1(s36,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1474, c_0_1463])])).
cnf(c_0_1487,plain,(m_main_v_state2(X1,c_n2)|~node14(s36,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1468]), c_0_1475]), c_0_1476]), c_0_1477]), c_0_1478]), c_0_1479])).
cnf(c_0_1488,plain,(node6(s36,s37)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1480, c_0_1481]), c_0_1482]), c_0_1483]), c_0_1484]), c_0_1485]), c_0_1486])).
cnf(c_0_1489,plain,(node17(s36,s37)|node16(s36)|node15(s36)), inference(spm,[status(thm)],[c_0_152, c_0_1469])).
cnf(c_0_1490,plain,(m_main_v_state2(X1,c_n2)|~trans(s36,X1)), inference(spm,[status(thm)],[c_0_1487, c_0_29])).
cnf(c_0_1491,plain,(m_main_v_state1(s37,X1)|~m_main_v_state1(s36,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1488])).
cnf(c_0_1492,plain,(m_main_v_state1(s36,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1462, c_0_1463])])).
cnf(c_0_1493,plain,(node16(s36)|node15(s36)|m_main_v_turn(s37,X1)|~m_main_v_turn(s36,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1489])).
cnf(c_0_1494,plain,(trans(s37,s38)), inference(spm,[status(thm)],[c_0_8, c_0_99])).
cnf(c_0_1495,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s36,X1)), inference(spm,[status(thm)],[c_0_63, c_0_1490])).
cnf(c_0_1496,plain,(~node15(s37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_1491]), c_0_1492])])).
cnf(c_0_1497,plain,(node16(s36)|node15(s36)|m_main_v_turn(s37,c_2)|~node15(s35)), inference(spm,[status(thm)],[c_0_1493, c_0_1464])).
cnf(c_0_1498,plain,(~node15(s36)|~node15(s35)), inference(spm,[status(thm)],[c_0_189, c_0_1462])).
cnf(c_0_1499,plain,(node17(s37,s38)|node16(s37)|node15(s37)), inference(spm,[status(thm)],[c_0_152, c_0_1494])).
cnf(c_0_1500,plain,(m_main_v_turn(X1,c_1)|~node18(s37,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1495, c_0_1491]), c_0_1469]), c_0_1492])]), c_0_1496])).
cnf(c_0_1501,plain,(node16(s36)|node15(s36)|m_main_v_turn(s37,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1497, c_0_1463])])).
cnf(c_0_1502,plain,(~node15(s36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1498, c_0_1463])])).
cnf(c_0_1503,plain,(node16(s37)|node15(s37)|m_main_v_turn(s37,X1)|~m_main_v_turn(s38,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1499])).
cnf(c_0_1504,plain,(m_main_v_turn(X1,c_1)|~trans(s37,X1)), inference(spm,[status(thm)],[c_0_1500, c_0_15])).
cnf(c_0_1505,plain,(node16(s36)|m_main_v_turn(s37,c_2)), inference(sr,[status(thm)],[c_0_1501, c_0_1502])).
cnf(c_0_1506,plain,(~node16(s36)), inference(spm,[status(thm)],[c_0_392, c_0_1468])).
cnf(c_0_1507,plain,(node16(s37)|m_main_v_turn(s37,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1503, c_0_1504]), c_0_1494])]), c_0_1496])).
cnf(c_0_1508,plain,(m_main_v_turn(s37,c_2)), inference(sr,[status(thm)],[c_0_1505, c_0_1506])).
cnf(c_0_1509,plain,(m_main_v_state2(s38,c_t2)|~node16(s37)), inference(spm,[status(thm)],[c_0_400, c_0_1494])).
cnf(c_0_1510,plain,(node3(s37)|node2(s37)|node1(s37)|m_main_v_state1(s38,c_c1)|~node16(s37)), inference(spm,[status(thm)],[c_0_401, c_0_1494])).
cnf(c_0_1511,plain,(node16(s37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_1507]), c_0_1508])])).
cnf(c_0_1512,plain,(~node3(s37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_1491]), c_0_1492])])).
cnf(c_0_1513,plain,(~node2(s37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_1491]), c_0_1492])])).
cnf(c_0_1514,plain,(~node1(s37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_1491]), c_0_1492])])).
cnf(c_0_1515,plain,(trans(s38,s39)), inference(spm,[status(thm)],[c_0_8, c_0_100])).
cnf(c_0_1516,plain,(~node16(s37)|~node8(s38)), inference(spm,[status(thm)],[c_0_168, c_0_1509])).
cnf(c_0_1517,plain,(m_main_v_state1(s38,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1510, c_0_1511])]), c_0_1512]), c_0_1513]), c_0_1514])).
cnf(c_0_1518,plain,(~node16(s37)|~node10(s38)), inference(spm,[status(thm)],[c_0_257, c_0_1509])).
cnf(c_0_1519,plain,(~node16(s37)|~node9(s38)), inference(spm,[status(thm)],[c_0_258, c_0_1509])).
cnf(c_0_1520,plain,(~node16(s37)|~m_main_v_state2(s38,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_1509])).
cnf(c_0_1521,plain,(node16(s37)|node15(s37)|m_main_v_turn(s38,X1)|~m_main_v_turn(s37,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1499])).
cnf(c_0_1522,plain,(node17(s38,s39)|node16(s38)|node15(s38)), inference(spm,[status(thm)],[c_0_152, c_0_1515])).
cnf(c_0_1523,plain,(node15(s37)|m_main_v_turn(s38,c_1)|~node16(s37)), inference(spm,[status(thm)],[c_0_433, c_0_1494])).
cnf(c_0_1524,plain,(~node16(s38)|~node16(s37)), inference(spm,[status(thm)],[c_0_188, c_0_1509])).
cnf(c_0_1525,plain,(node13(s38,s39)|node12(s38)|node11(s38)|node10(s38)|node9(s38)|node8(s38)|m_main_v_state2(s38,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1515])).
cnf(c_0_1526,plain,(~node8(s38)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1516, c_0_1511])])).
cnf(c_0_1527,plain,(~node12(s38)), inference(spm,[status(thm)],[c_0_428, c_0_1517])).
cnf(c_0_1528,plain,(~node11(s38)), inference(spm,[status(thm)],[c_0_429, c_0_1517])).
cnf(c_0_1529,plain,(~node10(s38)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1518, c_0_1511])])).
cnf(c_0_1530,plain,(~node9(s38)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1519, c_0_1511])])).
cnf(c_0_1531,plain,(~m_main_v_state2(s38,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1520, c_0_1511])])).
cnf(c_0_1532,plain,(node16(s37)|node15(s37)|m_main_v_turn(s38,c_2)), inference(spm,[status(thm)],[c_0_1521, c_0_1508])).
cnf(c_0_1533,plain,(trans(s39,s40)), inference(spm,[status(thm)],[c_0_8, c_0_101])).
cnf(c_0_1534,plain,(node16(s38)|node15(s38)|m_main_v_turn(s39,X1)|~m_main_v_turn(s38,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1522])).
cnf(c_0_1535,plain,(m_main_v_turn(s38,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1523, c_0_1511])]), c_0_1496])).
cnf(c_0_1536,plain,(~node16(s38)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1524, c_0_1511])])).
cnf(c_0_1537,plain,(node13(s38,s39)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1525, c_0_1526]), c_0_1527]), c_0_1528]), c_0_1529]), c_0_1530]), c_0_1531])).
cnf(c_0_1538,plain,(node16(s37)|m_main_v_turn(s38,c_2)), inference(sr,[status(thm)],[c_0_1532, c_0_1496])).
cnf(c_0_1539,plain,(node17(s39,s40)|node16(s39)|node15(s39)), inference(spm,[status(thm)],[c_0_152, c_0_1533])).
cnf(c_0_1540,plain,(node15(s38)|m_main_v_turn(s39,c_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1534, c_0_1535]), c_0_1536])).
cnf(c_0_1541,plain,(~node15(s38)), inference(spm,[status(thm)],[c_0_229, c_0_1517])).
cnf(c_0_1542,plain,(m_main_v_state2(s39,X1)|~m_main_v_state2(s38,X1)), inference(spm,[status(thm)],[c_0_43, c_0_1537])).
cnf(c_0_1543,plain,(m_main_v_state2(s38,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1509, c_0_1511])])).
cnf(c_0_1544,plain,(node16(s37)|~node5(s38)), inference(spm,[status(thm)],[c_0_354, c_0_1538])).
cnf(c_0_1545,plain,(node16(s39)|node15(s39)|m_main_v_turn(s40,X1)|~m_main_v_turn(s39,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1539])).
cnf(c_0_1546,plain,(m_main_v_turn(s39,c_1)), inference(sr,[status(thm)],[c_0_1540, c_0_1541])).
cnf(c_0_1547,plain,(~node16(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_1542]), c_0_1543])])).
cnf(c_0_1548,plain,(m_main_v_state1(s38,c_c1)|~node5(s38)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1510, c_0_1544]), c_0_1512]), c_0_1513]), c_0_1514])).
cnf(c_0_1549,plain,(~node16(s37)|~node4(s38)), inference(spm,[status(thm)],[c_0_200, c_0_1509])).
cnf(c_0_1550,plain,(~node16(s37)|~node2(s38)), inference(spm,[status(thm)],[c_0_239, c_0_1509])).
cnf(c_0_1551,plain,(node16(s39)|node15(s39)|m_main_v_turn(s40,c_1)), inference(spm,[status(thm)],[c_0_1545, c_0_1546])).
cnf(c_0_1552,plain,(node17(s39,s40)|node15(s39)), inference(sr,[status(thm)],[c_0_1539, c_0_1547])).
cnf(c_0_1553,plain,(m_main_v_turn(X1,c_2)|~node18(s39,X1)|~m_main_v_state1(s39,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1542]), c_0_1543])])).
cnf(c_0_1554,plain,(~node5(s38)), inference(spm,[status(thm)],[c_0_233, c_0_1548])).
cnf(c_0_1555,plain,(~node4(s38)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1549, c_0_1511])])).
cnf(c_0_1556,plain,(~node3(s38)), inference(spm,[status(thm)],[c_0_415, c_0_1517])).
cnf(c_0_1557,plain,(~node2(s38)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1550, c_0_1511])])).
cnf(c_0_1558,plain,(~node1(s38)), inference(spm,[status(thm)],[c_0_416, c_0_1517])).
cnf(c_0_1559,plain,(node15(s39)|m_main_v_turn(s40,c_1)), inference(sr,[status(thm)],[c_0_1551, c_0_1547])).
cnf(c_0_1560,plain,(node15(s39)|m_main_v_turn(s39,X1)|~m_main_v_turn(s40,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1552])).
cnf(c_0_1561,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s39,c_n1)|~trans(s39,X1)), inference(spm,[status(thm)],[c_0_1553, c_0_15])).
cnf(c_0_1562,plain,(~m_main_v_turn(s39,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1546])).
cnf(c_0_1563,plain,(m_main_v_state1(X1,c_n1)|~node7(s38,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1517]), c_0_1554]), c_0_1555]), c_0_1556]), c_0_1557]), c_0_1558])).
cnf(c_0_1564,plain,(node15(s39)|~m_main_v_turn(s40,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1559])).
cnf(c_0_1565,plain,(node15(s39)|~m_main_v_state1(s39,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1560, c_0_1561]), c_0_1533])]), c_0_1562])).
cnf(c_0_1566,plain,(m_main_v_state1(X1,c_n1)|~trans(s38,X1)), inference(spm,[status(thm)],[c_0_1563, c_0_7])).
cnf(c_0_1567,plain,(node10(s39)|node9(s39)|node8(s39)|m_main_v_state2(s40,c_c2)|~node15(s39)), inference(spm,[status(thm)],[c_0_467, c_0_1533])).
cnf(c_0_1568,plain,(node15(s39)|~node12(s40)), inference(spm,[status(thm)],[c_0_1564, c_0_35])).
cnf(c_0_1569,plain,(~node10(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_1542]), c_0_1543])])).
cnf(c_0_1570,plain,(~node9(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_1542]), c_0_1543])])).
cnf(c_0_1571,plain,(~node8(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_1542]), c_0_1543])])).
cnf(c_0_1572,plain,(m_main_v_state1(s40,c_t1)|~node15(s39)), inference(spm,[status(thm)],[c_0_474, c_0_1533])).
cnf(c_0_1573,plain,(node15(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1565, c_0_1566]), c_0_1515])])).
cnf(c_0_1574,plain,(m_main_v_turn(s40,c_2)|~node15(s39)), inference(spm,[status(thm)],[c_0_470, c_0_1533])).
cnf(c_0_1575,plain,(m_main_v_state2(s40,c_c2)|~node12(s40)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1567, c_0_1568]), c_0_1569]), c_0_1570]), c_0_1571])).
cnf(c_0_1576,plain,(~node15(s39)|~node11(s40)), inference(spm,[status(thm)],[c_0_165, c_0_1572])).
cnf(c_0_1577,plain,(~node15(s39)|~node9(s40)), inference(spm,[status(thm)],[c_0_167, c_0_1572])).
cnf(c_0_1578,plain,(m_main_v_state2(s40,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1567, c_0_1573])]), c_0_1569]), c_0_1570]), c_0_1571])).
cnf(c_0_1579,plain,(trans(s40,s41)), inference(spm,[status(thm)],[c_0_8, c_0_102])).
cnf(c_0_1580,plain,(~node15(s39)|~node1(s40)), inference(spm,[status(thm)],[c_0_203, c_0_1572])).
cnf(c_0_1581,plain,(~node15(s39)|~node5(s40)), inference(spm,[status(thm)],[c_0_354, c_0_1574])).
cnf(c_0_1582,plain,(~node15(s39)|~node3(s40)), inference(spm,[status(thm)],[c_0_201, c_0_1572])).
cnf(c_0_1583,plain,(~node15(s39)|~node2(s40)), inference(spm,[status(thm)],[c_0_202, c_0_1572])).
cnf(c_0_1584,plain,(~node15(s39)|~m_main_v_state1(s40,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_1572])).
cnf(c_0_1585,plain,(~node12(s40)), inference(spm,[status(thm)],[c_0_345, c_0_1575])).
cnf(c_0_1586,plain,(~node11(s40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1576, c_0_1573])])).
cnf(c_0_1587,plain,(~node9(s40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1577, c_0_1573])])).
cnf(c_0_1588,plain,(~node10(s40)), inference(spm,[status(thm)],[c_0_339, c_0_1578])).
cnf(c_0_1589,plain,(~node8(s40)), inference(spm,[status(thm)],[c_0_343, c_0_1578])).
cnf(c_0_1590,plain,(node6(s40,s41)|node5(s40)|node4(s40)|node3(s40)|node2(s40)|node1(s40)|m_main_v_state1(s40,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1579])).
cnf(c_0_1591,plain,(~node1(s40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1580, c_0_1573])])).
cnf(c_0_1592,plain,(~node5(s40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1581, c_0_1573])])).
cnf(c_0_1593,plain,(~node4(s40)), inference(spm,[status(thm)],[c_0_361, c_0_1578])).
cnf(c_0_1594,plain,(~node3(s40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1582, c_0_1573])])).
cnf(c_0_1595,plain,(~node2(s40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1583, c_0_1573])])).
cnf(c_0_1596,plain,(~m_main_v_state1(s40,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1584, c_0_1573])])).
cnf(c_0_1597,plain,(m_main_v_state2(X1,c_n2)|~node14(s40,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1578]), c_0_1585]), c_0_1586]), c_0_1587]), c_0_1588]), c_0_1589])).
cnf(c_0_1598,plain,(node6(s40,s41)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1590, c_0_1591]), c_0_1592]), c_0_1593]), c_0_1594]), c_0_1595]), c_0_1596])).
cnf(c_0_1599,plain,(m_main_v_state2(X1,c_n2)|~trans(s40,X1)), inference(spm,[status(thm)],[c_0_1597, c_0_29])).
cnf(c_0_1600,plain,(m_main_v_state1(s41,X1)|~m_main_v_state1(s40,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1598])).
cnf(c_0_1601,plain,(m_main_v_state1(s40,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1572, c_0_1573])])).
cnf(c_0_1602,plain,(node17(s40,s41)|node16(s40)|node15(s40)), inference(spm,[status(thm)],[c_0_152, c_0_1579])).
cnf(c_0_1603,plain,(trans(s41,s42)), inference(spm,[status(thm)],[c_0_8, c_0_103])).
cnf(c_0_1604,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s40,X1)), inference(spm,[status(thm)],[c_0_63, c_0_1599])).
cnf(c_0_1605,plain,(~node15(s41)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_1600]), c_0_1601])])).
cnf(c_0_1606,plain,(node16(s40)|node15(s40)|m_main_v_turn(s41,X1)|~m_main_v_turn(s40,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1602])).
cnf(c_0_1607,plain,(~node15(s40)|~node15(s39)), inference(spm,[status(thm)],[c_0_189, c_0_1572])).
cnf(c_0_1608,plain,(node17(s41,s42)|node16(s41)|node15(s41)), inference(spm,[status(thm)],[c_0_152, c_0_1603])).
cnf(c_0_1609,plain,(m_main_v_turn(X1,c_1)|~node18(s41,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1604, c_0_1600]), c_0_1579]), c_0_1601])]), c_0_1605])).
cnf(c_0_1610,plain,(node16(s40)|node15(s40)|m_main_v_turn(s41,c_2)|~node15(s39)), inference(spm,[status(thm)],[c_0_1606, c_0_1574])).
cnf(c_0_1611,plain,(~node15(s40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1607, c_0_1573])])).
cnf(c_0_1612,plain,(node16(s41)|node15(s41)|m_main_v_turn(s41,X1)|~m_main_v_turn(s42,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1608])).
cnf(c_0_1613,plain,(m_main_v_turn(X1,c_1)|~trans(s41,X1)), inference(spm,[status(thm)],[c_0_1609, c_0_15])).
cnf(c_0_1614,plain,(node16(s40)|m_main_v_turn(s41,c_2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1610, c_0_1573])]), c_0_1611])).
cnf(c_0_1615,plain,(~node16(s40)), inference(spm,[status(thm)],[c_0_392, c_0_1578])).
cnf(c_0_1616,plain,(node16(s41)|m_main_v_turn(s41,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1612, c_0_1613]), c_0_1603])]), c_0_1605])).
cnf(c_0_1617,plain,(m_main_v_turn(s41,c_2)), inference(sr,[status(thm)],[c_0_1614, c_0_1615])).
cnf(c_0_1618,plain,(trans(s42,s43)), inference(spm,[status(thm)],[c_0_8, c_0_104])).
cnf(c_0_1619,plain,(m_main_v_state2(s42,c_t2)|~node16(s41)), inference(spm,[status(thm)],[c_0_400, c_0_1603])).
cnf(c_0_1620,plain,(node3(s41)|node2(s41)|node1(s41)|m_main_v_state1(s42,c_c1)|~node16(s41)), inference(spm,[status(thm)],[c_0_401, c_0_1603])).
cnf(c_0_1621,plain,(node16(s41)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_1616]), c_0_1617])])).
cnf(c_0_1622,plain,(~node3(s41)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_1600]), c_0_1601])])).
cnf(c_0_1623,plain,(~node2(s41)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_1600]), c_0_1601])])).
cnf(c_0_1624,plain,(~node1(s41)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_1600]), c_0_1601])])).
cnf(c_0_1625,plain,(node17(s42,s43)|node16(s42)|node15(s42)), inference(spm,[status(thm)],[c_0_152, c_0_1618])).
cnf(c_0_1626,plain,(~node16(s41)|~node8(s42)), inference(spm,[status(thm)],[c_0_168, c_0_1619])).
cnf(c_0_1627,plain,(m_main_v_state1(s42,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1620, c_0_1621])]), c_0_1622]), c_0_1623]), c_0_1624])).
cnf(c_0_1628,plain,(~node16(s41)|~node10(s42)), inference(spm,[status(thm)],[c_0_257, c_0_1619])).
cnf(c_0_1629,plain,(~node16(s41)|~node9(s42)), inference(spm,[status(thm)],[c_0_258, c_0_1619])).
cnf(c_0_1630,plain,(~node16(s41)|~m_main_v_state2(s42,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_1619])).
cnf(c_0_1631,plain,(node16(s41)|node15(s41)|m_main_v_turn(s42,X1)|~m_main_v_turn(s41,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1608])).
cnf(c_0_1632,plain,(node16(s42)|node15(s42)|m_main_v_turn(s43,X1)|~m_main_v_turn(s42,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1625])).
cnf(c_0_1633,plain,(~node16(s42)|~node16(s41)), inference(spm,[status(thm)],[c_0_188, c_0_1619])).
cnf(c_0_1634,plain,(node13(s42,s43)|node12(s42)|node11(s42)|node10(s42)|node9(s42)|node8(s42)|m_main_v_state2(s42,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1618])).
cnf(c_0_1635,plain,(~node8(s42)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1626, c_0_1621])])).
cnf(c_0_1636,plain,(~node12(s42)), inference(spm,[status(thm)],[c_0_428, c_0_1627])).
cnf(c_0_1637,plain,(~node11(s42)), inference(spm,[status(thm)],[c_0_429, c_0_1627])).
cnf(c_0_1638,plain,(~node10(s42)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1628, c_0_1621])])).
cnf(c_0_1639,plain,(~node9(s42)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1629, c_0_1621])])).
cnf(c_0_1640,plain,(~m_main_v_state2(s42,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1630, c_0_1621])])).
cnf(c_0_1641,plain,(node16(s41)|node15(s41)|m_main_v_turn(s42,c_2)), inference(spm,[status(thm)],[c_0_1631, c_0_1617])).
cnf(c_0_1642,plain,(trans(s43,s44)), inference(spm,[status(thm)],[c_0_8, c_0_105])).
cnf(c_0_1643,plain,(node16(s42)|node15(s42)|m_main_v_turn(s43,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1632, c_0_1613]), c_0_1603])])).
cnf(c_0_1644,plain,(~node16(s42)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1633, c_0_1621])])).
cnf(c_0_1645,plain,(node13(s42,s43)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1634, c_0_1635]), c_0_1636]), c_0_1637]), c_0_1638]), c_0_1639]), c_0_1640])).
cnf(c_0_1646,plain,(node16(s41)|m_main_v_turn(s42,c_2)), inference(sr,[status(thm)],[c_0_1641, c_0_1605])).
cnf(c_0_1647,plain,(node17(s43,s44)|node16(s43)|node15(s43)), inference(spm,[status(thm)],[c_0_152, c_0_1642])).
cnf(c_0_1648,plain,(node15(s42)|m_main_v_turn(s43,c_1)), inference(sr,[status(thm)],[c_0_1643, c_0_1644])).
cnf(c_0_1649,plain,(~node15(s42)), inference(spm,[status(thm)],[c_0_229, c_0_1627])).
cnf(c_0_1650,plain,(m_main_v_state2(s43,X1)|~m_main_v_state2(s42,X1)), inference(spm,[status(thm)],[c_0_43, c_0_1645])).
cnf(c_0_1651,plain,(m_main_v_state2(s42,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1619, c_0_1621])])).
cnf(c_0_1652,plain,(node16(s41)|~node5(s42)), inference(spm,[status(thm)],[c_0_354, c_0_1646])).
cnf(c_0_1653,plain,(node16(s43)|node15(s43)|m_main_v_turn(s44,X1)|~m_main_v_turn(s43,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1647])).
cnf(c_0_1654,plain,(m_main_v_turn(s43,c_1)), inference(sr,[status(thm)],[c_0_1648, c_0_1649])).
cnf(c_0_1655,plain,(~node16(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_1650]), c_0_1651])])).
cnf(c_0_1656,plain,(m_main_v_state1(s42,c_c1)|~node5(s42)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1620, c_0_1652]), c_0_1622]), c_0_1623]), c_0_1624])).
cnf(c_0_1657,plain,(~node16(s41)|~node4(s42)), inference(spm,[status(thm)],[c_0_200, c_0_1619])).
cnf(c_0_1658,plain,(~node16(s41)|~node2(s42)), inference(spm,[status(thm)],[c_0_239, c_0_1619])).
cnf(c_0_1659,plain,(node16(s43)|node15(s43)|m_main_v_turn(s44,c_1)), inference(spm,[status(thm)],[c_0_1653, c_0_1654])).
cnf(c_0_1660,plain,(node17(s43,s44)|node15(s43)), inference(sr,[status(thm)],[c_0_1647, c_0_1655])).
cnf(c_0_1661,plain,(m_main_v_turn(X1,c_2)|~node18(s43,X1)|~m_main_v_state1(s43,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1650]), c_0_1651])])).
cnf(c_0_1662,plain,(~node5(s42)), inference(spm,[status(thm)],[c_0_233, c_0_1656])).
cnf(c_0_1663,plain,(~node4(s42)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1657, c_0_1621])])).
cnf(c_0_1664,plain,(~node3(s42)), inference(spm,[status(thm)],[c_0_415, c_0_1627])).
cnf(c_0_1665,plain,(~node2(s42)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1658, c_0_1621])])).
cnf(c_0_1666,plain,(~node1(s42)), inference(spm,[status(thm)],[c_0_416, c_0_1627])).
cnf(c_0_1667,plain,(node15(s43)|m_main_v_turn(s44,c_1)), inference(sr,[status(thm)],[c_0_1659, c_0_1655])).
cnf(c_0_1668,plain,(node15(s43)|m_main_v_turn(s43,X1)|~m_main_v_turn(s44,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1660])).
cnf(c_0_1669,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s43,c_n1)|~trans(s43,X1)), inference(spm,[status(thm)],[c_0_1661, c_0_15])).
cnf(c_0_1670,plain,(~m_main_v_turn(s43,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1654])).
cnf(c_0_1671,plain,(m_main_v_state1(X1,c_n1)|~node7(s42,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1627]), c_0_1662]), c_0_1663]), c_0_1664]), c_0_1665]), c_0_1666])).
cnf(c_0_1672,plain,(node15(s43)|~m_main_v_turn(s44,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1667])).
cnf(c_0_1673,plain,(node15(s43)|~m_main_v_state1(s43,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1668, c_0_1669]), c_0_1642])]), c_0_1670])).
cnf(c_0_1674,plain,(m_main_v_state1(X1,c_n1)|~trans(s42,X1)), inference(spm,[status(thm)],[c_0_1671, c_0_7])).
cnf(c_0_1675,plain,(node10(s43)|node9(s43)|node8(s43)|m_main_v_state2(s44,c_c2)|~node15(s43)), inference(spm,[status(thm)],[c_0_467, c_0_1642])).
cnf(c_0_1676,plain,(node15(s43)|~node12(s44)), inference(spm,[status(thm)],[c_0_1672, c_0_35])).
cnf(c_0_1677,plain,(~node10(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_1650]), c_0_1651])])).
cnf(c_0_1678,plain,(~node9(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_1650]), c_0_1651])])).
cnf(c_0_1679,plain,(~node8(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_1650]), c_0_1651])])).
cnf(c_0_1680,plain,(m_main_v_state1(s44,c_t1)|~node15(s43)), inference(spm,[status(thm)],[c_0_474, c_0_1642])).
cnf(c_0_1681,plain,(node15(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1673, c_0_1674]), c_0_1618])])).
cnf(c_0_1682,plain,(m_main_v_turn(s44,c_2)|~node15(s43)), inference(spm,[status(thm)],[c_0_470, c_0_1642])).
cnf(c_0_1683,plain,(m_main_v_state2(s44,c_c2)|~node12(s44)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1675, c_0_1676]), c_0_1677]), c_0_1678]), c_0_1679])).
cnf(c_0_1684,plain,(~node15(s43)|~node11(s44)), inference(spm,[status(thm)],[c_0_165, c_0_1680])).
cnf(c_0_1685,plain,(m_main_v_state2(s44,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1675, c_0_1681])]), c_0_1677]), c_0_1678]), c_0_1679])).
cnf(c_0_1686,plain,(~node15(s43)|~node9(s44)), inference(spm,[status(thm)],[c_0_167, c_0_1680])).
cnf(c_0_1687,plain,(trans(s44,s45)), inference(spm,[status(thm)],[c_0_8, c_0_106])).
cnf(c_0_1688,plain,(~node15(s43)|~node5(s44)), inference(spm,[status(thm)],[c_0_354, c_0_1682])).
cnf(c_0_1689,plain,(~node15(s43)|~node3(s44)), inference(spm,[status(thm)],[c_0_201, c_0_1680])).
cnf(c_0_1690,plain,(~node15(s43)|~node2(s44)), inference(spm,[status(thm)],[c_0_202, c_0_1680])).
cnf(c_0_1691,plain,(~node15(s43)|~node1(s44)), inference(spm,[status(thm)],[c_0_203, c_0_1680])).
cnf(c_0_1692,plain,(~node15(s43)|~m_main_v_state1(s44,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_1680])).
cnf(c_0_1693,plain,(~node12(s44)), inference(spm,[status(thm)],[c_0_345, c_0_1683])).
cnf(c_0_1694,plain,(~node11(s44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1684, c_0_1681])])).
cnf(c_0_1695,plain,(~node10(s44)), inference(spm,[status(thm)],[c_0_339, c_0_1685])).
cnf(c_0_1696,plain,(~node9(s44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1686, c_0_1681])])).
cnf(c_0_1697,plain,(~node8(s44)), inference(spm,[status(thm)],[c_0_343, c_0_1685])).
cnf(c_0_1698,plain,(node6(s44,s45)|node5(s44)|node4(s44)|node3(s44)|node2(s44)|node1(s44)|m_main_v_state1(s44,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1687])).
cnf(c_0_1699,plain,(~node5(s44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1688, c_0_1681])])).
cnf(c_0_1700,plain,(~node4(s44)), inference(spm,[status(thm)],[c_0_361, c_0_1685])).
cnf(c_0_1701,plain,(~node3(s44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1689, c_0_1681])])).
cnf(c_0_1702,plain,(~node2(s44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1690, c_0_1681])])).
cnf(c_0_1703,plain,(~node1(s44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1691, c_0_1681])])).
cnf(c_0_1704,plain,(~m_main_v_state1(s44,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1692, c_0_1681])])).
cnf(c_0_1705,plain,(m_main_v_state2(X1,c_n2)|~node14(s44,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1685]), c_0_1693]), c_0_1694]), c_0_1695]), c_0_1696]), c_0_1697])).
cnf(c_0_1706,plain,(node6(s44,s45)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1698, c_0_1699]), c_0_1700]), c_0_1701]), c_0_1702]), c_0_1703]), c_0_1704])).
cnf(c_0_1707,plain,(node17(s44,s45)|node16(s44)|node15(s44)), inference(spm,[status(thm)],[c_0_152, c_0_1687])).
cnf(c_0_1708,plain,(m_main_v_state2(X1,c_n2)|~trans(s44,X1)), inference(spm,[status(thm)],[c_0_1705, c_0_29])).
cnf(c_0_1709,plain,(m_main_v_state1(s45,X1)|~m_main_v_state1(s44,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1706])).
cnf(c_0_1710,plain,(m_main_v_state1(s44,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1680, c_0_1681])])).
cnf(c_0_1711,plain,(node16(s44)|node15(s44)|m_main_v_turn(s45,X1)|~m_main_v_turn(s44,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1707])).
cnf(c_0_1712,plain,(trans(s45,s46)), inference(spm,[status(thm)],[c_0_8, c_0_107])).
cnf(c_0_1713,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s44,X1)), inference(spm,[status(thm)],[c_0_63, c_0_1708])).
cnf(c_0_1714,plain,(~node15(s45)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_1709]), c_0_1710])])).
cnf(c_0_1715,plain,(node16(s44)|node15(s44)|m_main_v_turn(s45,c_2)|~node15(s43)), inference(spm,[status(thm)],[c_0_1711, c_0_1682])).
cnf(c_0_1716,plain,(~node15(s44)|~node15(s43)), inference(spm,[status(thm)],[c_0_189, c_0_1680])).
cnf(c_0_1717,plain,(node17(s45,s46)|node16(s45)|node15(s45)), inference(spm,[status(thm)],[c_0_152, c_0_1712])).
cnf(c_0_1718,plain,(m_main_v_turn(X1,c_1)|~node18(s45,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1713, c_0_1709]), c_0_1687]), c_0_1710])]), c_0_1714])).
cnf(c_0_1719,plain,(node16(s44)|node15(s44)|m_main_v_turn(s45,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1715, c_0_1681])])).
cnf(c_0_1720,plain,(~node15(s44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1716, c_0_1681])])).
cnf(c_0_1721,plain,(node16(s45)|node15(s45)|m_main_v_turn(s45,X1)|~m_main_v_turn(s46,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1717])).
cnf(c_0_1722,plain,(m_main_v_turn(X1,c_1)|~trans(s45,X1)), inference(spm,[status(thm)],[c_0_1718, c_0_15])).
cnf(c_0_1723,plain,(node16(s44)|m_main_v_turn(s45,c_2)), inference(sr,[status(thm)],[c_0_1719, c_0_1720])).
cnf(c_0_1724,plain,(~node16(s44)), inference(spm,[status(thm)],[c_0_392, c_0_1685])).
cnf(c_0_1725,plain,(node16(s45)|m_main_v_turn(s45,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1721, c_0_1722]), c_0_1712])]), c_0_1714])).
cnf(c_0_1726,plain,(m_main_v_turn(s45,c_2)), inference(sr,[status(thm)],[c_0_1723, c_0_1724])).
cnf(c_0_1727,plain,(trans(s46,s47)), inference(spm,[status(thm)],[c_0_8, c_0_108])).
cnf(c_0_1728,plain,(m_main_v_state2(s46,c_t2)|~node16(s45)), inference(spm,[status(thm)],[c_0_400, c_0_1712])).
cnf(c_0_1729,plain,(node3(s45)|node2(s45)|node1(s45)|m_main_v_state1(s46,c_c1)|~node16(s45)), inference(spm,[status(thm)],[c_0_401, c_0_1712])).
cnf(c_0_1730,plain,(node16(s45)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_1725]), c_0_1726])])).
cnf(c_0_1731,plain,(~node3(s45)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_1709]), c_0_1710])])).
cnf(c_0_1732,plain,(~node2(s45)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_1709]), c_0_1710])])).
cnf(c_0_1733,plain,(~node1(s45)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_1709]), c_0_1710])])).
cnf(c_0_1734,plain,(node16(s45)|node15(s45)|m_main_v_turn(s46,X1)|~m_main_v_turn(s45,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1717])).
cnf(c_0_1735,plain,(node17(s46,s47)|node16(s46)|node15(s46)), inference(spm,[status(thm)],[c_0_152, c_0_1727])).
cnf(c_0_1736,plain,(~node16(s45)|~node8(s46)), inference(spm,[status(thm)],[c_0_168, c_0_1728])).
cnf(c_0_1737,plain,(m_main_v_state1(s46,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1729, c_0_1730])]), c_0_1731]), c_0_1732]), c_0_1733])).
cnf(c_0_1738,plain,(~node16(s45)|~node10(s46)), inference(spm,[status(thm)],[c_0_257, c_0_1728])).
cnf(c_0_1739,plain,(~node16(s45)|~node9(s46)), inference(spm,[status(thm)],[c_0_258, c_0_1728])).
cnf(c_0_1740,plain,(~node16(s45)|~m_main_v_state2(s46,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_1728])).
cnf(c_0_1741,plain,(node16(s45)|node16(s44)|node15(s45)|m_main_v_turn(s46,c_2)), inference(spm,[status(thm)],[c_0_1734, c_0_1723])).
cnf(c_0_1742,plain,(node16(s46)|node15(s46)|m_main_v_turn(s47,X1)|~m_main_v_turn(s46,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1735])).
cnf(c_0_1743,plain,(~node16(s46)|~node16(s45)), inference(spm,[status(thm)],[c_0_188, c_0_1728])).
cnf(c_0_1744,plain,(node13(s46,s47)|node12(s46)|node11(s46)|node10(s46)|node9(s46)|node8(s46)|m_main_v_state2(s46,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1727])).
cnf(c_0_1745,plain,(~node8(s46)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1736, c_0_1730])])).
cnf(c_0_1746,plain,(~node12(s46)), inference(spm,[status(thm)],[c_0_428, c_0_1737])).
cnf(c_0_1747,plain,(~node11(s46)), inference(spm,[status(thm)],[c_0_429, c_0_1737])).
cnf(c_0_1748,plain,(~node10(s46)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1738, c_0_1730])])).
cnf(c_0_1749,plain,(~node9(s46)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1739, c_0_1730])])).
cnf(c_0_1750,plain,(~m_main_v_state2(s46,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1740, c_0_1730])])).
cnf(c_0_1751,plain,(node16(s45)|node15(s45)|m_main_v_turn(s46,c_2)), inference(sr,[status(thm)],[c_0_1741, c_0_1724])).
cnf(c_0_1752,plain,(trans(s47,s48)), inference(spm,[status(thm)],[c_0_8, c_0_109])).
cnf(c_0_1753,plain,(node16(s46)|node15(s46)|m_main_v_turn(s47,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1742, c_0_1722]), c_0_1712])])).
cnf(c_0_1754,plain,(~node16(s46)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1743, c_0_1730])])).
cnf(c_0_1755,plain,(node13(s46,s47)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1744, c_0_1745]), c_0_1746]), c_0_1747]), c_0_1748]), c_0_1749]), c_0_1750])).
cnf(c_0_1756,plain,(node16(s45)|m_main_v_turn(s46,c_2)), inference(sr,[status(thm)],[c_0_1751, c_0_1714])).
cnf(c_0_1757,plain,(node17(s47,s48)|node16(s47)|node15(s47)), inference(spm,[status(thm)],[c_0_152, c_0_1752])).
cnf(c_0_1758,plain,(node15(s46)|m_main_v_turn(s47,c_1)), inference(sr,[status(thm)],[c_0_1753, c_0_1754])).
cnf(c_0_1759,plain,(~node15(s46)), inference(spm,[status(thm)],[c_0_229, c_0_1737])).
cnf(c_0_1760,plain,(m_main_v_state2(s47,X1)|~m_main_v_state2(s46,X1)), inference(spm,[status(thm)],[c_0_43, c_0_1755])).
cnf(c_0_1761,plain,(m_main_v_state2(s46,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1728, c_0_1730])])).
cnf(c_0_1762,plain,(node16(s45)|~node5(s46)), inference(spm,[status(thm)],[c_0_354, c_0_1756])).
cnf(c_0_1763,plain,(node16(s47)|node15(s47)|m_main_v_turn(s48,X1)|~m_main_v_turn(s47,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1757])).
cnf(c_0_1764,plain,(m_main_v_turn(s47,c_1)), inference(sr,[status(thm)],[c_0_1758, c_0_1759])).
cnf(c_0_1765,plain,(~node16(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_1760]), c_0_1761])])).
cnf(c_0_1766,plain,(m_main_v_state1(s46,c_c1)|~node5(s46)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1729, c_0_1762]), c_0_1731]), c_0_1732]), c_0_1733])).
cnf(c_0_1767,plain,(~node16(s45)|~node4(s46)), inference(spm,[status(thm)],[c_0_200, c_0_1728])).
cnf(c_0_1768,plain,(~node16(s45)|~node2(s46)), inference(spm,[status(thm)],[c_0_239, c_0_1728])).
cnf(c_0_1769,plain,(node16(s47)|node15(s47)|m_main_v_turn(s48,c_1)), inference(spm,[status(thm)],[c_0_1763, c_0_1764])).
cnf(c_0_1770,plain,(node17(s47,s48)|node15(s47)), inference(sr,[status(thm)],[c_0_1757, c_0_1765])).
cnf(c_0_1771,plain,(m_main_v_turn(X1,c_2)|~node18(s47,X1)|~m_main_v_state1(s47,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1760]), c_0_1761])])).
cnf(c_0_1772,plain,(~node5(s46)), inference(spm,[status(thm)],[c_0_233, c_0_1766])).
cnf(c_0_1773,plain,(~node4(s46)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1767, c_0_1730])])).
cnf(c_0_1774,plain,(~node3(s46)), inference(spm,[status(thm)],[c_0_415, c_0_1737])).
cnf(c_0_1775,plain,(~node2(s46)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1768, c_0_1730])])).
cnf(c_0_1776,plain,(~node1(s46)), inference(spm,[status(thm)],[c_0_416, c_0_1737])).
cnf(c_0_1777,plain,(node15(s47)|m_main_v_turn(s48,c_1)), inference(sr,[status(thm)],[c_0_1769, c_0_1765])).
cnf(c_0_1778,plain,(node15(s47)|m_main_v_turn(s47,X1)|~m_main_v_turn(s48,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1770])).
cnf(c_0_1779,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s47,c_n1)|~trans(s47,X1)), inference(spm,[status(thm)],[c_0_1771, c_0_15])).
cnf(c_0_1780,plain,(~m_main_v_turn(s47,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1764])).
cnf(c_0_1781,plain,(m_main_v_state1(X1,c_n1)|~node7(s46,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1737]), c_0_1772]), c_0_1773]), c_0_1774]), c_0_1775]), c_0_1776])).
cnf(c_0_1782,plain,(node15(s47)|~m_main_v_turn(s48,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1777])).
cnf(c_0_1783,plain,(node15(s47)|~m_main_v_state1(s47,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1778, c_0_1779]), c_0_1752])]), c_0_1780])).
cnf(c_0_1784,plain,(m_main_v_state1(X1,c_n1)|~trans(s46,X1)), inference(spm,[status(thm)],[c_0_1781, c_0_7])).
cnf(c_0_1785,plain,(node10(s47)|node9(s47)|node8(s47)|m_main_v_state2(s48,c_c2)|~node15(s47)), inference(spm,[status(thm)],[c_0_467, c_0_1752])).
cnf(c_0_1786,plain,(node15(s47)|~node12(s48)), inference(spm,[status(thm)],[c_0_1782, c_0_35])).
cnf(c_0_1787,plain,(~node10(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_1760]), c_0_1761])])).
cnf(c_0_1788,plain,(~node9(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_1760]), c_0_1761])])).
cnf(c_0_1789,plain,(~node8(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_1760]), c_0_1761])])).
cnf(c_0_1790,plain,(m_main_v_state1(s48,c_t1)|~node15(s47)), inference(spm,[status(thm)],[c_0_474, c_0_1752])).
cnf(c_0_1791,plain,(node15(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1783, c_0_1784]), c_0_1727])])).
cnf(c_0_1792,plain,(m_main_v_turn(s48,c_2)|~node15(s47)), inference(spm,[status(thm)],[c_0_470, c_0_1752])).
cnf(c_0_1793,plain,(m_main_v_state2(s48,c_c2)|~node12(s48)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1785, c_0_1786]), c_0_1787]), c_0_1788]), c_0_1789])).
cnf(c_0_1794,plain,(~node15(s47)|~node11(s48)), inference(spm,[status(thm)],[c_0_165, c_0_1790])).
cnf(c_0_1795,plain,(m_main_v_state2(s48,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1785, c_0_1791])]), c_0_1787]), c_0_1788]), c_0_1789])).
cnf(c_0_1796,plain,(~node15(s47)|~node9(s48)), inference(spm,[status(thm)],[c_0_167, c_0_1790])).
cnf(c_0_1797,plain,(trans(s48,s49)), inference(spm,[status(thm)],[c_0_8, c_0_110])).
cnf(c_0_1798,plain,(~node15(s47)|~node1(s48)), inference(spm,[status(thm)],[c_0_203, c_0_1790])).
cnf(c_0_1799,plain,(~node15(s47)|~node5(s48)), inference(spm,[status(thm)],[c_0_354, c_0_1792])).
cnf(c_0_1800,plain,(~node15(s47)|~node3(s48)), inference(spm,[status(thm)],[c_0_201, c_0_1790])).
cnf(c_0_1801,plain,(~node15(s47)|~node2(s48)), inference(spm,[status(thm)],[c_0_202, c_0_1790])).
cnf(c_0_1802,plain,(~node15(s47)|~m_main_v_state1(s48,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_1790])).
cnf(c_0_1803,plain,(~node12(s48)), inference(spm,[status(thm)],[c_0_345, c_0_1793])).
cnf(c_0_1804,plain,(~node11(s48)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1794, c_0_1791])])).
cnf(c_0_1805,plain,(~node10(s48)), inference(spm,[status(thm)],[c_0_339, c_0_1795])).
cnf(c_0_1806,plain,(~node9(s48)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1796, c_0_1791])])).
cnf(c_0_1807,plain,(~node8(s48)), inference(spm,[status(thm)],[c_0_343, c_0_1795])).
cnf(c_0_1808,plain,(node6(s48,s49)|node5(s48)|node4(s48)|node3(s48)|node2(s48)|node1(s48)|m_main_v_state1(s48,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1797])).
cnf(c_0_1809,plain,(~node1(s48)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1798, c_0_1791])])).
cnf(c_0_1810,plain,(~node5(s48)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1799, c_0_1791])])).
cnf(c_0_1811,plain,(~node4(s48)), inference(spm,[status(thm)],[c_0_361, c_0_1795])).
cnf(c_0_1812,plain,(~node3(s48)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1800, c_0_1791])])).
cnf(c_0_1813,plain,(~node2(s48)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1801, c_0_1791])])).
cnf(c_0_1814,plain,(~m_main_v_state1(s48,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1802, c_0_1791])])).
cnf(c_0_1815,plain,(m_main_v_state2(X1,c_n2)|~node14(s48,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1795]), c_0_1803]), c_0_1804]), c_0_1805]), c_0_1806]), c_0_1807])).
cnf(c_0_1816,plain,(node6(s48,s49)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1808, c_0_1809]), c_0_1810]), c_0_1811]), c_0_1812]), c_0_1813]), c_0_1814])).
cnf(c_0_1817,plain,(node17(s48,s49)|node16(s48)|node15(s48)), inference(spm,[status(thm)],[c_0_152, c_0_1797])).
cnf(c_0_1818,plain,(m_main_v_state2(X1,c_n2)|~trans(s48,X1)), inference(spm,[status(thm)],[c_0_1815, c_0_29])).
cnf(c_0_1819,plain,(m_main_v_state1(s49,X1)|~m_main_v_state1(s48,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1816])).
cnf(c_0_1820,plain,(m_main_v_state1(s48,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1790, c_0_1791])])).
cnf(c_0_1821,plain,(node16(s48)|node15(s48)|m_main_v_turn(s49,X1)|~m_main_v_turn(s48,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1817])).
cnf(c_0_1822,negated_conjecture,(xuntil28(s0)|node26(s0)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_1823,negated_conjecture,(~node26(s0)), inference(spm,[status(thm)],[c_0_142, c_0_113])).
cnf(c_0_1824,plain,(trans(s49,s50)), inference(spm,[status(thm)],[c_0_8, c_0_114])).
cnf(c_0_1825,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s48,X1)), inference(spm,[status(thm)],[c_0_63, c_0_1818])).
cnf(c_0_1826,plain,(~node15(s49)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_1819]), c_0_1820])])).
cnf(c_0_1827,plain,(node16(s48)|node15(s48)|m_main_v_turn(s49,c_2)|~node15(s47)), inference(spm,[status(thm)],[c_0_1821, c_0_1792])).
cnf(c_0_1828,plain,(~node15(s48)|~node15(s47)), inference(spm,[status(thm)],[c_0_189, c_0_1790])).
cnf(c_0_1829,negated_conjecture,(xuntil28(s0)), inference(sr,[status(thm)],[c_0_1822, c_0_1823])).
cnf(c_0_1830,plain,(node17(s49,s50)|node16(s49)|node15(s49)), inference(spm,[status(thm)],[c_0_152, c_0_1824])).
cnf(c_0_1831,plain,(m_main_v_turn(X1,c_1)|~node18(s49,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1825, c_0_1819]), c_0_1797]), c_0_1820])]), c_0_1826])).
cnf(c_0_1832,plain,(node16(s48)|node15(s48)|m_main_v_turn(s49,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1827, c_0_1791])])).
cnf(c_0_1833,plain,(~node15(s48)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1828, c_0_1791])])).
cnf(c_0_1834,negated_conjecture,(until27(s2)|~xuntil28(s1)), inference(spm,[status(thm)],[c_0_115, c_0_34])).
cnf(c_0_1835,negated_conjecture,(until27(s1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_9]), c_0_1829])])).
cnf(c_0_1836,plain,(node16(s49)|node15(s49)|m_main_v_turn(s49,X1)|~m_main_v_turn(s50,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1830])).
cnf(c_0_1837,plain,(m_main_v_turn(X1,c_1)|~trans(s49,X1)), inference(spm,[status(thm)],[c_0_1831, c_0_15])).
cnf(c_0_1838,plain,(node16(s48)|m_main_v_turn(s49,c_2)), inference(sr,[status(thm)],[c_0_1832, c_0_1833])).
cnf(c_0_1839,plain,(~node16(s48)), inference(spm,[status(thm)],[c_0_392, c_0_1795])).
cnf(c_0_1840,negated_conjecture,(until27(s3)|~xuntil28(s2)), inference(spm,[status(thm)],[c_0_115, c_0_42])).
cnf(c_0_1841,negated_conjecture,(xuntil28(s2)|node26(s2)|~xuntil28(s1)), inference(spm,[status(thm)],[c_0_111, c_0_1834])).
cnf(c_0_1842,negated_conjecture,(xuntil28(s1)|node26(s1)), inference(spm,[status(thm)],[c_0_111, c_0_1835])).
cnf(c_0_1843,plain,(node16(s49)|m_main_v_turn(s49,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1836, c_0_1837]), c_0_1824])]), c_0_1826])).
cnf(c_0_1844,plain,(m_main_v_turn(s49,c_2)), inference(sr,[status(thm)],[c_0_1838, c_0_1839])).
cnf(c_0_1845,negated_conjecture,(until27(s4)|~xuntil28(s3)), inference(spm,[status(thm)],[c_0_115, c_0_54])).
cnf(c_0_1846,negated_conjecture,(xuntil28(s3)|node26(s3)|~xuntil28(s2)), inference(spm,[status(thm)],[c_0_111, c_0_1840])).
cnf(c_0_1847,negated_conjecture,(xuntil28(s2)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1841, c_0_1842])).
cnf(c_0_1848,plain,(m_main_v_state2(s50,c_t2)|~node16(s49)), inference(spm,[status(thm)],[c_0_400, c_0_1824])).
cnf(c_0_1849,plain,(node3(s49)|node2(s49)|node1(s49)|m_main_v_state1(s50,c_c1)|~node16(s49)), inference(spm,[status(thm)],[c_0_401, c_0_1824])).
cnf(c_0_1850,plain,(node16(s49)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_1843]), c_0_1844])])).
cnf(c_0_1851,plain,(~node3(s49)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_1819]), c_0_1820])])).
cnf(c_0_1852,plain,(~node2(s49)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_1819]), c_0_1820])])).
cnf(c_0_1853,plain,(~node1(s49)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_1819]), c_0_1820])])).
cnf(c_0_1854,negated_conjecture,(until27(s5)|~xuntil28(s4)), inference(spm,[status(thm)],[c_0_115, c_0_61])).
cnf(c_0_1855,negated_conjecture,(xuntil28(s4)|node26(s4)|~xuntil28(s3)), inference(spm,[status(thm)],[c_0_111, c_0_1845])).
cnf(c_0_1856,negated_conjecture,(xuntil28(s3)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1846, c_0_1847])).
cnf(c_0_1857,plain,(trans(s50,s51)), inference(spm,[status(thm)],[c_0_8, c_0_116])).
cnf(c_0_1858,plain,(~node16(s49)|~node8(s50)), inference(spm,[status(thm)],[c_0_168, c_0_1848])).
cnf(c_0_1859,plain,(m_main_v_state1(s50,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1849, c_0_1850])]), c_0_1851]), c_0_1852]), c_0_1853])).
cnf(c_0_1860,plain,(~node16(s49)|~node10(s50)), inference(spm,[status(thm)],[c_0_257, c_0_1848])).
cnf(c_0_1861,plain,(~node16(s49)|~node9(s50)), inference(spm,[status(thm)],[c_0_258, c_0_1848])).
cnf(c_0_1862,plain,(~node16(s49)|~m_main_v_state2(s50,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_1848])).
cnf(c_0_1863,plain,(node16(s49)|node15(s49)|m_main_v_turn(s50,X1)|~m_main_v_turn(s49,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1830])).
cnf(c_0_1864,negated_conjecture,(until27(s6)|~xuntil28(s5)), inference(spm,[status(thm)],[c_0_115, c_0_62])).
cnf(c_0_1865,negated_conjecture,(xuntil28(s5)|node26(s5)|~xuntil28(s4)), inference(spm,[status(thm)],[c_0_111, c_0_1854])).
cnf(c_0_1866,negated_conjecture,(xuntil28(s4)|node26(s4)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1855, c_0_1856])).
cnf(c_0_1867,plain,(node17(s50,s51)|node16(s50)|node15(s50)), inference(spm,[status(thm)],[c_0_152, c_0_1857])).
cnf(c_0_1868,plain,(node15(s49)|m_main_v_turn(s50,c_1)|~node16(s49)), inference(spm,[status(thm)],[c_0_433, c_0_1824])).
cnf(c_0_1869,plain,(~node16(s50)|~node16(s49)), inference(spm,[status(thm)],[c_0_188, c_0_1848])).
cnf(c_0_1870,plain,(node13(s50,s51)|node12(s50)|node11(s50)|node10(s50)|node9(s50)|node8(s50)|m_main_v_state2(s50,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1857])).
cnf(c_0_1871,plain,(~node8(s50)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1858, c_0_1850])])).
cnf(c_0_1872,plain,(~node12(s50)), inference(spm,[status(thm)],[c_0_428, c_0_1859])).
cnf(c_0_1873,plain,(~node11(s50)), inference(spm,[status(thm)],[c_0_429, c_0_1859])).
cnf(c_0_1874,plain,(~node10(s50)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1860, c_0_1850])])).
cnf(c_0_1875,plain,(~node9(s50)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1861, c_0_1850])])).
cnf(c_0_1876,plain,(~m_main_v_state2(s50,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1862, c_0_1850])])).
cnf(c_0_1877,plain,(node16(s49)|node15(s49)|m_main_v_turn(s50,c_2)), inference(spm,[status(thm)],[c_0_1863, c_0_1844])).
cnf(c_0_1878,negated_conjecture,(xuntil28(s6)|node26(s6)|~xuntil28(s5)), inference(spm,[status(thm)],[c_0_111, c_0_1864])).
cnf(c_0_1879,negated_conjecture,(xuntil28(s5)|node26(s5)|node26(s4)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1865, c_0_1866])).
cnf(c_0_1880,negated_conjecture,(~node26(X1)|~m_main_v_state2(X1,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_113])).
cnf(c_0_1881,plain,(trans(s51,s52)), inference(spm,[status(thm)],[c_0_8, c_0_117])).
cnf(c_0_1882,plain,(node16(s50)|node15(s50)|m_main_v_turn(s51,X1)|~m_main_v_turn(s50,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1867])).
cnf(c_0_1883,plain,(m_main_v_turn(s50,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1868, c_0_1850])]), c_0_1826])).
cnf(c_0_1884,plain,(~node16(s50)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1869, c_0_1850])])).
cnf(c_0_1885,plain,(node13(s50,s51)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1870, c_0_1871]), c_0_1872]), c_0_1873]), c_0_1874]), c_0_1875]), c_0_1876])).
cnf(c_0_1886,plain,(node16(s49)|m_main_v_turn(s50,c_2)), inference(sr,[status(thm)],[c_0_1877, c_0_1826])).
cnf(c_0_1887,negated_conjecture,(until27(s7)|~xuntil28(s6)), inference(spm,[status(thm)],[c_0_115, c_0_67])).
cnf(c_0_1888,negated_conjecture,(xuntil28(s6)|node26(s6)|node26(s5)|node26(s4)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1878, c_0_1879])).
cnf(c_0_1889,negated_conjecture,(~node26(s4)), inference(spm,[status(thm)],[c_0_1880, c_0_346])).
cnf(c_0_1890,plain,(node17(s51,s52)|node16(s51)|node15(s51)), inference(spm,[status(thm)],[c_0_152, c_0_1881])).
cnf(c_0_1891,plain,(node15(s50)|m_main_v_turn(s51,c_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1882, c_0_1883]), c_0_1884])).
cnf(c_0_1892,plain,(~node15(s50)), inference(spm,[status(thm)],[c_0_229, c_0_1859])).
cnf(c_0_1893,plain,(m_main_v_state2(s51,X1)|~m_main_v_state2(s50,X1)), inference(spm,[status(thm)],[c_0_43, c_0_1885])).
cnf(c_0_1894,plain,(m_main_v_state2(s50,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1848, c_0_1850])])).
cnf(c_0_1895,plain,(node16(s49)|~node5(s50)), inference(spm,[status(thm)],[c_0_354, c_0_1886])).
cnf(c_0_1896,negated_conjecture,(xuntil28(s7)|node26(s7)|~xuntil28(s6)), inference(spm,[status(thm)],[c_0_111, c_0_1887])).
cnf(c_0_1897,negated_conjecture,(xuntil28(s6)|node26(s6)|node26(s5)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[c_0_1888, c_0_1889])).
cnf(c_0_1898,negated_conjecture,(~node26(X1)|~trans(s4,X1)), inference(spm,[status(thm)],[c_0_367, c_0_113])).
cnf(c_0_1899,plain,(node16(s51)|node15(s51)|m_main_v_turn(s52,X1)|~m_main_v_turn(s51,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1890])).
cnf(c_0_1900,plain,(m_main_v_turn(s51,c_1)), inference(sr,[status(thm)],[c_0_1891, c_0_1892])).
cnf(c_0_1901,plain,(~node16(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_1893]), c_0_1894])])).
cnf(c_0_1902,plain,(m_main_v_state1(s50,c_c1)|~node5(s50)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1849, c_0_1895]), c_0_1851]), c_0_1852]), c_0_1853])).
cnf(c_0_1903,plain,(~node16(s49)|~node4(s50)), inference(spm,[status(thm)],[c_0_200, c_0_1848])).
cnf(c_0_1904,plain,(~node16(s49)|~node2(s50)), inference(spm,[status(thm)],[c_0_239, c_0_1848])).
cnf(c_0_1905,negated_conjecture,(until27(s8)|~xuntil28(s7)), inference(spm,[status(thm)],[c_0_115, c_0_69])).
cnf(c_0_1906,negated_conjecture,(xuntil28(s7)|node26(s7)|node26(s6)|node26(s5)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1896, c_0_1897])).
cnf(c_0_1907,negated_conjecture,(~node26(s5)), inference(spm,[status(thm)],[c_0_1898, c_0_357])).
cnf(c_0_1908,plain,(node16(s51)|node15(s51)|m_main_v_turn(s52,c_1)), inference(spm,[status(thm)],[c_0_1899, c_0_1900])).
cnf(c_0_1909,plain,(node17(s51,s52)|node15(s51)), inference(sr,[status(thm)],[c_0_1890, c_0_1901])).
cnf(c_0_1910,plain,(m_main_v_turn(X1,c_2)|~node18(s51,X1)|~m_main_v_state1(s51,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1893]), c_0_1894])])).
cnf(c_0_1911,plain,(~node5(s50)), inference(spm,[status(thm)],[c_0_233, c_0_1902])).
cnf(c_0_1912,plain,(~node4(s50)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1903, c_0_1850])])).
cnf(c_0_1913,plain,(~node3(s50)), inference(spm,[status(thm)],[c_0_415, c_0_1859])).
cnf(c_0_1914,plain,(~node2(s50)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1904, c_0_1850])])).
cnf(c_0_1915,plain,(~node1(s50)), inference(spm,[status(thm)],[c_0_416, c_0_1859])).
cnf(c_0_1916,negated_conjecture,(until27(s9)|~xuntil28(s8)), inference(spm,[status(thm)],[c_0_115, c_0_70])).
cnf(c_0_1917,negated_conjecture,(xuntil28(s8)|node26(s8)|~xuntil28(s7)), inference(spm,[status(thm)],[c_0_111, c_0_1905])).
cnf(c_0_1918,negated_conjecture,(xuntil28(s7)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[c_0_1906, c_0_1907])).
cnf(c_0_1919,plain,(node15(s51)|m_main_v_turn(s52,c_1)), inference(sr,[status(thm)],[c_0_1908, c_0_1901])).
cnf(c_0_1920,plain,(node15(s51)|m_main_v_turn(s51,X1)|~m_main_v_turn(s52,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1909])).
cnf(c_0_1921,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s51,c_n1)|~trans(s51,X1)), inference(spm,[status(thm)],[c_0_1910, c_0_15])).
cnf(c_0_1922,plain,(~m_main_v_turn(s51,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1900])).
cnf(c_0_1923,plain,(m_main_v_state1(X1,c_n1)|~node7(s50,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1859]), c_0_1911]), c_0_1912]), c_0_1913]), c_0_1914]), c_0_1915])).
cnf(c_0_1924,negated_conjecture,(xuntil28(s9)|node26(s9)|~xuntil28(s8)), inference(spm,[status(thm)],[c_0_111, c_0_1916])).
cnf(c_0_1925,negated_conjecture,(xuntil28(s8)|node26(s8)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1917, c_0_1918])).
cnf(c_0_1926,plain,(node15(s51)|~m_main_v_turn(s52,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1919])).
cnf(c_0_1927,plain,(node15(s51)|~m_main_v_state1(s51,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1920, c_0_1921]), c_0_1881])]), c_0_1922])).
cnf(c_0_1928,plain,(m_main_v_state1(X1,c_n1)|~trans(s50,X1)), inference(spm,[status(thm)],[c_0_1923, c_0_7])).
cnf(c_0_1929,negated_conjecture,(until27(s10)|~xuntil28(s9)), inference(spm,[status(thm)],[c_0_115, c_0_71])).
cnf(c_0_1930,negated_conjecture,(xuntil28(s9)|node26(s9)|node26(s8)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1924, c_0_1925])).
cnf(c_0_1931,negated_conjecture,(~node26(s8)), inference(spm,[status(thm)],[c_0_1880, c_0_499])).
cnf(c_0_1932,plain,(node10(s51)|node9(s51)|node8(s51)|m_main_v_state2(s52,c_c2)|~node15(s51)), inference(spm,[status(thm)],[c_0_467, c_0_1881])).
cnf(c_0_1933,plain,(node15(s51)|~node12(s52)), inference(spm,[status(thm)],[c_0_1926, c_0_35])).
cnf(c_0_1934,plain,(~node10(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_1893]), c_0_1894])])).
cnf(c_0_1935,plain,(~node9(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_1893]), c_0_1894])])).
cnf(c_0_1936,plain,(~node8(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_1893]), c_0_1894])])).
cnf(c_0_1937,plain,(m_main_v_state1(s52,c_t1)|~node15(s51)), inference(spm,[status(thm)],[c_0_474, c_0_1881])).
cnf(c_0_1938,plain,(node15(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1927, c_0_1928]), c_0_1857])])).
cnf(c_0_1939,plain,(m_main_v_turn(s52,c_2)|~node15(s51)), inference(spm,[status(thm)],[c_0_470, c_0_1881])).
cnf(c_0_1940,negated_conjecture,(xuntil28(s10)|node26(s10)|~xuntil28(s9)), inference(spm,[status(thm)],[c_0_111, c_0_1929])).
cnf(c_0_1941,negated_conjecture,(xuntil28(s9)|node26(s9)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[c_0_1930, c_0_1931])).
cnf(c_0_1942,negated_conjecture,(~node26(X1)|~trans(s8,X1)), inference(spm,[status(thm)],[c_0_564, c_0_113])).
cnf(c_0_1943,plain,(m_main_v_state2(s52,c_c2)|~node12(s52)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1932, c_0_1933]), c_0_1934]), c_0_1935]), c_0_1936])).
cnf(c_0_1944,plain,(~node15(s51)|~node11(s52)), inference(spm,[status(thm)],[c_0_165, c_0_1937])).
cnf(c_0_1945,plain,(~node15(s51)|~node9(s52)), inference(spm,[status(thm)],[c_0_167, c_0_1937])).
cnf(c_0_1946,plain,(m_main_v_state2(s52,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1932, c_0_1938])]), c_0_1934]), c_0_1935]), c_0_1936])).
cnf(c_0_1947,plain,(trans(s52,s53)), inference(spm,[status(thm)],[c_0_8, c_0_118])).
cnf(c_0_1948,plain,(~node15(s51)|~node5(s52)), inference(spm,[status(thm)],[c_0_354, c_0_1939])).
cnf(c_0_1949,plain,(~node15(s51)|~node3(s52)), inference(spm,[status(thm)],[c_0_201, c_0_1937])).
cnf(c_0_1950,plain,(~node15(s51)|~node2(s52)), inference(spm,[status(thm)],[c_0_202, c_0_1937])).
cnf(c_0_1951,plain,(~node15(s51)|~node1(s52)), inference(spm,[status(thm)],[c_0_203, c_0_1937])).
cnf(c_0_1952,plain,(~node15(s51)|~m_main_v_state1(s52,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_1937])).
cnf(c_0_1953,negated_conjecture,(until27(s11)|~xuntil28(s10)), inference(spm,[status(thm)],[c_0_115, c_0_72])).
cnf(c_0_1954,negated_conjecture,(xuntil28(s10)|node26(s10)|node26(s9)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1940, c_0_1941])).
cnf(c_0_1955,negated_conjecture,(~node26(s9)), inference(spm,[status(thm)],[c_0_1942, c_0_461])).
cnf(c_0_1956,plain,(~node12(s52)), inference(spm,[status(thm)],[c_0_345, c_0_1943])).
cnf(c_0_1957,plain,(~node11(s52)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1944, c_0_1938])])).
cnf(c_0_1958,plain,(~node9(s52)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1945, c_0_1938])])).
cnf(c_0_1959,plain,(~node10(s52)), inference(spm,[status(thm)],[c_0_339, c_0_1946])).
cnf(c_0_1960,plain,(~node8(s52)), inference(spm,[status(thm)],[c_0_343, c_0_1946])).
cnf(c_0_1961,plain,(node6(s52,s53)|node5(s52)|node4(s52)|node3(s52)|node2(s52)|node1(s52)|m_main_v_state1(s52,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1947])).
cnf(c_0_1962,plain,(~node5(s52)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1948, c_0_1938])])).
cnf(c_0_1963,plain,(~node4(s52)), inference(spm,[status(thm)],[c_0_361, c_0_1946])).
cnf(c_0_1964,plain,(~node3(s52)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1949, c_0_1938])])).
cnf(c_0_1965,plain,(~node2(s52)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1950, c_0_1938])])).
cnf(c_0_1966,plain,(~node1(s52)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1951, c_0_1938])])).
cnf(c_0_1967,plain,(~m_main_v_state1(s52,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1952, c_0_1938])])).
cnf(c_0_1968,negated_conjecture,(until27(s12)|~xuntil28(s11)), inference(spm,[status(thm)],[c_0_115, c_0_73])).
cnf(c_0_1969,negated_conjecture,(xuntil28(s11)|node26(s11)|~xuntil28(s10)), inference(spm,[status(thm)],[c_0_111, c_0_1953])).
cnf(c_0_1970,negated_conjecture,(xuntil28(s10)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[c_0_1954, c_0_1955])).
cnf(c_0_1971,plain,(m_main_v_state2(X1,c_n2)|~node14(s52,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1946]), c_0_1956]), c_0_1957]), c_0_1958]), c_0_1959]), c_0_1960])).
cnf(c_0_1972,plain,(node6(s52,s53)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1961, c_0_1962]), c_0_1963]), c_0_1964]), c_0_1965]), c_0_1966]), c_0_1967])).
cnf(c_0_1973,negated_conjecture,(until27(s13)|~xuntil28(s12)), inference(spm,[status(thm)],[c_0_115, c_0_74])).
cnf(c_0_1974,negated_conjecture,(xuntil28(s12)|node26(s12)|~xuntil28(s11)), inference(spm,[status(thm)],[c_0_111, c_0_1968])).
cnf(c_0_1975,negated_conjecture,(xuntil28(s11)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1969, c_0_1970])).
cnf(c_0_1976,negated_conjecture,(~node26(s12)), inference(spm,[status(thm)],[c_0_1880, c_0_676])).
cnf(c_0_1977,negated_conjecture,(~node26(X1)|~trans(s12,X1)), inference(spm,[status(thm)],[c_0_692, c_0_113])).
cnf(c_0_1978,plain,(m_main_v_state2(X1,c_n2)|~trans(s52,X1)), inference(spm,[status(thm)],[c_0_1971, c_0_29])).
cnf(c_0_1979,plain,(m_main_v_state1(s53,X1)|~m_main_v_state1(s52,X1)), inference(spm,[status(thm)],[c_0_23, c_0_1972])).
cnf(c_0_1980,plain,(m_main_v_state1(s52,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1937, c_0_1938])])).
cnf(c_0_1981,plain,(node17(s52,s53)|node16(s52)|node15(s52)), inference(spm,[status(thm)],[c_0_152, c_0_1947])).
cnf(c_0_1982,negated_conjecture,(until27(s14)|~xuntil28(s13)), inference(spm,[status(thm)],[c_0_115, c_0_75])).
cnf(c_0_1983,negated_conjecture,(xuntil28(s13)|node26(s13)|~xuntil28(s12)), inference(spm,[status(thm)],[c_0_111, c_0_1973])).
cnf(c_0_1984,negated_conjecture,(xuntil28(s12)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1974, c_0_1975]), c_0_1976])).
cnf(c_0_1985,negated_conjecture,(~node26(s13)), inference(spm,[status(thm)],[c_0_1977, c_0_686])).
cnf(c_0_1986,plain,(trans(s53,s54)), inference(spm,[status(thm)],[c_0_8, c_0_119])).
cnf(c_0_1987,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s52,X1)), inference(spm,[status(thm)],[c_0_63, c_0_1978])).
cnf(c_0_1988,plain,(~node15(s53)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_1979]), c_0_1980])])).
cnf(c_0_1989,plain,(node16(s52)|node15(s52)|m_main_v_turn(s53,X1)|~m_main_v_turn(s52,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1981])).
cnf(c_0_1990,plain,(~node15(s52)|~node15(s51)), inference(spm,[status(thm)],[c_0_189, c_0_1937])).
cnf(c_0_1991,negated_conjecture,(until27(s15)|~xuntil28(s14)), inference(spm,[status(thm)],[c_0_115, c_0_76])).
cnf(c_0_1992,negated_conjecture,(xuntil28(s14)|node26(s14)|~xuntil28(s13)), inference(spm,[status(thm)],[c_0_111, c_0_1982])).
cnf(c_0_1993,negated_conjecture,(xuntil28(s13)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1983, c_0_1984]), c_0_1985])).
cnf(c_0_1994,plain,(node17(s53,s54)|node16(s53)|node15(s53)), inference(spm,[status(thm)],[c_0_152, c_0_1986])).
cnf(c_0_1995,plain,(m_main_v_turn(X1,c_1)|~node18(s53,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1987, c_0_1979]), c_0_1947]), c_0_1980])]), c_0_1988])).
cnf(c_0_1996,plain,(node16(s52)|node15(s52)|m_main_v_turn(s53,c_2)|~node15(s51)), inference(spm,[status(thm)],[c_0_1989, c_0_1939])).
cnf(c_0_1997,plain,(~node15(s52)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1990, c_0_1938])])).
cnf(c_0_1998,negated_conjecture,(until27(s16)|~xuntil28(s15)), inference(spm,[status(thm)],[c_0_115, c_0_77])).
cnf(c_0_1999,negated_conjecture,(xuntil28(s15)|node26(s15)|~xuntil28(s14)), inference(spm,[status(thm)],[c_0_111, c_0_1991])).
cnf(c_0_2000,negated_conjecture,(xuntil28(s14)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1992, c_0_1993])).
cnf(c_0_2001,plain,(~m_main_v_state2(X1,c_t2)|~trans(s16,X1)), inference(spm,[status(thm)],[c_0_0, c_0_858])).
cnf(c_0_2002,plain,(node16(s53)|node15(s53)|m_main_v_turn(s53,X1)|~m_main_v_turn(s54,X1)), inference(spm,[status(thm)],[c_0_66, c_0_1994])).
cnf(c_0_2003,plain,(m_main_v_turn(X1,c_1)|~trans(s53,X1)), inference(spm,[status(thm)],[c_0_1995, c_0_15])).
cnf(c_0_2004,plain,(node16(s52)|m_main_v_turn(s53,c_2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1996, c_0_1938])]), c_0_1997])).
cnf(c_0_2005,plain,(~node16(s52)), inference(spm,[status(thm)],[c_0_392, c_0_1946])).
cnf(c_0_2006,negated_conjecture,(until27(s17)|~xuntil28(s16)), inference(spm,[status(thm)],[c_0_115, c_0_78])).
cnf(c_0_2007,negated_conjecture,(xuntil28(s16)|node26(s16)|~xuntil28(s15)), inference(spm,[status(thm)],[c_0_111, c_0_1998])).
cnf(c_0_2008,negated_conjecture,(xuntil28(s15)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_1999, c_0_2000])).
cnf(c_0_2009,negated_conjecture,(~node26(s16)), inference(spm,[status(thm)],[c_0_1880, c_0_827])).
cnf(c_0_2010,negated_conjecture,(~node26(X1)|~trans(s16,X1)), inference(spm,[status(thm)],[c_0_2001, c_0_113])).
cnf(c_0_2011,plain,(node16(s53)|m_main_v_turn(s53,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2002, c_0_2003]), c_0_1986])]), c_0_1988])).
cnf(c_0_2012,plain,(m_main_v_turn(s53,c_2)), inference(sr,[status(thm)],[c_0_2004, c_0_2005])).
cnf(c_0_2013,negated_conjecture,(until27(s18)|~xuntil28(s17)), inference(spm,[status(thm)],[c_0_115, c_0_79])).
cnf(c_0_2014,negated_conjecture,(xuntil28(s17)|node26(s17)|~xuntil28(s16)), inference(spm,[status(thm)],[c_0_111, c_0_2006])).
cnf(c_0_2015,negated_conjecture,(xuntil28(s16)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2007, c_0_2008]), c_0_2009])).
cnf(c_0_2016,negated_conjecture,(~node26(s17)), inference(spm,[status(thm)],[c_0_2010, c_0_799])).
cnf(c_0_2017,plain,(m_main_v_state2(s54,c_t2)|~node16(s53)), inference(spm,[status(thm)],[c_0_400, c_0_1986])).
cnf(c_0_2018,plain,(node3(s53)|node2(s53)|node1(s53)|m_main_v_state1(s54,c_c1)|~node16(s53)), inference(spm,[status(thm)],[c_0_401, c_0_1986])).
cnf(c_0_2019,plain,(node16(s53)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2011]), c_0_2012])])).
cnf(c_0_2020,plain,(~node3(s53)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_1979]), c_0_1980])])).
cnf(c_0_2021,plain,(~node2(s53)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_1979]), c_0_1980])])).
cnf(c_0_2022,plain,(~node1(s53)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_1979]), c_0_1980])])).
cnf(c_0_2023,negated_conjecture,(until27(s19)|~xuntil28(s18)), inference(spm,[status(thm)],[c_0_115, c_0_80])).
cnf(c_0_2024,negated_conjecture,(xuntil28(s18)|node26(s18)|~xuntil28(s17)), inference(spm,[status(thm)],[c_0_111, c_0_2013])).
cnf(c_0_2025,negated_conjecture,(xuntil28(s17)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2014, c_0_2015]), c_0_2016])).
cnf(c_0_2026,plain,(trans(s54,s55)), inference(spm,[status(thm)],[c_0_8, c_0_120])).
cnf(c_0_2027,plain,(~node16(s53)|~node8(s54)), inference(spm,[status(thm)],[c_0_168, c_0_2017])).
cnf(c_0_2028,plain,(m_main_v_state1(s54,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2018, c_0_2019])]), c_0_2020]), c_0_2021]), c_0_2022])).
cnf(c_0_2029,plain,(~node16(s53)|~node10(s54)), inference(spm,[status(thm)],[c_0_257, c_0_2017])).
cnf(c_0_2030,plain,(~node16(s53)|~node9(s54)), inference(spm,[status(thm)],[c_0_258, c_0_2017])).
cnf(c_0_2031,plain,(~node16(s53)|~m_main_v_state2(s54,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_2017])).
cnf(c_0_2032,plain,(node16(s53)|node15(s53)|m_main_v_turn(s54,X1)|~m_main_v_turn(s53,X1)), inference(spm,[status(thm)],[c_0_24, c_0_1994])).
cnf(c_0_2033,negated_conjecture,(until27(s20)|~xuntil28(s19)), inference(spm,[status(thm)],[c_0_115, c_0_81])).
cnf(c_0_2034,negated_conjecture,(xuntil28(s19)|node26(s19)|~xuntil28(s18)), inference(spm,[status(thm)],[c_0_111, c_0_2023])).
cnf(c_0_2035,negated_conjecture,(xuntil28(s18)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2024, c_0_2025])).
cnf(c_0_2036,plain,(node17(s54,s55)|node16(s54)|node15(s54)), inference(spm,[status(thm)],[c_0_152, c_0_2026])).
cnf(c_0_2037,plain,(node15(s53)|m_main_v_turn(s54,c_1)|~node16(s53)), inference(spm,[status(thm)],[c_0_433, c_0_1986])).
cnf(c_0_2038,plain,(~node16(s54)|~node16(s53)), inference(spm,[status(thm)],[c_0_188, c_0_2017])).
cnf(c_0_2039,plain,(node13(s54,s55)|node12(s54)|node11(s54)|node10(s54)|node9(s54)|node8(s54)|m_main_v_state2(s54,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_2026])).
cnf(c_0_2040,plain,(~node8(s54)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2027, c_0_2019])])).
cnf(c_0_2041,plain,(~node12(s54)), inference(spm,[status(thm)],[c_0_428, c_0_2028])).
cnf(c_0_2042,plain,(~node11(s54)), inference(spm,[status(thm)],[c_0_429, c_0_2028])).
cnf(c_0_2043,plain,(~node10(s54)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2029, c_0_2019])])).
cnf(c_0_2044,plain,(~node9(s54)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2030, c_0_2019])])).
cnf(c_0_2045,plain,(~m_main_v_state2(s54,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2031, c_0_2019])])).
cnf(c_0_2046,plain,(node16(s53)|node15(s53)|m_main_v_turn(s54,c_2)), inference(spm,[status(thm)],[c_0_2032, c_0_2012])).
cnf(c_0_2047,negated_conjecture,(until27(s21)|~xuntil28(s20)), inference(spm,[status(thm)],[c_0_115, c_0_82])).
cnf(c_0_2048,negated_conjecture,(xuntil28(s20)|node26(s20)|~xuntil28(s19)), inference(spm,[status(thm)],[c_0_111, c_0_2033])).
cnf(c_0_2049,negated_conjecture,(xuntil28(s19)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2034, c_0_2035])).
cnf(c_0_2050,negated_conjecture,(~node26(s20)), inference(spm,[status(thm)],[c_0_1880, c_0_948])).
cnf(c_0_2051,negated_conjecture,(~node26(X1)|~trans(s20,X1)), inference(spm,[status(thm)],[c_0_991, c_0_113])).
cnf(c_0_2052,plain,(trans(s55,s56)), inference(spm,[status(thm)],[c_0_8, c_0_121])).
cnf(c_0_2053,plain,(node16(s54)|node15(s54)|m_main_v_turn(s55,X1)|~m_main_v_turn(s54,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2036])).
cnf(c_0_2054,plain,(m_main_v_turn(s54,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2037, c_0_2019])]), c_0_1988])).
cnf(c_0_2055,plain,(~node16(s54)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2038, c_0_2019])])).
cnf(c_0_2056,plain,(node13(s54,s55)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2039, c_0_2040]), c_0_2041]), c_0_2042]), c_0_2043]), c_0_2044]), c_0_2045])).
cnf(c_0_2057,plain,(node16(s53)|m_main_v_turn(s54,c_2)), inference(sr,[status(thm)],[c_0_2046, c_0_1988])).
cnf(c_0_2058,negated_conjecture,(until27(s22)|~xuntil28(s21)), inference(spm,[status(thm)],[c_0_115, c_0_83])).
cnf(c_0_2059,negated_conjecture,(xuntil28(s21)|node26(s21)|~xuntil28(s20)), inference(spm,[status(thm)],[c_0_111, c_0_2047])).
cnf(c_0_2060,negated_conjecture,(xuntil28(s20)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2048, c_0_2049]), c_0_2050])).
cnf(c_0_2061,negated_conjecture,(~node26(s21)), inference(spm,[status(thm)],[c_0_2051, c_0_923])).
cnf(c_0_2062,plain,(node17(s55,s56)|node16(s55)|node15(s55)), inference(spm,[status(thm)],[c_0_152, c_0_2052])).
cnf(c_0_2063,plain,(node15(s54)|m_main_v_turn(s55,c_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2053, c_0_2054]), c_0_2055])).
cnf(c_0_2064,plain,(~node15(s54)), inference(spm,[status(thm)],[c_0_229, c_0_2028])).
cnf(c_0_2065,plain,(m_main_v_state2(s55,X1)|~m_main_v_state2(s54,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2056])).
cnf(c_0_2066,plain,(m_main_v_state2(s54,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2017, c_0_2019])])).
cnf(c_0_2067,plain,(node16(s53)|~node5(s54)), inference(spm,[status(thm)],[c_0_354, c_0_2057])).
cnf(c_0_2068,negated_conjecture,(until27(s23)|~xuntil28(s22)), inference(spm,[status(thm)],[c_0_115, c_0_84])).
cnf(c_0_2069,negated_conjecture,(xuntil28(s22)|node26(s22)|~xuntil28(s21)), inference(spm,[status(thm)],[c_0_111, c_0_2058])).
cnf(c_0_2070,negated_conjecture,(xuntil28(s21)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2059, c_0_2060]), c_0_2061])).
cnf(c_0_2071,plain,(node16(s55)|node15(s55)|m_main_v_turn(s56,X1)|~m_main_v_turn(s55,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2062])).
cnf(c_0_2072,plain,(m_main_v_turn(s55,c_1)), inference(sr,[status(thm)],[c_0_2063, c_0_2064])).
cnf(c_0_2073,plain,(~node16(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_2065]), c_0_2066])])).
cnf(c_0_2074,plain,(m_main_v_state1(s54,c_c1)|~node5(s54)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2018, c_0_2067]), c_0_2020]), c_0_2021]), c_0_2022])).
cnf(c_0_2075,plain,(~node16(s53)|~node4(s54)), inference(spm,[status(thm)],[c_0_200, c_0_2017])).
cnf(c_0_2076,plain,(~node16(s53)|~node2(s54)), inference(spm,[status(thm)],[c_0_239, c_0_2017])).
cnf(c_0_2077,negated_conjecture,(until27(s24)|~xuntil28(s23)), inference(spm,[status(thm)],[c_0_115, c_0_85])).
cnf(c_0_2078,negated_conjecture,(xuntil28(s23)|node26(s23)|~xuntil28(s22)), inference(spm,[status(thm)],[c_0_111, c_0_2068])).
cnf(c_0_2079,negated_conjecture,(xuntil28(s22)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2069, c_0_2070])).
cnf(c_0_2080,plain,(node16(s55)|node15(s55)|m_main_v_turn(s56,c_1)), inference(spm,[status(thm)],[c_0_2071, c_0_2072])).
cnf(c_0_2081,plain,(node17(s55,s56)|node15(s55)), inference(sr,[status(thm)],[c_0_2062, c_0_2073])).
cnf(c_0_2082,plain,(m_main_v_turn(X1,c_2)|~node18(s55,X1)|~m_main_v_state1(s55,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_2065]), c_0_2066])])).
cnf(c_0_2083,plain,(~node5(s54)), inference(spm,[status(thm)],[c_0_233, c_0_2074])).
cnf(c_0_2084,plain,(~node4(s54)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2075, c_0_2019])])).
cnf(c_0_2085,plain,(~node2(s54)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2076, c_0_2019])])).
cnf(c_0_2086,plain,(~node3(s54)), inference(spm,[status(thm)],[c_0_415, c_0_2028])).
cnf(c_0_2087,plain,(~node1(s54)), inference(spm,[status(thm)],[c_0_416, c_0_2028])).
cnf(c_0_2088,negated_conjecture,(until27(s25)|~xuntil28(s24)), inference(spm,[status(thm)],[c_0_115, c_0_86])).
cnf(c_0_2089,negated_conjecture,(xuntil28(s24)|node26(s24)|~xuntil28(s23)), inference(spm,[status(thm)],[c_0_111, c_0_2077])).
cnf(c_0_2090,negated_conjecture,(xuntil28(s23)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2078, c_0_2079])).
cnf(c_0_2091,negated_conjecture,(~node26(s24)), inference(spm,[status(thm)],[c_0_1880, c_0_1098])).
cnf(c_0_2092,negated_conjecture,(~node26(X1)|~trans(s24,X1)), inference(spm,[status(thm)],[c_0_1113, c_0_113])).
cnf(c_0_2093,plain,(node15(s55)|m_main_v_turn(s56,c_1)), inference(sr,[status(thm)],[c_0_2080, c_0_2073])).
cnf(c_0_2094,plain,(node15(s55)|m_main_v_turn(s55,X1)|~m_main_v_turn(s56,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2081])).
cnf(c_0_2095,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s55,c_n1)|~trans(s55,X1)), inference(spm,[status(thm)],[c_0_2082, c_0_15])).
cnf(c_0_2096,plain,(~m_main_v_turn(s55,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_2072])).
cnf(c_0_2097,plain,(m_main_v_state1(X1,c_n1)|~node7(s54,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_2028]), c_0_2083]), c_0_2084]), c_0_2085]), c_0_2086]), c_0_2087])).
cnf(c_0_2098,negated_conjecture,(until27(s26)|~xuntil28(s25)), inference(spm,[status(thm)],[c_0_115, c_0_87])).
cnf(c_0_2099,negated_conjecture,(xuntil28(s25)|node26(s25)|~xuntil28(s24)), inference(spm,[status(thm)],[c_0_111, c_0_2088])).
cnf(c_0_2100,negated_conjecture,(xuntil28(s24)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2089, c_0_2090]), c_0_2091])).
cnf(c_0_2101,negated_conjecture,(~node26(s25)), inference(spm,[status(thm)],[c_0_2092, c_0_1107])).
cnf(c_0_2102,plain,(node15(s55)|~m_main_v_turn(s56,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_2093])).
cnf(c_0_2103,plain,(node15(s55)|~m_main_v_state1(s55,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2094, c_0_2095]), c_0_2052])]), c_0_2096])).
cnf(c_0_2104,plain,(m_main_v_state1(X1,c_n1)|~trans(s54,X1)), inference(spm,[status(thm)],[c_0_2097, c_0_7])).
cnf(c_0_2105,negated_conjecture,(until27(s27)|~xuntil28(s26)), inference(spm,[status(thm)],[c_0_115, c_0_88])).
cnf(c_0_2106,negated_conjecture,(xuntil28(s26)|node26(s26)|~xuntil28(s25)), inference(spm,[status(thm)],[c_0_111, c_0_2098])).
cnf(c_0_2107,negated_conjecture,(xuntil28(s25)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2099, c_0_2100]), c_0_2101])).
cnf(c_0_2108,plain,(node10(s55)|node9(s55)|node8(s55)|m_main_v_state2(s56,c_c2)|~node15(s55)), inference(spm,[status(thm)],[c_0_467, c_0_2052])).
cnf(c_0_2109,plain,(node15(s55)|~node12(s56)), inference(spm,[status(thm)],[c_0_2102, c_0_35])).
cnf(c_0_2110,plain,(~node10(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_2065]), c_0_2066])])).
cnf(c_0_2111,plain,(~node9(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_2065]), c_0_2066])])).
cnf(c_0_2112,plain,(~node8(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_2065]), c_0_2066])])).
cnf(c_0_2113,plain,(m_main_v_state1(s56,c_t1)|~node15(s55)), inference(spm,[status(thm)],[c_0_474, c_0_2052])).
cnf(c_0_2114,plain,(node15(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2103, c_0_2104]), c_0_2026])])).
cnf(c_0_2115,plain,(m_main_v_turn(s56,c_2)|~node15(s55)), inference(spm,[status(thm)],[c_0_470, c_0_2052])).
cnf(c_0_2116,negated_conjecture,(until27(s28)|~xuntil28(s27)), inference(spm,[status(thm)],[c_0_115, c_0_89])).
cnf(c_0_2117,negated_conjecture,(xuntil28(s27)|node26(s27)|~xuntil28(s26)), inference(spm,[status(thm)],[c_0_111, c_0_2105])).
cnf(c_0_2118,negated_conjecture,(xuntil28(s26)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2106, c_0_2107])).
cnf(c_0_2119,plain,(~m_main_v_state2(X1,c_t2)|~trans(s28,X1)), inference(spm,[status(thm)],[c_0_0, c_0_1267])).
cnf(c_0_2120,plain,(m_main_v_state2(s56,c_c2)|~node12(s56)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2108, c_0_2109]), c_0_2110]), c_0_2111]), c_0_2112])).
cnf(c_0_2121,plain,(~node15(s55)|~node11(s56)), inference(spm,[status(thm)],[c_0_165, c_0_2113])).
cnf(c_0_2122,plain,(m_main_v_state2(s56,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2108, c_0_2114])]), c_0_2110]), c_0_2111]), c_0_2112])).
cnf(c_0_2123,plain,(~node15(s55)|~node9(s56)), inference(spm,[status(thm)],[c_0_167, c_0_2113])).
cnf(c_0_2124,plain,(trans(s56,s57)), inference(spm,[status(thm)],[c_0_8, c_0_122])).
cnf(c_0_2125,plain,(~node15(s55)|~node1(s56)), inference(spm,[status(thm)],[c_0_203, c_0_2113])).
cnf(c_0_2126,plain,(~node15(s55)|~node5(s56)), inference(spm,[status(thm)],[c_0_354, c_0_2115])).
cnf(c_0_2127,plain,(~node15(s55)|~node3(s56)), inference(spm,[status(thm)],[c_0_201, c_0_2113])).
cnf(c_0_2128,plain,(~node15(s55)|~node2(s56)), inference(spm,[status(thm)],[c_0_202, c_0_2113])).
cnf(c_0_2129,plain,(~node15(s55)|~m_main_v_state1(s56,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_2113])).
cnf(c_0_2130,negated_conjecture,(until27(s29)|~xuntil28(s28)), inference(spm,[status(thm)],[c_0_115, c_0_90])).
cnf(c_0_2131,negated_conjecture,(xuntil28(s28)|node26(s28)|~xuntil28(s27)), inference(spm,[status(thm)],[c_0_111, c_0_2116])).
cnf(c_0_2132,negated_conjecture,(xuntil28(s27)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2117, c_0_2118])).
cnf(c_0_2133,negated_conjecture,(~node26(s28)), inference(spm,[status(thm)],[c_0_1880, c_0_1242])).
cnf(c_0_2134,negated_conjecture,(~node26(X1)|~trans(s28,X1)), inference(spm,[status(thm)],[c_0_2119, c_0_113])).
cnf(c_0_2135,plain,(~node12(s56)), inference(spm,[status(thm)],[c_0_345, c_0_2120])).
cnf(c_0_2136,plain,(~node11(s56)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2121, c_0_2114])])).
cnf(c_0_2137,plain,(~node10(s56)), inference(spm,[status(thm)],[c_0_339, c_0_2122])).
cnf(c_0_2138,plain,(~node9(s56)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2123, c_0_2114])])).
cnf(c_0_2139,plain,(~node8(s56)), inference(spm,[status(thm)],[c_0_343, c_0_2122])).
cnf(c_0_2140,plain,(node6(s56,s57)|node5(s56)|node4(s56)|node3(s56)|node2(s56)|node1(s56)|m_main_v_state1(s56,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_2124])).
cnf(c_0_2141,plain,(~node1(s56)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2125, c_0_2114])])).
cnf(c_0_2142,plain,(~node5(s56)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2126, c_0_2114])])).
cnf(c_0_2143,plain,(~node4(s56)), inference(spm,[status(thm)],[c_0_361, c_0_2122])).
cnf(c_0_2144,plain,(~node3(s56)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2127, c_0_2114])])).
cnf(c_0_2145,plain,(~node2(s56)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2128, c_0_2114])])).
cnf(c_0_2146,plain,(~m_main_v_state1(s56,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2129, c_0_2114])])).
cnf(c_0_2147,negated_conjecture,(until27(s30)|~xuntil28(s29)), inference(spm,[status(thm)],[c_0_115, c_0_91])).
cnf(c_0_2148,negated_conjecture,(xuntil28(s29)|node26(s29)|~xuntil28(s28)), inference(spm,[status(thm)],[c_0_111, c_0_2130])).
cnf(c_0_2149,negated_conjecture,(xuntil28(s28)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2131, c_0_2132]), c_0_2133])).
cnf(c_0_2150,negated_conjecture,(~node26(s29)), inference(spm,[status(thm)],[c_0_2134, c_0_1222])).
cnf(c_0_2151,plain,(m_main_v_state2(X1,c_n2)|~node14(s56,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_2122]), c_0_2135]), c_0_2136]), c_0_2137]), c_0_2138]), c_0_2139])).
cnf(c_0_2152,plain,(node6(s56,s57)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2140, c_0_2141]), c_0_2142]), c_0_2143]), c_0_2144]), c_0_2145]), c_0_2146])).
cnf(c_0_2153,plain,(node17(s56,s57)|node16(s56)|node15(s56)), inference(spm,[status(thm)],[c_0_152, c_0_2124])).
cnf(c_0_2154,negated_conjecture,(until27(s31)|~xuntil28(s30)), inference(spm,[status(thm)],[c_0_115, c_0_92])).
cnf(c_0_2155,negated_conjecture,(xuntil28(s30)|node26(s30)|~xuntil28(s29)), inference(spm,[status(thm)],[c_0_111, c_0_2147])).
cnf(c_0_2156,negated_conjecture,(xuntil28(s29)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2148, c_0_2149]), c_0_2150])).
cnf(c_0_2157,plain,(m_main_v_state2(X1,c_n2)|~trans(s56,X1)), inference(spm,[status(thm)],[c_0_2151, c_0_29])).
cnf(c_0_2158,plain,(m_main_v_state1(s57,X1)|~m_main_v_state1(s56,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2152])).
cnf(c_0_2159,plain,(m_main_v_state1(s56,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2113, c_0_2114])])).
cnf(c_0_2160,plain,(node16(s56)|node15(s56)|m_main_v_turn(s57,X1)|~m_main_v_turn(s56,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2153])).
cnf(c_0_2161,negated_conjecture,(until27(s32)|~xuntil28(s31)), inference(spm,[status(thm)],[c_0_115, c_0_93])).
cnf(c_0_2162,negated_conjecture,(xuntil28(s31)|node26(s31)|~xuntil28(s30)), inference(spm,[status(thm)],[c_0_111, c_0_2154])).
cnf(c_0_2163,negated_conjecture,(xuntil28(s30)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2155, c_0_2156])).
cnf(c_0_2164,plain,(~m_main_v_state2(X1,c_t2)|~trans(s32,X1)), inference(spm,[status(thm)],[c_0_0, c_0_1379])).
cnf(c_0_2165,plain,(trans(s57,s58)), inference(spm,[status(thm)],[c_0_8, c_0_123])).
cnf(c_0_2166,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s56,X1)), inference(spm,[status(thm)],[c_0_63, c_0_2157])).
cnf(c_0_2167,plain,(~node15(s57)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_2158]), c_0_2159])])).
cnf(c_0_2168,plain,(node16(s56)|node15(s56)|m_main_v_turn(s57,c_2)|~node15(s55)), inference(spm,[status(thm)],[c_0_2160, c_0_2115])).
cnf(c_0_2169,plain,(~node15(s56)|~node15(s55)), inference(spm,[status(thm)],[c_0_189, c_0_2113])).
cnf(c_0_2170,negated_conjecture,(until27(s33)|~xuntil28(s32)), inference(spm,[status(thm)],[c_0_115, c_0_94])).
cnf(c_0_2171,negated_conjecture,(xuntil28(s32)|node26(s32)|~xuntil28(s31)), inference(spm,[status(thm)],[c_0_111, c_0_2161])).
cnf(c_0_2172,negated_conjecture,(xuntil28(s31)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2162, c_0_2163])).
cnf(c_0_2173,negated_conjecture,(~node26(s32)), inference(spm,[status(thm)],[c_0_1880, c_0_1356])).
cnf(c_0_2174,negated_conjecture,(~node26(X1)|~trans(s32,X1)), inference(spm,[status(thm)],[c_0_2164, c_0_113])).
cnf(c_0_2175,plain,(node17(s57,s58)|node16(s57)|node15(s57)), inference(spm,[status(thm)],[c_0_152, c_0_2165])).
cnf(c_0_2176,plain,(m_main_v_turn(X1,c_1)|~node18(s57,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2166, c_0_2158]), c_0_2124]), c_0_2159])]), c_0_2167])).
cnf(c_0_2177,plain,(node16(s56)|node15(s56)|m_main_v_turn(s57,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2168, c_0_2114])])).
cnf(c_0_2178,plain,(~node15(s56)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2169, c_0_2114])])).
cnf(c_0_2179,negated_conjecture,(until27(s34)|~xuntil28(s33)), inference(spm,[status(thm)],[c_0_115, c_0_95])).
cnf(c_0_2180,negated_conjecture,(xuntil28(s33)|node26(s33)|~xuntil28(s32)), inference(spm,[status(thm)],[c_0_111, c_0_2170])).
cnf(c_0_2181,negated_conjecture,(xuntil28(s32)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2171, c_0_2172]), c_0_2173])).
cnf(c_0_2182,negated_conjecture,(~node26(s33)), inference(spm,[status(thm)],[c_0_2174, c_0_1358])).
cnf(c_0_2183,plain,(node16(s57)|node15(s57)|m_main_v_turn(s57,X1)|~m_main_v_turn(s58,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2175])).
cnf(c_0_2184,plain,(m_main_v_turn(X1,c_1)|~trans(s57,X1)), inference(spm,[status(thm)],[c_0_2176, c_0_15])).
cnf(c_0_2185,plain,(node16(s56)|m_main_v_turn(s57,c_2)), inference(sr,[status(thm)],[c_0_2177, c_0_2178])).
cnf(c_0_2186,plain,(~node16(s56)), inference(spm,[status(thm)],[c_0_392, c_0_2122])).
cnf(c_0_2187,negated_conjecture,(until27(s35)|~xuntil28(s34)), inference(spm,[status(thm)],[c_0_115, c_0_96])).
cnf(c_0_2188,negated_conjecture,(xuntil28(s34)|node26(s34)|~xuntil28(s33)), inference(spm,[status(thm)],[c_0_111, c_0_2179])).
cnf(c_0_2189,negated_conjecture,(xuntil28(s33)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2180, c_0_2181]), c_0_2182])).
cnf(c_0_2190,plain,(node16(s57)|m_main_v_turn(s57,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2183, c_0_2184]), c_0_2165])]), c_0_2167])).
cnf(c_0_2191,plain,(m_main_v_turn(s57,c_2)), inference(sr,[status(thm)],[c_0_2185, c_0_2186])).
cnf(c_0_2192,negated_conjecture,(until27(s36)|~xuntil28(s35)), inference(spm,[status(thm)],[c_0_115, c_0_97])).
cnf(c_0_2193,negated_conjecture,(xuntil28(s35)|node26(s35)|~xuntil28(s34)), inference(spm,[status(thm)],[c_0_111, c_0_2187])).
cnf(c_0_2194,negated_conjecture,(xuntil28(s34)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2188, c_0_2189])).
cnf(c_0_2195,plain,(~m_main_v_state2(X1,c_t2)|~trans(s36,X1)), inference(spm,[status(thm)],[c_0_0, c_0_1490])).
cnf(c_0_2196,plain,(m_main_v_state2(s58,c_t2)|~node16(s57)), inference(spm,[status(thm)],[c_0_400, c_0_2165])).
cnf(c_0_2197,plain,(node3(s57)|node2(s57)|node1(s57)|m_main_v_state1(s58,c_c1)|~node16(s57)), inference(spm,[status(thm)],[c_0_401, c_0_2165])).
cnf(c_0_2198,plain,(node16(s57)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2190]), c_0_2191])])).
cnf(c_0_2199,plain,(~node3(s57)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_2158]), c_0_2159])])).
cnf(c_0_2200,plain,(~node2(s57)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_2158]), c_0_2159])])).
cnf(c_0_2201,plain,(~node1(s57)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_2158]), c_0_2159])])).
cnf(c_0_2202,negated_conjecture,(until27(s37)|~xuntil28(s36)), inference(spm,[status(thm)],[c_0_115, c_0_98])).
cnf(c_0_2203,negated_conjecture,(xuntil28(s36)|node26(s36)|~xuntil28(s35)), inference(spm,[status(thm)],[c_0_111, c_0_2192])).
cnf(c_0_2204,negated_conjecture,(xuntil28(s35)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2193, c_0_2194])).
cnf(c_0_2205,negated_conjecture,(~node26(s36)), inference(spm,[status(thm)],[c_0_1880, c_0_1468])).
cnf(c_0_2206,negated_conjecture,(~node26(X1)|~trans(s36,X1)), inference(spm,[status(thm)],[c_0_2195, c_0_113])).
cnf(c_0_2207,plain,(trans(s58,s59)), inference(spm,[status(thm)],[c_0_8, c_0_124])).
cnf(c_0_2208,plain,(~node16(s57)|~node8(s58)), inference(spm,[status(thm)],[c_0_168, c_0_2196])).
cnf(c_0_2209,plain,(m_main_v_state1(s58,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2197, c_0_2198])]), c_0_2199]), c_0_2200]), c_0_2201])).
cnf(c_0_2210,plain,(~node16(s57)|~node10(s58)), inference(spm,[status(thm)],[c_0_257, c_0_2196])).
cnf(c_0_2211,plain,(~node16(s57)|~node9(s58)), inference(spm,[status(thm)],[c_0_258, c_0_2196])).
cnf(c_0_2212,plain,(~node16(s57)|~m_main_v_state2(s58,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_2196])).
cnf(c_0_2213,plain,(node16(s57)|node15(s57)|m_main_v_turn(s58,X1)|~m_main_v_turn(s57,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2175])).
cnf(c_0_2214,negated_conjecture,(until27(s38)|~xuntil28(s37)), inference(spm,[status(thm)],[c_0_115, c_0_99])).
cnf(c_0_2215,negated_conjecture,(xuntil28(s37)|node26(s37)|~xuntil28(s36)), inference(spm,[status(thm)],[c_0_111, c_0_2202])).
cnf(c_0_2216,negated_conjecture,(xuntil28(s36)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2203, c_0_2204]), c_0_2205])).
cnf(c_0_2217,negated_conjecture,(~node26(s37)), inference(spm,[status(thm)],[c_0_2206, c_0_1469])).
cnf(c_0_2218,plain,(node17(s58,s59)|node16(s58)|node15(s58)), inference(spm,[status(thm)],[c_0_152, c_0_2207])).
cnf(c_0_2219,plain,(node15(s57)|m_main_v_turn(s58,c_1)|~node16(s57)), inference(spm,[status(thm)],[c_0_433, c_0_2165])).
cnf(c_0_2220,plain,(~node16(s58)|~node16(s57)), inference(spm,[status(thm)],[c_0_188, c_0_2196])).
cnf(c_0_2221,plain,(node13(s58,s59)|node12(s58)|node11(s58)|node10(s58)|node9(s58)|node8(s58)|m_main_v_state2(s58,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_2207])).
cnf(c_0_2222,plain,(~node8(s58)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2208, c_0_2198])])).
cnf(c_0_2223,plain,(~node12(s58)), inference(spm,[status(thm)],[c_0_428, c_0_2209])).
cnf(c_0_2224,plain,(~node11(s58)), inference(spm,[status(thm)],[c_0_429, c_0_2209])).
cnf(c_0_2225,plain,(~node10(s58)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2210, c_0_2198])])).
cnf(c_0_2226,plain,(~node9(s58)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2211, c_0_2198])])).
cnf(c_0_2227,plain,(~m_main_v_state2(s58,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2212, c_0_2198])])).
cnf(c_0_2228,plain,(node16(s57)|node15(s57)|m_main_v_turn(s58,c_2)), inference(spm,[status(thm)],[c_0_2213, c_0_2191])).
cnf(c_0_2229,negated_conjecture,(until27(s39)|~xuntil28(s38)), inference(spm,[status(thm)],[c_0_115, c_0_100])).
cnf(c_0_2230,negated_conjecture,(xuntil28(s38)|node26(s38)|~xuntil28(s37)), inference(spm,[status(thm)],[c_0_111, c_0_2214])).
cnf(c_0_2231,negated_conjecture,(xuntil28(s37)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2215, c_0_2216]), c_0_2217])).
cnf(c_0_2232,plain,(trans(s59,s60)), inference(spm,[status(thm)],[c_0_8, c_0_125])).
cnf(c_0_2233,plain,(node16(s58)|node15(s58)|m_main_v_turn(s59,X1)|~m_main_v_turn(s58,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2218])).
cnf(c_0_2234,plain,(m_main_v_turn(s58,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2219, c_0_2198])]), c_0_2167])).
cnf(c_0_2235,plain,(~node16(s58)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2220, c_0_2198])])).
cnf(c_0_2236,plain,(node13(s58,s59)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2221, c_0_2222]), c_0_2223]), c_0_2224]), c_0_2225]), c_0_2226]), c_0_2227])).
cnf(c_0_2237,plain,(node16(s57)|m_main_v_turn(s58,c_2)), inference(sr,[status(thm)],[c_0_2228, c_0_2167])).
cnf(c_0_2238,negated_conjecture,(until27(s40)|~xuntil28(s39)), inference(spm,[status(thm)],[c_0_115, c_0_101])).
cnf(c_0_2239,negated_conjecture,(xuntil28(s39)|node26(s39)|~xuntil28(s38)), inference(spm,[status(thm)],[c_0_111, c_0_2229])).
cnf(c_0_2240,negated_conjecture,(xuntil28(s38)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2230, c_0_2231])).
cnf(c_0_2241,plain,(~m_main_v_state2(X1,c_t2)|~trans(s40,X1)), inference(spm,[status(thm)],[c_0_0, c_0_1599])).
cnf(c_0_2242,plain,(node17(s59,s60)|node16(s59)|node15(s59)), inference(spm,[status(thm)],[c_0_152, c_0_2232])).
cnf(c_0_2243,plain,(node15(s58)|m_main_v_turn(s59,c_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2233, c_0_2234]), c_0_2235])).
cnf(c_0_2244,plain,(~node15(s58)), inference(spm,[status(thm)],[c_0_229, c_0_2209])).
cnf(c_0_2245,plain,(m_main_v_state2(s59,X1)|~m_main_v_state2(s58,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2236])).
cnf(c_0_2246,plain,(m_main_v_state2(s58,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2196, c_0_2198])])).
cnf(c_0_2247,plain,(node16(s57)|~node5(s58)), inference(spm,[status(thm)],[c_0_354, c_0_2237])).
cnf(c_0_2248,negated_conjecture,(until27(s41)|~xuntil28(s40)), inference(spm,[status(thm)],[c_0_115, c_0_102])).
cnf(c_0_2249,negated_conjecture,(xuntil28(s40)|node26(s40)|~xuntil28(s39)), inference(spm,[status(thm)],[c_0_111, c_0_2238])).
cnf(c_0_2250,negated_conjecture,(xuntil28(s39)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2239, c_0_2240])).
cnf(c_0_2251,negated_conjecture,(~node26(s40)), inference(spm,[status(thm)],[c_0_1880, c_0_1578])).
cnf(c_0_2252,negated_conjecture,(~node26(X1)|~trans(s40,X1)), inference(spm,[status(thm)],[c_0_2241, c_0_113])).
cnf(c_0_2253,plain,(node16(s59)|node15(s59)|m_main_v_turn(s60,X1)|~m_main_v_turn(s59,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2242])).
cnf(c_0_2254,plain,(m_main_v_turn(s59,c_1)), inference(sr,[status(thm)],[c_0_2243, c_0_2244])).
cnf(c_0_2255,plain,(~node16(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_2245]), c_0_2246])])).
cnf(c_0_2256,plain,(m_main_v_state1(s58,c_c1)|~node5(s58)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2197, c_0_2247]), c_0_2199]), c_0_2200]), c_0_2201])).
cnf(c_0_2257,plain,(~node16(s57)|~node4(s58)), inference(spm,[status(thm)],[c_0_200, c_0_2196])).
cnf(c_0_2258,plain,(~node16(s57)|~node2(s58)), inference(spm,[status(thm)],[c_0_239, c_0_2196])).
cnf(c_0_2259,negated_conjecture,(until27(s42)|~xuntil28(s41)), inference(spm,[status(thm)],[c_0_115, c_0_103])).
cnf(c_0_2260,negated_conjecture,(xuntil28(s41)|node26(s41)|~xuntil28(s40)), inference(spm,[status(thm)],[c_0_111, c_0_2248])).
cnf(c_0_2261,negated_conjecture,(xuntil28(s40)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2249, c_0_2250]), c_0_2251])).
cnf(c_0_2262,negated_conjecture,(~node26(s41)), inference(spm,[status(thm)],[c_0_2252, c_0_1579])).
cnf(c_0_2263,plain,(node16(s59)|node15(s59)|m_main_v_turn(s60,c_1)), inference(spm,[status(thm)],[c_0_2253, c_0_2254])).
cnf(c_0_2264,plain,(node17(s59,s60)|node15(s59)), inference(sr,[status(thm)],[c_0_2242, c_0_2255])).
cnf(c_0_2265,plain,(m_main_v_turn(X1,c_2)|~node18(s59,X1)|~m_main_v_state1(s59,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_2245]), c_0_2246])])).
cnf(c_0_2266,plain,(~node5(s58)), inference(spm,[status(thm)],[c_0_233, c_0_2256])).
cnf(c_0_2267,plain,(~node4(s58)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2257, c_0_2198])])).
cnf(c_0_2268,plain,(~node2(s58)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2258, c_0_2198])])).
cnf(c_0_2269,plain,(~node3(s58)), inference(spm,[status(thm)],[c_0_415, c_0_2209])).
cnf(c_0_2270,plain,(~node1(s58)), inference(spm,[status(thm)],[c_0_416, c_0_2209])).
cnf(c_0_2271,negated_conjecture,(until27(s43)|~xuntil28(s42)), inference(spm,[status(thm)],[c_0_115, c_0_104])).
cnf(c_0_2272,negated_conjecture,(xuntil28(s42)|node26(s42)|~xuntil28(s41)), inference(spm,[status(thm)],[c_0_111, c_0_2259])).
cnf(c_0_2273,negated_conjecture,(xuntil28(s41)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2260, c_0_2261]), c_0_2262])).
cnf(c_0_2274,plain,(node15(s59)|m_main_v_turn(s60,c_1)), inference(sr,[status(thm)],[c_0_2263, c_0_2255])).
cnf(c_0_2275,plain,(node15(s59)|m_main_v_turn(s59,X1)|~m_main_v_turn(s60,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2264])).
cnf(c_0_2276,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s59,c_n1)|~trans(s59,X1)), inference(spm,[status(thm)],[c_0_2265, c_0_15])).
cnf(c_0_2277,plain,(~m_main_v_turn(s59,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_2254])).
cnf(c_0_2278,plain,(m_main_v_state1(X1,c_n1)|~node7(s58,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_2209]), c_0_2266]), c_0_2267]), c_0_2268]), c_0_2269]), c_0_2270])).
cnf(c_0_2279,negated_conjecture,(until27(s44)|~xuntil28(s43)), inference(spm,[status(thm)],[c_0_115, c_0_105])).
cnf(c_0_2280,negated_conjecture,(xuntil28(s43)|node26(s43)|~xuntil28(s42)), inference(spm,[status(thm)],[c_0_111, c_0_2271])).
cnf(c_0_2281,negated_conjecture,(xuntil28(s42)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2272, c_0_2273])).
cnf(c_0_2282,plain,(~m_main_v_state2(X1,c_t2)|~trans(s44,X1)), inference(spm,[status(thm)],[c_0_0, c_0_1708])).
cnf(c_0_2283,plain,(node15(s59)|~m_main_v_turn(s60,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_2274])).
cnf(c_0_2284,plain,(node15(s59)|~m_main_v_state1(s59,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2275, c_0_2276]), c_0_2232])]), c_0_2277])).
cnf(c_0_2285,plain,(m_main_v_state1(X1,c_n1)|~trans(s58,X1)), inference(spm,[status(thm)],[c_0_2278, c_0_7])).
cnf(c_0_2286,negated_conjecture,(until27(s45)|~xuntil28(s44)), inference(spm,[status(thm)],[c_0_115, c_0_106])).
cnf(c_0_2287,negated_conjecture,(xuntil28(s44)|node26(s44)|~xuntil28(s43)), inference(spm,[status(thm)],[c_0_111, c_0_2279])).
cnf(c_0_2288,negated_conjecture,(xuntil28(s43)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2280, c_0_2281])).
cnf(c_0_2289,negated_conjecture,(~node26(s44)), inference(spm,[status(thm)],[c_0_1880, c_0_1685])).
cnf(c_0_2290,negated_conjecture,(~node26(X1)|~trans(s44,X1)), inference(spm,[status(thm)],[c_0_2282, c_0_113])).
cnf(c_0_2291,plain,(node10(s59)|node9(s59)|node8(s59)|m_main_v_state2(s60,c_c2)|~node15(s59)), inference(spm,[status(thm)],[c_0_467, c_0_2232])).
cnf(c_0_2292,plain,(node15(s59)|~node12(s60)), inference(spm,[status(thm)],[c_0_2283, c_0_35])).
cnf(c_0_2293,plain,(~node10(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_2245]), c_0_2246])])).
cnf(c_0_2294,plain,(~node9(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_2245]), c_0_2246])])).
cnf(c_0_2295,plain,(~node8(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_2245]), c_0_2246])])).
cnf(c_0_2296,plain,(m_main_v_state1(s60,c_t1)|~node15(s59)), inference(spm,[status(thm)],[c_0_474, c_0_2232])).
cnf(c_0_2297,plain,(node15(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2284, c_0_2285]), c_0_2207])])).
cnf(c_0_2298,plain,(m_main_v_turn(s60,c_2)|~node15(s59)), inference(spm,[status(thm)],[c_0_470, c_0_2232])).
cnf(c_0_2299,negated_conjecture,(until27(s46)|~xuntil28(s45)), inference(spm,[status(thm)],[c_0_115, c_0_107])).
cnf(c_0_2300,negated_conjecture,(xuntil28(s45)|node26(s45)|~xuntil28(s44)), inference(spm,[status(thm)],[c_0_111, c_0_2286])).
cnf(c_0_2301,negated_conjecture,(xuntil28(s44)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2287, c_0_2288]), c_0_2289])).
cnf(c_0_2302,negated_conjecture,(~node26(s45)), inference(spm,[status(thm)],[c_0_2290, c_0_1687])).
cnf(c_0_2303,plain,(m_main_v_state2(s60,c_c2)|~node12(s60)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2291, c_0_2292]), c_0_2293]), c_0_2294]), c_0_2295])).
cnf(c_0_2304,plain,(~node15(s59)|~node11(s60)), inference(spm,[status(thm)],[c_0_165, c_0_2296])).
cnf(c_0_2305,plain,(~node15(s59)|~node9(s60)), inference(spm,[status(thm)],[c_0_167, c_0_2296])).
cnf(c_0_2306,plain,(m_main_v_state2(s60,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2291, c_0_2297])]), c_0_2293]), c_0_2294]), c_0_2295])).
cnf(c_0_2307,plain,(trans(s60,s61)), inference(spm,[status(thm)],[c_0_8, c_0_126])).
cnf(c_0_2308,plain,(~node15(s59)|~node5(s60)), inference(spm,[status(thm)],[c_0_354, c_0_2298])).
cnf(c_0_2309,plain,(~node15(s59)|~node3(s60)), inference(spm,[status(thm)],[c_0_201, c_0_2296])).
cnf(c_0_2310,plain,(~node15(s59)|~node2(s60)), inference(spm,[status(thm)],[c_0_202, c_0_2296])).
cnf(c_0_2311,plain,(~node15(s59)|~node1(s60)), inference(spm,[status(thm)],[c_0_203, c_0_2296])).
cnf(c_0_2312,plain,(~node15(s59)|~m_main_v_state1(s60,c_c1)), inference(spm,[status(thm)],[c_0_26, c_0_2296])).
cnf(c_0_2313,negated_conjecture,(until27(s47)|~xuntil28(s46)), inference(spm,[status(thm)],[c_0_115, c_0_108])).
cnf(c_0_2314,negated_conjecture,(xuntil28(s46)|node26(s46)|~xuntil28(s45)), inference(spm,[status(thm)],[c_0_111, c_0_2299])).
cnf(c_0_2315,negated_conjecture,(xuntil28(s45)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2300, c_0_2301]), c_0_2302])).
cnf(c_0_2316,plain,(~node12(s60)), inference(spm,[status(thm)],[c_0_345, c_0_2303])).
cnf(c_0_2317,plain,(~node11(s60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2304, c_0_2297])])).
cnf(c_0_2318,plain,(~node9(s60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2305, c_0_2297])])).
cnf(c_0_2319,plain,(~node10(s60)), inference(spm,[status(thm)],[c_0_339, c_0_2306])).
cnf(c_0_2320,plain,(~node8(s60)), inference(spm,[status(thm)],[c_0_343, c_0_2306])).
cnf(c_0_2321,plain,(node6(s60,s61)|node5(s60)|node4(s60)|node3(s60)|node2(s60)|node1(s60)|m_main_v_state1(s60,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_2307])).
cnf(c_0_2322,plain,(~node5(s60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2308, c_0_2297])])).
cnf(c_0_2323,plain,(~node4(s60)), inference(spm,[status(thm)],[c_0_361, c_0_2306])).
cnf(c_0_2324,plain,(~node3(s60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2309, c_0_2297])])).
cnf(c_0_2325,plain,(~node2(s60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2310, c_0_2297])])).
cnf(c_0_2326,plain,(~node1(s60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2311, c_0_2297])])).
cnf(c_0_2327,plain,(~m_main_v_state1(s60,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2312, c_0_2297])])).
cnf(c_0_2328,negated_conjecture,(until27(s48)|~xuntil28(s47)), inference(spm,[status(thm)],[c_0_115, c_0_109])).
cnf(c_0_2329,negated_conjecture,(xuntil28(s47)|node26(s47)|~xuntil28(s46)), inference(spm,[status(thm)],[c_0_111, c_0_2313])).
cnf(c_0_2330,negated_conjecture,(xuntil28(s46)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2314, c_0_2315])).
cnf(c_0_2331,plain,(~m_main_v_state2(X1,c_t2)|~trans(s48,X1)), inference(spm,[status(thm)],[c_0_0, c_0_1818])).
cnf(c_0_2332,plain,(m_main_v_state2(X1,c_n2)|~node14(s60,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_2306]), c_0_2316]), c_0_2317]), c_0_2318]), c_0_2319]), c_0_2320])).
cnf(c_0_2333,plain,(node6(s60,s61)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2321, c_0_2322]), c_0_2323]), c_0_2324]), c_0_2325]), c_0_2326]), c_0_2327])).
cnf(c_0_2334,negated_conjecture,(until27(s49)|~xuntil28(s48)), inference(spm,[status(thm)],[c_0_115, c_0_110])).
cnf(c_0_2335,negated_conjecture,(xuntil28(s48)|node26(s48)|~xuntil28(s47)), inference(spm,[status(thm)],[c_0_111, c_0_2328])).
cnf(c_0_2336,negated_conjecture,(xuntil28(s47)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2329, c_0_2330])).
cnf(c_0_2337,negated_conjecture,(~node26(s48)), inference(spm,[status(thm)],[c_0_1880, c_0_1795])).
cnf(c_0_2338,negated_conjecture,(~node26(X1)|~trans(s48,X1)), inference(spm,[status(thm)],[c_0_2331, c_0_113])).
cnf(c_0_2339,plain,(m_main_v_state2(X1,c_n2)|~trans(s60,X1)), inference(spm,[status(thm)],[c_0_2332, c_0_29])).
cnf(c_0_2340,plain,(m_main_v_state1(s61,X1)|~m_main_v_state1(s60,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2333])).
cnf(c_0_2341,plain,(m_main_v_state1(s60,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2296, c_0_2297])])).
cnf(c_0_2342,negated_conjecture,(until27(s50)|~xuntil28(s49)), inference(spm,[status(thm)],[c_0_115, c_0_114])).
cnf(c_0_2343,negated_conjecture,(xuntil28(s49)|node26(s49)|~xuntil28(s48)), inference(spm,[status(thm)],[c_0_111, c_0_2334])).
cnf(c_0_2344,negated_conjecture,(xuntil28(s48)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2335, c_0_2336]), c_0_2337])).
cnf(c_0_2345,negated_conjecture,(~node26(s49)), inference(spm,[status(thm)],[c_0_2338, c_0_1797])).
cnf(c_0_2346,plain,(trans(s61,s62)), inference(spm,[status(thm)],[c_0_8, c_0_127])).
cnf(c_0_2347,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s60,X1)), inference(spm,[status(thm)],[c_0_63, c_0_2339])).
cnf(c_0_2348,plain,(~node15(s61)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_2340]), c_0_2341])])).
cnf(c_0_2349,plain,(node17(s60,s61)|node16(s60)|node15(s60)), inference(spm,[status(thm)],[c_0_152, c_0_2307])).
cnf(c_0_2350,plain,(~node15(s60)|~node15(s59)), inference(spm,[status(thm)],[c_0_189, c_0_2296])).
cnf(c_0_2351,negated_conjecture,(until27(s51)|~xuntil28(s50)), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_2352,negated_conjecture,(xuntil28(s50)|node26(s50)|~xuntil28(s49)), inference(spm,[status(thm)],[c_0_111, c_0_2342])).
cnf(c_0_2353,negated_conjecture,(xuntil28(s49)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2343, c_0_2344]), c_0_2345])).
cnf(c_0_2354,plain,(node17(s61,s62)|node16(s61)|node15(s61)), inference(spm,[status(thm)],[c_0_152, c_0_2346])).
cnf(c_0_2355,plain,(m_main_v_turn(X1,c_1)|~node18(s61,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2347, c_0_2340]), c_0_2307]), c_0_2341])]), c_0_2348])).
cnf(c_0_2356,plain,(node16(s60)|node15(s60)|m_main_v_turn(s61,X1)|~m_main_v_turn(s60,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2349])).
cnf(c_0_2357,plain,(m_main_v_turn(s60,c_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2298, c_0_2297])])).
cnf(c_0_2358,plain,(~node15(s60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2350, c_0_2297])])).
cnf(c_0_2359,negated_conjecture,(until27(s52)|~xuntil28(s51)), inference(spm,[status(thm)],[c_0_115, c_0_117])).
cnf(c_0_2360,negated_conjecture,(xuntil28(s51)|node26(s51)|~xuntil28(s50)), inference(spm,[status(thm)],[c_0_111, c_0_2351])).
cnf(c_0_2361,negated_conjecture,(xuntil28(s50)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2352, c_0_2353])).
cnf(c_0_2362,plain,(~m_main_v_state2(X1,c_t2)|~trans(s52,X1)), inference(spm,[status(thm)],[c_0_0, c_0_1978])).
cnf(c_0_2363,plain,(node16(s61)|node15(s61)|m_main_v_turn(s61,X1)|~m_main_v_turn(s62,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2354])).
cnf(c_0_2364,plain,(m_main_v_turn(X1,c_1)|~trans(s61,X1)), inference(spm,[status(thm)],[c_0_2355, c_0_15])).
cnf(c_0_2365,plain,(node16(s60)|m_main_v_turn(s61,c_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2356, c_0_2357]), c_0_2358])).
cnf(c_0_2366,plain,(~node16(s60)), inference(spm,[status(thm)],[c_0_392, c_0_2306])).
cnf(c_0_2367,negated_conjecture,(until27(s53)|~xuntil28(s52)), inference(spm,[status(thm)],[c_0_115, c_0_118])).
cnf(c_0_2368,negated_conjecture,(xuntil28(s52)|node26(s52)|~xuntil28(s51)), inference(spm,[status(thm)],[c_0_111, c_0_2359])).
cnf(c_0_2369,negated_conjecture,(xuntil28(s51)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2360, c_0_2361])).
cnf(c_0_2370,negated_conjecture,(~node26(s52)), inference(spm,[status(thm)],[c_0_1880, c_0_1946])).
cnf(c_0_2371,negated_conjecture,(~node26(X1)|~trans(s52,X1)), inference(spm,[status(thm)],[c_0_2362, c_0_113])).
cnf(c_0_2372,plain,(node16(s61)|m_main_v_turn(s61,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2363, c_0_2364]), c_0_2346])]), c_0_2348])).
cnf(c_0_2373,plain,(m_main_v_turn(s61,c_2)), inference(sr,[status(thm)],[c_0_2365, c_0_2366])).
cnf(c_0_2374,negated_conjecture,(until27(s54)|~xuntil28(s53)), inference(spm,[status(thm)],[c_0_115, c_0_119])).
cnf(c_0_2375,negated_conjecture,(xuntil28(s53)|node26(s53)|~xuntil28(s52)), inference(spm,[status(thm)],[c_0_111, c_0_2367])).
cnf(c_0_2376,negated_conjecture,(xuntil28(s52)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2368, c_0_2369]), c_0_2370])).
cnf(c_0_2377,negated_conjecture,(~node26(s53)), inference(spm,[status(thm)],[c_0_2371, c_0_1947])).
cnf(c_0_2378,plain,(m_main_v_state2(s62,c_t2)|~node16(s61)), inference(spm,[status(thm)],[c_0_400, c_0_2346])).
cnf(c_0_2379,plain,(node3(s61)|node2(s61)|node1(s61)|m_main_v_state1(s62,c_c1)|~node16(s61)), inference(spm,[status(thm)],[c_0_401, c_0_2346])).
cnf(c_0_2380,plain,(node16(s61)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2372]), c_0_2373])])).
cnf(c_0_2381,plain,(~node3(s61)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_2340]), c_0_2341])])).
cnf(c_0_2382,plain,(~node2(s61)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_2340]), c_0_2341])])).
cnf(c_0_2383,plain,(~node1(s61)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_2340]), c_0_2341])])).
cnf(c_0_2384,negated_conjecture,(until27(s55)|~xuntil28(s54)), inference(spm,[status(thm)],[c_0_115, c_0_120])).
cnf(c_0_2385,negated_conjecture,(xuntil28(s54)|node26(s54)|~xuntil28(s53)), inference(spm,[status(thm)],[c_0_111, c_0_2374])).
cnf(c_0_2386,negated_conjecture,(xuntil28(s53)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2375, c_0_2376]), c_0_2377])).
cnf(c_0_2387,plain,(trans(s62,s63)), inference(spm,[status(thm)],[c_0_8, c_0_128])).
cnf(c_0_2388,plain,(~node16(s61)|~node8(s62)), inference(spm,[status(thm)],[c_0_168, c_0_2378])).
cnf(c_0_2389,plain,(m_main_v_state1(s62,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2379, c_0_2380])]), c_0_2381]), c_0_2382]), c_0_2383])).
cnf(c_0_2390,plain,(~node16(s61)|~node10(s62)), inference(spm,[status(thm)],[c_0_257, c_0_2378])).
cnf(c_0_2391,plain,(~node16(s61)|~node9(s62)), inference(spm,[status(thm)],[c_0_258, c_0_2378])).
cnf(c_0_2392,plain,(~node16(s61)|~m_main_v_state2(s62,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_2378])).
cnf(c_0_2393,plain,(node16(s61)|node15(s61)|m_main_v_turn(s62,X1)|~m_main_v_turn(s61,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2354])).
cnf(c_0_2394,negated_conjecture,(until27(s56)|~xuntil28(s55)), inference(spm,[status(thm)],[c_0_115, c_0_121])).
cnf(c_0_2395,negated_conjecture,(xuntil28(s55)|node26(s55)|~xuntil28(s54)), inference(spm,[status(thm)],[c_0_111, c_0_2384])).
cnf(c_0_2396,negated_conjecture,(xuntil28(s54)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2385, c_0_2386])).
cnf(c_0_2397,plain,(~m_main_v_state2(X1,c_t2)|~trans(s56,X1)), inference(spm,[status(thm)],[c_0_0, c_0_2157])).
cnf(c_0_2398,plain,(node13(s62,s63)|node12(s62)|node11(s62)|node10(s62)|node9(s62)|node8(s62)|m_main_v_state2(s62,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_2387])).
cnf(c_0_2399,plain,(~node8(s62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2388, c_0_2380])])).
cnf(c_0_2400,plain,(~node12(s62)), inference(spm,[status(thm)],[c_0_428, c_0_2389])).
cnf(c_0_2401,plain,(~node11(s62)), inference(spm,[status(thm)],[c_0_429, c_0_2389])).
cnf(c_0_2402,plain,(~node10(s62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2390, c_0_2380])])).
cnf(c_0_2403,plain,(~node9(s62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2391, c_0_2380])])).
cnf(c_0_2404,plain,(~m_main_v_state2(s62,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2392, c_0_2380])])).
cnf(c_0_2405,plain,(node16(s61)|node15(s61)|m_main_v_turn(s62,c_2)), inference(spm,[status(thm)],[c_0_2393, c_0_2373])).
cnf(c_0_2406,negated_conjecture,(until27(s57)|~xuntil28(s56)), inference(spm,[status(thm)],[c_0_115, c_0_122])).
cnf(c_0_2407,negated_conjecture,(xuntil28(s56)|node26(s56)|~xuntil28(s55)), inference(spm,[status(thm)],[c_0_111, c_0_2394])).
cnf(c_0_2408,negated_conjecture,(xuntil28(s55)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2395, c_0_2396])).
cnf(c_0_2409,negated_conjecture,(~node26(s56)), inference(spm,[status(thm)],[c_0_1880, c_0_2122])).
cnf(c_0_2410,negated_conjecture,(~node26(X1)|~trans(s56,X1)), inference(spm,[status(thm)],[c_0_2397, c_0_113])).
cnf(c_0_2411,plain,(node13(s62,s63)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2398, c_0_2399]), c_0_2400]), c_0_2401]), c_0_2402]), c_0_2403]), c_0_2404])).
cnf(c_0_2412,plain,(node17(s62,s63)|node16(s62)|node15(s62)), inference(spm,[status(thm)],[c_0_152, c_0_2387])).
cnf(c_0_2413,plain,(node15(s61)|m_main_v_turn(s62,c_1)|~node16(s61)), inference(spm,[status(thm)],[c_0_433, c_0_2346])).
cnf(c_0_2414,plain,(~node16(s62)|~node16(s61)), inference(spm,[status(thm)],[c_0_188, c_0_2378])).
cnf(c_0_2415,plain,(node16(s61)|m_main_v_turn(s62,c_2)), inference(sr,[status(thm)],[c_0_2405, c_0_2348])).
cnf(c_0_2416,negated_conjecture,(until27(s58)|~xuntil28(s57)), inference(spm,[status(thm)],[c_0_115, c_0_123])).
cnf(c_0_2417,negated_conjecture,(xuntil28(s57)|node26(s57)|~xuntil28(s56)), inference(spm,[status(thm)],[c_0_111, c_0_2406])).
cnf(c_0_2418,negated_conjecture,(xuntil28(s56)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2407, c_0_2408]), c_0_2409])).
cnf(c_0_2419,negated_conjecture,(~node26(s57)), inference(spm,[status(thm)],[c_0_2410, c_0_2124])).
cnf(c_0_2420,plain,(trans(s63,s64)), inference(spm,[status(thm)],[c_0_8, c_0_129])).
cnf(c_0_2421,plain,(m_main_v_state2(s63,X1)|~m_main_v_state2(s62,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2411])).
cnf(c_0_2422,plain,(m_main_v_state2(s62,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2378, c_0_2380])])).
cnf(c_0_2423,plain,(node16(s62)|node15(s62)|m_main_v_turn(s63,X1)|~m_main_v_turn(s62,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2412])).
cnf(c_0_2424,plain,(m_main_v_turn(s62,c_1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2413, c_0_2380])]), c_0_2348])).
cnf(c_0_2425,plain,(~node16(s62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2414, c_0_2380])])).
cnf(c_0_2426,plain,(node16(s61)|~node5(s62)), inference(spm,[status(thm)],[c_0_354, c_0_2415])).
cnf(c_0_2427,negated_conjecture,(until27(s59)|~xuntil28(s58)), inference(spm,[status(thm)],[c_0_115, c_0_124])).
cnf(c_0_2428,negated_conjecture,(xuntil28(s58)|node26(s58)|~xuntil28(s57)), inference(spm,[status(thm)],[c_0_111, c_0_2416])).
cnf(c_0_2429,negated_conjecture,(xuntil28(s57)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2417, c_0_2418]), c_0_2419])).
cnf(c_0_2430,plain,(node17(s63,s64)|node16(s63)|node15(s63)), inference(spm,[status(thm)],[c_0_152, c_0_2420])).
cnf(c_0_2431,plain,(~node16(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_2421]), c_0_2422])])).
cnf(c_0_2432,plain,(node15(s62)|m_main_v_turn(s63,c_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2423, c_0_2424]), c_0_2425])).
cnf(c_0_2433,plain,(~node15(s62)), inference(spm,[status(thm)],[c_0_229, c_0_2389])).
cnf(c_0_2434,plain,(m_main_v_state1(s62,c_c1)|~node5(s62)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2379, c_0_2426]), c_0_2381]), c_0_2382]), c_0_2383])).
cnf(c_0_2435,plain,(~node16(s61)|~node4(s62)), inference(spm,[status(thm)],[c_0_200, c_0_2378])).
cnf(c_0_2436,plain,(~node16(s61)|~node2(s62)), inference(spm,[status(thm)],[c_0_239, c_0_2378])).
cnf(c_0_2437,negated_conjecture,(until27(s60)|~xuntil28(s59)), inference(spm,[status(thm)],[c_0_115, c_0_125])).
cnf(c_0_2438,negated_conjecture,(xuntil28(s59)|node26(s59)|~xuntil28(s58)), inference(spm,[status(thm)],[c_0_111, c_0_2427])).
cnf(c_0_2439,negated_conjecture,(xuntil28(s58)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2428, c_0_2429])).
cnf(c_0_2440,plain,(~m_main_v_state2(X1,c_t2)|~trans(s60,X1)), inference(spm,[status(thm)],[c_0_0, c_0_2339])).
cnf(c_0_2441,plain,(node17(s63,s64)|node15(s63)), inference(sr,[status(thm)],[c_0_2430, c_0_2431])).
cnf(c_0_2442,plain,(m_main_v_turn(X1,c_2)|~node18(s63,X1)|~m_main_v_state1(s63,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_2421]), c_0_2422])])).
cnf(c_0_2443,plain,(m_main_v_turn(s63,c_1)), inference(sr,[status(thm)],[c_0_2432, c_0_2433])).
cnf(c_0_2444,plain,(~node5(s62)), inference(spm,[status(thm)],[c_0_233, c_0_2434])).
cnf(c_0_2445,plain,(~node4(s62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2435, c_0_2380])])).
cnf(c_0_2446,plain,(~node2(s62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2436, c_0_2380])])).
cnf(c_0_2447,plain,(~node3(s62)), inference(spm,[status(thm)],[c_0_415, c_0_2389])).
cnf(c_0_2448,plain,(~node1(s62)), inference(spm,[status(thm)],[c_0_416, c_0_2389])).
cnf(c_0_2449,negated_conjecture,(until27(s61)|~xuntil28(s60)), inference(spm,[status(thm)],[c_0_115, c_0_126])).
cnf(c_0_2450,negated_conjecture,(xuntil28(s60)|node26(s60)|~xuntil28(s59)), inference(spm,[status(thm)],[c_0_111, c_0_2437])).
cnf(c_0_2451,negated_conjecture,(xuntil28(s59)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2438, c_0_2439])).
cnf(c_0_2452,negated_conjecture,(~node26(s60)), inference(spm,[status(thm)],[c_0_1880, c_0_2306])).
cnf(c_0_2453,negated_conjecture,(~node26(X1)|~trans(s60,X1)), inference(spm,[status(thm)],[c_0_2440, c_0_113])).
cnf(c_0_2454,plain,(node15(s63)|m_main_v_turn(s63,X1)|~m_main_v_turn(s64,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2441])).
cnf(c_0_2455,plain,(m_main_v_turn(X1,c_2)|~m_main_v_state1(s63,c_n1)|~trans(s63,X1)), inference(spm,[status(thm)],[c_0_2442, c_0_15])).
cnf(c_0_2456,plain,(~m_main_v_turn(s63,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_2443])).
cnf(c_0_2457,plain,(m_main_v_state1(X1,c_n1)|~node7(s62,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_2389]), c_0_2444]), c_0_2445]), c_0_2446]), c_0_2447]), c_0_2448])).
cnf(c_0_2458,negated_conjecture,(until27(s62)|~xuntil28(s61)), inference(spm,[status(thm)],[c_0_115, c_0_127])).
cnf(c_0_2459,negated_conjecture,(xuntil28(s61)|node26(s61)|~xuntil28(s60)), inference(spm,[status(thm)],[c_0_111, c_0_2449])).
cnf(c_0_2460,negated_conjecture,(xuntil28(s60)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2450, c_0_2451]), c_0_2452])).
cnf(c_0_2461,negated_conjecture,(~node26(s61)), inference(spm,[status(thm)],[c_0_2453, c_0_2307])).
cnf(c_0_2462,plain,(node15(s63)|~m_main_v_state1(s63,c_n1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2454, c_0_2455]), c_0_2420])]), c_0_2456])).
cnf(c_0_2463,plain,(m_main_v_state1(X1,c_n1)|~trans(s62,X1)), inference(spm,[status(thm)],[c_0_2457, c_0_7])).
cnf(c_0_2464,negated_conjecture,(until27(s63)|~xuntil28(s62)), inference(spm,[status(thm)],[c_0_115, c_0_128])).
cnf(c_0_2465,negated_conjecture,(xuntil28(s62)|node26(s62)|~xuntil28(s61)), inference(spm,[status(thm)],[c_0_111, c_0_2458])).
cnf(c_0_2466,negated_conjecture,(xuntil28(s61)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2459, c_0_2460]), c_0_2461])).
cnf(c_0_2467,plain,(node10(s63)|node9(s63)|node8(s63)|m_main_v_state2(s64,c_c2)|~node15(s63)), inference(spm,[status(thm)],[c_0_467, c_0_2420])).
cnf(c_0_2468,plain,(node15(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2462, c_0_2463]), c_0_2387])])).
cnf(c_0_2469,plain,(~node10(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_2421]), c_0_2422])])).
cnf(c_0_2470,plain,(~node9(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_2421]), c_0_2422])])).
cnf(c_0_2471,plain,(~node8(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_2421]), c_0_2422])])).
cnf(c_0_2472,negated_conjecture,(until27(s64)|~xuntil28(s63)), inference(spm,[status(thm)],[c_0_115, c_0_129])).
cnf(c_0_2473,negated_conjecture,(xuntil28(s63)|node26(s63)|~xuntil28(s62)), inference(spm,[status(thm)],[c_0_111, c_0_2464])).
cnf(c_0_2474,negated_conjecture,(xuntil28(s62)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2465, c_0_2466])).
cnf(c_0_2475,plain,(m_main_v_state2(s64,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2467, c_0_2468])]), c_0_2469]), c_0_2470]), c_0_2471])).
cnf(c_0_2476,negated_conjecture,(xuntil28(s64)|node26(s64)|~xuntil28(s63)), inference(spm,[status(thm)],[c_0_111, c_0_2472])).
cnf(c_0_2477,negated_conjecture,(xuntil28(s63)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_2473, c_0_2474])).
cnf(c_0_2478,negated_conjecture,(~node26(s64)), inference(spm,[status(thm)],[c_0_1880, c_0_2475])).
cnf(c_0_2479,negated_conjecture,(xuntil28(s64)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2476, c_0_2477]), c_0_2478])).
cnf(c_0_2480,negated_conjecture,(node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|loop), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_2479]), c_0_131])])).
cnf(c_0_2481,negated_conjecture,(node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s63)|loop), inference(spm,[status(thm)],[c_0_132, c_0_2480])).
cnf(c_0_2482,negated_conjecture,(~always25(s60)), inference(spm,[status(thm)],[c_0_133, c_0_2306])).
cnf(c_0_2483,negated_conjecture,(always25(s60)|~always25(s59)), inference(spm,[status(thm)],[c_0_134, c_0_2232])).
cnf(c_0_2484,negated_conjecture,(node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(spm,[status(thm)],[c_0_132, c_0_2481])).
cnf(c_0_2485,negated_conjecture,(~always25(s59)), inference(spm,[status(thm)],[c_0_2482, c_0_2483])).
cnf(c_0_2486,negated_conjecture,(always25(s59)|~always25(s58)), inference(spm,[status(thm)],[c_0_134, c_0_2207])).
cnf(c_0_2487,negated_conjecture,(node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2484]), c_0_2485])).
cnf(c_0_2488,negated_conjecture,(~always25(s58)), inference(spm,[status(thm)],[c_0_2485, c_0_2486])).
cnf(c_0_2489,negated_conjecture,(~always25(s56)), inference(spm,[status(thm)],[c_0_133, c_0_2122])).
cnf(c_0_2490,negated_conjecture,(always25(s56)|~always25(s55)), inference(spm,[status(thm)],[c_0_134, c_0_2052])).
cnf(c_0_2491,negated_conjecture,(node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2487]), c_0_2488])).
cnf(c_0_2492,negated_conjecture,(~always25(s55)), inference(spm,[status(thm)],[c_0_2489, c_0_2490])).
cnf(c_0_2493,negated_conjecture,(always25(s55)|~always25(s54)), inference(spm,[status(thm)],[c_0_134, c_0_2026])).
cnf(c_0_2494,negated_conjecture,(node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2491]), c_0_2492])).
cnf(c_0_2495,negated_conjecture,(~always25(s54)), inference(spm,[status(thm)],[c_0_2492, c_0_2493])).
cnf(c_0_2496,negated_conjecture,(~always25(s52)), inference(spm,[status(thm)],[c_0_133, c_0_1946])).
cnf(c_0_2497,negated_conjecture,(always25(s52)|~always25(s51)), inference(spm,[status(thm)],[c_0_134, c_0_1881])).
cnf(c_0_2498,negated_conjecture,(node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2494]), c_0_2495])).
cnf(c_0_2499,negated_conjecture,(~always25(s51)), inference(spm,[status(thm)],[c_0_2496, c_0_2497])).
cnf(c_0_2500,negated_conjecture,(always25(s51)|~always25(s50)), inference(spm,[status(thm)],[c_0_134, c_0_1857])).
cnf(c_0_2501,negated_conjecture,(node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2498]), c_0_2499])).
cnf(c_0_2502,negated_conjecture,(~always25(s50)), inference(spm,[status(thm)],[c_0_2499, c_0_2500])).
cnf(c_0_2503,negated_conjecture,(~always25(s48)), inference(spm,[status(thm)],[c_0_133, c_0_1795])).
cnf(c_0_2504,negated_conjecture,(always25(s48)|~always25(s47)), inference(spm,[status(thm)],[c_0_134, c_0_1752])).
cnf(c_0_2505,negated_conjecture,(node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2501]), c_0_2502])).
cnf(c_0_2506,negated_conjecture,(~always25(s47)), inference(spm,[status(thm)],[c_0_2503, c_0_2504])).
cnf(c_0_2507,negated_conjecture,(always25(s47)|~always25(s46)), inference(spm,[status(thm)],[c_0_134, c_0_1727])).
cnf(c_0_2508,negated_conjecture,(node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2505]), c_0_2506])).
cnf(c_0_2509,negated_conjecture,(~always25(s46)), inference(spm,[status(thm)],[c_0_2506, c_0_2507])).
cnf(c_0_2510,negated_conjecture,(~always25(s44)), inference(spm,[status(thm)],[c_0_133, c_0_1685])).
cnf(c_0_2511,negated_conjecture,(always25(s44)|~always25(s43)), inference(spm,[status(thm)],[c_0_134, c_0_1642])).
cnf(c_0_2512,negated_conjecture,(node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2508]), c_0_2509])).
cnf(c_0_2513,negated_conjecture,(~always25(s43)), inference(spm,[status(thm)],[c_0_2510, c_0_2511])).
cnf(c_0_2514,negated_conjecture,(always25(s43)|~always25(s42)), inference(spm,[status(thm)],[c_0_134, c_0_1618])).
cnf(c_0_2515,negated_conjecture,(node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2512]), c_0_2513])).
cnf(c_0_2516,negated_conjecture,(~always25(s42)), inference(spm,[status(thm)],[c_0_2513, c_0_2514])).
cnf(c_0_2517,negated_conjecture,(~always25(s40)), inference(spm,[status(thm)],[c_0_133, c_0_1578])).
cnf(c_0_2518,negated_conjecture,(always25(s40)|~always25(s39)), inference(spm,[status(thm)],[c_0_134, c_0_1533])).
cnf(c_0_2519,negated_conjecture,(node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2515]), c_0_2516])).
cnf(c_0_2520,negated_conjecture,(~always25(s39)), inference(spm,[status(thm)],[c_0_2517, c_0_2518])).
cnf(c_0_2521,negated_conjecture,(always25(s39)|~always25(s38)), inference(spm,[status(thm)],[c_0_134, c_0_1515])).
cnf(c_0_2522,negated_conjecture,(node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2519]), c_0_2520])).
cnf(c_0_2523,negated_conjecture,(~always25(s38)), inference(spm,[status(thm)],[c_0_2520, c_0_2521])).
cnf(c_0_2524,negated_conjecture,(~always25(s36)), inference(spm,[status(thm)],[c_0_133, c_0_1468])).
cnf(c_0_2525,negated_conjecture,(always25(s36)|~always25(s35)), inference(spm,[status(thm)],[c_0_134, c_0_1423])).
cnf(c_0_2526,negated_conjecture,(node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2522]), c_0_2523])).
cnf(c_0_2527,negated_conjecture,(~always25(s35)), inference(spm,[status(thm)],[c_0_2524, c_0_2525])).
cnf(c_0_2528,negated_conjecture,(always25(s35)|~always25(s34)), inference(spm,[status(thm)],[c_0_134, c_0_1405])).
cnf(c_0_2529,negated_conjecture,(node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2526]), c_0_2527])).
cnf(c_0_2530,negated_conjecture,(~always25(s34)), inference(spm,[status(thm)],[c_0_2527, c_0_2528])).
cnf(c_0_2531,negated_conjecture,(~always25(s32)), inference(spm,[status(thm)],[c_0_133, c_0_1356])).
cnf(c_0_2532,negated_conjecture,(always25(s32)|~always25(s31)), inference(spm,[status(thm)],[c_0_134, c_0_1313])).
cnf(c_0_2533,negated_conjecture,(node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2529]), c_0_2530])).
cnf(c_0_2534,negated_conjecture,(~always25(s31)), inference(spm,[status(thm)],[c_0_2531, c_0_2532])).
cnf(c_0_2535,negated_conjecture,(always25(s31)|~always25(s30)), inference(spm,[status(thm)],[c_0_134, c_0_1289])).
cnf(c_0_2536,negated_conjecture,(~always25(s28)|~node11(s27)), inference(spm,[status(thm)],[c_0_133, c_0_1238])).
cnf(c_0_2537,negated_conjecture,(always25(s28)|~always25(s27)), inference(spm,[status(thm)],[c_0_134, c_0_1177])).
cnf(c_0_2538,plain,(node11(s27)|~m_main_v_state2(s28,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_1215])).
cnf(c_0_2539,negated_conjecture,(node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2533]), c_0_2534])).
cnf(c_0_2540,negated_conjecture,(~always25(s30)), inference(spm,[status(thm)],[c_0_2534, c_0_2535])).
cnf(c_0_2541,negated_conjecture,(~always25(s27)|~node11(s27)), inference(spm,[status(thm)],[c_0_2536, c_0_2537])).
cnf(c_0_2542,plain,(node11(s27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2538, c_0_1242])])).
cnf(c_0_2543,negated_conjecture,(always25(s27)|~always25(s26)), inference(spm,[status(thm)],[c_0_134, c_0_1141])).
cnf(c_0_2544,negated_conjecture,(node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2539]), c_0_2540])).
cnf(c_0_2545,negated_conjecture,(~always25(s27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2541, c_0_2542])])).
cnf(c_0_2546,negated_conjecture,(~always25(s26)|~node11(s27)), inference(spm,[status(thm)],[c_0_2541, c_0_2543])).
cnf(c_0_2547,negated_conjecture,(node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2544]), c_0_2545])).
cnf(c_0_2548,negated_conjecture,(~always25(s26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2546, c_0_2542])])).
cnf(c_0_2549,negated_conjecture,(~always25(s24)), inference(spm,[status(thm)],[c_0_133, c_0_1098])).
cnf(c_0_2550,negated_conjecture,(always25(s24)|~always25(s23)), inference(spm,[status(thm)],[c_0_134, c_0_1040])).
cnf(c_0_2551,negated_conjecture,(node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2547]), c_0_2548])).
cnf(c_0_2552,negated_conjecture,(~always25(s23)), inference(spm,[status(thm)],[c_0_2549, c_0_2550])).
cnf(c_0_2553,negated_conjecture,(always25(s23)|~always25(s22)), inference(spm,[status(thm)],[c_0_134, c_0_1004])).
cnf(c_0_2554,negated_conjecture,(node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2551]), c_0_2552])).
cnf(c_0_2555,negated_conjecture,(~always25(s22)), inference(spm,[status(thm)],[c_0_2552, c_0_2553])).
cnf(c_0_2556,negated_conjecture,(~always25(s20)), inference(spm,[status(thm)],[c_0_133, c_0_948])).
cnf(c_0_2557,negated_conjecture,(always25(s20)|~always25(s19)), inference(spm,[status(thm)],[c_0_134, c_0_912])).
cnf(c_0_2558,plain,(node16(s63)|node15(s63)|m_main_v_turn(s64,X1)|~m_main_v_turn(s63,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2430])).
cnf(c_0_2559,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node5(X2)|~m_main_v_state1(X2,c_n1)), inference(spm,[status(thm)],[c_0_45, c_0_135])).
cnf(c_0_2560,plain,(node15(s62)|~m_main_v_turn(s63,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_2432])).
cnf(c_0_2561,negated_conjecture,(node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2554]), c_0_2555])).
cnf(c_0_2562,negated_conjecture,(~always25(s19)), inference(spm,[status(thm)],[c_0_2556, c_0_2557])).
cnf(c_0_2563,negated_conjecture,(always25(s19)|~always25(s18)), inference(spm,[status(thm)],[c_0_134, c_0_895])).
cnf(c_0_2564,negated_conjecture,(~always25(s16)|~node11(s15)), inference(spm,[status(thm)],[c_0_133, c_0_822])).
cnf(c_0_2565,negated_conjecture,(always25(s16)|~always25(s15)), inference(spm,[status(thm)],[c_0_134, c_0_753])).
cnf(c_0_2566,plain,(node11(s15)|~m_main_v_state2(s16,c_c2)), inference(spm,[status(thm)],[c_0_36, c_0_791])).
cnf(c_0_2567,plain,(node16(s63)|node15(s63)|m_main_v_turn(s64,c_1)), inference(spm,[status(thm)],[c_0_2558, c_0_2443])).
cnf(c_0_2568,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node5(X2)|~trans(s62,X2)), inference(spm,[status(thm)],[c_0_2559, c_0_2463])).
cnf(c_0_2569,plain,(node10(s62)|node9(s62)|node8(s62)|m_main_v_state2(s63,c_c2)|~node15(s62)), inference(spm,[status(thm)],[c_0_467, c_0_2387])).
cnf(c_0_2570,plain,(node15(s62)|~node12(s63)), inference(spm,[status(thm)],[c_0_2560, c_0_35])).
cnf(c_0_2571,negated_conjecture,(node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2561]), c_0_2562])).
cnf(c_0_2572,negated_conjecture,(~always25(s18)), inference(spm,[status(thm)],[c_0_2562, c_0_2563])).
cnf(c_0_2573,negated_conjecture,(~always25(s15)|~node11(s15)), inference(spm,[status(thm)],[c_0_2564, c_0_2565])).
cnf(c_0_2574,plain,(node11(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2566, c_0_827])])).
cnf(c_0_2575,negated_conjecture,(always25(s15)|~always25(s14)), inference(spm,[status(thm)],[c_0_134, c_0_720])).
cnf(c_0_2576,plain,(node15(s34)|~m_main_v_turn(s35,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1430])).
cnf(c_0_2577,plain,(node15(s38)|~m_main_v_turn(s39,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1540])).
cnf(c_0_2578,plain,(node15(s42)|~m_main_v_turn(s43,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1648])).
cnf(c_0_2579,plain,(node15(s46)|~m_main_v_turn(s47,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1758])).
cnf(c_0_2580,plain,(node15(s50)|~m_main_v_turn(s51,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_1891])).
cnf(c_0_2581,plain,(node15(s54)|~m_main_v_turn(s55,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_2063])).
cnf(c_0_2582,plain,(node15(s58)|~m_main_v_turn(s59,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_2243])).
cnf(c_0_2583,plain,(node15(s63)|m_main_v_turn(s64,c_1)), inference(sr,[status(thm)],[c_0_2567, c_0_2431])).
cnf(c_0_2584,plain,(m_main_v_turn(X1,c_2)|~node18(s63,X1)|~node5(s63)), inference(spm,[status(thm)],[c_0_2568, c_0_2387])).
cnf(c_0_2585,plain,(m_main_v_state2(s63,c_c2)|~node12(s63)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2569, c_0_2570]), c_0_2402]), c_0_2403]), c_0_2399])).
cnf(c_0_2586,negated_conjecture,(node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2571]), c_0_2572])).
cnf(c_0_2587,negated_conjecture,(~always25(s15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2573, c_0_2574])])).
cnf(c_0_2588,negated_conjecture,(~always25(s14)|~node11(s15)), inference(spm,[status(thm)],[c_0_2573, c_0_2575])).
cnf(c_0_2589,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node5(X2)|~trans(s34,X2)), inference(spm,[status(thm)],[c_0_2559, c_0_1456])).
cnf(c_0_2590,plain,(node10(s34)|node9(s34)|node8(s34)|m_main_v_state2(s35,c_c2)|~node15(s34)), inference(spm,[status(thm)],[c_0_467, c_0_1405])).
cnf(c_0_2591,plain,(node15(s34)|~node12(s35)), inference(spm,[status(thm)],[c_0_2576, c_0_35])).
cnf(c_0_2592,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node5(X2)|~trans(s38,X2)), inference(spm,[status(thm)],[c_0_2559, c_0_1566])).
cnf(c_0_2593,plain,(node10(s38)|node9(s38)|node8(s38)|m_main_v_state2(s39,c_c2)|~node15(s38)), inference(spm,[status(thm)],[c_0_467, c_0_1515])).
cnf(c_0_2594,plain,(node15(s38)|~node12(s39)), inference(spm,[status(thm)],[c_0_2577, c_0_35])).
cnf(c_0_2595,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node5(X2)|~trans(s42,X2)), inference(spm,[status(thm)],[c_0_2559, c_0_1674])).
cnf(c_0_2596,plain,(node10(s42)|node9(s42)|node8(s42)|m_main_v_state2(s43,c_c2)|~node15(s42)), inference(spm,[status(thm)],[c_0_467, c_0_1618])).
cnf(c_0_2597,plain,(node15(s42)|~node12(s43)), inference(spm,[status(thm)],[c_0_2578, c_0_35])).
cnf(c_0_2598,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node5(X2)|~trans(s46,X2)), inference(spm,[status(thm)],[c_0_2559, c_0_1784])).
cnf(c_0_2599,plain,(node10(s46)|node9(s46)|node8(s46)|m_main_v_state2(s47,c_c2)|~node15(s46)), inference(spm,[status(thm)],[c_0_467, c_0_1727])).
cnf(c_0_2600,plain,(node15(s46)|~node12(s47)), inference(spm,[status(thm)],[c_0_2579, c_0_35])).
cnf(c_0_2601,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node5(X2)|~trans(s50,X2)), inference(spm,[status(thm)],[c_0_2559, c_0_1928])).
cnf(c_0_2602,plain,(node10(s50)|node9(s50)|node8(s50)|m_main_v_state2(s51,c_c2)|~node15(s50)), inference(spm,[status(thm)],[c_0_467, c_0_1857])).
cnf(c_0_2603,plain,(node15(s50)|~node12(s51)), inference(spm,[status(thm)],[c_0_2580, c_0_35])).
cnf(c_0_2604,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node5(X2)|~trans(s54,X2)), inference(spm,[status(thm)],[c_0_2559, c_0_2104])).
cnf(c_0_2605,plain,(node10(s54)|node9(s54)|node8(s54)|m_main_v_state2(s55,c_c2)|~node15(s54)), inference(spm,[status(thm)],[c_0_467, c_0_2026])).
cnf(c_0_2606,plain,(node15(s54)|~node12(s55)), inference(spm,[status(thm)],[c_0_2581, c_0_35])).
cnf(c_0_2607,plain,(m_main_v_turn(X1,c_2)|~node18(X2,X1)|~node5(X2)|~trans(s58,X2)), inference(spm,[status(thm)],[c_0_2559, c_0_2285])).
cnf(c_0_2608,plain,(node10(s58)|node9(s58)|node8(s58)|m_main_v_state2(s59,c_c2)|~node15(s58)), inference(spm,[status(thm)],[c_0_467, c_0_2207])).
cnf(c_0_2609,plain,(node15(s58)|~node12(s59)), inference(spm,[status(thm)],[c_0_2582, c_0_35])).
cnf(c_0_2610,plain,(node15(s63)|~m_main_v_turn(s64,c_2)), inference(spm,[status(thm)],[c_0_30, c_0_2583])).
cnf(c_0_2611,plain,(m_main_v_turn(X1,c_2)|~node5(s63)|~trans(s63,X1)), inference(spm,[status(thm)],[c_0_2584, c_0_15])).
cnf(c_0_2612,plain,(node13(s63,s64)|node12(s63)|node11(s63)|node10(s63)|node9(s63)|node8(s63)|m_main_v_state2(s63,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_2420])).
cnf(c_0_2613,plain,(~node12(s63)), inference(spm,[status(thm)],[c_0_345, c_0_2585])).
cnf(c_0_2614,plain,(~m_main_v_state2(s63,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2421]), c_0_2422])])).
cnf(c_0_2615,negated_conjecture,(node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2586]), c_0_2587])).
cnf(c_0_2616,negated_conjecture,(~always25(s14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2588, c_0_2574])])).
cnf(c_0_2617,negated_conjecture,(~always25(s12)), inference(spm,[status(thm)],[c_0_133, c_0_676])).
cnf(c_0_2618,negated_conjecture,(always25(s12)|~always25(s11)), inference(spm,[status(thm)],[c_0_134, c_0_625])).
cnf(c_0_2619,plain,(m_main_v_turn(X1,c_2)|~node18(s35,X1)|~node5(s35)), inference(spm,[status(thm)],[c_0_2589, c_0_1405])).
cnf(c_0_2620,plain,(m_main_v_state2(s35,c_c2)|~node12(s35)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2590, c_0_2591]), c_0_1419]), c_0_1420]), c_0_1416])).
cnf(c_0_2621,plain,(m_main_v_turn(X1,c_2)|~node18(s39,X1)|~node5(s39)), inference(spm,[status(thm)],[c_0_2592, c_0_1515])).
cnf(c_0_2622,plain,(m_main_v_state2(s39,c_c2)|~node12(s39)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2593, c_0_2594]), c_0_1529]), c_0_1530]), c_0_1526])).
cnf(c_0_2623,plain,(m_main_v_turn(X1,c_2)|~node18(s43,X1)|~node5(s43)), inference(spm,[status(thm)],[c_0_2595, c_0_1618])).
cnf(c_0_2624,plain,(m_main_v_state2(s43,c_c2)|~node12(s43)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2596, c_0_2597]), c_0_1638]), c_0_1639]), c_0_1635])).
cnf(c_0_2625,plain,(m_main_v_turn(X1,c_2)|~node18(s47,X1)|~node5(s47)), inference(spm,[status(thm)],[c_0_2598, c_0_1727])).
cnf(c_0_2626,plain,(m_main_v_state2(s47,c_c2)|~node12(s47)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2599, c_0_2600]), c_0_1748]), c_0_1749]), c_0_1745])).
cnf(c_0_2627,plain,(m_main_v_turn(X1,c_2)|~node18(s51,X1)|~node5(s51)), inference(spm,[status(thm)],[c_0_2601, c_0_1857])).
cnf(c_0_2628,plain,(m_main_v_state2(s51,c_c2)|~node12(s51)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2602, c_0_2603]), c_0_1874]), c_0_1875]), c_0_1871])).
cnf(c_0_2629,plain,(m_main_v_turn(X1,c_2)|~node18(s55,X1)|~node5(s55)), inference(spm,[status(thm)],[c_0_2604, c_0_2026])).
cnf(c_0_2630,plain,(m_main_v_state2(s55,c_c2)|~node12(s55)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2605, c_0_2606]), c_0_2043]), c_0_2044]), c_0_2040])).
cnf(c_0_2631,plain,(m_main_v_turn(X1,c_2)|~node18(s59,X1)|~node5(s59)), inference(spm,[status(thm)],[c_0_2607, c_0_2207])).
cnf(c_0_2632,plain,(m_main_v_state2(s59,c_c2)|~node12(s59)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2608, c_0_2609]), c_0_2225]), c_0_2226]), c_0_2222])).
cnf(c_0_2633,plain,(node15(s63)|~node5(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2610, c_0_2611]), c_0_2420])])).
cnf(c_0_2634,plain,(m_main_v_state2(s62,X1)|~m_main_v_state2(s63,X1)), inference(spm,[status(thm)],[c_0_47, c_0_2411])).
cnf(c_0_2635,plain,(node13(s63,s64)|node11(s63)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2612, c_0_2613]), c_0_2469]), c_0_2470]), c_0_2471]), c_0_2614])).
cnf(c_0_2636,negated_conjecture,(node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2615]), c_0_2616])).
cnf(c_0_2637,negated_conjecture,(~always25(s11)), inference(spm,[status(thm)],[c_0_2617, c_0_2618])).
cnf(c_0_2638,negated_conjecture,(always25(s11)|~always25(s10)), inference(spm,[status(thm)],[c_0_134, c_0_585])).
cnf(c_0_2639,plain,(m_main_v_turn(X1,c_2)|~node5(s35)|~trans(s35,X1)), inference(spm,[status(thm)],[c_0_2619, c_0_15])).
cnf(c_0_2640,plain,(node13(s35,s36)|node12(s35)|node11(s35)|node10(s35)|node9(s35)|node8(s35)|m_main_v_state2(s35,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1423])).
cnf(c_0_2641,plain,(~node12(s35)), inference(spm,[status(thm)],[c_0_345, c_0_2620])).
cnf(c_0_2642,plain,(~m_main_v_state2(s35,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_1432]), c_0_1433])])).
cnf(c_0_2643,plain,(m_main_v_turn(X1,c_2)|~node5(s39)|~trans(s39,X1)), inference(spm,[status(thm)],[c_0_2621, c_0_15])).
cnf(c_0_2644,plain,(node13(s39,s40)|node12(s39)|node11(s39)|node10(s39)|node9(s39)|node8(s39)|m_main_v_state2(s39,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1533])).
cnf(c_0_2645,plain,(~node12(s39)), inference(spm,[status(thm)],[c_0_345, c_0_2622])).
cnf(c_0_2646,plain,(~m_main_v_state2(s39,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_1542]), c_0_1543])])).
cnf(c_0_2647,plain,(m_main_v_turn(X1,c_2)|~node5(s43)|~trans(s43,X1)), inference(spm,[status(thm)],[c_0_2623, c_0_15])).
cnf(c_0_2648,plain,(node13(s43,s44)|node12(s43)|node11(s43)|node10(s43)|node9(s43)|node8(s43)|m_main_v_state2(s43,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1642])).
cnf(c_0_2649,plain,(~node12(s43)), inference(spm,[status(thm)],[c_0_345, c_0_2624])).
cnf(c_0_2650,plain,(~m_main_v_state2(s43,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_1650]), c_0_1651])])).
cnf(c_0_2651,plain,(m_main_v_turn(X1,c_2)|~node5(s47)|~trans(s47,X1)), inference(spm,[status(thm)],[c_0_2625, c_0_15])).
cnf(c_0_2652,plain,(node13(s47,s48)|node12(s47)|node11(s47)|node10(s47)|node9(s47)|node8(s47)|m_main_v_state2(s47,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1752])).
cnf(c_0_2653,plain,(~node12(s47)), inference(spm,[status(thm)],[c_0_345, c_0_2626])).
cnf(c_0_2654,plain,(~m_main_v_state2(s47,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_1760]), c_0_1761])])).
cnf(c_0_2655,plain,(m_main_v_turn(X1,c_2)|~node5(s51)|~trans(s51,X1)), inference(spm,[status(thm)],[c_0_2627, c_0_15])).
cnf(c_0_2656,plain,(node13(s51,s52)|node12(s51)|node11(s51)|node10(s51)|node9(s51)|node8(s51)|m_main_v_state2(s51,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_1881])).
cnf(c_0_2657,plain,(~node12(s51)), inference(spm,[status(thm)],[c_0_345, c_0_2628])).
cnf(c_0_2658,plain,(~m_main_v_state2(s51,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_1893]), c_0_1894])])).
cnf(c_0_2659,plain,(m_main_v_turn(X1,c_2)|~node5(s55)|~trans(s55,X1)), inference(spm,[status(thm)],[c_0_2629, c_0_15])).
cnf(c_0_2660,plain,(node13(s55,s56)|node12(s55)|node11(s55)|node10(s55)|node9(s55)|node8(s55)|m_main_v_state2(s55,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_2052])).
cnf(c_0_2661,plain,(~node12(s55)), inference(spm,[status(thm)],[c_0_345, c_0_2630])).
cnf(c_0_2662,plain,(~m_main_v_state2(s55,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2065]), c_0_2066])])).
cnf(c_0_2663,plain,(m_main_v_turn(X1,c_2)|~node5(s59)|~trans(s59,X1)), inference(spm,[status(thm)],[c_0_2631, c_0_15])).
cnf(c_0_2664,plain,(node13(s59,s60)|node12(s59)|node11(s59)|node10(s59)|node9(s59)|node8(s59)|m_main_v_state2(s59,c_c2)), inference(spm,[status(thm)],[c_0_179, c_0_2232])).
cnf(c_0_2665,plain,(~node12(s59)), inference(spm,[status(thm)],[c_0_345, c_0_2632])).
cnf(c_0_2666,plain,(~m_main_v_state2(s59,c_c2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2245]), c_0_2246])])).
cnf(c_0_2667,plain,(node6(s63,s64)|node5(s63)|node4(s63)|node3(s63)|node2(s63)|node1(s63)|m_main_v_state1(s63,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_2420])).
cnf(c_0_2668,plain,(~node4(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_2421]), c_0_2422])])).
cnf(c_0_2669,plain,(~node5(s63)), inference(spm,[status(thm)],[c_0_805, c_0_2633])).
cnf(c_0_2670,plain,(~node3(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2634, c_0_12]), c_0_2404])).
cnf(c_0_2671,plain,(~node2(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_2421]), c_0_2422])])).
cnf(c_0_2672,plain,(node11(s63)|m_main_v_state2(s64,X1)|~m_main_v_state2(s63,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2635])).
cnf(c_0_2673,negated_conjecture,(node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2636]), c_0_2637])).
cnf(c_0_2674,negated_conjecture,(~always25(s10)), inference(spm,[status(thm)],[c_0_2637, c_0_2638])).
cnf(c_0_2675,negated_conjecture,(~always25(s8)), inference(spm,[status(thm)],[c_0_133, c_0_499])).
cnf(c_0_2676,negated_conjecture,(always25(s8)|~always25(s7)), inference(spm,[status(thm)],[c_0_134, c_0_445])).
cnf(c_0_2677,plain,(node15(s63)|~node12(s64)), inference(spm,[status(thm)],[c_0_2610, c_0_35])).
cnf(c_0_2678,plain,(m_main_v_state1(s64,c_t1)|~node15(s63)), inference(spm,[status(thm)],[c_0_474, c_0_2420])).
cnf(c_0_2679,plain,(node15(s35)|~node5(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1454, c_0_2639]), c_0_1423])])).
cnf(c_0_2680,plain,(m_main_v_state2(s34,X1)|~m_main_v_state2(s35,X1)), inference(spm,[status(thm)],[c_0_47, c_0_1427])).
cnf(c_0_2681,plain,(node13(s35,s36)|node11(s35)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2640, c_0_2641]), c_0_1459]), c_0_1460]), c_0_1461]), c_0_2642])).
cnf(c_0_2682,plain,(node15(s39)|~node5(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1564, c_0_2643]), c_0_1533])])).
cnf(c_0_2683,plain,(m_main_v_state2(s38,X1)|~m_main_v_state2(s39,X1)), inference(spm,[status(thm)],[c_0_47, c_0_1537])).
cnf(c_0_2684,plain,(node13(s39,s40)|node11(s39)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2644, c_0_2645]), c_0_1569]), c_0_1570]), c_0_1571]), c_0_2646])).
cnf(c_0_2685,plain,(node15(s43)|~node5(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1672, c_0_2647]), c_0_1642])])).
cnf(c_0_2686,plain,(m_main_v_state2(s42,X1)|~m_main_v_state2(s43,X1)), inference(spm,[status(thm)],[c_0_47, c_0_1645])).
cnf(c_0_2687,plain,(node13(s43,s44)|node11(s43)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2648, c_0_2649]), c_0_1677]), c_0_1678]), c_0_1679]), c_0_2650])).
cnf(c_0_2688,plain,(node15(s47)|~node5(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1782, c_0_2651]), c_0_1752])])).
cnf(c_0_2689,plain,(m_main_v_state2(s46,X1)|~m_main_v_state2(s47,X1)), inference(spm,[status(thm)],[c_0_47, c_0_1755])).
cnf(c_0_2690,plain,(node13(s47,s48)|node11(s47)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2652, c_0_2653]), c_0_1787]), c_0_1788]), c_0_1789]), c_0_2654])).
cnf(c_0_2691,plain,(node15(s51)|~node5(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1926, c_0_2655]), c_0_1881])])).
cnf(c_0_2692,plain,(m_main_v_state2(s50,X1)|~m_main_v_state2(s51,X1)), inference(spm,[status(thm)],[c_0_47, c_0_1885])).
cnf(c_0_2693,plain,(node13(s51,s52)|node11(s51)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2656, c_0_2657]), c_0_1934]), c_0_1935]), c_0_1936]), c_0_2658])).
cnf(c_0_2694,plain,(node15(s55)|~node5(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2102, c_0_2659]), c_0_2052])])).
cnf(c_0_2695,plain,(m_main_v_state2(s54,X1)|~m_main_v_state2(s55,X1)), inference(spm,[status(thm)],[c_0_47, c_0_2056])).
cnf(c_0_2696,plain,(node13(s55,s56)|node11(s55)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2660, c_0_2661]), c_0_2110]), c_0_2111]), c_0_2112]), c_0_2662])).
cnf(c_0_2697,plain,(node15(s59)|~node5(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2283, c_0_2663]), c_0_2232])])).
cnf(c_0_2698,plain,(m_main_v_state2(s58,X1)|~m_main_v_state2(s59,X1)), inference(spm,[status(thm)],[c_0_47, c_0_2236])).
cnf(c_0_2699,plain,(node13(s59,s60)|node11(s59)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2664, c_0_2665]), c_0_2293]), c_0_2294]), c_0_2295]), c_0_2666])).
cnf(c_0_2700,plain,(node6(s63,s64)|node1(s63)|m_main_v_state1(s63,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2667, c_0_2668]), c_0_2669]), c_0_2670]), c_0_2671])).
cnf(c_0_2701,plain,(node11(s63)|m_main_v_state2(s64,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2672, c_0_17]), c_0_2468])])).
cnf(c_0_2702,negated_conjecture,(node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2673]), c_0_2674])).
cnf(c_0_2703,negated_conjecture,(~always25(s7)), inference(spm,[status(thm)],[c_0_2675, c_0_2676])).
cnf(c_0_2704,negated_conjecture,(always25(s7)|~always25(s6)), inference(spm,[status(thm)],[c_0_134, c_0_426])).
cnf(c_0_2705,plain,(m_main_v_state2(s64,c_c2)|~node12(s64)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2467, c_0_2677]), c_0_2469]), c_0_2470]), c_0_2471])).
cnf(c_0_2706,plain,(~node15(s63)|~node11(s64)), inference(spm,[status(thm)],[c_0_165, c_0_2678])).
cnf(c_0_2707,plain,(~node15(s63)|~node9(s64)), inference(spm,[status(thm)],[c_0_167, c_0_2678])).
cnf(c_0_2708,plain,(node6(s35,s36)|node5(s35)|node4(s35)|node3(s35)|node2(s35)|node1(s35)|m_main_v_state1(s35,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1423])).
cnf(c_0_2709,plain,(~node4(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_1432]), c_0_1433])])).
cnf(c_0_2710,plain,(~node5(s35)), inference(spm,[status(thm)],[c_0_805, c_0_2679])).
cnf(c_0_2711,plain,(~node3(s35)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2680, c_0_12]), c_0_1421])).
cnf(c_0_2712,plain,(~node2(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_1432]), c_0_1433])])).
cnf(c_0_2713,plain,(node11(s35)|m_main_v_state2(s36,X1)|~m_main_v_state2(s35,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2681])).
cnf(c_0_2714,plain,(node6(s39,s40)|node5(s39)|node4(s39)|node3(s39)|node2(s39)|node1(s39)|m_main_v_state1(s39,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1533])).
cnf(c_0_2715,plain,(~node4(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_1542]), c_0_1543])])).
cnf(c_0_2716,plain,(~node5(s39)), inference(spm,[status(thm)],[c_0_805, c_0_2682])).
cnf(c_0_2717,plain,(~node3(s39)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2683, c_0_12]), c_0_1531])).
cnf(c_0_2718,plain,(~node2(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_1542]), c_0_1543])])).
cnf(c_0_2719,plain,(node11(s39)|m_main_v_state2(s40,X1)|~m_main_v_state2(s39,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2684])).
cnf(c_0_2720,plain,(node6(s43,s44)|node5(s43)|node4(s43)|node3(s43)|node2(s43)|node1(s43)|m_main_v_state1(s43,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1642])).
cnf(c_0_2721,plain,(~node4(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_1650]), c_0_1651])])).
cnf(c_0_2722,plain,(~node5(s43)), inference(spm,[status(thm)],[c_0_805, c_0_2685])).
cnf(c_0_2723,plain,(~node3(s43)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2686, c_0_12]), c_0_1640])).
cnf(c_0_2724,plain,(~node2(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_1650]), c_0_1651])])).
cnf(c_0_2725,plain,(node11(s43)|m_main_v_state2(s44,X1)|~m_main_v_state2(s43,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2687])).
cnf(c_0_2726,plain,(node6(s47,s48)|node5(s47)|node4(s47)|node3(s47)|node2(s47)|node1(s47)|m_main_v_state1(s47,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1752])).
cnf(c_0_2727,plain,(~node4(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_1760]), c_0_1761])])).
cnf(c_0_2728,plain,(~node5(s47)), inference(spm,[status(thm)],[c_0_805, c_0_2688])).
cnf(c_0_2729,plain,(~node3(s47)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2689, c_0_12]), c_0_1750])).
cnf(c_0_2730,plain,(~node2(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_1760]), c_0_1761])])).
cnf(c_0_2731,plain,(node11(s47)|m_main_v_state2(s48,X1)|~m_main_v_state2(s47,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2690])).
cnf(c_0_2732,plain,(node6(s51,s52)|node5(s51)|node4(s51)|node3(s51)|node2(s51)|node1(s51)|m_main_v_state1(s51,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_1881])).
cnf(c_0_2733,plain,(~node4(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_1893]), c_0_1894])])).
cnf(c_0_2734,plain,(~node5(s51)), inference(spm,[status(thm)],[c_0_805, c_0_2691])).
cnf(c_0_2735,plain,(~node3(s51)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2692, c_0_12]), c_0_1876])).
cnf(c_0_2736,plain,(~node2(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_1893]), c_0_1894])])).
cnf(c_0_2737,plain,(node11(s51)|m_main_v_state2(s52,X1)|~m_main_v_state2(s51,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2693])).
cnf(c_0_2738,plain,(node6(s55,s56)|node5(s55)|node4(s55)|node3(s55)|node2(s55)|node1(s55)|m_main_v_state1(s55,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_2052])).
cnf(c_0_2739,plain,(~node4(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_2065]), c_0_2066])])).
cnf(c_0_2740,plain,(~node5(s55)), inference(spm,[status(thm)],[c_0_805, c_0_2694])).
cnf(c_0_2741,plain,(~node3(s55)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2695, c_0_12]), c_0_2045])).
cnf(c_0_2742,plain,(~node2(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_2065]), c_0_2066])])).
cnf(c_0_2743,plain,(node11(s55)|m_main_v_state2(s56,X1)|~m_main_v_state2(s55,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2696])).
cnf(c_0_2744,plain,(node6(s59,s60)|node5(s59)|node4(s59)|node3(s59)|node2(s59)|node1(s59)|m_main_v_state1(s59,c_c1)), inference(spm,[status(thm)],[c_0_146, c_0_2232])).
cnf(c_0_2745,plain,(~node4(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_2245]), c_0_2246])])).
cnf(c_0_2746,plain,(~node5(s59)), inference(spm,[status(thm)],[c_0_805, c_0_2697])).
cnf(c_0_2747,plain,(~node3(s59)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2698, c_0_12]), c_0_2227])).
cnf(c_0_2748,plain,(~node2(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_2245]), c_0_2246])])).
cnf(c_0_2749,plain,(node11(s59)|m_main_v_state2(s60,X1)|~m_main_v_state2(s59,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2699])).
cnf(c_0_2750,plain,(node1(s63)|m_main_v_state1(s63,c_c1)|m_main_v_state1(s64,X1)|~m_main_v_state1(s63,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2700])).
cnf(c_0_2751,plain,(node11(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2701]), c_0_2475])])).
cnf(c_0_2752,negated_conjecture,(node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2702]), c_0_2703])).
cnf(c_0_2753,negated_conjecture,(~always25(s6)), inference(spm,[status(thm)],[c_0_2703, c_0_2704])).
cnf(c_0_2754,negated_conjecture,(~always25(s4)), inference(spm,[status(thm)],[c_0_133, c_0_346])).
cnf(c_0_2755,negated_conjecture,(always25(s4)|~always25(s3)), inference(spm,[status(thm)],[c_0_134, c_0_279])).
cnf(c_0_2756,plain,(~node12(s64)), inference(spm,[status(thm)],[c_0_345, c_0_2705])).
cnf(c_0_2757,plain,(~node11(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2706, c_0_2468])])).
cnf(c_0_2758,plain,(~node9(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2707, c_0_2468])])).
cnf(c_0_2759,plain,(~node10(s64)), inference(spm,[status(thm)],[c_0_339, c_0_2475])).
cnf(c_0_2760,plain,(~node8(s64)), inference(spm,[status(thm)],[c_0_343, c_0_2475])).
cnf(c_0_2761,plain,(node6(s35,s36)|node1(s35)|m_main_v_state1(s35,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2708, c_0_2709]), c_0_2710]), c_0_2711]), c_0_2712])).
cnf(c_0_2762,plain,(node11(s35)|m_main_v_state2(s36,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2713, c_0_17]), c_0_1463])])).
cnf(c_0_2763,plain,(node6(s39,s40)|node1(s39)|m_main_v_state1(s39,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2714, c_0_2715]), c_0_2716]), c_0_2717]), c_0_2718])).
cnf(c_0_2764,plain,(node11(s39)|m_main_v_state2(s40,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2719, c_0_17]), c_0_1573])])).
cnf(c_0_2765,plain,(node6(s43,s44)|node1(s43)|m_main_v_state1(s43,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2720, c_0_2721]), c_0_2722]), c_0_2723]), c_0_2724])).
cnf(c_0_2766,plain,(node11(s43)|m_main_v_state2(s44,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2725, c_0_17]), c_0_1681])])).
cnf(c_0_2767,plain,(node6(s47,s48)|node1(s47)|m_main_v_state1(s47,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2726, c_0_2727]), c_0_2728]), c_0_2729]), c_0_2730])).
cnf(c_0_2768,plain,(node11(s47)|m_main_v_state2(s48,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2731, c_0_17]), c_0_1791])])).
cnf(c_0_2769,plain,(node6(s51,s52)|node1(s51)|m_main_v_state1(s51,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2732, c_0_2733]), c_0_2734]), c_0_2735]), c_0_2736])).
cnf(c_0_2770,plain,(node11(s51)|m_main_v_state2(s52,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2737, c_0_17]), c_0_1938])])).
cnf(c_0_2771,plain,(node6(s55,s56)|node1(s55)|m_main_v_state1(s55,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2738, c_0_2739]), c_0_2740]), c_0_2741]), c_0_2742])).
cnf(c_0_2772,plain,(node11(s55)|m_main_v_state2(s56,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2743, c_0_17]), c_0_2114])])).
cnf(c_0_2773,plain,(node6(s59,s60)|node1(s59)|m_main_v_state1(s59,c_c1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2744, c_0_2745]), c_0_2746]), c_0_2747]), c_0_2748])).
cnf(c_0_2774,plain,(node11(s59)|m_main_v_state2(s60,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2749, c_0_17]), c_0_2297])])).
cnf(c_0_2775,plain,(node1(s63)|m_main_v_state1(s63,c_c1)|m_main_v_state1(s64,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2750, c_0_25]), c_0_2751])])).
cnf(c_0_2776,negated_conjecture,(node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2752]), c_0_2753])).
cnf(c_0_2777,negated_conjecture,(~always25(s3)), inference(spm,[status(thm)],[c_0_2754, c_0_2755])).
cnf(c_0_2778,negated_conjecture,(always25(s3)|~always25(s2)), inference(spm,[status(thm)],[c_0_134, c_0_204])).
cnf(c_0_2779,plain,(m_main_v_state2(X1,c_n2)|~node14(s64,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_2475]), c_0_2756]), c_0_2757]), c_0_2758]), c_0_2759]), c_0_2760])).
cnf(c_0_2780,plain,(node13(s22,s23)|node12(s22)|m_main_v_state2(s22,c_c2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1017, c_0_1036]), c_0_1037])).
cnf(c_0_2781,plain,(node1(s35)|m_main_v_state1(s35,c_c1)|m_main_v_state1(s36,X1)|~m_main_v_state1(s35,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2761])).
cnf(c_0_2782,plain,(node11(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2762]), c_0_1468])])).
cnf(c_0_2783,plain,(node1(s39)|m_main_v_state1(s39,c_c1)|m_main_v_state1(s40,X1)|~m_main_v_state1(s39,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2763])).
cnf(c_0_2784,plain,(node11(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2764]), c_0_1578])])).
cnf(c_0_2785,plain,(node1(s43)|m_main_v_state1(s43,c_c1)|m_main_v_state1(s44,X1)|~m_main_v_state1(s43,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2765])).
cnf(c_0_2786,plain,(node11(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2766]), c_0_1685])])).
cnf(c_0_2787,plain,(node1(s47)|m_main_v_state1(s47,c_c1)|m_main_v_state1(s48,X1)|~m_main_v_state1(s47,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2767])).
cnf(c_0_2788,plain,(node11(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2768]), c_0_1795])])).
cnf(c_0_2789,plain,(node1(s51)|m_main_v_state1(s51,c_c1)|m_main_v_state1(s52,X1)|~m_main_v_state1(s51,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2769])).
cnf(c_0_2790,plain,(node11(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2770]), c_0_1946])])).
cnf(c_0_2791,plain,(node1(s55)|m_main_v_state1(s55,c_c1)|m_main_v_state1(s56,X1)|~m_main_v_state1(s55,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2771])).
cnf(c_0_2792,plain,(node11(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2772]), c_0_2122])])).
cnf(c_0_2793,plain,(node1(s59)|m_main_v_state1(s59,c_c1)|m_main_v_state1(s60,X1)|~m_main_v_state1(s59,X1)), inference(spm,[status(thm)],[c_0_23, c_0_2773])).
cnf(c_0_2794,plain,(node11(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2774]), c_0_2306])])).
cnf(c_0_2795,plain,(m_main_v_turn(s64,c_2)|~node1(s63)), inference(spm,[status(thm)],[c_0_232, c_0_2420])).
cnf(c_0_2796,plain,(node1(s63)|m_main_v_state1(s64,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2775]), c_0_2468])])).
cnf(c_0_2797,plain,(m_main_v_state1(s64,c_t1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2678, c_0_2468])])).
cnf(c_0_2798,plain,(node8(X1)|m_main_v_state2(X2,c_t2)|~node15(X1)|~node14(X1,X2)|~m_main_v_state2(X1,c_n2)), inference(spm,[status(thm)],[c_0_20, c_0_38])).
cnf(c_0_2799,negated_conjecture,(node26(s2)|node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2776]), c_0_2777])).
cnf(c_0_2800,negated_conjecture,(~always25(s2)), inference(spm,[status(thm)],[c_0_2777, c_0_2778])).
cnf(c_0_2801,negated_conjecture,(always25(s2)|~always25(s1)), inference(spm,[status(thm)],[c_0_134, c_0_180])).
cnf(c_0_2802,plain,(m_main_v_state2(X1,c_n2)|~trans(s64,X1)), inference(spm,[status(thm)],[c_0_2779, c_0_29])).
cnf(c_0_2803,plain,(node17(s4,s5)|node16(s4)), inference(sr,[status(thm)],[c_0_378, c_0_391])).
cnf(c_0_2804,plain,(node17(s8,s9)|node16(s8)), inference(sr,[status(thm)],[c_0_469, c_0_496])).
cnf(c_0_2805,plain,(node12(s10)|m_main_v_state2(s10,c_c2)|m_main_v_state2(s11,X1)|~m_main_v_state2(s10,X1)), inference(spm,[status(thm)],[c_0_43, c_0_620])).
cnf(c_0_2806,plain,(m_main_v_state2(X1,c_t2)|~trans(s9,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_563, c_0_557]), c_0_461])])).
cnf(c_0_2807,plain,(node17(s12,s13)|node16(s12)), inference(sr,[status(thm)],[c_0_700, c_0_715])).
cnf(c_0_2808,plain,(node17(s16,s17)|node16(s16)), inference(sr,[status(thm)],[c_0_837, c_0_856])).
cnf(c_0_2809,plain,(node17(s20,s21)|node16(s20)), inference(sr,[status(thm)],[c_0_927, c_0_946])).
cnf(c_0_2810,plain,(node13(s22,s23)|m_main_v_state2(s22,c_c2)), inference(sr,[status(thm)],[c_0_2780, c_0_1038])).
cnf(c_0_2811,plain,(node17(s24,s25)|node16(s24)), inference(sr,[status(thm)],[c_0_1124, c_0_1135])).
cnf(c_0_2812,plain,(node17(s28,s29)|node16(s28)), inference(sr,[status(thm)],[c_0_1273, c_0_1279])).
cnf(c_0_2813,plain,(node17(s32,s33)|node16(s32)), inference(sr,[status(thm)],[c_0_1378, c_0_1391])).
cnf(c_0_2814,plain,(node1(s35)|m_main_v_state1(s35,c_c1)|m_main_v_state1(s36,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2781, c_0_25]), c_0_2782])])).
cnf(c_0_2815,plain,(node17(s36,s37)|node16(s36)), inference(sr,[status(thm)],[c_0_1489, c_0_1502])).
cnf(c_0_2816,plain,(node1(s39)|m_main_v_state1(s39,c_c1)|m_main_v_state1(s40,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2783, c_0_25]), c_0_2784])])).
cnf(c_0_2817,plain,(node17(s40,s41)|node16(s40)), inference(sr,[status(thm)],[c_0_1602, c_0_1611])).
cnf(c_0_2818,plain,(node1(s43)|m_main_v_state1(s43,c_c1)|m_main_v_state1(s44,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2785, c_0_25]), c_0_2786])])).
cnf(c_0_2819,plain,(node17(s44,s45)|node16(s44)), inference(sr,[status(thm)],[c_0_1707, c_0_1720])).
cnf(c_0_2820,plain,(node1(s47)|m_main_v_state1(s47,c_c1)|m_main_v_state1(s48,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2787, c_0_25]), c_0_2788])])).
cnf(c_0_2821,plain,(node17(s48,s49)|node16(s48)), inference(sr,[status(thm)],[c_0_1817, c_0_1833])).
cnf(c_0_2822,plain,(node1(s51)|m_main_v_state1(s51,c_c1)|m_main_v_state1(s52,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2789, c_0_25]), c_0_2790])])).
cnf(c_0_2823,plain,(node17(s52,s53)|node16(s52)), inference(sr,[status(thm)],[c_0_1981, c_0_1997])).
cnf(c_0_2824,plain,(node1(s55)|m_main_v_state1(s55,c_c1)|m_main_v_state1(s56,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2791, c_0_25]), c_0_2792])])).
cnf(c_0_2825,plain,(node17(s56,s57)|node16(s56)), inference(sr,[status(thm)],[c_0_2153, c_0_2178])).
cnf(c_0_2826,plain,(node1(s59)|m_main_v_state1(s59,c_c1)|m_main_v_state1(s60,c_n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2793, c_0_25]), c_0_2794])])).
cnf(c_0_2827,plain,(node17(s60,s61)|node16(s60)), inference(sr,[status(thm)],[c_0_2349, c_0_2358])).
cnf(c_0_2828,plain,(node11(s64)|node10(s64)|node9(s64)|node8(s64)|m_main_v_state2(X1,c_c2)|~node14(s64,X1)|~node1(s63)|~m_main_v_state2(s64,c_t2)|~m_main_v_state1(s64,c_t1)), inference(spm,[status(thm)],[c_0_136, c_0_2795])).
cnf(c_0_2829,plain,(node1(s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_2796]), c_0_2797])])).
cnf(c_0_2830,plain,(m_main_v_state2(X1,c_t2)|~node14(s63,X1)|~m_main_v_state2(s62,c_n2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2798, c_0_2421]), c_0_2468])]), c_0_2471])).
cnf(c_0_2831,negated_conjecture,(node26(s1)|always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2799]), c_0_2800])).
cnf(c_0_2832,negated_conjecture,(~always25(s1)), inference(spm,[status(thm)],[c_0_2800, c_0_2801])).
cnf(c_0_2833,plain,(node15(X1)|m_main_v_turn(X2,c_1)|~node18(X1,X2)|~m_main_v_state1(X1,c_t1)|~trans(s64,X1)), inference(spm,[status(thm)],[c_0_63, c_0_2802])).
cnf(c_0_2834,plain,(node17(s4,s5)), inference(sr,[status(thm)],[c_0_2803, c_0_398])).
cnf(c_0_2835,plain,(node17(s8,s9)), inference(sr,[status(thm)],[c_0_2804, c_0_615])).
cnf(c_0_2836,plain,(node12(s10)|m_main_v_state2(s10,c_c2)|m_main_v_state2(s11,c_t2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2805, c_0_2806]), c_0_528])])).
cnf(c_0_2837,plain,(node17(s12,s13)), inference(sr,[status(thm)],[c_0_2807, c_0_719])).
cnf(c_0_2838,plain,(node17(s16,s17)), inference(sr,[status(thm)],[c_0_2808, c_0_862])).
cnf(c_0_2839,plain,(node17(s20,s21)), inference(sr,[status(thm)],[c_0_2809, c_0_1031])).
cnf(c_0_2840,plain,(node13(s22,s23)), inference(sr,[status(thm)],[c_0_2810, c_0_1052])).
cnf(c_0_2841,plain,(node17(s24,s25)), inference(sr,[status(thm)],[c_0_2811, c_0_1139])).
cnf(c_0_2842,plain,(node17(s28,s29)), inference(sr,[status(thm)],[c_0_2812, c_0_1283])).
cnf(c_0_2843,plain,(node17(s32,s33)), inference(sr,[status(thm)],[c_0_2813, c_0_1395])).
cnf(c_0_2844,plain,(node1(s35)|m_main_v_state1(s35,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_2814]), c_0_1492])])).
cnf(c_0_2845,plain,(node17(s36,s37)), inference(sr,[status(thm)],[c_0_2815, c_0_1506])).
cnf(c_0_2846,plain,(node1(s39)|m_main_v_state1(s39,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_2816]), c_0_1601])])).
cnf(c_0_2847,plain,(node17(s40,s41)), inference(sr,[status(thm)],[c_0_2817, c_0_1615])).
cnf(c_0_2848,plain,(node1(s43)|m_main_v_state1(s43,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_2818]), c_0_1710])])).
cnf(c_0_2849,plain,(node17(s44,s45)), inference(sr,[status(thm)],[c_0_2819, c_0_1724])).
cnf(c_0_2850,plain,(node1(s47)|m_main_v_state1(s47,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_2820]), c_0_1820])])).
cnf(c_0_2851,plain,(node17(s48,s49)), inference(sr,[status(thm)],[c_0_2821, c_0_1839])).
cnf(c_0_2852,plain,(node1(s51)|m_main_v_state1(s51,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_2822]), c_0_1980])])).
cnf(c_0_2853,plain,(node17(s52,s53)), inference(sr,[status(thm)],[c_0_2823, c_0_2005])).
cnf(c_0_2854,plain,(node1(s55)|m_main_v_state1(s55,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_2824]), c_0_2159])])).
cnf(c_0_2855,plain,(node17(s56,s57)), inference(sr,[status(thm)],[c_0_2825, c_0_2186])).
cnf(c_0_2856,plain,(node1(s59)|m_main_v_state1(s59,c_c1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_2826]), c_0_2341])])).
cnf(c_0_2857,plain,(node17(s60,s61)), inference(sr,[status(thm)],[c_0_2827, c_0_2366])).
cnf(c_0_2858,plain,(m_main_v_state2(X1,c_c2)|~node14(s64,X1)|~m_main_v_state2(s64,c_t2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2828, c_0_2797])]), c_0_2757]), c_0_2758]), c_0_2829])]), c_0_2759]), c_0_2760])).
cnf(c_0_2859,plain,(m_main_v_state2(X1,c_t2)|~m_main_v_state2(s62,c_n2)|~trans(s63,X1)), inference(spm,[status(thm)],[c_0_2830, c_0_29])).
cnf(c_0_2860,negated_conjecture,(always25(s64)|~always25(s63)), inference(spm,[status(thm)],[c_0_134, c_0_2420])).
cnf(c_0_2861,negated_conjecture,(always25(s63)|~always25(s62)), inference(spm,[status(thm)],[c_0_134, c_0_2387])).
cnf(c_0_2862,negated_conjecture,(always25(s62)|always25(s63)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2831]), c_0_2832])).
cnf(c_0_2863,negated_conjecture,(~always25(s64)), inference(spm,[status(thm)],[c_0_133, c_0_2475])).
cnf(c_0_2864,plain,(m_main_v_turn(X1,c_1)|~node18(s4,X1)|~trans(s64,s4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_383]), c_0_391])).
cnf(c_0_2865,plain,(m_main_v_turn(s5,X1)|~m_main_v_turn(s4,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2834])).
cnf(c_0_2866,plain,(~m_main_v_state2(s7,c_t2)|~m_main_v_state2(s6,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_455])).
cnf(c_0_2867,plain,(m_main_v_turn(X1,c_1)|~node18(s8,X1)|~trans(s64,s8)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_531]), c_0_496])).
cnf(c_0_2868,plain,(m_main_v_turn(s9,X1)|~m_main_v_turn(s8,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2835])).
cnf(c_0_2869,plain,(m_main_v_turn(s9,c_2)), inference(sr,[status(thm)],[c_0_507, c_0_615])).
cnf(c_0_2870,plain,(m_main_v_state2(s11,c_t2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_2836, c_0_621]), c_0_622])).
cnf(c_0_2871,plain,(m_main_v_turn(X1,c_1)|~node18(s12,X1)|~trans(s64,s12)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_709]), c_0_715])).
cnf(c_0_2872,plain,(m_main_v_turn(s13,X1)|~m_main_v_turn(s12,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2837])).
cnf(c_0_2873,plain,(~m_main_v_state2(s15,c_t2)|~m_main_v_state2(s14,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_755])).
cnf(c_0_2874,plain,(m_main_v_turn(X1,c_1)|~node18(s16,X1)|~trans(s64,s16)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_819]), c_0_856])).
cnf(c_0_2875,plain,(m_main_v_turn(s17,X1)|~m_main_v_turn(s16,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2838])).
cnf(c_0_2876,plain,(~m_main_v_state2(s19,c_t2)|~m_main_v_state2(s18,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_919])).
cnf(c_0_2877,plain,(m_main_v_turn(X1,c_1)|~node18(s20,X1)|~trans(s64,s20)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_969]), c_0_946])).
cnf(c_0_2878,plain,(m_main_v_turn(s21,X1)|~m_main_v_turn(s20,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2839])).
cnf(c_0_2879,plain,(m_main_v_turn(s21,c_2)), inference(sr,[status(thm)],[c_0_953, c_0_1031])).
cnf(c_0_2880,plain,(m_main_v_state2(s23,X1)|~m_main_v_state2(s22,X1)), inference(spm,[status(thm)],[c_0_43, c_0_2840])).
cnf(c_0_2881,plain,(m_main_v_turn(X1,c_1)|~node18(s24,X1)|~trans(s64,s24)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_1129]), c_0_1135])).
cnf(c_0_2882,plain,(m_main_v_turn(s25,X1)|~m_main_v_turn(s24,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2841])).
cnf(c_0_2883,plain,(~m_main_v_state2(s27,c_t2)|~m_main_v_state2(s26,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_1180])).
cnf(c_0_2884,plain,(m_main_v_turn(X1,c_1)|~node18(s28,X1)|~trans(s64,s28)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_1252]), c_0_1279])).
cnf(c_0_2885,plain,(m_main_v_turn(s29,X1)|~m_main_v_turn(s28,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2842])).
cnf(c_0_2886,plain,(~m_main_v_state2(s31,c_t2)|~m_main_v_state2(s30,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_1321])).
cnf(c_0_2887,plain,(m_main_v_turn(X1,c_1)|~node18(s32,X1)|~trans(s64,s32)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_1381]), c_0_1391])).
cnf(c_0_2888,plain,(m_main_v_turn(s33,X1)|~m_main_v_turn(s32,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2843])).
cnf(c_0_2889,plain,(~m_main_v_state2(s35,c_t2)|~m_main_v_state2(s34,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_1432])).
cnf(c_0_2890,plain,(node1(s35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2844]), c_0_1463])])).
cnf(c_0_2891,plain,(m_main_v_turn(X1,c_1)|~node18(s36,X1)|~trans(s64,s36)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_1492]), c_0_1502])).
cnf(c_0_2892,plain,(m_main_v_turn(s37,X1)|~m_main_v_turn(s36,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2845])).
cnf(c_0_2893,plain,(~m_main_v_state2(s39,c_t2)|~m_main_v_state2(s38,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_1542])).
cnf(c_0_2894,plain,(node1(s39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2846]), c_0_1573])])).
cnf(c_0_2895,plain,(m_main_v_turn(X1,c_1)|~node18(s40,X1)|~trans(s64,s40)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_1601]), c_0_1611])).
cnf(c_0_2896,plain,(m_main_v_turn(s41,X1)|~m_main_v_turn(s40,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2847])).
cnf(c_0_2897,plain,(~m_main_v_state2(s43,c_t2)|~m_main_v_state2(s42,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_1650])).
cnf(c_0_2898,plain,(node1(s43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2848]), c_0_1681])])).
cnf(c_0_2899,plain,(m_main_v_turn(X1,c_1)|~node18(s44,X1)|~trans(s64,s44)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_1710]), c_0_1720])).
cnf(c_0_2900,plain,(m_main_v_turn(s45,X1)|~m_main_v_turn(s44,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2849])).
cnf(c_0_2901,plain,(~m_main_v_state2(s47,c_t2)|~m_main_v_state2(s46,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_1760])).
cnf(c_0_2902,plain,(node1(s47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2850]), c_0_1791])])).
cnf(c_0_2903,plain,(m_main_v_turn(X1,c_1)|~node18(s48,X1)|~trans(s64,s48)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_1820]), c_0_1833])).
cnf(c_0_2904,plain,(m_main_v_turn(s49,X1)|~m_main_v_turn(s48,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2851])).
cnf(c_0_2905,plain,(~m_main_v_state2(s51,c_t2)|~m_main_v_state2(s50,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_1893])).
cnf(c_0_2906,plain,(node1(s51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2852]), c_0_1938])])).
cnf(c_0_2907,plain,(m_main_v_turn(X1,c_1)|~node18(s52,X1)|~trans(s64,s52)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_1980]), c_0_1997])).
cnf(c_0_2908,plain,(m_main_v_turn(s53,X1)|~m_main_v_turn(s52,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2853])).
cnf(c_0_2909,plain,(~m_main_v_state2(s55,c_t2)|~m_main_v_state2(s54,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_2065])).
cnf(c_0_2910,plain,(node1(s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2854]), c_0_2114])])).
cnf(c_0_2911,plain,(m_main_v_turn(X1,c_1)|~node18(s56,X1)|~trans(s64,s56)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_2159]), c_0_2178])).
cnf(c_0_2912,plain,(m_main_v_turn(s57,X1)|~m_main_v_turn(s56,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2855])).
cnf(c_0_2913,plain,(~m_main_v_state2(s59,c_t2)|~m_main_v_state2(s58,c_n2)), inference(spm,[status(thm)],[c_0_0, c_0_2245])).
cnf(c_0_2914,plain,(node1(s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2856]), c_0_2297])])).
cnf(c_0_2915,plain,(m_main_v_turn(X1,c_1)|~node18(s60,X1)|~trans(s64,s60)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2833, c_0_2341]), c_0_2358])).
cnf(c_0_2916,plain,(m_main_v_turn(s61,X1)|~m_main_v_turn(s60,X1)), inference(spm,[status(thm)],[c_0_24, c_0_2857])).
cnf(c_0_2917,plain,(m_main_v_state2(X1,c_c2)|~m_main_v_state2(s64,c_t2)|~trans(s64,X1)), inference(spm,[status(thm)],[c_0_2858, c_0_29])).
cnf(c_0_2918,plain,(m_main_v_state2(X1,c_t2)|~trans(s64,s62)|~trans(s63,X1)), inference(spm,[status(thm)],[c_0_2859, c_0_2802])).
cnf(c_0_2919,negated_conjecture,(always25(s64)|~always25(s62)), inference(spm,[status(thm)],[c_0_2860, c_0_2861])).
cnf(c_0_2920,negated_conjecture,(always25(s62)|loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2860, c_0_2862]), c_0_2863])).
cnf(c_0_2921,plain,(m_main_v_turn(s4,X1)|~m_main_v_turn(s5,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2834])).
cnf(c_0_2922,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s4)|~trans(s4,X1)), inference(spm,[status(thm)],[c_0_2864, c_0_15])).
cnf(c_0_2923,plain,(~m_main_v_turn(s4,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2865]), c_0_403])])).
cnf(c_0_2924,plain,(~m_main_v_state2(s6,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2866, c_0_17]), c_0_486])])).
cnf(c_0_2925,plain,(m_main_v_turn(s8,X1)|~m_main_v_turn(s9,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2835])).
cnf(c_0_2926,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s8)|~trans(s8,X1)), inference(spm,[status(thm)],[c_0_2867, c_0_15])).
cnf(c_0_2927,plain,(~m_main_v_turn(s8,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2868]), c_0_2869])])).
cnf(c_0_2928,plain,(~m_main_v_state2(s10,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_634]), c_0_2870])])).
cnf(c_0_2929,plain,(m_main_v_state2(s10,X1)|~m_main_v_state2(s11,X1)), inference(spm,[status(thm)],[c_0_47, c_0_628])).
cnf(c_0_2930,plain,(m_main_v_turn(s12,X1)|~m_main_v_turn(s13,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2837])).
cnf(c_0_2931,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s12)|~trans(s12,X1)), inference(spm,[status(thm)],[c_0_2871, c_0_15])).
cnf(c_0_2932,plain,(~m_main_v_turn(s12,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2872]), c_0_723])])).
cnf(c_0_2933,plain,(~m_main_v_state2(s14,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2873, c_0_755]), c_0_756])])).
cnf(c_0_2934,plain,(m_main_v_turn(s16,X1)|~m_main_v_turn(s17,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2838])).
cnf(c_0_2935,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s16)|~trans(s16,X1)), inference(spm,[status(thm)],[c_0_2874, c_0_15])).
cnf(c_0_2936,plain,(~m_main_v_turn(s16,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2875]), c_0_867])])).
cnf(c_0_2937,plain,(~m_main_v_state2(s18,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2876, c_0_17]), c_0_939])])).
cnf(c_0_2938,plain,(m_main_v_state2(s18,X1)|~m_main_v_state2(s19,X1)), inference(spm,[status(thm)],[c_0_47, c_0_914])).
cnf(c_0_2939,plain,(m_main_v_turn(s20,X1)|~m_main_v_turn(s21,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2839])).
cnf(c_0_2940,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s20)|~trans(s20,X1)), inference(spm,[status(thm)],[c_0_2877, c_0_15])).
cnf(c_0_2941,plain,(~m_main_v_turn(s20,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2878]), c_0_2879])])).
cnf(c_0_2942,plain,(~m_main_v_state2(s22,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_2880]), c_0_1062])])).
cnf(c_0_2943,plain,(m_main_v_state2(s22,X1)|~m_main_v_state2(s23,X1)), inference(spm,[status(thm)],[c_0_47, c_0_2840])).
cnf(c_0_2944,plain,(m_main_v_turn(s24,X1)|~m_main_v_turn(s25,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2841])).
cnf(c_0_2945,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s24)|~trans(s24,X1)), inference(spm,[status(thm)],[c_0_2881, c_0_15])).
cnf(c_0_2946,plain,(~m_main_v_turn(s24,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2882]), c_0_1144])])).
cnf(c_0_2947,plain,(~m_main_v_state2(s26,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2883, c_0_1180]), c_0_1181])])).
cnf(c_0_2948,plain,(m_main_v_state2(s26,X1)|~m_main_v_state2(s27,X1)), inference(spm,[status(thm)],[c_0_47, c_0_1173])).
cnf(c_0_2949,plain,(m_main_v_turn(s28,X1)|~m_main_v_turn(s29,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2842])).
cnf(c_0_2950,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s28)|~trans(s28,X1)), inference(spm,[status(thm)],[c_0_2884, c_0_15])).
cnf(c_0_2951,plain,(~m_main_v_turn(s28,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2885]), c_0_1285])])).
cnf(c_0_2952,plain,(~m_main_v_state2(s30,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2886, c_0_17]), c_0_1352])])).
cnf(c_0_2953,plain,(m_main_v_state2(s30,X1)|~m_main_v_state2(s31,X1)), inference(spm,[status(thm)],[c_0_47, c_0_1316])).
cnf(c_0_2954,plain,(m_main_v_turn(s32,X1)|~m_main_v_turn(s33,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2843])).
cnf(c_0_2955,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s32)|~trans(s32,X1)), inference(spm,[status(thm)],[c_0_2887, c_0_15])).
cnf(c_0_2956,plain,(~m_main_v_turn(s32,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2888]), c_0_1397])])).
cnf(c_0_2957,plain,(~m_main_v_state2(s34,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2889, c_0_5]), c_0_2890])])).
cnf(c_0_2958,plain,(m_main_v_turn(s36,X1)|~m_main_v_turn(s37,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2845])).
cnf(c_0_2959,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s36)|~trans(s36,X1)), inference(spm,[status(thm)],[c_0_2891, c_0_15])).
cnf(c_0_2960,plain,(~m_main_v_turn(s36,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2892]), c_0_1508])])).
cnf(c_0_2961,plain,(~m_main_v_state2(s38,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2893, c_0_5]), c_0_2894])])).
cnf(c_0_2962,plain,(m_main_v_turn(s40,X1)|~m_main_v_turn(s41,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2847])).
cnf(c_0_2963,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s40)|~trans(s40,X1)), inference(spm,[status(thm)],[c_0_2895, c_0_15])).
cnf(c_0_2964,plain,(~m_main_v_turn(s40,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2896]), c_0_1617])])).
cnf(c_0_2965,plain,(~m_main_v_state2(s42,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2897, c_0_5]), c_0_2898])])).
cnf(c_0_2966,plain,(m_main_v_turn(s44,X1)|~m_main_v_turn(s45,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2849])).
cnf(c_0_2967,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s44)|~trans(s44,X1)), inference(spm,[status(thm)],[c_0_2899, c_0_15])).
cnf(c_0_2968,plain,(~m_main_v_turn(s44,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2900]), c_0_1726])])).
cnf(c_0_2969,plain,(~m_main_v_state2(s46,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2901, c_0_5]), c_0_2902])])).
cnf(c_0_2970,plain,(m_main_v_turn(s48,X1)|~m_main_v_turn(s49,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2851])).
cnf(c_0_2971,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s48)|~trans(s48,X1)), inference(spm,[status(thm)],[c_0_2903, c_0_15])).
cnf(c_0_2972,plain,(~m_main_v_turn(s48,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2904]), c_0_1844])])).
cnf(c_0_2973,plain,(~m_main_v_state2(s50,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2905, c_0_5]), c_0_2906])])).
cnf(c_0_2974,plain,(m_main_v_turn(s52,X1)|~m_main_v_turn(s53,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2853])).
cnf(c_0_2975,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s52)|~trans(s52,X1)), inference(spm,[status(thm)],[c_0_2907, c_0_15])).
cnf(c_0_2976,plain,(~m_main_v_turn(s52,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2908]), c_0_2012])])).
cnf(c_0_2977,plain,(~m_main_v_state2(s54,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2909, c_0_5]), c_0_2910])])).
cnf(c_0_2978,plain,(m_main_v_state2(s54,c_n2)|~trans(s64,s55)), inference(spm,[status(thm)],[c_0_2695, c_0_2802])).
cnf(c_0_2979,plain,(m_main_v_turn(s56,X1)|~m_main_v_turn(s57,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2855])).
cnf(c_0_2980,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s56)|~trans(s56,X1)), inference(spm,[status(thm)],[c_0_2911, c_0_15])).
cnf(c_0_2981,plain,(~m_main_v_turn(s56,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2912]), c_0_2191])])).
cnf(c_0_2982,plain,(~m_main_v_state2(s58,c_n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2913, c_0_5]), c_0_2914])])).
cnf(c_0_2983,plain,(m_main_v_state2(s58,c_n2)|~trans(s64,s59)), inference(spm,[status(thm)],[c_0_2698, c_0_2802])).
cnf(c_0_2984,plain,(m_main_v_turn(s60,X1)|~m_main_v_turn(s61,X1)), inference(spm,[status(thm)],[c_0_66, c_0_2857])).
cnf(c_0_2985,plain,(m_main_v_turn(X1,c_1)|~trans(s64,s60)|~trans(s60,X1)), inference(spm,[status(thm)],[c_0_2915, c_0_15])).
cnf(c_0_2986,plain,(~m_main_v_turn(s60,c_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_2916]), c_0_2373])])).
cnf(c_0_2987,plain,(m_main_v_state2(X1,c_c2)|~trans(s64,s62)|~trans(s64,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2917, c_0_2918]), c_0_2420])])).
cnf(c_0_2988,plain,(m_main_v_state2(s62,c_n2)|~trans(s64,s63)), inference(spm,[status(thm)],[c_0_2634, c_0_2802])).
cnf(c_0_2989,plain,(~m_main_v_state2(X1,c_c2)|~trans(s64,X1)), inference(spm,[status(thm)],[c_0_4, c_0_2802])).
cnf(c_0_2990,negated_conjecture,(loop), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2919, c_0_2920]), c_0_2863])).
cnf(c_0_2991,plain,(~trans(s64,s1)), inference(spm,[status(thm)],[c_0_284, c_0_2802])).
cnf(c_0_2992,plain,(~trans(s64,s2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_2802]), c_0_284])).
cnf(c_0_2993,plain,(~trans(s64,s3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_2802]), c_0_286])).
cnf(c_0_2994,plain,(~trans(s64,s4)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2921, c_0_2922]), c_0_357])]), c_0_2923])).
cnf(c_0_2995,plain,(~trans(s64,s6)), inference(spm,[status(thm)],[c_0_2924, c_0_2802])).
cnf(c_0_2996,plain,(~trans(s64,s7)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_505, c_0_2802]), c_0_2924])).
cnf(c_0_2997,plain,(~trans(s64,s8)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2925, c_0_2926]), c_0_461])]), c_0_2927])).
cnf(c_0_2998,plain,(~trans(s64,s10)), inference(spm,[status(thm)],[c_0_2928, c_0_2802])).
cnf(c_0_2999,plain,(~trans(s64,s11)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2929, c_0_2802]), c_0_2928])).
cnf(c_0_3000,plain,(~trans(s64,s12)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2930, c_0_2931]), c_0_686])]), c_0_2932])).
cnf(c_0_3001,plain,(~trans(s64,s14)), inference(spm,[status(thm)],[c_0_2933, c_0_2802])).
cnf(c_0_3002,plain,(~trans(s64,s15)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_783, c_0_2802]), c_0_2933])).
cnf(c_0_3003,plain,(~trans(s64,s16)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2934, c_0_2935]), c_0_799])]), c_0_2936])).
cnf(c_0_3004,plain,(~trans(s64,s18)), inference(spm,[status(thm)],[c_0_2937, c_0_2802])).
cnf(c_0_3005,plain,(~trans(s64,s19)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2938, c_0_2802]), c_0_2937])).
cnf(c_0_3006,plain,(~trans(s64,s20)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2939, c_0_2940]), c_0_923])]), c_0_2941])).
cnf(c_0_3007,plain,(~trans(s64,s22)), inference(spm,[status(thm)],[c_0_2942, c_0_2802])).
cnf(c_0_3008,plain,(~trans(s64,s23)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2943, c_0_2802]), c_0_2942])).
cnf(c_0_3009,plain,(~trans(s64,s24)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2944, c_0_2945]), c_0_1107])]), c_0_2946])).
cnf(c_0_3010,plain,(~trans(s64,s26)), inference(spm,[status(thm)],[c_0_2947, c_0_2802])).
cnf(c_0_3011,plain,(~trans(s64,s27)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2948, c_0_2802]), c_0_2947])).
cnf(c_0_3012,plain,(~trans(s64,s28)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2949, c_0_2950]), c_0_1222])]), c_0_2951])).
cnf(c_0_3013,plain,(~trans(s64,s30)), inference(spm,[status(thm)],[c_0_2952, c_0_2802])).
cnf(c_0_3014,plain,(~trans(s64,s31)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2953, c_0_2802]), c_0_2952])).
cnf(c_0_3015,plain,(~trans(s64,s32)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2954, c_0_2955]), c_0_1358])]), c_0_2956])).
cnf(c_0_3016,plain,(~trans(s64,s34)), inference(spm,[status(thm)],[c_0_2957, c_0_2802])).
cnf(c_0_3017,plain,(~trans(s64,s35)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2680, c_0_2802]), c_0_2957])).
cnf(c_0_3018,plain,(~trans(s64,s36)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2958, c_0_2959]), c_0_1469])]), c_0_2960])).
cnf(c_0_3019,plain,(~trans(s64,s38)), inference(spm,[status(thm)],[c_0_2961, c_0_2802])).
cnf(c_0_3020,plain,(~trans(s64,s39)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2683, c_0_2802]), c_0_2961])).
cnf(c_0_3021,plain,(~trans(s64,s40)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2962, c_0_2963]), c_0_1579])]), c_0_2964])).
cnf(c_0_3022,plain,(~trans(s64,s42)), inference(spm,[status(thm)],[c_0_2965, c_0_2802])).
cnf(c_0_3023,plain,(~trans(s64,s43)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2686, c_0_2802]), c_0_2965])).
cnf(c_0_3024,plain,(~trans(s64,s44)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2966, c_0_2967]), c_0_1687])]), c_0_2968])).
cnf(c_0_3025,plain,(~trans(s64,s46)), inference(spm,[status(thm)],[c_0_2969, c_0_2802])).
cnf(c_0_3026,plain,(~trans(s64,s47)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2689, c_0_2802]), c_0_2969])).
cnf(c_0_3027,plain,(~trans(s64,s48)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2970, c_0_2971]), c_0_1797])]), c_0_2972])).
cnf(c_0_3028,plain,(~trans(s64,s50)), inference(spm,[status(thm)],[c_0_2973, c_0_2802])).
cnf(c_0_3029,plain,(~trans(s64,s51)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2692, c_0_2802]), c_0_2973])).
cnf(c_0_3030,plain,(~trans(s64,s52)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2974, c_0_2975]), c_0_1947])]), c_0_2976])).
cnf(c_0_3031,plain,(~trans(s64,s54)), inference(spm,[status(thm)],[c_0_2977, c_0_2802])).
cnf(c_0_3032,plain,(~trans(s64,s55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_2978]), c_0_2066])])).
cnf(c_0_3033,plain,(~trans(s64,s56)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2979, c_0_2980]), c_0_2124])]), c_0_2981])).
cnf(c_0_3034,plain,(~trans(s64,s58)), inference(spm,[status(thm)],[c_0_2982, c_0_2802])).
cnf(c_0_3035,plain,(~trans(s64,s59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_2983]), c_0_2246])])).
cnf(c_0_3036,plain,(~trans(s64,s60)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2984, c_0_2985]), c_0_2307])]), c_0_2986])).
cnf(c_0_3037,plain,(~trans(s64,s62)), inference(spm,[status(thm)],[c_0_2404, c_0_2987])).
cnf(c_0_3038,plain,(~trans(s64,s63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_2988]), c_0_2422])])).
cnf(c_0_3039,plain,(~trans(s64,s64)), inference(spm,[status(thm)],[c_0_2989, c_0_2475])).
cnf(c_0_3040,plain,(trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|trans(s64,s45)|trans(s64,s49)|trans(s64,s53)|trans(s64,s57)|trans(s64,s61)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_137, c_0_2990])]), c_0_2991]), c_0_2992]), c_0_2993]), c_0_2994]), c_0_2995]), c_0_2996]), c_0_2997]), c_0_2998]), c_0_2999]), c_0_3000]), c_0_3001]), c_0_3002]), c_0_3003]), c_0_3004]), c_0_3005]), c_0_3006]), c_0_3007]), c_0_3008]), c_0_3009]), c_0_3010]), c_0_3011]), c_0_3012]), c_0_3013]), c_0_3014]), c_0_3015]), c_0_3016]), c_0_3017]), c_0_3018]), c_0_3019]), c_0_3020]), c_0_3021]), c_0_3022]), c_0_3023]), c_0_3024]), c_0_3025]), c_0_3026]), c_0_3027]), c_0_3028]), c_0_3029]), c_0_3030]), c_0_3031]), c_0_3032]), c_0_3033]), c_0_3034]), c_0_3035]), c_0_3036]), c_0_3037]), c_0_3038]), c_0_3039])).
cnf(c_0_3041,negated_conjecture,(until2p29(s61)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|trans(s64,s45)|trans(s64,s49)|trans(s64,s53)|trans(s64,s57)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3040]), c_0_131])])).
cnf(c_0_3042,negated_conjecture,(until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|trans(s64,s45)|trans(s64,s49)|trans(s64,s53)|trans(s64,s57)), inference(spm,[status(thm)],[c_0_3041, c_0_2479])).
cnf(c_0_3043,negated_conjecture,(until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|trans(s64,s45)|trans(s64,s49)|trans(s64,s53)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3042]), c_0_131])])).
cnf(c_0_3044,negated_conjecture,(until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|trans(s64,s45)|trans(s64,s49)|trans(s64,s53)), inference(spm,[status(thm)],[c_0_3043, c_0_2479])).
cnf(c_0_3045,negated_conjecture,(until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|trans(s64,s45)|trans(s64,s49)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3044]), c_0_131])])).
cnf(c_0_3046,negated_conjecture,(until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|trans(s64,s45)|trans(s64,s49)), inference(spm,[status(thm)],[c_0_3045, c_0_2479])).
cnf(c_0_3047,negated_conjecture,(until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|trans(s64,s45)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3046]), c_0_131])])).
cnf(c_0_3048,negated_conjecture,(until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|trans(s64,s45)), inference(spm,[status(thm)],[c_0_3047, c_0_2479])).
cnf(c_0_3049,negated_conjecture,(until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3048]), c_0_131])])).
cnf(c_0_3050,negated_conjecture,(until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|trans(s64,s41)), inference(spm,[status(thm)],[c_0_3049, c_0_2479])).
cnf(c_0_3051,negated_conjecture,(until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3050]), c_0_131])])).
cnf(c_0_3052,negated_conjecture,(until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|trans(s64,s37)), inference(spm,[status(thm)],[c_0_3051, c_0_2479])).
cnf(c_0_3053,negated_conjecture,(until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3052]), c_0_131])])).
cnf(c_0_3054,negated_conjecture,(until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|trans(s64,s33)), inference(spm,[status(thm)],[c_0_3053, c_0_2479])).
cnf(c_0_3055,negated_conjecture,(until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3054]), c_0_131])])).
cnf(c_0_3056,negated_conjecture,(until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|trans(s64,s29)), inference(spm,[status(thm)],[c_0_3055, c_0_2479])).
cnf(c_0_3057,negated_conjecture,(until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3056]), c_0_131])])).
cnf(c_0_3058,negated_conjecture,(until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|trans(s64,s25)), inference(spm,[status(thm)],[c_0_3057, c_0_2479])).
cnf(c_0_3059,negated_conjecture,(until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3058]), c_0_131])])).
cnf(c_0_3060,negated_conjecture,(until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|trans(s64,s21)), inference(spm,[status(thm)],[c_0_3059, c_0_2479])).
cnf(c_0_3061,negated_conjecture,(until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3060]), c_0_131])])).
cnf(c_0_3062,negated_conjecture,(until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|trans(s64,s17)), inference(spm,[status(thm)],[c_0_3061, c_0_2479])).
cnf(c_0_3063,negated_conjecture,(until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3062]), c_0_131])])).
cnf(c_0_3064,negated_conjecture,(until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|trans(s64,s13)), inference(spm,[status(thm)],[c_0_3063, c_0_2479])).
cnf(c_0_3065,negated_conjecture,(until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3064]), c_0_131])])).
cnf(c_0_3066,negated_conjecture,(until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|trans(s64,s9)), inference(spm,[status(thm)],[c_0_3065, c_0_2479])).
cnf(c_0_3067,negated_conjecture,(until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3066]), c_0_131])])).
cnf(c_0_3068,negated_conjecture,(until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|trans(s64,s5)), inference(spm,[status(thm)],[c_0_3067, c_0_2479])).
cnf(c_0_3069,negated_conjecture,(until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3068]), c_0_131])])).
cnf(c_0_3070,negated_conjecture,(until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|trans(s64,s0)), inference(spm,[status(thm)],[c_0_3069, c_0_2479])).
cnf(c_0_3071,negated_conjecture,(until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|~xuntil28(s64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_3070]), c_0_131])])).
cnf(c_0_3072,negated_conjecture,(until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|until2p29(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3071, c_0_2479])).
cnf(c_0_3073,negated_conjecture,(xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|until2p29(s57)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3072]), c_0_2461])).
cnf(c_0_3074,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|until2p29(s53)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3073]), c_0_2419])).
cnf(c_0_3075,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|until2p29(s49)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3074]), c_0_2377])).
cnf(c_0_3076,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|until2p29(s45)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3075]), c_0_2345])).
cnf(c_0_3077,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|until2p29(s41)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3076]), c_0_2302])).
cnf(c_0_3078,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|until2p29(s37)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3077]), c_0_2262])).
cnf(c_0_3079,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|until2p29(s33)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3078]), c_0_2217])).
cnf(c_0_3080,negated_conjecture,(xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|until2p29(s29)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3079]), c_0_2182])).
cnf(c_0_3081,negated_conjecture,(xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|until2p29(s25)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3080]), c_0_2150])).
cnf(c_0_3082,negated_conjecture,(xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|until2p29(s21)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3081]), c_0_2101])).
cnf(c_0_3083,negated_conjecture,(xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|until2p29(s17)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3082]), c_0_2061])).
cnf(c_0_3084,negated_conjecture,(xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|until2p29(s13)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3083]), c_0_2016])).
cnf(c_0_3085,negated_conjecture,(xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|until2p29(s9)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3084]), c_0_1985])).
cnf(c_0_3086,negated_conjecture,(xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|until2p29(s5)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3085]), c_0_1955])).
cnf(c_0_3087,negated_conjecture,(until2p29(s62)|~xuntil2p30(s61)), inference(spm,[status(thm)],[c_0_140, c_0_127])).
cnf(c_0_3088,negated_conjecture,(xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|until2p29(s0)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3086]), c_0_1907])).
cnf(c_0_3089,negated_conjecture,(until2p29(s63)|~xuntil2p30(s62)), inference(spm,[status(thm)],[c_0_140, c_0_128])).
cnf(c_0_3090,negated_conjecture,(xuntil2p30(s62)|node26(s62)|~xuntil2p30(s61)), inference(spm,[status(thm)],[c_0_139, c_0_3087])).
cnf(c_0_3091,negated_conjecture,(xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|xuntil2p30(s61)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_3088]), c_0_1823])).
cnf(c_0_3092,negated_conjecture,(until2p29(s64)|~xuntil2p30(s63)), inference(spm,[status(thm)],[c_0_140, c_0_129])).
cnf(c_0_3093,negated_conjecture,(xuntil2p30(s63)|node26(s63)|~xuntil2p30(s62)), inference(spm,[status(thm)],[c_0_139, c_0_3089])).
cnf(c_0_3094,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|node26(s62)|node26(s63)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3091])).
cnf(c_0_3095,negated_conjecture,(until2p29(s58)|~xuntil2p30(s57)), inference(spm,[status(thm)],[c_0_140, c_0_123])).
cnf(c_0_3096,negated_conjecture,(xuntil2p30(s64)|node26(s64)|~xuntil2p30(s63)), inference(spm,[status(thm)],[c_0_139, c_0_3092])).
cnf(c_0_3097,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3094])).
cnf(c_0_3098,negated_conjecture,(until2p29(s59)|~xuntil2p30(s58)), inference(spm,[status(thm)],[c_0_140, c_0_124])).
cnf(c_0_3099,negated_conjecture,(xuntil2p30(s58)|node26(s58)|~xuntil2p30(s57)), inference(spm,[status(thm)],[c_0_139, c_0_3095])).
cnf(c_0_3100,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|xuntil2p30(s57)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3097]), c_0_2478])).
cnf(c_0_3101,negated_conjecture,(until2p29(s60)|~xuntil2p30(s59)), inference(spm,[status(thm)],[c_0_140, c_0_125])).
cnf(c_0_3102,negated_conjecture,(xuntil2p30(s59)|node26(s59)|~xuntil2p30(s58)), inference(spm,[status(thm)],[c_0_139, c_0_3098])).
cnf(c_0_3103,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|node26(s58)|node26(s63)|node26(s62)|node26(s59)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3100])).
cnf(c_0_3104,negated_conjecture,(until2p29(s61)|~xuntil2p30(s60)), inference(spm,[status(thm)],[c_0_140, c_0_126])).
cnf(c_0_3105,negated_conjecture,(xuntil2p30(s60)|node26(s60)|~xuntil2p30(s59)), inference(spm,[status(thm)],[c_0_139, c_0_3101])).
cnf(c_0_3106,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|node26(s59)|node26(s58)|node26(s63)|node26(s62)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3103])).
cnf(c_0_3107,negated_conjecture,(xuntil2p30(s61)|node26(s61)|~xuntil2p30(s60)), inference(spm,[status(thm)],[c_0_139, c_0_3104])).
cnf(c_0_3108,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|node26(s59)|node26(s58)|node26(s63)|node26(s62)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3106]), c_0_2452])).
cnf(c_0_3109,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|node26(s59)|node26(s58)|node26(s63)|node26(s62)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3108]), c_0_2461])).
cnf(c_0_3110,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|node26(s62)|node26(s59)|node26(s58)|node26(s63)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3109])).
cnf(c_0_3111,negated_conjecture,(until2p29(s54)|~xuntil2p30(s53)), inference(spm,[status(thm)],[c_0_140, c_0_119])).
cnf(c_0_3112,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3110])).
cnf(c_0_3113,negated_conjecture,(until2p29(s55)|~xuntil2p30(s54)), inference(spm,[status(thm)],[c_0_140, c_0_120])).
cnf(c_0_3114,negated_conjecture,(xuntil2p30(s54)|node26(s54)|~xuntil2p30(s53)), inference(spm,[status(thm)],[c_0_139, c_0_3111])).
cnf(c_0_3115,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|xuntil2p30(s53)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3112]), c_0_2478])).
cnf(c_0_3116,negated_conjecture,(until2p29(s56)|~xuntil2p30(s55)), inference(spm,[status(thm)],[c_0_140, c_0_121])).
cnf(c_0_3117,negated_conjecture,(xuntil2p30(s55)|node26(s55)|~xuntil2p30(s54)), inference(spm,[status(thm)],[c_0_139, c_0_3113])).
cnf(c_0_3118,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s54)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3115])).
cnf(c_0_3119,negated_conjecture,(until2p29(s57)|~xuntil2p30(s56)), inference(spm,[status(thm)],[c_0_140, c_0_122])).
cnf(c_0_3120,negated_conjecture,(xuntil2p30(s56)|node26(s56)|~xuntil2p30(s55)), inference(spm,[status(thm)],[c_0_139, c_0_3116])).
cnf(c_0_3121,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s55)|node26(s54)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3118])).
cnf(c_0_3122,negated_conjecture,(xuntil2p30(s57)|node26(s57)|~xuntil2p30(s56)), inference(spm,[status(thm)],[c_0_139, c_0_3119])).
cnf(c_0_3123,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s55)|node26(s54)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3121]), c_0_2409])).
cnf(c_0_3124,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s55)|node26(s54)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3123]), c_0_2419])).
cnf(c_0_3125,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s58)|node26(s55)|node26(s54)|node26(s63)|node26(s62)|node26(s59)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3124])).
cnf(c_0_3126,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s63)|node26(s62)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3125])).
cnf(c_0_3127,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s63)|node26(s62)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3126]), c_0_2452])).
cnf(c_0_3128,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s63)|node26(s62)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3127]), c_0_2461])).
cnf(c_0_3129,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s63)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3128])).
cnf(c_0_3130,negated_conjecture,(until2p29(s50)|~xuntil2p30(s49)), inference(spm,[status(thm)],[c_0_140, c_0_114])).
cnf(c_0_3131,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3129])).
cnf(c_0_3132,negated_conjecture,(until2p29(s51)|~xuntil2p30(s50)), inference(spm,[status(thm)],[c_0_140, c_0_116])).
cnf(c_0_3133,negated_conjecture,(xuntil2p30(s50)|node26(s50)|~xuntil2p30(s49)), inference(spm,[status(thm)],[c_0_139, c_0_3130])).
cnf(c_0_3134,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|xuntil2p30(s49)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3131]), c_0_2478])).
cnf(c_0_3135,negated_conjecture,(until2p29(s52)|~xuntil2p30(s51)), inference(spm,[status(thm)],[c_0_140, c_0_117])).
cnf(c_0_3136,negated_conjecture,(xuntil2p30(s51)|node26(s51)|~xuntil2p30(s50)), inference(spm,[status(thm)],[c_0_139, c_0_3132])).
cnf(c_0_3137,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s50)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3134])).
cnf(c_0_3138,negated_conjecture,(until2p29(s53)|~xuntil2p30(s52)), inference(spm,[status(thm)],[c_0_140, c_0_118])).
cnf(c_0_3139,negated_conjecture,(xuntil2p30(s52)|node26(s52)|~xuntil2p30(s51)), inference(spm,[status(thm)],[c_0_139, c_0_3135])).
cnf(c_0_3140,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3137])).
cnf(c_0_3141,negated_conjecture,(xuntil2p30(s53)|node26(s53)|~xuntil2p30(s52)), inference(spm,[status(thm)],[c_0_139, c_0_3138])).
cnf(c_0_3142,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3140]), c_0_2370])).
cnf(c_0_3143,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3142]), c_0_2377])).
cnf(c_0_3144,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s54)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3143])).
cnf(c_0_3145,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3144])).
cnf(c_0_3146,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3145]), c_0_2409])).
cnf(c_0_3147,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3146]), c_0_2419])).
cnf(c_0_3148,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s59)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3147])).
cnf(c_0_3149,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3148])).
cnf(c_0_3150,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3149]), c_0_2452])).
cnf(c_0_3151,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s63)|node26(s62)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3150]), c_0_2461])).
cnf(c_0_3152,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s63)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3151])).
cnf(c_0_3153,negated_conjecture,(until2p29(s46)|~xuntil2p30(s45)), inference(spm,[status(thm)],[c_0_140, c_0_107])).
cnf(c_0_3154,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3152])).
cnf(c_0_3155,negated_conjecture,(until2p29(s47)|~xuntil2p30(s46)), inference(spm,[status(thm)],[c_0_140, c_0_108])).
cnf(c_0_3156,negated_conjecture,(xuntil2p30(s46)|node26(s46)|~xuntil2p30(s45)), inference(spm,[status(thm)],[c_0_139, c_0_3153])).
cnf(c_0_3157,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|xuntil2p30(s45)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3154]), c_0_2478])).
cnf(c_0_3158,negated_conjecture,(until2p29(s48)|~xuntil2p30(s47)), inference(spm,[status(thm)],[c_0_140, c_0_109])).
cnf(c_0_3159,negated_conjecture,(xuntil2p30(s47)|node26(s47)|~xuntil2p30(s46)), inference(spm,[status(thm)],[c_0_139, c_0_3155])).
cnf(c_0_3160,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3157])).
cnf(c_0_3161,negated_conjecture,(until2p29(s49)|~xuntil2p30(s48)), inference(spm,[status(thm)],[c_0_140, c_0_110])).
cnf(c_0_3162,negated_conjecture,(xuntil2p30(s48)|node26(s48)|~xuntil2p30(s47)), inference(spm,[status(thm)],[c_0_139, c_0_3158])).
cnf(c_0_3163,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3160])).
cnf(c_0_3164,negated_conjecture,(xuntil2p30(s49)|node26(s49)|~xuntil2p30(s48)), inference(spm,[status(thm)],[c_0_139, c_0_3161])).
cnf(c_0_3165,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3163]), c_0_2337])).
cnf(c_0_3166,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3165]), c_0_2345])).
cnf(c_0_3167,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3166])).
cnf(c_0_3168,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3167])).
cnf(c_0_3169,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3168]), c_0_2370])).
cnf(c_0_3170,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3169]), c_0_2377])).
cnf(c_0_3171,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3170])).
cnf(c_0_3172,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3171])).
cnf(c_0_3173,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3172]), c_0_2409])).
cnf(c_0_3174,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3173]), c_0_2419])).
cnf(c_0_3175,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s59)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3174])).
cnf(c_0_3176,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3175])).
cnf(c_0_3177,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3176]), c_0_2452])).
cnf(c_0_3178,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s62)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3177]), c_0_2461])).
cnf(c_0_3179,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s63)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3178])).
cnf(c_0_3180,negated_conjecture,(until2p29(s42)|~xuntil2p30(s41)), inference(spm,[status(thm)],[c_0_140, c_0_103])).
cnf(c_0_3181,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3179])).
cnf(c_0_3182,negated_conjecture,(until2p29(s43)|~xuntil2p30(s42)), inference(spm,[status(thm)],[c_0_140, c_0_104])).
cnf(c_0_3183,negated_conjecture,(xuntil2p30(s42)|node26(s42)|~xuntil2p30(s41)), inference(spm,[status(thm)],[c_0_139, c_0_3180])).
cnf(c_0_3184,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|xuntil2p30(s41)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3181]), c_0_2478])).
cnf(c_0_3185,negated_conjecture,(until2p29(s44)|~xuntil2p30(s43)), inference(spm,[status(thm)],[c_0_140, c_0_105])).
cnf(c_0_3186,negated_conjecture,(xuntil2p30(s43)|node26(s43)|~xuntil2p30(s42)), inference(spm,[status(thm)],[c_0_139, c_0_3182])).
cnf(c_0_3187,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3184])).
cnf(c_0_3188,negated_conjecture,(until2p29(s45)|~xuntil2p30(s44)), inference(spm,[status(thm)],[c_0_140, c_0_106])).
cnf(c_0_3189,negated_conjecture,(xuntil2p30(s44)|node26(s44)|~xuntil2p30(s43)), inference(spm,[status(thm)],[c_0_139, c_0_3185])).
cnf(c_0_3190,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3187])).
cnf(c_0_3191,negated_conjecture,(xuntil2p30(s45)|node26(s45)|~xuntil2p30(s44)), inference(spm,[status(thm)],[c_0_139, c_0_3188])).
cnf(c_0_3192,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3190]), c_0_2289])).
cnf(c_0_3193,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3192]), c_0_2302])).
cnf(c_0_3194,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3193])).
cnf(c_0_3195,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3194])).
cnf(c_0_3196,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3195]), c_0_2337])).
cnf(c_0_3197,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3196]), c_0_2345])).
cnf(c_0_3198,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3197])).
cnf(c_0_3199,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3198])).
cnf(c_0_3200,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3199]), c_0_2370])).
cnf(c_0_3201,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3200]), c_0_2377])).
cnf(c_0_3202,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3201])).
cnf(c_0_3203,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3202])).
cnf(c_0_3204,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3203]), c_0_2409])).
cnf(c_0_3205,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3204]), c_0_2419])).
cnf(c_0_3206,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s59)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3205])).
cnf(c_0_3207,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3206])).
cnf(c_0_3208,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3207]), c_0_2452])).
cnf(c_0_3209,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s62)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3208]), c_0_2461])).
cnf(c_0_3210,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s63)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3209])).
cnf(c_0_3211,negated_conjecture,(until2p29(s38)|~xuntil2p30(s37)), inference(spm,[status(thm)],[c_0_140, c_0_99])).
cnf(c_0_3212,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3210])).
cnf(c_0_3213,negated_conjecture,(until2p29(s39)|~xuntil2p30(s38)), inference(spm,[status(thm)],[c_0_140, c_0_100])).
cnf(c_0_3214,negated_conjecture,(xuntil2p30(s38)|node26(s38)|~xuntil2p30(s37)), inference(spm,[status(thm)],[c_0_139, c_0_3211])).
cnf(c_0_3215,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|xuntil2p30(s37)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3212]), c_0_2478])).
cnf(c_0_3216,negated_conjecture,(until2p29(s40)|~xuntil2p30(s39)), inference(spm,[status(thm)],[c_0_140, c_0_101])).
cnf(c_0_3217,negated_conjecture,(xuntil2p30(s39)|node26(s39)|~xuntil2p30(s38)), inference(spm,[status(thm)],[c_0_139, c_0_3213])).
cnf(c_0_3218,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3214, c_0_3215])).
cnf(c_0_3219,negated_conjecture,(until2p29(s41)|~xuntil2p30(s40)), inference(spm,[status(thm)],[c_0_140, c_0_102])).
cnf(c_0_3220,negated_conjecture,(xuntil2p30(s40)|node26(s40)|~xuntil2p30(s39)), inference(spm,[status(thm)],[c_0_139, c_0_3216])).
cnf(c_0_3221,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3217, c_0_3218])).
cnf(c_0_3222,negated_conjecture,(xuntil2p30(s41)|node26(s41)|~xuntil2p30(s40)), inference(spm,[status(thm)],[c_0_139, c_0_3219])).
cnf(c_0_3223,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3221]), c_0_2251])).
cnf(c_0_3224,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3223]), c_0_2262])).
cnf(c_0_3225,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3224])).
cnf(c_0_3226,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3225])).
cnf(c_0_3227,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3226]), c_0_2289])).
cnf(c_0_3228,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3227]), c_0_2302])).
cnf(c_0_3229,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3228])).
cnf(c_0_3230,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3229])).
cnf(c_0_3231,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3230]), c_0_2337])).
cnf(c_0_3232,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3231]), c_0_2345])).
cnf(c_0_3233,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3232])).
cnf(c_0_3234,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3233])).
cnf(c_0_3235,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3234]), c_0_2370])).
cnf(c_0_3236,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3235]), c_0_2377])).
cnf(c_0_3237,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3236])).
cnf(c_0_3238,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3237])).
cnf(c_0_3239,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3238]), c_0_2409])).
cnf(c_0_3240,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3239]), c_0_2419])).
cnf(c_0_3241,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s59)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3240])).
cnf(c_0_3242,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3241])).
cnf(c_0_3243,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3242]), c_0_2452])).
cnf(c_0_3244,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s62)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3243]), c_0_2461])).
cnf(c_0_3245,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s63)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3244])).
cnf(c_0_3246,negated_conjecture,(until2p29(s34)|~xuntil2p30(s33)), inference(spm,[status(thm)],[c_0_140, c_0_95])).
cnf(c_0_3247,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3245])).
cnf(c_0_3248,negated_conjecture,(until2p29(s35)|~xuntil2p30(s34)), inference(spm,[status(thm)],[c_0_140, c_0_96])).
cnf(c_0_3249,negated_conjecture,(xuntil2p30(s34)|node26(s34)|~xuntil2p30(s33)), inference(spm,[status(thm)],[c_0_139, c_0_3246])).
cnf(c_0_3250,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|xuntil2p30(s33)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3247]), c_0_2478])).
cnf(c_0_3251,negated_conjecture,(until2p29(s36)|~xuntil2p30(s35)), inference(spm,[status(thm)],[c_0_140, c_0_97])).
cnf(c_0_3252,negated_conjecture,(xuntil2p30(s35)|node26(s35)|~xuntil2p30(s34)), inference(spm,[status(thm)],[c_0_139, c_0_3248])).
cnf(c_0_3253,negated_conjecture,(xuntil2p30(s34)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3249, c_0_3250])).
cnf(c_0_3254,negated_conjecture,(until2p29(s37)|~xuntil2p30(s36)), inference(spm,[status(thm)],[c_0_140, c_0_98])).
cnf(c_0_3255,negated_conjecture,(xuntil2p30(s36)|node26(s36)|~xuntil2p30(s35)), inference(spm,[status(thm)],[c_0_139, c_0_3251])).
cnf(c_0_3256,negated_conjecture,(xuntil2p30(s35)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3252, c_0_3253])).
cnf(c_0_3257,negated_conjecture,(xuntil2p30(s37)|node26(s37)|~xuntil2p30(s36)), inference(spm,[status(thm)],[c_0_139, c_0_3254])).
cnf(c_0_3258,negated_conjecture,(xuntil2p30(s36)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3255, c_0_3256]), c_0_2205])).
cnf(c_0_3259,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3257, c_0_3258]), c_0_2217])).
cnf(c_0_3260,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3214, c_0_3259])).
cnf(c_0_3261,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3217, c_0_3260])).
cnf(c_0_3262,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3261]), c_0_2251])).
cnf(c_0_3263,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3262]), c_0_2262])).
cnf(c_0_3264,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3263])).
cnf(c_0_3265,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3264])).
cnf(c_0_3266,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3265]), c_0_2289])).
cnf(c_0_3267,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3266]), c_0_2302])).
cnf(c_0_3268,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3267])).
cnf(c_0_3269,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3268])).
cnf(c_0_3270,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3269]), c_0_2337])).
cnf(c_0_3271,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3270]), c_0_2345])).
cnf(c_0_3272,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3271])).
cnf(c_0_3273,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3272])).
cnf(c_0_3274,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3273]), c_0_2370])).
cnf(c_0_3275,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3274]), c_0_2377])).
cnf(c_0_3276,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3275])).
cnf(c_0_3277,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3276])).
cnf(c_0_3278,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3277]), c_0_2409])).
cnf(c_0_3279,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3278]), c_0_2419])).
cnf(c_0_3280,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s59)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3279])).
cnf(c_0_3281,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3280])).
cnf(c_0_3282,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3281]), c_0_2452])).
cnf(c_0_3283,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s62)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3282]), c_0_2461])).
cnf(c_0_3284,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s63)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3283])).
cnf(c_0_3285,negated_conjecture,(until2p29(s30)|~xuntil2p30(s29)), inference(spm,[status(thm)],[c_0_140, c_0_91])).
cnf(c_0_3286,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3284])).
cnf(c_0_3287,negated_conjecture,(until2p29(s31)|~xuntil2p30(s30)), inference(spm,[status(thm)],[c_0_140, c_0_92])).
cnf(c_0_3288,negated_conjecture,(xuntil2p30(s30)|node26(s30)|~xuntil2p30(s29)), inference(spm,[status(thm)],[c_0_139, c_0_3285])).
cnf(c_0_3289,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|xuntil2p30(s29)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3286]), c_0_2478])).
cnf(c_0_3290,negated_conjecture,(until2p29(s32)|~xuntil2p30(s31)), inference(spm,[status(thm)],[c_0_140, c_0_93])).
cnf(c_0_3291,negated_conjecture,(xuntil2p30(s31)|node26(s31)|~xuntil2p30(s30)), inference(spm,[status(thm)],[c_0_139, c_0_3287])).
cnf(c_0_3292,negated_conjecture,(xuntil2p30(s30)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3288, c_0_3289])).
cnf(c_0_3293,negated_conjecture,(until2p29(s33)|~xuntil2p30(s32)), inference(spm,[status(thm)],[c_0_140, c_0_94])).
cnf(c_0_3294,negated_conjecture,(xuntil2p30(s32)|node26(s32)|~xuntil2p30(s31)), inference(spm,[status(thm)],[c_0_139, c_0_3290])).
cnf(c_0_3295,negated_conjecture,(xuntil2p30(s31)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3291, c_0_3292])).
cnf(c_0_3296,negated_conjecture,(xuntil2p30(s33)|node26(s33)|~xuntil2p30(s32)), inference(spm,[status(thm)],[c_0_139, c_0_3293])).
cnf(c_0_3297,negated_conjecture,(xuntil2p30(s32)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3294, c_0_3295]), c_0_2173])).
cnf(c_0_3298,negated_conjecture,(xuntil2p30(s33)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3296, c_0_3297]), c_0_2182])).
cnf(c_0_3299,negated_conjecture,(xuntil2p30(s34)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3249, c_0_3298])).
cnf(c_0_3300,negated_conjecture,(xuntil2p30(s35)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3252, c_0_3299])).
cnf(c_0_3301,negated_conjecture,(xuntil2p30(s36)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3255, c_0_3300]), c_0_2205])).
cnf(c_0_3302,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3257, c_0_3301]), c_0_2217])).
cnf(c_0_3303,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3214, c_0_3302])).
cnf(c_0_3304,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3217, c_0_3303])).
cnf(c_0_3305,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3304]), c_0_2251])).
cnf(c_0_3306,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3305]), c_0_2262])).
cnf(c_0_3307,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3306])).
cnf(c_0_3308,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3307])).
cnf(c_0_3309,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3308]), c_0_2289])).
cnf(c_0_3310,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3309]), c_0_2302])).
cnf(c_0_3311,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3310])).
cnf(c_0_3312,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3311])).
cnf(c_0_3313,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3312]), c_0_2337])).
cnf(c_0_3314,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3313]), c_0_2345])).
cnf(c_0_3315,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3314])).
cnf(c_0_3316,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3315])).
cnf(c_0_3317,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3316]), c_0_2370])).
cnf(c_0_3318,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3317]), c_0_2377])).
cnf(c_0_3319,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3318])).
cnf(c_0_3320,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3319])).
cnf(c_0_3321,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3320]), c_0_2409])).
cnf(c_0_3322,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3321]), c_0_2419])).
cnf(c_0_3323,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s59)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3322])).
cnf(c_0_3324,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3323])).
cnf(c_0_3325,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3324]), c_0_2452])).
cnf(c_0_3326,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s62)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3325]), c_0_2461])).
cnf(c_0_3327,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s63)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3326])).
cnf(c_0_3328,negated_conjecture,(until2p29(s26)|~xuntil2p30(s25)), inference(spm,[status(thm)],[c_0_140, c_0_87])).
cnf(c_0_3329,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3327])).
cnf(c_0_3330,negated_conjecture,(until2p29(s27)|~xuntil2p30(s26)), inference(spm,[status(thm)],[c_0_140, c_0_88])).
cnf(c_0_3331,negated_conjecture,(xuntil2p30(s26)|node26(s26)|~xuntil2p30(s25)), inference(spm,[status(thm)],[c_0_139, c_0_3328])).
cnf(c_0_3332,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|xuntil2p30(s25)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3329]), c_0_2478])).
cnf(c_0_3333,negated_conjecture,(until2p29(s28)|~xuntil2p30(s27)), inference(spm,[status(thm)],[c_0_140, c_0_89])).
cnf(c_0_3334,negated_conjecture,(xuntil2p30(s27)|node26(s27)|~xuntil2p30(s26)), inference(spm,[status(thm)],[c_0_139, c_0_3330])).
cnf(c_0_3335,negated_conjecture,(xuntil2p30(s26)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3331, c_0_3332])).
cnf(c_0_3336,negated_conjecture,(until2p29(s29)|~xuntil2p30(s28)), inference(spm,[status(thm)],[c_0_140, c_0_90])).
cnf(c_0_3337,negated_conjecture,(xuntil2p30(s28)|node26(s28)|~xuntil2p30(s27)), inference(spm,[status(thm)],[c_0_139, c_0_3333])).
cnf(c_0_3338,negated_conjecture,(xuntil2p30(s27)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3334, c_0_3335])).
cnf(c_0_3339,negated_conjecture,(xuntil2p30(s29)|node26(s29)|~xuntil2p30(s28)), inference(spm,[status(thm)],[c_0_139, c_0_3336])).
cnf(c_0_3340,negated_conjecture,(xuntil2p30(s28)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3337, c_0_3338]), c_0_2133])).
cnf(c_0_3341,negated_conjecture,(xuntil2p30(s29)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3339, c_0_3340]), c_0_2150])).
cnf(c_0_3342,negated_conjecture,(xuntil2p30(s30)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3288, c_0_3341])).
cnf(c_0_3343,negated_conjecture,(xuntil2p30(s31)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3291, c_0_3342])).
cnf(c_0_3344,negated_conjecture,(xuntil2p30(s32)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3294, c_0_3343]), c_0_2173])).
cnf(c_0_3345,negated_conjecture,(xuntil2p30(s33)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3296, c_0_3344]), c_0_2182])).
cnf(c_0_3346,negated_conjecture,(xuntil2p30(s34)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3249, c_0_3345])).
cnf(c_0_3347,negated_conjecture,(xuntil2p30(s35)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3252, c_0_3346])).
cnf(c_0_3348,negated_conjecture,(xuntil2p30(s36)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3255, c_0_3347]), c_0_2205])).
cnf(c_0_3349,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3257, c_0_3348]), c_0_2217])).
cnf(c_0_3350,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3214, c_0_3349])).
cnf(c_0_3351,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3217, c_0_3350])).
cnf(c_0_3352,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3351]), c_0_2251])).
cnf(c_0_3353,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3352]), c_0_2262])).
cnf(c_0_3354,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3353])).
cnf(c_0_3355,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3354])).
cnf(c_0_3356,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3355]), c_0_2289])).
cnf(c_0_3357,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3356]), c_0_2302])).
cnf(c_0_3358,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3357])).
cnf(c_0_3359,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3358])).
cnf(c_0_3360,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3359]), c_0_2337])).
cnf(c_0_3361,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3360]), c_0_2345])).
cnf(c_0_3362,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3361])).
cnf(c_0_3363,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3362])).
cnf(c_0_3364,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3363]), c_0_2370])).
cnf(c_0_3365,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3364]), c_0_2377])).
cnf(c_0_3366,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3365])).
cnf(c_0_3367,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3366])).
cnf(c_0_3368,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3367]), c_0_2409])).
cnf(c_0_3369,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3368]), c_0_2419])).
cnf(c_0_3370,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s59)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3369])).
cnf(c_0_3371,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3370])).
cnf(c_0_3372,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3371]), c_0_2452])).
cnf(c_0_3373,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s62)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3372]), c_0_2461])).
cnf(c_0_3374,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s63)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3373])).
cnf(c_0_3375,negated_conjecture,(until2p29(s22)|~xuntil2p30(s21)), inference(spm,[status(thm)],[c_0_140, c_0_83])).
cnf(c_0_3376,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3374])).
cnf(c_0_3377,negated_conjecture,(until2p29(s23)|~xuntil2p30(s22)), inference(spm,[status(thm)],[c_0_140, c_0_84])).
cnf(c_0_3378,negated_conjecture,(xuntil2p30(s22)|node26(s22)|~xuntil2p30(s21)), inference(spm,[status(thm)],[c_0_139, c_0_3375])).
cnf(c_0_3379,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|xuntil2p30(s21)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3376]), c_0_2478])).
cnf(c_0_3380,negated_conjecture,(until2p29(s24)|~xuntil2p30(s23)), inference(spm,[status(thm)],[c_0_140, c_0_85])).
cnf(c_0_3381,negated_conjecture,(xuntil2p30(s23)|node26(s23)|~xuntil2p30(s22)), inference(spm,[status(thm)],[c_0_139, c_0_3377])).
cnf(c_0_3382,negated_conjecture,(xuntil2p30(s22)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3378, c_0_3379])).
cnf(c_0_3383,negated_conjecture,(until2p29(s25)|~xuntil2p30(s24)), inference(spm,[status(thm)],[c_0_140, c_0_86])).
cnf(c_0_3384,negated_conjecture,(xuntil2p30(s24)|node26(s24)|~xuntil2p30(s23)), inference(spm,[status(thm)],[c_0_139, c_0_3380])).
cnf(c_0_3385,negated_conjecture,(xuntil2p30(s23)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3381, c_0_3382])).
cnf(c_0_3386,negated_conjecture,(xuntil2p30(s25)|node26(s25)|~xuntil2p30(s24)), inference(spm,[status(thm)],[c_0_139, c_0_3383])).
cnf(c_0_3387,negated_conjecture,(xuntil2p30(s24)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3384, c_0_3385]), c_0_2091])).
cnf(c_0_3388,negated_conjecture,(xuntil2p30(s25)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3386, c_0_3387]), c_0_2101])).
cnf(c_0_3389,negated_conjecture,(xuntil2p30(s26)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3331, c_0_3388])).
cnf(c_0_3390,negated_conjecture,(xuntil2p30(s27)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3334, c_0_3389])).
cnf(c_0_3391,negated_conjecture,(xuntil2p30(s28)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3337, c_0_3390]), c_0_2133])).
cnf(c_0_3392,negated_conjecture,(xuntil2p30(s29)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3339, c_0_3391]), c_0_2150])).
cnf(c_0_3393,negated_conjecture,(xuntil2p30(s30)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3288, c_0_3392])).
cnf(c_0_3394,negated_conjecture,(xuntil2p30(s31)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3291, c_0_3393])).
cnf(c_0_3395,negated_conjecture,(xuntil2p30(s32)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3294, c_0_3394]), c_0_2173])).
cnf(c_0_3396,negated_conjecture,(xuntil2p30(s33)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3296, c_0_3395]), c_0_2182])).
cnf(c_0_3397,negated_conjecture,(xuntil2p30(s34)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3249, c_0_3396])).
cnf(c_0_3398,negated_conjecture,(xuntil2p30(s35)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3252, c_0_3397])).
cnf(c_0_3399,negated_conjecture,(xuntil2p30(s36)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3255, c_0_3398]), c_0_2205])).
cnf(c_0_3400,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3257, c_0_3399]), c_0_2217])).
cnf(c_0_3401,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3214, c_0_3400])).
cnf(c_0_3402,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3217, c_0_3401])).
cnf(c_0_3403,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3402]), c_0_2251])).
cnf(c_0_3404,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3403]), c_0_2262])).
cnf(c_0_3405,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3404])).
cnf(c_0_3406,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3405])).
cnf(c_0_3407,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3406]), c_0_2289])).
cnf(c_0_3408,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3407]), c_0_2302])).
cnf(c_0_3409,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3408])).
cnf(c_0_3410,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3409])).
cnf(c_0_3411,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3410]), c_0_2337])).
cnf(c_0_3412,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3411]), c_0_2345])).
cnf(c_0_3413,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3412])).
cnf(c_0_3414,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3413])).
cnf(c_0_3415,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3414]), c_0_2370])).
cnf(c_0_3416,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3415]), c_0_2377])).
cnf(c_0_3417,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3416])).
cnf(c_0_3418,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3417])).
cnf(c_0_3419,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3418]), c_0_2409])).
cnf(c_0_3420,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3419]), c_0_2419])).
cnf(c_0_3421,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s59)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3420])).
cnf(c_0_3422,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3421])).
cnf(c_0_3423,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3422]), c_0_2452])).
cnf(c_0_3424,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s62)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3423]), c_0_2461])).
cnf(c_0_3425,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s63)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3424])).
cnf(c_0_3426,negated_conjecture,(until2p29(s18)|~xuntil2p30(s17)), inference(spm,[status(thm)],[c_0_140, c_0_79])).
cnf(c_0_3427,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3425])).
cnf(c_0_3428,negated_conjecture,(until2p29(s19)|~xuntil2p30(s18)), inference(spm,[status(thm)],[c_0_140, c_0_80])).
cnf(c_0_3429,negated_conjecture,(xuntil2p30(s18)|node26(s18)|~xuntil2p30(s17)), inference(spm,[status(thm)],[c_0_139, c_0_3426])).
cnf(c_0_3430,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|xuntil2p30(s17)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3427]), c_0_2478])).
cnf(c_0_3431,negated_conjecture,(until2p29(s20)|~xuntil2p30(s19)), inference(spm,[status(thm)],[c_0_140, c_0_81])).
cnf(c_0_3432,negated_conjecture,(xuntil2p30(s19)|node26(s19)|~xuntil2p30(s18)), inference(spm,[status(thm)],[c_0_139, c_0_3428])).
cnf(c_0_3433,negated_conjecture,(xuntil2p30(s18)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3429, c_0_3430])).
cnf(c_0_3434,negated_conjecture,(until2p29(s21)|~xuntil2p30(s20)), inference(spm,[status(thm)],[c_0_140, c_0_82])).
cnf(c_0_3435,negated_conjecture,(xuntil2p30(s20)|node26(s20)|~xuntil2p30(s19)), inference(spm,[status(thm)],[c_0_139, c_0_3431])).
cnf(c_0_3436,negated_conjecture,(xuntil2p30(s19)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3432, c_0_3433])).
cnf(c_0_3437,negated_conjecture,(xuntil2p30(s21)|node26(s21)|~xuntil2p30(s20)), inference(spm,[status(thm)],[c_0_139, c_0_3434])).
cnf(c_0_3438,negated_conjecture,(xuntil2p30(s20)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3435, c_0_3436]), c_0_2050])).
cnf(c_0_3439,negated_conjecture,(xuntil2p30(s21)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3437, c_0_3438]), c_0_2061])).
cnf(c_0_3440,negated_conjecture,(xuntil2p30(s22)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3378, c_0_3439])).
cnf(c_0_3441,negated_conjecture,(xuntil2p30(s23)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3381, c_0_3440])).
cnf(c_0_3442,negated_conjecture,(xuntil2p30(s24)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3384, c_0_3441]), c_0_2091])).
cnf(c_0_3443,negated_conjecture,(xuntil2p30(s25)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3386, c_0_3442]), c_0_2101])).
cnf(c_0_3444,negated_conjecture,(xuntil2p30(s26)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3331, c_0_3443])).
cnf(c_0_3445,negated_conjecture,(xuntil2p30(s27)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3334, c_0_3444])).
cnf(c_0_3446,negated_conjecture,(xuntil2p30(s28)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3337, c_0_3445]), c_0_2133])).
cnf(c_0_3447,negated_conjecture,(xuntil2p30(s29)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3339, c_0_3446]), c_0_2150])).
cnf(c_0_3448,negated_conjecture,(xuntil2p30(s30)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3288, c_0_3447])).
cnf(c_0_3449,negated_conjecture,(xuntil2p30(s31)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3291, c_0_3448])).
cnf(c_0_3450,negated_conjecture,(xuntil2p30(s32)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3294, c_0_3449]), c_0_2173])).
cnf(c_0_3451,negated_conjecture,(xuntil2p30(s33)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3296, c_0_3450]), c_0_2182])).
cnf(c_0_3452,negated_conjecture,(xuntil2p30(s34)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3249, c_0_3451])).
cnf(c_0_3453,negated_conjecture,(xuntil2p30(s35)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3252, c_0_3452])).
cnf(c_0_3454,negated_conjecture,(xuntil2p30(s36)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3255, c_0_3453]), c_0_2205])).
cnf(c_0_3455,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3257, c_0_3454]), c_0_2217])).
cnf(c_0_3456,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3214, c_0_3455])).
cnf(c_0_3457,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3217, c_0_3456])).
cnf(c_0_3458,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3457]), c_0_2251])).
cnf(c_0_3459,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3458]), c_0_2262])).
cnf(c_0_3460,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3459])).
cnf(c_0_3461,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3460])).
cnf(c_0_3462,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3461]), c_0_2289])).
cnf(c_0_3463,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3462]), c_0_2302])).
cnf(c_0_3464,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3463])).
cnf(c_0_3465,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3464])).
cnf(c_0_3466,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3465]), c_0_2337])).
cnf(c_0_3467,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3466]), c_0_2345])).
cnf(c_0_3468,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3467])).
cnf(c_0_3469,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3468])).
cnf(c_0_3470,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3469]), c_0_2370])).
cnf(c_0_3471,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3470]), c_0_2377])).
cnf(c_0_3472,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3471])).
cnf(c_0_3473,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3472])).
cnf(c_0_3474,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3473]), c_0_2409])).
cnf(c_0_3475,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3474]), c_0_2419])).
cnf(c_0_3476,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s59)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3475])).
cnf(c_0_3477,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3476])).
cnf(c_0_3478,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3477]), c_0_2452])).
cnf(c_0_3479,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s62)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3478]), c_0_2461])).
cnf(c_0_3480,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s63)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3479])).
cnf(c_0_3481,negated_conjecture,(until2p29(s14)|~xuntil2p30(s13)), inference(spm,[status(thm)],[c_0_140, c_0_75])).
cnf(c_0_3482,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3480])).
cnf(c_0_3483,negated_conjecture,(until2p29(s15)|~xuntil2p30(s14)), inference(spm,[status(thm)],[c_0_140, c_0_76])).
cnf(c_0_3484,negated_conjecture,(xuntil2p30(s14)|node26(s14)|~xuntil2p30(s13)), inference(spm,[status(thm)],[c_0_139, c_0_3481])).
cnf(c_0_3485,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|xuntil2p30(s13)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3482]), c_0_2478])).
cnf(c_0_3486,negated_conjecture,(until2p29(s16)|~xuntil2p30(s15)), inference(spm,[status(thm)],[c_0_140, c_0_77])).
cnf(c_0_3487,negated_conjecture,(xuntil2p30(s15)|node26(s15)|~xuntil2p30(s14)), inference(spm,[status(thm)],[c_0_139, c_0_3483])).
cnf(c_0_3488,negated_conjecture,(xuntil2p30(s14)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3484, c_0_3485])).
cnf(c_0_3489,negated_conjecture,(until2p29(s17)|~xuntil2p30(s16)), inference(spm,[status(thm)],[c_0_140, c_0_78])).
cnf(c_0_3490,negated_conjecture,(xuntil2p30(s16)|node26(s16)|~xuntil2p30(s15)), inference(spm,[status(thm)],[c_0_139, c_0_3486])).
cnf(c_0_3491,negated_conjecture,(xuntil2p30(s15)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3487, c_0_3488])).
cnf(c_0_3492,negated_conjecture,(xuntil2p30(s17)|node26(s17)|~xuntil2p30(s16)), inference(spm,[status(thm)],[c_0_139, c_0_3489])).
cnf(c_0_3493,negated_conjecture,(xuntil2p30(s16)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3490, c_0_3491]), c_0_2009])).
cnf(c_0_3494,negated_conjecture,(xuntil2p30(s17)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3492, c_0_3493]), c_0_2016])).
cnf(c_0_3495,negated_conjecture,(xuntil2p30(s18)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3429, c_0_3494])).
cnf(c_0_3496,negated_conjecture,(xuntil2p30(s19)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3432, c_0_3495])).
cnf(c_0_3497,negated_conjecture,(xuntil2p30(s20)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3435, c_0_3496]), c_0_2050])).
cnf(c_0_3498,negated_conjecture,(xuntil2p30(s21)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3437, c_0_3497]), c_0_2061])).
cnf(c_0_3499,negated_conjecture,(xuntil2p30(s22)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3378, c_0_3498])).
cnf(c_0_3500,negated_conjecture,(xuntil2p30(s23)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3381, c_0_3499])).
cnf(c_0_3501,negated_conjecture,(xuntil2p30(s24)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3384, c_0_3500]), c_0_2091])).
cnf(c_0_3502,negated_conjecture,(xuntil2p30(s25)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3386, c_0_3501]), c_0_2101])).
cnf(c_0_3503,negated_conjecture,(xuntil2p30(s26)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3331, c_0_3502])).
cnf(c_0_3504,negated_conjecture,(xuntil2p30(s27)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3334, c_0_3503])).
cnf(c_0_3505,negated_conjecture,(xuntil2p30(s28)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3337, c_0_3504]), c_0_2133])).
cnf(c_0_3506,negated_conjecture,(xuntil2p30(s29)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3339, c_0_3505]), c_0_2150])).
cnf(c_0_3507,negated_conjecture,(xuntil2p30(s30)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3288, c_0_3506])).
cnf(c_0_3508,negated_conjecture,(xuntil2p30(s31)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3291, c_0_3507])).
cnf(c_0_3509,negated_conjecture,(xuntil2p30(s32)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3294, c_0_3508]), c_0_2173])).
cnf(c_0_3510,negated_conjecture,(xuntil2p30(s33)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3296, c_0_3509]), c_0_2182])).
cnf(c_0_3511,negated_conjecture,(xuntil2p30(s34)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3249, c_0_3510])).
cnf(c_0_3512,negated_conjecture,(xuntil2p30(s35)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3252, c_0_3511])).
cnf(c_0_3513,negated_conjecture,(xuntil2p30(s36)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3255, c_0_3512]), c_0_2205])).
cnf(c_0_3514,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3257, c_0_3513]), c_0_2217])).
cnf(c_0_3515,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3214, c_0_3514])).
cnf(c_0_3516,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3217, c_0_3515])).
cnf(c_0_3517,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3516]), c_0_2251])).
cnf(c_0_3518,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3517]), c_0_2262])).
cnf(c_0_3519,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3518])).
cnf(c_0_3520,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3519])).
cnf(c_0_3521,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3520]), c_0_2289])).
cnf(c_0_3522,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3521]), c_0_2302])).
cnf(c_0_3523,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3522])).
cnf(c_0_3524,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3523])).
cnf(c_0_3525,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3524]), c_0_2337])).
cnf(c_0_3526,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3525]), c_0_2345])).
cnf(c_0_3527,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3526])).
cnf(c_0_3528,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3527])).
cnf(c_0_3529,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3528]), c_0_2370])).
cnf(c_0_3530,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3529]), c_0_2377])).
cnf(c_0_3531,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3530])).
cnf(c_0_3532,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3531])).
cnf(c_0_3533,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3532]), c_0_2409])).
cnf(c_0_3534,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3533]), c_0_2419])).
cnf(c_0_3535,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s59)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3534])).
cnf(c_0_3536,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3535])).
cnf(c_0_3537,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3536]), c_0_2452])).
cnf(c_0_3538,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s62)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3537]), c_0_2461])).
cnf(c_0_3539,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s63)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3538])).
cnf(c_0_3540,negated_conjecture,(until2p29(s10)|~xuntil2p30(s9)), inference(spm,[status(thm)],[c_0_140, c_0_71])).
cnf(c_0_3541,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3539])).
cnf(c_0_3542,negated_conjecture,(until2p29(s11)|~xuntil2p30(s10)), inference(spm,[status(thm)],[c_0_140, c_0_72])).
cnf(c_0_3543,negated_conjecture,(xuntil2p30(s10)|node26(s10)|~xuntil2p30(s9)), inference(spm,[status(thm)],[c_0_139, c_0_3540])).
cnf(c_0_3544,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|xuntil2p30(s9)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3541]), c_0_2478])).
cnf(c_0_3545,negated_conjecture,(until2p29(s12)|~xuntil2p30(s11)), inference(spm,[status(thm)],[c_0_140, c_0_73])).
cnf(c_0_3546,negated_conjecture,(xuntil2p30(s11)|node26(s11)|~xuntil2p30(s10)), inference(spm,[status(thm)],[c_0_139, c_0_3542])).
cnf(c_0_3547,negated_conjecture,(xuntil2p30(s10)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3543, c_0_3544])).
cnf(c_0_3548,negated_conjecture,(until2p29(s13)|~xuntil2p30(s12)), inference(spm,[status(thm)],[c_0_140, c_0_74])).
cnf(c_0_3549,negated_conjecture,(xuntil2p30(s12)|node26(s12)|~xuntil2p30(s11)), inference(spm,[status(thm)],[c_0_139, c_0_3545])).
cnf(c_0_3550,negated_conjecture,(xuntil2p30(s11)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3546, c_0_3547])).
cnf(c_0_3551,negated_conjecture,(xuntil2p30(s13)|node26(s13)|~xuntil2p30(s12)), inference(spm,[status(thm)],[c_0_139, c_0_3548])).
cnf(c_0_3552,negated_conjecture,(xuntil2p30(s12)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3549, c_0_3550]), c_0_1976])).
cnf(c_0_3553,negated_conjecture,(xuntil2p30(s13)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3551, c_0_3552]), c_0_1985])).
cnf(c_0_3554,negated_conjecture,(xuntil2p30(s14)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3484, c_0_3553])).
cnf(c_0_3555,negated_conjecture,(xuntil2p30(s15)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3487, c_0_3554])).
cnf(c_0_3556,negated_conjecture,(xuntil2p30(s16)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3490, c_0_3555]), c_0_2009])).
cnf(c_0_3557,negated_conjecture,(xuntil2p30(s17)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3492, c_0_3556]), c_0_2016])).
cnf(c_0_3558,negated_conjecture,(xuntil2p30(s18)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3429, c_0_3557])).
cnf(c_0_3559,negated_conjecture,(xuntil2p30(s19)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3432, c_0_3558])).
cnf(c_0_3560,negated_conjecture,(xuntil2p30(s20)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3435, c_0_3559]), c_0_2050])).
cnf(c_0_3561,negated_conjecture,(xuntil2p30(s21)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3437, c_0_3560]), c_0_2061])).
cnf(c_0_3562,negated_conjecture,(xuntil2p30(s22)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3378, c_0_3561])).
cnf(c_0_3563,negated_conjecture,(xuntil2p30(s23)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3381, c_0_3562])).
cnf(c_0_3564,negated_conjecture,(xuntil2p30(s24)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3384, c_0_3563]), c_0_2091])).
cnf(c_0_3565,negated_conjecture,(xuntil2p30(s25)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3386, c_0_3564]), c_0_2101])).
cnf(c_0_3566,negated_conjecture,(xuntil2p30(s26)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3331, c_0_3565])).
cnf(c_0_3567,negated_conjecture,(xuntil2p30(s27)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3334, c_0_3566])).
cnf(c_0_3568,negated_conjecture,(xuntil2p30(s28)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3337, c_0_3567]), c_0_2133])).
cnf(c_0_3569,negated_conjecture,(xuntil2p30(s29)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3339, c_0_3568]), c_0_2150])).
cnf(c_0_3570,negated_conjecture,(xuntil2p30(s30)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3288, c_0_3569])).
cnf(c_0_3571,negated_conjecture,(xuntil2p30(s31)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3291, c_0_3570])).
cnf(c_0_3572,negated_conjecture,(xuntil2p30(s32)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3294, c_0_3571]), c_0_2173])).
cnf(c_0_3573,negated_conjecture,(xuntil2p30(s33)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3296, c_0_3572]), c_0_2182])).
cnf(c_0_3574,negated_conjecture,(xuntil2p30(s34)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3249, c_0_3573])).
cnf(c_0_3575,negated_conjecture,(xuntil2p30(s35)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3252, c_0_3574])).
cnf(c_0_3576,negated_conjecture,(xuntil2p30(s36)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3255, c_0_3575]), c_0_2205])).
cnf(c_0_3577,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3257, c_0_3576]), c_0_2217])).
cnf(c_0_3578,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3214, c_0_3577])).
cnf(c_0_3579,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3217, c_0_3578])).
cnf(c_0_3580,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3579]), c_0_2251])).
cnf(c_0_3581,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3580]), c_0_2262])).
cnf(c_0_3582,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3581])).
cnf(c_0_3583,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3582])).
cnf(c_0_3584,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3583]), c_0_2289])).
cnf(c_0_3585,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3584]), c_0_2302])).
cnf(c_0_3586,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3585])).
cnf(c_0_3587,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3586])).
cnf(c_0_3588,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3587]), c_0_2337])).
cnf(c_0_3589,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3588]), c_0_2345])).
cnf(c_0_3590,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3589])).
cnf(c_0_3591,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3590])).
cnf(c_0_3592,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3591]), c_0_2370])).
cnf(c_0_3593,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3592]), c_0_2377])).
cnf(c_0_3594,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3593])).
cnf(c_0_3595,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3594])).
cnf(c_0_3596,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3595]), c_0_2409])).
cnf(c_0_3597,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3596]), c_0_2419])).
cnf(c_0_3598,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s59)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3597])).
cnf(c_0_3599,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3598])).
cnf(c_0_3600,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3599]), c_0_2452])).
cnf(c_0_3601,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s62)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3600]), c_0_2461])).
cnf(c_0_3602,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s63)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3601])).
cnf(c_0_3603,negated_conjecture,(until2p29(s6)|~xuntil2p30(s5)), inference(spm,[status(thm)],[c_0_140, c_0_62])).
cnf(c_0_3604,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3602])).
cnf(c_0_3605,negated_conjecture,(until2p29(s7)|~xuntil2p30(s6)), inference(spm,[status(thm)],[c_0_140, c_0_67])).
cnf(c_0_3606,negated_conjecture,(xuntil2p30(s6)|node26(s6)|~xuntil2p30(s5)), inference(spm,[status(thm)],[c_0_139, c_0_3603])).
cnf(c_0_3607,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|xuntil2p30(s5)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3604]), c_0_2478])).
cnf(c_0_3608,negated_conjecture,(until2p29(s8)|~xuntil2p30(s7)), inference(spm,[status(thm)],[c_0_140, c_0_69])).
cnf(c_0_3609,negated_conjecture,(xuntil2p30(s7)|node26(s7)|~xuntil2p30(s6)), inference(spm,[status(thm)],[c_0_139, c_0_3605])).
cnf(c_0_3610,negated_conjecture,(xuntil2p30(s6)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3606, c_0_3607])).
cnf(c_0_3611,negated_conjecture,(until2p29(s9)|~xuntil2p30(s8)), inference(spm,[status(thm)],[c_0_140, c_0_70])).
cnf(c_0_3612,negated_conjecture,(xuntil2p30(s8)|node26(s8)|~xuntil2p30(s7)), inference(spm,[status(thm)],[c_0_139, c_0_3608])).
cnf(c_0_3613,negated_conjecture,(xuntil2p30(s7)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3609, c_0_3610])).
cnf(c_0_3614,negated_conjecture,(xuntil2p30(s9)|node26(s9)|~xuntil2p30(s8)), inference(spm,[status(thm)],[c_0_139, c_0_3611])).
cnf(c_0_3615,negated_conjecture,(xuntil2p30(s8)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3612, c_0_3613]), c_0_1931])).
cnf(c_0_3616,negated_conjecture,(xuntil2p30(s9)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3614, c_0_3615]), c_0_1955])).
cnf(c_0_3617,negated_conjecture,(xuntil2p30(s10)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3543, c_0_3616])).
cnf(c_0_3618,negated_conjecture,(xuntil2p30(s11)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3546, c_0_3617])).
cnf(c_0_3619,negated_conjecture,(xuntil2p30(s12)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3549, c_0_3618]), c_0_1976])).
cnf(c_0_3620,negated_conjecture,(xuntil2p30(s13)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3551, c_0_3619]), c_0_1985])).
cnf(c_0_3621,negated_conjecture,(xuntil2p30(s14)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3484, c_0_3620])).
cnf(c_0_3622,negated_conjecture,(xuntil2p30(s15)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3487, c_0_3621])).
cnf(c_0_3623,negated_conjecture,(xuntil2p30(s16)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3490, c_0_3622]), c_0_2009])).
cnf(c_0_3624,negated_conjecture,(xuntil2p30(s17)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3492, c_0_3623]), c_0_2016])).
cnf(c_0_3625,negated_conjecture,(xuntil2p30(s18)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3429, c_0_3624])).
cnf(c_0_3626,negated_conjecture,(xuntil2p30(s19)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3432, c_0_3625])).
cnf(c_0_3627,negated_conjecture,(xuntil2p30(s20)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3435, c_0_3626]), c_0_2050])).
cnf(c_0_3628,negated_conjecture,(xuntil2p30(s21)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3437, c_0_3627]), c_0_2061])).
cnf(c_0_3629,negated_conjecture,(xuntil2p30(s22)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3378, c_0_3628])).
cnf(c_0_3630,negated_conjecture,(xuntil2p30(s23)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3381, c_0_3629])).
cnf(c_0_3631,negated_conjecture,(xuntil2p30(s24)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3384, c_0_3630]), c_0_2091])).
cnf(c_0_3632,negated_conjecture,(xuntil2p30(s25)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3386, c_0_3631]), c_0_2101])).
cnf(c_0_3633,negated_conjecture,(xuntil2p30(s26)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3331, c_0_3632])).
cnf(c_0_3634,negated_conjecture,(xuntil2p30(s27)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3334, c_0_3633])).
cnf(c_0_3635,negated_conjecture,(xuntil2p30(s28)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3337, c_0_3634]), c_0_2133])).
cnf(c_0_3636,negated_conjecture,(xuntil2p30(s29)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3339, c_0_3635]), c_0_2150])).
cnf(c_0_3637,negated_conjecture,(xuntil2p30(s30)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3288, c_0_3636])).
cnf(c_0_3638,negated_conjecture,(xuntil2p30(s31)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3291, c_0_3637])).
cnf(c_0_3639,negated_conjecture,(xuntil2p30(s32)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3294, c_0_3638]), c_0_2173])).
cnf(c_0_3640,negated_conjecture,(xuntil2p30(s33)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3296, c_0_3639]), c_0_2182])).
cnf(c_0_3641,negated_conjecture,(xuntil2p30(s34)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3249, c_0_3640])).
cnf(c_0_3642,negated_conjecture,(xuntil2p30(s35)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3252, c_0_3641])).
cnf(c_0_3643,negated_conjecture,(xuntil2p30(s36)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3255, c_0_3642]), c_0_2205])).
cnf(c_0_3644,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3257, c_0_3643]), c_0_2217])).
cnf(c_0_3645,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3214, c_0_3644])).
cnf(c_0_3646,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3217, c_0_3645])).
cnf(c_0_3647,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3646]), c_0_2251])).
cnf(c_0_3648,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3647]), c_0_2262])).
cnf(c_0_3649,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3183, c_0_3648])).
cnf(c_0_3650,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3186, c_0_3649])).
cnf(c_0_3651,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3650]), c_0_2289])).
cnf(c_0_3652,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3651]), c_0_2302])).
cnf(c_0_3653,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3156, c_0_3652])).
cnf(c_0_3654,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3159, c_0_3653])).
cnf(c_0_3655,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3654]), c_0_2337])).
cnf(c_0_3656,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3655]), c_0_2345])).
cnf(c_0_3657,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3133, c_0_3656])).
cnf(c_0_3658,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3136, c_0_3657])).
cnf(c_0_3659,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3658]), c_0_2370])).
cnf(c_0_3660,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3659]), c_0_2377])).
cnf(c_0_3661,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3114, c_0_3660])).
cnf(c_0_3662,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3117, c_0_3661])).
cnf(c_0_3663,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3662]), c_0_2409])).
cnf(c_0_3664,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3663]), c_0_2419])).
cnf(c_0_3665,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s59)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3099, c_0_3664])).
cnf(c_0_3666,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3102, c_0_3665])).
cnf(c_0_3667,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3666]), c_0_2452])).
cnf(c_0_3668,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s62)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3667]), c_0_2461])).
cnf(c_0_3669,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s63)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3090, c_0_3668])).
cnf(c_0_3670,negated_conjecture,(until2p29(s1)|~xuntil2p30(s0)), inference(spm,[status(thm)],[c_0_140, c_0_9])).
cnf(c_0_3671,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|xuntil2p30(s0)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3669])).
cnf(c_0_3672,negated_conjecture,(until2p29(s2)|~xuntil2p30(s1)), inference(spm,[status(thm)],[c_0_140, c_0_34])).
cnf(c_0_3673,negated_conjecture,(xuntil2p30(s1)|node26(s1)|~xuntil2p30(s0)), inference(spm,[status(thm)],[c_0_139, c_0_3670])).
cnf(c_0_3674,negated_conjecture,(xuntil2p30(s64)|xuntil2p30(s0)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3671]), c_0_2478])).
cnf(c_0_3675,negated_conjecture,(until2p29(s3)|~xuntil2p30(s2)), inference(spm,[status(thm)],[c_0_140, c_0_42])).
cnf(c_0_3676,negated_conjecture,(xuntil2p30(s2)|node26(s2)|~xuntil2p30(s1)), inference(spm,[status(thm)],[c_0_139, c_0_3672])).
cnf(c_0_3677,negated_conjecture,(xuntil2p30(s1)|xuntil2p30(s64)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)), inference(spm,[status(thm)],[c_0_3673, c_0_3674])).
cnf(c_0_3678,negated_conjecture,(until2p29(s4)|~xuntil2p30(s3)), inference(spm,[status(thm)],[c_0_140, c_0_54])).
cnf(c_0_3679,negated_conjecture,(xuntil2p30(s3)|node26(s3)|~xuntil2p30(s2)), inference(spm,[status(thm)],[c_0_139, c_0_3675])).
cnf(c_0_3680,negated_conjecture,(xuntil2p30(s2)|xuntil2p30(s64)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)), inference(spm,[status(thm)],[c_0_3676, c_0_3677])).
cnf(c_0_3681,negated_conjecture,(until2p29(s5)|~xuntil2p30(s4)), inference(spm,[status(thm)],[c_0_140, c_0_61])).
cnf(c_0_3682,negated_conjecture,(xuntil2p30(s4)|node26(s4)|~xuntil2p30(s3)), inference(spm,[status(thm)],[c_0_139, c_0_3678])).
cnf(c_0_3683,negated_conjecture,(xuntil2p30(s3)|xuntil2p30(s64)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)), inference(spm,[status(thm)],[c_0_3679, c_0_3680])).
cnf(c_0_3684,negated_conjecture,(xuntil2p30(s5)|node26(s5)|~xuntil2p30(s4)), inference(spm,[status(thm)],[c_0_139, c_0_3681])).
cnf(c_0_3685,negated_conjecture,(xuntil2p30(s4)|xuntil2p30(s64)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3682, c_0_3683]), c_0_1889])).
cnf(c_0_3686,negated_conjecture,(xuntil2p30(s5)|xuntil2p30(s64)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3684, c_0_3685]), c_0_1907])).
cnf(c_0_3687,negated_conjecture,(xuntil2p30(s6)|xuntil2p30(s64)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)), inference(spm,[status(thm)],[c_0_3606, c_0_3686])).
cnf(c_0_3688,negated_conjecture,(xuntil2p30(s7)|xuntil2p30(s64)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)), inference(spm,[status(thm)],[c_0_3609, c_0_3687])).
cnf(c_0_3689,negated_conjecture,(xuntil2p30(s8)|xuntil2p30(s64)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3612, c_0_3688]), c_0_1931])).
cnf(c_0_3690,negated_conjecture,(xuntil2p30(s9)|xuntil2p30(s64)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3614, c_0_3689]), c_0_1955])).
cnf(c_0_3691,negated_conjecture,(xuntil2p30(s10)|xuntil2p30(s64)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)), inference(spm,[status(thm)],[c_0_3543, c_0_3690])).
cnf(c_0_3692,negated_conjecture,(xuntil2p30(s11)|xuntil2p30(s64)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)), inference(spm,[status(thm)],[c_0_3546, c_0_3691])).
cnf(c_0_3693,negated_conjecture,(xuntil2p30(s12)|xuntil2p30(s64)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3549, c_0_3692]), c_0_1976])).
cnf(c_0_3694,negated_conjecture,(xuntil2p30(s13)|xuntil2p30(s64)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3551, c_0_3693]), c_0_1985])).
cnf(c_0_3695,negated_conjecture,(xuntil2p30(s14)|xuntil2p30(s64)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)), inference(spm,[status(thm)],[c_0_3484, c_0_3694])).
cnf(c_0_3696,negated_conjecture,(xuntil2p30(s15)|xuntil2p30(s64)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)), inference(spm,[status(thm)],[c_0_3487, c_0_3695])).
cnf(c_0_3697,negated_conjecture,(xuntil2p30(s16)|xuntil2p30(s64)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3490, c_0_3696]), c_0_2009])).
cnf(c_0_3698,negated_conjecture,(xuntil2p30(s17)|xuntil2p30(s64)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3492, c_0_3697]), c_0_2016])).
cnf(c_0_3699,negated_conjecture,(xuntil2p30(s18)|xuntil2p30(s64)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)), inference(spm,[status(thm)],[c_0_3429, c_0_3698])).
cnf(c_0_3700,negated_conjecture,(xuntil2p30(s19)|xuntil2p30(s64)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)), inference(spm,[status(thm)],[c_0_3432, c_0_3699])).
cnf(c_0_3701,negated_conjecture,(xuntil2p30(s20)|xuntil2p30(s64)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3435, c_0_3700]), c_0_2050])).
cnf(c_0_3702,negated_conjecture,(xuntil2p30(s21)|xuntil2p30(s64)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3437, c_0_3701]), c_0_2061])).
cnf(c_0_3703,negated_conjecture,(xuntil2p30(s22)|xuntil2p30(s64)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)), inference(spm,[status(thm)],[c_0_3378, c_0_3702])).
cnf(c_0_3704,negated_conjecture,(xuntil2p30(s23)|xuntil2p30(s64)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)), inference(spm,[status(thm)],[c_0_3381, c_0_3703])).
cnf(c_0_3705,negated_conjecture,(xuntil2p30(s24)|xuntil2p30(s64)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3384, c_0_3704]), c_0_2091])).
cnf(c_0_3706,negated_conjecture,(xuntil2p30(s25)|xuntil2p30(s64)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3386, c_0_3705]), c_0_2101])).
cnf(c_0_3707,negated_conjecture,(xuntil2p30(s26)|xuntil2p30(s64)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)), inference(spm,[status(thm)],[c_0_3331, c_0_3706])).
cnf(c_0_3708,negated_conjecture,(xuntil2p30(s27)|xuntil2p30(s64)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)), inference(spm,[status(thm)],[c_0_3334, c_0_3707])).
cnf(c_0_3709,negated_conjecture,(xuntil2p30(s28)|xuntil2p30(s64)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3337, c_0_3708]), c_0_2133])).
cnf(c_0_3710,negated_conjecture,(xuntil2p30(s29)|xuntil2p30(s64)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3339, c_0_3709]), c_0_2150])).
cnf(c_0_3711,negated_conjecture,(xuntil2p30(s30)|xuntil2p30(s64)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)), inference(spm,[status(thm)],[c_0_3288, c_0_3710])).
cnf(c_0_3712,negated_conjecture,(xuntil2p30(s31)|xuntil2p30(s64)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)), inference(spm,[status(thm)],[c_0_3291, c_0_3711])).
cnf(c_0_3713,negated_conjecture,(xuntil2p30(s32)|xuntil2p30(s64)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3294, c_0_3712]), c_0_2173])).
cnf(c_0_3714,negated_conjecture,(xuntil2p30(s33)|xuntil2p30(s64)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3296, c_0_3713]), c_0_2182])).
cnf(c_0_3715,negated_conjecture,(xuntil2p30(s34)|xuntil2p30(s64)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)), inference(spm,[status(thm)],[c_0_3249, c_0_3714])).
cnf(c_0_3716,negated_conjecture,(xuntil2p30(s35)|xuntil2p30(s64)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)), inference(spm,[status(thm)],[c_0_3252, c_0_3715])).
cnf(c_0_3717,negated_conjecture,(xuntil2p30(s36)|xuntil2p30(s64)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3255, c_0_3716]), c_0_2205])).
cnf(c_0_3718,negated_conjecture,(xuntil2p30(s37)|xuntil2p30(s64)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3257, c_0_3717]), c_0_2217])).
cnf(c_0_3719,negated_conjecture,(xuntil2p30(s38)|xuntil2p30(s64)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)), inference(spm,[status(thm)],[c_0_3214, c_0_3718])).
cnf(c_0_3720,negated_conjecture,(xuntil2p30(s39)|xuntil2p30(s64)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)), inference(spm,[status(thm)],[c_0_3217, c_0_3719])).
cnf(c_0_3721,negated_conjecture,(xuntil2p30(s40)|xuntil2p30(s64)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3220, c_0_3720]), c_0_2251])).
cnf(c_0_3722,negated_conjecture,(xuntil2p30(s41)|xuntil2p30(s64)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3222, c_0_3721]), c_0_2262])).
cnf(c_0_3723,negated_conjecture,(xuntil2p30(s42)|xuntil2p30(s64)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)), inference(spm,[status(thm)],[c_0_3183, c_0_3722])).
cnf(c_0_3724,negated_conjecture,(xuntil2p30(s43)|xuntil2p30(s64)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)), inference(spm,[status(thm)],[c_0_3186, c_0_3723])).
cnf(c_0_3725,negated_conjecture,(xuntil2p30(s44)|xuntil2p30(s64)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3189, c_0_3724]), c_0_2289])).
cnf(c_0_3726,negated_conjecture,(xuntil2p30(s45)|xuntil2p30(s64)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3191, c_0_3725]), c_0_2302])).
cnf(c_0_3727,negated_conjecture,(xuntil2p30(s46)|xuntil2p30(s64)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)), inference(spm,[status(thm)],[c_0_3156, c_0_3726])).
cnf(c_0_3728,negated_conjecture,(xuntil2p30(s47)|xuntil2p30(s64)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)), inference(spm,[status(thm)],[c_0_3159, c_0_3727])).
cnf(c_0_3729,negated_conjecture,(xuntil2p30(s48)|xuntil2p30(s64)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3162, c_0_3728]), c_0_2337])).
cnf(c_0_3730,negated_conjecture,(xuntil2p30(s49)|xuntil2p30(s64)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3164, c_0_3729]), c_0_2345])).
cnf(c_0_3731,negated_conjecture,(xuntil2p30(s50)|xuntil2p30(s64)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)), inference(spm,[status(thm)],[c_0_3133, c_0_3730])).
cnf(c_0_3732,negated_conjecture,(xuntil2p30(s51)|xuntil2p30(s64)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)), inference(spm,[status(thm)],[c_0_3136, c_0_3731])).
cnf(c_0_3733,negated_conjecture,(xuntil2p30(s52)|xuntil2p30(s64)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3139, c_0_3732]), c_0_2370])).
cnf(c_0_3734,negated_conjecture,(xuntil2p30(s53)|xuntil2p30(s64)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3141, c_0_3733]), c_0_2377])).
cnf(c_0_3735,negated_conjecture,(xuntil2p30(s54)|xuntil2p30(s64)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)), inference(spm,[status(thm)],[c_0_3114, c_0_3734])).
cnf(c_0_3736,negated_conjecture,(xuntil2p30(s55)|xuntil2p30(s64)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)), inference(spm,[status(thm)],[c_0_3117, c_0_3735])).
cnf(c_0_3737,negated_conjecture,(xuntil2p30(s56)|xuntil2p30(s64)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3120, c_0_3736]), c_0_2409])).
cnf(c_0_3738,negated_conjecture,(xuntil2p30(s57)|xuntil2p30(s64)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)|node26(s58)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3122, c_0_3737]), c_0_2419])).
cnf(c_0_3739,negated_conjecture,(xuntil2p30(s58)|xuntil2p30(s64)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)|node26(s59)), inference(spm,[status(thm)],[c_0_3099, c_0_3738])).
cnf(c_0_3740,negated_conjecture,(xuntil2p30(s59)|xuntil2p30(s64)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)), inference(spm,[status(thm)],[c_0_3102, c_0_3739])).
cnf(c_0_3741,negated_conjecture,(xuntil2p30(s60)|xuntil2p30(s64)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3105, c_0_3740]), c_0_2452])).
cnf(c_0_3742,negated_conjecture,(xuntil2p30(s61)|xuntil2p30(s64)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)|node26(s62)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3107, c_0_3741]), c_0_2461])).
cnf(c_0_3743,negated_conjecture,(xuntil2p30(s62)|xuntil2p30(s64)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|node26(s63)), inference(spm,[status(thm)],[c_0_3090, c_0_3742])).
cnf(c_0_3744,negated_conjecture,(xuntil2p30(s63)|xuntil2p30(s64)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(spm,[status(thm)],[c_0_3093, c_0_3743])).
cnf(c_0_3745,negated_conjecture,(xuntil2p30(s64)|node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_3096, c_0_3744]), c_0_2478])).
cnf(c_0_3746,negated_conjecture,(node26(s63)|node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_3745]), c_0_131])])).
cnf(c_0_3747,negated_conjecture,(node26(s62)|node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s63)), inference(spm,[status(thm)],[c_0_132, c_0_3746])).
cnf(c_0_3748,negated_conjecture,(node26(s59)|node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(spm,[status(thm)],[c_0_132, c_0_3747])).
cnf(c_0_3749,negated_conjecture,(node26(s58)|node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3748]), c_0_2485])).
cnf(c_0_3750,negated_conjecture,(node26(s55)|node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3749]), c_0_2488])).
cnf(c_0_3751,negated_conjecture,(node26(s54)|node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3750]), c_0_2492])).
cnf(c_0_3752,negated_conjecture,(node26(s51)|node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3751]), c_0_2495])).
cnf(c_0_3753,negated_conjecture,(node26(s50)|node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3752]), c_0_2499])).
cnf(c_0_3754,negated_conjecture,(node26(s47)|node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3753]), c_0_2502])).
cnf(c_0_3755,negated_conjecture,(node26(s46)|node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3754]), c_0_2506])).
cnf(c_0_3756,negated_conjecture,(node26(s43)|node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3755]), c_0_2509])).
cnf(c_0_3757,negated_conjecture,(node26(s42)|node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3756]), c_0_2513])).
cnf(c_0_3758,negated_conjecture,(node26(s39)|node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3757]), c_0_2516])).
cnf(c_0_3759,negated_conjecture,(node26(s38)|node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3758]), c_0_2520])).
cnf(c_0_3760,negated_conjecture,(node26(s35)|node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3759]), c_0_2523])).
cnf(c_0_3761,negated_conjecture,(node26(s34)|node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3760]), c_0_2527])).
cnf(c_0_3762,negated_conjecture,(node26(s31)|node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3761]), c_0_2530])).
cnf(c_0_3763,negated_conjecture,(node26(s30)|node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3762]), c_0_2534])).
cnf(c_0_3764,negated_conjecture,(node26(s27)|node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3763]), c_0_2540])).
cnf(c_0_3765,negated_conjecture,(node26(s26)|node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3764]), c_0_2545])).
cnf(c_0_3766,negated_conjecture,(node26(s23)|node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3765]), c_0_2548])).
cnf(c_0_3767,negated_conjecture,(node26(s22)|node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3766]), c_0_2552])).
cnf(c_0_3768,negated_conjecture,(node26(s19)|node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3767]), c_0_2555])).
cnf(c_0_3769,negated_conjecture,(node26(s18)|node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3768]), c_0_2562])).
cnf(c_0_3770,negated_conjecture,(node26(s15)|node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3769]), c_0_2572])).
cnf(c_0_3771,negated_conjecture,(node26(s14)|node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3770]), c_0_2587])).
cnf(c_0_3772,negated_conjecture,(node26(s11)|node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3771]), c_0_2616])).
cnf(c_0_3773,negated_conjecture,(node26(s10)|node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3772]), c_0_2637])).
cnf(c_0_3774,negated_conjecture,(node26(s7)|node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3773]), c_0_2674])).
cnf(c_0_3775,negated_conjecture,(node26(s6)|node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3774]), c_0_2703])).
cnf(c_0_3776,negated_conjecture,(node26(s3)|node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3775]), c_0_2753])).
cnf(c_0_3777,negated_conjecture,(node26(s2)|node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3776]), c_0_2777])).
cnf(c_0_3778,negated_conjecture,(node26(s1)|always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3777]), c_0_2800])).
cnf(c_0_3779,negated_conjecture,(always25(s62)|always25(s63)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_3778]), c_0_2832])).
cnf(c_0_3780,negated_conjecture,(always25(s62)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2860, c_0_3779]), c_0_2863])).
cnf(c_0_3781,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2919, c_0_3780])]), c_0_2863]), ['proof']).
# SZS output end CNFRefutation.
EOF
