% Problem    : PUZ017-1 : TPTP v6.1.0. Released v1.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n105.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:16:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C02_F1_PI_AE_Q12_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,negated_conjecture,(~hasperson(n1,X1)|~hasperson(n2,X2)|~hasperson(n3,X3)|~hasperson(n4,X4)|~hasperson(n5,X5)|~hascolor(n1,X6)|~hascolor(n2,X7)|~hascolor(n3,X8)|~hascolor(n4,X9)|~hascolor(n5,X10)|~hasdrink(n1,X11)|~hasdrink(n2,X12)|~hasdrink(n3,X13)|~hasdrink(n4,X14)|~hasdrink(n5,X15)|~hasgame(n1,X16)|~hasgame(n2,X17)|~hasgame(n3,X18)|~hasgame(n4,X19)|~hasgame(n5,X20)|~haspet(n1,X21)|~haspet(n2,X22)|~haspet(n3,X23)|~haspet(n4,X24)|~haspet(n5,X25)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', find_out_house_details)).
cnf(c_0_1,axiom,(haspet(X1,guppy)|haspet(X1,toad)|haspet(X1,camel)|haspet(X1,rat)|haspet(X1,no_pet)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', every_house_has_a_pet)).
cnf(c_0_2,hypothesis,(hascolor(X2,blue)|~hasperson(X1,russian)|~samehouse(n1,X1)|~nextto(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', c20)).
cnf(c_0_3,hypothesis,(hasperson(n1,russian)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', house_1_has_russian)).
cnf(c_0_4,axiom,(samehouse(X1,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', reflexivity_for_samehouse)).
cnf(c_0_5,axiom,(nextto(X1,X2)|~left(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', left_means_nextto)).
cnf(c_0_6,hypothesis,(left(n1,n2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', house_1_left_of_2)).
cnf(c_0_7,hypothesis,(hasdrink(n3,milk)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', house_3_has_milk)).
cnf(c_0_8,axiom,(hasgame(X1,backgammon)|hasgame(X1,racquetball)|hasgame(X1,quoits)|hasgame(X1,solitaire)|hasgame(X1,charades)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', every_house_has_a_game)).
cnf(c_0_9,axiom,(hasdrink(X1,lemonade)|hasdrink(X1,coffee)|hasdrink(X1,milk)|hasdrink(X1,vodka)|hasdrink(X1,unknown_drink)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', every_house_has_a_drink)).
cnf(c_0_10,axiom,(hascolor(X1,red)|hascolor(X1,white)|hascolor(X1,green)|hascolor(X1,yellow)|hascolor(X1,blue)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', every_house_has_color)).
cnf(c_0_11,axiom,(hasperson(X1,englishman)|hasperson(X1,italian)|hasperson(X1,swede)|hasperson(X1,russian)|hasperson(X1,american)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', every_house_has_a_national)).
fof(c_0_12, plain, ((~epred1_0<=>![X21]:~haspet(n1,X21))), introduced(definition)).
fof(c_0_13, plain, ((~epred2_0<=>![X22]:~haspet(n2,X22))), introduced(definition)).
fof(c_0_14, plain, ((~epred3_0<=>![X23]:~haspet(n3,X23))), introduced(definition)).
fof(c_0_15, plain, ((~epred4_0<=>![X24]:~haspet(n4,X24))), introduced(definition)).
fof(c_0_16, plain, ((~epred5_0<=>![X25]:~haspet(n5,X25))), introduced(definition)).
fof(c_0_17, plain, ((~epred6_0<=>![X16]:~hasgame(n1,X16))), introduced(definition)).
fof(c_0_18, plain, ((~epred7_0<=>![X17]:~hasgame(n2,X17))), introduced(definition)).
fof(c_0_19, plain, ((~epred8_0<=>![X18]:~hasgame(n3,X18))), introduced(definition)).
fof(c_0_20, plain, ((~epred9_0<=>![X19]:~hasgame(n4,X19))), introduced(definition)).
fof(c_0_21, plain, ((~epred10_0<=>![X20]:~hasgame(n5,X20))), introduced(definition)).
fof(c_0_22, plain, ((~epred11_0<=>![X11]:~hasdrink(n1,X11))), introduced(definition)).
fof(c_0_23, plain, ((~epred12_0<=>![X12]:~hasdrink(n2,X12))), introduced(definition)).
fof(c_0_24, plain, ((~epred13_0<=>![X13]:~hasdrink(n3,X13))), introduced(definition)).
fof(c_0_25, plain, ((~epred14_0<=>![X14]:~hasdrink(n4,X14))), introduced(definition)).
fof(c_0_26, plain, ((~epred15_0<=>![X15]:~hasdrink(n5,X15))), introduced(definition)).
fof(c_0_27, plain, ((~epred16_0<=>![X6]:~hascolor(n1,X6))), introduced(definition)).
fof(c_0_28, plain, ((~epred17_0<=>![X7]:~hascolor(n2,X7))), introduced(definition)).
fof(c_0_29, plain, ((~epred18_0<=>![X8]:~hascolor(n3,X8))), introduced(definition)).
fof(c_0_30, plain, ((~epred19_0<=>![X9]:~hascolor(n4,X9))), introduced(definition)).
fof(c_0_31, plain, ((~epred20_0<=>![X10]:~hascolor(n5,X10))), introduced(definition)).
fof(c_0_32, plain, ((~epred21_0<=>![X1]:~hasperson(n1,X1))), introduced(definition)).
fof(c_0_33, plain, ((~epred22_0<=>![X2]:~hasperson(n2,X2))), introduced(definition)).
fof(c_0_34, plain, ((~epred23_0<=>![X3]:~hasperson(n3,X3))), introduced(definition)).
fof(c_0_35, plain, ((~epred24_0<=>![X4]:~hasperson(n4,X4))), introduced(definition)).
fof(c_0_36, plain, ((~epred25_0<=>![X5]:~hasperson(n5,X5))), introduced(definition)).
cnf(c_0_37,negated_conjecture,(epred1_0|~haspet(n1,X1)), inference(split_equiv,[status(thm)],[c_0_12])).
cnf(c_0_38,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred21_0|~epred20_0|~epred19_0|~epred18_0|~epred17_0|~epred16_0|~epred15_0|~epred14_0|~epred13_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0|~epred8_0|~epred7_0|~epred6_0|~epred5_0|~epred4_0|~epred3_0|~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_0, c_0_12]), c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_39,negated_conjecture,(epred1_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1]), c_0_37]), c_0_37]), c_0_37]), c_0_37])).
cnf(c_0_40,negated_conjecture,(epred2_0|~haspet(n2,X1)), inference(split_equiv,[status(thm)],[c_0_13])).
cnf(c_0_41,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred21_0|~epred20_0|~epred19_0|~epred18_0|~epred17_0|~epred16_0|~epred15_0|~epred14_0|~epred13_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0|~epred8_0|~epred7_0|~epred6_0|~epred5_0|~epred4_0|~epred3_0|~epred2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39])])).
cnf(c_0_42,negated_conjecture,(epred2_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1]), c_0_40]), c_0_40]), c_0_40]), c_0_40])).
cnf(c_0_43,negated_conjecture,(epred3_0|~haspet(n3,X1)), inference(split_equiv,[status(thm)],[c_0_14])).
cnf(c_0_44,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred21_0|~epred20_0|~epred19_0|~epred18_0|~epred17_0|~epred16_0|~epred15_0|~epred14_0|~epred13_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0|~epred8_0|~epred7_0|~epred6_0|~epred5_0|~epred4_0|~epred3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42])])).
cnf(c_0_45,negated_conjecture,(epred3_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_1]), c_0_43]), c_0_43]), c_0_43]), c_0_43])).
cnf(c_0_46,negated_conjecture,(epred4_0|~haspet(n4,X1)), inference(split_equiv,[status(thm)],[c_0_15])).
cnf(c_0_47,hypothesis,(hascolor(X1,blue)|~nextto(n1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_3]), c_0_4])])).
cnf(c_0_48,hypothesis,(nextto(n1,n2)), inference(spm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_49,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred21_0|~epred20_0|~epred19_0|~epred18_0|~epred17_0|~epred16_0|~epred15_0|~epred14_0|~epred13_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0|~epred8_0|~epred7_0|~epred6_0|~epred5_0|~epred4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45])])).
cnf(c_0_50,negated_conjecture,(epred4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_1]), c_0_46]), c_0_46]), c_0_46]), c_0_46])).
cnf(c_0_51,negated_conjecture,(epred13_0|~hasdrink(n3,X1)), inference(split_equiv,[status(thm)],[c_0_24])).
cnf(c_0_52,negated_conjecture,(epred21_0|~hasperson(n1,X1)), inference(split_equiv,[status(thm)],[c_0_32])).
cnf(c_0_53,negated_conjecture,(epred17_0|~hascolor(n2,X1)), inference(split_equiv,[status(thm)],[c_0_28])).
cnf(c_0_54,hypothesis,(hascolor(n2,blue)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_55,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred21_0|~epred20_0|~epred19_0|~epred18_0|~epred17_0|~epred16_0|~epred15_0|~epred14_0|~epred13_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0|~epred8_0|~epred7_0|~epred6_0|~epred5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_56,hypothesis,(epred13_0), inference(spm,[status(thm)],[c_0_51, c_0_7])).
cnf(c_0_57,hypothesis,(epred21_0), inference(spm,[status(thm)],[c_0_52, c_0_3])).
cnf(c_0_58,hypothesis,(epred17_0), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59,negated_conjecture,(epred5_0|~haspet(n5,X1)), inference(split_equiv,[status(thm)],[c_0_16])).
cnf(c_0_60,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred20_0|~epred19_0|~epred18_0|~epred16_0|~epred15_0|~epred14_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0|~epred8_0|~epred7_0|~epred6_0|~epred5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56])]), c_0_57]), c_0_58])])).
cnf(c_0_61,negated_conjecture,(epred5_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1]), c_0_59]), c_0_59]), c_0_59]), c_0_59])).
cnf(c_0_62,negated_conjecture,(epred6_0|~hasgame(n1,X1)), inference(split_equiv,[status(thm)],[c_0_17])).
cnf(c_0_63,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred20_0|~epred19_0|~epred18_0|~epred16_0|~epred15_0|~epred14_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0|~epred8_0|~epred7_0|~epred6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61])])).
cnf(c_0_64,negated_conjecture,(epred6_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_8]), c_0_62]), c_0_62]), c_0_62]), c_0_62])).
cnf(c_0_65,negated_conjecture,(epred7_0|~hasgame(n2,X1)), inference(split_equiv,[status(thm)],[c_0_18])).
cnf(c_0_66,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred20_0|~epred19_0|~epred18_0|~epred16_0|~epred15_0|~epred14_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0|~epred8_0|~epred7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_67,negated_conjecture,(epred7_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_8]), c_0_65]), c_0_65]), c_0_65]), c_0_65])).
cnf(c_0_68,negated_conjecture,(epred8_0|~hasgame(n3,X1)), inference(split_equiv,[status(thm)],[c_0_19])).
cnf(c_0_69,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred20_0|~epred19_0|~epred18_0|~epred16_0|~epred15_0|~epred14_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0|~epred8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67])])).
cnf(c_0_70,negated_conjecture,(epred8_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_8]), c_0_68]), c_0_68]), c_0_68]), c_0_68])).
cnf(c_0_71,negated_conjecture,(epred9_0|~hasgame(n4,X1)), inference(split_equiv,[status(thm)],[c_0_20])).
cnf(c_0_72,negated_conjecture,(epred15_0|~hasdrink(n5,X1)), inference(split_equiv,[status(thm)],[c_0_26])).
cnf(c_0_73,negated_conjecture,(epred14_0|~hasdrink(n4,X1)), inference(split_equiv,[status(thm)],[c_0_25])).
cnf(c_0_74,negated_conjecture,(epred12_0|~hasdrink(n2,X1)), inference(split_equiv,[status(thm)],[c_0_23])).
cnf(c_0_75,negated_conjecture,(epred11_0|~hasdrink(n1,X1)), inference(split_equiv,[status(thm)],[c_0_22])).
cnf(c_0_76,negated_conjecture,(epred10_0|~hasgame(n5,X1)), inference(split_equiv,[status(thm)],[c_0_21])).
cnf(c_0_77,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred20_0|~epred19_0|~epred18_0|~epred16_0|~epred15_0|~epred14_0|~epred12_0|~epred11_0|~epred10_0|~epred9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_70])])).
cnf(c_0_78,negated_conjecture,(epred9_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_8]), c_0_71]), c_0_71]), c_0_71]), c_0_71])).
cnf(c_0_79,negated_conjecture,(epred15_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_9]), c_0_72]), c_0_72]), c_0_72]), c_0_72])).
cnf(c_0_80,negated_conjecture,(epred14_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_9]), c_0_73]), c_0_73]), c_0_73]), c_0_73])).
cnf(c_0_81,negated_conjecture,(epred12_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_9]), c_0_74]), c_0_74]), c_0_74]), c_0_74])).
cnf(c_0_82,negated_conjecture,(epred11_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_9]), c_0_75]), c_0_75]), c_0_75]), c_0_75])).
cnf(c_0_83,negated_conjecture,(epred10_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_8]), c_0_76]), c_0_76]), c_0_76]), c_0_76])).
cnf(c_0_84,negated_conjecture,(epred16_0|~hascolor(n1,X1)), inference(split_equiv,[status(thm)],[c_0_27])).
cnf(c_0_85,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred20_0|~epred19_0|~epred18_0|~epred16_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_78])]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83])])).
cnf(c_0_86,negated_conjecture,(epred16_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_10]), c_0_84]), c_0_84]), c_0_84]), c_0_84])).
cnf(c_0_87,negated_conjecture,(epred18_0|~hascolor(n3,X1)), inference(split_equiv,[status(thm)],[c_0_29])).
cnf(c_0_88,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred20_0|~epred19_0|~epred18_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86])])).
cnf(c_0_89,negated_conjecture,(epred18_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_10]), c_0_87]), c_0_87]), c_0_87]), c_0_87])).
cnf(c_0_90,negated_conjecture,(epred19_0|~hascolor(n4,X1)), inference(split_equiv,[status(thm)],[c_0_30])).
cnf(c_0_91,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred20_0|~epred19_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_89])])).
cnf(c_0_92,negated_conjecture,(epred19_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_10]), c_0_90]), c_0_90]), c_0_90]), c_0_90])).
cnf(c_0_93,negated_conjecture,(epred20_0|~hascolor(n5,X1)), inference(split_equiv,[status(thm)],[c_0_31])).
cnf(c_0_94,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0|~epred20_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_92])])).
cnf(c_0_95,negated_conjecture,(epred20_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_10]), c_0_93]), c_0_93]), c_0_93]), c_0_93])).
cnf(c_0_96,negated_conjecture,(epred22_0|~hasperson(n2,X1)), inference(split_equiv,[status(thm)],[c_0_33])).
cnf(c_0_97,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0|~epred22_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_98,negated_conjecture,(epred22_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_11]), c_0_96]), c_0_96]), c_0_96]), c_0_96])).
cnf(c_0_99,negated_conjecture,(epred23_0|~hasperson(n3,X1)), inference(split_equiv,[status(thm)],[c_0_34])).
cnf(c_0_100,negated_conjecture,(~epred25_0|~epred24_0|~epred23_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_98])])).
cnf(c_0_101,negated_conjecture,(epred23_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_11]), c_0_99]), c_0_99]), c_0_99]), c_0_99])).
cnf(c_0_102,negated_conjecture,(epred24_0|~hasperson(n4,X1)), inference(split_equiv,[status(thm)],[c_0_35])).
cnf(c_0_103,negated_conjecture,(~epred25_0|~epred24_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_101])])).
cnf(c_0_104,negated_conjecture,(epred24_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_11]), c_0_102]), c_0_102]), c_0_102]), c_0_102])).
cnf(c_0_105,negated_conjecture,(epred25_0|~hasperson(n5,X1)), inference(split_equiv,[status(thm)],[c_0_36])).
cnf(c_0_106,negated_conjecture,(~epred25_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_107,negated_conjecture,(hasperson(n5,englishman)|hasperson(n5,italian)|hasperson(n5,swede)|hasperson(n5,russian)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_11]), c_0_106])).
cnf(c_0_108,negated_conjecture,(hasperson(n5,russian)|hasperson(n5,italian)|hasperson(n5,englishman)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_107]), c_0_106])).
cnf(c_0_109,negated_conjecture,(hasperson(n5,englishman)|hasperson(n5,russian)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_108]), c_0_106])).
cnf(c_0_110,negated_conjecture,(hasperson(n5,russian)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_109]), c_0_106])).
cnf(c_0_111,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_110]), c_0_106]), ['proof']).
# SZS output end CNFRefutation.
EOF
