cnf(c_0_1,axiom,(add(X1,X2)=add(X2,X1)), file('Axioms/BOO004-0.ax', commutativity_of_add)).
cnf(c_0_2,axiom,(multiply(X1,X2)=multiply(X2,X1)), file('Axioms/BOO004-0.ax', commutativity_of_multiply)).
cnf(c_0_3,axiom,(add(X1,multiply(X2,X3))=multiply(add(X1,X2),add(X1,X3))), file('Axioms/BOO004-0.ax', distributivity1)).
cnf(c_0_4,axiom,(multiply(X1,add(X2,X3))=add(multiply(X1,X2),multiply(X1,X3))), file('Axioms/BOO004-0.ax', distributivity2)).
cnf(c_0_5,axiom,(add(X1,additive_identity)=X1), file('Axioms/BOO004-0.ax', additive_id1)).
cnf(c_0_6,axiom,(multiply(X1,multiplicative_identity)=X1), file('Axioms/BOO004-0.ax', multiplicative_id1)).
cnf(c_0_7,axiom,(add(X1,inverse(X1))=multiplicative_identity), file('Axioms/BOO004-0.ax', additive_inverse1)).
cnf(c_0_8,axiom,(multiply(X1,inverse(X1))=additive_identity), file('Axioms/BOO004-0.ax', multiplicative_inverse1)).
cnf(c_0_9,negated_conjecture,(multiply(a,a)!=a), file('<stdin>', prove_a_times_a_is_a)).
# Initializing proof state
# Scanning for AC axioms
# add is commutative
# multiply is commutative
# AC handling enabled
cnf(c_0_10,plain,(add(additive_identity,X1)=X1),inference(pm,[status(thm)],[c_0_5,c_0_1,theory(equality)])).
cnf(c_0_11,plain,(add(additive_identity,X1)=X1),inference(pm,[status(thm)],[c_0_5,c_0_1,theory(equality)])).
cnf(c_0_12,plain,(inverse(additive_identity)=multiplicative_identity),inference(pm,[status(thm)],[c_0_7,c_0_10,theory(equality)])).
cnf(c_0_13,plain,(X1=add(X1,additive_identity)),inference(pm,[status(thm)],[c_0_1,c_0_10,theory(equality)])).
cnf(c_0_14,plain,(X1=add(X1,additive_identity)),inference(pm,[status(thm)],[c_0_1,c_0_10,theory(equality)])).
cnf(c_0_15,plain,(add(additive_identity,multiplicative_identity)=multiplicative_identity),inference(pm,[status(thm)],[c_0_7,c_0_12,theory(equality)])).
cnf(c_0_16,plain,(multiply(additive_identity,multiplicative_identity)=additive_identity),inference(pm,[status(thm)],[c_0_8,c_0_12,theory(equality)])).
cnf(c_0_17,negated_conjecture,(multiply(a,a)!=a),inference(pm,[status(thm)],[c_0_9,c_0_2,theory(equality)])).
cnf(c_0_18,plain,(multiply(multiplicative_identity,X1)=X1),inference(pm,[status(thm)],[c_0_6,c_0_2,theory(equality)])).
cnf(c_0_19,negated_conjecture,(multiply(a,a)!=a),inference(pm,[status(thm)],[c_0_9,c_0_2,theory(equality)])).
cnf(c_0_20,plain,(multiply(multiplicative_identity,X1)=X1),inference(pm,[status(thm)],[c_0_6,c_0_2,theory(equality)])).
cnf(c_0_21,plain,(inverse(multiplicative_identity)=additive_identity),inference(pm,[status(thm)],[c_0_8,c_0_18,theory(equality)])).
cnf(c_0_22,plain,(X1=multiply(X1,multiplicative_identity)),inference(pm,[status(thm)],[c_0_2,c_0_18,theory(equality)])).
cnf(c_0_23,plain,(X1=multiply(X1,multiplicative_identity)),inference(pm,[status(thm)],[c_0_2,c_0_18,theory(equality)])).
cnf(c_0_24,plain,(add(multiplicative_identity,additive_identity)=multiplicative_identity),inference(pm,[status(thm)],[c_0_7,c_0_21,theory(equality)])).
cnf(c_0_25,plain,(multiply(multiplicative_identity,additive_identity)=additive_identity),inference(pm,[status(thm)],[c_0_8,c_0_21,theory(equality)])).
cnf(c_0_26,plain,(multiply(X1,add(X2,X3))=add(multiply(X1,X3),multiply(X1,X2))),inference(pm,[status(thm)],[c_0_1,c_0_4,theory(equality)])).
cnf(c_0_27,plain,(multiply(X1,add(X3,X2))=add(multiply(X1,X2),multiply(X1,X3))),inference(pm,[status(thm)],[c_0_1,c_0_4,theory(equality)])).
cnf(c_0_28,plain,(add(multiply(multiplicative_identity,X1),X2)=multiply(multiplicative_identity,add(X1,X2))),inference(pm,[status(thm)],[c_0_4,c_0_18,theory(equality)])).
cnf(c_0_29,plain,(add(multiply(X1,X2),additive_identity)=multiply(X1,add(X2,inverse(X1)))),inference(pm,[status(thm)],[c_0_4,c_0_8,theory(equality)])).
cnf(c_0_30,plain,(add(multiply(X1,X2),X1)=multiply(X1,add(X2,multiplicative_identity))),inference(pm,[status(thm)],[c_0_4,c_0_6,theory(equality)])).
cnf(c_0_31,plain,(add(multiply(X1,X2),multiply(X3,X1))=multiply(X1,add(X2,X3))),inference(pm,[status(thm)],[c_0_4,c_0_2,theory(equality)])).
cnf(c_0_32,plain,(add(multiply(X1,X2),multiply(X3,X1))=multiply(X1,add(X2,X3))),inference(pm,[status(thm)],[c_0_4,c_0_2,theory(equality)])).
cnf(c_0_33,plain,(add(X1,multiply(multiplicative_identity,X2))=multiply(multiplicative_identity,add(X1,X2))),inference(pm,[status(thm)],[c_0_4,c_0_18,theory(equality)])).
cnf(c_0_34,plain,(add(additive_identity,multiply(X1,X2))=multiply(X1,add(inverse(X1),X2))),inference(pm,[status(thm)],[c_0_4,c_0_8,theory(equality)])).
cnf(c_0_35,plain,(add(X1,multiply(X1,X2))=multiply(X1,add(multiplicative_identity,X2))),inference(pm,[status(thm)],[c_0_4,c_0_6,theory(equality)])).
cnf(c_0_36,plain,(add(multiply(X2,X1),multiply(X1,X3))=multiply(X1,add(X2,X3))),inference(pm,[status(thm)],[c_0_4,c_0_2,theory(equality)])).
cnf(c_0_37,plain,(add(multiply(X2,X1),multiply(X1,X3))=multiply(X1,add(X2,X3))),inference(pm,[status(thm)],[c_0_4,c_0_2,theory(equality)])).
cnf(c_0_38,plain,(multiply(multiplicative_identity,X1)=add(X1,inverse(multiplicative_identity))),inference(pm,[status(thm)],[c_0_18,c_0_29,theory(equality)])).
cnf(c_0_39,plain,(add(multiply(X1,X2),multiply(X1,X3))=multiply(X1,add(add(X2,inverse(X1)),X3))),inference(pm,[status(thm)],[c_0_4,c_0_29,theory(equality)])).
cnf(c_0_40,plain,(add(multiply(X1,X2),multiply(X1,X3))=multiply(X1,add(X2,add(X3,inverse(X1))))),inference(pm,[status(thm)],[c_0_4,c_0_29,theory(equality)])).
cnf(c_0_41,plain,(multiply(multiplicative_identity,add(X1,additive_identity))=multiply(multiplicative_identity,X1)),inference(pm,[status(thm)],[c_0_29,c_0_21,theory(equality)])).
cnf(c_0_42,plain,(multiply(additive_identity,add(X1,multiplicative_identity))=multiply(additive_identity,X1)),inference(pm,[status(thm)],[c_0_29,c_0_12,theory(equality)])).
cnf(c_0_43,plain,(multiply(X1,inverse(X1))=multiply(X1,additive_identity)),inference(pm,[status(thm)],[c_0_29,c_0_10,theory(equality)])).
cnf(c_0_44,plain,(multiply(X1,add(inverse(X1),X2))=multiply(X1,X2)),inference(pm,[status(thm)],[c_0_29,c_0_1,theory(equality)])).
cnf(c_0_45,plain,(multiply(X1,add(inverse(X1),X2))=multiply(X1,X2)),inference(pm,[status(thm)],[c_0_29,c_0_1,theory(equality)])).
cnf(c_0_46,plain,(multiply(X1,multiplicative_identity)=multiply(X1,X1)),inference(pm,[status(thm)],[c_0_29,c_0_7,theory(equality)])).
cnf(c_0_47,plain,(X1=multiply(X1,X1)),inference(pm,[status(thm)],[c_0_2,c_0_46,theory(equality)])).
cnf(c_0_48,plain,(add(X1,multiply(X1,X2))=multiply(X1,add(X1,X2))),inference(pm,[status(thm)],[c_0_4,c_0_46,theory(equality)])).
cnf(c_0_49,plain,(X1=multiply(X1,X1)),inference(pm,[status(thm)],[c_0_2,c_0_46,theory(equality)])).
cnf(c_0_50,plain,(add(multiply(X1,X2),X1)=multiply(X1,add(X2,X1))),inference(pm,[status(thm)],[c_0_4,c_0_46,theory(equality)])).
cnf(c_0_51,negated_conjecture,($false),inference(cn,[status(thm)],[c_0_9, theory(equality,[symmetry])])).
cnf(c_0_52,negated_conjecture,($false), c_0_51,['proof']).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(multiply(X1,add(X2,X3))=add(multiply(X1,X2),multiply(X1,X3))), file('Axioms/BOO004-0.ax', distributivity2)).
cnf(c_0_1,axiom,(multiply(X1,inverse(X1))=additive_identity), file('Axioms/BOO004-0.ax', multiplicative_inverse1)).
cnf(c_0_2,axiom,(add(X1,additive_identity)=X1), file('Axioms/BOO004-0.ax', additive_id1)).
cnf(c_0_3,axiom,(add(X1,inverse(X1))=multiplicative_identity), file('Axioms/BOO004-0.ax', additive_inverse1)).
cnf(c_0_4,axiom,(multiply(X1,multiplicative_identity)=X1), file('Axioms/BOO004-0.ax', multiplicative_id1)).
cnf(c_0_5,negated_conjecture,(multiply(a,a)!=a), file('<stdin>', prove_a_times_a_is_a)).
cnf(c_0_6,axiom,(add(multiply(X1,X2),multiply(X1,X3))=multiply(X1,add(X2,X3))), c_0_0).
cnf(c_0_7,axiom,(multiply(X1,inverse(X1))=additive_identity), c_0_1).
cnf(c_0_8,axiom,(add(X1,additive_identity)=X1), c_0_2).
cnf(c_0_9,plain,(add(multiply(X1,X2),multiply(X1,X3))=multiply(X1,add(X2,X3))), c_0_6).
cnf(c_0_10,plain,(multiply(X1,inverse(X1))=additive_identity), c_0_7).
cnf(c_0_11,plain,(add(X1,additive_identity)=X1), c_0_8).
cnf(c_0_12,axiom,(add(X1,inverse(X1))=multiplicative_identity), c_0_3).
cnf(c_0_13,axiom,(multiply(X1,multiplicative_identity)=X1), c_0_4).
cnf(c_0_14,negated_conjecture,(multiply(a,a)!=a), c_0_5).
cnf(c_0_15,plain,(multiply(X1,add(X2,inverse(X1)))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_9, c_0_10, theory(equality)]), c_0_11, theory(equality)])).
cnf(c_0_16,plain,(add(X1,inverse(X1))=multiplicative_identity), c_0_12).
cnf(c_0_17,plain,(multiply(X1,multiplicative_identity)=X1), c_0_13).
cnf(c_0_18,negated_conjecture,(multiply(a,a)!=a), c_0_14).
cnf(c_0_19,plain,(multiply(X1,X1)=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_15, c_0_16, theory(equality)]), c_0_17, theory(equality)])).
cnf(c_0_20,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19, theory(equality)]), theory(equality,[symmetry])]), ['proof']).
# SZS output end CNFRefutation.
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 9
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 9
# Processed clauses                    : 18
# ...of these trivial                  : 2
# ...subsumed                          : 2
# ...remaining for further processing  : 14
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 41
# ...of the previous two non-trivial   : 25
# Contextual simplify-reflections      : 0
# Paramodulations                      : 41
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 13
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 16
# ...number of literals in the above   : 16
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
