cnf(c_0_1,axiom,(apply(apply(u,X1),X2)=apply(X2,apply(apply(X1,X1),X2))), file('<stdin>', u_definition)).
cnf(c_0_2,negated_conjecture,(apply(X1,f(X1))!=apply(f(X1),apply(X1,f(X1)))), file('<stdin>', prove_fixed_point)).
# Initializing proof state
# Scanning for AC axioms
cnf(c_0_3,plain,(apply(apply(apply(X1,X1),apply(X2,X2)),apply(apply(u,X1),apply(X2,X2)))=apply(apply(u,X2),apply(apply(X1,X1),apply(X2,X2)))),inference(pm,[status(thm)],[c_0_1,c_0_1,theory(equality)])).
cnf(c_0_4,plain,(apply(apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),apply(X2,X2)),apply(apply(u,apply(apply(u,u),apply(X1,X1))),apply(X2,X2)))=apply(apply(u,X2),apply(apply(apply(apply(u,u),apply(X1,X1)),apply(apply(u,u),apply(X1,X1))),apply(X2,X2)))),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_5,plain,(apply(X1,apply(apply(apply(u,X2),apply(apply(u,u),apply(X2,X2))),X1))=apply(apply(u,apply(apply(u,u),apply(X2,X2))),X1)),inference(pm,[status(thm)],[c_0_1,c_0_3,theory(equality)])).
cnf(c_0_6,plain,(apply(apply(apply(X1,X1),apply(apply(u,X2),apply(apply(u,u),apply(X2,X2)))),apply(apply(u,X1),apply(apply(apply(u,u),apply(X2,X2)),apply(apply(u,u),apply(X2,X2)))))=apply(apply(u,apply(apply(u,u),apply(X2,X2))),apply(apply(X1,X1),apply(apply(apply(u,u),apply(X2,X2)),apply(apply(u,u),apply(X2,X2)))))),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_7,plain,(apply(apply(apply(X1,X1),apply(apply(apply(u,u),apply(X2,X2)),apply(apply(u,u),apply(X2,X2)))),apply(apply(u,X1),apply(apply(u,X2),apply(apply(u,u),apply(X2,X2)))))=apply(apply(u,apply(apply(u,u),apply(X2,X2))),apply(apply(X1,X1),apply(apply(apply(u,u),apply(X2,X2)),apply(apply(u,u),apply(X2,X2)))))),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_8,plain,(apply(apply(apply(u,X1),apply(X1,X1)),apply(apply(u,X1),apply(apply(X1,X1),apply(X1,X1))))=apply(apply(u,apply(X1,X1)),apply(apply(u,X1),apply(X1,X1)))),inference(pm,[status(thm)],[c_0_1,c_0_3,theory(equality)])).
cnf(c_0_9,plain,(apply(apply(apply(u,u),apply(apply(u,u),apply(u,u))),apply(apply(u,apply(u,u)),apply(apply(u,u),apply(u,u))))=apply(apply(u,apply(u,u)),apply(apply(apply(u,u),apply(u,u)),apply(apply(u,u),apply(u,u))))),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_10,plain,(apply(apply(apply(apply(apply(X1,X1),u),apply(apply(X1,X1),u)),apply(X2,X2)),apply(apply(apply(u,X1),u),apply(X2,X2)))=apply(apply(u,X2),apply(apply(apply(apply(X1,X1),u),apply(apply(X1,X1),u)),apply(X2,X2)))),inference(pm,[status(thm)],[c_0_3,c_0_1,theory(equality)])).
cnf(c_0_11,plain,(apply(apply(apply(u,u),apply(apply(u,u),apply(u,u))),apply(apply(u,u),apply(u,u)))=apply(apply(u,apply(u,u)),apply(apply(u,u),apply(apply(u,u),apply(u,u))))),inference(pm,[status(thm)],[c_0_3,c_0_1,theory(equality)])).
cnf(c_0_12,plain,(apply(apply(apply(u,X1),apply(X1,X1)),apply(apply(u,X1),apply(apply(X1,X1),apply(X1,X1))))=apply(apply(u,apply(X1,X1)),apply(apply(X1,X1),apply(apply(X1,X1),apply(X1,X1))))),inference(pm,[status(thm)],[c_0_3,c_0_1,theory(equality)])).
cnf(c_0_13,plain,(apply(apply(apply(apply(u,u),apply(u,u)),apply(apply(u,u),apply(u,u))),apply(apply(u,u),apply(u,u)))=apply(apply(u,apply(u,u)),apply(apply(apply(u,u),apply(u,u)),apply(apply(u,u),apply(u,u))))),inference(pm,[status(thm)],[c_0_3,c_0_11,theory(equality)])).
cnf(c_0_14,negated_conjecture,(apply(apply(u,X1),f(apply(X1,X1)))!=apply(apply(X1,X1),f(apply(X1,X1)))),inference(pm,[status(thm)],[c_0_2,c_0_1,theory(equality)])).
cnf(c_0_15,negated_conjecture,(apply(apply(u,apply(apply(u,u),apply(X1,X1))),f(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1)))))!=apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),f(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1)))))),inference(pm,[status(thm)],[c_0_2,c_0_5,theory(equality)])).
cnf(c_0_16,plain,(apply(apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),apply(apply(u,X2),apply(apply(u,u),apply(X2,X2)))),apply(apply(u,apply(apply(u,u),apply(X1,X1))),apply(apply(u,X2),apply(apply(u,u),apply(X2,X2)))))=apply(apply(u,apply(apply(u,u),apply(X2,X2))),apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),apply(apply(u,X2),apply(apply(u,u),apply(X2,X2)))))),inference(pm,[status(thm)],[c_0_5,c_0_5,theory(equality)])).
cnf(c_0_17,plain,(apply(apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),apply(X2,X2)),apply(apply(u,apply(apply(u,u),apply(X1,X1))),apply(X2,X2)))=apply(apply(u,X2),apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),apply(X2,X2)))),inference(pm,[status(thm)],[c_0_1,c_0_5,theory(equality)])).
cnf(c_0_18,plain,(apply(apply(apply(apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),u),apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),u)),apply(X2,X2)),apply(apply(apply(u,apply(apply(u,u),apply(X1,X1))),u),apply(X2,X2)))=apply(apply(u,X2),apply(apply(apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),u),apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),u)),apply(X2,X2)))),inference(pm,[status(thm)],[c_0_3,c_0_5,theory(equality)])).
cnf(c_0_19,plain,(apply(X1,apply(apply(apply(apply(u,apply(apply(u,u),apply(X2,X2))),u),apply(apply(u,u),apply(apply(apply(apply(u,X2),apply(apply(u,u),apply(X2,X2))),u),apply(apply(apply(u,X2),apply(apply(u,u),apply(X2,X2))),u)))),X1))=apply(apply(u,apply(apply(u,u),apply(apply(apply(apply(u,X2),apply(apply(u,u),apply(X2,X2))),u),apply(apply(apply(u,X2),apply(apply(u,u),apply(X2,X2))),u)))),X1)),inference(pm,[status(thm)],[c_0_5,c_0_5,theory(equality)])).
cnf(c_0_20,plain,(apply(X1,apply(apply(apply(u,apply(apply(u,u),apply(X2,X2))),apply(apply(u,u),apply(apply(u,X2),apply(apply(u,u),apply(X2,X2))))),X1))=apply(apply(u,apply(apply(u,u),apply(apply(apply(u,u),apply(X2,X2)),apply(apply(u,u),apply(X2,X2))))),X1)),inference(pm,[status(thm)],[c_0_5,c_0_3,theory(equality)])).
cnf(c_0_21,plain,(apply(X1,apply(apply(apply(u,apply(u,u)),apply(apply(u,u),apply(u,u))),X1))=apply(apply(u,apply(apply(u,u),apply(apply(u,u),apply(u,u)))),X1)),inference(pm,[status(thm)],[c_0_5,c_0_1,theory(equality)])).
cnf(c_0_22,plain,(apply(X1,apply(apply(apply(apply(u,X2),u),apply(apply(u,u),apply(apply(apply(X2,X2),u),apply(apply(X2,X2),u)))),X1))=apply(apply(u,apply(apply(u,u),apply(apply(apply(X2,X2),u),apply(apply(X2,X2),u)))),X1)),inference(pm,[status(thm)],[c_0_5,c_0_1,theory(equality)])).
cnf(c_0_23,plain,(apply(X1,apply(apply(apply(u,u),apply(u,u)),X1))=apply(apply(u,apply(apply(u,u),apply(u,u))),X1)),inference(pm,[status(thm)],[c_0_5,c_0_1,theory(equality)])).
cnf(c_0_24,plain,(apply(apply(apply(X1,X1),apply(apply(u,X2),apply(apply(u,u),apply(X2,X2)))),apply(apply(u,X1),apply(apply(u,X2),apply(apply(u,u),apply(X2,X2)))))=apply(apply(u,apply(apply(u,u),apply(X2,X2))),apply(apply(X1,X1),apply(apply(u,X2),apply(apply(u,u),apply(X2,X2)))))),inference(pm,[status(thm)],[c_0_5,c_0_1,theory(equality)])).
cnf(c_0_25,plain,(apply(apply(apply(u,u),apply(apply(u,u),apply(u,u))),apply(apply(u,apply(u,u)),apply(apply(u,u),apply(apply(u,u),apply(u,u)))))=apply(apply(u,apply(apply(u,u),apply(u,u))),apply(apply(u,u),apply(apply(u,u),apply(u,u))))),inference(pm,[status(thm)],[c_0_5,c_0_3,theory(equality)])).
cnf(c_0_26,negated_conjecture,(apply(f(apply(u,apply(apply(u,u),apply(u,u)))),apply(apply(u,apply(u,u)),f(apply(u,apply(apply(u,u),apply(u,u))))))!=apply(apply(u,apply(apply(u,u),apply(u,u))),f(apply(u,apply(apply(u,u),apply(u,u)))))),inference(pm,[status(thm)],[c_0_2,c_0_21,theory(equality)])).
cnf(c_0_27,plain,(apply(apply(u,apply(u,u)),apply(apply(X1,X1),apply(u,apply(apply(u,u),apply(u,u)))))=apply(apply(u,X1),apply(u,apply(apply(u,u),apply(u,u))))),inference(pm,[status(thm)],[c_0_1,c_0_21,theory(equality)])).
cnf(c_0_28,plain,(apply(apply(u,apply(u,u)),apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),apply(u,apply(apply(u,u),apply(u,u)))))=apply(apply(u,apply(apply(u,u),apply(X1,X1))),apply(u,apply(apply(u,u),apply(u,u))))),inference(pm,[status(thm)],[c_0_5,c_0_21,theory(equality)])).
cnf(c_0_29,plain,(apply(apply(apply(apply(u,apply(u,u)),apply(u,apply(apply(u,u),apply(u,u)))),apply(X1,X1)),apply(apply(u,apply(u,apply(apply(u,u),apply(u,u)))),apply(X1,X1)))=apply(apply(u,X1),apply(apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u)))),apply(X1,X1)))),inference(pm,[status(thm)],[c_0_3,c_0_21,theory(equality)])).
cnf(c_0_30,plain,(apply(X1,apply(apply(apply(u,apply(u,u)),apply(u,apply(apply(u,u),apply(u,u)))),X1))=apply(apply(u,apply(u,apply(apply(u,u),apply(u,u)))),X1)),inference(pm,[status(thm)],[c_0_1,c_0_21,theory(equality)])).
cnf(c_0_31,plain,(apply(apply(apply(X1,X1),apply(apply(u,apply(u,u)),apply(u,apply(apply(u,u),apply(u,u))))),apply(apply(u,X1),apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u))))))=apply(apply(u,apply(u,apply(apply(u,u),apply(u,u)))),apply(apply(X1,X1),apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u))))))),inference(pm,[status(thm)],[c_0_3,c_0_21,theory(equality)])).
cnf(c_0_32,plain,(apply(apply(apply(X1,X1),apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u))))),apply(apply(u,X1),apply(apply(u,apply(u,u)),apply(u,apply(apply(u,u),apply(u,u))))))=apply(apply(u,apply(u,apply(apply(u,u),apply(u,u)))),apply(apply(X1,X1),apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u))))))),inference(pm,[status(thm)],[c_0_3,c_0_21,theory(equality)])).
cnf(c_0_33,plain,(apply(X1,apply(apply(apply(u,apply(u,apply(apply(u,u),apply(u,u)))),apply(apply(u,u),apply(apply(u,apply(u,u)),apply(u,apply(apply(u,u),apply(u,u)))))),X1))=apply(apply(u,apply(apply(u,u),apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u)))))),X1)),inference(pm,[status(thm)],[c_0_5,c_0_21,theory(equality)])).
cnf(c_0_34,plain,(apply(X1,apply(apply(apply(u,apply(u,u)),apply(apply(u,u),apply(apply(apply(u,u),apply(u,u)),apply(apply(u,u),apply(u,u))))),X1))=apply(apply(u,apply(apply(u,u),apply(apply(apply(u,u),apply(u,u)),apply(apply(u,u),apply(u,u))))),X1)),inference(pm,[status(thm)],[c_0_5,c_0_21,theory(equality)])).
cnf(c_0_35,plain,(apply(apply(apply(apply(apply(u,u),apply(u,u)),apply(apply(u,u),apply(u,u))),apply(X1,X1)),apply(apply(u,apply(u,u)),apply(X1,X1)))=apply(apply(u,X1),apply(apply(apply(apply(u,u),apply(u,u)),apply(apply(u,u),apply(u,u))),apply(X1,X1)))),inference(pm,[status(thm)],[c_0_3,c_0_21,theory(equality)])).
cnf(c_0_36,negated_conjecture,($false),inference(er,[status(thm)],[c_0_14,theory(equality)])).
cnf(c_0_37,negated_conjecture,(apply(apply(apply(apply(u,u),apply(X1,X1)),apply(apply(u,u),apply(X1,X1))),f(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1)))))!=apply(apply(u,apply(apply(u,u),apply(X1,X1))),f(apply(apply(apply(u,u),apply(X1,X1)),apply(apply(u,u),apply(X1,X1)))))),inference(pm,[status(thm)],[c_0_14,c_0_3,theory(equality)])).
cnf(c_0_38,negated_conjecture,(apply(apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u)))),f(apply(apply(u,apply(u,u)),apply(u,apply(apply(u,u),apply(u,u))))))!=apply(apply(u,apply(u,apply(apply(u,u),apply(u,u)))),f(apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u))))))),inference(pm,[status(thm)],[c_0_14,c_0_21,theory(equality)])).
cnf(c_0_39,negated_conjecture,(apply(apply(apply(u,X1),apply(apply(u,u),apply(X1,X1))),f(apply(apply(apply(u,u),apply(X1,X1)),apply(apply(u,u),apply(X1,X1)))))!=apply(apply(u,apply(apply(u,u),apply(X1,X1))),f(apply(apply(apply(u,u),apply(X1,X1)),apply(apply(u,u),apply(X1,X1)))))),inference(pm,[status(thm)],[c_0_14,c_0_3,theory(equality)])).
cnf(c_0_40,negated_conjecture,(apply(apply(apply(u,apply(u,u)),apply(u,apply(apply(u,u),apply(u,u)))),f(apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u))))))!=apply(apply(u,apply(u,apply(apply(u,u),apply(u,u)))),f(apply(apply(u,apply(apply(u,u),apply(u,u))),apply(u,apply(apply(u,u),apply(u,u))))))),inference(pm,[status(thm)],[c_0_14,c_0_21,theory(equality)])).
cnf(c_0_41,negated_conjecture,($false), c_0_36,['proof']).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,negated_conjecture,(apply(X1,f(X1))!=apply(f(X1),apply(X1,f(X1)))), file('<stdin>', prove_fixed_point)).
cnf(c_0_1,axiom,(apply(apply(u,X1),X2)=apply(X2,apply(apply(X1,X1),X2))), file('<stdin>', u_definition)).
cnf(c_0_2,negated_conjecture,(apply(f(X1),apply(X1,f(X1)))!=apply(X1,f(X1))), c_0_0).
cnf(c_0_3,axiom,(apply(X2,apply(apply(X1,X1),X2))=apply(apply(u,X1),X2)), c_0_1).
cnf(c_0_4,negated_conjecture,(apply(f(X1),apply(X1,f(X1)))!=apply(X1,f(X1))), c_0_2).
cnf(c_0_5,plain,(apply(X1,apply(apply(X2,X2),X1))=apply(apply(u,X2),X1)), c_0_3).
cnf(c_0_6,negated_conjecture,(apply(apply(u,X1),f(apply(X1,X1)))!=apply(apply(X1,X1),f(apply(X1,X1)))), inference(pm,[status(thm)],[c_0_4, c_0_5, theory(equality)])).
cnf(c_0_7,negated_conjecture,($false), inference(er,[status(thm)],[c_0_6, theory(equality)]), ['proof']).
# SZS output end CNFRefutation.
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 2
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 2
# Processed clauses                    : 8
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 7
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 38
# ...of the previous two non-trivial   : 30
# Contextual simplify-reflections      : 0
# Paramodulations                      : 37
# Factorizations                       : 0
# Equation resolutions                 : 1
# Current number of processed clauses  : 7
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 24
# ...number of literals in the above   : 24
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
