cnf(c_0_1,axiom,(divide(divide(X1,divide(divide(X1,X2),X3)),X2)=X3), file('<stdin>', single_axiom)).
cnf(c_0_2,axiom,(multiply(X1,X2)=divide(X1,divide(divide(X3,X3),X2))), file('<stdin>', multiply)).
cnf(c_0_3,axiom,(inverse(X1)=divide(divide(X2,X2),X1)), file('<stdin>', inverse)).
cnf(c_0_4,axiom,(identity=divide(X1,X1)), file('<stdin>', identity)).
cnf(c_0_5,negated_conjecture,(multiply(inverse(a1),a1)!=multiply(inverse(b1),b1)), file('<stdin>', prove_these_axioms_1)).
# Initializing proof state
# Scanning for AC axioms
cnf(c_0_6,plain,(inverse(identity)=identity),inference(pm,[status(thm)],[c_0_4,c_0_3,theory(equality)])).
cnf(c_0_7,negated_conjecture,($false),inference(cn,[status(thm)],[c_0_5, theory(equality,[symmetry])])).
cnf(c_0_8,negated_conjecture,($false), c_0_7,['proof']).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(identity=divide(X1,X1)), file('<stdin>', identity)).
cnf(c_0_1,axiom,(inverse(X1)=divide(divide(X2,X2),X1)), file('<stdin>', inverse)).
cnf(c_0_2,negated_conjecture,(multiply(inverse(a1),a1)!=multiply(inverse(b1),b1)), file('<stdin>', prove_these_axioms_1)).
cnf(c_0_3,axiom,(multiply(X1,X2)=divide(X1,divide(divide(X3,X3),X2))), file('<stdin>', multiply)).
cnf(c_0_4,axiom,(divide(X1,X1)=identity), c_0_0).
cnf(c_0_5,axiom,(divide(divide(X2,X2),X1)=inverse(X1)), c_0_1).
cnf(c_0_6,plain,(divide(X1,X1)=identity), c_0_4).
cnf(c_0_7,negated_conjecture,(multiply(inverse(b1),b1)!=multiply(inverse(a1),a1)), c_0_2).
cnf(c_0_8,axiom,(divide(X1,divide(divide(X3,X3),X2))=multiply(X1,X2)), c_0_3).
cnf(c_0_9,plain,(divide(identity,X1)=inverse(X1)), inference(rw,[status(thm)],[c_0_5, c_0_6, theory(equality)])).
cnf(c_0_10,negated_conjecture,(multiply(inverse(a1),a1)!=multiply(inverse(b1),b1)), c_0_7).
cnf(c_0_11,plain,(multiply(X1,X2)=divide(X1,inverse(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_6, theory(equality)]), c_0_9, theory(equality)])).
cnf(c_0_12,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11, theory(equality)]), c_0_6, theory(equality)]), c_0_11, theory(equality)]), c_0_6, theory(equality)]), theory(equality,[symmetry])]), ['proof']).
# SZS output end CNFRefutation.
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 5
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 5
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 1
# ...of the previous two non-trivial   : 1
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 4
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 1
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
