cnf(c_0_1,axiom,(f(X1)=g(X2)), file('<stdin>', clause_1)).
cnf(c_0_2,negated_conjecture,(f(f(a))!=f(g(b))), file('<stdin>', prove_this)).
# Initializing proof state
# Scanning for AC axioms
cnf(c_0_3,plain,(g(X3)=g(X1)),inference(pm,[status(thm)],[c_0_1,c_0_1,theory(equality)])).
cnf(c_0_4,plain,(f(X3)=f(X2)),inference(pm,[status(thm)],[c_0_1,c_0_1,theory(equality)])).
cnf(c_0_5,plain,(g(X3)=g(X1)),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_6,plain,(g(X3)=f(X2)),inference(pm,[status(thm)],[c_0_1,c_0_3,theory(equality)])).
cnf(c_0_7,plain,(g(X3)=g(X2)),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_8,plain,(g(X3)=g(X1)),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_9,plain,(g(X3)=f(X2)),inference(pm,[status(thm)],[c_0_1,c_0_3,theory(equality)])).
cnf(c_0_10,plain,(g(X3)=g(X2)),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_11,plain,(f(X3)=f(X2)),inference(pm,[status(thm)],[c_0_4,c_0_4,theory(equality)])).
cnf(c_0_12,plain,(f(X3)=g(X1)),inference(pm,[status(thm)],[c_0_1,c_0_4,theory(equality)])).
cnf(c_0_13,plain,(f(X3)=f(X1)),inference(pm,[status(thm)],[c_0_4,c_0_4,theory(equality)])).
cnf(c_0_14,plain,(f(X3)=f(X2)),inference(pm,[status(thm)],[c_0_4,c_0_4,theory(equality)])).
cnf(c_0_15,plain,(f(X3)=g(X1)),inference(pm,[status(thm)],[c_0_1,c_0_4,theory(equality)])).
cnf(c_0_16,plain,(f(X3)=f(X1)),inference(pm,[status(thm)],[c_0_4,c_0_4,theory(equality)])).
cnf(c_0_17,negated_conjecture,($false),inference(sr,[status(thm)],[c_0_2,c_0_4,theory(equality)])).
cnf(c_0_18,negated_conjecture,($false), c_0_17,['proof']).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(f(X1)=g(X2)), file('<stdin>', clause_1)).
cnf(c_0_1,negated_conjecture,(f(f(a))!=f(g(b))), file('<stdin>', prove_this)).
cnf(c_0_2,axiom,(g(X2)=f(X1)), c_0_0).
cnf(c_0_3,plain,(g(X1)=f(X2)), c_0_2).
cnf(c_0_4,negated_conjecture,(f(g(b))!=f(f(a))), c_0_1).
cnf(c_0_5,plain,(f(X1)=f(X2)), inference(pm,[status(thm)],[c_0_3, c_0_3, theory(equality)])).
cnf(c_0_6,negated_conjecture,($false), inference(sr,[status(thm)],[c_0_4, c_0_5, theory(equality)]), ['proof']).
# SZS output end CNFRefutation.
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 2
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 2
# Processed clauses                    : 4
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 3
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 14
# ...of the previous two non-trivial   : 14
# Contextual simplify-reflections      : 0
# Paramodulations                      : 14
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 3
#    Positive orientable unit clauses  : 0
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 12
# ...number of literals in the above   : 12
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 34
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
