fof(c_0_1, axiom, ((![X1]:f(X1,g(X1))=g(X1)&![X2]:![X3]:h(f(g(X2),X3))=g(X3))), file('<stdin>', clause_1)).
fof(c_0_2, conjecture, (h(g(g(c)))=g(g(g(c)))), file('<stdin>', clause_1)).
fof(c_0_3, negated_conjecture, (~(h(g(g(c)))=g(g(g(c))))),inference(assume_negation, [status(cth)],[c_0_2])).
fof(c_0_4, negated_conjecture, (h(g(g(c)))!=g(g(g(c)))),inference(fof_simplification, [status(thm)],[c_0_3])).
fof(c_0_5, plain, ((![X4]:f(X4,g(X4))=g(X4)&![X5]:![X6]:h(f(g(X5),X6))=g(X6))),inference(variable_rename, [status(thm)],[c_0_1])).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:(f(X4,g(X4))=g(X4)&h(f(g(X5),X6))=g(X6))),inference(shift_quantors, [status(thm)],[c_0_5])).
cnf(c_0_7,plain,(h(f(g(X1),X2))=g(X2)),inference(split_conjunct, [status(thm)],[c_0_6])).
cnf(c_0_8,plain,(f(X1,g(X1))=g(X1)),inference(split_conjunct, [status(thm)],[c_0_6])).
cnf(c_0_9,negated_conjecture,(h(g(g(c)))!=g(g(g(c)))),inference(split_conjunct, [status(thm)],[c_0_4])).
# Initializing proof state
# Scanning for AC axioms
cnf(c_0_10,plain,(h(g(g(X1)))=g(g(g(X1)))),inference(pm,[status(thm)],[c_0_7,c_0_8,theory(equality)])).
cnf(c_0_11,plain,(f(g(g(X1)),h(g(g(X1))))=g(g(g(X1)))),inference(pm,[status(thm)],[c_0_8,c_0_10,theory(equality)])).
cnf(c_0_12,plain,(h(f(h(g(g(X1))),X2))=g(X2)),inference(pm,[status(thm)],[c_0_7,c_0_10,theory(equality)])).
cnf(c_0_13,plain,(g(g(h(g(g(X1)))))=h(g(g(g(g(X1)))))),inference(pm,[status(thm)],[c_0_10,c_0_10,theory(equality)])).
cnf(c_0_14,plain,(g(h(g(g(X1))))=h(g(g(g(X1))))),inference(pm,[status(thm)],[c_0_10,c_0_10,theory(equality)])).
cnf(c_0_15,negated_conjecture,($false),inference(cn,[status(thm)],[c_0_9, theory(equality,[symmetry])])).
cnf(c_0_16,negated_conjecture,($false), c_0_15,['proof']).

