cnf(c_0_1,axiom,(response(response(lark,X1),X2)=response(X1,response(X2,X2))), file('<stdin>', lark_exists)).
cnf(c_0_2,negated_conjecture,(response(X1,X1)!=X1), file('<stdin>', prove_the_bird_exists)).
# Initializing proof state
# Scanning for AC axioms
cnf(c_0_3,plain,(response(response(response(lark,lark),X1),X2)=response(response(X1,X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_1,c_0_1,theory(equality)])).
cnf(c_0_4,plain,(response(X1,response(response(lark,response(X2,X2)),X2))=response(response(lark,X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_1,c_0_1,theory(equality)])).
cnf(c_0_5,plain,(response(X1,response(response(lark,response(lark,response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),response(X2,X2)))=response(response(lark,X1),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_4,c_0_4,theory(equality)])).
cnf(c_0_6,plain,(response(response(response(lark,lark),response(X1,X1)),X2)=response(response(response(lark,response(X1,X1)),X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_1,c_0_4,theory(equality)])).
cnf(c_0_7,plain,(response(X1,response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2)))=response(response(lark,X1),response(response(lark,response(X2,X2)),X2))),inference(pm,[status(thm)],[c_0_1,c_0_4,theory(equality)])).
cnf(c_0_8,plain,(response(X1,response(response(lark,response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2))),response(response(lark,response(X2,X2)),X2)))=response(response(lark,X1),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_4,c_0_4,theory(equality)])).
cnf(c_0_9,plain,(response(X1,response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2)))=response(response(lark,X1),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_4,c_0_1,theory(equality)])).
cnf(c_0_10,plain,(response(X1,response(response(response(lark,lark),X2),X2))=response(response(lark,X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_4,c_0_1,theory(equality)])).
cnf(c_0_11,plain,(response(X1,response(response(lark,response(lark,response(response(X2,X2),response(X2,X2)))),X2))=response(response(lark,X1),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_4,c_0_1,theory(equality)])).
cnf(c_0_12,plain,(response(X1,response(response(X2,X2),response(X2,X2)))=response(response(lark,X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_4,c_0_1,theory(equality)])).
cnf(c_0_13,plain,(response(response(response(response(lark,lark),X1),X1),response(X2,X2))=response(response(response(lark,lark),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_14,plain,(response(X1,response(response(response(lark,lark),X2),X2))=response(response(lark,X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_1,c_0_3,theory(equality)])).
cnf(c_0_15,plain,(response(response(X1,X1),response(response(response(lark,lark),X2),X2))=response(response(response(lark,lark),X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_3,c_0_3,theory(equality)])).
cnf(c_0_16,plain,(response(X1,response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2)))=response(response(lark,X1),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_4,c_0_3,theory(equality)])).
cnf(c_0_17,plain,(response(response(response(lark,lark),X1),X2)=response(response(lark,response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_1,c_0_3,theory(equality)])).
cnf(c_0_18,plain,(response(response(X1,X1),response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2)))=response(response(response(lark,lark),X1),response(response(lark,response(X2,X2)),X2))),inference(pm,[status(thm)],[c_0_3,c_0_4,theory(equality)])).
cnf(c_0_19,plain,(response(response(X1,X1),response(response(lark,response(X2,X2)),X2))=response(response(response(lark,lark),X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_3,c_0_1,theory(equality)])).
cnf(c_0_20,plain,(response(response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(lark,response(X1,X1)),X1)),X2)),inference(pm,[status(thm)],[c_0_3,c_0_4,theory(equality)])).
cnf(c_0_21,plain,(response(response(response(lark,response(X1,X1)),X1),response(X2,X2))=response(response(response(lark,lark),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_3,c_0_1,theory(equality)])).
cnf(c_0_22,plain,(response(response(lark,response(X1,X1)),X2)=response(response(X1,X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_3,c_0_1,theory(equality)])).
cnf(c_0_23,negated_conjecture,(response(response(response(lark,lark),X1),X1)!=response(X1,X1)),inference(pm,[status(thm)],[c_0_2,c_0_3,theory(equality)])).
cnf(c_0_24,negated_conjecture,(response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1))!=response(response(lark,response(X1,X1)),X1)),inference(pm,[status(thm)],[c_0_2,c_0_4,theory(equality)])).
cnf(c_0_25,negated_conjecture,(response(response(lark,response(X1,X1)),X1)!=response(X1,X1)),inference(pm,[status(thm)],[c_0_2,c_0_1,theory(equality)])).
cnf(c_0_26,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(response(lark,lark),X2),X2))),response(X2,X2)))=response(response(lark,X1),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_10,c_0_10,theory(equality)])).
cnf(c_0_27,plain,(response(X1,response(response(lark,response(lark,response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),response(X2,X2)))=response(response(lark,X1),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_4,c_0_10,theory(equality)])).
cnf(c_0_28,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(X2,X2)),response(response(response(lark,lark),X2),X2)))=response(response(lark,X1),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_10,c_0_10,theory(equality)])).
cnf(c_0_29,plain,(response(response(response(lark,lark),response(X1,X1)),X2)=response(response(response(response(lark,lark),X1),X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_1,c_0_10,theory(equality)])).
cnf(c_0_30,negated_conjecture,(response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1))!=response(response(response(lark,lark),X1),X1)),inference(pm,[status(thm)],[c_0_2,c_0_10,theory(equality)])).
cnf(c_0_31,plain,(response(response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,lark),X1),X1)),X2)),inference(pm,[status(thm)],[c_0_3,c_0_10,theory(equality)])).
cnf(c_0_32,plain,(response(X1,response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2)))=response(response(lark,X1),response(response(response(lark,lark),X2),X2))),inference(pm,[status(thm)],[c_0_1,c_0_10,theory(equality)])).
cnf(c_0_33,plain,(response(response(X1,X1),response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2)))=response(response(response(lark,lark),X1),response(response(response(lark,lark),X2),X2))),inference(pm,[status(thm)],[c_0_3,c_0_10,theory(equality)])).
cnf(c_0_34,plain,(response(X1,response(response(lark,response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2))),response(response(response(lark,lark),X2),X2)))=response(response(lark,X1),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_4,c_0_10,theory(equality)])).
cnf(c_0_35,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(X2,X2)),response(response(lark,response(X2,X2)),X2)))=response(response(lark,X1),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_10,c_0_4,theory(equality)])).
cnf(c_0_36,plain,(response(X1,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))=response(response(lark,X1),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_10,c_0_1,theory(equality)])).
cnf(c_0_37,plain,(response(X1,response(response(response(response(lark,lark),lark),X2),response(X2,X2)))=response(response(lark,X1),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_10,c_0_3,theory(equality)])).
cnf(c_0_38,plain,(response(X1,response(response(lark,response(X2,X2)),X2))=response(response(lark,X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_10,c_0_1,theory(equality)])).
cnf(c_0_39,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(lark,response(X2,X2)),X2))),response(X2,X2)))=response(response(lark,X1),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_10,c_0_4,theory(equality)])).
cnf(c_0_40,plain,(response(X1,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))=response(response(lark,X1),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_10,c_0_1,theory(equality)])).
cnf(c_0_41,plain,(response(X1,response(response(response(lark,lark),response(lark,lark)),lark))=response(response(lark,X1),response(response(lark,lark),response(lark,lark)))),inference(pm,[status(thm)],[c_0_10,c_0_3,theory(equality)])).
cnf(c_0_42,plain,(response(X1,response(response(X2,X2),response(X2,X2)))=response(response(lark,X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_10,c_0_3,theory(equality)])).
cnf(c_0_43,negated_conjecture,(response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1))!=response(response(X1,X1),response(X1,X1))),inference(pm,[status(thm)],[c_0_25,c_0_3,theory(equality)])).
cnf(c_0_44,negated_conjecture,(response(response(lark,response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1))),response(response(lark,response(X1,X1)),X1))!=response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_25,c_0_4,theory(equality)])).
cnf(c_0_45,negated_conjecture,(response(response(lark,response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1))),response(response(response(lark,lark),X1),X1))!=response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_25,c_0_10,theory(equality)])).
cnf(c_0_46,negated_conjecture,(response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1))!=response(response(X1,X1),response(X1,X1))),inference(pm,[status(thm)],[c_0_25,c_0_1,theory(equality)])).
cnf(c_0_47,negated_conjecture,(response(response(response(lark,lark),X1),X1)!=response(X1,X1)),inference(pm,[status(thm)],[c_0_25,c_0_1,theory(equality)])).
cnf(c_0_48,negated_conjecture,(response(response(lark,response(lark,response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1)))),response(X1,X1))!=response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_25,c_0_4,theory(equality)])).
cnf(c_0_49,negated_conjecture,(response(response(lark,response(lark,response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1)))),response(X1,X1))!=response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_25,c_0_10,theory(equality)])).
cnf(c_0_50,negated_conjecture,(response(response(lark,response(lark,response(response(X1,X1),response(X1,X1)))),X1)!=response(response(X1,X1),response(X1,X1))),inference(pm,[status(thm)],[c_0_25,c_0_1,theory(equality)])).
cnf(c_0_51,negated_conjecture,(response(response(X1,X1),response(X1,X1))!=response(X1,X1)),inference(pm,[status(thm)],[c_0_25,c_0_1,theory(equality)])).
cnf(c_0_52,plain,(response(response(response(lark,lark),response(X1,X1)),X2)=response(response(lark,response(response(lark,response(X1,X1)),X1)),X2)),inference(pm,[status(thm)],[c_0_1,c_0_6,theory(equality)])).
cnf(c_0_53,plain,(response(response(response(lark,response(X1,X1)),X1),response(response(response(lark,lark),X2),X2))=response(response(response(lark,lark),response(X1,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_6,c_0_3,theory(equality)])).
cnf(c_0_54,plain,(response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(lark,response(X2,X2)),X2))),inference(pm,[status(thm)],[c_0_6,c_0_4,theory(equality)])).
cnf(c_0_55,plain,(response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(response(lark,lark),X2),X2))),inference(pm,[status(thm)],[c_0_6,c_0_10,theory(equality)])).
cnf(c_0_56,plain,(response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X2,X2)),X2))=response(response(response(lark,lark),response(X1,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_6,c_0_1,theory(equality)])).
cnf(c_0_57,plain,(response(response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_3,theory(equality)])).
cnf(c_0_58,plain,(response(response(response(lark,response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1))),response(response(lark,response(X1,X1)),X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_4,theory(equality)])).
cnf(c_0_59,plain,(response(response(response(lark,response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1))),response(response(response(lark,lark),X1),X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_10,theory(equality)])).
cnf(c_0_60,plain,(response(response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_1,theory(equality)])).
cnf(c_0_61,plain,(response(response(response(response(lark,lark),X1),X1),response(X2,X2))=response(response(response(lark,lark),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_6,c_0_1,theory(equality)])).
cnf(c_0_62,plain,(response(response(response(lark,response(lark,response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1)))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_4,theory(equality)])).
cnf(c_0_63,plain,(response(response(response(lark,response(lark,response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1)))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_10,theory(equality)])).
cnf(c_0_64,plain,(response(response(response(lark,response(lark,response(response(X1,X1),response(X1,X1)))),X1),response(X2,X2))=response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_1,theory(equality)])).
cnf(c_0_65,plain,(response(response(response(X1,X1),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_6,c_0_1,theory(equality)])).
cnf(c_0_66,plain,(response(response(response(lark,lark),response(X1,X1)),X2)=response(response(lark,response(response(response(lark,lark),X1),X1)),X2)),inference(pm,[status(thm)],[c_0_1,c_0_13,theory(equality)])).
cnf(c_0_67,plain,(response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X2),X2))=response(response(response(lark,lark),response(X1,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_13,c_0_3,theory(equality)])).
cnf(c_0_68,plain,(response(response(response(response(lark,lark),X1),X1),response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(lark,response(X2,X2)),X2))),inference(pm,[status(thm)],[c_0_13,c_0_4,theory(equality)])).
cnf(c_0_69,plain,(response(response(response(response(lark,lark),X1),X1),response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(response(lark,lark),X2),X2))),inference(pm,[status(thm)],[c_0_13,c_0_10,theory(equality)])).
cnf(c_0_70,plain,(response(response(response(response(lark,lark),X1),X1),response(response(lark,response(X2,X2)),X2))=response(response(response(lark,lark),response(X1,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_13,c_0_1,theory(equality)])).
cnf(c_0_71,plain,(response(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(X1,X1)),X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_4,theory(equality)])).
cnf(c_0_72,plain,(response(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,lark),X1),X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_10,theory(equality)])).
cnf(c_0_73,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_1,theory(equality)])).
cnf(c_0_74,plain,(response(response(response(response(response(lark,lark),lark),X1),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_3,theory(equality)])).
cnf(c_0_75,plain,(response(response(response(lark,response(X1,X1)),X1),response(X2,X2))=response(response(response(lark,lark),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_13,c_0_1,theory(equality)])).
cnf(c_0_76,plain,(response(response(response(lark,response(response(lark,lark),response(response(lark,response(X1,X1)),X1))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_4,theory(equality)])).
cnf(c_0_77,plain,(response(response(response(lark,response(response(lark,lark),response(response(response(lark,lark),X1),X1))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_10,theory(equality)])).
cnf(c_0_78,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(X2,X2))=response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_1,theory(equality)])).
cnf(c_0_79,plain,(response(response(response(response(lark,lark),response(lark,lark)),lark),response(X1,X1))=response(response(response(lark,lark),response(response(lark,lark),response(lark,lark))),X1)),inference(pm,[status(thm)],[c_0_13,c_0_3,theory(equality)])).
cnf(c_0_80,plain,(response(response(response(X1,X1),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_13,c_0_3,theory(equality)])).
cnf(c_0_81,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(X1,X1)),X1))!=response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_23,c_0_4,theory(equality)])).
cnf(c_0_82,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,lark),X1),X1))!=response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_23,c_0_10,theory(equality)])).
cnf(c_0_83,negated_conjecture,(response(response(response(lark,response(lark,lark)),X1),response(X1,X1))!=response(response(X1,X1),response(X1,X1))),inference(pm,[status(thm)],[c_0_23,c_0_1,theory(equality)])).
cnf(c_0_84,negated_conjecture,(response(response(response(response(lark,lark),lark),X1),response(X1,X1))!=response(response(X1,X1),response(X1,X1))),inference(pm,[status(thm)],[c_0_23,c_0_3,theory(equality)])).
cnf(c_0_85,negated_conjecture,(response(response(lark,response(X1,X1)),X1)!=response(X1,X1)),inference(pm,[status(thm)],[c_0_23,c_0_1,theory(equality)])).
cnf(c_0_86,negated_conjecture,(response(response(lark,response(response(lark,lark),response(response(lark,response(X1,X1)),X1))),response(X1,X1))!=response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_23,c_0_4,theory(equality)])).
cnf(c_0_87,negated_conjecture,(response(response(lark,response(response(lark,lark),response(response(response(lark,lark),X1),X1))),response(X1,X1))!=response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_23,c_0_10,theory(equality)])).
cnf(c_0_88,negated_conjecture,(response(response(lark,response(response(lark,lark),response(X1,X1))),X1)!=response(response(X1,X1),response(X1,X1))),inference(pm,[status(thm)],[c_0_23,c_0_1,theory(equality)])).
cnf(c_0_89,negated_conjecture,(response(response(response(lark,lark),response(lark,lark)),lark)!=response(response(lark,lark),response(lark,lark))),inference(pm,[status(thm)],[c_0_23,c_0_3,theory(equality)])).
cnf(c_0_90,negated_conjecture,(response(response(X1,X1),response(X1,X1))!=response(X1,X1)),inference(pm,[status(thm)],[c_0_23,c_0_3,theory(equality)])).
cnf(c_0_91,plain,(response(X1,response(response(lark,response(lark,response(response(lark,lark),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),response(X2,X2)))=response(response(lark,X1),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_10,c_0_11,theory(equality)])).
cnf(c_0_92,plain,(response(response(response(lark,response(lark,response(response(lark,lark),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_11,theory(equality)])).
cnf(c_0_93,negated_conjecture,(response(response(lark,response(lark,response(response(lark,lark),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)))),response(X1,X1))!=response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_23,c_0_11,theory(equality)])).
cnf(c_0_94,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2))))),response(X2,X2)))=response(response(lark,X1),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_4,c_0_11,theory(equality)])).
cnf(c_0_95,plain,(response(response(response(lark,response(lark,response(lark,response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_11,theory(equality)])).
cnf(c_0_96,negated_conjecture,(response(response(lark,response(lark,response(lark,response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))))),response(X1,X1))!=response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_25,c_0_11,theory(equality)])).
cnf(c_0_97,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_11,c_0_11,theory(equality)])).
cnf(c_0_98,plain,(response(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_11,theory(equality)])).
cnf(c_0_99,negated_conjecture,(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))!=response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_23,c_0_11,theory(equality)])).
cnf(c_0_100,plain,(response(X1,response(response(response(lark,response(lark,response(lark,lark))),response(X2,X2)),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))=response(response(lark,X1),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_10,c_0_11,theory(equality)])).
cnf(c_0_101,plain,(response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)=response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_1,c_0_11,theory(equality)])).
cnf(c_0_102,negated_conjecture,(response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1))!=response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),inference(pm,[status(thm)],[c_0_2,c_0_11,theory(equality)])).
cnf(c_0_103,plain,(response(response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),X2)),inference(pm,[status(thm)],[c_0_3,c_0_11,theory(equality)])).
cnf(c_0_104,plain,(response(response(response(lark,response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1))),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_11,theory(equality)])).
cnf(c_0_105,negated_conjecture,(response(response(lark,response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1))),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))!=response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_25,c_0_11,theory(equality)])).
cnf(c_0_106,plain,(response(response(response(response(lark,lark),X1),X1),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),inference(pm,[status(thm)],[c_0_13,c_0_11,theory(equality)])).
cnf(c_0_107,plain,(response(X1,response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),response(X2,X2)))=response(response(lark,X1),response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),inference(pm,[status(thm)],[c_0_1,c_0_11,theory(equality)])).
cnf(c_0_108,plain,(response(response(X1,X1),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),response(X2,X2)))=response(response(response(lark,lark),X1),response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),inference(pm,[status(thm)],[c_0_3,c_0_11,theory(equality)])).
cnf(c_0_109,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),response(X2,X2))),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))=response(response(lark,X1),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_4,c_0_11,theory(equality)])).
cnf(c_0_110,plain,(response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),inference(pm,[status(thm)],[c_0_6,c_0_11,theory(equality)])).
cnf(c_0_111,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),response(X2,X2)))),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_11,c_0_11,theory(equality)])).
cnf(c_0_112,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(response(lark,lark),X2),X2))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_11,c_0_3,theory(equality)])).
cnf(c_0_113,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2)))),response(response(lark,response(X2,X2)),X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_11,c_0_4,theory(equality)])).
cnf(c_0_114,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2)))),response(response(response(lark,lark),X2),X2)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_11,c_0_10,theory(equality)])).
cnf(c_0_115,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(lark,response(X2,X2)),X2))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_11,c_0_1,theory(equality)])).
cnf(c_0_116,plain,(response(X1,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))=response(response(lark,response(lark,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_11,c_0_1,theory(equality)])).
cnf(c_0_117,plain,(response(X1,response(response(lark,response(response(response(lark,lark),lark),X2)),X2))=response(response(lark,response(lark,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_11,c_0_3,theory(equality)])).
cnf(c_0_118,plain,(response(X1,response(response(response(lark,lark),response(lark,lark)),lark))=response(response(lark,response(lark,X1)),response(lark,lark))),inference(pm,[status(thm)],[c_0_11,c_0_1,theory(equality)])).
cnf(c_0_119,plain,(response(X1,response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2))))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_11,c_0_4,theory(equality)])).
cnf(c_0_120,plain,(response(X1,response(response(lark,response(lark,response(response(lark,lark),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2))))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_11,c_0_10,theory(equality)])).
cnf(c_0_121,plain,(response(X1,response(response(lark,response(lark,response(response(lark,lark),response(response(X2,X2),response(X2,X2))))),X2))=response(response(lark,response(lark,X1)),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_11,c_0_1,theory(equality)])).
cnf(c_0_122,plain,(response(X1,response(response(response(lark,lark),response(X2,X2)),response(X2,X2)))=response(response(lark,response(lark,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_11,c_0_1,theory(equality)])).
cnf(c_0_123,plain,(response(X1,response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),response(X2,X2)))=response(response(lark,X1),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_10,c_0_36,theory(equality)])).
cnf(c_0_124,plain,(response(response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_36,theory(equality)])).
cnf(c_0_125,negated_conjecture,(response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))))),response(X1,X1))!=response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_23,c_0_36,theory(equality)])).
cnf(c_0_126,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))))),response(X2,X2)))=response(response(lark,X1),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_4,c_0_36,theory(equality)])).
cnf(c_0_127,plain,(response(response(response(lark,response(lark,response(lark,response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_36,theory(equality)])).
cnf(c_0_128,negated_conjecture,(response(response(lark,response(lark,response(lark,response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))))),response(X1,X1))!=response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_25,c_0_36,theory(equality)])).
cnf(c_0_129,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_11,c_0_36,theory(equality)])).
cnf(c_0_130,plain,(response(X1,response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X2,X2)),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_36,c_0_36,theory(equality)])).
cnf(c_0_131,plain,(response(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),response(X2,X2))=response(response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_36,theory(equality)])).
cnf(c_0_132,negated_conjecture,(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))!=response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_23,c_0_36,theory(equality)])).
cnf(c_0_133,plain,(response(X1,response(response(response(lark,response(lark,response(lark,lark))),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))=response(response(lark,X1),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_10,c_0_36,theory(equality)])).
cnf(c_0_134,plain,(response(response(lark,response(lark,response(response(lark,response(X1,X1)),X1))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))=response(response(response(lark,lark),response(X1,X1)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_6,c_0_36,theory(equality)])).
cnf(c_0_135,plain,(response(response(lark,response(lark,response(response(response(lark,lark),X1),X1))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))=response(response(response(lark,lark),response(X1,X1)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_13,c_0_36,theory(equality)])).
cnf(c_0_136,plain,(response(X1,response(response(lark,response(lark,response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))=response(response(lark,response(lark,X1)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_36,c_0_36,theory(equality)])).
cnf(c_0_137,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))=response(response(lark,response(lark,X1)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_11,c_0_36,theory(equality)])).
cnf(c_0_138,plain,(response(response(lark,response(lark,response(X1,X1))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))=response(response(response(lark,lark),X1),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_3,c_0_36,theory(equality)])).
cnf(c_0_139,plain,(response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)=response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_1,c_0_36,theory(equality)])).
cnf(c_0_140,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))=response(response(lark,X1),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_4,c_0_36,theory(equality)])).
cnf(c_0_141,plain,(response(response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(X1,X1))=response(response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),X1)),inference(pm,[status(thm)],[c_0_6,c_0_36,theory(equality)])).
cnf(c_0_142,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))!=response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_25,c_0_36,theory(equality)])).
cnf(c_0_143,negated_conjecture,(response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1))!=response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),inference(pm,[status(thm)],[c_0_2,c_0_36,theory(equality)])).
cnf(c_0_144,plain,(response(response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_3,c_0_36,theory(equality)])).
cnf(c_0_145,plain,(response(response(response(lark,response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1))),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),response(X2,X2))=response(response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_36,theory(equality)])).
cnf(c_0_146,negated_conjecture,(response(response(lark,response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1))),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))!=response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_25,c_0_36,theory(equality)])).
cnf(c_0_147,plain,(response(response(response(response(lark,lark),X1),X1),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_13,c_0_36,theory(equality)])).
cnf(c_0_148,plain,(response(X1,response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),response(X2,X2)))=response(response(lark,X1),response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_1,c_0_36,theory(equality)])).
cnf(c_0_149,plain,(response(response(X1,X1),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),response(X2,X2)))=response(response(response(lark,lark),X1),response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_3,c_0_36,theory(equality)])).
cnf(c_0_150,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),response(X2,X2))),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))=response(response(lark,X1),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_4,c_0_36,theory(equality)])).
cnf(c_0_151,plain,(response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_6,c_0_36,theory(equality)])).
cnf(c_0_152,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),response(X2,X2)))),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_11,c_0_36,theory(equality)])).
cnf(c_0_153,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),response(X2,X2))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_36,c_0_36,theory(equality)])).
cnf(c_0_154,plain,(response(response(lark,response(lark,X1)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))=response(response(lark,X1),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_1,c_0_36,theory(equality)])).
cnf(c_0_155,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(X2,X2)),response(response(response(lark,lark),X2),X2)))=response(response(lark,response(lark,X1)),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_36,c_0_3,theory(equality)])).
cnf(c_0_156,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(response(lark,response(X2,X2)),X2)),response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2))))=response(response(lark,response(lark,X1)),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_36,c_0_4,theory(equality)])).
cnf(c_0_157,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(response(response(lark,lark),X2),X2)),response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_36,c_0_10,theory(equality)])).
cnf(c_0_158,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),response(X2,X2))))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_36,c_0_11,theory(equality)])).
cnf(c_0_159,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(X2,X2)),response(response(lark,response(X2,X2)),X2)))=response(response(lark,response(lark,X1)),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_36,c_0_1,theory(equality)])).
cnf(c_0_160,plain,(response(X1,response(response(response(lark,response(lark,response(lark,lark))),response(X2,X2)),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2))))=response(response(lark,response(lark,X1)),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_36,c_0_4,theory(equality)])).
cnf(c_0_161,plain,(response(X1,response(response(response(lark,response(lark,response(lark,lark))),response(X2,X2)),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_36,c_0_10,theory(equality)])).
cnf(c_0_162,plain,(response(X1,response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X2,X2)),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2))))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_36,c_0_11,theory(equality)])).
cnf(c_0_163,plain,(response(X1,response(response(response(lark,response(lark,response(lark,lark))),X2),response(response(X2,X2),response(X2,X2))))=response(response(lark,response(lark,X1)),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_36,c_0_1,theory(equality)])).
cnf(c_0_164,plain,(response(X1,response(response(response(lark,lark),response(X2,X2)),response(X2,X2)))=response(response(lark,response(lark,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_36,c_0_1,theory(equality)])).
cnf(c_0_165,plain,(response(X1,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))=response(response(lark,response(lark,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_36,c_0_1,theory(equality)])).
cnf(c_0_166,plain,(response(X1,response(response(response(lark,lark),response(lark,response(lark,lark))),response(lark,response(lark,lark))))=response(response(lark,response(lark,X1)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_36,c_0_3,theory(equality)])).
cnf(c_0_167,plain,(response(X1,response(response(response(lark,lark),response(lark,lark)),lark))=response(response(lark,response(lark,X1)),response(lark,lark))),inference(pm,[status(thm)],[c_0_36,c_0_6,theory(equality)])).
cnf(c_0_168,plain,(response(X1,response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,lark),response(lark,lark)),lark)))),response(lark,lark)))=response(response(lark,X1),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_10,c_0_41,theory(equality)])).
cnf(c_0_169,plain,(response(response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,lark),response(lark,lark)),lark)))),response(lark,lark)),response(X1,X1))=response(response(response(lark,lark),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),X1)),inference(pm,[status(thm)],[c_0_13,c_0_41,theory(equality)])).
cnf(c_0_170,negated_conjecture,(response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,lark),response(lark,lark)),lark)))),response(lark,lark))!=response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_23,c_0_41,theory(equality)])).
cnf(c_0_171,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))))),response(lark,lark)))=response(response(lark,X1),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_4,c_0_41,theory(equality)])).
cnf(c_0_172,plain,(response(response(response(lark,response(lark,response(lark,response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))))),response(lark,lark)),response(X1,X1))=response(response(response(lark,lark),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),X1)),inference(pm,[status(thm)],[c_0_6,c_0_41,theory(equality)])).
cnf(c_0_173,negated_conjecture,(response(response(lark,response(lark,response(lark,response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))))),response(lark,lark))!=response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_25,c_0_41,theory(equality)])).
cnf(c_0_174,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(lark,lark),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))))),response(lark,lark)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_11,c_0_41,theory(equality)])).
cnf(c_0_175,plain,(response(X1,response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(lark,lark)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_36,c_0_41,theory(equality)])).
cnf(c_0_176,plain,(response(response(response(response(lark,response(lark,response(lark,lark))),response(lark,lark)),response(response(response(lark,lark),response(lark,lark)),lark)),response(X1,X1))=response(response(response(lark,lark),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),X1)),inference(pm,[status(thm)],[c_0_13,c_0_41,theory(equality)])).
cnf(c_0_177,negated_conjecture,(response(response(response(lark,response(lark,response(lark,lark))),response(lark,lark)),response(response(response(lark,lark),response(lark,lark)),lark))!=response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_23,c_0_41,theory(equality)])).
cnf(c_0_178,plain,(response(X1,response(response(response(lark,response(lark,response(lark,lark))),response(lark,lark)),response(response(response(lark,lark),response(lark,lark)),lark)))=response(response(lark,X1),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_10,c_0_41,theory(equality)])).
cnf(c_0_179,plain,(response(response(response(lark,response(lark,lark)),response(lark,lark)),X1)=response(response(response(response(lark,lark),response(lark,lark)),lark),response(X1,X1))),inference(pm,[status(thm)],[c_0_1,c_0_41,theory(equality)])).
cnf(c_0_180,negated_conjecture,(response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark))!=response(response(response(lark,lark),response(lark,lark)),lark)),inference(pm,[status(thm)],[c_0_2,c_0_41,theory(equality)])).
cnf(c_0_181,plain,(response(response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)),response(X1,X1))=response(response(response(lark,lark),response(response(response(lark,lark),response(lark,lark)),lark)),X1)),inference(pm,[status(thm)],[c_0_3,c_0_41,theory(equality)])).
cnf(c_0_182,plain,(response(response(response(lark,response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark))),response(response(response(lark,lark),response(lark,lark)),lark)),response(X1,X1))=response(response(response(lark,lark),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),X1)),inference(pm,[status(thm)],[c_0_6,c_0_41,theory(equality)])).
cnf(c_0_183,negated_conjecture,(response(response(lark,response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark))),response(response(response(lark,lark),response(lark,lark)),lark))!=response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_25,c_0_41,theory(equality)])).
cnf(c_0_184,plain,(response(response(response(response(lark,lark),X1),X1),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))=response(response(response(lark,lark),response(X1,X1)),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_13,c_0_41,theory(equality)])).
cnf(c_0_185,plain,(response(X1,response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))=response(response(lark,X1),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_1,c_0_41,theory(equality)])).
cnf(c_0_186,plain,(response(response(X1,X1),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))=response(response(response(lark,lark),X1),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_3,c_0_41,theory(equality)])).
cnf(c_0_187,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark))),response(response(response(lark,lark),response(lark,lark)),lark)))=response(response(lark,X1),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_4,c_0_41,theory(equality)])).
cnf(c_0_188,plain,(response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))=response(response(response(lark,lark),response(X1,X1)),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_6,c_0_41,theory(equality)])).
cnf(c_0_189,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))),response(response(response(lark,lark),response(lark,lark)),lark)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_11,c_0_41,theory(equality)])).
cnf(c_0_190,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(response(response(lark,lark),response(lark,lark)),lark)),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_36,c_0_41,theory(equality)])).
cnf(c_0_191,plain,(response(X1,response(response(lark,response(lark,response(response(lark,lark),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),response(X2,X2)))=response(response(lark,X1),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_10,c_0_116,theory(equality)])).
cnf(c_0_192,plain,(response(response(response(lark,response(lark,response(response(lark,lark),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_116,theory(equality)])).
cnf(c_0_193,negated_conjecture,(response(response(lark,response(lark,response(response(lark,lark),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)))),response(X1,X1))!=response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_23,c_0_116,theory(equality)])).
cnf(c_0_194,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))))),response(X2,X2)))=response(response(lark,X1),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_4,c_0_116,theory(equality)])).
cnf(c_0_195,plain,(response(response(response(lark,response(lark,response(lark,response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_116,theory(equality)])).
cnf(c_0_196,negated_conjecture,(response(response(lark,response(lark,response(lark,response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))))),response(X1,X1))!=response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_25,c_0_116,theory(equality)])).
cnf(c_0_197,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_11,c_0_116,theory(equality)])).
cnf(c_0_198,plain,(response(X1,response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X2,X2)),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_36,c_0_116,theory(equality)])).
cnf(c_0_199,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X2,X2))),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_116,c_0_116,theory(equality)])).
cnf(c_0_200,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(lark,response(lark,lark)),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_116,c_0_116,theory(equality)])).
cnf(c_0_201,plain,(response(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_13,c_0_116,theory(equality)])).
cnf(c_0_202,negated_conjecture,(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))!=response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_23,c_0_116,theory(equality)])).
cnf(c_0_203,plain,(response(X1,response(response(response(lark,response(lark,response(lark,lark))),response(X2,X2)),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))=response(response(lark,X1),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_10,c_0_116,theory(equality)])).
cnf(c_0_204,plain,(response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)=response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(X2,X2))),inference(pm,[status(thm)],[c_0_1,c_0_116,theory(equality)])).
cnf(c_0_205,negated_conjecture,(response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1))!=response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)),inference(pm,[status(thm)],[c_0_2,c_0_116,theory(equality)])).
cnf(c_0_206,plain,(response(response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1)),response(X2,X2))=response(response(response(lark,lark),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)),X2)),inference(pm,[status(thm)],[c_0_3,c_0_116,theory(equality)])).
cnf(c_0_207,plain,(response(response(response(lark,response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1))),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)),response(X2,X2))=response(response(response(lark,lark),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_6,c_0_116,theory(equality)])).
cnf(c_0_208,negated_conjecture,(response(response(lark,response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1))),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))!=response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_25,c_0_116,theory(equality)])).
cnf(c_0_209,plain,(response(response(response(response(lark,lark),X1),X1),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),inference(pm,[status(thm)],[c_0_13,c_0_116,theory(equality)])).
cnf(c_0_210,plain,(response(X1,response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),response(X2,X2)))=response(response(lark,X1),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),inference(pm,[status(thm)],[c_0_1,c_0_116,theory(equality)])).
cnf(c_0_211,plain,(response(response(X1,X1),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),response(X2,X2)))=response(response(response(lark,lark),X1),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),inference(pm,[status(thm)],[c_0_3,c_0_116,theory(equality)])).
cnf(c_0_212,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),response(X2,X2))),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))=response(response(lark,X1),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_4,c_0_116,theory(equality)])).
cnf(c_0_213,plain,(response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),response(X2,X2)))=response(response(response(lark,lark),response(X1,X1)),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),inference(pm,[status(thm)],[c_0_6,c_0_116,theory(equality)])).
cnf(c_0_214,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),response(X2,X2)))),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_11,c_0_116,theory(equality)])).
cnf(c_0_215,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),response(X2,X2))))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,response(lark,lark)),X2)),X2),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2)))),inference(pm,[status(thm)],[c_0_36,c_0_116,theory(equality)])).
cnf(c_0_216,plain,(response(response(lark,response(lark,X1)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))=response(response(lark,X1),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_4,c_0_116,theory(equality)])).
cnf(c_0_217,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X2,X2))),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_116,c_0_36,theory(equality)])).
cnf(c_0_218,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(lark,response(lark,lark)))),response(lark,lark))),response(response(response(lark,lark),response(lark,lark)),lark)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_116,c_0_41,theory(equality)])).
cnf(c_0_219,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(lark,lark))),response(X2,X2))),response(response(lark,response(X2,X2)),X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_116,c_0_4,theory(equality)])).
cnf(c_0_220,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(lark,lark))),response(X2,X2))),response(response(response(lark,lark),X2),X2)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_116,c_0_10,theory(equality)])).
cnf(c_0_221,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X2,X2))),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_116,c_0_11,theory(equality)])).
cnf(c_0_222,plain,(response(X1,response(response(lark,response(response(lark,response(lark,response(lark,lark))),X2)),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_116,c_0_1,theory(equality)])).
cnf(c_0_223,plain,(response(X1,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))=response(response(lark,response(lark,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_116,c_0_1,theory(equality)])).
cnf(c_0_224,plain,(response(X1,response(response(response(lark,lark),response(lark,response(lark,lark))),response(lark,response(lark,lark))))=response(response(lark,response(lark,X1)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_116,c_0_1,theory(equality)])).
cnf(c_0_225,plain,(response(X1,response(response(response(lark,lark),response(lark,lark)),lark))=response(response(lark,response(lark,X1)),response(lark,lark))),inference(pm,[status(thm)],[c_0_116,c_0_4,theory(equality)])).
cnf(c_0_226,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(lark,response(lark,lark)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,response(lark,lark)),X2),response(X2,X2)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2))))),inference(pm,[status(thm)],[c_0_116,c_0_36,theory(equality)])).
cnf(c_0_227,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(lark,response(lark,lark)),response(response(response(lark,lark),response(lark,lark)),lark))))),response(lark,lark)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))),inference(pm,[status(thm)],[c_0_116,c_0_41,theory(equality)])).
cnf(c_0_228,plain,(response(X1,response(response(lark,response(lark,response(response(lark,response(lark,lark)),response(response(lark,response(X2,X2)),X2)))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(X2,X2)),X2),response(response(lark,response(X2,X2)),X2)))),inference(pm,[status(thm)],[c_0_116,c_0_4,theory(equality)])).
cnf(c_0_229,plain,(response(X1,response(response(lark,response(lark,response(response(lark,response(lark,lark)),response(response(response(lark,lark),X2),X2)))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(response(lark,lark),X2),X2),response(response(response(lark,lark),X2),X2)))),inference(pm,[status(thm)],[c_0_116,c_0_10,theory(equality)])).
cnf(c_0_230,plain,(response(X1,response(response(lark,response(lark,response(lark,response(response(lark,response(lark,lark)),response(response(lark,response(response(lark,lark),response(X2,X2))),X2))))),response(X2,X2)))=response(response(lark,response(lark,X1)),response(response(response(lark,response(response(lark,lark),response(X2,X2))),X2),response(response(lark,response(response(lark,lark),response(X2,X2))),X2)))),inference(pm,[status(thm)],[c_0_116,c_0_11,theory(equality)])).
cnf(c_0_231,plain,(response(X1,response(response(lark,response(lark,response(response(lark,response(lark,lark)),response(X2,X2)))),X2))=response(response(lark,response(lark,X1)),response(response(X2,X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_116,c_0_1,theory(equality)])).
cnf(c_0_232,plain,(response(X1,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))=response(response(lark,response(lark,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_116,c_0_1,theory(equality)])).
cnf(c_0_233,negated_conjecture,(response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),response(response(response(lark,lark),X1),X1))!=response(response(lark,response(response(X1,X1),response(X1,X1))),response(X1,X1))),inference(pm,[status(thm)],[c_0_24,c_0_3,theory(equality)])).
cnf(c_0_234,negated_conjecture,(response(response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1)))!=response(response(lark,response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_24,c_0_36,theory(equality)])).
cnf(c_0_235,negated_conjecture,(response(response(response(lark,lark),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))!=response(response(lark,response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_24,c_0_41,theory(equality)])).
cnf(c_0_236,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1)))!=response(response(lark,response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_24,c_0_4,theory(equality)])).
cnf(c_0_237,negated_conjecture,(response(response(response(lark,lark),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1)))!=response(response(lark,response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_24,c_0_10,theory(equality)])).
cnf(c_0_238,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1)))!=response(response(lark,response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_24,c_0_11,theory(equality)])).
cnf(c_0_239,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1)))!=response(response(lark,response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_24,c_0_116,theory(equality)])).
cnf(c_0_240,negated_conjecture,(response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),response(response(lark,response(X1,X1)),X1))!=response(response(lark,response(response(X1,X1),response(X1,X1))),response(X1,X1))),inference(pm,[status(thm)],[c_0_24,c_0_1,theory(equality)])).
cnf(c_0_241,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,lark),X1),X1)),response(response(X1,X1),response(X1,X1)))!=response(response(lark,response(response(X1,X1),response(X1,X1))),response(X1,X1))),inference(pm,[status(thm)],[c_0_24,c_0_3,theory(equality)])).
cnf(c_0_242,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1))),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))))!=response(response(lark,response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_24,c_0_36,theory(equality)])).
cnf(c_0_243,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark))),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))!=response(response(lark,response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_24,c_0_41,theory(equality)])).
cnf(c_0_244,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1))),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1)))!=response(response(lark,response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_24,c_0_4,theory(equality)])).
cnf(c_0_245,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1))),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1)))!=response(response(lark,response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_24,c_0_10,theory(equality)])).
cnf(c_0_246,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1))),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)))!=response(response(lark,response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_24,c_0_11,theory(equality)])).
cnf(c_0_247,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1))),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)))!=response(response(lark,response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_24,c_0_116,theory(equality)])).
cnf(c_0_248,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(X1,X1)),X1)),response(response(X1,X1),response(X1,X1)))!=response(response(lark,response(response(X1,X1),response(X1,X1))),response(X1,X1))),inference(pm,[status(thm)],[c_0_24,c_0_1,theory(equality)])).
cnf(c_0_249,negated_conjecture,(response(response(response(lark,response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))!=response(response(lark,response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_24,c_0_36,theory(equality)])).
cnf(c_0_250,negated_conjecture,(response(response(response(lark,response(lark,lark)),X1),response(X1,X1))!=response(response(lark,response(X1,X1)),X1)),inference(pm,[status(thm)],[c_0_24,c_0_1,theory(equality)])).
cnf(c_0_251,negated_conjecture,(response(response(response(response(lark,lark),lark),X1),response(X1,X1))!=response(response(lark,response(X1,X1)),X1)),inference(pm,[status(thm)],[c_0_24,c_0_3,theory(equality)])).
cnf(c_0_252,negated_conjecture,(response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))!=response(response(lark,response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_24,c_0_36,theory(equality)])).
cnf(c_0_253,negated_conjecture,(response(response(lark,response(response(lark,lark),response(X1,X1))),X1)!=response(response(lark,response(X1,X1)),X1)),inference(pm,[status(thm)],[c_0_24,c_0_1,theory(equality)])).
cnf(c_0_254,negated_conjecture,(response(response(response(lark,lark),response(lark,lark)),lark)!=response(response(lark,response(lark,lark)),lark)),inference(pm,[status(thm)],[c_0_24,c_0_3,theory(equality)])).
cnf(c_0_255,negated_conjecture,(response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),response(response(response(lark,lark),X1),X1))!=response(response(response(lark,lark),response(X1,X1)),response(X1,X1))),inference(pm,[status(thm)],[c_0_30,c_0_3,theory(equality)])).
cnf(c_0_256,negated_conjecture,(response(response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1)))!=response(response(response(lark,lark),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_30,c_0_36,theory(equality)])).
cnf(c_0_257,negated_conjecture,(response(response(response(lark,lark),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))!=response(response(response(lark,lark),response(response(response(lark,lark),response(lark,lark)),lark)),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_30,c_0_41,theory(equality)])).
cnf(c_0_258,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1)))!=response(response(response(lark,lark),response(response(lark,response(X1,X1)),X1)),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_30,c_0_4,theory(equality)])).
cnf(c_0_259,negated_conjecture,(response(response(response(lark,lark),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1)))!=response(response(response(lark,lark),response(response(response(lark,lark),X1),X1)),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_30,c_0_10,theory(equality)])).
cnf(c_0_260,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1)))!=response(response(response(lark,lark),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_30,c_0_11,theory(equality)])).
cnf(c_0_261,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1)))!=response(response(response(lark,lark),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_30,c_0_116,theory(equality)])).
cnf(c_0_262,negated_conjecture,(response(response(response(lark,lark),response(response(X1,X1),response(X1,X1))),response(response(lark,response(X1,X1)),X1))!=response(response(response(lark,lark),response(X1,X1)),response(X1,X1))),inference(pm,[status(thm)],[c_0_30,c_0_1,theory(equality)])).
cnf(c_0_263,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,lark),X1),X1)),response(response(X1,X1),response(X1,X1)))!=response(response(response(lark,lark),response(X1,X1)),response(X1,X1))),inference(pm,[status(thm)],[c_0_30,c_0_3,theory(equality)])).
cnf(c_0_264,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1))),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))))!=response(response(response(lark,lark),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_30,c_0_36,theory(equality)])).
cnf(c_0_265,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark))),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))!=response(response(response(lark,lark),response(response(response(lark,lark),response(lark,lark)),lark)),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_30,c_0_41,theory(equality)])).
cnf(c_0_266,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1))),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1)))!=response(response(response(lark,lark),response(response(lark,response(X1,X1)),X1)),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_30,c_0_4,theory(equality)])).
cnf(c_0_267,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1))),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1)))!=response(response(response(lark,lark),response(response(response(lark,lark),X1),X1)),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_30,c_0_10,theory(equality)])).
cnf(c_0_268,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1))),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)))!=response(response(response(lark,lark),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_30,c_0_11,theory(equality)])).
cnf(c_0_269,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1))),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)))!=response(response(response(lark,lark),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_30,c_0_116,theory(equality)])).
cnf(c_0_270,negated_conjecture,(response(response(response(lark,lark),response(response(lark,response(X1,X1)),X1)),response(response(X1,X1),response(X1,X1)))!=response(response(response(lark,lark),response(X1,X1)),response(X1,X1))),inference(pm,[status(thm)],[c_0_30,c_0_1,theory(equality)])).
cnf(c_0_271,negated_conjecture,(response(response(response(lark,response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))!=response(response(response(lark,lark),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_30,c_0_36,theory(equality)])).
cnf(c_0_272,negated_conjecture,(response(response(response(lark,response(lark,lark)),X1),response(X1,X1))!=response(response(response(lark,lark),X1),X1)),inference(pm,[status(thm)],[c_0_30,c_0_1,theory(equality)])).
cnf(c_0_273,negated_conjecture,(response(response(response(response(lark,lark),lark),X1),response(X1,X1))!=response(response(response(lark,lark),X1),X1)),inference(pm,[status(thm)],[c_0_30,c_0_3,theory(equality)])).
cnf(c_0_274,negated_conjecture,(response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))!=response(response(response(lark,lark),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_30,c_0_36,theory(equality)])).
cnf(c_0_275,negated_conjecture,(response(response(lark,response(response(lark,lark),response(X1,X1))),X1)!=response(response(response(lark,lark),X1),X1)),inference(pm,[status(thm)],[c_0_30,c_0_1,theory(equality)])).
cnf(c_0_276,negated_conjecture,(response(response(response(lark,lark),response(lark,lark)),lark)!=response(response(response(lark,lark),lark),lark)),inference(pm,[status(thm)],[c_0_30,c_0_3,theory(equality)])).
cnf(c_0_277,plain,(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))=response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_36,c_0_64,theory(equality)])).
cnf(c_0_278,plain,(response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)=response(response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),X2)),inference(pm,[status(thm)],[c_0_1,c_0_64,theory(equality)])).
cnf(c_0_279,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(response(lark,lark),X2),X2))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_64,c_0_3,theory(equality)])).
cnf(c_0_280,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_64,c_0_36,theory(equality)])).
cnf(c_0_281,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_64,c_0_41,theory(equality)])).
cnf(c_0_282,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(X2,X2)),X2))),inference(pm,[status(thm)],[c_0_64,c_0_4,theory(equality)])).
cnf(c_0_283,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,lark),X2),X2))),inference(pm,[status(thm)],[c_0_64,c_0_10,theory(equality)])).
cnf(c_0_284,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),inference(pm,[status(thm)],[c_0_64,c_0_11,theory(equality)])).
cnf(c_0_285,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),inference(pm,[status(thm)],[c_0_64,c_0_116,theory(equality)])).
cnf(c_0_286,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(X2,X2)),X2))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_64,c_0_1,theory(equality)])).
cnf(c_0_287,plain,(response(response(response(lark,response(response(lark,lark),response(response(response(lark,lark),X1),X1))),response(X1,X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_3,theory(equality)])).
cnf(c_0_288,plain,(response(response(response(lark,response(response(lark,lark),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1)))),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_36,theory(equality)])).
cnf(c_0_289,plain,(response(response(response(lark,response(response(lark,lark),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))),response(response(response(lark,lark),response(lark,lark)),lark)),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),X1)),inference(pm,[status(thm)],[c_0_64,c_0_41,theory(equality)])).
cnf(c_0_290,plain,(response(response(response(lark,response(response(lark,lark),response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1)))),response(response(lark,response(X1,X1)),X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_4,theory(equality)])).
cnf(c_0_291,plain,(response(response(response(lark,response(response(lark,lark),response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1)))),response(response(response(lark,lark),X1),X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_10,theory(equality)])).
cnf(c_0_292,plain,(response(response(response(lark,response(response(lark,lark),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1)))),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_11,theory(equality)])).
cnf(c_0_293,plain,(response(response(response(lark,response(response(lark,lark),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1)))),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_116,theory(equality)])).
cnf(c_0_294,plain,(response(response(response(lark,response(response(lark,lark),response(response(lark,response(X1,X1)),X1))),response(X1,X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_1,theory(equality)])).
cnf(c_0_295,plain,(response(response(response(lark,response(response(lark,response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),X1)),inference(pm,[status(thm)],[c_0_64,c_0_36,theory(equality)])).
cnf(c_0_296,plain,(response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_64,c_0_1,theory(equality)])).
cnf(c_0_297,plain,(response(response(response(lark,response(response(response(lark,lark),lark),X1)),X1),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_64,c_0_3,theory(equality)])).
cnf(c_0_298,plain,(response(response(response(response(lark,lark),response(lark,lark)),lark),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(lark,lark)),X1)),inference(pm,[status(thm)],[c_0_64,c_0_1,theory(equality)])).
cnf(c_0_299,plain,(response(response(response(lark,response(lark,response(lark,response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))))))),response(X1,X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_36,theory(equality)])).
cnf(c_0_300,plain,(response(response(response(lark,response(lark,response(lark,response(response(lark,lark),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))))),response(lark,lark)),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),X1)),inference(pm,[status(thm)],[c_0_64,c_0_41,theory(equality)])).
cnf(c_0_301,plain,(response(response(response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))))),response(X1,X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_4,theory(equality)])).
cnf(c_0_302,plain,(response(response(response(lark,response(lark,response(response(lark,lark),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))))),response(X1,X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_10,theory(equality)])).
cnf(c_0_303,plain,(response(response(response(lark,response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)))))),response(X1,X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_11,theory(equality)])).
cnf(c_0_304,plain,(response(response(response(lark,response(lark,response(lark,response(response(lark,lark),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)))))),response(X1,X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_116,theory(equality)])).
cnf(c_0_305,plain,(response(response(response(lark,response(lark,response(response(lark,lark),response(response(X1,X1),response(X1,X1))))),X1),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_64,c_0_1,theory(equality)])).
cnf(c_0_306,plain,(response(response(response(response(lark,lark),response(X1,X1)),response(X1,X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_64,c_0_1,theory(equality)])).
cnf(c_0_307,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))!=response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_2,c_0_73,theory(equality)])).
cnf(c_0_308,plain,(response(response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(X1,X1))=response(response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),X1)),inference(pm,[status(thm)],[c_0_3,c_0_73,theory(equality)])).
cnf(c_0_309,plain,(response(response(response(lark,response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(X1,X1))=response(response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),X1)),inference(pm,[status(thm)],[c_0_6,c_0_73,theory(equality)])).
cnf(c_0_310,plain,(response(response(response(response(lark,response(lark,lark)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),X1)),inference(pm,[status(thm)],[c_0_73,c_0_73,theory(equality)])).
cnf(c_0_311,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))!=response(response(lark,response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_24,c_0_73,theory(equality)])).
cnf(c_0_312,negated_conjecture,(response(response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))!=response(response(lark,response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_24,c_0_73,theory(equality)])).
cnf(c_0_313,negated_conjecture,(response(response(lark,response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))!=response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_25,c_0_73,theory(equality)])).
cnf(c_0_314,negated_conjecture,(response(response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))!=response(response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_30,c_0_73,theory(equality)])).
cnf(c_0_315,negated_conjecture,(response(response(response(lark,lark),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))!=response(response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_30,c_0_73,theory(equality)])).
cnf(c_0_316,plain,(response(response(response(lark,response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),X1)),inference(pm,[status(thm)],[c_0_64,c_0_73,theory(equality)])).
cnf(c_0_317,plain,(response(response(response(response(lark,lark),X1),X1),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))=response(response(response(lark,lark),response(X1,X1)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_13,c_0_73,theory(equality)])).
cnf(c_0_318,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))=response(response(lark,X1),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_1,c_0_73,theory(equality)])).
cnf(c_0_319,plain,(response(response(X1,X1),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))=response(response(response(lark,lark),X1),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_3,c_0_73,theory(equality)])).
cnf(c_0_320,plain,(response(X1,response(response(lark,response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))=response(response(lark,X1),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))),inference(pm,[status(thm)],[c_0_4,c_0_73,theory(equality)])).
cnf(c_0_321,plain,(response(response(response(lark,response(X1,X1)),X1),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))=response(response(response(lark,lark),response(X1,X1)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_6,c_0_73,theory(equality)])).
cnf(c_0_322,plain,(response(X1,response(response(lark,response(response(lark,lark),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))),inference(pm,[status(thm)],[c_0_11,c_0_73,theory(equality)])).
cnf(c_0_323,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_73,c_0_73,theory(equality)])).
cnf(c_0_324,plain,(response(X1,response(response(response(lark,response(lark,lark)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))=response(response(lark,response(lark,X1)),response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))))),inference(pm,[status(thm)],[c_0_36,c_0_73,theory(equality)])).
cnf(c_0_325,plain,(response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_64,c_0_73,theory(equality)])).
cnf(c_0_326,plain,(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))=response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_36,c_0_73,theory(equality)])).
cnf(c_0_327,plain,(response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)=response(response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_1,c_0_73,theory(equality)])).
cnf(c_0_328,plain,(response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),X1)=response(response(response(lark,lark),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),X1)),inference(pm,[status(thm)],[c_0_3,c_0_73,theory(equality)])).
cnf(c_0_329,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,lark),X2),X2))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_73,c_0_3,theory(equality)])).
cnf(c_0_330,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,response(lark,lark)),X2),response(X2,X2)))),inference(pm,[status(thm)],[c_0_73,c_0_36,theory(equality)])).
cnf(c_0_331,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_73,c_0_41,theory(equality)])).
cnf(c_0_332,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(lark,response(response(lark,response(X2,X2)),X2)),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(X2,X2)),X2))),inference(pm,[status(thm)],[c_0_73,c_0_4,theory(equality)])).
cnf(c_0_333,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(lark,response(response(response(lark,lark),X2),X2)),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,lark),X2),X2))),inference(pm,[status(thm)],[c_0_73,c_0_10,theory(equality)])).
cnf(c_0_334,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(response(lark,lark),response(X2,X2))),X2))),inference(pm,[status(thm)],[c_0_73,c_0_11,theory(equality)])).
cnf(c_0_335,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),response(X2,X2)))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(response(lark,response(lark,lark)),X2)),X2))),inference(pm,[status(thm)],[c_0_73,c_0_116,theory(equality)])).
cnf(c_0_336,plain,(response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(lark,response(X2,X2)),X2))=response(response(response(lark,response(lark,lark)),response(X1,X1)),response(X2,X2))),inference(pm,[status(thm)],[c_0_73,c_0_1,theory(equality)])).
cnf(c_0_337,plain,(response(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,lark),X1),X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_3,theory(equality)])).
cnf(c_0_338,plain,(response(response(response(response(lark,response(lark,lark)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_36,theory(equality)])).
cnf(c_0_339,plain,(response(response(response(response(lark,response(lark,lark)),response(response(response(lark,lark),response(lark,lark)),lark)),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark))),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),X1)),inference(pm,[status(thm)],[c_0_73,c_0_41,theory(equality)])).
cnf(c_0_340,plain,(response(response(response(response(lark,response(lark,lark)),response(response(lark,response(X1,X1)),X1)),response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_4,theory(equality)])).
cnf(c_0_341,plain,(response(response(response(response(lark,response(lark,lark)),response(response(response(lark,lark),X1),X1)),response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_10,theory(equality)])).
cnf(c_0_342,plain,(response(response(response(response(lark,response(lark,lark)),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_11,theory(equality)])).
cnf(c_0_343,plain,(response(response(response(response(lark,response(lark,lark)),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_116,theory(equality)])).
cnf(c_0_344,plain,(response(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(X1,X1)),X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_1,theory(equality)])).
cnf(c_0_345,plain,(response(response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X1,X1)),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_36,theory(equality)])).
cnf(c_0_346,plain,(response(response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(lark,lark)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),X1)),inference(pm,[status(thm)],[c_0_73,c_0_41,theory(equality)])).
cnf(c_0_347,plain,(response(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_4,theory(equality)])).
cnf(c_0_348,plain,(response(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_10,theory(equality)])).
cnf(c_0_349,plain,(response(response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X1,X1)),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_11,theory(equality)])).
cnf(c_0_350,plain,(response(response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X1,X1)),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_116,theory(equality)])).
cnf(c_0_351,plain,(response(response(response(response(lark,response(lark,response(lark,lark))),X1),response(response(X1,X1),response(X1,X1))),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(response(X1,X1),response(X1,X1))),X2)),inference(pm,[status(thm)],[c_0_73,c_0_1,theory(equality)])).
cnf(c_0_352,plain,(response(response(response(response(lark,lark),response(X1,X1)),response(X1,X1)),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_73,c_0_1,theory(equality)])).
cnf(c_0_353,plain,(response(response(response(lark,response(lark,response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),X1)),inference(pm,[status(thm)],[c_0_73,c_0_36,theory(equality)])).
cnf(c_0_354,plain,(response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(X2,X2))=response(response(response(lark,response(lark,lark)),response(X1,X1)),X2)),inference(pm,[status(thm)],[c_0_73,c_0_1,theory(equality)])).
cnf(c_0_355,plain,(response(response(response(response(lark,lark),response(lark,response(lark,lark))),response(lark,response(lark,lark))),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),X1)),inference(pm,[status(thm)],[c_0_73,c_0_3,theory(equality)])).
cnf(c_0_356,plain,(response(response(response(response(lark,lark),response(lark,lark)),lark),response(X1,X1))=response(response(response(lark,response(lark,lark)),response(lark,lark)),X1)),inference(pm,[status(thm)],[c_0_73,c_0_6,theory(equality)])).
cnf(c_0_357,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))!=response(response(lark,response(response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),inference(pm,[status(thm)],[c_0_250,c_0_73,theory(equality)])).
cnf(c_0_358,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(response(lark,lark),X1),X1))!=response(response(lark,response(response(X1,X1),response(X1,X1))),response(X1,X1))),inference(pm,[status(thm)],[c_0_250,c_0_3,theory(equality)])).
cnf(c_0_359,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))),response(response(lark,response(lark,response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(X1,X1)))!=response(response(lark,response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_250,c_0_36,theory(equality)])).
cnf(c_0_360,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(response(response(lark,lark),response(lark,lark)),lark)),response(response(lark,response(lark,response(response(response(lark,lark),response(lark,lark)),lark))),response(lark,lark)))!=response(response(lark,response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_250,c_0_41,theory(equality)])).
cnf(c_0_361,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(response(lark,response(X1,X1)),X1)),response(response(lark,response(response(lark,response(X1,X1)),X1)),response(X1,X1)))!=response(response(lark,response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_250,c_0_4,theory(equality)])).
cnf(c_0_362,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(response(response(lark,lark),X1),X1)),response(response(lark,response(response(response(lark,lark),X1),X1)),response(X1,X1)))!=response(response(lark,response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_250,c_0_10,theory(equality)])).
cnf(c_0_363,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)),response(response(lark,response(lark,response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(X1,X1)))!=response(response(lark,response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_250,c_0_11,theory(equality)])).
cnf(c_0_364,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)),response(response(lark,response(lark,response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(X1,X1)))!=response(response(lark,response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_250,c_0_116,theory(equality)])).
cnf(c_0_365,negated_conjecture,(response(response(response(lark,response(lark,lark)),response(X1,X1)),response(response(lark,response(X1,X1)),X1))!=response(response(lark,response(response(X1,X1),response(X1,X1))),response(X1,X1))),inference(pm,[status(thm)],[c_0_250,c_0_1,theory(equality)])).
cnf(c_0_366,negated_conjecture,(response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X1,X1)),response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1))))!=response(response(lark,response(response(response(response(lark,response(lark,lark)),X1),response(X1,X1)),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),response(response(response(lark,response(lark,lark)),X1),response(X1,X1)))),inference(pm,[status(thm)],[c_0_250,c_0_36,theory(equality)])).
cnf(c_0_367,negated_conjecture,(response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(lark,lark)),response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark)))!=response(response(lark,response(response(response(response(lark,lark),response(lark,lark)),lark),response(response(response(lark,lark),response(lark,lark)),lark))),response(response(response(lark,lark),response(lark,lark)),lark))),inference(pm,[status(thm)],[c_0_250,c_0_41,theory(equality)])).
cnf(c_0_368,negated_conjecture,(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1)))!=response(response(lark,response(response(response(lark,response(X1,X1)),X1),response(response(lark,response(X1,X1)),X1))),response(response(lark,response(X1,X1)),X1))),inference(pm,[status(thm)],[c_0_250,c_0_4,theory(equality)])).
cnf(c_0_369,negated_conjecture,(response(response(response(lark,response(lark,response(lark,lark))),response(X1,X1)),response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1)))!=response(response(lark,response(response(response(response(lark,lark),X1),X1),response(response(response(lark,lark),X1),X1))),response(response(response(lark,lark),X1),X1))),inference(pm,[status(thm)],[c_0_250,c_0_10,theory(equality)])).
cnf(c_0_370,negated_conjecture,(response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X1,X1)),response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1)))!=response(response(lark,response(response(response(lark,response(response(lark,lark),response(X1,X1))),X1),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),response(response(lark,response(response(lark,lark),response(X1,X1))),X1))),inference(pm,[status(thm)],[c_0_250,c_0_11,theory(equality)])).
cnf(c_0_371,negated_conjecture,(response(response(response(lark,response(lark,response(lark,response(lark,lark)))),response(X1,X1)),response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)))!=response(response(lark,response(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),response(response(lark,response(response(lark,response(lark,lark)),X1)),X1))),inference(pm,[status(thm)],[c_0_250,c_0_116,theory(equality)])).
cnf(c_0_372,negated_conjecture,(response(response(response(lark,response(lark,response(lark,lark))),X1),response(response(X1,X1),response(X1,X1)))!=response(response(lark,response(response(X1,X1),response(X1,X1))),response(X1,X1))),inference(pm,[status(thm)],[c_0_250,c_0_1,theory(equality)])).
cnf(c_0_373,negated_conjecture,(response(response(response(lark,lark),response(X1,X1)),response(X1,X1))!=response(response(lark,response(X1,X1)),X1)),inference(pm,[status(thm)],[c_0_250,c_0_1,theory(equality)])).
cnf(c_0_374,negated_conjecture,(response(response(lark,response(lark,response(response(lark,response(lark,lark)),response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))!=response(response(lark,response(response(response(lark,response(lark,lark)),response(lark,response(lark,lark))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),response(response(lark,response(lark,lark)),response(lark,response(lark,lark))))),inference(pm,[status(thm)],[c_0_250,c_0_36,theory(equality)])).
cnf(c_0_375,negated_conjecture,(response(response(lark,response(response(lark,response(lark,lark)),X1)),X1)!=response(response(lark,response(X1,X1)),X1)),inference(pm,[status(thm)],[c_0_250,c_0_1,theory(equality)])).
cnf(c_0_376,negated_conjecture,(response(response(response(lark,lark),response(lark,response(lark,lark))),response(lark,response(lark,lark)))!=response(response(lark,response(response(lark,response(lark,lark)),response(lark,response(lark,lark)))),response(lark,response(lark,lark)))),inference(pm,[status(thm)],[c_0_250,c_0_3,theory(equality)])).
cnf(c_0_377,negated_conjecture,(response(response(response(lark,lark),response(lark,lark)),lark)!=response(response(lark,response(lark,lark)),lark)),inference(pm,[status(thm)],[c_0_250,c_0_6,theory(equality)])).
cnf(c_0_378,negated_conjecture,($false),inference(cn,[status(thm)],[c_0_376, theory(equality,[symmetry])])).
cnf(c_0_379,negated_conjecture,($false), c_0_378,['proof']).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(response(response(lark,X1),X2)=response(X1,response(X2,X2))), file('<stdin>', lark_exists)).
cnf(c_0_1,negated_conjecture,(response(X1,X1)!=X1), file('<stdin>', prove_the_bird_exists)).
cnf(c_0_2,axiom,(response(response(lark,X1),X2)=response(X1,response(X2,X2))), c_0_0).
cnf(c_0_3,negated_conjecture,(response(X1,X1)!=X1), c_0_1).
cnf(c_0_4,plain,(response(response(lark,X1),X2)=response(X1,response(X2,X2))), c_0_2).
cnf(c_0_5,negated_conjecture,(response(X1,X1)!=X1), c_0_3).
cnf(c_0_6,plain,(response(X1,response(response(lark,response(X2,X2)),X2))=response(response(lark,X1),response(X2,X2))), inference(pm,[status(thm)],[c_0_4, c_0_4, theory(equality)])).
cnf(c_0_7,negated_conjecture,(response(response(response(lark,lark),response(X1,X1)),response(X1,X1))!=response(response(lark,response(X1,X1)),X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_6, theory(equality)]), c_0_6, theory(equality)])).
cnf(c_0_8,negated_conjecture,(response(response(response(lark,response(lark,lark)),X1),response(X1,X1))!=response(response(lark,response(X1,X1)),X1)), inference(pm,[status(thm)],[c_0_7, c_0_4, theory(equality)])).
cnf(c_0_9,plain,(response(response(response(lark,lark),X1),X2)=response(response(X1,X1),response(X2,X2))), inference(pm,[status(thm)],[c_0_4, c_0_4, theory(equality)])).
cnf(c_0_10,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_9, theory(equality)]), c_0_4, theory(equality)]), theory(equality,[symmetry])]), ['proof']).
# SZS output end CNFRefutation.
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 2
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 2
# Processed clauses                    : 48
# ...of these trivial                  : 18
# ...subsumed                          : 12
# ...remaining for further processing  : 18
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 375
# ...of the previous two non-trivial   : 328
# Contextual simplify-reflections      : 0
# Paramodulations                      : 375
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 18
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 282
# ...number of literals in the above   : 282
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 17
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
