% Problem    : SYN547-1 : TPTP v6.1.0. Released v2.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n118.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:37:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____011_C07_F1_PI_AE_SP_S0V
# and selection function PSelectComplexExceptRRHorn.
#
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___207_C18_F1_AE_Q7_CS_SP_PI_CO_PS_S00
# and selection function NoSelection.
#
# Presaturation interreduction done
# No success with AutoSched1
# Trying AutoSched2 for 23 seconds
# AutoSched2-Mode selected heuristic H_____011_C18_F1_PI_SE_SP_S2S
# and selection function SelectNewComplexAHP.
#
# No success with AutoSched2
# Trying AutoSched3 for 17 seconds
# AutoSched3-Mode selected heuristic H_____011_C07_F1_PI_AE_Q4_CS_SP_PS_SOV
# and selection function PSelectComplexExceptRRHorn.
#
# Presaturation interreduction done

# No proof found!
# SZS status Satisfiable
# SZS output start Saturation.
cnf(c_0_0,negated_conjecture,(ndr1_0|~ssSkC26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause2)).
cnf(c_0_1,negated_conjecture,(ndr1_0|ssSkC26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause6)).
cnf(c_0_2,negated_conjecture,(c8_0|~c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause41)).
cnf(c_0_3,negated_conjecture,(~c8_0|~c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause109)).
cnf(c_0_4,negated_conjecture,(c5_1(X1)|c8_2(X2,X3)|c2_2(X2,X3)|c1_2(X2,a557)|c7_1(X2)|c5_0|~c6_1(X1)|~c3_1(X1)|~ndr1_0|~ndr1_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause419)).
cnf(c_0_5,negated_conjecture,(c6_1(a565)|ssSkC20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause48)).
cnf(c_0_6,negated_conjecture,(c3_1(a565)|ssSkC20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause49)).
cnf(c_0_7,negated_conjecture,(ndr1_1(X1)|ssSkC8|~ssSkP3(X1)|~ssSkP4(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause265)).
cnf(c_0_8,negated_conjecture,(ndr1_1(X1)|ssSkP3(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause71)).
cnf(c_0_9,negated_conjecture,(ndr1_1(X1)|ssSkP4(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause70)).
cnf(c_0_10,negated_conjecture,(c5_0|c8_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause110)).
cnf(c_0_11,negated_conjecture,(c8_0|c5_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause113)).
cnf(c_0_12,negated_conjecture,(c3_2(X1,X2)|c6_2(X1,X2)|c3_2(X1,X3)|c5_1(X1)|c4_1(a531)|~c7_2(X1,X2)|~ndr1_1(X1)|~c8_2(X1,X3)|~c1_2(X1,X3)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause442)).
cnf(c_0_13,negated_conjecture,(c7_2(X1,X2)|c6_2(X1,X2)|c1_2(X1,X2)|c1_1(X1)|c5_1(a501)|c7_0|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause377)).
cnf(c_0_14,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|c4_1(a458)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause378)).
cnf(c_0_15,negated_conjecture,(c7_0|c2_0|c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause42)).
cnf(c_0_16,negated_conjecture,(c2_1(a565)|ssSkC20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause50)).
cnf(c_0_17,negated_conjecture,(c6_2(a552,X1)|c8_2(a552,X1)|c2_2(a552,X1)|ssSkC18|~ndr1_1(a552)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause328)).
cnf(c_0_18,negated_conjecture,(c5_2(a592,X1)|c6_2(a592,X1)|ssSkC24|~c1_2(a592,X1)|~ndr1_1(a592)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause346)).
cnf(c_0_19,negated_conjecture,(c5_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause43)).
cnf(c_0_20,negated_conjecture,(c7_2(X1,a526)|ssSkP6(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause124)).
cnf(c_0_21,negated_conjecture,(ssSkP6(X1)|~c3_2(X1,a526)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause186)).
cnf(c_0_22,negated_conjecture,(c2_1(a517)|ssSkC11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause65)).
cnf(c_0_23,negated_conjecture,(c7_2(X1,a515)|ssSkP5(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause126)).
cnf(c_0_24,negated_conjecture,(ssSkP5(X1)|~c3_2(X1,a515)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause188)).
cnf(c_0_25,negated_conjecture,(c4_2(X1,a487)|c5_1(X1)|ssSkC5|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause278)).
cnf(c_0_26,negated_conjecture,(c7_2(a542,X1)|ssSkC16|~c4_2(a542,X1)|~ndr1_1(a542)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause296)).
cnf(c_0_27,negated_conjecture,(ndr1_1(a542)|ssSkC16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause56)).
cnf(c_0_28,negated_conjecture,(c3_2(a592,X1)|ssSkC24|~c5_2(a592,X1)|~ndr1_1(a592)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause295)).
cnf(c_0_29,negated_conjecture,(c7_2(a600,X1)|ssSkC26|~c1_2(a600,X1)|~ndr1_1(a600)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause294)).
cnf(c_0_30,negated_conjecture,(c3_2(X1,X2)|c2_2(X1,X2)|c4_2(X1,a527)|ssSkC12|~ssSkP6(X1)|~c1_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause373)).
cnf(c_0_31,negated_conjecture,(ndr1_1(X1)|ssSkP6(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause62)).
cnf(c_0_32,negated_conjecture,(ndr1_1(a512)|ssSkC9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause68)).
cnf(c_0_33,negated_conjecture,(ndr1_1(a520)|c5_0|c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause93)).
cnf(c_0_34,negated_conjecture,(c2_1(a542)|ssSkC16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause57)).
cnf(c_0_35,negated_conjecture,(c3_2(X1,X2)|c1_2(X1,a493)|ssSkC7|~ssSkP2(X1)|~c5_2(X1,X2)|~c8_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause382)).
cnf(c_0_36,negated_conjecture,(c4_2(X1,a492)|ssSkP2(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause130)).
cnf(c_0_37,negated_conjecture,(ssSkP2(X1)|~c7_2(X1,a492)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause197)).
cnf(c_0_38,negated_conjecture,(c2_2(X1,X2)|c1_2(X1,X2)|c5_2(X1,X3)|c4_2(X1,X3)|c8_1(X1)|c1_1(a573)|c1_0|~c3_2(X1,X2)|~ndr1_1(X1)|~c7_2(X1,X3)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause440)).
cnf(c_0_39,negated_conjecture,(c2_2(X1,X2)|c1_2(X1,X2)|c5_2(X1,X3)|c4_2(X1,X3)|c8_1(X1)|c3_1(a573)|c1_0|~c3_2(X1,X2)|~ndr1_1(X1)|~c7_2(X1,X3)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause441)).
cnf(c_0_40,negated_conjecture,(c2_2(X1,X2)|c5_2(X1,X2)|c2_2(X1,a572)|c1_0|c5_0|~ndr1_1(X1)|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause363)).
cnf(c_0_41,negated_conjecture,(c6_1(a552)|ssSkC18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause54)).
cnf(c_0_42,negated_conjecture,(c6_1(a512)|ssSkC9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause69)).
cnf(c_0_43,negated_conjecture,(c3_1(X1)|c8_1(X1)|c1_2(X2,X3)|c3_2(X2,X3)|c8_1(X2)|c4_1(X2)|c1_0|~c1_1(X1)|~ndr1_0|~c8_2(X2,X3)|~ndr1_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause436)).
cnf(c_0_44,negated_conjecture,(c3_2(X1,X2)|c1_2(X1,X2)|c6_2(X1,a479)|ssSkC3|~c8_2(X1,X2)|~ndr1_1(X1)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause385)).
cnf(c_0_45,negated_conjecture,(ndr1_1(a546)|ssSkC17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause55)).
cnf(c_0_46,negated_conjecture,(ssSkP8(X1)|~c1_2(X1,a588)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause180)).
cnf(c_0_47,negated_conjecture,(ssSkP1(X1)|~c1_2(X1,a485)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause200)).
cnf(c_0_48,negated_conjecture,(ndr1_1(X1)|ssSkP5(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause66)).
cnf(c_0_49,negated_conjecture,(ssSkP3(X1)|~c8_2(X1,a502)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause195)).
cnf(c_0_50,negated_conjecture,(c4_1(a570)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause31)).
cnf(c_0_51,negated_conjecture,(ndr1_1(a570)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause30)).
cnf(c_0_52,negated_conjecture,(c4_1(a537)|ssSkC15), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause58)).
cnf(c_0_53,negated_conjecture,(c4_1(a517)|ssSkC11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause64)).
cnf(c_0_54,negated_conjecture,(c8_2(a542,a543)|ssSkC16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause121)).
cnf(c_0_55,negated_conjecture,(c2_2(X1,X2)|c1_2(X1,X2)|c5_2(X1,X3)|c4_2(X1,X3)|c8_1(X1)|c5_2(a573,X4)|c1_0|~c3_2(X1,X2)|~ndr1_1(X1)|~c7_2(X1,X3)|~ndr1_0|~c6_2(a573,X4)|~c2_2(a573,X4)|~ndr1_1(a573)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause475)).
cnf(c_0_56,negated_conjecture,(c5_2(X1,a588)|ssSkP8(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause116)).
cnf(c_0_57,negated_conjecture,(c5_0|c1_0|~c5_1(a520)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause155)).
cnf(c_0_58,negated_conjecture,(ssSkC17|~c4_1(a546)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause120)).
cnf(c_0_59,negated_conjecture,(c1_2(a587,X1)|ssSkC23|~c5_2(a587,X1)|~c8_2(a587,X1)|~ndr1_1(a587)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause352)).
cnf(c_0_60,negated_conjecture,(ndr1_1(X1)|~ssSkC0|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause239)).
cnf(c_0_61,negated_conjecture,(c1_0|c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause5)).
cnf(c_0_62,negated_conjecture,(c1_2(X1,a503)|ssSkP4(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause129)).
cnf(c_0_63,negated_conjecture,(ssSkP4(X1)|~c3_2(X1,a503)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause191)).
cnf(c_0_64,negated_conjecture,(ssSkP1(X1)|~c3_2(X1,a485)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause199)).
cnf(c_0_65,negated_conjecture,(c1_2(X1,X6)|c7_2(X1,X6)|c3_2(X1,X6)|c1_1(X1)|c2_1(X1)|c1_2(X2,X3)|c8_1(X2)|c5_2(X4,X5)|c8_1(X4)|~ndr1_1(X1)|~ndr1_0|~c6_2(X2,X3)|~c2_2(X2,X3)|~ndr1_1(X2)|~c1_1(X2)|~c6_2(X4,X5)|~c7_2(X4,X5)|~ndr1_1(X4)|~c1_2(X4,a465)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause482)).
cnf(c_0_66,negated_conjecture,(c1_2(X1,a460)|~ssSkC0|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause262)).
cnf(c_0_67,negated_conjecture,(c1_2(X1,X6)|c7_2(X1,X6)|c3_2(X1,X6)|c1_1(X1)|c2_1(X1)|c1_2(X2,X3)|c8_1(X2)|c5_2(X4,X5)|c8_1(X4)|~ndr1_1(X1)|~ndr1_0|~c6_2(X2,X3)|~c2_2(X2,X3)|~ndr1_1(X2)|~c1_1(X2)|~c6_2(X4,X5)|~c7_2(X4,X5)|~ndr1_1(X4)|~c5_2(X4,a465)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause483)).
cnf(c_0_68,negated_conjecture,(ssSkC18|~c8_1(a552)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause118)).
cnf(c_0_69,negated_conjecture,(c1_2(X1,a487)|c5_1(X1)|ssSkC5|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause277)).
cnf(c_0_70,negated_conjecture,(c3_2(X1,X2)|~ssSkC17|~c6_2(X1,X2)|~c8_2(X1,X2)|~ndr1_1(X1)|~c7_2(X1,a548)|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause409)).
cnf(c_0_71,negated_conjecture,(c4_2(X1,X2)|c8_1(X1)|c1_2(X3,X4)|c4_2(X3,X4)|c6_1(X3)|c6_2(X5,a474)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c7_1(X3)|~c5_1(X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause479)).
cnf(c_0_72,negated_conjecture,(c4_2(X1,X2)|c8_1(X1)|c1_2(X3,X4)|c4_2(X3,X4)|c6_1(X3)|c7_2(X5,a474)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c7_1(X3)|~c5_1(X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause480)).
cnf(c_0_73,negated_conjecture,(c1_1(a508)|~c4_2(X1,X2)|~c3_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c4_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause421)).
cnf(c_0_74,negated_conjecture,(c4_2(X1,X2)|c8_1(X1)|c1_2(X3,X4)|c4_2(X3,X4)|c6_1(X3)|c1_2(X5,a474)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c7_1(X3)|~c5_1(X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause481)).
cnf(c_0_75,negated_conjecture,(c4_2(X1,X2)|c8_1(X1)|c1_2(X3,X4)|c4_2(X3,X4)|c6_1(X3)|ndr1_1(X5)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c7_1(X3)|~c5_1(X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause476)).
cnf(c_0_76,negated_conjecture,(c3_2(X1,X2)|c7_2(X1,X3)|c2_2(X1,X3)|c8_2(X1,X3)|c4_2(a559,a560)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_1(X1)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause455)).
cnf(c_0_77,negated_conjecture,(c3_2(X1,X2)|c7_2(X1,X3)|c2_2(X1,X3)|c8_2(X1,X3)|c7_2(a559,a560)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_1(X1)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause456)).
cnf(c_0_78,negated_conjecture,(c3_2(X1,X2)|c7_2(X1,X3)|c2_2(X1,X3)|c8_2(X1,X3)|c7_1(a559)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_1(X1)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause452)).
cnf(c_0_79,negated_conjecture,(c7_2(a470,X1)|ssSkC2|~c8_2(a470,X1)|~ndr1_1(a470)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause297)).
cnf(c_0_80,negated_conjecture,(c3_2(a536,X1)|c5_2(a536,X1)|ssSkC14|~c7_2(a536,X1)|~ndr1_1(a536)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause348)).
cnf(c_0_81,negated_conjecture,(c3_2(X1,X2)|c7_2(X1,X3)|c2_2(X1,X3)|c8_2(X1,X3)|ndr1_1(a559)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_1(X1)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause451)).
cnf(c_0_82,negated_conjecture,(c7_2(X1,X2)|c4_1(X1)|c7_1(X1)|c5_1(X3)|c7_1(X3)|~c2_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c4_2(X3,a583)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause438)).
cnf(c_0_83,negated_conjecture,(c6_2(a491,X1)|c4_2(a491,X1)|c3_2(a491,X1)|ssSkC6|~ndr1_1(a491)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause329)).
cnf(c_0_84,negated_conjecture,(c6_2(a568,X1)|c4_2(a568,X1)|ssSkC21|~ndr1_1(a568)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause271)).
cnf(c_0_85,negated_conjecture,(c5_2(X1,a460)|~ssSkC0|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause261)).
cnf(c_0_86,negated_conjecture,(c1_2(X1,X2)|c7_2(X1,X2)|c6_1(X1)|c4_2(a534,a535)|~c2_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,X3)|~c3_2(X1,X3)|~c6_2(X1,X3)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause466)).
cnf(c_0_87,negated_conjecture,(c1_2(X1,X2)|c7_2(X1,X2)|c6_1(X1)|c8_2(a534,a535)|~c2_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,X3)|~c3_2(X1,X3)|~c6_2(X1,X3)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause467)).
cnf(c_0_88,negated_conjecture,(c1_2(X1,X2)|c7_2(X1,X2)|c6_1(X1)|c8_1(a534)|~c2_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,X3)|~c3_2(X1,X3)|~c6_2(X1,X3)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause459)).
cnf(c_0_89,negated_conjecture,(c1_2(X1,X2)|c7_2(X1,X2)|c6_1(X1)|ndr1_1(a534)|~c2_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,X3)|~c3_2(X1,X3)|~c6_2(X1,X3)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause458)).
cnf(c_0_90,negated_conjecture,(c4_2(X1,a576)|c7_1(X1)|ssSkC22|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause274)).
cnf(c_0_91,negated_conjecture,(ndr1_1(X1)|c7_1(X1)|ssSkC22|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause266)).
cnf(c_0_92,negated_conjecture,(c3_2(a498,a499)|c1_0|c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause163)).
cnf(c_0_93,negated_conjecture,(c7_2(a498,a499)|c1_0|c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause164)).
cnf(c_0_94,negated_conjecture,(ndr1_1(a498)|c1_0|c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause96)).
cnf(c_0_95,negated_conjecture,(c8_2(X1,a460)|~ssSkC0|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause260)).
cnf(c_0_96,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c6_2(X1,X2)|c5_1(X1)|ndr1_1(a574)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause387)).
cnf(c_0_97,negated_conjecture,(c3_2(a564,X1)|c8_2(a564,X1)|ssSkC19|~c1_2(a564,X1)|~ndr1_1(a564)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause347)).
cnf(c_0_98,negated_conjecture,(c8_2(X1,a526)|ssSkP6(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause125)).
cnf(c_0_99,negated_conjecture,(c8_2(X1,a456)|ssSkP0(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause134)).
cnf(c_0_100,negated_conjecture,(ndr1_1(X1)|ssSkP0(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause75)).
cnf(c_0_101,negated_conjecture,(c8_2(a498,a500)|c1_0|c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause161)).
cnf(c_0_102,negated_conjecture,(ndr1_1(X1)|c5_1(X1)|ssSkC5|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause268)).
cnf(c_0_103,negated_conjecture,(ndr1_1(a482)|c5_0|~ssSkC4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause102)).
cnf(c_0_104,negated_conjecture,(c3_2(a496,X1)|~c8_2(a496,X1)|~c2_2(a496,X1)|~ndr1_1(a496)|~c4_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause371)).
cnf(c_0_105,negated_conjecture,(ndr1_1(a496)|~c4_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause225)).
cnf(c_0_106,negated_conjecture,(c7_2(X1,a459)|ssSkC0|~c3_1(X1)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause300)).
cnf(c_0_107,negated_conjecture,(ndr1_1(X1)|ssSkC0|~c3_1(X1)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause279)).
cnf(c_0_108,negated_conjecture,(ndr1_1(a544)|~ssSkC16|~c8_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause273)).
cnf(c_0_109,negated_conjecture,(c3_1(X1)|~ssSkC3|~c6_2(X1,a480)|~c2_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause341)).
cnf(c_0_110,negated_conjecture,(c3_2(X1,a480)|c3_1(X1)|~ssSkC3|~c2_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause316)).
cnf(c_0_111,negated_conjecture,(c4_2(X1,a480)|c3_1(X1)|~ssSkC3|~c2_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause317)).
cnf(c_0_112,negated_conjecture,(ndr1_1(X1)|c3_1(X1)|~ssSkC3|~c2_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause291)).
cnf(c_0_113,negated_conjecture,(~c4_2(X1,X2)|~c3_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c4_1(X1)|~ndr1_0|~c2_1(a508)|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause435)).
cnf(c_0_114,negated_conjecture,(c8_2(X1,X2)|c8_1(X1)|c8_2(X3,a467)|~c4_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0|~c6_1(X3)|~c8_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause448)).
cnf(c_0_115,negated_conjecture,(c8_2(X1,X2)|c8_1(X1)|c3_2(X3,a467)|~c4_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0|~c6_1(X3)|~c8_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause450)).
cnf(c_0_116,negated_conjecture,(c8_2(X1,X2)|c8_1(X1)|c2_2(X3,a467)|~c4_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0|~c6_1(X3)|~c8_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause449)).
cnf(c_0_117,negated_conjecture,(c1_1(a466)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause106)).
cnf(c_0_118,negated_conjecture,(~c8_1(a466)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause175)).
cnf(c_0_119,negated_conjecture,(c2_0|~c7_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause179)).
cnf(c_0_120,negated_conjecture,(c3_2(X1,X2)|~ssSkC17|~c6_2(X1,X2)|~c8_2(X1,X2)|~ndr1_1(X1)|~c4_2(X1,a548)|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause408)).
cnf(c_0_121,negated_conjecture,(c6_2(X1,a456)|ssSkP0(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause133)).
cnf(c_0_122,negated_conjecture,(ssSkP0(X1)|~c4_2(X1,a456)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause201)).
cnf(c_0_123,negated_conjecture,(c1_1(X1)|c2_1(X1)|c1_2(X2,X3)|c2_2(X2,X3)|c2_2(X2,X4)|c2_0|~c4_1(X1)|~ndr1_0|~c5_2(X2,X3)|~ndr1_1(X2)|~c7_2(X2,X4)|~c1_2(X2,X4)|~c5_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause472)).
cnf(c_0_124,negated_conjecture,(ssSkP3(X1)|~c1_2(X1,a502)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause193)).
cnf(c_0_125,negated_conjecture,(ssSkP3(X1)|~c7_2(X1,a502)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause194)).
cnf(c_0_126,negated_conjecture,(c7_2(a546,a547)|ssSkC17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause119)).
cnf(c_0_127,negated_conjecture,(c5_2(X1,X2)|~ssSkP5(X1)|~ssSkC9|~ssSkC10|~c4_2(X1,X2)|~ndr1_1(X1)|~c2_2(X1,a516)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause369)).
cnf(c_0_128,negated_conjecture,(c5_2(X1,X2)|~ssSkP5(X1)|~ssSkC9|~ssSkC10|~c4_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,a516)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause368)).
cnf(c_0_129,negated_conjecture,(ndr1_1(a505)|~ssSkC8|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause160)).
cnf(c_0_130,negated_conjecture,(c6_2(a496,a497)|~c4_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause257)).
cnf(c_0_131,negated_conjecture,(c3_2(a496,a497)|~c4_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause256)).
cnf(c_0_132,negated_conjecture,(~c4_2(a496,a497)|~c4_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause264)).
cnf(c_0_133,negated_conjecture,(c7_1(X1)|ssSkC22|~c8_2(X1,a576)|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause299)).
cnf(c_0_134,negated_conjecture,(c3_2(a512,a513)|ssSkC9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause128)).
cnf(c_0_135,negated_conjecture,(ssSkC9|~c2_2(a512,a513)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause190)).
cnf(c_0_136,negated_conjecture,(ssSkC9|~c1_2(a512,a513)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause189)).
cnf(c_0_137,negated_conjecture,(c5_2(a482,a484)|c5_0|~ssSkC4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause167)).
cnf(c_0_138,negated_conjecture,(c8_1(a581)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause83)).
cnf(c_0_139,negated_conjecture,(ssSkC6|~c5_1(a491)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause132)).
cnf(c_0_140,negated_conjecture,(ssSkP4(X1)|~c4_2(X1,a503)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause192)).
cnf(c_0_141,negated_conjecture,(c6_2(a549,a550)|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause245)).
cnf(c_0_142,negated_conjecture,(c7_2(X1,X2)|c4_1(X1)|c7_1(X1)|c7_2(X3,a583)|c5_1(X3)|c7_1(X3)|~c2_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause429)).
cnf(c_0_143,negated_conjecture,(c7_2(X1,X2)|c4_1(X1)|c7_1(X1)|c2_2(X3,a583)|c5_1(X3)|c7_1(X3)|~c2_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause430)).
cnf(c_0_144,negated_conjecture,(c6_2(X1,X2)|ssSkC1|~c2_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,a462)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause422)).
cnf(c_0_145,negated_conjecture,(c3_2(X1,X2)|c2_2(X1,X2)|ssSkC12|~ssSkP6(X1)|~c1_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,a527)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause381)).
cnf(c_0_146,negated_conjecture,(c7_2(X1,X2)|c4_1(X1)|c7_1(X1)|ndr1_1(X3)|c5_1(X3)|c7_1(X3)|~c2_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause417)).
cnf(c_0_147,negated_conjecture,(c1_1(a532)|ssSkC13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause61)).
cnf(c_0_148,negated_conjecture,(c8_2(X1,X2)|c8_1(X1)|ndr1_1(X3)|~c4_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0|~c6_1(X3)|~c8_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause444)).
cnf(c_0_149,negated_conjecture,(c3_2(a509,a511)|c7_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause219)).
cnf(c_0_150,negated_conjecture,(~c7_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause112)).
cnf(c_0_151,negated_conjecture,(ndr1_1(a549)|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause212)).
cnf(c_0_152,negated_conjecture,(c8_2(a549,a551)|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause242)).
cnf(c_0_153,negated_conjecture,(c4_2(X1,X2)|c8_2(X1,a481)|c8_1(X1)|ssSkC4|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause384)).
cnf(c_0_154,negated_conjecture,(c4_2(X1,X2)|c2_2(X1,a481)|c8_1(X1)|ssSkC4|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause383)).
cnf(c_0_155,negated_conjecture,(c8_2(X1,X3)|c7_2(X1,X3)|c6_2(X1,X3)|c5_1(X1)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c4_2(a574,X2)|~c2_2(a574,X2)|~c3_2(a574,X2)|~ndr1_1(a574)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause473)).
cnf(c_0_156,negated_conjecture,(c4_2(a509,a510)|c7_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause220)).
cnf(c_0_157,negated_conjecture,(ndr1_1(a509)|c7_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause159)).
cnf(c_0_158,negated_conjecture,(c7_0|~c8_2(a509,a510)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause252)).
cnf(c_0_159,negated_conjecture,(c5_2(a566,X1)|~ssSkC19|~ssSkC20|~c2_2(a566,X1)|~c1_2(a566,X1)|~ndr1_1(a566)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause358)).
cnf(c_0_160,negated_conjecture,(ndr1_1(a566)|~ssSkC19|~ssSkC20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause87)).
cnf(c_0_161,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c4_1(X1)|~ndr1_0|~c6_2(X3,a525)|~c3_1(X3)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause460)).
cnf(c_0_162,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c4_1(X1)|~ndr1_0|~c2_2(X3,a525)|~c3_1(X3)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause461)).
cnf(c_0_163,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c7_2(X3,a525)|c6_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c4_1(X1)|~ndr1_0|~c3_1(X3)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause453)).
cnf(c_0_164,negated_conjecture,(c2_2(a482,X1)|c5_0|~ssSkC4|~c6_2(a482,X1)|~c8_2(a482,X1)|~ndr1_1(a482)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause360)).
cnf(c_0_165,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|ndr1_1(X3)|c6_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c4_1(X1)|~ndr1_0|~c3_1(X3)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause447)).
cnf(c_0_166,negated_conjecture,(c2_2(X1,X2)|c5_2(X1,X2)|c1_0|c5_0|~ndr1_1(X1)|~c5_2(X1,a572)|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause367)).
cnf(c_0_167,negated_conjecture,(c8_1(X1)|c1_1(X1)|~c6_1(X1)|~ndr1_0|~c6_2(a530,X2)|~c8_2(a530,X2)|~c3_2(a530,X2)|~ndr1_1(a530)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause443)).
cnf(c_0_168,negated_conjecture,(c8_1(X1)|c1_1(X1)|c7_2(a530,X2)|~c6_1(X1)|~ndr1_0|~c6_2(a530,X2)|~c4_2(a530,X2)|~ndr1_1(a530)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause439)).
cnf(c_0_169,negated_conjecture,(ssSkC8|~ssSkP3(X1)|~ssSkP4(X1)|~c5_2(X1,a504)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause275)).
cnf(c_0_170,negated_conjecture,(c7_2(a470,X1)|ssSkC2|~c1_2(a470,X1)|~c4_2(a470,X1)|~ndr1_1(a470)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause356)).
cnf(c_0_171,negated_conjecture,(c6_2(a514,X1)|ssSkC10|~c4_2(a514,X1)|~c8_2(a514,X1)|~ndr1_1(a514)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause355)).
cnf(c_0_172,negated_conjecture,(c5_2(a549,a551)|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause243)).
cnf(c_0_173,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c6_2(X1,X2)|c5_1(X1)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c6_2(a574,a575)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause425)).
cnf(c_0_174,negated_conjecture,(c7_2(a532,X1)|ssSkC13|~c4_2(a532,X1)|~c8_2(a532,X1)|~ndr1_1(a532)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause354)).
cnf(c_0_175,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c6_2(X1,X2)|c5_1(X1)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c1_2(a574,a575)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause426)).
cnf(c_0_176,negated_conjecture,(c8_1(X1)|c1_1(X1)|c8_2(a530,X2)|c4_2(a530,X2)|c7_2(a530,X2)|~c6_1(X1)|~ndr1_0|~ndr1_1(a530)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause420)).
cnf(c_0_177,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c6_2(X1,X2)|c5_1(X1)|c4_2(a574,a575)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause406)).
cnf(c_0_178,negated_conjecture,(c4_2(a549,a550)|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause244)).
cnf(c_0_179,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c6_2(X1,X2)|c5_1(X1)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c3_1(a574)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause407)).
cnf(c_0_180,negated_conjecture,(c4_1(a488)|c1_0|~ssSkC5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause101)).
cnf(c_0_181,negated_conjecture,(ssSkC25|~c4_2(a596,X1)|~c1_2(a596,X1)|~ndr1_1(a596)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause321)).
cnf(c_0_182,negated_conjecture,(c8_2(a546,X1)|ssSkC17|~c1_2(a546,X1)|~c4_2(a546,X1)|~ndr1_1(a546)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause353)).
cnf(c_0_183,negated_conjecture,(c2_2(X2,X3)|c6_2(X2,X3)|c8_2(X2,X4)|c6_1(X2)|c5_2(X1,X5)|~ssSkP8(X1)|~ssSkC23|~c8_2(X2,X3)|~ndr1_1(X2)|~c4_2(X2,X4)|~c2_2(X2,X4)|~ndr1_0|~c2_2(X1,X5)|~ndr1_1(X1)|~c5_2(X1,a589)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause478)).
cnf(c_0_184,negated_conjecture,(c1_2(X1,X6)|c7_2(X1,X6)|c3_2(X1,X6)|c1_1(X1)|c2_1(X1)|c1_2(X2,X3)|c8_1(X2)|c5_2(X4,X5)|c8_1(X4)|~ndr1_1(X1)|~ndr1_0|~c6_2(X2,X3)|~c2_2(X2,X3)|~ndr1_1(X2)|~c1_1(X2)|~c6_2(X4,X5)|~c7_2(X4,X5)|~ndr1_1(X4)|~c7_2(X4,a465)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause484)).
cnf(c_0_185,negated_conjecture,(c2_2(X2,X3)|c6_2(X2,X3)|c8_2(X2,X4)|c6_1(X2)|c5_2(X1,X5)|~ssSkP8(X1)|~ssSkC23|~c8_2(X2,X3)|~ndr1_1(X2)|~c4_2(X2,X4)|~c2_2(X2,X4)|~ndr1_0|~c2_2(X1,X5)|~ndr1_1(X1)|~c7_2(X1,a589)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause477)).
cnf(c_0_186,negated_conjecture,(c2_2(X2,X3)|c6_2(X2,X3)|c8_2(X2,X4)|c6_1(X2)|c5_2(X1,X5)|c8_2(X1,a589)|~ssSkP8(X1)|~ssSkC23|~c8_2(X2,X3)|~ndr1_1(X2)|~c4_2(X2,X4)|~c2_2(X2,X4)|~ndr1_0|~c2_2(X1,X5)|~ndr1_1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause474)).
cnf(c_0_187,negated_conjecture,(c6_2(X1,X4)|c2_2(X1,X4)|c5_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c8_2(X1,X3)|~ssSkC22|~ndr1_1(X1)|~c1_2(X1,X2)|~c2_2(X1,X3)|~ndr1_0|~c6_2(a577,a579)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause469)).
cnf(c_0_188,negated_conjecture,(c6_2(X1,X4)|c2_2(X1,X4)|c5_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c8_2(X1,X3)|~ssSkC22|~ndr1_1(X1)|~c1_2(X1,X2)|~c2_2(X1,X3)|~ndr1_0|~c8_2(a577,a579)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause470)).
cnf(c_0_189,negated_conjecture,(c6_2(X1,X4)|c2_2(X1,X4)|c5_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c8_2(X1,X3)|c6_2(a577,a578)|~ssSkC22|~ndr1_1(X1)|~c1_2(X1,X2)|~c2_2(X1,X3)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause464)).
cnf(c_0_190,negated_conjecture,(c6_2(X1,X4)|c2_2(X1,X4)|c5_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c8_2(X1,X3)|c8_2(a577,a578)|~ssSkC22|~ndr1_1(X1)|~c1_2(X1,X2)|~c2_2(X1,X3)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause462)).
cnf(c_0_191,negated_conjecture,(c6_2(X1,X4)|c2_2(X1,X4)|c5_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c8_2(X1,X3)|c1_2(a577,a578)|~ssSkC22|~ndr1_1(X1)|~c1_2(X1,X2)|~c2_2(X1,X3)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause463)).
cnf(c_0_192,negated_conjecture,(c6_2(X1,X4)|c2_2(X1,X4)|c5_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c8_2(X1,X3)|~ssSkC22|~ndr1_1(X1)|~c1_2(X1,X2)|~c2_2(X1,X3)|~ndr1_0|~c7_1(a577)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause465)).
cnf(c_0_193,negated_conjecture,(c6_2(X1,X4)|c2_2(X1,X4)|c5_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c8_2(X1,X3)|ndr1_1(a577)|~ssSkC22|~ndr1_1(X1)|~c1_2(X1,X2)|~c2_2(X1,X3)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause454)).
cnf(c_0_194,negated_conjecture,(c1_2(X1,X2)|c7_2(X1,X2)|c6_1(X1)|~c2_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,X3)|~c3_2(X1,X3)|~c6_2(X1,X3)|~ndr1_0|~c6_2(a534,a535)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause471)).
cnf(c_0_195,negated_conjecture,(c1_2(X1,X2)|c7_2(X1,X2)|c6_1(X1)|~c2_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,X3)|~c3_2(X1,X3)|~c6_2(X1,X3)|~ndr1_0|~c2_1(a534)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause468)).
cnf(c_0_196,negated_conjecture,(c3_2(X1,X2)|c7_2(X1,X3)|c2_2(X1,X3)|c8_2(X1,X3)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_1(X1)|~ndr1_0|~c6_1(a559)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause457)).
cnf(c_0_197,negated_conjecture,(c3_2(X1,X2)|c6_2(X1,X2)|c3_2(X1,X3)|c5_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)|~c8_2(X1,X3)|~c1_2(X1,X3)|~ndr1_0|~c6_1(a531)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause446)).
cnf(c_0_198,negated_conjecture,(c3_2(X1,X2)|c6_2(X1,X2)|c3_2(X1,X3)|c5_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)|~c8_2(X1,X3)|~c1_2(X1,X3)|~ndr1_0|~c5_1(a531)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause445)).
cnf(c_0_199,negated_conjecture,(c4_2(X3,a473)|~c8_2(X1,X2)|~c4_2(X1,X2)|~c1_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c3_1(X3)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause428)).
cnf(c_0_200,negated_conjecture,(c8_2(X3,a473)|~c8_2(X1,X2)|~c4_2(X1,X2)|~c1_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c3_1(X3)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause427)).
cnf(c_0_201,negated_conjecture,(c5_1(X1)|c8_2(X2,X3)|c2_2(X2,X3)|c7_1(X2)|c5_0|~c6_1(X1)|~c3_1(X1)|~ndr1_0|~ndr1_1(X2)|~c6_2(X2,a557)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause433)).
cnf(c_0_202,negated_conjecture,(c5_1(X1)|c8_2(X2,X3)|c2_2(X2,X3)|c7_1(X2)|c5_0|~c6_1(X1)|~c3_1(X1)|~ndr1_0|~ndr1_1(X2)|~c2_2(X2,a557)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause434)).
cnf(c_0_203,negated_conjecture,(c6_2(X1,X2)|ssSkC1|~c2_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,a462)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause424)).
cnf(c_0_204,negated_conjecture,(c6_2(X1,X2)|ssSkC1|~c2_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c7_2(X1,a462)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause423)).
cnf(c_0_205,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X2)|c7_0|~ssSkC24|~c7_2(X1,X2)|~ndr1_1(X1)|~c4_2(X1,a593)|~c6_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause405)).
cnf(c_0_206,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X2)|c7_0|~ssSkC24|~c7_2(X1,X2)|~ndr1_1(X1)|~c7_2(X1,a593)|~c6_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause404)).
cnf(c_0_207,negated_conjecture,(ndr1_1(X3)|~c8_2(X1,X2)|~c4_2(X1,X2)|~c1_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c3_1(X3)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause414)).
cnf(c_0_208,negated_conjecture,(c3_2(X1,X2)|c8_2(X1,X2)|c2_1(X1)|~ssSkC12|~c7_2(X1,X2)|~ndr1_1(X1)|~c5_2(X1,a528)|~ndr1_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause411)).
cnf(c_0_209,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X2)|c8_2(X1,a593)|c7_0|~ssSkC24|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause386)).
cnf(c_0_210,negated_conjecture,(c3_2(X1,X2)|ssSkC7|~ssSkP2(X1)|~c5_2(X1,X2)|~c8_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,a493)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause401)).
cnf(c_0_211,negated_conjecture,(c3_2(X1,X2)|ssSkC7|~ssSkP2(X1)|~c5_2(X1,X2)|~c8_2(X1,X2)|~ndr1_1(X1)|~c2_2(X1,a493)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause402)).
cnf(c_0_212,negated_conjecture,(c7_2(X1,X2)|c2_1(X1)|c4_1(X1)|~ssSkC13|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c3_1(a533)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause410)).
cnf(c_0_213,negated_conjecture,(c3_2(X1,X2)|c2_2(X1,a548)|~ssSkC17|~c6_2(X1,X2)|~c8_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause388)).
cnf(c_0_214,negated_conjecture,(c3_2(X1,X2)|c2_2(X1,X2)|ssSkC12|~ssSkP6(X1)|~c1_2(X1,X2)|~ndr1_1(X1)|~c2_2(X1,a527)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause380)).
cnf(c_0_215,negated_conjecture,(c3_2(X1,X2)|c8_2(X1,X2)|c6_2(X1,a528)|c2_1(X1)|~ssSkC12|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause389)).
cnf(c_0_216,negated_conjecture,(c3_2(X1,X2)|c8_2(X1,X2)|c2_2(X1,a528)|c2_1(X1)|~ssSkC12|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause390)).
cnf(c_0_217,negated_conjecture,(c7_2(X1,X2)|c6_2(X1,X2)|c1_2(X1,X2)|c1_1(X1)|c7_0|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0|~c8_1(a501)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause392)).
cnf(c_0_218,negated_conjecture,(c7_2(X1,X2)|c6_2(X1,X2)|c1_2(X1,X2)|c1_1(X1)|c7_0|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0|~c6_1(a501)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause391)).
cnf(c_0_219,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c2_1(a458)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause399)).
cnf(c_0_220,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c6_1(a458)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause398)).
cnf(c_0_221,negated_conjecture,(c3_2(X1,X2)|c1_2(X1,X2)|ssSkC3|~c8_2(X1,X2)|~ndr1_1(X1)|~c7_2(X1,a479)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause403)).
cnf(c_0_222,negated_conjecture,(c1_2(a544,X2)|c4_2(a544,X2)|~ssSkC16|~c8_1(X1)|~c2_1(X1)|~ndr1_0|~ndr1_1(a544)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause365)).
cnf(c_0_223,negated_conjecture,(c2_2(a488,X1)|c5_2(a488,X1)|c1_0|~ssSkC5|~c6_2(a488,X1)|~ndr1_1(a488)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause351)).
cnf(c_0_224,negated_conjecture,(c4_2(a475,X1)|c5_2(a475,X1)|c7_0|~c8_2(a475,X1)|~ndr1_1(a475)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause361)).
cnf(c_0_225,negated_conjecture,(ndr1_1(a475)|c7_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause171)).
cnf(c_0_226,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|~ndr1_0|~c5_2(a539,a540)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause339)).
cnf(c_0_227,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|~ndr1_0|~c5_2(a539,a541)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause338)).
cnf(c_0_228,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|~ndr1_0|~c2_2(a539,a541)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause337)).
cnf(c_0_229,negated_conjecture,(c7_0|~ssSkC18|~c3_2(X1,a553)|~c4_1(X1)|~c5_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause336)).
cnf(c_0_230,negated_conjecture,(c7_0|~ssSkC18|~c2_2(X1,a553)|~c4_1(X1)|~c5_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause335)).
cnf(c_0_231,negated_conjecture,(c2_2(X1,X2)|c5_2(X1,X2)|c1_0|c5_0|~ndr1_1(X1)|~c6_2(X1,a572)|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause366)).
cnf(c_0_232,negated_conjecture,(ssSkC0|~c6_2(X1,a459)|~c3_1(X1)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause330)).
cnf(c_0_233,negated_conjecture,(ssSkC0|~c3_2(X1,a459)|~c3_1(X1)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause331)).
cnf(c_0_234,negated_conjecture,(ssSkC8|~ssSkP3(X1)|~ssSkP4(X1)|~c3_2(X1,a504)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause276)).
cnf(c_0_235,negated_conjecture,(c3_1(X1)|c1_1(X1)|~c5_2(X1,a582)|~ndr1_0|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause332)).
cnf(c_0_236,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|c8_2(a539,a540)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause309)).
cnf(c_0_237,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|c8_2(a539,a541)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause308)).
cnf(c_0_238,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|c3_2(a539,a540)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause310)).
cnf(c_0_239,negated_conjecture,(c2_1(X1)|c7_1(X1)|c1_0|~c7_2(X1,a602)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause301)).
cnf(c_0_240,negated_conjecture,(c7_2(X1,a553)|c7_0|~ssSkC18|~c4_1(X1)|~c5_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause307)).
cnf(c_0_241,negated_conjecture,(c3_2(X1,a582)|c3_1(X1)|c1_1(X1)|~ndr1_0|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause302)).
cnf(c_0_242,negated_conjecture,(c2_2(X1,a582)|c3_1(X1)|c1_1(X1)|~ndr1_0|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause303)).
cnf(c_0_243,negated_conjecture,(c8_2(X1,a504)|ssSkC8|~ssSkP3(X1)|~ssSkP4(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause267)).
cnf(c_0_244,negated_conjecture,(c4_2(X1,a602)|c2_1(X1)|c7_1(X1)|c1_0|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause281)).
cnf(c_0_245,negated_conjecture,(c3_2(X1,a602)|c2_1(X1)|c7_1(X1)|c1_0|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause280)).
cnf(c_0_246,negated_conjecture,(ndr1_1(X1)|c7_0|~ssSkC18|~c4_1(X1)|~c5_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause286)).
cnf(c_0_247,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|c6_1(a539)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause288)).
cnf(c_0_248,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|ndr1_1(a539)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause287)).
cnf(c_0_249,negated_conjecture,(ndr1_1(X1)|c3_1(X1)|c1_1(X1)|~ndr1_0|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause282)).
cnf(c_0_250,negated_conjecture,(c1_0|c2_0|~c8_2(a498,a499)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause224)).
cnf(c_0_251,negated_conjecture,(c1_0|c2_0|~c7_2(a498,a500)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause223)).
cnf(c_0_252,negated_conjecture,(ndr1_1(X1)|c2_1(X1)|c7_1(X1)|c1_0|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause269)).
cnf(c_0_253,negated_conjecture,(c8_2(a482,a483)|c5_0|~ssSkC4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause170)).
cnf(c_0_254,negated_conjecture,(c3_2(a482,a483)|c5_0|~ssSkC4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause168)).
cnf(c_0_255,negated_conjecture,(c2_2(a482,a483)|c5_0|~ssSkC4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause169)).
cnf(c_0_256,negated_conjecture,(c4_2(a498,a500)|c1_0|c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause162)).
cnf(c_0_257,negated_conjecture,(c7_1(a498)|c1_0|c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause97)).
cnf(c_0_258,negated_conjecture,(~ssSkC16|~c8_1(X1)|~c2_1(X1)|~ndr1_0|~c4_2(a544,a545)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause326)).
cnf(c_0_259,negated_conjecture,(~ssSkC16|~c8_1(X1)|~c2_1(X1)|~ndr1_0|~c7_2(a544,a545)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause325)).
cnf(c_0_260,negated_conjecture,(~ssSkC16|~c8_1(X1)|~c2_1(X1)|~ndr1_0|~c3_2(a544,a545)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause327)).
cnf(c_0_261,negated_conjecture,(c7_0|~c6_2(a509,a511)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause251)).
cnf(c_0_262,negated_conjecture,(c7_0|~c5_2(a509,a511)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause250)).
cnf(c_0_263,negated_conjecture,(c7_0|~c1_2(a509,a510)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause253)).
cnf(c_0_264,negated_conjecture,(c5_0|~ssSkC4|~c6_2(a482,a484)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause227)).
cnf(c_0_265,negated_conjecture,(c5_0|~ssSkC4|~c4_2(a482,a484)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause228)).
cnf(c_0_266,negated_conjecture,(~c2_1(a496)|~c4_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause258)).
cnf(c_0_267,negated_conjecture,(c1_0|~ssSkC5|~c3_1(a488)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause166)).
cnf(c_0_268,negated_conjecture,(~c8_2(X1,X2)|~c4_2(X1,X2)|~c1_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c3_2(X3,a473)|~c3_1(X3)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause437)).
cnf(c_0_269,negated_conjecture,(c5_0|c1_0|~c4_2(a520,a521)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause216)).
cnf(c_0_270,negated_conjecture,(c5_0|c1_0|~c1_2(a520,a521)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause217)).
cnf(c_0_271,negated_conjecture,(c5_0|c1_0|~c2_1(a520)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause156)).
cnf(c_0_272,negated_conjecture,(c2_2(a520,a521)|c5_0|c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause154)).
cnf(c_0_273,negated_conjecture,(c3_2(a570,a571)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause86)).
cnf(c_0_274,negated_conjecture,(c1_2(a570,a571)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause85)).
cnf(c_0_275,negated_conjecture,(c2_2(a570,a571)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause84)).
cnf(c_0_276,negated_conjecture,(c5_1(a570)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause32)).
cnf(c_0_277,negated_conjecture,(ssSkP5(X1)|~c2_2(X1,a515)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause187)).
cnf(c_0_278,negated_conjecture,(ndr1_1(X1)|ssSkP8(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause46)).
cnf(c_0_279,negated_conjecture,(ndr1_1(X1)|ssSkP1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause73)).
cnf(c_0_280,negated_conjecture,(c1_1(a538)|~ssSkC14|~ssSkC15), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause91)).
cnf(c_0_281,negated_conjecture,(c1_1(a494)|~ssSkC6|~ssSkC7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause99)).
cnf(c_0_282,negated_conjecture,(c1_1(a558)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause89)).
cnf(c_0_283,negated_conjecture,(ssSkC16|~c4_2(a542,a543)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause185)).
cnf(c_0_284,negated_conjecture,(c6_2(a566,a567)|~ssSkC19|~ssSkC20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause148)).
cnf(c_0_285,negated_conjecture,(c6_2(a505,a506)|~ssSkC8|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause221)).
cnf(c_0_286,negated_conjecture,(~ssSkC19|~ssSkC20|~c1_2(a566,a567)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause208)).
cnf(c_0_287,negated_conjecture,(~ssSkC8|~c1_2(a505,a506)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause255)).
cnf(c_0_288,negated_conjecture,(ssSkC14|~c7_1(a536)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause123)).
cnf(c_0_289,negated_conjecture,(ssSkC17|~c2_2(a546,a547)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause183)).
cnf(c_0_290,negated_conjecture,(ssSkC21|~c4_1(a568)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause117)).
cnf(c_0_291,negated_conjecture,(ssSkC6|~c4_1(a491)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause131)).
cnf(c_0_292,negated_conjecture,(c1_1(a566)|~ssSkC19|~ssSkC20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause88)).
cnf(c_0_293,negated_conjecture,(c8_2(a494,a495)|~ssSkC6|~ssSkC7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause165)).
cnf(c_0_294,negated_conjecture,(ndr1_1(a494)|~ssSkC6|~ssSkC7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause98)).
cnf(c_0_295,negated_conjecture,(~ssSkC6|~ssSkC7|~c3_2(a494,a495)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause226)).
cnf(c_0_296,negated_conjecture,(c6_1(a517)|ssSkC11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause63)).
cnf(c_0_297,negated_conjecture,(c6_1(a558)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause90)).
cnf(c_0_298,negated_conjecture,(c6_1(a468)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause105)).
cnf(c_0_299,negated_conjecture,(ssSkP2(X1)|~c6_2(X1,a492)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause196)).
cnf(c_0_300,negated_conjecture,(ssSkC15|~c3_2(a537,X1)|~c7_2(a537,X1)|~c2_2(a537,X1)|~ndr1_1(a537)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause362)).
cnf(c_0_301,negated_conjecture,(c4_1(a494)|~ssSkC6|~ssSkC7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause100)).
cnf(c_0_302,negated_conjecture,(c4_1(a569)|~ssSkC21|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause146)).
cnf(c_0_303,negated_conjecture,(c7_2(a581,X1)|c8_2(a581,X1)|c2_2(a581,X1)|~ndr1_1(a581)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause333)).
cnf(c_0_304,negated_conjecture,(c3_1(a569)|~ssSkC21|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause147)).
cnf(c_0_305,negated_conjecture,(c3_1(a600)|ssSkC26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause44)).
cnf(c_0_306,negated_conjecture,(c3_1(a587)|ssSkC23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause47)).
cnf(c_0_307,negated_conjecture,(c3_1(a532)|ssSkC13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause60)).
cnf(c_0_308,negated_conjecture,(~c6_1(a581)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause145)).
cnf(c_0_309,negated_conjecture,(c8_2(a538,X1)|c7_2(a538,X1)|~ssSkC14|~ssSkC15|~c1_2(a538,X1)|~ndr1_1(a538)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause350)).
cnf(c_0_310,negated_conjecture,(c6_2(a569,X1)|c3_2(a569,X1)|~ssSkC21|~c4_2(a569,X1)|~ndr1_1(a569)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause357)).
cnf(c_0_311,negated_conjecture,(c2_2(a538,X1)|c3_2(a538,X1)|~ssSkC14|~ssSkC15|~c7_2(a538,X1)|~ndr1_1(a538)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause349)).
cnf(c_0_312,negated_conjecture,(c8_2(a561,X1)|c4_2(a561,X1)|~ndr1_1(a561)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause272)).
cnf(c_0_313,negated_conjecture,(ssSkP1(X1)|~c2_2(X1,a485)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause198)).
cnf(c_0_314,negated_conjecture,(c7_0|~c3_2(a475,a476)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause259)).
cnf(c_0_315,negated_conjecture,(ssSkP7(X1)|~c6_2(X1,a554)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause181)).
cnf(c_0_316,negated_conjecture,(ssSkP7(X1)|~c3_2(X1,a554)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause182)).
cnf(c_0_317,negated_conjecture,(c2_2(a549,a551)|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause241)).
cnf(c_0_318,negated_conjecture,(c6_2(a475,a477)|c7_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause230)).
cnf(c_0_319,negated_conjecture,(c4_2(a475,a477)|c7_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause229)).
cnf(c_0_320,negated_conjecture,(c7_2(a475,a476)|c7_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause233)).
cnf(c_0_321,negated_conjecture,(c7_2(a475,a477)|c7_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause231)).
cnf(c_0_322,negated_conjecture,(c5_2(a475,a476)|c7_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause232)).
cnf(c_0_323,negated_conjecture,(ssSkC17|~c8_2(a546,a547)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause184)).
cnf(c_0_324,negated_conjecture,(ssSkC10|~c1_1(a514)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause127)).
cnf(c_0_325,negated_conjecture,(ssSkC15|~c5_1(a537)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause122)).
cnf(c_0_326,negated_conjecture,(ssSkC24|~c6_1(a592)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause115)).
cnf(c_0_327,negated_conjecture,(ssSkC25|~c6_1(a596)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause114)).
cnf(c_0_328,negated_conjecture,(c5_1(a601)|~ssSkC26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause25)).
cnf(c_0_329,negated_conjecture,(c7_1(a601)|~ssSkC26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause24)).
cnf(c_0_330,negated_conjecture,(c7_1(a470)|ssSkC2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause74)).
cnf(c_0_331,negated_conjecture,(c5_1(a514)|ssSkC10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause67)).
cnf(c_0_332,negated_conjecture,(c5_1(a536)|ssSkC14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause59)).
cnf(c_0_333,negated_conjecture,(~ssSkC8|~c8_2(a505,a506)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause254)).
cnf(c_0_334,negated_conjecture,(~ssSkC19|~ssSkC20|~c8_2(a566,a567)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause207)).
cnf(c_0_335,negated_conjecture,(~c1_2(a549,a550)|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause263)).
cnf(c_0_336,negated_conjecture,(c5_1(a564)|ssSkC19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause52)).
cnf(c_0_337,negated_conjecture,(c7_1(a564)|ssSkC19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause51)).
cnf(c_0_338,negated_conjecture,(c7_1(a596)|ssSkC25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause45)).
cnf(c_0_339,negated_conjecture,(ndr1_1(X1)|ssSkP2(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause72)).
cnf(c_0_340,negated_conjecture,(~ssSkC8|~c2_1(a505)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause222)).
cnf(c_0_341,negated_conjecture,(~c2_1(a549)|~c4_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause246)).
cnf(c_0_342,negated_conjecture,(~c2_1(a561)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause151)).
cnf(c_0_343,negated_conjecture,(~c4_1(a468)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause173)).
cnf(c_0_344,negated_conjecture,(~c3_1(a468)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause174)).
cnf(c_0_345,negated_conjecture,(ndr1_1(X1)|ssSkP7(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause53)).
cnf(c_0_346,negated_conjecture,(~c5_1(a561)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause150)).
cnf(c_0_347,negated_conjecture,(~ssSkC26|~c1_1(a601)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause76)).
cnf(c_0_348,negated_conjecture,(c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1)).
cnf(c_0_349,negated_conjecture,(ndr1_0), inference(csr,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_350,negated_conjecture,(~c5_0), inference(csr,[status(thm)],[c_0_2, c_0_3])).
cnf(c_0_351,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(X3)|c7_1(X1)|~c6_1(X3)|~ndr1_1(X1)|~c3_1(X3)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_349])]), c_0_350])).
cnf(c_0_352,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(a565)|c7_1(X1)|ssSkC20|~ndr1_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_5]), c_0_6]), ['final']).
cnf(c_0_353,negated_conjecture,(ssSkC8|ndr1_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_349])]), c_0_8]), c_0_9])).
cnf(c_0_354,negated_conjecture,(c8_0|~c2_0), inference(sr,[status(thm)],[c_0_10, c_0_350])).
cnf(c_0_355,negated_conjecture,(c8_0|~c7_0), inference(sr,[status(thm)],[c_0_11, c_0_350])).
cnf(c_0_356,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c4_1(a531)|c5_1(X1)|~c8_2(X1,X3)|~c7_2(X1,X2)|~c1_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_349])])).
cnf(c_0_357,negated_conjecture,(ssSkC8|c8_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(a565)|c7_1(X1)|ssSkC20), inference(spm,[status(thm)],[c_0_352, c_0_353]), ['final']).
cnf(c_0_358,negated_conjecture,(c6_2(X1,X2)|c7_2(X1,X2)|c1_2(X1,X2)|c1_1(X1)|c7_0|c5_1(a501)|~c2_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_349])])).
cnf(c_0_359,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|c4_1(a458)|~c7_2(X1,X2)|~c6_1(X1)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_349])])).
cnf(c_0_360,negated_conjecture,(c8_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_15, c_0_354]), c_0_355]), ['final']).
cnf(c_0_361,negated_conjecture,(ssSkC8|c6_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c1_2(X1,a557)|c2_2(X1,X3)|c4_1(a531)|c5_1(a565)|c5_1(X1)|c7_1(X1)|ssSkC20|~c7_2(X1,X2)|~c1_2(X1,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_357]), c_0_353]), ['final']).
cnf(c_0_362,negated_conjecture,(c6_2(a565,X1)|c7_2(a565,X1)|c1_2(a565,X1)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20|~ndr1_1(a565)), inference(spm,[status(thm)],[c_0_358, c_0_16]), ['final']).
cnf(c_0_363,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c8_2(a552,X1)|c2_2(a552,X1)|ssSkC18), inference(spm,[status(thm)],[c_0_17, c_0_353]), ['final']).
cnf(c_0_364,negated_conjecture,(ssSkC8|c6_2(a592,X1)|c5_2(a592,X1)|ssSkC24|~c1_2(a592,X1)), inference(spm,[status(thm)],[c_0_18, c_0_353]), ['final']).
cnf(c_0_365,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|c4_1(a458)|~c7_2(X1,X2)|~c6_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_359, c_0_360])]), ['final']).
cnf(c_0_366,negated_conjecture,(~c1_0), inference(sr,[status(thm)],[c_0_19, c_0_350])).
cnf(c_0_367,negated_conjecture,(ssSkC8|c6_2(X1,a526)|c3_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|ssSkP6(X1)|c4_1(a531)|c5_1(a565)|c5_1(X1)|c7_1(X1)|ssSkC20|~c1_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_20]), c_0_21]), ['final']).
cnf(c_0_368,negated_conjecture,(ssSkC8|c6_2(a565,X1)|c7_2(a565,X1)|c1_2(a565,X1)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20), inference(spm,[status(thm)],[c_0_362, c_0_353]), ['final']).
cnf(c_0_369,negated_conjecture,(c6_2(a517,X1)|c7_2(a517,X1)|c1_2(a517,X1)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11|~ndr1_1(a517)), inference(spm,[status(thm)],[c_0_358, c_0_22]), ['final']).
cnf(c_0_370,negated_conjecture,(ssSkC8|c6_2(X1,a515)|c3_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|ssSkP5(X1)|c4_1(a531)|c5_1(a565)|c5_1(X1)|c7_1(X1)|ssSkC20|~c1_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_23]), c_0_24]), ['final']).
cnf(c_0_371,negated_conjecture,(c4_2(X1,a487)|ssSkC5|c5_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_349])])).
cnf(c_0_372,negated_conjecture,(c7_2(a542,X1)|ssSkC16|~c4_2(a542,X1)), inference(csr,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_373,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c6_2(a552,X2)|c3_2(a552,X2)|c3_2(a552,X1)|c2_2(a552,X1)|c4_1(a531)|c5_1(a552)|ssSkC18|~c7_2(a552,X2)|~c1_2(a552,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_363]), c_0_353]), ['final']).
cnf(c_0_374,negated_conjecture,(ssSkC8|c3_2(a592,X1)|ssSkC24|~c5_2(a592,X1)), inference(spm,[status(thm)],[c_0_28, c_0_353]), ['final']).
cnf(c_0_375,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c8_2(a592,X1)|c5_2(a592,a557)|c2_2(a592,X1)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(spm,[status(thm)],[c_0_364, c_0_357]), ['final']).
cnf(c_0_376,negated_conjecture,(ssSkC8|c7_2(a600,X1)|ssSkC26|~c1_2(a600,X1)), inference(spm,[status(thm)],[c_0_29, c_0_353]), ['final']).
cnf(c_0_377,negated_conjecture,(ssSkC12|c4_2(X1,a527)|c3_2(X1,X2)|c2_2(X1,X2)|~c1_2(X1,X2)|~ssSkP6(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_349])])).
cnf(c_0_378,negated_conjecture,(ssSkP6(X1)|c2_1(X1)|c4_1(a458)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_20]), c_0_31]), c_0_21]), ['final']).
cnf(c_0_379,negated_conjecture,(c8_2(a512,X1)|c1_2(a512,a557)|c2_2(a512,X1)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20), inference(spm,[status(thm)],[c_0_352, c_0_32]), ['final']).
cnf(c_0_380,negated_conjecture,(ndr1_1(a520)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_33, c_0_350]), c_0_366])).
cnf(c_0_381,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c7_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_367, c_0_368]), ['final']).
cnf(c_0_382,negated_conjecture,(ssSkC8|c6_2(a517,X1)|c7_2(a517,X1)|c1_2(a517,X1)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11), inference(spm,[status(thm)],[c_0_369, c_0_353]), ['final']).
cnf(c_0_383,negated_conjecture,(c8_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_352, c_0_27]), ['final']).
cnf(c_0_384,negated_conjecture,(c6_2(a542,X1)|c7_2(a542,X1)|c1_2(a542,X1)|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_358, c_0_34]), c_0_27]), ['final']).
cnf(c_0_385,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c7_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|ssSkP5(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_370, c_0_368]), ['final']).
cnf(c_0_386,negated_conjecture,(c4_2(a542,a487)|ssSkC5|c5_1(a542)|ssSkC16), inference(spm,[status(thm)],[c_0_371, c_0_34]), ['final']).
cnf(c_0_387,negated_conjecture,(c3_2(X1,X2)|c1_2(X1,a493)|ssSkC7|~c8_2(X1,X2)|~c5_2(X1,X2)|~ssSkP2(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_349])])).
cnf(c_0_388,negated_conjecture,(ssSkP2(a542)|ssSkC16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_36]), c_0_37]), ['final']).
cnf(c_0_389,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,X1)|c3_2(a552,X1)|c2_2(a552,X1)|ssSkP6(a552)|c4_1(a531)|c5_1(a552)|ssSkC18|~c1_2(a552,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_20]), c_0_21]), ['final']).
cnf(c_0_390,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,X1)|c3_2(a552,X1)|c2_2(a552,X1)|ssSkP5(a552)|c4_1(a531)|c5_1(a552)|ssSkC18|~c1_2(a552,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_23]), c_0_24]), ['final']).
cnf(c_0_391,negated_conjecture,(c4_2(X1,X2)|c5_2(X1,X2)|c1_2(X1,X3)|c2_2(X1,X3)|c8_1(X1)|c1_1(a573)|~c7_2(X1,X2)|~c3_2(X1,X3)|~ndr1_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_349])]), c_0_366])).
cnf(c_0_392,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c8_2(a592,X1)|c3_2(a592,a557)|c2_2(a592,X1)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(spm,[status(thm)],[c_0_374, c_0_375]), ['final']).
cnf(c_0_393,negated_conjecture,(c4_2(X1,X2)|c5_2(X1,X2)|c1_2(X1,X3)|c2_2(X1,X3)|c8_1(X1)|c3_1(a573)|~c7_2(X1,X2)|~c3_2(X1,X3)|~ndr1_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_349])]), c_0_366])).
cnf(c_0_394,negated_conjecture,(ssSkC8|c8_2(a600,X1)|c7_2(a600,a557)|c2_2(a600,X1)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26), inference(spm,[status(thm)],[c_0_376, c_0_357]), ['final']).
cnf(c_0_395,negated_conjecture,(c5_2(X1,X2)|c2_2(X1,a572)|c2_2(X1,X2)|~c4_1(X1)|~ndr1_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_349])]), c_0_350]), c_0_366])).
cnf(c_0_396,negated_conjecture,(ssSkC12|ssSkC8|c4_2(X1,a527)|c8_2(X1,X2)|c3_2(X1,a557)|c2_2(X1,a557)|c2_2(X1,X2)|c5_1(a565)|c7_1(X1)|ssSkC20|~ssSkP6(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_357]), c_0_353]), ['final']).
cnf(c_0_397,negated_conjecture,(ssSkP6(a552)|c2_1(a552)|c4_1(a458)|ssSkC18), inference(spm,[status(thm)],[c_0_378, c_0_41]), ['final']).
cnf(c_0_398,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c8_2(a512,X1)|c3_2(a512,a557)|c2_2(a512,a557)|c2_2(a512,X1)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkP6(a512)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_379]), c_0_32]), ['final']).
cnf(c_0_399,negated_conjecture,(ssSkP6(a512)|c2_1(a512)|c4_1(a458)|ssSkC9), inference(spm,[status(thm)],[c_0_378, c_0_42]), ['final']).
cnf(c_0_400,negated_conjecture,(c3_2(X1,X2)|c1_2(X1,X2)|c8_1(X3)|c8_1(X1)|c4_1(X1)|c3_1(X3)|~c8_2(X1,X2)|~c1_1(X3)|~ndr1_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_349])]), c_0_366])).
cnf(c_0_401,negated_conjecture,(c8_2(a520,X1)|c1_2(a520,a557)|c2_2(a520,X1)|c5_1(a565)|c7_1(a520)|ssSkC20), inference(spm,[status(thm)],[c_0_352, c_0_380]), ['final']).
cnf(c_0_402,negated_conjecture,(ssSkC3|c6_2(X1,a479)|c3_2(X1,X2)|c1_2(X1,X2)|~c8_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_349])])).
cnf(c_0_403,negated_conjecture,(c8_2(a546,X1)|c1_2(a546,a557)|c2_2(a546,X1)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20), inference(spm,[status(thm)],[c_0_352, c_0_45]), ['final']).
cnf(c_0_404,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c3_2(a565,X2)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|c2_2(a565,X2)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c1_2(a565,X2)), inference(spm,[status(thm)],[c_0_361, c_0_381]), ['final']).
cnf(c_0_405,negated_conjecture,(ssSkC8|c6_2(a517,a588)|c7_2(a517,a588)|c1_1(a517)|ssSkP8(a517)|c7_0|c5_1(a501)|ssSkC11), inference(spm,[status(thm)],[c_0_46, c_0_382]), ['final']).
cnf(c_0_406,negated_conjecture,(c6_2(a542,X1)|c3_2(a542,X1)|c3_2(a542,X2)|c1_2(a542,a557)|c2_2(a542,X2)|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20|~c7_2(a542,X1)|~c1_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_383]), c_0_27]), ['final']).
cnf(c_0_407,negated_conjecture,(c6_2(a542,a588)|c7_2(a542,a588)|c1_1(a542)|ssSkP8(a542)|c7_0|c5_1(a501)|ssSkC16), inference(spm,[status(thm)],[c_0_46, c_0_384]), ['final']).
cnf(c_0_408,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c3_2(a565,X2)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|c2_2(a565,X2)|ssSkP5(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c1_2(a565,X2)), inference(spm,[status(thm)],[c_0_361, c_0_385]), ['final']).
cnf(c_0_409,negated_conjecture,(c7_2(a542,a487)|ssSkC5|c5_1(a542)|ssSkC16), inference(spm,[status(thm)],[c_0_372, c_0_386]), ['final']).
cnf(c_0_410,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c7_2(a517,a485)|ssSkP1(a517)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11), inference(spm,[status(thm)],[c_0_47, c_0_382]), ['final']).
cnf(c_0_411,negated_conjecture,(ssSkC8|c6_2(a565,a588)|c7_2(a565,a588)|c1_1(a565)|ssSkP8(a565)|c7_0|c5_1(a501)|ssSkC20), inference(spm,[status(thm)],[c_0_46, c_0_368]), ['final']).
cnf(c_0_412,negated_conjecture,(c6_2(a542,a485)|c7_2(a542,a485)|ssSkP1(a542)|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16), inference(spm,[status(thm)],[c_0_47, c_0_384]), ['final']).
cnf(c_0_413,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,a493)|c2_2(a542,X1)|ssSkC7|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c5_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_383]), c_0_27]), c_0_388]), ['final']).
cnf(c_0_414,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c8_2(a552,X1)|c3_2(a552,a557)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP6(a552)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(spm,[status(thm)],[c_0_389, c_0_357]), ['final']).
cnf(c_0_415,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c7_2(a565,a485)|ssSkP1(a565)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20), inference(spm,[status(thm)],[c_0_47, c_0_368]), ['final']).
cnf(c_0_416,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c8_2(a552,X1)|c3_2(a552,a557)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP5(a552)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(spm,[status(thm)],[c_0_390, c_0_357]), ['final']).
cnf(c_0_417,negated_conjecture,(c4_2(X1,a515)|c5_2(X1,a515)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(X1)|ssSkP5(X1)|c1_1(a573)|~c3_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_23]), c_0_48]), ['final']).
cnf(c_0_418,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c2_2(a592,a502)|ssSkP3(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(spm,[status(thm)],[c_0_49, c_0_392]), ['final']).
cnf(c_0_419,negated_conjecture,(c4_2(X1,a515)|c5_2(X1,a515)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(X1)|ssSkP5(X1)|c3_1(a573)|~c3_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_23]), c_0_48]), ['final']).
cnf(c_0_420,negated_conjecture,(c4_1(a570)), inference(sr,[status(thm)],[c_0_50, c_0_350])).
cnf(c_0_421,negated_conjecture,(ndr1_1(a570)), inference(sr,[status(thm)],[c_0_51, c_0_350])).
cnf(c_0_422,negated_conjecture,(ssSkC8|c6_2(a600,X1)|c7_2(a600,a557)|c3_2(a600,X1)|c3_2(a600,X2)|c2_2(a600,X2)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26|~c7_2(a600,X1)|~c1_2(a600,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_394]), c_0_353]), ['final']).
cnf(c_0_423,negated_conjecture,(c5_2(a537,X1)|c2_2(a537,a572)|c2_2(a537,X1)|ssSkC15|~ndr1_1(a537)), inference(spm,[status(thm)],[c_0_395, c_0_52]), ['final']).
cnf(c_0_424,negated_conjecture,(c5_2(a517,X1)|c2_2(a517,a572)|c2_2(a517,X1)|ssSkC11|~ndr1_1(a517)), inference(spm,[status(thm)],[c_0_395, c_0_53]), ['final']).
cnf(c_0_425,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c8_2(a552,X1)|c3_2(a552,a557)|c2_2(a552,a557)|c2_2(a552,X1)|c2_1(a552)|c4_1(a458)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20), inference(spm,[status(thm)],[c_0_396, c_0_397]), ['final']).
cnf(c_0_426,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X2)|c4_1(a531)|c5_1(X1)|~c8_2(X1,X2)|~c7_2(X1,X2)|~c1_2(X1,X2)|~ndr1_1(X1)), inference(ef,[status(thm)],[c_0_356]), ['final']).
cnf(c_0_427,negated_conjecture,(c6_2(a542,X1)|c3_2(a542,a543)|c3_2(a542,X1)|c4_1(a531)|c5_1(a542)|ssSkC16|~c7_2(a542,X1)|~c1_2(a542,a543)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_54]), c_0_27]), ['final']).
cnf(c_0_428,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c8_2(a512,X1)|c3_2(a512,a557)|c2_2(a512,a557)|c2_2(a512,X1)|c2_1(a512)|c4_1(a458)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20), inference(spm,[status(thm)],[c_0_398, c_0_399]), ['final']).
cnf(c_0_429,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c6_2(a592,X1)|c5_2(a592,a557)|c3_2(a592,X1)|c3_2(a592,X2)|c2_2(a592,X2)|c4_1(a531)|c5_1(a565)|c5_1(a592)|c7_1(a592)|ssSkC20|ssSkC24|~c7_2(a592,X1)|~c1_2(a592,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_375]), c_0_353]), ['final']).
cnf(c_0_430,negated_conjecture,(c4_2(X1,X2)|c5_2(a573,X3)|c5_2(X1,X2)|c1_2(X1,X4)|c2_2(X1,X4)|c8_1(X1)|~c6_2(a573,X3)|~c7_2(X1,X2)|~c3_2(X1,X4)|~c2_2(a573,X3)|~ndr1_1(a573)|~ndr1_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_349])]), c_0_366])).
cnf(c_0_431,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a557)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(X3)|c8_1(X1)|c4_1(X1)|c3_1(X3)|c5_1(a565)|c7_1(X1)|ssSkC20|~c1_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_357]), c_0_353]), ['final']).
cnf(c_0_432,negated_conjecture,(ssSkC8|c3_2(a592,a588)|ssSkP8(a592)|ssSkC24), inference(spm,[status(thm)],[c_0_374, c_0_56]), ['final']).
cnf(c_0_433,negated_conjecture,(~c5_1(a520)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_57, c_0_350]), c_0_366])).
cnf(c_0_434,negated_conjecture,(c3_2(a520,X1)|c1_2(a520,a557)|c1_2(a520,X1)|c2_2(a520,X1)|c8_1(a520)|c8_1(X2)|c4_1(a520)|c3_1(X2)|c5_1(a565)|c7_1(a520)|ssSkC20|~c1_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_401]), c_0_380])]), ['final']).
cnf(c_0_435,negated_conjecture,(ssSkC3|ssSkC8|c6_2(X1,a479)|c3_2(X1,X2)|c1_2(X1,a557)|c1_2(X1,X2)|c2_2(X1,X2)|c5_1(a565)|c7_1(X1)|ssSkC20|~c8_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_357]), c_0_353]), ['final']).
cnf(c_0_436,negated_conjecture,(c3_2(a546,X1)|c1_2(a546,a557)|c1_2(a546,X1)|c2_2(a546,X1)|c8_1(a546)|c8_1(X2)|c3_1(X2)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c1_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_403]), c_0_45]), c_0_58]), ['final']).
cnf(c_0_437,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|c8_1(X2)|c4_1(a542)|c3_1(X2)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c1_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_383]), c_0_27]), ['final']).
cnf(c_0_438,negated_conjecture,(c3_2(a512,X1)|c1_2(a512,a557)|c1_2(a512,X1)|c2_2(a512,X1)|c8_1(a512)|c8_1(X2)|c4_1(a512)|c3_1(X2)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c1_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_379]), c_0_32]), ['final']).
cnf(c_0_439,negated_conjecture,(ssSkC8|c1_2(a587,X1)|ssSkC23|~c8_2(a587,X1)|~c5_2(a587,X1)), inference(spm,[status(thm)],[c_0_59, c_0_353]), ['final']).
cnf(c_0_440,negated_conjecture,(ndr1_1(X1)|~ssSkC0|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_349])])).
cnf(c_0_441,negated_conjecture,(c6_0), inference(sr,[status(thm)],[c_0_61, c_0_366])).
cnf(c_0_442,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,a503)|c2_2(a565,X1)|ssSkP4(a565)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_62]), c_0_63]), ['final']).
cnf(c_0_443,negated_conjecture,(ssSkC8|c6_2(a517,a588)|c3_2(a517,a588)|c3_2(a517,X1)|c1_2(a517,a557)|c2_2(a517,X1)|c1_1(a517)|ssSkP8(a517)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20|~c1_2(a517,X1)), inference(spm,[status(thm)],[c_0_361, c_0_405]), ['final']).
cnf(c_0_444,negated_conjecture,(c6_2(a542,a588)|c3_2(a542,a588)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|c1_1(a542)|ssSkP8(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c1_2(a542,X1)), inference(spm,[status(thm)],[c_0_406, c_0_407]), ['final']).
cnf(c_0_445,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,a503)|c2_2(a565,X1)|ssSkP4(a565)|ssSkP5(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_62]), c_0_63]), ['final']).
cnf(c_0_446,negated_conjecture,(c6_2(a542,a487)|c3_2(a542,a487)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|ssSkC5|c4_1(a531)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c1_2(a542,X1)), inference(spm,[status(thm)],[c_0_406, c_0_409]), ['final']).
cnf(c_0_447,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c3_2(a517,X1)|c1_2(a517,a557)|c2_2(a517,X1)|ssSkP1(a517)|c1_1(a517)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20|~c1_2(a517,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_410]), c_0_64]), ['final']).
cnf(c_0_448,negated_conjecture,(ssSkC8|c6_2(a565,a588)|c3_2(a565,a588)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|c1_1(a565)|ssSkP8(a565)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c1_2(a565,X1)), inference(spm,[status(thm)],[c_0_361, c_0_411]), ['final']).
cnf(c_0_449,negated_conjecture,(c4_2(X1,a526)|c5_2(X1,a526)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(X1)|ssSkP6(X1)|c3_1(a573)|~c3_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_20]), c_0_31]), ['final']).
cnf(c_0_450,negated_conjecture,(c4_2(X1,a526)|c5_2(X1,a526)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(X1)|ssSkP6(X1)|c1_1(a573)|~c3_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_20]), c_0_31]), ['final']).
cnf(c_0_451,negated_conjecture,(c6_2(a542,a485)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|ssSkP1(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c1_2(a542,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_412]), c_0_64]), ['final']).
cnf(c_0_452,negated_conjecture,(c6_2(a542,a526)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|ssSkP6(a542)|c4_1(a531)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c1_2(a542,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_20]), c_0_21]), ['final']).
cnf(c_0_453,negated_conjecture,(c3_2(a542,a588)|c1_2(a542,a493)|c1_2(a542,a557)|c2_2(a542,a588)|ssSkP8(a542)|ssSkC7|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_413, c_0_56]), ['final']).
cnf(c_0_454,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c3_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP6(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20), inference(spm,[status(thm)],[c_0_49, c_0_414]), ['final']).
cnf(c_0_455,negated_conjecture,(c6_2(a542,a515)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|ssSkP5(a542)|c4_1(a531)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c1_2(a542,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_23]), c_0_24]), ['final']).
cnf(c_0_456,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|ssSkP1(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c1_2(a565,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_415]), c_0_64]), ['final']).
cnf(c_0_457,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c3_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP5(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20), inference(spm,[status(thm)],[c_0_49, c_0_416]), ['final']).
cnf(c_0_458,negated_conjecture,(c7_2(X1,X2)|c5_2(X3,X4)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X5,X6)|c8_1(X5)|c8_1(X3)|c1_1(X1)|c2_1(X1)|~c6_2(X3,X4)|~c6_2(X5,X6)|~c7_2(X3,X4)|~c1_2(X3,a465)|~c2_2(X5,X6)|~c1_1(X5)|~ndr1_1(X3)|~ndr1_1(X5)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_349])])).
cnf(c_0_459,negated_conjecture,(c1_2(X1,a460)|~ssSkC0|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_349])])).
cnf(c_0_460,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c4_2(a592,a515)|c5_2(a592,a515)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c1_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(spm,[status(thm)],[c_0_417, c_0_418])).
cnf(c_0_461,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c4_2(a592,a515)|c5_2(a592,a515)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c3_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(spm,[status(thm)],[c_0_419, c_0_418])).
cnf(c_0_462,negated_conjecture,(c7_2(X1,X2)|c5_2(X3,X4)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X5,X6)|c8_1(X5)|c8_1(X3)|c1_1(X1)|c2_1(X1)|~c6_2(X3,X4)|~c6_2(X5,X6)|~c7_2(X3,X4)|~c5_2(X3,a465)|~c2_2(X5,X6)|~c1_1(X5)|~ndr1_1(X3)|~ndr1_1(X5)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_349])])).
cnf(c_0_463,negated_conjecture,(c5_2(a570,X1)|c2_2(a570,a572)|c2_2(a570,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_420]), c_0_421])]), ['final']).
cnf(c_0_464,negated_conjecture,(ssSkC8|c6_2(a600,a526)|c7_2(a600,a557)|c3_2(a600,X1)|c2_2(a600,X1)|ssSkP6(a600)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c1_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_20]), c_0_21]), ['final']).
cnf(c_0_465,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c6_2(a552,X1)|c3_2(a552,a557)|c3_2(a552,X1)|c3_2(a552,X2)|c2_2(a552,a557)|c2_2(a552,X2)|ssSkP6(a552)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~c7_2(a552,X1)|~c1_2(a552,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_414]), c_0_353])).
cnf(c_0_466,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c6_2(a552,X1)|c3_2(a552,a557)|c3_2(a552,X1)|c3_2(a552,X2)|c2_2(a552,a557)|c2_2(a552,X2)|ssSkP5(a552)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~c7_2(a552,X1)|~c1_2(a552,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_416]), c_0_353])).
cnf(c_0_467,negated_conjecture,(ssSkC8|c6_2(a600,a515)|c7_2(a600,a557)|c3_2(a600,X1)|c2_2(a600,X1)|ssSkP5(a600)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c1_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_23]), c_0_24]), ['final']).
cnf(c_0_468,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(X2)|ssSkP6(a552)|c4_1(a531)|c4_1(a552)|c3_1(X2)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c1_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_414]), c_0_353]), c_0_68]), ['final']).
cnf(c_0_469,negated_conjecture,(c1_2(X1,a487)|ssSkC5|c5_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_349])])).
cnf(c_0_470,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(X2)|ssSkP5(a552)|c4_1(a531)|c4_1(a552)|c3_1(X2)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c1_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_416]), c_0_353]), c_0_68]), ['final']).
cnf(c_0_471,negated_conjecture,(ssSkC8|c5_2(a537,X1)|c2_2(a537,a572)|c2_2(a537,X1)|ssSkC15), inference(spm,[status(thm)],[c_0_423, c_0_353]), ['final']).
cnf(c_0_472,negated_conjecture,(ssSkC8|c5_2(a517,X1)|c2_2(a517,a572)|c2_2(a517,X1)|ssSkC11), inference(spm,[status(thm)],[c_0_424, c_0_353]), ['final']).
cnf(c_0_473,negated_conjecture,(c3_2(X1,X2)|~c6_2(X1,X2)|~c8_2(X1,X2)|~c7_2(X1,a548)|~c1_1(X1)|~ndr1_1(X1)|~ssSkC17), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_349])])).
cnf(c_0_474,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(X2)|c2_1(a552)|c4_1(a458)|c4_1(a552)|c3_1(X2)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c1_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_425]), c_0_353]), c_0_68]), ['final']).
cnf(c_0_475,negated_conjecture,(ssSkC8|c6_2(a600,X1)|c7_2(a600,a557)|c3_2(a600,X1)|c2_2(a600,X1)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26|~c1_2(a600,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_394]), c_0_353]), c_0_376]), ['final']).
cnf(c_0_476,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,X1)|c7_2(a542,a543)|c3_2(a542,a543)|c3_2(a542,X1)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c7_2(a542,X1)), inference(spm,[status(thm)],[c_0_427, c_0_384]), ['final']).
cnf(c_0_477,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c2_2(a600,a502)|ssSkP3(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26), inference(spm,[status(thm)],[c_0_49, c_0_394]), ['final']).
cnf(c_0_478,negated_conjecture,(ssSkC8|c6_2(X1,X2)|c3_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c4_1(a531)|c5_1(a565)|c5_1(X1)|c7_1(X1)|ssSkC20|~c7_2(X1,X2)|~c1_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_357]), c_0_353]), ['final']).
cnf(c_0_479,negated_conjecture,(ssSkC8|c7_2(a600,a503)|ssSkP4(a600)|ssSkC26), inference(spm,[status(thm)],[c_0_376, c_0_62]), ['final']).
cnf(c_0_480,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c6_2(a592,X1)|c3_2(a592,a557)|c3_2(a592,X1)|c3_2(a592,X2)|c2_2(a592,X2)|c4_1(a531)|c5_1(a565)|c5_1(a592)|c7_1(a592)|ssSkC20|ssSkC24|~c7_2(a592,X1)|~c1_2(a592,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_392]), c_0_353]), ['final']).
cnf(c_0_481,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c1_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c8_1(a512)|c8_1(X2)|c2_1(a512)|c4_1(a458)|c4_1(a512)|c3_1(X2)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c1_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_428]), c_0_32]), ['final']).
cnf(c_0_482,negated_conjecture,(c6_2(X1,a474)|c4_2(X2,X3)|c4_2(X4,X5)|c1_2(X4,X5)|c8_1(X2)|c6_1(X4)|~c6_2(X4,X5)|~c6_2(X2,X3)|~c3_2(X2,X3)|~ndr1_1(X4)|~ndr1_1(X2)|~c3_1(X2)|~c5_1(X1)|~c7_1(X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_349])])).
cnf(c_0_483,negated_conjecture,(ssSkC8|c6_2(a592,a526)|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|ssSkP6(a592)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c1_2(a592,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_20]), c_0_21]), ['final']).
cnf(c_0_484,negated_conjecture,(ssSkC8|c6_2(a592,a515)|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|ssSkP5(a592)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c1_2(a592,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_23]), c_0_24]), ['final']).
cnf(c_0_485,negated_conjecture,(c4_2(X1,X2)|c4_2(X3,X4)|c7_2(X5,a474)|c1_2(X3,X4)|c8_1(X1)|c6_1(X3)|~c6_2(X3,X4)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X3)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X5)|~c7_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_349])])).
cnf(c_0_486,negated_conjecture,(c1_1(a508)|~c4_2(X1,X2)|~c7_2(X1,X2)|~c3_2(X1,X2)|~c4_1(X1)|~ndr1_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_349])])).
cnf(c_0_487,negated_conjecture,(c4_2(X1,X2)|c4_2(X3,X4)|c1_2(X5,a474)|c1_2(X3,X4)|c8_1(X1)|c6_1(X3)|~c6_2(X3,X4)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X3)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X5)|~c7_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_349])])).
cnf(c_0_488,negated_conjecture,(ssSkC8|c4_2(X1,X2)|c5_2(a573,X3)|c5_2(X1,X2)|c1_2(X1,X4)|c2_2(X1,X4)|c8_1(X1)|~c6_2(a573,X3)|~c7_2(X1,X2)|~c3_2(X1,X4)|~c2_2(a573,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_353]), c_0_353]), ['final']).
cnf(c_0_489,negated_conjecture,(c4_2(X1,X2)|c4_2(X3,X4)|c1_2(X3,X4)|c8_1(X1)|c6_1(X3)|ndr1_1(X5)|~c6_2(X3,X4)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X3)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X5)|~c7_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_349])])).
cnf(c_0_490,negated_conjecture,(ssSkC8|c3_2(X1,a557)|c1_2(X1,a557)|c2_2(X1,a557)|c8_1(X1)|c8_1(X2)|c4_1(X1)|c3_1(X2)|c5_1(a565)|c7_1(X1)|ssSkC20|~c1_1(X2)), inference(ef,[status(thm)],[c_0_431]), ['final']).
cnf(c_0_491,negated_conjecture,(c6_2(a546,X1)|c3_2(a546,X1)|c3_2(a546,X2)|c1_2(a546,a557)|c2_2(a546,X2)|c4_1(a531)|c5_1(a565)|c5_1(a546)|c7_1(a546)|ssSkC17|ssSkC20|~c7_2(a546,X1)|~c1_2(a546,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_403]), c_0_45]), ['final']).
cnf(c_0_492,negated_conjecture,(c6_2(a512,X1)|c3_2(a512,X1)|c3_2(a512,X2)|c1_2(a512,a557)|c2_2(a512,X2)|c4_1(a531)|c5_1(a565)|c5_1(a512)|c7_1(a512)|ssSkC9|ssSkC20|~c7_2(a512,X1)|~c1_2(a512,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_379]), c_0_32]), ['final']).
cnf(c_0_493,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a592)|c8_1(X2)|c4_1(a592)|c3_1(X2)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c1_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_392]), c_0_353]), ['final']).
cnf(c_0_494,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a592)|c8_1(X2)|c4_1(a592)|c3_1(X2)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c1_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_375]), c_0_353]), ['final']).
cnf(c_0_495,negated_conjecture,(ssSkC8|c4_2(a592,a515)|c5_2(a592,a515)|c2_2(a592,a588)|c8_1(a592)|ssSkP5(a592)|c1_1(a573)|ssSkP8(a592)|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_432]), c_0_46])).
cnf(c_0_496,negated_conjecture,(c4_2(a559,a560)|c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_349])])).
cnf(c_0_497,negated_conjecture,(ssSkC8|c4_2(a592,a515)|c5_2(a592,a515)|c2_2(a592,a588)|c8_1(a592)|ssSkP5(a592)|ssSkP8(a592)|c3_1(a573)|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_432]), c_0_46])).
cnf(c_0_498,negated_conjecture,(c8_2(X1,X2)|c7_2(a559,a560)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_349])])).
cnf(c_0_499,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|c7_1(a559)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_349])])).
cnf(c_0_500,negated_conjecture,(c6_2(a520,X1)|c3_2(a520,X1)|c3_2(a520,X2)|c1_2(a520,a557)|c2_2(a520,X2)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20|~c7_2(a520,X1)|~c1_2(a520,X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_401]), c_0_380])]), c_0_433]), ['final']).
cnf(c_0_501,negated_conjecture,(ssSkC8|c7_2(a470,X1)|ssSkC2|~c8_2(a470,X1)), inference(spm,[status(thm)],[c_0_79, c_0_353]), ['final']).
cnf(c_0_502,negated_conjecture,(ssSkC8|c5_2(a536,X1)|c3_2(a536,X1)|ssSkC14|~c7_2(a536,X1)), inference(spm,[status(thm)],[c_0_80, c_0_353]), ['final']).
cnf(c_0_503,negated_conjecture,(c3_2(a520,a557)|c1_2(a520,a557)|c2_2(a520,a557)|c8_1(a520)|c8_1(X1)|c4_1(a520)|c3_1(X1)|c5_1(a565)|c7_1(a520)|ssSkC20|~c1_1(X1)), inference(ef,[status(thm)],[c_0_434]), ['final']).
cnf(c_0_504,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|ndr1_1(a559)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_349])])).
cnf(c_0_505,negated_conjecture,(ssSkC3|ssSkC8|c6_2(X1,a479)|c3_2(X1,a557)|c1_2(X1,a557)|c2_2(X1,a557)|c5_1(a565)|c7_1(X1)|ssSkC20|~c8_1(X1)), inference(ef,[status(thm)],[c_0_435]), ['final']).
cnf(c_0_506,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c3_2(a600,X1)|c1_2(a600,X1)|c2_2(a600,X1)|c8_1(a600)|c8_1(X2)|c4_1(a600)|c3_1(X2)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c1_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_394]), c_0_353]), ['final']).
cnf(c_0_507,negated_conjecture,(c3_2(a546,a557)|c1_2(a546,a557)|c2_2(a546,a557)|c8_1(a546)|c8_1(X1)|c3_1(X1)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c1_1(X1)), inference(ef,[status(thm)],[c_0_436]), ['final']).
cnf(c_0_508,negated_conjecture,(c3_2(a542,a557)|c1_2(a542,a557)|c2_2(a542,a557)|c8_1(a542)|c8_1(X1)|c4_1(a542)|c3_1(X1)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c1_1(X1)), inference(ef,[status(thm)],[c_0_437]), ['final']).
cnf(c_0_509,negated_conjecture,(c3_2(a512,a557)|c1_2(a512,a557)|c2_2(a512,a557)|c8_1(a512)|c8_1(X1)|c4_1(a512)|c3_1(X1)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c1_1(X1)), inference(ef,[status(thm)],[c_0_438]), ['final']).
cnf(c_0_510,negated_conjecture,(c7_2(X1,X2)|c4_1(X1)|c5_1(X3)|c7_1(X1)|c7_1(X3)|~c6_2(X1,X2)|~c4_2(X3,a583)|~c2_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_349])])).
cnf(c_0_511,negated_conjecture,(ssSkC8|c6_2(a491,X1)|c4_2(a491,X1)|c3_2(a491,X1)|ssSkC6), inference(spm,[status(thm)],[c_0_83, c_0_353]), ['final']).
cnf(c_0_512,negated_conjecture,(ssSkC8|c6_2(a592,a503)|c5_2(a592,a503)|ssSkP4(a592)|ssSkC24), inference(spm,[status(thm)],[c_0_364, c_0_62])).
cnf(c_0_513,negated_conjecture,(ssSkC8|c6_2(a568,X1)|c4_2(a568,X1)|ssSkC21), inference(spm,[status(thm)],[c_0_84, c_0_353]), ['final']).
cnf(c_0_514,negated_conjecture,(c5_2(X1,a460)|~ssSkC0|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_349])])).
cnf(c_0_515,negated_conjecture,(c4_2(a534,a535)|c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_349])])).
cnf(c_0_516,negated_conjecture,(c8_2(a534,a535)|c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_349])])).
cnf(c_0_517,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|c8_1(a534)|c6_1(X1)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_349])])).
cnf(c_0_518,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|ndr1_1(a534)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_349])])).
cnf(c_0_519,negated_conjecture,(ssSkC22|c4_2(X1,a576)|c7_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_349])])).
cnf(c_0_520,negated_conjecture,(ssSkC22|ndr1_1(X1)|c7_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_349])])).
cnf(c_0_521,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,X1)|c8_1(X2)|c4_1(a552)|c3_1(X2)|ssSkC18|~c1_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_363]), c_0_353]), c_0_68]), ['final']).
cnf(c_0_522,negated_conjecture,(ssSkC8|c1_2(a587,a557)|c1_2(a587,X1)|c2_2(a587,X1)|c5_1(a565)|c7_1(a587)|ssSkC20|ssSkC23|~c5_2(a587,X1)), inference(spm,[status(thm)],[c_0_439, c_0_357]), ['final']).
cnf(c_0_523,negated_conjecture,(c3_2(a498,a499)|c2_0), inference(sr,[status(thm)],[c_0_92, c_0_366])).
cnf(c_0_524,negated_conjecture,(c7_2(a498,a499)|c2_0), inference(sr,[status(thm)],[c_0_93, c_0_366])).
cnf(c_0_525,negated_conjecture,(ndr1_1(a498)|c2_0), inference(sr,[status(thm)],[c_0_94, c_0_366])).
cnf(c_0_526,negated_conjecture,(ssSkC12|c4_2(X1,a527)|c2_2(X1,a503)|ssSkP4(X1)|~ssSkP6(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_62]), c_0_9]), c_0_63]), ['final']).
cnf(c_0_527,negated_conjecture,(c8_2(X1,a460)|~ssSkC0|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_349])])).
cnf(c_0_528,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|ndr1_1(a574)|c5_1(X1)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_349])])).
cnf(c_0_529,negated_conjecture,(ssSkC8|c8_2(a564,X1)|c3_2(a564,X1)|ssSkC19|~c1_2(a564,X1)), inference(spm,[status(thm)],[c_0_97, c_0_353]), ['final']).
cnf(c_0_530,negated_conjecture,(c1_2(X1,a526)|c8_1(X2)|c8_1(X1)|ssSkP6(X1)|c4_1(X1)|c3_1(X2)|~c1_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_98]), c_0_31]), c_0_21]), ['final']).
cnf(c_0_531,negated_conjecture,(c3_2(X1,a456)|c1_2(X1,a456)|c8_1(X2)|c8_1(X1)|ssSkP0(X1)|c4_1(X1)|c3_1(X2)|~c1_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_99]), c_0_100]), ['final']).
cnf(c_0_532,negated_conjecture,(c3_2(a542,a543)|c1_2(a542,a543)|c8_1(a542)|c8_1(X1)|c4_1(a542)|c3_1(X1)|ssSkC16|~c1_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_54]), c_0_27]), ['final']).
cnf(c_0_533,negated_conjecture,(c8_2(a498,a500)|c2_0), inference(sr,[status(thm)],[c_0_101, c_0_366])).
cnf(c_0_534,negated_conjecture,(ssSkC5|ndr1_1(X1)|c5_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_349])])).
cnf(c_0_535,negated_conjecture,(ndr1_1(a482)|~ssSkC4), inference(sr,[status(thm)],[c_0_103, c_0_350])).
cnf(c_0_536,negated_conjecture,(ssSkC3|c6_2(X1,a479)|c1_2(X1,a526)|ssSkP6(X1)|~c8_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_98]), c_0_31]), c_0_21]), ['final']).
cnf(c_0_537,negated_conjecture,(ssSkC3|c6_2(X1,a479)|c3_2(X1,a456)|c1_2(X1,a456)|ssSkP0(X1)|~c8_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_99]), c_0_100]), ['final']).
cnf(c_0_538,negated_conjecture,(c3_2(a496,X1)|~c8_2(a496,X1)|~c2_2(a496,X1)|~c4_0|~c6_0), inference(csr,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_539,negated_conjecture,(ssSkC0|c7_2(X1,a459)|~c8_1(X1)|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_349])])).
cnf(c_0_540,negated_conjecture,(ssSkC0|ndr1_1(X1)|~c8_1(X1)|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_349])])).
cnf(c_0_541,negated_conjecture,(ndr1_1(X1)|~ssSkC0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_440, c_0_441])])).
cnf(c_0_542,negated_conjecture,(ndr1_1(a544)|~c8_1(X1)|~c2_1(X1)|~ssSkC16), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_349])])).
fof(c_0_543, plain, ((~epred1_0<=>![X2]:![X1]:((((~ndr1_1(X1)|~c7_0)|~c1_2(X1,X2))|~c8_2(X1,X2))|~c4_2(X1,X2)))), introduced(definition)).
cnf(c_0_544,negated_conjecture,(c3_1(X1)|~ssSkC3|~c6_2(X1,a480)|~c2_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_349])])).
cnf(c_0_545,negated_conjecture,(c3_2(X1,a480)|c3_1(X1)|~ssSkC3|~c2_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_349])])).
cnf(c_0_546,negated_conjecture,(c4_2(X1,a480)|c3_1(X1)|~ssSkC3|~c2_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_349])])).
fof(c_0_547, plain, ((~epred2_0<=>![X3]:(~c3_1(X3)|~c3_2(X3,a473)))), introduced(definition)).
cnf(c_0_548,negated_conjecture,(ndr1_1(X1)|c3_1(X1)|~ssSkC3|~c2_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_349])])).
cnf(c_0_549,negated_conjecture,(~c4_2(X1,X2)|~c7_2(X1,X2)|~c3_2(X1,X2)|~c2_1(a508)|~c4_1(X1)|~ndr1_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_349])])).
cnf(c_0_550,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a557)|c1_2(X1,a493)|c2_2(X1,X2)|ssSkC7|c5_1(a565)|c7_1(X1)|ssSkC20|~c5_2(X1,X2)|~ssSkP2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_357]), c_0_353]), ['final']).
cnf(c_0_551,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,a515)|c5_2(a565,a515)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|ssSkP6(a565)|c1_1(a565)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_417, c_0_442]), ['final']).
cnf(c_0_552,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,a515)|c5_2(a565,a515)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_419, c_0_442]), ['final']).
cnf(c_0_553,negated_conjecture,(ssSkC8|c6_2(a517,a588)|c6_2(a517,X1)|c7_2(a517,X1)|c3_2(a517,a588)|c3_2(a517,X1)|c1_2(a517,a557)|c2_2(a517,X1)|c1_1(a517)|ssSkP8(a517)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20), inference(spm,[status(thm)],[c_0_443, c_0_382]), ['final']).
cnf(c_0_554,negated_conjecture,(c6_2(a542,a588)|c6_2(a542,X1)|c7_2(a542,X1)|c3_2(a542,a588)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|c1_1(a542)|ssSkP8(a542)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_444, c_0_384]), ['final']).
cnf(c_0_555,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,a515)|c5_2(a565,a515)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|c1_1(a565)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_417, c_0_445]), ['final']).
cnf(c_0_556,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,a515)|c5_2(a565,a515)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|c1_1(a565)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_419, c_0_445]), ['final']).
cnf(c_0_557,negated_conjecture,(c6_2(a542,a487)|c6_2(a542,X1)|c7_2(a542,X1)|c3_2(a542,a487)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|c1_1(a542)|ssSkC5|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_446, c_0_384]), ['final']).
cnf(c_0_558,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c6_2(a517,X1)|c7_2(a517,X1)|c3_2(a517,X1)|c1_2(a517,a557)|c2_2(a517,X1)|ssSkP1(a517)|c1_1(a517)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20), inference(spm,[status(thm)],[c_0_447, c_0_382]), ['final']).
cnf(c_0_559,negated_conjecture,(ssSkC8|c6_2(a565,a588)|c6_2(a565,X1)|c7_2(a565,X1)|c3_2(a565,a588)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|c1_1(a565)|ssSkP8(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_448, c_0_368]), ['final']).
cnf(c_0_560,negated_conjecture,(ssSkC8|c6_2(a517,a526)|c6_2(a517,X1)|c7_2(a517,X1)|c3_2(a517,X1)|c1_2(a517,a557)|c2_2(a517,X1)|ssSkP6(a517)|c1_1(a517)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20), inference(spm,[status(thm)],[c_0_367, c_0_382]), ['final']).
cnf(c_0_561,negated_conjecture,(ssSkC8|c6_2(a517,a515)|c6_2(a517,X1)|c7_2(a517,X1)|c3_2(a517,X1)|c1_2(a517,a557)|c2_2(a517,X1)|ssSkP5(a517)|c1_1(a517)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20), inference(spm,[status(thm)],[c_0_370, c_0_382]), ['final']).
cnf(c_0_562,negated_conjecture,(c1_2(X1,a493)|ssSkP6(X1)|ssSkC7|~c5_2(X1,a526)|~ssSkP2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_98]), c_0_31]), c_0_21]), ['final']).
cnf(c_0_563,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,a526)|c5_2(a565,a526)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_449, c_0_445]), ['final']).
cnf(c_0_564,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,a526)|c5_2(a565,a526)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|ssSkP6(a565)|c1_1(a565)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_450, c_0_445]), ['final']).
cnf(c_0_565,negated_conjecture,(c6_2(a542,a485)|c6_2(a542,X1)|c7_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|ssSkP1(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_451, c_0_384]), ['final']).
cnf(c_0_566,negated_conjecture,(c6_2(a542,a526)|c6_2(a542,X1)|c7_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|ssSkP6(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_452, c_0_384]), ['final']).
cnf(c_0_567,negated_conjecture,(c6_2(a542,a485)|c3_2(a542,a588)|c3_2(a542,a493)|c1_2(a542,a557)|c2_2(a542,a588)|c2_2(a542,a493)|ssSkP1(a542)|c1_1(a542)|ssSkP8(a542)|ssSkC7|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_451, c_0_453]), ['final']).
cnf(c_0_568,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c4_2(a552,a515)|c5_2(a552,a515)|c1_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP5(a552)|ssSkP6(a552)|c1_1(a573)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_454]), c_0_68]), ['final']).
cnf(c_0_569,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,a526)|c5_2(a565,a526)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_449, c_0_442]), ['final']).
cnf(c_0_570,negated_conjecture,(c6_2(a542,a515)|c6_2(a542,X1)|c7_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|ssSkP5(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_455, c_0_384]), ['final']).
cnf(c_0_571,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c4_2(a552,a515)|c5_2(a552,a515)|c1_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP5(a552)|ssSkP6(a552)|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_454]), c_0_68]), ['final']).
cnf(c_0_572,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c6_2(a565,X1)|c7_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|ssSkP1(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_456, c_0_368]), ['final']).
cnf(c_0_573,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,a526)|c5_2(a565,a526)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP6(a565)|c1_1(a565)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_450, c_0_442]), ['final']).
cnf(c_0_574,negated_conjecture,(c6_2(a542,a588)|c3_2(a542,a493)|c3_2(a542,a588)|c1_2(a542,a557)|c2_2(a542,a588)|c2_2(a542,a493)|c1_1(a542)|ssSkP8(a542)|ssSkC7|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_444, c_0_453]), ['final']).
cnf(c_0_575,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c4_2(a552,a515)|c5_2(a552,a515)|c1_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP5(a552)|c1_1(a573)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_457]), c_0_68]), ['final']).
cnf(c_0_576,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c4_2(a552,a515)|c5_2(a552,a515)|c1_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP5(a552)|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_457]), c_0_68]), ['final']).
cnf(c_0_577,negated_conjecture,(c6_2(a542,a487)|c3_2(a542,a588)|c3_2(a542,a493)|c3_2(a542,a487)|c1_2(a542,a557)|c2_2(a542,a588)|c2_2(a542,a493)|ssSkP8(a542)|ssSkC5|ssSkC7|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_446, c_0_453]), ['final']).
cnf(c_0_578,negated_conjecture,(c6_2(a542,a465)|c7_2(a542,a465)|c7_2(X1,X2)|c5_2(a542,X3)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X4,X5)|c8_1(a542)|c8_1(X4)|c1_1(a542)|c1_1(X1)|c2_1(X1)|c7_0|c5_1(a501)|ssSkC16|~c6_2(a542,X3)|~c6_2(X4,X5)|~c7_2(a542,X3)|~c2_2(X4,X5)|~c1_1(X4)|~ndr1_1(X4)|~ndr1_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_458, c_0_384]), c_0_27]), ['final']).
cnf(c_0_579,negated_conjecture,(ssSkC8|c6_2(a565,a465)|c7_2(a565,a465)|c7_2(X1,X2)|c5_2(a565,X3)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X4,X5)|c8_1(a565)|c8_1(X4)|c1_1(a565)|c1_1(X1)|c2_1(X1)|c7_0|c5_1(a501)|ssSkC20|~c6_2(a565,X3)|~c6_2(X4,X5)|~c7_2(a565,X3)|~c2_2(X4,X5)|~c1_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_458, c_0_368]), c_0_353]), c_0_353]), c_0_353]), ['final']).
cnf(c_0_580,negated_conjecture,(ssSkC8|c6_2(a517,a465)|c7_2(a517,a465)|c7_2(X1,X2)|c5_2(a517,X3)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X4,X5)|c8_1(a517)|c8_1(X4)|c1_1(a517)|c1_1(X1)|c2_1(X1)|c7_0|c5_1(a501)|ssSkC11|~c6_2(a517,X3)|~c6_2(X4,X5)|~c7_2(a517,X3)|~c2_2(X4,X5)|~c1_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_458, c_0_382]), c_0_353]), c_0_353]), c_0_353]), ['final']).
cnf(c_0_581,negated_conjecture,(c1_2(X1,a460)|~ssSkC0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_459, c_0_441])])).
cnf(c_0_582,negated_conjecture,(c6_2(a542,a515)|c3_2(a542,a588)|c3_2(a542,a493)|c1_2(a542,a557)|c2_2(a542,a588)|c2_2(a542,a493)|ssSkP5(a542)|ssSkP8(a542)|ssSkC7|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_455, c_0_453]), ['final']).
cnf(c_0_583,negated_conjecture,(c8_2(X1,a467)|c8_2(X2,X3)|c8_1(X2)|~c4_2(X2,X3)|~c7_2(X2,X3)|~c8_1(X1)|~c2_1(X2)|~c6_1(X1)|~ndr1_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_349])])).
cnf(c_0_584,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c4_2(a592,a515)|c1_2(a592,a557)|c2_2(a592,a557)|c2_2(a592,a502)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c1_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_460]), c_0_24]), ['final']).
cnf(c_0_585,negated_conjecture,(c8_2(X1,X2)|c3_2(X3,a467)|c8_1(X1)|~c4_2(X1,X2)|~c7_2(X1,X2)|~c8_1(X3)|~c2_1(X1)|~c6_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_349])])).
cnf(c_0_586,negated_conjecture,(c8_2(X1,X2)|c2_2(X3,a467)|c8_1(X1)|~c4_2(X1,X2)|~c7_2(X1,X2)|~c8_1(X3)|~c2_1(X1)|~c6_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_349])])).
cnf(c_0_587,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c4_2(a592,a515)|c1_2(a592,a557)|c2_2(a592,a557)|c2_2(a592,a502)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c3_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_461]), c_0_24]), ['final']).
cnf(c_0_588,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c4_2(a552,a526)|c5_2(a552,a526)|c1_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP5(a552)|ssSkP6(a552)|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_457]), c_0_68]), ['final']).
cnf(c_0_589,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c4_2(a552,a526)|c5_2(a552,a526)|c1_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP5(a552)|ssSkP6(a552)|c1_1(a573)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_457]), c_0_68]), ['final']).
cnf(c_0_590,negated_conjecture,(c7_2(X1,X2)|c5_2(a570,X3)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X4,X5)|c2_2(a570,a465)|c2_2(a570,a572)|c8_1(a570)|c8_1(X4)|c1_1(X1)|c2_1(X1)|~c6_2(a570,X3)|~c6_2(X4,X5)|~c7_2(a570,X3)|~c2_2(X4,X5)|~c1_1(X4)|~ndr1_1(X4)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_462, c_0_463]), c_0_421])]), ['final']).
cnf(c_0_591,negated_conjecture,(ssSkC8|c6_2(a600,a526)|c7_2(a600,a557)|c2_2(a600,a503)|ssSkP4(a600)|ssSkP6(a600)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_62]), c_0_63]), ['final']).
cnf(c_0_592,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c3_2(a552,a557)|c3_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP6(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c1_2(a552,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_465, c_0_20]), c_0_21]), ['final']).
cnf(c_0_593,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c3_2(a552,a557)|c3_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP5(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c1_2(a552,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_466, c_0_23]), c_0_24]), ['final']).
cnf(c_0_594,negated_conjecture,(ssSkC8|c6_2(a600,a515)|c7_2(a600,a557)|c2_2(a600,a503)|ssSkP4(a600)|ssSkP5(a600)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_467, c_0_62]), c_0_63]), ['final']).
cnf(c_0_595,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP6(a552)|c4_1(a552)|c4_1(a531)|c3_1(a466)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_117]), c_0_118]), ['final']).
cnf(c_0_596,negated_conjecture,(c2_0|~c7_0), inference(csr,[status(thm)],[c_0_119, c_0_355])).
cnf(c_0_597,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c4_2(a552,a526)|c5_2(a552,a526)|c1_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP6(a552)|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_454]), c_0_68]), ['final']).
cnf(c_0_598,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c4_2(a552,a526)|c5_2(a552,a526)|c1_2(a552,a557)|c2_2(a552,a502)|c2_2(a552,a557)|ssSkP3(a552)|ssSkP6(a552)|c1_1(a573)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_454]), c_0_68]), ['final']).
cnf(c_0_599,negated_conjecture,(c1_2(a517,a487)|ssSkC5|c5_1(a517)|ssSkC11), inference(spm,[status(thm)],[c_0_469, c_0_22]), ['final']).
cnf(c_0_600,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP5(a552)|c4_1(a552)|c4_1(a531)|c3_1(a466)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_470, c_0_117]), c_0_118]), ['final']).
cnf(c_0_601,negated_conjecture,(ssSkC8|c7_2(X1,X2)|c5_2(a537,X3)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X4,X5)|c2_2(a537,a465)|c2_2(a537,a572)|c8_1(a537)|c8_1(X4)|c1_1(X1)|c2_1(X1)|ssSkC15|~c6_2(a537,X3)|~c6_2(X4,X5)|~c7_2(a537,X3)|~c2_2(X4,X5)|~c1_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_462, c_0_471]), c_0_353]), c_0_353]), c_0_353]), ['final']).
cnf(c_0_602,negated_conjecture,(ssSkC8|c7_2(X1,X2)|c5_2(a517,X3)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X4,X5)|c2_2(a517,a465)|c2_2(a517,a572)|c8_1(a517)|c8_1(X4)|c1_1(X1)|c2_1(X1)|ssSkC11|~c6_2(a517,X3)|~c6_2(X4,X5)|~c7_2(a517,X3)|~c2_2(X4,X5)|~c1_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_462, c_0_472]), c_0_353]), c_0_353]), c_0_353]), ['final']).
cnf(c_0_603,negated_conjecture,(c3_2(X1,X2)|~c6_2(X1,X2)|~c4_2(X1,a548)|~c8_2(X1,X2)|~c1_1(X1)|~ndr1_1(X1)|~ssSkC17), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_349])])).
cnf(c_0_604,negated_conjecture,(ssSkP6(X1)|~c6_2(X1,a526)|~c7_2(X1,a548)|~c1_1(X1)|~ssSkC17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_98]), c_0_31]), c_0_21]), ['final']).
cnf(c_0_605,negated_conjecture,(c1_2(a542,a487)|ssSkC5|c5_1(a542)|ssSkC16), inference(spm,[status(thm)],[c_0_469, c_0_34]), ['final']).
cnf(c_0_606,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c2_1(a552)|c4_1(a552)|c4_1(a458)|c3_1(a466)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_474, c_0_117]), c_0_118]), ['final']).
cnf(c_0_607,negated_conjecture,(ssSkC8|c6_2(a600,a503)|c7_2(a600,a557)|c2_2(a600,a503)|ssSkP4(a600)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_475, c_0_62]), c_0_63]), ['final']).
cnf(c_0_608,negated_conjecture,(c6_2(a542,a588)|c6_2(a542,a543)|c7_2(a542,a543)|c3_2(a542,a588)|c3_2(a542,a543)|c1_1(a542)|ssSkP8(a542)|c7_0|c4_1(a531)|c5_1(a542)|c5_1(a501)|ssSkC16), inference(spm,[status(thm)],[c_0_476, c_0_407]), ['final']).
cnf(c_0_609,negated_conjecture,(ssSkC8|c6_2(a600,a557)|c3_2(a600,a557)|c3_2(a600,X1)|c1_2(a600,a557)|c2_2(a600,a502)|c2_2(a600,X1)|ssSkP3(a600)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c1_2(a600,X1)), inference(spm,[status(thm)],[c_0_361, c_0_477]), ['final']).
cnf(c_0_610,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c1_2(a565,X1)), inference(spm,[status(thm)],[c_0_478, c_0_381]), ['final']).
cnf(c_0_611,negated_conjecture,(c6_2(a542,a485)|c6_2(a542,a543)|c7_2(a542,a543)|c3_2(a542,a543)|ssSkP1(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a542)|c5_1(a501)|ssSkC16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_412]), c_0_64]), ['final']).
cnf(c_0_612,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,X1)|ssSkP5(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c1_2(a565,X1)), inference(spm,[status(thm)],[c_0_478, c_0_385]), ['final']).
cnf(c_0_613,negated_conjecture,(c1_2(a565,a487)|ssSkC5|c5_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_469, c_0_16]), ['final']).
cnf(c_0_614,negated_conjecture,(c6_2(a542,a526)|c6_2(a542,a543)|c7_2(a542,a543)|c3_2(a542,a543)|ssSkP6(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a542)|c5_1(a501)|ssSkC16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_20]), c_0_21]), ['final']).
cnf(c_0_615,negated_conjecture,(ssSkC8|c6_2(a600,a503)|c3_2(a600,X1)|c1_2(a600,a557)|c2_2(a600,X1)|ssSkP4(a600)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c1_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_479]), c_0_63]), ['final']).
cnf(c_0_616,negated_conjecture,(c6_2(a542,a515)|c6_2(a542,a543)|c7_2(a542,a543)|c3_2(a542,a543)|ssSkP5(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a542)|c5_1(a501)|ssSkC16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_23]), c_0_24]), ['final']).
cnf(c_0_617,negated_conjecture,(ssSkC8|c6_2(X1,a526)|c1_2(X1,a557)|c2_2(X1,a503)|ssSkP4(X1)|ssSkP6(X1)|c4_1(a531)|c5_1(a565)|c5_1(X1)|c7_1(X1)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_62]), c_0_63]), ['final']).
cnf(c_0_618,negated_conjecture,(c6_2(a542,a487)|c6_2(a542,a543)|c7_2(a542,a543)|c3_2(a542,a487)|c3_2(a542,a543)|c1_1(a542)|ssSkC5|c7_0|c4_1(a531)|c5_1(a542)|c5_1(a501)|ssSkC16), inference(spm,[status(thm)],[c_0_476, c_0_409]), ['final']).
cnf(c_0_619,negated_conjecture,(ssSkC8|c6_2(X1,a515)|c1_2(X1,a557)|c2_2(X1,a503)|ssSkP4(X1)|ssSkP5(X1)|c4_1(a531)|c5_1(a565)|c5_1(X1)|c7_1(X1)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_370, c_0_62]), c_0_63]), ['final']).
cnf(c_0_620,negated_conjecture,(ssSkC8|c6_2(a592,a526)|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|ssSkP6(a592)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c1_2(a592,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_20]), c_0_21]), ['final']).
cnf(c_0_621,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c1_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c8_1(a512)|c2_1(a512)|c4_1(a512)|c4_1(a458)|c3_1(a466)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_117]), c_0_118]), ['final']).
cnf(c_0_622,negated_conjecture,(ssSkC8|c6_2(a592,a515)|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|ssSkP5(a592)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c1_2(a592,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_23]), c_0_24]), ['final']).
cnf(c_0_623,negated_conjecture,(c6_2(X1,a474)|c4_2(X2,X3)|c1_2(X4,a456)|c8_1(X2)|ssSkP0(X4)|c6_1(X4)|~c6_2(X2,X3)|~c3_2(X2,X3)|~ndr1_1(X2)|~c3_1(X2)|~c5_1(X1)|~c7_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_624,negated_conjecture,(c1_2(X1,X2)|c2_2(X1,X2)|c2_2(X1,X3)|c1_1(X4)|c2_1(X4)|c2_0|~c7_2(X1,X3)|~c5_2(X1,X2)|~c1_2(X1,X3)|~c4_1(X4)|~ndr1_1(X1)|~c5_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_349])])).
cnf(c_0_625,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c6_2(a592,a526)|c5_2(a592,a557)|c2_2(a592,a503)|ssSkP4(a592)|ssSkP6(a592)|c4_1(a531)|c5_1(a565)|c5_1(a592)|c7_1(a592)|ssSkC20|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_483, c_0_62]), c_0_63]), ['final']).
cnf(c_0_626,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c6_2(a592,a515)|c5_2(a592,a557)|c2_2(a592,a503)|ssSkP4(a592)|ssSkP5(a592)|c4_1(a531)|c5_1(a565)|c5_1(a592)|c7_1(a592)|ssSkC20|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_484, c_0_62]), c_0_63]), ['final']).
cnf(c_0_627,negated_conjecture,(c4_2(X1,X2)|c7_2(X3,a474)|c1_2(X4,a456)|c8_1(X1)|ssSkP0(X4)|c6_1(X4)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_628,negated_conjecture,(c6_2(a542,a502)|ssSkP3(a542)|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_384]), c_0_125]), ['final']).
cnf(c_0_629,negated_conjecture,(c1_1(a508)|~c4_2(X1,X2)|~c7_2(X1,X2)|~c3_2(X1,X2)|~c4_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_486, c_0_441])])).
cnf(c_0_630,negated_conjecture,(c4_2(X1,X2)|c1_2(X3,a456)|c1_2(X4,a474)|c8_1(X1)|ssSkP0(X3)|c6_1(X3)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X4)|~c7_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_631,negated_conjecture,(ssSkC8|c4_2(a573,X1)|c5_2(a573,X1)|c1_2(a573,X2)|c2_2(a573,X2)|c8_1(a573)|~c6_2(a573,X1)|~c7_2(a573,X1)|~c3_2(a573,X2)|~c2_2(a573,X1)), inference(ef,[status(thm)],[c_0_488]), ['final']).
cnf(c_0_632,negated_conjecture,(c4_2(X1,X2)|c1_2(X3,a456)|c8_1(X1)|ssSkP0(X3)|c6_1(X3)|ndr1_1(X4)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X4)|~c7_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_633,negated_conjecture,(ssSkC8|c6_2(a600,a503)|c7_2(a600,a557)|c3_2(a600,X1)|c2_2(a600,X1)|ssSkP4(a600)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c1_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_479]), c_0_63]), ['final']).
cnf(c_0_634,negated_conjecture,(ssSkC8|c6_2(a565,a502)|ssSkP3(a565)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_368]), c_0_125]), ['final']).
cnf(c_0_635,negated_conjecture,(ssSkC8|c6_2(a517,a502)|ssSkP3(a517)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_382]), c_0_125]), ['final']).
cnf(c_0_636,negated_conjecture,(ssSkC8|c3_2(X1,a557)|c1_2(X1,a557)|c2_2(X1,a557)|c8_1(X1)|c4_1(X1)|c3_1(a466)|c5_1(a565)|c7_1(X1)|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_490, c_0_117]), c_0_118]), ['final']).
cnf(c_0_637,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a557)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(X1)|c4_1(X1)|c3_1(a466)|c5_1(a565)|c7_1(X1)|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_117]), c_0_118]), ['final']).
cnf(c_0_638,negated_conjecture,(c6_2(X1,a474)|c4_2(X2,X3)|c1_2(X2,X3)|c8_1(X2)|c6_1(X2)|~c6_2(X2,X3)|~c3_2(X2,X3)|~ndr1_1(X2)|~c3_1(X2)|~c5_1(X1)|~c7_1(X2)), inference(ef,[status(thm)],[c_0_482]), ['final']).
cnf(c_0_639,negated_conjecture,(c4_2(X1,X2)|c7_2(X3,a474)|c1_2(X1,X2)|c8_1(X1)|c6_1(X1)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(X1)), inference(ef,[status(thm)],[c_0_485]), ['final']).
cnf(c_0_640,negated_conjecture,(c4_2(X1,X2)|c1_2(X3,a474)|c1_2(X1,X2)|c8_1(X1)|c6_1(X1)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(X1)), inference(ef,[status(thm)],[c_0_487]), ['final']).
cnf(c_0_641,negated_conjecture,(c6_2(a546,a526)|c3_2(a546,X1)|c1_2(a546,a557)|c2_2(a546,X1)|ssSkP6(a546)|c4_1(a531)|c5_1(a546)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c1_2(a546,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_20]), c_0_21]), ['final']).
cnf(c_0_642,negated_conjecture,(c6_2(a546,a515)|c3_2(a546,X1)|c1_2(a546,a557)|c2_2(a546,X1)|ssSkP5(a546)|c4_1(a531)|c5_1(a546)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c1_2(a546,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_23]), c_0_24]), ['final']).
cnf(c_0_643,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c6_2(a592,X1)|c5_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|c4_1(a531)|c5_1(a565)|c5_1(a592)|c7_1(a592)|ssSkC20|ssSkC24|~c7_2(a592,X1)|~c1_2(a592,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_375]), c_0_353]), ['final']).
cnf(c_0_644,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c6_2(a592,X1)|c3_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|c4_1(a531)|c5_1(a565)|c5_1(a592)|c7_1(a592)|ssSkC20|ssSkC24|~c7_2(a592,X1)|~c1_2(a592,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_392]), c_0_353]), ['final']).
cnf(c_0_645,negated_conjecture,(c6_2(a512,a526)|c3_2(a512,X1)|c1_2(a512,a557)|c2_2(a512,X1)|ssSkP6(a512)|c4_1(a531)|c5_1(a512)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c1_2(a512,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_20]), c_0_21]), ['final']).
cnf(c_0_646,negated_conjecture,(c6_2(a512,a515)|c3_2(a512,X1)|c1_2(a512,a557)|c2_2(a512,X1)|ssSkP5(a512)|c4_1(a531)|c5_1(a512)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c1_2(a512,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_23]), c_0_24]), ['final']).
cnf(c_0_647,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a592)|c4_1(a592)|c3_1(a466)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_117]), c_0_118]), ['final']).
cnf(c_0_648,negated_conjecture,(c6_2(a546,a547)|c3_2(a546,a547)|c3_2(a546,X1)|c1_2(a546,a557)|c2_2(a546,X1)|c4_1(a531)|c5_1(a546)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c1_2(a546,X1)), inference(spm,[status(thm)],[c_0_491, c_0_126]), ['final']).
cnf(c_0_649,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a592)|c4_1(a592)|c3_1(a466)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_117]), c_0_118]), ['final']).
cnf(c_0_650,negated_conjecture,(c4_2(X1,X2)|c1_2(X1,X2)|c8_1(X1)|c6_1(X1)|ndr1_1(X3)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(X1)), inference(ef,[status(thm)],[c_0_489]), ['final']).
cnf(c_0_651,negated_conjecture,(ssSkC8|c4_2(a592,a515)|c2_2(a592,a588)|c8_1(a592)|ssSkP5(a592)|c1_1(a573)|ssSkP8(a592)|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_495]), c_0_24]), ['final']).
cnf(c_0_652,negated_conjecture,(c4_2(a559,a560)|c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_496, c_0_360])]), ['final']).
cnf(c_0_653,negated_conjecture,(c5_2(X1,X2)|~c4_2(X1,X2)|~c2_2(X1,a516)|~ssSkP5(X1)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_349])])).
cnf(c_0_654,negated_conjecture,(ssSkC8|c4_2(a592,a515)|c2_2(a592,a588)|c8_1(a592)|ssSkP5(a592)|ssSkP8(a592)|c3_1(a573)|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_497]), c_0_24]), ['final']).
cnf(c_0_655,negated_conjecture,(c8_2(X1,X2)|c7_2(a559,a560)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_498, c_0_360])]), ['final']).
cnf(c_0_656,negated_conjecture,(c5_2(X1,X2)|~c6_2(X1,a516)|~c4_2(X1,X2)|~ssSkP5(X1)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_349])])).
cnf(c_0_657,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|c7_1(a559)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_499, c_0_360])]), ['final']).
cnf(c_0_658,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c4_2(a592,a526)|c5_2(a592,a526)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c8_1(a592)|ssSkP3(a592)|ssSkP6(a592)|c1_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(spm,[status(thm)],[c_0_450, c_0_418])).
cnf(c_0_659,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c4_2(a592,a526)|c5_2(a592,a526)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c8_1(a592)|ssSkP3(a592)|ssSkP6(a592)|c3_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(spm,[status(thm)],[c_0_449, c_0_418])).
cnf(c_0_660,negated_conjecture,(ndr1_1(a505)|~c4_0), inference(csr,[status(thm)],[c_0_129, c_0_353]), ['final']).
cnf(c_0_661,negated_conjecture,(c6_2(a496,a497)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_441])])).
cnf(c_0_662,negated_conjecture,(ndr1_1(a496)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_441])])).
cnf(c_0_663,negated_conjecture,(c3_2(a496,a497)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_441])])).
cnf(c_0_664,negated_conjecture,(~c4_2(a496,a497)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_132, c_0_441])])).
cnf(c_0_665,negated_conjecture,(ssSkC22|c7_1(X1)|~c8_2(X1,a576)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_349])])).
cnf(c_0_666,negated_conjecture,(c6_2(a520,a526)|c3_2(a520,X1)|c1_2(a520,a557)|c2_2(a520,X1)|ssSkP6(a520)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20|~c1_2(a520,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_20]), c_0_21]), ['final']).
cnf(c_0_667,negated_conjecture,(c6_2(a520,a515)|c3_2(a520,X1)|c1_2(a520,a557)|c2_2(a520,X1)|ssSkP5(a520)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20|~c1_2(a520,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_23]), c_0_24]), ['final']).
cnf(c_0_668,negated_conjecture,(ssSkC8|c7_2(a470,a456)|ssSkP0(a470)|ssSkC2), inference(spm,[status(thm)],[c_0_501, c_0_99]), ['final']).
cnf(c_0_669,negated_conjecture,(c4_2(a512,a515)|c5_2(a512,a515)|c8_1(a512)|ssSkP5(a512)|c3_1(a573)|ssSkC9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_134]), c_0_135]), c_0_136]), ['final']).
cnf(c_0_670,negated_conjecture,(c4_2(a512,a515)|c5_2(a512,a515)|c8_1(a512)|ssSkP5(a512)|c1_1(a573)|ssSkC9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_134]), c_0_135]), c_0_136]), ['final']).
cnf(c_0_671,negated_conjecture,(ssSkC8|c5_2(a536,a515)|ssSkP5(a536)|ssSkC14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_23]), c_0_24]), ['final']).
cnf(c_0_672,negated_conjecture,(c5_2(a482,a484)|~ssSkC4), inference(sr,[status(thm)],[c_0_137, c_0_350])).
cnf(c_0_673,negated_conjecture,(c3_2(a520,X1)|c1_2(a520,a557)|c1_2(a520,X1)|c2_2(a520,X1)|c8_1(a520)|c4_1(a520)|c3_1(a466)|c5_1(a565)|c7_1(a520)|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_434, c_0_117]), c_0_118]), ['final']).
cnf(c_0_674,negated_conjecture,(c3_2(a520,a557)|c1_2(a520,a557)|c2_2(a520,a557)|c8_1(a520)|c4_1(a520)|c3_1(a466)|c5_1(a565)|c7_1(a520)|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_503, c_0_117]), c_0_118]), ['final']).
cnf(c_0_675,negated_conjecture,(c6_2(a546,X1)|c3_2(a546,X1)|c1_2(a546,a557)|c2_2(a546,X1)|c4_1(a531)|c5_1(a565)|c5_1(a546)|c7_1(a546)|ssSkC17|ssSkC20|~c7_2(a546,X1)|~c1_2(a546,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_403]), c_0_45]), ['final']).
cnf(c_0_676,negated_conjecture,(c6_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20|~c7_2(a542,X1)|~c1_2(a542,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_383]), c_0_27]), ['final']).
cnf(c_0_677,negated_conjecture,(c6_2(a512,X1)|c3_2(a512,X1)|c1_2(a512,a557)|c2_2(a512,X1)|c4_1(a531)|c5_1(a565)|c5_1(a512)|c7_1(a512)|ssSkC9|ssSkC20|~c7_2(a512,X1)|~c1_2(a512,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_379]), c_0_32]), ['final']).
cnf(c_0_678,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|ndr1_1(a559)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_504, c_0_360])]), ['final']).
cnf(c_0_679,negated_conjecture,(ssSkC3|ssSkC8|c6_2(a581,a479)|c3_2(a581,a557)|c1_2(a581,a557)|c2_2(a581,a557)|c5_1(a565)|c7_1(a581)|ssSkC20|~c2_0), inference(spm,[status(thm)],[c_0_505, c_0_138]), ['final']).
cnf(c_0_680,negated_conjecture,(ssSkC3|ssSkC8|c6_2(a581,a479)|c3_2(a581,X1)|c1_2(a581,a557)|c1_2(a581,X1)|c2_2(a581,X1)|c5_1(a565)|c7_1(a581)|ssSkC20|~c2_0), inference(spm,[status(thm)],[c_0_435, c_0_138]), ['final']).
cnf(c_0_681,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c3_2(a600,X1)|c1_2(a600,X1)|c2_2(a600,X1)|c8_1(a600)|c4_1(a600)|c3_1(a466)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_506, c_0_117]), c_0_118]), ['final']).
cnf(c_0_682,negated_conjecture,(c6_2(a542,a543)|c3_2(a542,a543)|c4_1(a531)|c5_1(a542)|ssSkC16|~c7_2(a542,a543)|~c1_2(a542,a543)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_54]), c_0_27]), ['final']).
cnf(c_0_683,negated_conjecture,(c3_2(a546,a557)|c1_2(a546,a557)|c2_2(a546,a557)|c8_1(a546)|c3_1(a466)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_507, c_0_117]), c_0_118]), ['final']).
cnf(c_0_684,negated_conjecture,(c3_2(a546,X1)|c1_2(a546,a557)|c1_2(a546,X1)|c2_2(a546,X1)|c8_1(a546)|c3_1(a466)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_436, c_0_117]), c_0_118]), ['final']).
cnf(c_0_685,negated_conjecture,(c3_2(a542,a557)|c1_2(a542,a557)|c2_2(a542,a557)|c8_1(a542)|c4_1(a542)|c3_1(a466)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_508, c_0_117]), c_0_118]), ['final']).
cnf(c_0_686,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|c4_1(a542)|c3_1(a466)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_437, c_0_117]), c_0_118]), ['final']).
cnf(c_0_687,negated_conjecture,(c3_2(a512,a557)|c1_2(a512,a557)|c2_2(a512,a557)|c8_1(a512)|c4_1(a512)|c3_1(a466)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_509, c_0_117]), c_0_118]), ['final']).
cnf(c_0_688,negated_conjecture,(c3_2(a512,X1)|c1_2(a512,a557)|c1_2(a512,X1)|c2_2(a512,X1)|c8_1(a512)|c4_1(a512)|c3_1(a466)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_438, c_0_117]), c_0_118]), ['final']).
cnf(c_0_689,negated_conjecture,(c6_2(a520,X1)|c3_2(a520,X1)|c1_2(a520,a557)|c2_2(a520,X1)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20|~c7_2(a520,X1)|~c1_2(a520,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_401]), c_0_380])]), c_0_433]), ['final']).
cnf(c_0_690,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c7_2(X1,X2)|c3_2(a491,a583)|c4_1(X1)|c7_1(a491)|c7_1(X1)|ssSkC6|~c6_2(X1,X2)|~c2_2(X1,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_510, c_0_511]), c_0_353]), c_0_139]), ['final']).
cnf(c_0_691,negated_conjecture,(ssSkC8|c6_2(a592,a503)|ssSkP4(a592)|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_512]), c_0_63]), ['final']).
cnf(c_0_692,negated_conjecture,(ssSkC8|c6_2(a568,a503)|ssSkP4(a568)|ssSkC21), inference(spm,[status(thm)],[c_0_140, c_0_513]), ['final']).
cnf(c_0_693,negated_conjecture,(ssSkC8|c6_2(a491,a503)|ssSkP4(a491)|ssSkC6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_511]), c_0_63]), ['final']).
cnf(c_0_694,negated_conjecture,(c6_2(a549,a550)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_141, c_0_360])]), ['final']).
cnf(c_0_695,negated_conjecture,(c7_2(X1,a583)|c7_2(X2,X3)|c4_1(X2)|c5_1(X1)|c7_1(X2)|c7_1(X1)|~c6_2(X2,X3)|~c2_2(X2,X3)|~ndr1_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_142, c_0_349])])).
cnf(c_0_696,negated_conjecture,(c7_2(X1,X2)|c2_2(X3,a583)|c4_1(X1)|c5_1(X3)|c7_1(X1)|c7_1(X3)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_143, c_0_349])])).
cnf(c_0_697,negated_conjecture,(c6_2(X1,X2)|ssSkC1|~c4_2(X1,X2)|~c1_2(X1,a462)|~c2_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_144, c_0_349])])).
cnf(c_0_698,negated_conjecture,(ssSkC12|c3_2(X1,X2)|c2_2(X1,X2)|~c1_2(X1,a527)|~c1_2(X1,X2)|~ssSkP6(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_145, c_0_349])])).
cnf(c_0_699,negated_conjecture,(c5_2(X1,a460)|~ssSkC0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_514, c_0_441])])).
cnf(c_0_700,negated_conjecture,(c4_2(a534,a535)|c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_515, c_0_360])]), ['final']).
cnf(c_0_701,negated_conjecture,(c8_2(a534,a535)|c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_516, c_0_360])]), ['final']).
cnf(c_0_702,negated_conjecture,(c7_2(X1,X2)|c4_1(X1)|ndr1_1(X3)|c5_1(X3)|c7_1(X1)|c7_1(X3)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_146, c_0_349])])).
cnf(c_0_703,negated_conjecture,(ssSkC8|c6_2(a568,a583)|c7_2(X1,X2)|c4_1(X1)|c5_1(a568)|c7_1(a568)|c7_1(X1)|ssSkC21|~c6_2(X1,X2)|~c2_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_510, c_0_513]), c_0_353]), ['final']).
cnf(c_0_704,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|c8_1(a534)|c6_1(X1)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_517, c_0_360])]), ['final']).
cnf(c_0_705,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|ndr1_1(a534)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_518, c_0_360])]), ['final']).
cnf(c_0_706,negated_conjecture,(ssSkC22|c4_2(a532,a576)|c7_1(a532)|ssSkC13), inference(spm,[status(thm)],[c_0_519, c_0_147]), ['final']).
cnf(c_0_707,negated_conjecture,(ssSkC22|ndr1_1(a532)|c7_1(a532)|ssSkC13), inference(spm,[status(thm)],[c_0_520, c_0_147]), ['final']).
cnf(c_0_708,negated_conjecture,(c4_2(a565,a487)|ssSkC5|c5_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_371, c_0_16]), ['final']).
cnf(c_0_709,negated_conjecture,(c4_2(a517,a487)|ssSkC5|c5_1(a517)|ssSkC11), inference(spm,[status(thm)],[c_0_371, c_0_22]), ['final']).
cnf(c_0_710,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,X1)|c4_1(a552)|c3_1(a466)|ssSkC18|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_521, c_0_117]), c_0_118]), ['final']).
cnf(c_0_711,negated_conjecture,(ssSkC8|c1_2(a587,a557)|c2_2(a587,a557)|c5_1(a565)|c7_1(a587)|ssSkC20|ssSkC23|~c5_2(a587,a557)), inference(ef,[status(thm)],[c_0_522]), ['final']).
cnf(c_0_712,negated_conjecture,(c8_2(X1,X2)|c8_1(X1)|ndr1_1(X3)|~c4_2(X1,X2)|~c7_2(X1,X2)|~c8_1(X3)|~c2_1(X1)|~c6_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_148, c_0_349])])).
cnf(c_0_713,negated_conjecture,(ssSkC8|c1_2(a587,a557)|c2_2(a587,a588)|ssSkP8(a587)|c5_1(a565)|c7_1(a587)|ssSkC20|ssSkC23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_522, c_0_56]), c_0_46]), ['final']).
cnf(c_0_714,negated_conjecture,(ssSkC8|c4_2(a592,a526)|c5_2(a592,a526)|c2_2(a592,a588)|c8_1(a592)|ssSkP6(a592)|c1_1(a573)|ssSkP8(a592)|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_432]), c_0_46])).
cnf(c_0_715,negated_conjecture,(c8_2(a570,X1)|c1_2(a570,a557)|c2_2(a570,X1)|c5_1(a565)|c7_1(a570)|ssSkC20), inference(spm,[status(thm)],[c_0_352, c_0_421]), ['final']).
cnf(c_0_716,negated_conjecture,(c4_2(a498,a526)|c5_2(a498,a526)|c1_2(a498,a499)|c2_2(a498,a499)|c8_1(a498)|ssSkP6(a498)|c1_1(a573)|c2_0), inference(spm,[status(thm)],[c_0_450, c_0_523]), ['final']).
cnf(c_0_717,negated_conjecture,(c3_2(a509,a511)|~c4_0), inference(csr,[status(thm)],[c_0_149, c_0_150])).
cnf(c_0_718,negated_conjecture,(c4_2(a512,a526)|c5_2(a512,a526)|c8_1(a512)|ssSkP6(a512)|c1_1(a573)|ssSkC9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_134]), c_0_135]), c_0_136]), ['final']).
cnf(c_0_719,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c3_2(a552,X1)|c2_2(a552,X1)|c4_1(a531)|c5_1(a552)|ssSkC18|~c7_2(a552,X1)|~c1_2(a552,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_363]), c_0_353]), ['final']).
cnf(c_0_720,negated_conjecture,(c4_2(a498,a526)|c5_2(a498,a526)|c1_2(a498,a499)|c2_2(a498,a499)|c8_1(a498)|ssSkP6(a498)|c2_0|c3_1(a573)), inference(spm,[status(thm)],[c_0_449, c_0_523]), ['final']).
cnf(c_0_721,negated_conjecture,(ssSkC8|c4_2(a592,a526)|c5_2(a592,a526)|c2_2(a592,a588)|c8_1(a592)|ssSkP6(a592)|ssSkP8(a592)|c3_1(a573)|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_432]), c_0_46])).
cnf(c_0_722,negated_conjecture,(c4_2(a512,a526)|c5_2(a512,a526)|c8_1(a512)|ssSkP6(a512)|c3_1(a573)|ssSkC9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_134]), c_0_135]), c_0_136]), ['final']).
cnf(c_0_723,negated_conjecture,(c4_2(a498,a499)|c5_2(a498,a499)|c1_2(a498,X1)|c2_2(a498,X1)|c8_1(a498)|c1_1(a573)|c2_0|~c3_2(a498,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_524]), c_0_525]), ['final']).
cnf(c_0_724,negated_conjecture,(c4_2(a498,a499)|c5_2(a498,a499)|c1_2(a498,X1)|c2_2(a498,X1)|c8_1(a498)|c2_0|c3_1(a573)|~c3_2(a498,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_524]), c_0_525]), ['final']).
cnf(c_0_725,negated_conjecture,(ndr1_1(a549)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_151, c_0_360])]), ['final']).
cnf(c_0_726,negated_conjecture,(ssSkC12|c4_2(a552,a527)|c2_2(a552,a503)|ssSkP4(a552)|c2_1(a552)|c4_1(a458)|ssSkC18), inference(spm,[status(thm)],[c_0_526, c_0_397]), ['final']).
cnf(c_0_727,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c2_2(a512,a503)|ssSkP4(a512)|c2_1(a512)|c4_1(a458)|ssSkC9), inference(spm,[status(thm)],[c_0_526, c_0_399]), ['final']).
cnf(c_0_728,negated_conjecture,(c8_2(a549,a551)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_360])]), ['final']).
cnf(c_0_729,negated_conjecture,(c8_2(X1,a460)|~ssSkC0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_527, c_0_441])])).
cnf(c_0_730,negated_conjecture,(c4_2(X1,X2)|c8_2(X1,a481)|ssSkC4|c8_1(X1)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_153, c_0_349])])).
cnf(c_0_731,negated_conjecture,(c4_2(X1,X2)|ssSkC4|c2_2(X1,a481)|c8_1(X1)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_154, c_0_349])])).
cnf(c_0_732,negated_conjecture,(ssSkC8|c1_2(a587,a526)|ssSkP6(a587)|ssSkC23|~c5_2(a587,a526)), inference(spm,[status(thm)],[c_0_439, c_0_98]), ['final']).
cnf(c_0_733,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c4_2(a574,X3)|~c3_2(a574,X3)|~c2_2(a574,X3)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_155, c_0_349])]), c_0_528])).
cnf(c_0_734,negated_conjecture,(c4_2(a509,a510)|~c4_0), inference(csr,[status(thm)],[c_0_156, c_0_150])).
cnf(c_0_735,negated_conjecture,(ndr1_1(a509)|~c4_0), inference(csr,[status(thm)],[c_0_157, c_0_150])).
cnf(c_0_736,negated_conjecture,(~c8_2(a509,a510)|~c4_0), inference(csr,[status(thm)],[c_0_158, c_0_150])).
cnf(c_0_737,negated_conjecture,(c3_2(X1,a456)|c1_2(X1,a493)|ssSkP0(X1)|ssSkC7|~c5_2(X1,a456)|~ssSkP2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_99]), c_0_100]), ['final']).
cnf(c_0_738,negated_conjecture,(c5_2(a566,X1)|~c1_2(a566,X1)|~c2_2(a566,X1)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[c_0_159, c_0_160])).
cnf(c_0_739,negated_conjecture,(ssSkC8|c8_2(a564,a503)|ssSkP4(a564)|ssSkC19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_529, c_0_62]), c_0_63]), ['final']).
cnf(c_0_740,negated_conjecture,(c1_2(X1,a526)|c8_1(X1)|ssSkP6(X1)|c4_1(X1)|c3_1(a466)|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_530, c_0_117]), c_0_118]), ['final']).
cnf(c_0_741,negated_conjecture,(c3_2(X1,a456)|c1_2(X1,a456)|c8_1(X1)|ssSkP0(X1)|c4_1(X1)|c3_1(a466)|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_531, c_0_117]), c_0_118]), ['final']).
cnf(c_0_742,negated_conjecture,(c3_2(a542,a543)|c1_2(a542,a543)|c8_1(a542)|c4_1(a542)|c3_1(a466)|ssSkC16|~c2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_117]), c_0_118]), ['final']).
cnf(c_0_743,negated_conjecture,(c3_2(a498,a500)|c1_2(a498,a500)|c8_1(a498)|c8_1(X1)|c4_1(a498)|c2_0|c3_1(X1)|~c1_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_533]), c_0_525]), ['final']).
cnf(c_0_744,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|~c6_2(X3,a525)|~c6_2(X1,X2)|~c4_2(X1,X2)|~c4_1(X1)|~ndr1_1(X1)|~c2_0|~c3_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_161, c_0_349])])).
cnf(c_0_745,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~c2_2(X3,a525)|~c4_1(X1)|~ndr1_1(X1)|~c2_0|~c3_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_162, c_0_349])])).
cnf(c_0_746,negated_conjecture,(ssSkC8|c5_2(a536,a526)|ssSkP6(a536)|ssSkC14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_20]), c_0_21]), ['final']).
cnf(c_0_747,negated_conjecture,(c7_2(X1,a525)|c3_2(X2,X3)|c8_1(X2)|c6_1(X1)|~c6_2(X2,X3)|~c4_2(X2,X3)|~c4_1(X2)|~ndr1_1(X2)|~c2_0|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_163, c_0_349])])).
cnf(c_0_748,negated_conjecture,(ssSkC5|ndr1_1(a565)|c5_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_534, c_0_16]), ['final']).
cnf(c_0_749,negated_conjecture,(ssSkC5|ndr1_1(a517)|c5_1(a517)|ssSkC11), inference(spm,[status(thm)],[c_0_534, c_0_22]), ['final']).
cnf(c_0_750,negated_conjecture,(c2_2(a482,X1)|~c6_2(a482,X1)|~c8_2(a482,X1)|~ssSkC4), inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_164, c_0_350]), c_0_535])).
cnf(c_0_751,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|ndr1_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~c4_1(X1)|~ndr1_1(X1)|~c2_0|~c3_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_165, c_0_349])])).
cnf(c_0_752,negated_conjecture,(c5_2(X1,X2)|c2_2(X1,X2)|~c5_2(X1,a572)|~c4_1(X1)|~ndr1_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_166, c_0_349])]), c_0_350]), c_0_366])).
cnf(c_0_753,negated_conjecture,(c8_1(X1)|c1_1(X1)|~c6_2(a530,X2)|~c8_2(a530,X2)|~c3_2(a530,X2)|~c6_1(X1)|~ndr1_1(a530)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_167, c_0_349])])).
cnf(c_0_754,negated_conjecture,(ssSkC3|c6_2(a581,a479)|c1_2(a581,a526)|ssSkP6(a581)|~c2_0), inference(spm,[status(thm)],[c_0_536, c_0_138]), ['final']).
cnf(c_0_755,negated_conjecture,(ssSkP5(X1)|c2_1(X1)|c4_1(a458)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_23]), c_0_48]), c_0_24]), ['final']).
cnf(c_0_756,negated_conjecture,(c7_2(a530,X1)|c8_1(X2)|c1_1(X2)|~c6_2(a530,X1)|~c4_2(a530,X1)|~c6_1(X2)|~ndr1_1(a530)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_168, c_0_349])])).
cnf(c_0_757,negated_conjecture,(ssSkC8|~c5_2(X1,a504)|~ssSkP3(X1)|~ssSkP4(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_169, c_0_349])])).
cnf(c_0_758,negated_conjecture,(ssSkC8|c7_2(a470,X1)|ssSkC2|~c4_2(a470,X1)|~c1_2(a470,X1)), inference(spm,[status(thm)],[c_0_170, c_0_353]), ['final']).
cnf(c_0_759,negated_conjecture,(ssSkC8|c6_2(a514,X1)|ssSkC10|~c4_2(a514,X1)|~c8_2(a514,X1)), inference(spm,[status(thm)],[c_0_171, c_0_353]), ['final']).
cnf(c_0_760,negated_conjecture,(ssSkC3|c6_2(a581,a479)|c3_2(a581,a456)|c1_2(a581,a456)|ssSkP0(a581)|~c2_0), inference(spm,[status(thm)],[c_0_537, c_0_138]), ['final']).
cnf(c_0_761,negated_conjecture,(c5_2(a549,a551)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_172, c_0_360])]), ['final']).
cnf(c_0_762,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c6_2(a574,a575)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_173, c_0_349])])).
cnf(c_0_763,negated_conjecture,(ssSkC8|c7_2(a532,X1)|ssSkC13|~c4_2(a532,X1)|~c8_2(a532,X1)), inference(spm,[status(thm)],[c_0_174, c_0_353]), ['final']).
cnf(c_0_764,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c1_2(a574,a575)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_175, c_0_349])])).
cnf(c_0_765,negated_conjecture,(c4_2(a530,X1)|c8_2(a530,X1)|c7_2(a530,X1)|c8_1(X2)|c1_1(X2)|~c6_1(X2)|~ndr1_1(a530)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_176, c_0_349])])).
cnf(c_0_766,negated_conjecture,(c6_2(X1,X2)|c4_2(a574,a575)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_177, c_0_349])])).
cnf(c_0_767,negated_conjecture,(c4_2(a549,a550)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_178, c_0_360])]), ['final']).
cnf(c_0_768,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c6_1(X1)|~ndr1_1(X1)|~c2_0|~c3_1(a574)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_179, c_0_349])])).
cnf(c_0_769,negated_conjecture,(c4_1(a488)|~ssSkC5), inference(sr,[status(thm)],[c_0_180, c_0_366])).
cnf(c_0_770,negated_conjecture,(ssSkC8|ssSkC25|~c4_2(a596,X1)|~c1_2(a596,X1)), inference(spm,[status(thm)],[c_0_181, c_0_353]), ['final']).
cnf(c_0_771,negated_conjecture,(c3_2(a496,X1)|~c8_2(a496,X1)|~c2_2(a496,X1)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_538, c_0_441])])).
cnf(c_0_772,negated_conjecture,(c8_2(a546,X1)|ssSkC17|~c4_2(a546,X1)|~c1_2(a546,X1)), inference(csr,[status(thm)],[c_0_182, c_0_45])).
cnf(c_0_773,negated_conjecture,(ssSkC0|c7_2(a581,a459)|~c2_0|~c3_1(a581)), inference(spm,[status(thm)],[c_0_539, c_0_138]), ['final']).
cnf(c_0_774,negated_conjecture,(ndr1_1(X1)|~c8_1(X1)|~c3_1(X1)), inference(csr,[status(thm)],[c_0_540, c_0_541]), ['final']).
cnf(c_0_775,negated_conjecture,(ssSkC22|c4_2(a466,a576)|c7_1(a466)|~c2_0), inference(spm,[status(thm)],[c_0_519, c_0_117]), ['final']).
cnf(c_0_776,negated_conjecture,(ssSkC22|ndr1_1(a466)|c7_1(a466)|~c2_0), inference(spm,[status(thm)],[c_0_520, c_0_117]), ['final']).
cnf(c_0_777,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X4,X5)|c2_2(X1,X2)|c6_1(X1)|~c4_2(X1,X3)|~c8_2(X1,X2)|~c5_2(X4,a589)|~c2_2(X1,X3)|~c2_2(X4,X5)|~ssSkP8(X4)|~ndr1_1(X1)|~ndr1_1(X4)|~ssSkC23), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_183, c_0_349])])).
cnf(c_0_778,negated_conjecture,(c7_2(X1,X2)|c5_2(X3,X4)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X5,X6)|c8_1(X5)|c8_1(X3)|c1_1(X1)|c2_1(X1)|~c6_2(X3,X4)|~c6_2(X5,X6)|~c7_2(X3,a465)|~c7_2(X3,X4)|~c2_2(X5,X6)|~c1_1(X5)|~ndr1_1(X3)|~ndr1_1(X5)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_184, c_0_349])])).
cnf(c_0_779,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X4,X5)|c2_2(X1,X2)|c6_1(X1)|~c4_2(X1,X3)|~c8_2(X1,X2)|~c7_2(X4,a589)|~c2_2(X1,X3)|~c2_2(X4,X5)|~ssSkP8(X4)|~ndr1_1(X1)|~ndr1_1(X4)|~ssSkC23), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_185, c_0_349])])).
cnf(c_0_780,negated_conjecture,(c6_2(X1,X2)|c8_2(X3,a589)|c8_2(X1,X4)|c5_2(X3,X5)|c2_2(X1,X2)|c6_1(X1)|~c4_2(X1,X4)|~c8_2(X1,X2)|~c2_2(X1,X4)|~c2_2(X3,X5)|~ssSkP8(X3)|~ndr1_1(X1)|~ndr1_1(X3)|~ssSkC23), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_186, c_0_349])])).
cnf(c_0_781,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X4)|c3_2(X1,X3)|c2_2(X1,X2)|~ssSkC22|~c6_2(a577,a579)|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_187, c_0_349])])).
cnf(c_0_782,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X4)|c3_2(X1,X3)|c2_2(X1,X2)|~ssSkC22|~c8_2(a577,a579)|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_188, c_0_349])])).
cnf(c_0_783,negated_conjecture,(c6_2(a577,a578)|c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X4)|c3_2(X1,X3)|c2_2(X1,X2)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_189, c_0_349])])).
cnf(c_0_784,negated_conjecture,(c6_2(X1,X2)|c8_2(a577,a578)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X4)|c3_2(X1,X3)|c2_2(X1,X2)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_190, c_0_349])])).
cnf(c_0_785,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X4)|c3_2(X1,X3)|c1_2(a577,a578)|c2_2(X1,X2)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_191, c_0_349])])).
cnf(c_0_786,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X4)|c3_2(X1,X3)|c2_2(X1,X2)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)|~c7_1(a577)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_192, c_0_349])])).
cnf(c_0_787,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X4)|c3_2(X1,X3)|c2_2(X1,X2)|ndr1_1(a577)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_193, c_0_349])])).
cnf(c_0_788,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|~c6_2(a534,a535)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_194, c_0_349])])).
cnf(c_0_789,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~c2_1(a534)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_195, c_0_349])])).
cnf(c_0_790,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~c6_1(a559)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_196, c_0_349])])).
cnf(c_0_791,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c5_1(X1)|~c8_2(X1,X3)|~c7_2(X1,X2)|~c1_2(X1,X3)|~c6_1(a531)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_197, c_0_349])])).
cnf(c_0_792,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X2)|c3_2(X1,X3)|c5_1(X1)|~c8_2(X1,X3)|~c7_2(X1,X2)|~c1_2(X1,X3)|~ndr1_1(X1)|~c5_1(a531)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_198, c_0_349])])).
cnf(c_0_793,negated_conjecture,(c4_2(X1,a473)|~c4_2(X2,X3)|~c8_2(X2,X3)|~c1_2(X2,X3)|~c7_0|~ndr1_1(X2)|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_199, c_0_349])])).
cnf(c_0_794,negated_conjecture,(c8_2(X1,a473)|~c4_2(X2,X3)|~c8_2(X2,X3)|~c1_2(X2,X3)|~c7_0|~ndr1_1(X2)|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_200, c_0_349])])).
cnf(c_0_795,negated_conjecture,(c8_2(X1,X2)|c2_2(X1,X2)|c5_1(X3)|c7_1(X1)|~c6_2(X1,a557)|~c6_1(X3)|~ndr1_1(X1)|~c3_1(X3)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_201, c_0_349])]), c_0_350])).
cnf(c_0_796,negated_conjecture,(c8_2(X1,X2)|c2_2(X1,X2)|c5_1(X3)|c7_1(X1)|~c2_2(X1,a557)|~c6_1(X3)|~ndr1_1(X1)|~c3_1(X3)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_202, c_0_349])]), c_0_350])).
cnf(c_0_797,negated_conjecture,(c6_2(X1,X2)|ssSkC1|~c6_2(X1,a462)|~c4_2(X1,X2)|~c2_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_203, c_0_349])])).
cnf(c_0_798,negated_conjecture,(c6_2(X1,X2)|ssSkC1|~c4_2(X1,X2)|~c7_2(X1,a462)|~c2_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_204, c_0_349])])).
cnf(c_0_799,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X2)|c7_0|~c4_2(X1,a593)|~c7_2(X1,X2)|~c6_1(X1)|~ndr1_1(X1)|~ssSkC24), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_205, c_0_349])])).
cnf(c_0_800,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X2)|c7_0|~c7_2(X1,a593)|~c7_2(X1,X2)|~c6_1(X1)|~ndr1_1(X1)|~ssSkC24), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_206, c_0_349])])).
cnf(c_0_801,negated_conjecture,(ndr1_1(X1)|~c4_2(X2,X3)|~c8_2(X2,X3)|~c1_2(X2,X3)|~c7_0|~ndr1_1(X2)|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_207, c_0_349])])).
cnf(c_0_802,negated_conjecture,(c8_2(X1,X2)|c3_2(X1,X2)|c2_1(X1)|~ssSkC12|~c7_2(X1,X2)|~c5_2(X1,a528)|~c4_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_208, c_0_349])])).
cnf(c_0_803,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,a593)|c2_2(X1,X2)|c7_0|~c7_2(X1,X2)|~c6_1(X1)|~ndr1_1(X1)|~ssSkC24), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_209, c_0_349])])).
cnf(c_0_804,negated_conjecture,(c3_2(X1,X2)|ssSkC7|~c6_2(X1,a493)|~c8_2(X1,X2)|~c5_2(X1,X2)|~ssSkP2(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_210, c_0_349])])).
cnf(c_0_805,negated_conjecture,(c3_2(X1,X2)|ssSkC7|~c8_2(X1,X2)|~c5_2(X1,X2)|~c2_2(X1,a493)|~ssSkP2(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_211, c_0_349])])).
cnf(c_0_806,negated_conjecture,(c7_2(X1,X2)|c2_1(X1)|c4_1(X1)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)|~c3_1(a533)|~ssSkC13), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_212, c_0_349])])).
cnf(c_0_807,negated_conjecture,(c3_2(X1,X2)|c2_2(X1,a548)|~c6_2(X1,X2)|~c8_2(X1,X2)|~c1_1(X1)|~ndr1_1(X1)|~ssSkC17), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_213, c_0_349])])).
cnf(c_0_808,negated_conjecture,(ssSkC12|c3_2(X1,X2)|c2_2(X1,X2)|~c1_2(X1,X2)|~c2_2(X1,a527)|~ssSkP6(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_214, c_0_349])])).
cnf(c_0_809,negated_conjecture,(c6_2(X1,a528)|c8_2(X1,X2)|c3_2(X1,X2)|c2_1(X1)|~ssSkC12|~c7_2(X1,X2)|~c4_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_215, c_0_349])])).
cnf(c_0_810,negated_conjecture,(c8_2(X1,X2)|c3_2(X1,X2)|c2_2(X1,a528)|c2_1(X1)|~ssSkC12|~c7_2(X1,X2)|~c4_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_216, c_0_349])])).
cnf(c_0_811,negated_conjecture,(c6_2(X1,X2)|c7_2(X1,X2)|c1_2(X1,X2)|c1_1(X1)|c7_0|~c8_1(a501)|~c2_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_217, c_0_349])])).
cnf(c_0_812,negated_conjecture,(c6_2(X1,X2)|c7_2(X1,X2)|c1_2(X1,X2)|c1_1(X1)|c7_0|~c2_1(X1)|~c6_1(a501)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_218, c_0_349])])).
cnf(c_0_813,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|~c7_2(X1,X2)|~c2_1(a458)|~c6_1(X1)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_219, c_0_349])])).
cnf(c_0_814,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|~c7_2(X1,X2)|~c6_1(a458)|~c6_1(X1)|~ndr1_1(X1)|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_220, c_0_349])])).
cnf(c_0_815,negated_conjecture,(ssSkC3|c3_2(X1,X2)|c1_2(X1,X2)|~c8_2(X1,X2)|~c7_2(X1,a479)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_221, c_0_349])])).
cnf(c_0_816,negated_conjecture,(c4_2(a544,X1)|c1_2(a544,X1)|~c8_1(X2)|~c2_1(X2)|~ssSkC16), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_222, c_0_349])]), c_0_542])).
cnf(c_0_817,negated_conjecture,(c5_2(a488,X1)|c2_2(a488,X1)|~c6_2(a488,X1)|~ssSkC5|~ndr1_1(a488)), inference(sr,[status(thm)],[c_0_223, c_0_366])).
cnf(c_0_818,negated_conjecture,(epred1_0|~c4_2(X1,X2)|~c8_2(X1,X2)|~c1_2(X1,X2)|~c7_0|~ndr1_1(X1)), inference(split_equiv,[status(thm)],[c_0_543])).
cnf(c_0_819,negated_conjecture,(c4_2(a475,X1)|c5_2(a475,X1)|c7_0|~c8_2(a475,X1)|~c2_0), inference(csr,[status(thm)],[c_0_224, c_0_225])).
cnf(c_0_820,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|~c5_2(a539,a540)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_226, c_0_349])])).
cnf(c_0_821,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|~c5_2(a539,a541)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_227, c_0_349])])).
cnf(c_0_822,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|~c2_2(a539,a541)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_228, c_0_349])])).
cnf(c_0_823,negated_conjecture,(c7_0|~c3_2(X1,a553)|~c4_1(X1)|~c5_1(X1)|~ssSkC18), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_229, c_0_349])])).
cnf(c_0_824,negated_conjecture,(c7_0|~c2_2(X1,a553)|~c4_1(X1)|~c5_1(X1)|~ssSkC18), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_230, c_0_349])])).
cnf(c_0_825,negated_conjecture,(c5_2(X1,X2)|c2_2(X1,X2)|~c6_2(X1,a572)|~c4_1(X1)|~ndr1_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_231, c_0_349])]), c_0_350]), c_0_366])).
cnf(c_0_826,negated_conjecture,(ssSkC0|~c6_2(X1,a459)|~c8_1(X1)|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_232, c_0_349])])).
cnf(c_0_827,negated_conjecture,(ssSkC0|~c3_2(X1,a459)|~c8_1(X1)|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_233, c_0_349])])).
cnf(c_0_828,negated_conjecture,(ssSkC8|~c3_2(X1,a504)|~ssSkP3(X1)|~ssSkP4(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_234, c_0_349])])).
cnf(c_0_829,negated_conjecture,(c1_1(X1)|c3_1(X1)|~c5_2(X1,a582)|~c4_0|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_235, c_0_349])])).
cnf(c_0_830,negated_conjecture,(c8_2(a539,a540)|c1_1(X1)|c6_1(X1)|c5_1(X1)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_236, c_0_349])])).
cnf(c_0_831,negated_conjecture,(c8_2(a539,a541)|c1_1(X1)|c6_1(X1)|c5_1(X1)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_237, c_0_349])])).
cnf(c_0_832,negated_conjecture,(c3_1(X1)|~ssSkC3|~c6_2(X1,a480)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_544, c_0_441])])).
cnf(c_0_833,negated_conjecture,(c3_2(a539,a540)|c1_1(X1)|c6_1(X1)|c5_1(X1)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_238, c_0_349])])).
cnf(c_0_834,negated_conjecture,(c2_1(X1)|c7_1(X1)|~c7_2(X1,a602)|~c7_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_239, c_0_349])]), c_0_366])).
cnf(c_0_835,negated_conjecture,(c7_2(X1,a553)|c7_0|~c4_1(X1)|~c5_1(X1)|~ssSkC18), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_240, c_0_349])])).
cnf(c_0_836,negated_conjecture,(c3_2(X1,a582)|c1_1(X1)|c3_1(X1)|~c4_0|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_241, c_0_349])])).
cnf(c_0_837,negated_conjecture,(c2_2(X1,a582)|c1_1(X1)|c3_1(X1)|~c4_0|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_242, c_0_349])])).
cnf(c_0_838,negated_conjecture,(ssSkC8|c8_2(X1,a504)|~ssSkP3(X1)|~ssSkP4(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_243, c_0_349])])).
cnf(c_0_839,negated_conjecture,(c3_2(X1,a480)|c3_1(X1)|~ssSkC3|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_545, c_0_441])])).
cnf(c_0_840,negated_conjecture,(c4_2(X1,a480)|c3_1(X1)|~ssSkC3|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_546, c_0_441])])).
cnf(c_0_841,negated_conjecture,(c4_2(X1,a602)|c2_1(X1)|c7_1(X1)|~c7_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_244, c_0_349])]), c_0_366])).
cnf(c_0_842,negated_conjecture,(c3_2(X1,a602)|c2_1(X1)|c7_1(X1)|~c7_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_245, c_0_349])]), c_0_366])).
cnf(c_0_843,negated_conjecture,(epred2_0|~c3_2(X1,a473)|~c3_1(X1)), inference(split_equiv,[status(thm)],[c_0_547])).
cnf(c_0_844,negated_conjecture,(c7_0|ndr1_1(X1)|~c4_1(X1)|~c5_1(X1)|~ssSkC18), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_246, c_0_349])])).
cnf(c_0_845,negated_conjecture,(c1_1(X1)|c6_1(a539)|c6_1(X1)|c5_1(X1)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_247, c_0_349])])).
cnf(c_0_846,negated_conjecture,(c1_1(X1)|c6_1(X1)|ndr1_1(a539)|c5_1(X1)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_248, c_0_349])])).
cnf(c_0_847,negated_conjecture,(c1_1(X1)|ndr1_1(X1)|c3_1(X1)|~c4_0|~c8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_249, c_0_349])])).
cnf(c_0_848,negated_conjecture,(ndr1_1(X1)|c3_1(X1)|~ssSkC3|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_548, c_0_441])])).
cnf(c_0_849,negated_conjecture,(c2_0|~c8_2(a498,a499)), inference(sr,[status(thm)],[c_0_250, c_0_366])).
cnf(c_0_850,negated_conjecture,(c2_0|~c7_2(a498,a500)), inference(sr,[status(thm)],[c_0_251, c_0_366])).
cnf(c_0_851,negated_conjecture,(c2_1(X1)|ndr1_1(X1)|c7_1(X1)|~c7_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_252, c_0_349])]), c_0_366])).
cnf(c_0_852,negated_conjecture,(c8_2(a482,a483)|~ssSkC4), inference(sr,[status(thm)],[c_0_253, c_0_350])).
cnf(c_0_853,negated_conjecture,(c3_2(a482,a483)|~ssSkC4), inference(sr,[status(thm)],[c_0_254, c_0_350])).
cnf(c_0_854,negated_conjecture,(c2_2(a482,a483)|~ssSkC4), inference(sr,[status(thm)],[c_0_255, c_0_350])).
cnf(c_0_855,negated_conjecture,(c4_2(a498,a500)|c2_0), inference(sr,[status(thm)],[c_0_256, c_0_366])).
cnf(c_0_856,negated_conjecture,(c2_0|c7_1(a498)), inference(sr,[status(thm)],[c_0_257, c_0_366])).
cnf(c_0_857,negated_conjecture,(~c4_2(X1,X2)|~c7_2(X1,X2)|~c3_2(X1,X2)|~c2_1(a508)|~c4_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_549, c_0_441])])).
cnf(c_0_858,negated_conjecture,(~c4_2(a544,a545)|~c8_1(X1)|~c2_1(X1)|~ssSkC16), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_258, c_0_349])])).
cnf(c_0_859,negated_conjecture,(~c7_2(a544,a545)|~c8_1(X1)|~c2_1(X1)|~ssSkC16), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_259, c_0_349])])).
cnf(c_0_860,negated_conjecture,(~c3_2(a544,a545)|~c8_1(X1)|~c2_1(X1)|~ssSkC16), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_260, c_0_349])])).
cnf(c_0_861,negated_conjecture,(~c6_2(a509,a511)|~c4_0), inference(csr,[status(thm)],[c_0_261, c_0_150])).
cnf(c_0_862,negated_conjecture,(~c5_2(a509,a511)|~c4_0), inference(csr,[status(thm)],[c_0_262, c_0_150])).
cnf(c_0_863,negated_conjecture,(~c1_2(a509,a510)|~c4_0), inference(csr,[status(thm)],[c_0_263, c_0_150])).
cnf(c_0_864,negated_conjecture,(~c6_2(a482,a484)|~ssSkC4), inference(sr,[status(thm)],[c_0_264, c_0_350])).
cnf(c_0_865,negated_conjecture,(~c4_2(a482,a484)|~ssSkC4), inference(sr,[status(thm)],[c_0_265, c_0_350])).
cnf(c_0_866,negated_conjecture,(~c2_1(a496)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_266, c_0_441])])).
cnf(c_0_867,negated_conjecture,(~ssSkC5|~c3_1(a488)), inference(sr,[status(thm)],[c_0_267, c_0_366])).
cnf(c_0_868,negated_conjecture,(~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_268, c_0_349])]), c_0_543]), c_0_547])).
cnf(c_0_869,negated_conjecture,(~c4_2(a520,a521)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_269, c_0_350]), c_0_366])).
cnf(c_0_870,negated_conjecture,(~c1_2(a520,a521)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_270, c_0_350]), c_0_366])).
cnf(c_0_871,negated_conjecture,(~c2_1(a520)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_271, c_0_350]), c_0_366])).
cnf(c_0_872,negated_conjecture,(c2_2(a520,a521)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_272, c_0_350]), c_0_366])).
cnf(c_0_873,negated_conjecture,(c3_2(a570,a571)), inference(sr,[status(thm)],[c_0_273, c_0_350])).
cnf(c_0_874,negated_conjecture,(c1_2(a570,a571)), inference(sr,[status(thm)],[c_0_274, c_0_350])).
cnf(c_0_875,negated_conjecture,(c2_2(a570,a571)), inference(sr,[status(thm)],[c_0_275, c_0_350])).
cnf(c_0_876,negated_conjecture,(c5_1(a570)), inference(sr,[status(thm)],[c_0_276, c_0_350])).
cnf(c_0_877,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,a515)|c1_2(a565,a493)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|ssSkP6(a565)|c1_1(a573)|c1_1(a565)|ssSkC7|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~ssSkP2(a565)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_551]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_878,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,a515)|c1_2(a565,a493)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|ssSkP6(a565)|c1_1(a565)|ssSkC7|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~ssSkP2(a565)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_552]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_879,negated_conjecture,(ssSkC8|c6_2(a517,a588)|c6_2(a517,X1)|c4_2(a517,X1)|c5_2(a517,X1)|c3_2(a517,a588)|c3_2(a517,X1)|c1_2(a517,a557)|c1_2(a517,X2)|c2_2(a517,X1)|c2_2(a517,X2)|c8_1(a517)|c1_1(a517)|c1_1(a573)|ssSkP8(a517)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20|~c3_2(a517,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_553]), c_0_353]), ['final']).
cnf(c_0_880,negated_conjecture,(ssSkC8|c6_2(a517,a588)|c6_2(a517,X1)|c4_2(a517,X1)|c5_2(a517,X1)|c3_2(a517,a588)|c3_2(a517,X1)|c1_2(a517,a557)|c1_2(a517,X2)|c2_2(a517,X1)|c2_2(a517,X2)|c8_1(a517)|c1_1(a517)|ssSkP8(a517)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20|~c3_2(a517,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_553]), c_0_353]), ['final']).
cnf(c_0_881,negated_conjecture,(c6_2(a542,a588)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,a588)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|c1_1(a542)|c1_1(a573)|ssSkP8(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_554]), c_0_278]), ['final']).
cnf(c_0_882,negated_conjecture,(c6_2(a542,a588)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,a588)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|c1_1(a542)|ssSkP8(a542)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_554]), c_0_278]), ['final']).
cnf(c_0_883,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,a515)|c1_2(a565,a493)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|c1_1(a573)|c1_1(a565)|ssSkC7|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~ssSkP2(a565)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_555]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_884,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,a515)|c1_2(a565,a493)|c1_2(a565,a557)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|c1_1(a565)|ssSkC7|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~ssSkP2(a565)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_556]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_885,negated_conjecture,(c6_2(a542,a487)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,a487)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|c1_1(a542)|c1_1(a573)|ssSkC5|c7_0|c4_1(a531)|c5_1(a542)|c5_1(a565)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_557]), c_0_27]), ['final']).
cnf(c_0_886,negated_conjecture,(c6_2(a542,a487)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,a487)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|c1_1(a542)|ssSkC5|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a542)|c5_1(a565)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_557]), c_0_27]), ['final']).
cnf(c_0_887,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c6_2(a517,X1)|c4_2(a517,X1)|c5_2(a517,X1)|c3_2(a517,X1)|c1_2(a517,a557)|c1_2(a517,X2)|c2_2(a517,X1)|c2_2(a517,X2)|c8_1(a517)|ssSkP1(a517)|c1_1(a517)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20|~c3_2(a517,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_558]), c_0_353]), ['final']).
cnf(c_0_888,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c6_2(a517,X1)|c4_2(a517,X1)|c5_2(a517,X1)|c3_2(a517,X1)|c1_2(a517,a557)|c1_2(a517,X2)|c2_2(a517,X1)|c2_2(a517,X2)|c8_1(a517)|ssSkP1(a517)|c1_1(a517)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20|~c3_2(a517,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_558]), c_0_353]), ['final']).
cnf(c_0_889,negated_conjecture,(ssSkC8|c6_2(a565,a588)|c6_2(a565,X1)|c4_2(a565,X1)|c5_2(a565,X1)|c3_2(a565,a588)|c3_2(a565,X1)|c1_2(a565,a557)|c1_2(a565,X2)|c2_2(a565,X1)|c2_2(a565,X2)|c8_1(a565)|c1_1(a565)|c1_1(a573)|ssSkP8(a565)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c3_2(a565,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_559]), c_0_353]), ['final']).
cnf(c_0_890,negated_conjecture,(ssSkC8|c6_2(a565,a588)|c6_2(a565,X1)|c4_2(a565,X1)|c5_2(a565,X1)|c3_2(a565,a588)|c3_2(a565,X1)|c1_2(a565,a557)|c1_2(a565,X2)|c2_2(a565,X1)|c2_2(a565,X2)|c8_1(a565)|c1_1(a565)|ssSkP8(a565)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c3_2(a565,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_559]), c_0_353]), ['final']).
cnf(c_0_891,negated_conjecture,(ssSkC8|c6_2(a517,a526)|c6_2(a517,X1)|c4_2(a517,X1)|c5_2(a517,X1)|c3_2(a517,X1)|c1_2(a517,a557)|c1_2(a517,X2)|c2_2(a517,X1)|c2_2(a517,X2)|c8_1(a517)|ssSkP6(a517)|c1_1(a517)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20|~c3_2(a517,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_560]), c_0_353]), ['final']).
cnf(c_0_892,negated_conjecture,(ssSkC8|c6_2(a517,a526)|c6_2(a517,X1)|c4_2(a517,X1)|c5_2(a517,X1)|c3_2(a517,X1)|c1_2(a517,a557)|c1_2(a517,X2)|c2_2(a517,X1)|c2_2(a517,X2)|c8_1(a517)|ssSkP6(a517)|c1_1(a517)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20|~c3_2(a517,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_560]), c_0_353]), ['final']).
cnf(c_0_893,negated_conjecture,(ssSkC8|c6_2(a517,a515)|c6_2(a517,X1)|c4_2(a517,X1)|c5_2(a517,X1)|c3_2(a517,X1)|c1_2(a517,a557)|c1_2(a517,X2)|c2_2(a517,X1)|c2_2(a517,X2)|c8_1(a517)|ssSkP5(a517)|c1_1(a517)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20|~c3_2(a517,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_561]), c_0_353]), ['final']).
cnf(c_0_894,negated_conjecture,(ssSkC8|c6_2(a517,a515)|c6_2(a517,X1)|c4_2(a517,X1)|c5_2(a517,X1)|c3_2(a517,X1)|c1_2(a517,a557)|c1_2(a517,X2)|c2_2(a517,X1)|c2_2(a517,X2)|c8_1(a517)|ssSkP5(a517)|c1_1(a517)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20|~c3_2(a517,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_561]), c_0_353]), ['final']).
cnf(c_0_895,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,a526)|c1_2(a565,a557)|c1_2(a565,a493)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|ssSkP6(a565)|c1_1(a565)|ssSkC7|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkP2(a565)), inference(spm,[status(thm)],[c_0_562, c_0_563]), ['final']).
cnf(c_0_896,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,a526)|c1_2(a565,a557)|c1_2(a565,a493)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP5(a565)|ssSkP6(a565)|c1_1(a573)|c1_1(a565)|ssSkC7|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkP2(a565)), inference(spm,[status(thm)],[c_0_562, c_0_564]), ['final']).
cnf(c_0_897,negated_conjecture,(c6_2(a542,a485)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|ssSkP1(a542)|c1_1(a542)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_565]), c_0_279]), ['final']).
cnf(c_0_898,negated_conjecture,(c6_2(a542,a485)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|ssSkP1(a542)|c1_1(a542)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_565]), c_0_279]), ['final']).
cnf(c_0_899,negated_conjecture,(c6_2(a542,a526)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|ssSkP6(a542)|c1_1(a542)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a542)|c5_1(a565)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_566]), c_0_31]), ['final']).
cnf(c_0_900,negated_conjecture,(c6_2(a542,a526)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|ssSkP6(a542)|c1_1(a542)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a542)|c5_1(a565)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_566]), c_0_31]), ['final']).
cnf(c_0_901,negated_conjecture,(ssSkC12|c6_2(a542,a485)|c4_2(a542,a527)|c3_2(a542,a493)|c3_2(a542,a588)|c3_2(a542,a557)|c2_2(a542,a493)|c2_2(a542,a588)|c2_2(a542,a557)|ssSkP1(a542)|c1_1(a542)|ssSkP8(a542)|ssSkC7|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkP6(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_567]), c_0_279]), ['final']).
cnf(c_0_902,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c4_2(a552,a515)|c1_2(a552,a493)|c1_2(a552,a557)|c2_2(a552,a557)|c2_2(a552,a502)|ssSkP3(a552)|ssSkP5(a552)|ssSkP6(a552)|c1_1(a573)|ssSkC7|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkP2(a552)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_568]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_903,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,a526)|c1_2(a565,a557)|c1_2(a565,a493)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP6(a565)|c1_1(a565)|ssSkC7|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkP2(a565)), inference(spm,[status(thm)],[c_0_562, c_0_569]), ['final']).
cnf(c_0_904,negated_conjecture,(c6_2(a542,a515)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|ssSkP5(a542)|c1_1(a542)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a542)|c5_1(a565)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_570]), c_0_48]), ['final']).
cnf(c_0_905,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c4_2(a552,a515)|c1_2(a552,a493)|c1_2(a552,a557)|c2_2(a552,a557)|c2_2(a552,a502)|ssSkP3(a552)|ssSkP5(a552)|ssSkP6(a552)|ssSkC7|c4_1(a531)|c3_1(a573)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkP2(a552)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_571]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_906,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c6_2(a565,X1)|c4_2(a565,X1)|c5_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c1_2(a565,X2)|c2_2(a565,X1)|c2_2(a565,X2)|c8_1(a565)|ssSkP1(a565)|c1_1(a565)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c3_2(a565,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_572]), c_0_353]), ['final']).
cnf(c_0_907,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c6_2(a565,X1)|c4_2(a565,X1)|c5_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c1_2(a565,X2)|c2_2(a565,X1)|c2_2(a565,X2)|c8_1(a565)|ssSkP1(a565)|c1_1(a565)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a565)|c7_1(a565)|ssSkC20|~c3_2(a565,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_572]), c_0_353]), ['final']).
cnf(c_0_908,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,a526)|c1_2(a565,a557)|c1_2(a565,a493)|c1_2(a565,X1)|c2_2(a565,a503)|c2_2(a565,X1)|c8_1(a565)|ssSkP4(a565)|ssSkP6(a565)|c1_1(a573)|c1_1(a565)|ssSkC7|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkP2(a565)), inference(spm,[status(thm)],[c_0_562, c_0_573]), ['final']).
cnf(c_0_909,negated_conjecture,(c6_2(a542,a515)|c6_2(a542,X1)|c4_2(a542,X1)|c5_2(a542,X1)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X2)|c2_2(a542,X1)|c2_2(a542,X2)|c8_1(a542)|ssSkP5(a542)|c1_1(a542)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a542)|c5_1(a565)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20|~c3_2(a542,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_570]), c_0_48]), ['final']).
cnf(c_0_910,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,X1)|c5_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c1_2(a565,X2)|c2_2(a565,X1)|c2_2(a565,X2)|c8_1(a565)|ssSkP6(a565)|c1_1(a565)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~c3_2(a565,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_381]), c_0_353]), ['final']).
cnf(c_0_911,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c4_2(a565,X1)|c5_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c1_2(a565,X2)|c2_2(a565,X1)|c2_2(a565,X2)|c8_1(a565)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~c3_2(a565,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_381]), c_0_353]), ['final']).
cnf(c_0_912,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,X1)|c5_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c1_2(a565,X2)|c2_2(a565,X1)|c2_2(a565,X2)|c8_1(a565)|ssSkP5(a565)|c1_1(a565)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~c3_2(a565,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_385]), c_0_353]), ['final']).
cnf(c_0_913,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c4_2(a565,X1)|c5_2(a565,X1)|c3_2(a565,X1)|c1_2(a565,a557)|c1_2(a565,X2)|c2_2(a565,X1)|c2_2(a565,X2)|c8_1(a565)|ssSkP5(a565)|c1_1(a565)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~c3_2(a565,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_385]), c_0_353]), ['final']).
cnf(c_0_914,negated_conjecture,(ssSkC12|c6_2(a542,a588)|c4_2(a542,a527)|c3_2(a542,a588)|c3_2(a542,a493)|c3_2(a542,a557)|c2_2(a542,a493)|c2_2(a542,a588)|c2_2(a542,a557)|c1_1(a542)|ssSkP8(a542)|ssSkC7|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkP6(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_574]), c_0_278]), ['final']).
cnf(c_0_915,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c4_2(a552,a515)|c1_2(a552,a493)|c1_2(a552,a557)|c2_2(a552,a557)|c2_2(a552,a502)|ssSkP3(a552)|ssSkP5(a552)|c1_1(a573)|ssSkC7|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkP2(a552)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_575]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_916,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c4_2(a552,a515)|c1_2(a552,a493)|c1_2(a552,a557)|c2_2(a552,a557)|c2_2(a552,a502)|ssSkP3(a552)|ssSkP5(a552)|ssSkC7|c4_1(a531)|c3_1(a573)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkP2(a552)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_576]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_917,negated_conjecture,(ssSkC12|c6_2(a542,a487)|c4_2(a542,a527)|c3_2(a542,a487)|c3_2(a542,a493)|c3_2(a542,a588)|c3_2(a542,a557)|c2_2(a542,a493)|c2_2(a542,a588)|c2_2(a542,a557)|ssSkP8(a542)|ssSkC5|ssSkC7|c4_1(a531)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkP6(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_577]), c_0_278]), ['final']).
cnf(c_0_918,negated_conjecture,(c6_2(a542,a465)|c7_2(a542,a465)|c7_2(X1,X2)|c5_2(a542,a456)|c3_2(X1,X2)|c1_2(X3,X4)|c1_2(X1,X2)|c8_1(a542)|c8_1(X3)|ssSkP0(a542)|c1_1(a542)|c1_1(X1)|c2_1(X1)|c7_0|c5_1(a501)|ssSkC16|~c6_2(X3,X4)|~c7_2(a542,a456)|~c2_2(X3,X4)|~c1_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_578, c_0_121]), ['final']).
cnf(c_0_919,negated_conjecture,(ssSkC8|c6_2(a565,a465)|c7_2(a565,a465)|c7_2(X1,X2)|c5_2(a565,a456)|c3_2(X1,X2)|c1_2(X3,X4)|c1_2(X1,X2)|c8_1(a565)|c8_1(X3)|ssSkP0(a565)|c1_1(a565)|c1_1(X1)|c2_1(X1)|c7_0|c5_1(a501)|ssSkC20|~c6_2(X3,X4)|~c7_2(a565,a456)|~c2_2(X3,X4)|~c1_1(X3)), inference(spm,[status(thm)],[c_0_579, c_0_121]), ['final']).
cnf(c_0_920,negated_conjecture,(ssSkC8|c6_2(a517,a465)|c7_2(a517,a465)|c7_2(X1,X2)|c5_2(a517,a456)|c3_2(X1,X2)|c1_2(X3,X4)|c1_2(X1,X2)|c8_1(a517)|c8_1(X3)|ssSkP0(a517)|c1_1(a517)|c1_1(X1)|c2_1(X1)|c7_0|c5_1(a501)|ssSkC11|~c6_2(X3,X4)|~c7_2(a517,a456)|~c2_2(X3,X4)|~c1_1(X3)), inference(spm,[status(thm)],[c_0_580, c_0_121]), ['final']).
cnf(c_0_921,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c6_2(a565,X1)|c3_2(a565,a460)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,a460)|c2_2(a565,X1)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_404, c_0_581]), ['final']).
cnf(c_0_922,negated_conjecture,(ssSkC12|c6_2(a542,a515)|c4_2(a542,a527)|c3_2(a542,a493)|c3_2(a542,a588)|c3_2(a542,a557)|c2_2(a542,a493)|c2_2(a542,a588)|c2_2(a542,a557)|ssSkP5(a542)|ssSkP8(a542)|ssSkC7|c4_1(a531)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkP6(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_582]), c_0_278]), ['final']).
cnf(c_0_923,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c8_2(a592,a515)|c8_2(X1,a467)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c1_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_584]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_924,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c8_2(a592,a515)|c3_2(X1,a467)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c1_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_584]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_925,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c8_2(a592,a515)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c2_2(X1,a467)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c1_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_584]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_926,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c6_2(a565,X1)|c3_2(a565,a460)|c3_2(a565,X1)|c1_2(a565,a557)|c2_2(a565,a460)|c2_2(a565,X1)|ssSkP5(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_408, c_0_581]), ['final']).
cnf(c_0_927,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c8_2(a592,a515)|c8_2(X1,a467)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c3_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_587]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_928,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c8_2(a592,a515)|c3_2(X1,a467)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c3_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_587]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_929,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c4_2(a552,a526)|c1_2(a552,a557)|c1_2(a552,a493)|c2_2(a552,a557)|c2_2(a552,a502)|ssSkP3(a552)|ssSkP5(a552)|ssSkP6(a552)|ssSkC7|c4_1(a531)|c3_1(a573)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkP2(a552)), inference(spm,[status(thm)],[c_0_562, c_0_588]), ['final']).
cnf(c_0_930,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c4_2(a552,a526)|c1_2(a552,a557)|c1_2(a552,a493)|c2_2(a552,a557)|c2_2(a552,a502)|ssSkP3(a552)|ssSkP5(a552)|ssSkP6(a552)|c1_1(a573)|ssSkC7|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkP2(a552)), inference(spm,[status(thm)],[c_0_562, c_0_589]), ['final']).
cnf(c_0_931,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c8_2(a592,a515)|c1_2(a592,a557)|c2_2(a592,a502)|c2_2(a592,a557)|c2_2(X1,a467)|c8_1(a592)|ssSkP3(a592)|ssSkP5(a592)|c3_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_587]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_932,negated_conjecture,(c7_2(X1,X2)|c5_2(a570,a456)|c3_2(X1,X2)|c1_2(X3,X4)|c1_2(X1,X2)|c2_2(a570,a572)|c2_2(a570,a465)|c8_1(a570)|c8_1(X3)|ssSkP0(a570)|c1_1(X1)|c2_1(X1)|~c6_2(X3,X4)|~c7_2(a570,a456)|~c2_2(X3,X4)|~c1_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_590, c_0_121]), ['final']).
cnf(c_0_933,negated_conjecture,(ssSkC8|c6_2(a600,a526)|c4_2(a600,a557)|c5_2(a600,a557)|c1_2(a600,X1)|c2_2(a600,a503)|c2_2(a600,X1)|c8_1(a600)|ssSkP4(a600)|ssSkP6(a600)|c1_1(a573)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_591]), c_0_353]), ['final']).
cnf(c_0_934,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c3_2(a552,a460)|c3_2(a552,a557)|c2_2(a552,a460)|c2_2(a552,a557)|ssSkP6(a552)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_592, c_0_581]), ['final']).
cnf(c_0_935,negated_conjecture,(ssSkC8|c6_2(a600,a526)|c4_2(a600,a557)|c5_2(a600,a557)|c1_2(a600,X1)|c2_2(a600,a503)|c2_2(a600,X1)|c8_1(a600)|ssSkP4(a600)|ssSkP6(a600)|c4_1(a531)|c3_1(a573)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_591]), c_0_353]), ['final']).
cnf(c_0_936,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c3_2(a552,a460)|c3_2(a552,a557)|c2_2(a552,a460)|c2_2(a552,a557)|ssSkP5(a552)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_593, c_0_581]), ['final']).
cnf(c_0_937,negated_conjecture,(ssSkC8|c6_2(a600,a515)|c4_2(a600,a557)|c5_2(a600,a557)|c1_2(a600,X1)|c2_2(a600,a503)|c2_2(a600,X1)|c8_1(a600)|ssSkP4(a600)|ssSkP5(a600)|c1_1(a573)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_594]), c_0_353]), ['final']).
cnf(c_0_938,negated_conjecture,(ssSkC8|c6_2(a600,a515)|c4_2(a600,a557)|c5_2(a600,a557)|c1_2(a600,X1)|c2_2(a600,a503)|c2_2(a600,X1)|c8_1(a600)|ssSkP4(a600)|ssSkP5(a600)|c4_1(a531)|c3_1(a573)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_594]), c_0_353]), ['final']).
cnf(c_0_939,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(a538)|ssSkP6(a552)|c4_1(a552)|c4_1(a531)|c3_1(a538)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_468, c_0_280]), ['final']).
cnf(c_0_940,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(a494)|ssSkP6(a552)|c4_1(a552)|c4_1(a531)|c3_1(a494)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_468, c_0_281]), ['final']).
cnf(c_0_941,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(a558)|ssSkP6(a552)|c4_1(a552)|c4_1(a531)|c3_1(a558)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_468, c_0_282]), ['final']).
cnf(c_0_942,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP6(a552)|c4_1(a531)|c4_1(a552)|c3_1(a466)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_595, c_0_596]), ['final']).
cnf(c_0_943,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c4_2(a552,a526)|c1_2(a552,a557)|c1_2(a552,a493)|c2_2(a552,a557)|c2_2(a552,a502)|ssSkP3(a552)|ssSkP6(a552)|ssSkC7|c4_1(a531)|c3_1(a573)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkP2(a552)), inference(spm,[status(thm)],[c_0_562, c_0_597]), ['final']).
cnf(c_0_944,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c4_2(a552,a526)|c1_2(a552,a557)|c1_2(a552,a493)|c2_2(a552,a557)|c2_2(a552,a502)|ssSkP3(a552)|ssSkP6(a552)|c1_1(a573)|ssSkC7|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkP2(a552)), inference(spm,[status(thm)],[c_0_562, c_0_598]), ['final']).
cnf(c_0_945,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,a493)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP6(a552)|ssSkC7|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c5_2(a552,X1)|~ssSkP2(a552)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_414]), c_0_353]), ['final']).
cnf(c_0_946,negated_conjecture,(ssSkC8|c6_2(a517,a588)|c3_2(a517,a588)|c1_2(a517,a557)|c2_2(a517,a503)|ssSkP4(a517)|c1_1(a517)|ssSkP8(a517)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_443, c_0_62]), c_0_63]), ['final']).
cnf(c_0_947,negated_conjecture,(ssSkC8|c6_2(a517,a588)|c3_2(a517,a460)|c3_2(a517,a588)|c1_2(a517,a557)|c2_2(a517,a460)|c1_1(a517)|ssSkP8(a517)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_443, c_0_581]), ['final']).
cnf(c_0_948,negated_conjecture,(ssSkC8|c6_2(a517,a588)|c3_2(a517,a487)|c3_2(a517,a588)|c1_2(a517,a557)|c2_2(a517,a487)|c1_1(a517)|ssSkP8(a517)|ssSkC5|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20), inference(spm,[status(thm)],[c_0_443, c_0_599]), ['final']).
cnf(c_0_949,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(a538)|ssSkP5(a552)|c4_1(a552)|c4_1(a531)|c3_1(a538)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_470, c_0_280]), ['final']).
cnf(c_0_950,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(a494)|ssSkP5(a552)|c4_1(a552)|c4_1(a531)|c3_1(a494)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_470, c_0_281]), ['final']).
cnf(c_0_951,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(a558)|ssSkP5(a552)|c4_1(a552)|c4_1(a531)|c3_1(a558)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_470, c_0_282]), ['final']).
cnf(c_0_952,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP5(a552)|c4_1(a531)|c4_1(a552)|c3_1(a466)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_600, c_0_596]), ['final']).
cnf(c_0_953,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c1_2(a517,a557)|c2_2(a517,a503)|ssSkP1(a517)|ssSkP4(a517)|c1_1(a517)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_447, c_0_62]), c_0_63]), ['final']).
cnf(c_0_954,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c3_2(a517,a460)|c1_2(a517,a557)|c2_2(a517,a460)|ssSkP1(a517)|c1_1(a517)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_447, c_0_581]), ['final']).
cnf(c_0_955,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c3_2(a517,a487)|c1_2(a517,a557)|c2_2(a517,a487)|ssSkP1(a517)|c1_1(a517)|ssSkC5|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a517)|c5_1(a501)|c7_1(a517)|ssSkC11|ssSkC20), inference(spm,[status(thm)],[c_0_447, c_0_599]), ['final']).
cnf(c_0_956,negated_conjecture,(ssSkC8|c7_2(X1,X2)|c5_2(a537,a456)|c3_2(X1,X2)|c1_2(X3,X4)|c1_2(X1,X2)|c2_2(a537,a572)|c2_2(a537,a465)|c8_1(a537)|c8_1(X3)|ssSkP0(a537)|c1_1(X1)|c2_1(X1)|ssSkC15|~c6_2(X3,X4)|~c7_2(a537,a456)|~c2_2(X3,X4)|~c1_1(X3)), inference(spm,[status(thm)],[c_0_601, c_0_121]), ['final']).
cnf(c_0_957,negated_conjecture,(ssSkC8|c7_2(X1,X2)|c5_2(a517,a456)|c3_2(X1,X2)|c1_2(X3,X4)|c1_2(X1,X2)|c2_2(a517,a572)|c2_2(a517,a465)|c8_1(a517)|c8_1(X3)|ssSkP0(a517)|c1_1(X1)|c2_1(X1)|ssSkC11|~c6_2(X3,X4)|~c7_2(a517,a456)|~c2_2(X3,X4)|~c1_1(X3)), inference(spm,[status(thm)],[c_0_602, c_0_121]), ['final']).
cnf(c_0_958,negated_conjecture,(ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP6(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c6_2(a552,X1)|~c4_2(a552,a548)|~c1_1(a552)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_414]), c_0_353]), ['final']).
cnf(c_0_959,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP6(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c6_2(a552,X1)|~c7_2(a552,a548)|~c1_1(a552)|~ssSkC17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_414]), c_0_353]), c_0_604]), ['final']).
cnf(c_0_960,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP5(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c6_2(a552,X1)|~c4_2(a552,a548)|~c1_1(a552)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_416]), c_0_353]), ['final']).
cnf(c_0_961,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP5(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c6_2(a552,X1)|~c7_2(a552,a548)|~c1_1(a552)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_416]), c_0_353]), ['final']).
cnf(c_0_962,negated_conjecture,(ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,a493)|c2_2(a552,a557)|c2_2(a552,X1)|ssSkP5(a552)|ssSkC7|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c5_2(a552,X1)|~ssSkP2(a552)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_416]), c_0_353]), ['final']).
cnf(c_0_963,negated_conjecture,(c6_2(a542,a588)|c3_2(a542,a588)|c1_2(a542,a557)|c2_2(a542,a503)|ssSkP4(a542)|c1_1(a542)|ssSkP8(a542)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_62]), c_0_63]), ['final']).
cnf(c_0_964,negated_conjecture,(c6_2(a542,a588)|c3_2(a542,a460)|c3_2(a542,a588)|c1_2(a542,a557)|c2_2(a542,a460)|c1_1(a542)|ssSkP8(a542)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_444, c_0_581]), ['final']).
cnf(c_0_965,negated_conjecture,(c6_2(a542,a588)|c3_2(a542,a487)|c3_2(a542,a588)|c1_2(a542,a557)|c2_2(a542,a487)|c1_1(a542)|ssSkP8(a542)|ssSkC5|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_444, c_0_605]), ['final']).
cnf(c_0_966,negated_conjecture,(ssSkC12|ssSkC8|c6_2(a552,X1)|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c3_2(a552,X2)|c2_2(a552,a557)|c2_2(a552,X2)|c2_1(a552)|c4_1(a458)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20|~c7_2(a552,X1)|~c1_2(a552,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_425]), c_0_353]), ['final']).
cnf(c_0_967,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(a538)|c2_1(a552)|c4_1(a552)|c4_1(a458)|c3_1(a538)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_474, c_0_280]), ['final']).
cnf(c_0_968,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(a494)|c2_1(a552)|c4_1(a552)|c4_1(a458)|c3_1(a494)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_474, c_0_281]), ['final']).
cnf(c_0_969,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c8_1(a558)|c2_1(a552)|c4_1(a552)|c4_1(a458)|c3_1(a558)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_474, c_0_282]), ['final']).
cnf(c_0_970,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c2_1(a552)|c4_1(a458)|c4_1(a552)|c3_1(a466)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_606, c_0_596]), ['final']).
cnf(c_0_971,negated_conjecture,(c6_2(a542,a485)|c1_2(a542,a557)|c2_2(a542,a503)|ssSkP1(a542)|ssSkP4(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_62]), c_0_63]), ['final']).
cnf(c_0_972,negated_conjecture,(c6_2(a542,a485)|c3_2(a542,a460)|c1_2(a542,a557)|c2_2(a542,a460)|ssSkP1(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_451, c_0_581]), ['final']).
cnf(c_0_973,negated_conjecture,(c6_2(a542,a485)|c3_2(a542,a487)|c1_2(a542,a557)|c2_2(a542,a487)|ssSkP1(a542)|c1_1(a542)|ssSkC5|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a542)|c5_1(a501)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_451, c_0_605]), ['final']).
cnf(c_0_974,negated_conjecture,(ssSkC8|c6_2(a600,a503)|c4_2(a600,a557)|c5_2(a600,a557)|c1_2(a600,X1)|c2_2(a600,a503)|c2_2(a600,X1)|c8_1(a600)|ssSkP4(a600)|c1_1(a573)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_607]), c_0_353]), ['final']).
cnf(c_0_975,negated_conjecture,(ssSkC8|c6_2(a600,a503)|c4_2(a600,a557)|c5_2(a600,a557)|c1_2(a600,X1)|c2_2(a600,a503)|c2_2(a600,X1)|c8_1(a600)|ssSkP4(a600)|c4_1(a531)|c3_1(a573)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_607]), c_0_353]), ['final']).
cnf(c_0_976,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a588)|c5_2(a542,a543)|c3_2(a542,a543)|c3_2(a542,a588)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|c1_1(a542)|c1_1(a573)|ssSkP8(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_608]), c_0_278]), c_0_283]), ['final']).
cnf(c_0_977,negated_conjecture,(ssSkC8|c6_2(a600,a557)|c3_2(a600,a557)|c1_2(a600,a557)|c2_2(a600,a503)|c2_2(a600,a502)|ssSkP3(a600)|ssSkP4(a600)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_609, c_0_62]), c_0_63]), ['final']).
cnf(c_0_978,negated_conjecture,(ssSkC8|c6_2(a600,a557)|c3_2(a600,a460)|c3_2(a600,a557)|c1_2(a600,a557)|c2_2(a600,a460)|c2_2(a600,a502)|ssSkP3(a600)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26|~ssSkC0), inference(spm,[status(thm)],[c_0_609, c_0_581]), ['final']).
cnf(c_0_979,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a588)|c5_2(a542,a543)|c3_2(a542,a543)|c3_2(a542,a588)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|c1_1(a542)|ssSkP8(a542)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_608]), c_0_278]), c_0_283]), ['final']).
cnf(c_0_980,negated_conjecture,(ssSkC8|c6_2(a565,a503)|c6_2(a565,a526)|c1_2(a565,a557)|c2_2(a565,a503)|ssSkP4(a565)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_610, c_0_62]), c_0_63]), ['final']).
cnf(c_0_981,negated_conjecture,(ssSkC8|c6_2(a565,a460)|c6_2(a565,a526)|c3_2(a565,a460)|c1_2(a565,a557)|c2_2(a565,a460)|ssSkP6(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_610, c_0_581]), ['final']).
cnf(c_0_982,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a485)|c5_2(a542,a543)|c3_2(a542,a543)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|ssSkP1(a542)|c1_1(a542)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_611]), c_0_279]), c_0_283]), ['final']).
cnf(c_0_983,negated_conjecture,(ssSkC8|c6_2(a565,a503)|c6_2(a565,a515)|c1_2(a565,a557)|c2_2(a565,a503)|ssSkP4(a565)|ssSkP5(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_612, c_0_62]), c_0_63]), ['final']).
cnf(c_0_984,negated_conjecture,(ssSkC8|c6_2(a565,a460)|c6_2(a565,a515)|c3_2(a565,a460)|c1_2(a565,a557)|c2_2(a565,a460)|ssSkP5(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_612, c_0_581]), ['final']).
cnf(c_0_985,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a485)|c5_2(a542,a543)|c3_2(a542,a543)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|ssSkP1(a542)|c1_1(a542)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_611]), c_0_279]), c_0_283]), ['final']).
cnf(c_0_986,negated_conjecture,(ssSkC8|c6_2(a565,a588)|c3_2(a565,a588)|c1_2(a565,a557)|c2_2(a565,a503)|ssSkP4(a565)|c1_1(a565)|ssSkP8(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_448, c_0_62]), c_0_63]), ['final']).
cnf(c_0_987,negated_conjecture,(ssSkC8|c6_2(a565,a588)|c3_2(a565,a487)|c3_2(a565,a588)|c1_2(a565,a557)|c2_2(a565,a487)|c1_1(a565)|ssSkP8(a565)|ssSkC5|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_448, c_0_613]), ['final']).
cnf(c_0_988,negated_conjecture,(ssSkC8|c6_2(a565,a588)|c3_2(a565,a460)|c3_2(a565,a588)|c1_2(a565,a557)|c2_2(a565,a460)|c1_1(a565)|ssSkP8(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_448, c_0_581]), ['final']).
cnf(c_0_989,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c1_2(a565,a557)|c2_2(a565,a503)|ssSkP1(a565)|ssSkP4(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_62]), c_0_63]), ['final']).
cnf(c_0_990,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c3_2(a565,a460)|c1_2(a565,a557)|c2_2(a565,a460)|ssSkP1(a565)|c1_1(a565)|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_456, c_0_581]), ['final']).
cnf(c_0_991,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c3_2(a565,a487)|c1_2(a565,a557)|c2_2(a565,a487)|ssSkP1(a565)|c1_1(a565)|ssSkC5|c7_0|c4_1(a531)|c5_1(a565)|c5_1(a501)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_456, c_0_613]), ['final']).
cnf(c_0_992,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a526)|c5_2(a542,a543)|c3_2(a542,a543)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|ssSkP6(a542)|c1_1(a542)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_614]), c_0_31]), c_0_283]), ['final']).
cnf(c_0_993,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a526)|c5_2(a542,a543)|c3_2(a542,a543)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|ssSkP6(a542)|c1_1(a542)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_614]), c_0_31]), c_0_283]), ['final']).
cnf(c_0_994,negated_conjecture,(ssSkC8|c6_2(a600,a503)|c3_2(a600,a460)|c1_2(a600,a557)|c2_2(a600,a460)|ssSkP4(a600)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26|~ssSkC0), inference(spm,[status(thm)],[c_0_615, c_0_581]), ['final']).
cnf(c_0_995,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a515)|c5_2(a542,a543)|c3_2(a542,a543)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|ssSkP5(a542)|c1_1(a542)|c1_1(a573)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_616]), c_0_48]), c_0_283]), ['final']).
cnf(c_0_996,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a515)|c5_2(a542,a543)|c3_2(a542,a543)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|ssSkP5(a542)|c1_1(a542)|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_616]), c_0_48]), c_0_283]), ['final']).
cnf(c_0_997,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a526)|c3_2(a552,a557)|c2_2(a552,a503)|c2_2(a552,a557)|ssSkP4(a552)|ssSkP6(a552)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(spm,[status(thm)],[c_0_389, c_0_617]), ['final']).
cnf(c_0_998,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a487)|c5_2(a542,a543)|c3_2(a542,a543)|c3_2(a542,a487)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|c1_1(a542)|c1_1(a573)|ssSkC5|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_618]), c_0_27]), c_0_283]), ['final']).
cnf(c_0_999,negated_conjecture,(ssSkC8|c6_2(a517,a526)|c3_2(a517,a487)|c1_2(a517,a557)|c2_2(a517,a487)|ssSkP6(a517)|ssSkC5|c4_1(a531)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20), inference(spm,[status(thm)],[c_0_367, c_0_599]), ['final']).
cnf(c_0_1000,negated_conjecture,(c6_2(a542,a543)|c6_2(a542,a487)|c5_2(a542,a543)|c3_2(a542,a543)|c3_2(a542,a487)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|c1_1(a542)|ssSkC5|c7_0|c4_1(a531)|c3_1(a573)|c5_1(a501)|c5_1(a542)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_618]), c_0_27]), c_0_283]), ['final']).
cnf(c_0_1001,negated_conjecture,(ssSkC8|c6_2(a565,a526)|c3_2(a565,a487)|c1_2(a565,a557)|c2_2(a565,a487)|ssSkP6(a565)|ssSkC5|c4_1(a531)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_367, c_0_613]), ['final']).
cnf(c_0_1002,negated_conjecture,(ssSkC8|c6_2(X1,a526)|c3_2(X1,a460)|c1_2(X1,a557)|c2_2(X1,a460)|ssSkP6(X1)|c4_1(a531)|c5_1(a565)|c5_1(X1)|c7_1(X1)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_367, c_0_581]), ['final']).
cnf(c_0_1003,negated_conjecture,(ssSkC8|c6_2(a552,a557)|c6_2(a552,a515)|c3_2(a552,a557)|c2_2(a552,a503)|c2_2(a552,a557)|ssSkP4(a552)|ssSkP5(a552)|c4_1(a531)|c5_1(a565)|c5_1(a552)|c7_1(a552)|ssSkC18|ssSkC20), inference(spm,[status(thm)],[c_0_390, c_0_619]), ['final']).
cnf(c_0_1004,negated_conjecture,(ssSkC12|c6_2(a512,X1)|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c3_2(a512,X2)|c2_2(a512,a557)|c2_2(a512,X2)|c2_1(a512)|c4_1(a458)|c4_1(a531)|c5_1(a565)|c5_1(a512)|c7_1(a512)|ssSkC9|ssSkC20|~c7_2(a512,X1)|~c1_2(a512,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_428]), c_0_32]), ['final']).
cnf(c_0_1005,negated_conjecture,(ssSkC8|c6_2(X1,a515)|c3_2(X1,a460)|c1_2(X1,a557)|c2_2(X1,a460)|ssSkP5(X1)|c4_1(a531)|c5_1(a565)|c5_1(X1)|c7_1(X1)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_370, c_0_581]), ['final']).
cnf(c_0_1006,negated_conjecture,(ssSkC8|c6_2(a517,a515)|c3_2(a517,a487)|c1_2(a517,a557)|c2_2(a517,a487)|ssSkP5(a517)|ssSkC5|c4_1(a531)|c5_1(a517)|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20), inference(spm,[status(thm)],[c_0_370, c_0_599]), ['final']).
cnf(c_0_1007,negated_conjecture,(ssSkC12|c6_2(a512,X1)|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c2_1(a512)|c4_1(a458)|c4_1(a531)|c5_1(a565)|c5_1(a512)|c7_1(a512)|ssSkC9|ssSkC20|~c7_2(a512,X1)|~c1_2(a512,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_428]), c_0_32]), ['final']).
cnf(c_0_1008,negated_conjecture,(ssSkC8|c6_2(a565,a515)|c3_2(a565,a487)|c1_2(a565,a557)|c2_2(a565,a487)|ssSkP5(a565)|ssSkC5|c4_1(a531)|c5_1(a565)|c7_1(a565)|ssSkC20), inference(spm,[status(thm)],[c_0_370, c_0_613]), ['final']).
cnf(c_0_1009,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c1_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c8_1(a538)|c8_1(a512)|c2_1(a512)|c4_1(a512)|c4_1(a458)|c3_1(a538)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_481, c_0_280]), ['final']).
cnf(c_0_1010,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c1_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c8_1(a494)|c8_1(a512)|c2_1(a512)|c4_1(a512)|c4_1(a458)|c3_1(a494)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_481, c_0_281]), ['final']).
cnf(c_0_1011,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c6_2(a592,a526)|c3_2(a592,a557)|c2_2(a592,a557)|ssSkP6(a592)|c4_1(a531)|c5_1(a565)|c5_1(a592)|c7_1(a592)|ssSkC20|ssSkC24|~c1_2(a592,a557)), inference(ef,[status(thm)],[c_0_620]), ['final']).
cnf(c_0_1012,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c1_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c8_1(a558)|c8_1(a512)|c2_1(a512)|c4_1(a512)|c4_1(a458)|c3_1(a558)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_481, c_0_282]), ['final']).
cnf(c_0_1013,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c1_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c8_1(a512)|c2_1(a512)|c4_1(a458)|c4_1(a512)|c3_1(a466)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_621, c_0_596]), ['final']).
cnf(c_0_1014,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c6_2(a592,a515)|c3_2(a592,a557)|c2_2(a592,a557)|ssSkP5(a592)|c4_1(a531)|c5_1(a565)|c5_1(a592)|c7_1(a592)|ssSkC20|ssSkC24|~c1_2(a592,a557)), inference(ef,[status(thm)],[c_0_622]), ['final']).
cnf(c_0_1015,negated_conjecture,(c6_2(X1,a474)|c1_2(X2,a456)|c8_1(X3)|ssSkP0(X3)|ssSkP0(X2)|c6_1(X2)|~c3_2(X3,a456)|~c3_1(X3)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_623, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1016,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c1_2(a552,a493)|c2_2(a552,a557)|c2_2(a552,X1)|c2_1(a552)|ssSkC7|c4_1(a458)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c5_2(a552,X1)|~ssSkP2(a552)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_425]), c_0_353]), ['final']).
cnf(c_0_1017,negated_conjecture,(ssSkC8|c1_2(a600,X1)|c2_2(a600,a502)|c2_2(a600,a557)|c2_2(a600,X1)|ssSkP3(a600)|c1_1(X2)|c2_1(X2)|c2_0|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c5_2(a600,X1)|~c1_2(a600,a557)|~c4_1(X2)|~c5_1(a600)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_477]), c_0_353]), ['final']).
cnf(c_0_1018,negated_conjecture,(ssSkC8|c6_2(a592,a526)|c6_2(a592,a557)|c3_2(a592,a557)|c2_2(a592,a503)|ssSkP4(a592)|ssSkP6(a592)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(spm,[status(thm)],[c_0_374, c_0_625]), ['final']).
cnf(c_0_1019,negated_conjecture,(ssSkC8|c6_2(a592,a515)|c6_2(a592,a557)|c3_2(a592,a557)|c2_2(a592,a503)|ssSkP4(a592)|ssSkP5(a592)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(spm,[status(thm)],[c_0_374, c_0_626]), ['final']).
cnf(c_0_1020,negated_conjecture,(c7_2(X1,a474)|c1_2(X2,a456)|c8_1(X3)|ssSkP0(X3)|ssSkP0(X2)|c6_1(X2)|~c3_2(X3,a456)|~c3_1(X3)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1021,negated_conjecture,(c6_2(X1,a474)|c4_2(a542,a502)|c4_2(X2,X3)|c8_1(X2)|ssSkP3(a542)|c1_1(a542)|c6_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c6_2(X2,X3)|~c3_2(X2,X3)|~ndr1_1(X2)|~c3_1(X2)|~c5_1(X1)|~c7_1(a542)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_628]), c_0_8]), c_0_124]), ['final']).
cnf(c_0_1022,negated_conjecture,(ssSkC12|ssSkC8|c8_2(a552,X1)|c3_2(a552,a557)|c2_2(a552,a557)|c2_2(a552,X1)|c1_1(a508)|c2_1(a552)|c4_1(a458)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c7_2(a552,a527)|~c3_2(a552,a527)|~c4_1(a552)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_425]), c_0_353]), ['final']).
cnf(c_0_1023,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c1_2(a512,a493)|c2_2(a512,a557)|c2_2(a512,X1)|c2_1(a512)|ssSkC7|c4_1(a458)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c5_2(a512,X1)|~ssSkP2(a512)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_428]), c_0_32]), ['final']).
cnf(c_0_1024,negated_conjecture,(c4_2(a542,a502)|c4_2(X1,X2)|c7_2(X3,a474)|c8_1(X1)|ssSkP3(a542)|c1_1(a542)|c6_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a542)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_628]), c_0_8]), c_0_124]), ['final']).
cnf(c_0_1025,negated_conjecture,(c1_2(X1,a474)|c1_2(X2,a456)|c8_1(X3)|ssSkP0(X3)|ssSkP0(X2)|c6_1(X2)|~c3_2(X3,a456)|~c3_1(X3)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_630, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1026,negated_conjecture,(ssSkC8|c5_2(a573,a456)|c1_2(a573,X1)|c2_2(a573,X1)|c8_1(a573)|ssSkP0(a573)|~c7_2(a573,a456)|~c3_2(a573,X1)|~c2_2(a573,a456)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_631, c_0_121]), c_0_122]), ['final']).
cnf(c_0_1027,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c2_1(a552)|c4_1(a458)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c6_2(a552,X1)|~c4_2(a552,a548)|~c1_1(a552)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_425]), c_0_353]), ['final']).
cnf(c_0_1028,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a552,a527)|c3_2(a552,a557)|c3_2(a552,X1)|c2_2(a552,a557)|c2_2(a552,X1)|c2_1(a552)|c4_1(a458)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c6_2(a552,X1)|~c7_2(a552,a548)|~c1_1(a552)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_425]), c_0_353]), ['final']).
cnf(c_0_1029,negated_conjecture,(ssSkC12|c8_2(a512,X1)|c3_2(a512,a557)|c2_2(a512,a557)|c2_2(a512,X1)|c1_1(a508)|c2_1(a512)|c4_1(a458)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c7_2(a512,a527)|~c3_2(a512,a527)|~c4_1(a512)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_428]), c_0_32]), ['final']).
cnf(c_0_1030,negated_conjecture,(ssSkC12|ssSkC3|c6_2(a512,a479)|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c1_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c2_1(a512)|c4_1(a458)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c8_1(a512)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_428]), c_0_32]), ['final']).
cnf(c_0_1031,negated_conjecture,(c4_2(a542,a502)|c4_2(X1,X2)|c1_2(X3,a474)|c8_1(X1)|ssSkP3(a542)|c1_1(a542)|c6_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a542)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_628]), c_0_8]), c_0_124]), ['final']).
cnf(c_0_1032,negated_conjecture,(c1_2(X1,a456)|c8_1(X2)|ssSkP0(X2)|ssSkP0(X1)|c6_1(X1)|ndr1_1(X3)|~c3_2(X2,a456)|~c3_1(X2)|~c5_1(X3)|~c7_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_632, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1033,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c6_2(a592,X1)|c3_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,a557)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c7_2(a592,X1)|~c1_2(a592,a557)), inference(ef,[status(thm)],[c_0_480]), ['final']).
cnf(c_0_1034,negated_conjecture,(ssSkC8|c6_2(a600,a503)|c7_2(a600,a557)|c3_2(a600,a460)|c2_2(a600,a460)|ssSkP4(a600)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26|~ssSkC0), inference(spm,[status(thm)],[c_0_633, c_0_581]), ['final']).
cnf(c_0_1035,negated_conjecture,(ssSkC8|c6_2(a600,a526)|c7_2(a600,a557)|c3_2(a600,a460)|c2_2(a600,a460)|ssSkP6(a600)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26|~ssSkC0), inference(spm,[status(thm)],[c_0_464, c_0_581]), ['final']).
cnf(c_0_1036,negated_conjecture,(c4_2(a542,a502)|c4_2(X1,X2)|c8_1(X1)|ssSkP3(a542)|c1_1(a542)|c6_1(a542)|c7_0|ndr1_1(X3)|c5_1(a501)|ssSkC16|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a542)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_628]), c_0_8]), c_0_124]), ['final']).
cnf(c_0_1037,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c7_2(a592,a557)|~c1_2(a592,X1)), inference(ef,[status(thm)],[c_0_480]), ['final']).
cnf(c_0_1038,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c2_1(a512)|c4_1(a458)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c6_2(a512,X1)|~c4_2(a512,a548)|~c1_1(a512)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_428]), c_0_32]), ['final']).
cnf(c_0_1039,negated_conjecture,(ssSkC12|c4_2(a512,a527)|c3_2(a512,a557)|c3_2(a512,X1)|c2_2(a512,a557)|c2_2(a512,X1)|c2_1(a512)|c4_1(a458)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c6_2(a512,X1)|~c7_2(a512,a548)|~c1_1(a512)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_428]), c_0_32]), ['final']).
cnf(c_0_1040,negated_conjecture,(ssSkC8|c6_2(a600,a515)|c7_2(a600,a557)|c3_2(a600,a460)|c2_2(a600,a460)|ssSkP5(a600)|c4_1(a531)|c5_1(a565)|c5_1(a600)|c7_1(a600)|ssSkC20|ssSkC26|~ssSkC0), inference(spm,[status(thm)],[c_0_467, c_0_581]), ['final']).
cnf(c_0_1041,negated_conjecture,(ssSkC8|c6_2(X1,a474)|c4_2(a565,a502)|c1_2(X2,a456)|c8_1(a565)|ssSkP0(X2)|ssSkP3(a565)|c1_1(a565)|c6_1(X2)|c7_0|c5_1(a501)|ssSkC20|~c3_2(a565,a502)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_623, c_0_634]), c_0_6]), c_0_353]), ['final']).
cnf(c_0_1042,negated_conjecture,(ssSkC8|c6_2(X1,a474)|c4_2(a517,a502)|c1_2(X2,a456)|c8_1(a517)|ssSkP0(X2)|ssSkP3(a517)|c1_1(a517)|c6_1(X2)|c7_0|c5_1(a501)|ssSkC11|~c3_2(a517,a502)|~c3_1(a517)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_623, c_0_635]), c_0_353]), ['final']).
cnf(c_0_1043,negated_conjecture,(ssSkC8|c4_2(a565,a502)|c7_2(X1,a474)|c1_2(X2,a456)|c8_1(a565)|ssSkP0(X2)|ssSkP3(a565)|c1_1(a565)|c6_1(X2)|c7_0|c5_1(a501)|ssSkC20|~c3_2(a565,a502)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_634]), c_0_6]), c_0_353]), ['final']).
cnf(c_0_1044,negated_conjecture,(ssSkC8|c4_2(a517,a502)|c7_2(X1,a474)|c1_2(X2,a456)|c8_1(a517)|ssSkP0(X2)|ssSkP3(a517)|c1_1(a517)|c6_1(X2)|c7_0|c5_1(a501)|ssSkC11|~c3_2(a517,a502)|~c3_1(a517)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_635]), c_0_353]), ['final']).
cnf(c_0_1045,negated_conjecture,(ssSkC8|c4_2(a565,a502)|c1_2(X1,a474)|c1_2(X2,a456)|c8_1(a565)|ssSkP0(X2)|ssSkP3(a565)|c1_1(a565)|c6_1(X2)|c7_0|c5_1(a501)|ssSkC20|~c3_2(a565,a502)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_630, c_0_634]), c_0_6]), c_0_353]), ['final']).
cnf(c_0_1046,negated_conjecture,(ssSkC8|c4_2(a517,a502)|c1_2(X1,a474)|c1_2(X2,a456)|c8_1(a517)|ssSkP0(X2)|ssSkP3(a517)|c1_1(a517)|c6_1(X2)|c7_0|c5_1(a501)|ssSkC11|~c3_2(a517,a502)|~c3_1(a517)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_630, c_0_635]), c_0_353]), ['final']).
cnf(c_0_1047,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a557)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(a538)|c8_1(X1)|c4_1(X1)|c3_1(a538)|c5_1(a565)|c7_1(X1)|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_431, c_0_280]), ['final']).
cnf(c_0_1048,negated_conjecture,(ssSkC8|c3_2(X1,a557)|c1_2(X1,a557)|c2_2(X1,a557)|c8_1(a538)|c8_1(X1)|c4_1(X1)|c3_1(a538)|c5_1(a565)|c7_1(X1)|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_490, c_0_280]), ['final']).
cnf(c_0_1049,negated_conjecture,(ssSkC8|c3_2(X1,a557)|c1_2(X1,a557)|c2_2(X1,a557)|c8_1(a494)|c8_1(X1)|c4_1(X1)|c3_1(a494)|c5_1(a565)|c7_1(X1)|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_490, c_0_281]), ['final']).
cnf(c_0_1050,negated_conjecture,(ssSkC8|c3_2(X1,a557)|c1_2(X1,a557)|c2_2(X1,a557)|c8_1(a558)|c8_1(X1)|c4_1(X1)|c3_1(a558)|c5_1(a565)|c7_1(X1)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_490, c_0_282]), ['final']).
cnf(c_0_1051,negated_conjecture,(ssSkC8|c3_2(X1,a557)|c1_2(X1,a557)|c2_2(X1,a557)|c8_1(X1)|c4_1(X1)|c3_1(a466)|c5_1(a565)|c7_1(X1)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_636, c_0_596]), ['final']).
cnf(c_0_1052,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a557)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(a494)|c8_1(X1)|c4_1(X1)|c3_1(a494)|c5_1(a565)|c7_1(X1)|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_431, c_0_281]), ['final']).
cnf(c_0_1053,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a557)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(a558)|c8_1(X1)|c4_1(X1)|c3_1(a558)|c5_1(a565)|c7_1(X1)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_431, c_0_282]), ['final']).
cnf(c_0_1054,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a557)|c1_2(X1,X2)|c2_2(X1,X2)|c8_1(X1)|c4_1(X1)|c3_1(a466)|c5_1(a565)|c7_1(X1)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_637, c_0_596]), ['final']).
cnf(c_0_1055,negated_conjecture,(c6_2(X1,a474)|c1_2(X2,a456)|c8_1(X2)|ssSkP0(X2)|c6_1(X2)|~c3_2(X2,a456)|~c3_1(X2)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_638, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1056,negated_conjecture,(c6_2(X1,a474)|c4_2(a542,a502)|c8_1(a542)|ssSkP3(a542)|c1_1(a542)|c6_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c3_2(a542,a502)|~c3_1(a542)|~c5_1(X1)|~c7_1(a542)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_638, c_0_628]), c_0_8]), c_0_124]), ['final']).
cnf(c_0_1057,negated_conjecture,(c4_2(a542,a502)|c7_2(X1,a474)|c8_1(a542)|ssSkP3(a542)|c1_1(a542)|c6_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c3_2(a542,a502)|~c3_1(a542)|~c5_1(X1)|~c7_1(a542)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_639, c_0_628]), c_0_8]), c_0_124]), ['final']).
cnf(c_0_1058,negated_conjecture,(c4_2(a542,a502)|c1_2(X1,a474)|c8_1(a542)|ssSkP3(a542)|c1_1(a542)|c6_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c3_2(a542,a502)|~c3_1(a542)|~c5_1(X1)|~c7_1(a542)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_640, c_0_628]), c_0_8]), c_0_124]), ['final']).
cnf(c_0_1059,negated_conjecture,(c6_2(a546,a526)|c1_2(a546,a557)|c2_2(a546,a503)|ssSkP4(a546)|ssSkP6(a546)|c4_1(a531)|c5_1(a565)|c5_1(a546)|c7_1(a546)|ssSkC17|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_641, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1060,negated_conjecture,(c6_2(a546,a526)|c3_2(a546,a460)|c1_2(a546,a557)|c2_2(a546,a460)|ssSkP6(a546)|c4_1(a531)|c5_1(a565)|c5_1(a546)|c7_1(a546)|ssSkC17|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_641, c_0_581]), ['final']).
cnf(c_0_1061,negated_conjecture,(c6_2(X1,a474)|c4_2(a542,a502)|c1_2(X2,a456)|c8_1(a542)|ssSkP0(X2)|ssSkP3(a542)|c1_1(a542)|c6_1(X2)|c7_0|c5_1(a501)|ssSkC16|~c3_2(a542,a502)|~c3_1(a542)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_623, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1062,negated_conjecture,(c4_2(a542,a502)|c7_2(X1,a474)|c1_2(X2,a456)|c8_1(a542)|ssSkP0(X2)|ssSkP3(a542)|c1_1(a542)|c6_1(X2)|c7_0|c5_1(a501)|ssSkC16|~c3_2(a542,a502)|~c3_1(a542)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1063,negated_conjecture,(c4_2(a542,a502)|c1_2(X1,a474)|c1_2(X2,a456)|c8_1(a542)|ssSkP0(X2)|ssSkP3(a542)|c1_1(a542)|c6_1(X2)|c7_0|c5_1(a501)|ssSkC16|~c3_2(a542,a502)|~c3_1(a542)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_630, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1064,negated_conjecture,(ssSkC8|c6_2(a600,a503)|c1_2(a600,a557)|c2_2(a600,a503)|ssSkP4(a600)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_478, c_0_479]), c_0_62]), c_0_63]), ['final']).
cnf(c_0_1065,negated_conjecture,(ssSkC8|c6_2(X1,a515)|c1_2(X1,a557)|ssSkP5(X1)|c4_1(a531)|c5_1(a565)|c5_1(X1)|c7_1(X1)|ssSkC20|~c1_2(X1,a515)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_478, c_0_23]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1066,negated_conjecture,(c6_2(a546,a515)|c1_2(a546,a557)|c2_2(a546,a503)|ssSkP4(a546)|ssSkP5(a546)|c4_1(a531)|c5_1(a565)|c5_1(a546)|c7_1(a546)|ssSkC17|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_642, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1067,negated_conjecture,(c6_2(a546,a515)|c3_2(a546,a460)|c1_2(a546,a557)|c2_2(a546,a460)|ssSkP5(a546)|c4_1(a531)|c5_1(a565)|c5_1(a546)|c7_1(a546)|ssSkC17|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_642, c_0_581]), ['final']).
cnf(c_0_1068,negated_conjecture,(c6_2(a542,a526)|c3_2(a542,a588)|c3_2(a542,a493)|c1_2(a542,a557)|c2_2(a542,a588)|c2_2(a542,a493)|ssSkP6(a542)|ssSkP8(a542)|ssSkC7|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_452, c_0_453]), ['final']).
cnf(c_0_1069,negated_conjecture,(ssSkC8|c6_2(a592,a515)|c6_2(a592,a557)|c5_2(a592,a557)|ssSkP5(a592)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c1_2(a592,a515)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_643, c_0_23]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1070,negated_conjecture,(ssSkC8|c6_2(a592,a515)|c6_2(a592,a557)|c3_2(a592,a557)|ssSkP5(a592)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c1_2(a592,a515)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_644, c_0_23]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1071,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c2_2(a592,a557)|c4_1(a531)|c5_1(a592)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c7_2(a592,a557)|~c1_2(a592,a557)), inference(ef,[status(thm)],[c_0_644]), ['final']).
cnf(c_0_1072,negated_conjecture,(ssSkC8|c4_2(a600,a557)|c5_2(a600,a557)|c1_2(a600,X1)|c2_2(a600,a502)|c2_2(a600,X1)|c8_1(a600)|ssSkP3(a600)|c1_1(a573)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_477]), c_0_353]), ['final']).
cnf(c_0_1073,negated_conjecture,(ssSkC8|c4_2(a600,a557)|c5_2(a600,a557)|c1_2(a600,X1)|c2_2(a600,a502)|c2_2(a600,X1)|c8_1(a600)|ssSkP3(a600)|c3_1(a573)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_477]), c_0_353]), ['final']).
cnf(c_0_1074,negated_conjecture,(c6_2(a542,a526)|c1_2(a542,a557)|c2_2(a542,a503)|ssSkP4(a542)|ssSkP6(a542)|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_452, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1075,negated_conjecture,(c6_2(a542,a526)|c3_2(a542,a460)|c1_2(a542,a557)|c2_2(a542,a460)|ssSkP6(a542)|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_452, c_0_581]), ['final']).
cnf(c_0_1076,negated_conjecture,(c6_2(a542,a526)|c3_2(a542,a487)|c1_2(a542,a557)|c2_2(a542,a487)|ssSkP6(a542)|ssSkC5|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_452, c_0_605]), ['final']).
cnf(c_0_1077,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c4_2(a565,a485)|c5_2(a565,a485)|c1_2(a565,X1)|c2_2(a565,X1)|c8_1(a565)|ssSkP1(a565)|c1_1(a565)|c1_1(a573)|c7_0|c5_1(a501)|ssSkC20|~c3_2(a565,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_415]), c_0_353]), ['final']).
cnf(c_0_1078,negated_conjecture,(ssSkC8|c6_2(a565,a485)|c4_2(a565,a485)|c5_2(a565,a485)|c1_2(a565,X1)|c2_2(a565,X1)|c8_1(a565)|ssSkP1(a565)|c1_1(a565)|c7_0|c3_1(a573)|c5_1(a501)|ssSkC20|~c3_2(a565,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_415]), c_0_353]), ['final']).
cnf(c_0_1079,negated_conjecture,(c6_2(a542,a515)|c1_2(a542,a557)|c2_2(a542,a503)|ssSkP4(a542)|ssSkP5(a542)|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_455, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1080,negated_conjecture,(c6_2(a542,a515)|c3_2(a542,a460)|c1_2(a542,a557)|c2_2(a542,a460)|ssSkP5(a542)|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_455, c_0_581]), ['final']).
cnf(c_0_1081,negated_conjecture,(c6_2(a542,a515)|c3_2(a542,a487)|c1_2(a542,a557)|c2_2(a542,a487)|ssSkP5(a542)|ssSkC5|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(spm,[status(thm)],[c_0_455, c_0_605]), ['final']).
cnf(c_0_1082,negated_conjecture,(c6_2(a512,a526)|c1_2(a512,a557)|c2_2(a512,a503)|ssSkP4(a512)|ssSkP6(a512)|c4_1(a531)|c5_1(a565)|c5_1(a512)|c7_1(a512)|ssSkC9|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_645, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1083,negated_conjecture,(c6_2(a512,a526)|c3_2(a512,a460)|c1_2(a512,a557)|c2_2(a512,a460)|ssSkP6(a512)|c4_1(a531)|c5_1(a565)|c5_1(a512)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_645, c_0_581]), ['final']).
cnf(c_0_1084,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c4_2(a517,a485)|c5_2(a517,a485)|c1_2(a517,X1)|c2_2(a517,X1)|c8_1(a517)|ssSkP1(a517)|c1_1(a517)|c1_1(a573)|c7_0|c5_1(a501)|ssSkC11|~c3_2(a517,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_410]), c_0_353]), ['final']).
cnf(c_0_1085,negated_conjecture,(ssSkC8|c6_2(a517,a485)|c4_2(a517,a485)|c5_2(a517,a485)|c1_2(a517,X1)|c2_2(a517,X1)|c8_1(a517)|ssSkP1(a517)|c1_1(a517)|c7_0|c3_1(a573)|c5_1(a501)|ssSkC11|~c3_2(a517,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_410]), c_0_353]), ['final']).
cnf(c_0_1086,negated_conjecture,(c4_2(a542,a502)|c1_2(X1,a456)|c8_1(a542)|ssSkP0(X1)|ssSkP3(a542)|c1_1(a542)|c6_1(X1)|c7_0|ndr1_1(X2)|c5_1(a501)|ssSkC16|~c3_2(a542,a502)|~c3_1(a542)|~c5_1(X2)|~c7_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_632, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1087,negated_conjecture,(c6_2(a512,a515)|c1_2(a512,a557)|c2_2(a512,a503)|ssSkP4(a512)|ssSkP5(a512)|c4_1(a531)|c5_1(a565)|c5_1(a512)|c7_1(a512)|ssSkC9|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_646, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1088,negated_conjecture,(c6_2(a512,a515)|c3_2(a512,a460)|c1_2(a512,a557)|c2_2(a512,a460)|ssSkP5(a512)|c4_1(a531)|c5_1(a565)|c5_1(a512)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_646, c_0_581]), ['final']).
cnf(c_0_1089,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a538)|c8_1(a592)|c4_1(a592)|c3_1(a538)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_493, c_0_280]), ['final']).
cnf(c_0_1090,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a494)|c8_1(a592)|c4_1(a592)|c3_1(a494)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_493, c_0_281]), ['final']).
cnf(c_0_1091,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a558)|c8_1(a592)|c4_1(a592)|c3_1(a558)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c7_0), inference(spm,[status(thm)],[c_0_493, c_0_282]), ['final']).
cnf(c_0_1092,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a592)|c4_1(a592)|c3_1(a466)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c7_0), inference(spm,[status(thm)],[c_0_647, c_0_596]), ['final']).
cnf(c_0_1093,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a538)|c8_1(a592)|c4_1(a592)|c3_1(a538)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_494, c_0_280]), ['final']).
cnf(c_0_1094,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a494)|c8_1(a592)|c4_1(a592)|c3_1(a494)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_494, c_0_281]), ['final']).
cnf(c_0_1095,negated_conjecture,(c6_2(a542,a487)|c3_2(a542,a487)|c1_2(a542,a557)|c2_2(a542,a503)|ssSkP4(a542)|ssSkC5|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_446, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1096,negated_conjecture,(c6_2(a542,a487)|c3_2(a542,a460)|c3_2(a542,a487)|c1_2(a542,a557)|c2_2(a542,a460)|ssSkC5|c4_1(a531)|c5_1(a565)|c5_1(a542)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_446, c_0_581]), ['final']).
cnf(c_0_1097,negated_conjecture,(c6_2(a546,a547)|c3_2(a546,a547)|c1_2(a546,a557)|c2_2(a546,a503)|ssSkP4(a546)|c4_1(a531)|c5_1(a565)|c5_1(a546)|c7_1(a546)|ssSkC17|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_648, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1098,negated_conjecture,(c6_2(a546,a547)|c3_2(a546,a460)|c3_2(a546,a547)|c1_2(a546,a557)|c2_2(a546,a460)|c4_1(a531)|c5_1(a565)|c5_1(a546)|c7_1(a546)|ssSkC17|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_648, c_0_581]), ['final']).
cnf(c_0_1099,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a558)|c8_1(a592)|c4_1(a592)|c3_1(a558)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c7_0), inference(spm,[status(thm)],[c_0_494, c_0_282]), ['final']).
cnf(c_0_1100,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c8_1(a592)|c4_1(a592)|c3_1(a466)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c7_0), inference(spm,[status(thm)],[c_0_649, c_0_596]), ['final']).
cnf(c_0_1101,negated_conjecture,(c6_2(a542,a485)|c4_2(a542,a485)|c5_2(a542,a485)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|ssSkP1(a542)|c1_1(a542)|c1_1(a573)|c7_0|c5_1(a501)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_412]), c_0_279]), ['final']).
cnf(c_0_1102,negated_conjecture,(c6_2(a542,a485)|c4_2(a542,a485)|c5_2(a542,a485)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|ssSkP1(a542)|c1_1(a542)|c7_0|c3_1(a573)|c5_1(a501)|ssSkC16|~c3_2(a542,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_412]), c_0_279]), ['final']).
cnf(c_0_1103,negated_conjecture,(c4_2(a542,a502)|c8_1(a542)|ssSkP3(a542)|c1_1(a542)|c6_1(a542)|c7_0|ndr1_1(X1)|c5_1(a501)|ssSkC16|~c3_2(a542,a502)|~c3_1(a542)|~c5_1(X1)|~c7_1(a542)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_650, c_0_628]), c_0_8]), c_0_124]), ['final']).
cnf(c_0_1104,negated_conjecture,(c7_2(X1,a474)|c1_2(X2,a456)|c8_1(X2)|ssSkP0(X2)|c6_1(X2)|~c3_2(X2,a456)|~c3_1(X2)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_639, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1105,negated_conjecture,(ssSkC8|c8_2(a592,a515)|c8_2(X1,a467)|c2_2(a592,a588)|c8_1(a592)|ssSkP5(a592)|c1_1(a573)|ssSkP8(a592)|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_651]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_1106,negated_conjecture,(ssSkC8|c8_2(a592,a515)|c3_2(X1,a467)|c2_2(a592,a588)|c8_1(a592)|ssSkP5(a592)|c1_1(a573)|ssSkP8(a592)|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_651]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_1107,negated_conjecture,(ssSkC8|c4_2(a559,a560)|c8_2(a565,X1)|c7_2(a565,X1)|c3_2(a565,a502)|c2_2(a565,X1)|ssSkP3(a565)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20|~c2_2(a565,a502)|~c8_1(a565)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_634]), c_0_353]), ['final']).
cnf(c_0_1108,negated_conjecture,(c7_2(X1,X2)|c5_2(X3,X4)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X5,X6)|c8_1(X5)|c8_1(X3)|c1_1(X1)|c2_1(X1)|~c6_2(X3,X4)|~c6_2(X5,X6)|~c4_2(X3,a465)|~c7_2(X3,X4)|~c2_2(X3,a516)|~c2_2(X5,X6)|~ssSkP5(X3)|~c1_1(X5)|~ndr1_1(X3)|~ndr1_1(X5)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), inference(spm,[status(thm)],[c_0_462, c_0_653]), ['final']).
cnf(c_0_1109,negated_conjecture,(c1_2(X1,a456)|c1_2(X2,a474)|c8_1(X1)|ssSkP0(X1)|c6_1(X1)|~c3_2(X1,a456)|~c3_1(X1)|~c5_1(X2)|~c7_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_640, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1110,negated_conjecture,(ssSkC8|c8_2(a592,a515)|c2_2(a592,a588)|c2_2(X1,a467)|c8_1(a592)|ssSkP5(a592)|c1_1(a573)|ssSkP8(a592)|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_651]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_1111,negated_conjecture,(ssSkC8|c8_2(a592,a515)|c8_2(X1,a467)|c2_2(a592,a588)|c8_1(a592)|ssSkP5(a592)|ssSkP8(a592)|c3_1(a573)|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_654]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_1112,negated_conjecture,(ssSkC8|c8_2(a592,a515)|c3_2(X1,a467)|c2_2(a592,a588)|c8_1(a592)|ssSkP5(a592)|ssSkP8(a592)|c3_1(a573)|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_654]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_1113,negated_conjecture,(ssSkC8|c8_2(a592,a515)|c2_2(a592,a588)|c2_2(X1,a467)|c8_1(a592)|ssSkP5(a592)|ssSkP8(a592)|c3_1(a573)|ssSkC24|~c8_1(X1)|~c2_1(a592)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_654]), c_0_353]), c_0_23]), ['final']).
cnf(c_0_1114,negated_conjecture,(ssSkC8|c8_2(a565,X1)|c7_2(a559,a560)|c7_2(a565,X1)|c3_2(a565,a502)|c2_2(a565,X1)|ssSkP3(a565)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20|~c2_2(a565,a502)|~c8_1(a565)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_634]), c_0_353]), ['final']).
cnf(c_0_1115,negated_conjecture,(ssSkC8|c4_2(a559,a560)|c8_2(a517,X1)|c7_2(a517,X1)|c3_2(a517,a502)|c2_2(a517,X1)|ssSkP3(a517)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11|~c2_2(a517,a502)|~c8_1(a517)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_635]), c_0_353]), ['final']).
cnf(c_0_1116,negated_conjecture,(ssSkC8|c8_2(a517,X1)|c7_2(a559,a560)|c7_2(a517,X1)|c3_2(a517,a502)|c2_2(a517,X1)|ssSkP3(a517)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11|~c2_2(a517,a502)|~c8_1(a517)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_635]), c_0_353]), ['final']).
cnf(c_0_1117,negated_conjecture,(c7_2(X1,X2)|c5_2(X3,X4)|c3_2(X1,X2)|c1_2(X1,X2)|c1_2(X5,X6)|c8_1(X5)|c8_1(X3)|c1_1(X1)|c2_1(X1)|~c6_2(X3,a516)|~c6_2(X3,X4)|~c6_2(X5,X6)|~c4_2(X3,a465)|~c7_2(X3,X4)|~c2_2(X5,X6)|~ssSkP5(X3)|~c1_1(X5)|~ndr1_1(X3)|~ndr1_1(X5)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), inference(spm,[status(thm)],[c_0_462, c_0_656]), ['final']).
cnf(c_0_1118,negated_conjecture,(ssSkC8|c6_2(a600,a460)|c7_2(a600,a557)|c3_2(a600,a460)|c2_2(a600,a460)|c4_1(a531)|c5_1(a600)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~ssSkC0), inference(spm,[status(thm)],[c_0_475, c_0_581]), ['final']).
cnf(c_0_1119,negated_conjecture,(ssSkC8|c8_2(a565,X1)|c7_2(a565,X1)|c3_2(a565,a502)|c2_2(a565,X1)|ssSkP3(a565)|c1_1(a565)|c7_0|c5_1(a501)|c7_1(a559)|ssSkC20|~c2_2(a565,a502)|~c8_1(a565)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_634]), c_0_353]), ['final']).
cnf(c_0_1120,negated_conjecture,(ssSkC8|c8_2(a517,X1)|c7_2(a517,X1)|c3_2(a517,a502)|c2_2(a517,X1)|ssSkP3(a517)|c1_1(a517)|c7_0|c5_1(a501)|c7_1(a559)|ssSkC11|~c2_2(a517,a502)|~c8_1(a517)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_635]), c_0_353]), ['final']).
cnf(c_0_1121,negated_conjecture,(c4_2(a559,a560)|c8_2(a542,X1)|c7_2(a542,X1)|c3_2(a542,a502)|c2_2(a542,X1)|ssSkP3(a542)|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c2_2(a542,a502)|~c8_1(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1122,negated_conjecture,(c8_2(a542,X1)|c7_2(a559,a560)|c7_2(a542,X1)|c3_2(a542,a502)|c2_2(a542,X1)|ssSkP3(a542)|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c2_2(a542,a502)|~c8_1(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1123,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c4_2(a592,a526)|c1_2(a592,a557)|c2_2(a592,a557)|c2_2(a592,a502)|c8_1(a592)|ssSkP3(a592)|ssSkP6(a592)|c1_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_658]), c_0_21]), ['final']).
cnf(c_0_1124,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c4_2(a592,a526)|c1_2(a592,a557)|c2_2(a592,a557)|c2_2(a592,a502)|c8_1(a592)|ssSkP3(a592)|ssSkP6(a592)|c3_1(a573)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_659]), c_0_21]), ['final']).
cnf(c_0_1125,negated_conjecture,(c6_2(X1,a474)|c4_2(a566,a567)|c1_2(X2,a456)|c8_1(a566)|ssSkP0(X2)|c6_1(X2)|~c3_2(a566,a567)|~c3_1(a566)|~c5_1(X1)|~c7_1(X2)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_623, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1126,negated_conjecture,(c6_2(X1,a474)|c4_2(a505,a506)|c1_2(X2,a456)|c8_1(a505)|ssSkP0(X2)|c6_1(X2)|~ssSkC8|~c3_2(a505,a506)|~c4_0|~c3_1(a505)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_623, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1127,negated_conjecture,(c6_2(X1,a474)|c1_2(X2,a456)|c8_1(a496)|ssSkP0(X2)|c6_1(X2)|~c4_0|~c3_1(a496)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_623, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1128,negated_conjecture,(ssSkC22|ssSkC8|c6_2(a552,a526)|c6_2(a552,a557)|c3_2(a552,a557)|c2_2(a552,a576)|c2_2(a552,a557)|ssSkP6(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c1_1(a552)), inference(spm,[status(thm)],[c_0_665, c_0_414]), ['final']).
cnf(c_0_1129,negated_conjecture,(ssSkC22|ssSkC8|c6_2(a552,a515)|c6_2(a552,a557)|c3_2(a552,a557)|c2_2(a552,a576)|c2_2(a552,a557)|ssSkP5(a552)|c4_1(a531)|c5_1(a552)|c5_1(a565)|c7_1(a552)|ssSkC18|ssSkC20|~c1_1(a552)), inference(spm,[status(thm)],[c_0_665, c_0_416]), ['final']).
cnf(c_0_1130,negated_conjecture,(c6_2(a520,a526)|c1_2(a520,a557)|c2_2(a520,a503)|ssSkP4(a520)|ssSkP6(a520)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_666, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1131,negated_conjecture,(c6_2(a520,a526)|c3_2(a520,a460)|c1_2(a520,a557)|c2_2(a520,a460)|ssSkP6(a520)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_666, c_0_581]), ['final']).
cnf(c_0_1132,negated_conjecture,(c6_2(a520,a515)|c1_2(a520,a557)|c2_2(a520,a503)|ssSkP4(a520)|ssSkP5(a520)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_667, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1133,negated_conjecture,(c6_2(a520,a515)|c3_2(a520,a460)|c1_2(a520,a557)|c2_2(a520,a460)|ssSkP5(a520)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_667, c_0_581]), ['final']).
cnf(c_0_1134,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a493)|c1_2(a542,a557)|c2_2(a542,X1)|ssSkC7|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c4_2(a542,X1)|~c2_2(a542,a516)|~ssSkP5(a542)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_653]), c_0_27]), ['final']).
cnf(c_0_1135,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a493)|c1_2(a542,a557)|c2_2(a542,X1)|ssSkC7|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c6_2(a542,a516)|~c4_2(a542,X1)|~ssSkP5(a542)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_656]), c_0_27]), ['final']).
cnf(c_0_1136,negated_conjecture,(ssSkC12|c4_2(a542,a527)|c3_2(a542,a588)|c3_2(a542,a493)|c1_2(a542,a557)|c2_2(a542,a588)|c2_2(a542,a493)|ssSkP8(a542)|ssSkC7|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkP6(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_453]), c_0_278]), ['final']).
cnf(c_0_1137,negated_conjecture,(c6_2(X1,a474)|c4_2(a566,a567)|c8_1(a566)|c6_1(a566)|~c3_2(a566,a567)|~c3_1(a566)|~c5_1(X1)|~c7_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_638, c_0_284]), c_0_160]), c_0_286]), ['final']).
cnf(c_0_1138,negated_conjecture,(c6_2(X1,a474)|c4_2(a505,a506)|c8_1(a505)|c6_1(a505)|~ssSkC8|~c3_2(a505,a506)|~c4_0|~c3_1(a505)|~c5_1(X1)|~c7_1(a505)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_638, c_0_285]), c_0_660]), c_0_287]), ['final']).
cnf(c_0_1139,negated_conjecture,(c6_2(X1,a474)|c1_2(a496,a497)|c8_1(a496)|c6_1(a496)|~c4_0|~c3_1(a496)|~c5_1(X1)|~c7_1(a496)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_638, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1140,negated_conjecture,(c1_2(X1,a456)|c8_1(X1)|ssSkP0(X1)|c6_1(X1)|ndr1_1(X2)|~c3_2(X1,a456)|~c3_1(X1)|~c5_1(X2)|~c7_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_650, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1141,negated_conjecture,(ssSkC8|c1_2(a470,X1)|c2_2(a470,a456)|c2_2(a470,X1)|ssSkP0(a470)|c1_1(X2)|c2_1(X2)|c2_0|ssSkC2|~c5_2(a470,X1)|~c1_2(a470,a456)|~c4_1(X2)|~c5_1(a470)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_668]), c_0_353]), ['final']).
cnf(c_0_1142,negated_conjecture,(c4_2(a566,a567)|c7_2(X1,a474)|c1_2(X2,a456)|c8_1(a566)|ssSkP0(X2)|c6_1(X2)|~c3_2(a566,a567)|~c3_1(a566)|~c5_1(X1)|~c7_1(X2)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1143,negated_conjecture,(c4_2(a505,a506)|c7_2(X1,a474)|c1_2(X2,a456)|c8_1(a505)|ssSkP0(X2)|c6_1(X2)|~ssSkC8|~c3_2(a505,a506)|~c4_0|~c3_1(a505)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1144,negated_conjecture,(c7_2(X1,a474)|c1_2(X2,a456)|c8_1(a496)|ssSkP0(X2)|c6_1(X2)|~c4_0|~c3_1(a496)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1145,negated_conjecture,(ssSkC8|c4_2(a512,a515)|c1_2(a512,a493)|c1_2(a512,a557)|c8_1(a512)|ssSkP5(a512)|ssSkC7|c3_1(a573)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkP2(a512)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_669]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1146,negated_conjecture,(ssSkC8|c4_2(a512,a515)|c1_2(a512,a493)|c1_2(a512,a557)|c8_1(a512)|ssSkP5(a512)|c1_1(a573)|ssSkC7|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkP2(a512)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_670]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1147,negated_conjecture,(ssSkC8|c1_2(a600,X1)|c2_2(a600,a503)|c2_2(a600,X1)|ssSkP4(a600)|c1_1(X2)|c2_1(X2)|c2_0|ssSkC26|~c5_2(a600,X1)|~c4_1(X2)|~c5_1(a600)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_479]), c_0_353]), c_0_62]), ['final']).
cnf(c_0_1148,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a493)|c1_2(X1,a557)|c2_2(X1,X2)|ssSkC7|c5_1(a565)|c7_1(X1)|ssSkC20|~c4_2(X1,X2)|~c2_2(X1,a516)|~ssSkP2(X1)|~ssSkP5(X1)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_653]), c_0_353]), ['final']).
cnf(c_0_1149,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a493)|c1_2(X1,a557)|c2_2(X1,X2)|ssSkC7|c5_1(a565)|c7_1(X1)|ssSkC20|~c6_2(X1,a516)|~c4_2(X1,X2)|~ssSkP2(X1)|~ssSkP5(X1)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_656]), c_0_353]), ['final']).
cnf(c_0_1150,negated_conjecture,(ssSkC8|c3_2(a537,X1)|c1_2(a537,a493)|c1_2(a537,a557)|c2_2(a537,a572)|c2_2(a537,X1)|ssSkC7|c5_1(a565)|c7_1(a537)|ssSkC15|ssSkC20|~ssSkP2(a537)), inference(spm,[status(thm)],[c_0_550, c_0_471]), ['final']).
cnf(c_0_1151,negated_conjecture,(ssSkC8|c3_2(X1,a588)|c1_2(X1,a493)|c1_2(X1,a557)|c2_2(X1,a588)|ssSkP8(X1)|ssSkC7|c5_1(a565)|c7_1(X1)|ssSkC20|~ssSkP2(X1)), inference(spm,[status(thm)],[c_0_550, c_0_56]), ['final']).
cnf(c_0_1152,negated_conjecture,(ssSkC8|c3_2(a517,X1)|c1_2(a517,a493)|c1_2(a517,a557)|c2_2(a517,a572)|c2_2(a517,X1)|ssSkC7|c5_1(a565)|c7_1(a517)|ssSkC11|ssSkC20|~ssSkP2(a517)), inference(spm,[status(thm)],[c_0_550, c_0_472]), ['final']).
cnf(c_0_1153,negated_conjecture,(ssSkC8|c3_2(a570,X1)|c1_2(a570,a493)|c1_2(a570,a557)|c2_2(a570,a572)|c2_2(a570,X1)|ssSkC7|c5_1(a565)|c7_1(a570)|ssSkC20|~ssSkP2(a570)), inference(spm,[status(thm)],[c_0_550, c_0_463]), ['final']).
cnf(c_0_1154,negated_conjecture,(ssSkC8|c1_2(a536,a493)|c1_2(a536,a557)|ssSkP5(a536)|ssSkC7|c5_1(a565)|ssSkC14|ssSkC20|~ssSkP2(a536)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_671]), c_0_288]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1155,negated_conjecture,(ssSkC8|c3_2(a482,a484)|c1_2(a482,a493)|c1_2(a482,a557)|c2_2(a482,a484)|ssSkC7|c5_1(a565)|c7_1(a482)|ssSkC20|~ssSkC4|~ssSkP2(a482)), inference(spm,[status(thm)],[c_0_550, c_0_672]), ['final']).
cnf(c_0_1156,negated_conjecture,(c4_2(a566,a567)|c1_2(X1,a474)|c1_2(X2,a456)|c8_1(a566)|ssSkP0(X2)|c6_1(X2)|~c3_2(a566,a567)|~c3_1(a566)|~c5_1(X1)|~c7_1(X2)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_630, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1157,negated_conjecture,(c4_2(a566,a567)|c7_2(X1,a474)|c8_1(a566)|c6_1(a566)|~c3_2(a566,a567)|~c3_1(a566)|~c5_1(X1)|~c7_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_639, c_0_284]), c_0_160]), c_0_286]), ['final']).
cnf(c_0_1158,negated_conjecture,(c4_2(a505,a506)|c7_2(X1,a474)|c8_1(a505)|c6_1(a505)|~ssSkC8|~c3_2(a505,a506)|~c4_0|~c3_1(a505)|~c5_1(X1)|~c7_1(a505)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_639, c_0_285]), c_0_660]), c_0_287]), ['final']).
cnf(c_0_1159,negated_conjecture,(c7_2(X1,a474)|c1_2(a496,a497)|c8_1(a496)|c6_1(a496)|~c4_0|~c3_1(a496)|~c5_1(X1)|~c7_1(a496)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_639, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1160,negated_conjecture,(c4_2(a505,a506)|c1_2(X1,a474)|c1_2(X2,a456)|c8_1(a505)|ssSkP0(X2)|c6_1(X2)|~ssSkC8|~c3_2(a505,a506)|~c4_0|~c3_1(a505)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_630, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1161,negated_conjecture,(c1_2(X1,a474)|c1_2(X2,a456)|c8_1(a496)|ssSkP0(X2)|c6_1(X2)|~c4_0|~c3_1(a496)|~c5_1(X1)|~c7_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_630, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1162,negated_conjecture,(c3_2(a520,X1)|c1_2(a520,a557)|c1_2(a520,X1)|c2_2(a520,X1)|c8_1(a538)|c8_1(a520)|c4_1(a520)|c3_1(a538)|c5_1(a565)|c7_1(a520)|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_434, c_0_280]), ['final']).
cnf(c_0_1163,negated_conjecture,(c3_2(a520,X1)|c1_2(a520,a557)|c1_2(a520,X1)|c2_2(a520,X1)|c8_1(a494)|c8_1(a520)|c4_1(a520)|c3_1(a494)|c5_1(a565)|c7_1(a520)|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_434, c_0_281]), ['final']).
cnf(c_0_1164,negated_conjecture,(c3_2(a520,X1)|c1_2(a520,a557)|c1_2(a520,X1)|c2_2(a520,X1)|c8_1(a558)|c8_1(a520)|c4_1(a520)|c3_1(a558)|c5_1(a565)|c7_1(a520)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_434, c_0_282]), ['final']).
cnf(c_0_1165,negated_conjecture,(c3_2(a520,X1)|c1_2(a520,a557)|c1_2(a520,X1)|c2_2(a520,X1)|c8_1(a520)|c4_1(a520)|c3_1(a466)|c5_1(a565)|c7_1(a520)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_673, c_0_596]), ['final']).
cnf(c_0_1166,negated_conjecture,(c3_2(a520,a557)|c1_2(a520,a557)|c2_2(a520,a557)|c8_1(a538)|c8_1(a520)|c4_1(a520)|c3_1(a538)|c5_1(a565)|c7_1(a520)|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_503, c_0_280]), ['final']).
cnf(c_0_1167,negated_conjecture,(c3_2(a520,a557)|c1_2(a520,a557)|c2_2(a520,a557)|c8_1(a494)|c8_1(a520)|c4_1(a520)|c3_1(a494)|c5_1(a565)|c7_1(a520)|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_503, c_0_281]), ['final']).
cnf(c_0_1168,negated_conjecture,(c3_2(a520,a557)|c1_2(a520,a557)|c2_2(a520,a557)|c8_1(a558)|c8_1(a520)|c4_1(a520)|c3_1(a558)|c5_1(a565)|c7_1(a520)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_503, c_0_282]), ['final']).
cnf(c_0_1169,negated_conjecture,(c3_2(a520,a557)|c1_2(a520,a557)|c2_2(a520,a557)|c8_1(a520)|c4_1(a520)|c3_1(a466)|c5_1(a565)|c7_1(a520)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_674, c_0_596]), ['final']).
cnf(c_0_1170,negated_conjecture,(c6_2(X1,a474)|c4_2(a566,a567)|c4_2(X2,X3)|c8_1(X2)|c6_1(a566)|~c6_2(X2,X3)|~c3_2(X2,X3)|~ndr1_1(X2)|~c3_1(X2)|~c5_1(X1)|~c7_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_284]), c_0_160]), c_0_286]), ['final']).
cnf(c_0_1171,negated_conjecture,(c6_2(X1,a474)|c4_2(a505,a506)|c4_2(X2,X3)|c8_1(X2)|c6_1(a505)|~ssSkC8|~c6_2(X2,X3)|~c3_2(X2,X3)|~c4_0|~ndr1_1(X2)|~c3_1(X2)|~c5_1(X1)|~c7_1(a505)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_285]), c_0_660]), c_0_287]), ['final']).
cnf(c_0_1172,negated_conjecture,(c4_2(a566,a567)|c4_2(X1,X2)|c7_2(X3,a474)|c8_1(X1)|c6_1(a566)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_284]), c_0_160]), c_0_286]), ['final']).
cnf(c_0_1173,negated_conjecture,(c4_2(a505,a506)|c4_2(X1,X2)|c7_2(X3,a474)|c8_1(X1)|c6_1(a505)|~ssSkC8|~c6_2(X1,X2)|~c3_2(X1,X2)|~c4_0|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a505)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_285]), c_0_660]), c_0_287]), ['final']).
cnf(c_0_1174,negated_conjecture,(ssSkC3|ssSkC8|c6_2(a592,a557)|c6_2(a592,a479)|c3_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c8_1(a592)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_392]), c_0_353]), ['final']).
cnf(c_0_1175,negated_conjecture,(ssSkC3|ssSkC8|c6_2(a592,a557)|c6_2(a592,a479)|c5_2(a592,a557)|c3_2(a592,X1)|c1_2(a592,X1)|c2_2(a592,X1)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c8_1(a592)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_375]), c_0_353]), ['final']).
cnf(c_0_1176,negated_conjecture,(c6_2(a546,a515)|c1_2(a546,a557)|ssSkP5(a546)|c4_1(a531)|c5_1(a546)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c1_2(a546,a515)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_675, c_0_23]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1177,negated_conjecture,(c6_2(a542,a515)|c1_2(a542,a557)|ssSkP5(a542)|c4_1(a531)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c1_2(a542,a515)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_676, c_0_23]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1178,negated_conjecture,(c6_2(a512,a515)|c1_2(a512,a557)|ssSkP5(a512)|c4_1(a531)|c5_1(a512)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c1_2(a512,a515)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_677, c_0_23]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1179,negated_conjecture,(c4_2(a566,a567)|c4_2(X1,X2)|c1_2(X3,a474)|c8_1(X1)|c6_1(a566)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_284]), c_0_160]), c_0_286]), ['final']).
cnf(c_0_1180,negated_conjecture,(c6_2(a542,a487)|c3_2(a542,a487)|c1_2(a542,a557)|c2_2(a542,a487)|ssSkC5|c4_1(a531)|c5_1(a542)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_676, c_0_409]), c_0_605]), ['final']).
cnf(c_0_1181,negated_conjecture,(c6_2(a546,a547)|c3_2(a546,a547)|c1_2(a546,a557)|c4_1(a531)|c5_1(a546)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c1_2(a546,a547)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_675, c_0_126]), c_0_289]), ['final']).
cnf(c_0_1182,negated_conjecture,(c8_2(a542,X1)|c7_2(a542,X1)|c3_2(a542,a502)|c2_2(a542,X1)|ssSkP3(a542)|c1_1(a542)|c7_0|ndr1_1(a559)|c5_1(a501)|ssSkC16|~c2_2(a542,a502)|~c8_1(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1183,negated_conjecture,(c8_2(a542,X1)|c7_2(a542,X1)|c3_2(a542,a502)|c2_2(a542,X1)|ssSkP3(a542)|c1_1(a542)|c7_0|c5_1(a501)|c7_1(a559)|ssSkC16|~c2_2(a542,a502)|~c8_1(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1184,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c3_2(a600,X1)|c1_2(a600,X1)|c2_2(a600,X1)|c8_1(a538)|c8_1(a600)|c4_1(a600)|c3_1(a538)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_506, c_0_280]), ['final']).
cnf(c_0_1185,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c3_2(a600,X1)|c1_2(a600,X1)|c2_2(a600,X1)|c8_1(a494)|c8_1(a600)|c4_1(a600)|c3_1(a494)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_506, c_0_281]), ['final']).
cnf(c_0_1186,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c3_2(a600,X1)|c1_2(a600,X1)|c2_2(a600,X1)|c8_1(a558)|c8_1(a600)|c4_1(a600)|c3_1(a558)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c7_0), inference(spm,[status(thm)],[c_0_506, c_0_282]), ['final']).
cnf(c_0_1187,negated_conjecture,(ssSkC3|ssSkC8|c6_2(a581,a479)|c3_2(a581,a557)|c1_2(a581,a557)|c2_2(a581,a557)|c5_1(a565)|c7_1(a581)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_679, c_0_596]), ['final']).
cnf(c_0_1188,negated_conjecture,(ssSkC3|ssSkC8|c6_2(a581,a479)|c3_2(a581,X1)|c1_2(a581,a557)|c1_2(a581,X1)|c2_2(a581,X1)|c5_1(a565)|c7_1(a581)|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_680, c_0_596]), ['final']).
cnf(c_0_1189,negated_conjecture,(c4_2(a505,a506)|c4_2(X1,X2)|c1_2(X3,a474)|c8_1(X1)|c6_1(a505)|~ssSkC8|~c6_2(X1,X2)|~c3_2(X1,X2)|~c4_0|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a505)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_285]), c_0_660]), c_0_287]), ['final']).
cnf(c_0_1190,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c3_2(a600,X1)|c1_2(a600,X1)|c2_2(a600,X1)|c8_1(a600)|c4_1(a600)|c3_1(a466)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c7_0), inference(spm,[status(thm)],[c_0_681, c_0_596]), ['final']).
cnf(c_0_1191,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c6_2(a592,X1)|~c4_2(a592,a548)|~c1_1(a592)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_392]), c_0_353]), ['final']).
cnf(c_0_1192,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c6_2(a592,X1)|~c7_2(a592,a548)|~c1_1(a592)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_392]), c_0_353]), ['final']).
cnf(c_0_1193,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c6_2(a592,X1)|~c4_2(a592,a548)|~c1_1(a592)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_375]), c_0_353]), ['final']).
cnf(c_0_1194,negated_conjecture,(c4_2(a566,a567)|c1_2(X1,a474)|c8_1(a566)|c6_1(a566)|~c3_2(a566,a567)|~c3_1(a566)|~c5_1(X1)|~c7_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_640, c_0_284]), c_0_160]), c_0_286]), ['final']).
cnf(c_0_1195,negated_conjecture,(c4_2(a505,a506)|c1_2(X1,a474)|c8_1(a505)|c6_1(a505)|~ssSkC8|~c3_2(a505,a506)|~c4_0|~c3_1(a505)|~c5_1(X1)|~c7_1(a505)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_640, c_0_285]), c_0_660]), c_0_287]), ['final']).
cnf(c_0_1196,negated_conjecture,(c1_2(a496,a497)|c1_2(X1,a474)|c8_1(a496)|c6_1(a496)|~c4_0|~c3_1(a496)|~c5_1(X1)|~c7_1(a496)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_640, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1197,negated_conjecture,(c6_2(a542,a588)|c6_2(a542,a543)|c3_2(a542,a588)|c3_2(a542,a543)|c1_1(a542)|ssSkP8(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c1_2(a542,a543)), inference(spm,[status(thm)],[c_0_682, c_0_608]), ['final']).
cnf(c_0_1198,negated_conjecture,(c6_2(a542,a485)|c6_2(a542,a543)|c3_2(a542,a543)|ssSkP1(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c1_2(a542,a543)), inference(spm,[status(thm)],[c_0_682, c_0_611]), ['final']).
cnf(c_0_1199,negated_conjecture,(c6_2(a542,a526)|c6_2(a542,a543)|c3_2(a542,a543)|ssSkP6(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c1_2(a542,a543)), inference(spm,[status(thm)],[c_0_682, c_0_614]), ['final']).
cnf(c_0_1200,negated_conjecture,(c4_2(a566,a567)|c1_2(X1,a456)|c8_1(a566)|ssSkP0(X1)|c6_1(X1)|ndr1_1(X2)|~c3_2(a566,a567)|~c3_1(a566)|~c5_1(X2)|~c7_1(X1)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_632, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1201,negated_conjecture,(c4_2(a505,a506)|c1_2(X1,a456)|c8_1(a505)|ssSkP0(X1)|c6_1(X1)|ndr1_1(X2)|~c3_2(a505,a506)|~c4_0|~c3_1(a505)|~c5_1(X2)|~c7_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_632, c_0_285]), c_0_353]), c_0_660]), ['final']).
cnf(c_0_1202,negated_conjecture,(c1_2(X1,a456)|c8_1(a496)|ssSkP0(X1)|c6_1(X1)|ndr1_1(X2)|~c4_0|~c3_1(a496)|~c5_1(X2)|~c7_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_632, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1203,negated_conjecture,(c6_2(a542,a515)|c6_2(a542,a543)|c3_2(a542,a543)|ssSkP5(a542)|c1_1(a542)|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c1_2(a542,a543)), inference(spm,[status(thm)],[c_0_682, c_0_616]), ['final']).
cnf(c_0_1204,negated_conjecture,(c6_2(a542,a487)|c6_2(a542,a543)|c3_2(a542,a487)|c3_2(a542,a543)|c1_1(a542)|ssSkC5|c7_0|c4_1(a531)|c5_1(a501)|c5_1(a542)|ssSkC16|~c1_2(a542,a543)), inference(spm,[status(thm)],[c_0_682, c_0_618]), ['final']).
cnf(c_0_1205,negated_conjecture,(ssSkC8|c6_2(a592,a557)|c5_2(a592,a557)|c3_2(a592,X1)|c2_2(a592,X1)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c6_2(a592,X1)|~c7_2(a592,a548)|~c1_1(a592)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_375]), c_0_353]), ['final']).
cnf(c_0_1206,negated_conjecture,(c4_2(a566,a567)|c4_2(X1,X2)|c8_1(X1)|c6_1(a566)|ndr1_1(X3)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_284]), c_0_160]), c_0_286]), ['final']).
cnf(c_0_1207,negated_conjecture,(c6_2(X1,a474)|c4_2(X2,X3)|c1_2(a496,a497)|c8_1(X2)|c6_1(a496)|~c6_2(X2,X3)|~c3_2(X2,X3)|~c4_0|~ndr1_1(X2)|~c3_1(X2)|~c5_1(X1)|~c7_1(a496)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_661]), c_0_662]), c_0_664]), ['final']).
cnf(c_0_1208,negated_conjecture,(c4_2(X1,X2)|c7_2(X3,a474)|c1_2(a496,a497)|c8_1(X1)|c6_1(a496)|~c6_2(X1,X2)|~c3_2(X1,X2)|~c4_0|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a496)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_661]), c_0_662]), c_0_664]), ['final']).
cnf(c_0_1209,negated_conjecture,(c4_2(X1,X2)|c1_2(a496,a497)|c1_2(X3,a474)|c8_1(X1)|c6_1(a496)|~c6_2(X1,X2)|~c3_2(X1,X2)|~c4_0|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a496)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_661]), c_0_662]), c_0_664]), ['final']).
cnf(c_0_1210,negated_conjecture,(c3_2(a546,a557)|c1_2(a546,a557)|c2_2(a546,a557)|c8_1(a538)|c8_1(a546)|c3_1(a538)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_507, c_0_280]), ['final']).
cnf(c_0_1211,negated_conjecture,(c3_2(a546,a557)|c1_2(a546,a557)|c2_2(a546,a557)|c8_1(a494)|c8_1(a546)|c3_1(a494)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_507, c_0_281]), ['final']).
cnf(c_0_1212,negated_conjecture,(c3_2(a546,a557)|c1_2(a546,a557)|c2_2(a546,a557)|c8_1(a558)|c8_1(a546)|c3_1(a558)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_507, c_0_282]), ['final']).
cnf(c_0_1213,negated_conjecture,(c3_2(a546,a557)|c1_2(a546,a557)|c2_2(a546,a557)|c8_1(a546)|c3_1(a466)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_683, c_0_596]), ['final']).
cnf(c_0_1214,negated_conjecture,(c4_2(a505,a506)|c4_2(X1,X2)|c1_2(a505,a506)|c8_1(X1)|c6_1(a505)|ndr1_1(X3)|~c6_2(X1,X2)|~c3_2(X1,X2)|~c4_0|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a505)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_285]), c_0_353]), c_0_660]), ['final']).
cnf(c_0_1215,negated_conjecture,(c3_2(a546,X1)|c1_2(a546,a557)|c1_2(a546,X1)|c2_2(a546,X1)|c8_1(a538)|c8_1(a546)|c3_1(a538)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_436, c_0_280]), ['final']).
cnf(c_0_1216,negated_conjecture,(c3_2(a546,X1)|c1_2(a546,a557)|c1_2(a546,X1)|c2_2(a546,X1)|c8_1(a494)|c8_1(a546)|c3_1(a494)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_436, c_0_281]), ['final']).
cnf(c_0_1217,negated_conjecture,(c3_2(a546,X1)|c1_2(a546,a557)|c1_2(a546,X1)|c2_2(a546,X1)|c8_1(a558)|c8_1(a546)|c3_1(a558)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_436, c_0_282]), ['final']).
cnf(c_0_1218,negated_conjecture,(c3_2(a546,X1)|c1_2(a546,a557)|c1_2(a546,X1)|c2_2(a546,X1)|c8_1(a546)|c3_1(a466)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_684, c_0_596]), ['final']).
cnf(c_0_1219,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a538)|c8_1(a542)|c4_1(a542)|c3_1(a538)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_437, c_0_280]), ['final']).
cnf(c_0_1220,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a494)|c8_1(a542)|c4_1(a542)|c3_1(a494)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_437, c_0_281]), ['final']).
cnf(c_0_1221,negated_conjecture,(c3_2(a542,a557)|c1_2(a542,a557)|c2_2(a542,a557)|c8_1(a538)|c8_1(a542)|c4_1(a542)|c3_1(a538)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_508, c_0_280]), ['final']).
cnf(c_0_1222,negated_conjecture,(c3_2(a542,a557)|c1_2(a542,a557)|c2_2(a542,a557)|c8_1(a494)|c8_1(a542)|c4_1(a542)|c3_1(a494)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_508, c_0_281]), ['final']).
cnf(c_0_1223,negated_conjecture,(c3_2(a542,a557)|c1_2(a542,a557)|c2_2(a542,a557)|c8_1(a558)|c8_1(a542)|c4_1(a542)|c3_1(a558)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_508, c_0_282]), ['final']).
cnf(c_0_1224,negated_conjecture,(c3_2(a542,a557)|c1_2(a542,a557)|c2_2(a542,a557)|c8_1(a542)|c4_1(a542)|c3_1(a466)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_685, c_0_596]), ['final']).
cnf(c_0_1225,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a558)|c8_1(a542)|c4_1(a542)|c3_1(a558)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_437, c_0_282]), ['final']).
cnf(c_0_1226,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X1)|c2_2(a542,X1)|c8_1(a542)|c4_1(a542)|c3_1(a466)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_686, c_0_596]), ['final']).
cnf(c_0_1227,negated_conjecture,(c4_2(a566,a567)|c8_1(a566)|c6_1(a566)|ndr1_1(X1)|~c3_2(a566,a567)|~c3_1(a566)|~c5_1(X1)|~c7_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_650, c_0_284]), c_0_160]), c_0_286]), ['final']).
cnf(c_0_1228,negated_conjecture,(c4_2(a505,a506)|c1_2(a505,a506)|c8_1(a505)|c6_1(a505)|ndr1_1(X1)|~c3_2(a505,a506)|~c4_0|~c3_1(a505)|~c5_1(X1)|~c7_1(a505)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_650, c_0_285]), c_0_353]), c_0_660]), ['final']).
cnf(c_0_1229,negated_conjecture,(c3_2(a512,X1)|c1_2(a512,a557)|c1_2(a512,X1)|c2_2(a512,X1)|c8_1(a538)|c8_1(a512)|c4_1(a512)|c3_1(a538)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_438, c_0_280]), ['final']).
cnf(c_0_1230,negated_conjecture,(c1_2(a496,a497)|c8_1(a496)|c6_1(a496)|ndr1_1(X1)|~c4_0|~c3_1(a496)|~c5_1(X1)|~c7_1(a496)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_650, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1231,negated_conjecture,(c3_2(a512,a557)|c1_2(a512,a557)|c2_2(a512,a557)|c8_1(a538)|c8_1(a512)|c4_1(a512)|c3_1(a538)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_509, c_0_280]), ['final']).
cnf(c_0_1232,negated_conjecture,(c3_2(a512,a557)|c1_2(a512,a557)|c2_2(a512,a557)|c8_1(a494)|c8_1(a512)|c4_1(a512)|c3_1(a494)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_509, c_0_281]), ['final']).
cnf(c_0_1233,negated_conjecture,(c3_2(a512,a557)|c1_2(a512,a557)|c2_2(a512,a557)|c8_1(a558)|c8_1(a512)|c4_1(a512)|c3_1(a558)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_509, c_0_282]), ['final']).
cnf(c_0_1234,negated_conjecture,(c3_2(a512,a557)|c1_2(a512,a557)|c2_2(a512,a557)|c8_1(a512)|c4_1(a512)|c3_1(a466)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_687, c_0_596]), ['final']).
cnf(c_0_1235,negated_conjecture,(c3_2(a512,X1)|c1_2(a512,a557)|c1_2(a512,X1)|c2_2(a512,X1)|c8_1(a494)|c8_1(a512)|c4_1(a512)|c3_1(a494)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_438, c_0_281]), ['final']).
cnf(c_0_1236,negated_conjecture,(c3_2(a512,X1)|c1_2(a512,a557)|c1_2(a512,X1)|c2_2(a512,X1)|c8_1(a558)|c8_1(a512)|c4_1(a512)|c3_1(a558)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_438, c_0_282]), ['final']).
cnf(c_0_1237,negated_conjecture,(c3_2(a512,X1)|c1_2(a512,a557)|c1_2(a512,X1)|c2_2(a512,X1)|c8_1(a512)|c4_1(a512)|c3_1(a466)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c7_0), inference(spm,[status(thm)],[c_0_688, c_0_596]), ['final']).
cnf(c_0_1238,negated_conjecture,(c6_2(a520,a526)|c1_2(a520,a557)|c2_2(a520,a526)|ssSkP6(a520)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20|~c1_2(a520,a526)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_689, c_0_20]), c_0_21]), ['final']).
cnf(c_0_1239,negated_conjecture,(c6_2(a520,a515)|c1_2(a520,a557)|ssSkP5(a520)|c4_1(a531)|c5_1(a565)|c7_1(a520)|ssSkC20|~c1_2(a520,a515)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_689, c_0_23]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1240,negated_conjecture,(ssSkC8|c4_2(X1,X2)|c5_2(a573,a456)|c5_2(X1,X2)|c1_2(X1,X3)|c2_2(X1,X3)|c8_1(X1)|ssSkP0(a573)|~c7_2(X1,X2)|~c3_2(X1,X3)|~c2_2(a573,a456)), inference(spm,[status(thm)],[c_0_488, c_0_121]), ['final']).
cnf(c_0_1241,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c3_2(a491,a583)|ssSkP3(a565)|c1_1(a565)|c7_0|c4_1(a565)|c5_1(a501)|c7_1(a565)|c7_1(a491)|ssSkC6|ssSkC20|~c2_2(a565,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_634]), c_0_125]), ['final']).
cnf(c_0_1242,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c3_2(a491,a583)|ssSkP3(a542)|c1_1(a542)|c7_0|c4_1(a542)|c5_1(a501)|c7_1(a542)|c7_1(a491)|ssSkC6|ssSkC16|~c2_2(a542,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_628]), c_0_125]), ['final']).
cnf(c_0_1243,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c7_2(X1,a456)|c3_2(a491,a583)|ssSkP0(X1)|c4_1(X1)|c7_1(a491)|c7_1(X1)|ssSkC6|~c2_2(X1,a456)), inference(spm,[status(thm)],[c_0_690, c_0_121]), ['final']).
cnf(c_0_1244,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c7_2(a592,a503)|c3_2(a491,a583)|ssSkP4(a592)|c4_1(a592)|c7_1(a592)|c7_1(a491)|ssSkC6|ssSkC24|~c2_2(a592,a503)), inference(spm,[status(thm)],[c_0_690, c_0_691]), ['final']).
cnf(c_0_1245,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c7_2(a568,a503)|c3_2(a491,a583)|ssSkP4(a568)|c7_1(a568)|c7_1(a491)|ssSkC6|ssSkC21|~c2_2(a568,a503)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_692]), c_0_290]), ['final']).
cnf(c_0_1246,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c7_2(a566,a567)|c3_2(a491,a583)|c4_1(a566)|c7_1(a566)|c7_1(a491)|ssSkC6|~c2_2(a566,a567)|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_690, c_0_284]), ['final']).
cnf(c_0_1247,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c7_2(a491,a503)|c3_2(a491,a583)|ssSkP4(a491)|c7_1(a491)|ssSkC6|~c2_2(a491,a503)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_693]), c_0_291]), ['final']).
cnf(c_0_1248,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c7_2(a549,a550)|c3_2(a491,a583)|c4_1(a549)|c7_1(a549)|c7_1(a491)|ssSkC6|~c2_2(a549,a550)|~c4_0), inference(spm,[status(thm)],[c_0_690, c_0_694]), ['final']).
cnf(c_0_1249,negated_conjecture,(ssSkC8|c6_2(a491,a583)|c7_2(a496,a497)|c3_2(a491,a583)|c4_1(a496)|c7_1(a496)|c7_1(a491)|ssSkC6|~c2_2(a496,a497)|~c4_0), inference(spm,[status(thm)],[c_0_690, c_0_661]), ['final']).
cnf(c_0_1250,negated_conjecture,(ssSkC8|c7_2(X1,a583)|ssSkP3(a565)|c1_1(a565)|c7_0|c4_1(a565)|c5_1(a501)|c5_1(X1)|c7_1(a565)|c7_1(X1)|ssSkC20|~c2_2(a565,a502)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_634]), c_0_353]), c_0_125]), ['final']).
cnf(c_0_1251,negated_conjecture,(ssSkC8|c2_2(X1,a583)|ssSkP3(a565)|c1_1(a565)|c7_0|c4_1(a565)|c5_1(a501)|c5_1(X1)|c7_1(a565)|c7_1(X1)|ssSkC20|~c2_2(a565,a502)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_634]), c_0_353]), c_0_125]), ['final']).
cnf(c_0_1252,negated_conjecture,(ssSkC8|c6_2(a565,a462)|c6_2(a565,X1)|c7_2(a565,a462)|ssSkC1|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20|~c4_2(a565,X1)|~c2_2(a565,X1)|~c8_1(a565)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_697, c_0_368]), c_0_353]), ['final']).
cnf(c_0_1253,negated_conjecture,(ssSkC8|c6_2(a552,a503)|c6_2(a552,a526)|c2_2(a552,a503)|ssSkP4(a552)|ssSkP6(a552)|c4_1(a531)|c5_1(a552)|ssSkC18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1254,negated_conjecture,(ssSkC8|c6_2(a552,a460)|c6_2(a552,a526)|c3_2(a552,a460)|c2_2(a552,a460)|ssSkP6(a552)|c4_1(a531)|c5_1(a552)|ssSkC18|~ssSkC0), inference(spm,[status(thm)],[c_0_389, c_0_581]), ['final']).
cnf(c_0_1255,negated_conjecture,(ssSkC8|c6_2(a552,a503)|c6_2(a552,a515)|c2_2(a552,a503)|ssSkP4(a552)|ssSkP5(a552)|c4_1(a531)|c5_1(a552)|ssSkC18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_62]), c_0_63]), ['final']).
cnf(c_0_1256,negated_conjecture,(ssSkC8|c6_2(a552,a460)|c6_2(a552,a515)|c3_2(a552,a460)|c2_2(a552,a460)|ssSkP5(a552)|c4_1(a531)|c5_1(a552)|ssSkC18|~ssSkC0), inference(spm,[status(thm)],[c_0_390, c_0_581]), ['final']).
cnf(c_0_1257,negated_conjecture,(c1_2(X1,X2)|c2_2(X1,a526)|c2_2(X1,X2)|ssSkP6(X1)|c1_1(X3)|c2_1(X3)|c2_0|~c5_2(X1,X2)|~c1_2(X1,a526)|~c4_1(X3)|~c5_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_20]), c_0_31]), ['final']).
cnf(c_0_1258,negated_conjecture,(ssSkC8|c6_2(a517,a462)|c6_2(a517,X1)|c7_2(a517,a462)|ssSkC1|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11|~c4_2(a517,X1)|~c2_2(a517,X1)|~c8_1(a517)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_697, c_0_382]), c_0_353]), ['final']).
cnf(c_0_1259,negated_conjecture,(c4_2(X1,X2)|c1_2(a496,a497)|c8_1(X1)|c6_1(a496)|ndr1_1(X3)|~c6_2(X1,X2)|~c3_2(X1,X2)|~c4_0|~ndr1_1(X1)|~c3_1(X1)|~c5_1(X3)|~c7_1(a496)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_661]), c_0_662]), c_0_664]), ['final']).
cnf(c_0_1260,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c3_2(a600,X1)|c1_2(a600,a493)|c2_2(a600,X1)|ssSkC7|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c5_2(a600,X1)|~ssSkP2(a600)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_394]), c_0_353]), ['final']).
cnf(c_0_1261,negated_conjecture,(ssSkC12|ssSkC8|c6_2(a565,a527)|c7_2(a565,a527)|c3_2(a565,X1)|c2_2(a565,X1)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20|~c1_2(a565,X1)|~ssSkP6(a565)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_698, c_0_368]), c_0_353]), ['final']).
cnf(c_0_1262,negated_conjecture,(ssSkC12|ssSkC8|c6_2(a517,a527)|c7_2(a517,a527)|c3_2(a517,X1)|c2_2(a517,X1)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11|~c1_2(a517,X1)|~ssSkP6(a517)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_698, c_0_382]), c_0_353]), ['final']).
cnf(c_0_1263,negated_conjecture,(c1_2(X1,X2)|c2_2(X1,X2)|ssSkP5(X1)|c1_1(X3)|c2_1(X3)|c2_0|~c5_2(X1,X2)|~c1_2(X1,a515)|~c4_1(X3)|~c5_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_23]), c_0_48]), c_0_277]), ['final']).
cnf(c_0_1264,negated_conjecture,(ssSkC8|c1_2(a587,a557)|c1_2(a587,X1)|c2_2(a587,X1)|c5_1(a565)|c7_1(a587)|ssSkC20|ssSkC23|~c4_2(a587,X1)|~c2_2(a587,a516)|~ssSkP5(a587)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_522, c_0_653]), c_0_353]), ['final']).
cnf(c_0_1265,negated_conjecture,(ssSkC8|c1_2(a587,a557)|c1_2(a587,X1)|c2_2(a587,X1)|c5_1(a565)|c7_1(a587)|ssSkC20|ssSkC23|~c6_2(a587,a516)|~c4_2(a587,X1)|~ssSkP5(a587)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_522, c_0_656]), c_0_353]), ['final']).
cnf(c_0_1266,negated_conjecture,(c7_2(X1,a583)|ssSkP3(a542)|c1_1(a542)|c7_0|c4_1(a542)|c5_1(a501)|c5_1(X1)|c7_1(a542)|c7_1(X1)|ssSkC16|~c2_2(a542,a502)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_628]), c_0_8]), c_0_125]), ['final']).
cnf(c_0_1267,negated_conjecture,(c2_2(X1,a583)|ssSkP3(a542)|c1_1(a542)|c7_0|c4_1(a542)|c5_1(a501)|c5_1(X1)|c7_1(a542)|c7_1(X1)|ssSkC16|~c2_2(a542,a502)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_628]), c_0_8]), c_0_125]), ['final']).
cnf(c_0_1268,negated_conjecture,(c6_2(a542,a462)|c6_2(a542,X1)|c7_2(a542,a462)|ssSkC1|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c4_2(a542,X1)|~c2_2(a542,X1)|~c8_1(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_697, c_0_384]), c_0_27]), ['final']).
cnf(c_0_1269,negated_conjecture,(ssSkC3|ssSkC8|c6_2(a600,a479)|c7_2(a600,a557)|c3_2(a600,X1)|c1_2(a600,X1)|c2_2(a600,X1)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c8_1(a600)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_394]), c_0_353]), ['final']).
cnf(c_0_1270,negated_conjecture,(ssSkC12|c6_2(a542,a527)|c7_2(a542,a527)|c3_2(a542,X1)|c2_2(a542,X1)|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c1_2(a542,X1)|~ssSkP6(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_698, c_0_384]), c_0_27]), ['final']).
cnf(c_0_1271,negated_conjecture,(c3_2(a546,X1)|c1_2(a546,a557)|c1_2(a546,a493)|c2_2(a546,X1)|ssSkC7|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c5_2(a546,X1)|~ssSkP2(a546)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_403]), c_0_45]), ['final']).
cnf(c_0_1272,negated_conjecture,(c3_2(a542,a460)|c1_2(a542,a493)|c1_2(a542,a557)|c2_2(a542,a460)|ssSkC7|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_413, c_0_699]), ['final']).
cnf(c_0_1273,negated_conjecture,(c1_2(a546,X1)|c2_2(a546,X1)|c1_1(X2)|c2_1(X2)|c2_0|ssSkC17|~c5_2(a546,X1)|~c1_2(a546,a547)|~c4_1(X2)|~c5_1(a546)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_126]), c_0_45]), c_0_289]), ['final']).
cnf(c_0_1274,negated_conjecture,(c3_2(a512,X1)|c1_2(a512,a557)|c1_2(a512,a493)|c2_2(a512,X1)|ssSkC7|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c5_2(a512,X1)|~ssSkP2(a512)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_379]), c_0_32]), ['final']).
cnf(c_0_1275,negated_conjecture,(ssSkC12|ssSkC8|c6_2(a565,X1)|c4_2(a565,a527)|c7_2(a565,X1)|c3_2(a565,X1)|c2_2(a565,X1)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20|~ssSkP6(a565)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_368]), c_0_353]), ['final']).
cnf(c_0_1276,negated_conjecture,(ssSkC12|ssSkC8|c6_2(a517,X1)|c4_2(a517,a527)|c7_2(a517,X1)|c3_2(a517,X1)|c2_2(a517,X1)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11|~ssSkP6(a517)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_382]), c_0_353]), ['final']).
cnf(c_0_1277,negated_conjecture,(ssSkC8|c6_2(a491,X1)|c8_2(a491,X1)|c8_2(X2,a467)|c3_2(a491,X1)|c8_1(a491)|ssSkC6|~c7_2(a491,X1)|~c8_1(X2)|~c2_1(a491)|~c6_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_511]), c_0_353]), ['final']).
cnf(c_0_1278,negated_conjecture,(ssSkC8|c4_2(a559,a560)|c8_2(a592,X1)|c7_2(a592,X1)|c2_2(a592,X1)|ssSkP4(a592)|ssSkC24|~c2_2(a592,a503)|~c8_1(a592)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_691]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1279,negated_conjecture,(ssSkC8|c8_2(a592,X1)|c7_2(a559,a560)|c7_2(a592,X1)|c2_2(a592,X1)|ssSkP4(a592)|ssSkC24|~c2_2(a592,a503)|~c8_1(a592)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_691]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1280,negated_conjecture,(ssSkC8|c4_2(a559,a560)|c8_2(a491,X1)|c7_2(a491,X1)|c2_2(a491,X1)|ssSkP4(a491)|ssSkC6|~c2_2(a491,a503)|~c8_1(a491)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_693]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1281,negated_conjecture,(ssSkC8|c8_2(a491,X1)|c7_2(a559,a560)|c7_2(a491,X1)|c2_2(a491,X1)|ssSkP4(a491)|ssSkC6|~c2_2(a491,a503)|~c8_1(a491)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_693]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1282,negated_conjecture,(ssSkC8|c4_2(a559,a560)|c8_2(a568,X1)|c7_2(a568,X1)|c2_2(a568,X1)|ssSkP4(a568)|ssSkC21|~c2_2(a568,a503)|~c8_1(a568)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_692]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1283,negated_conjecture,(ssSkC8|c8_2(a568,X1)|c7_2(a559,a560)|c7_2(a568,X1)|c2_2(a568,X1)|ssSkP4(a568)|ssSkC21|~c2_2(a568,a503)|~c8_1(a568)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_692]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1284,negated_conjecture,(c4_2(a534,a535)|c7_2(X1,X2)|c1_2(X1,X2)|ssSkP0(X1)|c6_1(X1)|~c3_2(X1,a456)|~c1_2(X1,a456)|~c2_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_700, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1285,negated_conjecture,(ssSkC8|c6_2(a491,X1)|c8_2(a491,X1)|c3_2(a491,X1)|c3_2(X2,a467)|c8_1(a491)|ssSkC6|~c7_2(a491,X1)|~c8_1(X2)|~c2_1(a491)|~c6_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_511]), c_0_353]), ['final']).
cnf(c_0_1286,negated_conjecture,(ssSkC8|c6_2(a491,X1)|c8_2(a491,X1)|c3_2(a491,X1)|c2_2(X2,a467)|c8_1(a491)|ssSkC6|~c7_2(a491,X1)|~c8_1(X2)|~c2_1(a491)|~c6_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_511]), c_0_353]), ['final']).
cnf(c_0_1287,negated_conjecture,(c8_2(a534,a535)|c7_2(X1,X2)|c1_2(X1,X2)|ssSkP0(X1)|c6_1(X1)|~c3_2(X1,a456)|~c1_2(X1,a456)|~c2_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_701, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1288,negated_conjecture,(c1_2(a498,X1)|c2_2(a498,a499)|c2_2(a498,X1)|c1_1(X2)|c2_1(X2)|c2_0|~c5_2(a498,X1)|~c1_2(a498,a499)|~c4_1(X2)|~c5_1(a498)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_524]), c_0_525]), ['final']).
cnf(c_0_1289,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c3_2(a600,X1)|c2_2(a600,X1)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c6_2(a600,X1)|~c4_2(a600,a548)|~c1_1(a600)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_394]), c_0_353]), ['final']).
cnf(c_0_1290,negated_conjecture,(ssSkC8|c7_2(a600,a557)|c3_2(a600,X1)|c2_2(a600,X1)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c6_2(a600,X1)|~c7_2(a600,a548)|~c1_1(a600)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_394]), c_0_353]), ['final']).
cnf(c_0_1291,negated_conjecture,(ssSkC12|c6_2(a542,X1)|c4_2(a542,a527)|c7_2(a542,X1)|c3_2(a542,X1)|c2_2(a542,X1)|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16|~ssSkP6(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_384]), c_0_27]), ['final']).
cnf(c_0_1292,negated_conjecture,(ssSkP3(a542)|c1_1(a542)|c7_0|c4_1(a542)|ndr1_1(X1)|c5_1(a501)|c5_1(X1)|c7_1(a542)|c7_1(X1)|ssSkC16|~c2_2(a542,a502)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_628]), c_0_8]), c_0_125]), ['final']).
cnf(c_0_1293,negated_conjecture,(ssSkC3|c6_2(a546,a479)|c3_2(a546,X1)|c1_2(a546,a557)|c1_2(a546,X1)|c2_2(a546,X1)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~c8_1(a546)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_403]), c_0_45]), ['final']).
cnf(c_0_1294,negated_conjecture,(ssSkC8|c6_2(a568,a583)|ssSkP3(a565)|c1_1(a565)|c7_0|c4_1(a565)|c5_1(a501)|c5_1(a568)|c7_1(a565)|c7_1(a568)|ssSkC20|ssSkC21|~c2_2(a565,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_703, c_0_634]), c_0_125]), ['final']).
cnf(c_0_1295,negated_conjecture,(ssSkC8|c6_2(a568,a583)|ssSkP3(a542)|c1_1(a542)|c7_0|c4_1(a542)|c5_1(a501)|c5_1(a568)|c7_1(a542)|c7_1(a568)|ssSkC16|ssSkC21|~c2_2(a542,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_703, c_0_628]), c_0_125]), ['final']).
cnf(c_0_1296,negated_conjecture,(ssSkC8|c6_2(a568,a583)|c7_2(a592,a503)|ssSkP4(a592)|c4_1(a592)|c5_1(a568)|c7_1(a592)|c7_1(a568)|ssSkC21|ssSkC24|~c2_2(a592,a503)), inference(spm,[status(thm)],[c_0_703, c_0_691]), ['final']).
cnf(c_0_1297,negated_conjecture,(ssSkC8|c6_2(a568,a583)|c7_2(a491,a503)|ssSkP4(a491)|c5_1(a568)|c7_1(a491)|c7_1(a568)|ssSkC6|ssSkC21|~c2_2(a491,a503)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_703, c_0_693]), c_0_291]), ['final']).
cnf(c_0_1298,negated_conjecture,(c4_2(a559,a560)|c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,a456)|c2_2(X1,X2)|ssSkP0(X1)|~c2_2(X1,a456)|~c8_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1299,negated_conjecture,(ssSkC8|c6_2(a568,a583)|c7_2(X1,a456)|ssSkP0(X1)|c4_1(X1)|c5_1(a568)|c7_1(a568)|c7_1(X1)|ssSkC21|~c2_2(X1,a456)), inference(spm,[status(thm)],[c_0_703, c_0_121]), ['final']).
cnf(c_0_1300,negated_conjecture,(ssSkC8|c6_2(a568,a583)|c7_2(a568,a503)|ssSkP4(a568)|c5_1(a568)|c7_1(a568)|ssSkC21|~c2_2(a568,a503)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_703, c_0_692]), c_0_290]), ['final']).
cnf(c_0_1301,negated_conjecture,(ssSkC8|c6_2(a568,a583)|c7_2(a566,a567)|c4_1(a566)|c5_1(a568)|c7_1(a566)|c7_1(a568)|ssSkC21|~c2_2(a566,a567)|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_703, c_0_284]), ['final']).
cnf(c_0_1302,negated_conjecture,(ssSkC8|c6_2(a568,a583)|c7_2(a549,a550)|c4_1(a549)|c5_1(a568)|c7_1(a549)|c7_1(a568)|ssSkC21|~c2_2(a549,a550)|~c4_0), inference(spm,[status(thm)],[c_0_703, c_0_694]), ['final']).
cnf(c_0_1303,negated_conjecture,(ssSkC8|c6_2(a568,a583)|c7_2(a496,a497)|c4_1(a496)|c5_1(a568)|c7_1(a496)|c7_1(a568)|ssSkC21|~c2_2(a496,a497)|~c4_0), inference(spm,[status(thm)],[c_0_703, c_0_661]), ['final']).
cnf(c_0_1304,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|c8_1(a534)|ssSkP0(X1)|c6_1(X1)|~c3_2(X1,a456)|~c1_2(X1,a456)|~c2_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_704, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1305,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|ssSkP0(X1)|c6_1(X1)|ndr1_1(a534)|~c3_2(X1,a456)|~c1_2(X1,a456)|~c2_2(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_705, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1306,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c6_2(a542,X1)|~c4_2(a542,a548)|~c1_1(a542)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_383]), c_0_27]), ['final']).
cnf(c_0_1307,negated_conjecture,(ssSkC3|c6_2(a542,a479)|c3_2(a542,X1)|c1_2(a542,a557)|c1_2(a542,X1)|c2_2(a542,X1)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c8_1(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_383]), c_0_27]), ['final']).
cnf(c_0_1308,negated_conjecture,(ssSkC3|c6_2(a512,a479)|c3_2(a512,X1)|c1_2(a512,a557)|c1_2(a512,X1)|c2_2(a512,X1)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c8_1(a512)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_379]), c_0_32]), ['final']).
cnf(c_0_1309,negated_conjecture,(ssSkC8|c8_2(a592,X1)|c7_2(a592,X1)|c2_2(a592,X1)|ssSkP4(a592)|c7_1(a559)|ssSkC24|~c2_2(a592,a503)|~c8_1(a592)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_691]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1310,negated_conjecture,(ssSkC8|c8_2(a491,X1)|c7_2(a491,X1)|c2_2(a491,X1)|ssSkP4(a491)|c7_1(a559)|ssSkC6|~c2_2(a491,a503)|~c8_1(a491)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_693]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1311,negated_conjecture,(ssSkC8|c8_2(a568,X1)|c7_2(a568,X1)|c2_2(a568,X1)|ssSkP4(a568)|c7_1(a559)|ssSkC21|~c2_2(a568,a503)|~c8_1(a568)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_692]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1312,negated_conjecture,(c8_2(X1,X2)|c7_2(a559,a560)|c7_2(X1,X2)|c3_2(X1,a456)|c2_2(X1,X2)|ssSkP0(X1)|~c2_2(X1,a456)|~c8_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1313,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X2)|ssSkP6(X1)|c4_1(a531)|c5_1(X1)|~c7_2(X1,X2)|~c1_2(X1,a526)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_98]), c_0_31]), c_0_21]), ['final']).
cnf(c_0_1314,negated_conjecture,(c3_2(a542,X1)|c1_2(a542,a557)|c2_2(a542,X1)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~c6_2(a542,X1)|~c7_2(a542,a548)|~c1_1(a542)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_383]), c_0_27]), ['final']).
cnf(c_0_1315,negated_conjecture,(c3_2(a512,X1)|c1_2(a512,a557)|c2_2(a512,X1)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c6_2(a512,X1)|~c4_2(a512,a548)|~c1_1(a512)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_379]), c_0_32]), ['final']).
cnf(c_0_1316,negated_conjecture,(c3_2(a512,X1)|c1_2(a512,a557)|c2_2(a512,X1)|c5_1(a565)|c7_1(a512)|ssSkC9|ssSkC20|~c6_2(a512,X1)|~c7_2(a512,a548)|~c1_1(a512)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_379]), c_0_32]), ['final']).
cnf(c_0_1317,negated_conjecture,(ssSkC8|c5_2(a470,a456)|c1_2(a470,X1)|c2_2(a470,X1)|c8_1(a470)|ssSkP0(a470)|c1_1(a573)|ssSkC2|~c3_2(a470,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_668]), c_0_353]), c_0_122]), ['final']).
cnf(c_0_1318,negated_conjecture,(ssSkC8|c5_2(a470,a456)|c1_2(a470,X1)|c2_2(a470,X1)|c8_1(a470)|ssSkP0(a470)|c3_1(a573)|ssSkC2|~c3_2(a470,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_668]), c_0_353]), c_0_122]), ['final']).
cnf(c_0_1319,negated_conjecture,(ssSkC8|c5_2(a600,a503)|c1_2(a600,X1)|c2_2(a600,X1)|c8_1(a600)|ssSkP4(a600)|c1_1(a573)|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_479]), c_0_353]), c_0_140]), ['final']).
cnf(c_0_1320,negated_conjecture,(ssSkC8|c5_2(a600,a503)|c1_2(a600,X1)|c2_2(a600,X1)|c8_1(a600)|ssSkP4(a600)|c3_1(a573)|ssSkC26|~c3_2(a600,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_479]), c_0_353]), c_0_140]), ['final']).
cnf(c_0_1321,negated_conjecture,(ssSkC22|c8_2(a532,a576)|c8_2(X1,a467)|c8_1(a532)|c7_1(a532)|ssSkC13|~c7_2(a532,a576)|~c8_1(X1)|~c2_1(a532)|~c6_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_706]), c_0_707]), ['final']).
cnf(c_0_1322,negated_conjecture,(ssSkC22|c8_2(a532,a576)|c3_2(X1,a467)|c8_1(a532)|c7_1(a532)|ssSkC13|~c7_2(a532,a576)|~c8_1(X1)|~c2_1(a532)|~c6_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_706]), c_0_707]), ['final']).
cnf(c_0_1323,negated_conjecture,(ssSkC22|c8_2(a532,a576)|c2_2(X1,a467)|c8_1(a532)|c7_1(a532)|ssSkC13|~c7_2(a532,a576)|~c8_1(X1)|~c2_1(a532)|~c6_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_706]), c_0_707]), ['final']).
cnf(c_0_1324,negated_conjecture,(c8_2(a565,a487)|c8_2(X1,a467)|c8_1(a565)|ssSkC5|c5_1(a565)|ssSkC20|~c7_2(a565,a487)|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_708]), c_0_534]), c_0_16]), ['final']).
cnf(c_0_1325,negated_conjecture,(c8_2(a565,a487)|c3_2(X1,a467)|c8_1(a565)|ssSkC5|c5_1(a565)|ssSkC20|~c7_2(a565,a487)|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_708]), c_0_534]), c_0_16]), ['final']).
cnf(c_0_1326,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,a456)|c3_2(X1,X2)|ssSkP0(X1)|c4_1(a531)|c5_1(X1)|~c7_2(X1,X2)|~c1_2(X1,a456)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_99]), c_0_100]), ['final']).
cnf(c_0_1327,negated_conjecture,(c8_2(a565,a487)|c2_2(X1,a467)|c8_1(a565)|ssSkC5|c5_1(a565)|ssSkC20|~c7_2(a565,a487)|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_708]), c_0_534]), c_0_16]), ['final']).
cnf(c_0_1328,negated_conjecture,(c8_2(a542,a487)|c8_2(X1,a467)|c8_1(a542)|ssSkC5|c5_1(a542)|ssSkC16|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_386]), c_0_534]), c_0_34]), c_0_409]), ['final']).
cnf(c_0_1329,negated_conjecture,(c8_2(a542,a487)|c3_2(X1,a467)|c8_1(a542)|ssSkC5|c5_1(a542)|ssSkC16|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_386]), c_0_534]), c_0_34]), c_0_409]), ['final']).
cnf(c_0_1330,negated_conjecture,(c8_2(a542,a487)|c2_2(X1,a467)|c8_1(a542)|ssSkC5|c5_1(a542)|ssSkC16|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_386]), c_0_534]), c_0_34]), c_0_409]), ['final']).
cnf(c_0_1331,negated_conjecture,(c4_2(X1,X2)|c5_2(a573,X3)|c5_2(X1,X2)|c1_2(X1,X4)|c2_2(X1,X4)|c8_1(X1)|~ssSkC0|~c6_2(a573,X3)|~c7_2(X1,X2)|~c3_2(X1,X4)|~c2_2(a573,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_541]), c_0_541]), ['final']).
cnf(c_0_1332,negated_conjecture,(c8_2(a517,a487)|c8_2(X1,a467)|c8_1(a517)|ssSkC5|c5_1(a517)|ssSkC11|~c7_2(a517,a487)|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_709]), c_0_534]), c_0_22]), ['final']).
cnf(c_0_1333,negated_conjecture,(c8_2(a517,a487)|c3_2(X1,a467)|c8_1(a517)|ssSkC5|c5_1(a517)|ssSkC11|~c7_2(a517,a487)|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_709]), c_0_534]), c_0_22]), ['final']).
cnf(c_0_1334,negated_conjecture,(c8_2(a517,a487)|c2_2(X1,a467)|c8_1(a517)|ssSkC5|c5_1(a517)|ssSkC11|~c7_2(a517,a487)|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_709]), c_0_534]), c_0_22]), ['final']).
cnf(c_0_1335,negated_conjecture,(ssSkC12|c4_2(a546,a527)|c8_2(a546,X1)|c3_2(a546,a557)|c2_2(a546,a557)|c2_2(a546,X1)|c5_1(a565)|c7_1(a546)|ssSkC17|ssSkC20|~ssSkP6(a546)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_403]), c_0_45]), ['final']).
cnf(c_0_1336,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,X1)|c8_1(a566)|c4_1(a552)|c3_1(a566)|ssSkC18|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_521, c_0_292]), ['final']).
cnf(c_0_1337,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,X1)|c8_1(a538)|c4_1(a552)|c3_1(a538)|ssSkC18|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_521, c_0_280]), ['final']).
cnf(c_0_1338,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,X1)|c8_1(a494)|c4_1(a552)|c3_1(a494)|ssSkC18|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_521, c_0_281]), ['final']).
cnf(c_0_1339,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,X1)|c8_1(a558)|c4_1(a552)|c3_1(a558)|ssSkC18|~c7_0), inference(spm,[status(thm)],[c_0_521, c_0_282]), ['final']).
cnf(c_0_1340,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c3_2(a552,X1)|c1_2(a552,X1)|c2_2(a552,X1)|c4_1(a552)|c3_1(a466)|ssSkC18|~c7_0), inference(spm,[status(thm)],[c_0_710, c_0_596]), ['final']).
cnf(c_0_1341,negated_conjecture,(ssSkC8|c6_2(a568,X1)|c8_2(a568,X1)|c8_2(X2,a467)|c8_1(a568)|ssSkC21|~c7_2(a568,X1)|~c8_1(X2)|~c2_1(a568)|~c6_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_513]), c_0_353]), ['final']).
cnf(c_0_1342,negated_conjecture,(ssSkC8|c6_2(a568,X1)|c8_2(a568,X1)|c3_2(X2,a467)|c8_1(a568)|ssSkC21|~c7_2(a568,X1)|~c8_1(X2)|~c2_1(a568)|~c6_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_513]), c_0_353]), ['final']).
cnf(c_0_1343,negated_conjecture,(ssSkC8|c6_2(a568,X1)|c8_2(a568,X1)|c2_2(X2,a467)|c8_1(a568)|ssSkC21|~c7_2(a568,X1)|~c8_1(X2)|~c2_1(a568)|~c6_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_513]), c_0_353]), ['final']).
cnf(c_0_1344,negated_conjecture,(ssSkC12|c4_2(a542,a527)|c8_2(a542,X1)|c3_2(a542,a557)|c2_2(a542,a557)|c2_2(a542,X1)|c5_1(a565)|c7_1(a542)|ssSkC16|ssSkC20|~ssSkP6(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_383]), c_0_27]), ['final']).
cnf(c_0_1345,negated_conjecture,(ssSkC8|c1_2(a587,a557)|c2_2(a587,a557)|c5_1(a565)|c7_1(a587)|ssSkC20|ssSkC23|~c4_2(a587,a557)|~c2_2(a587,a516)|~ssSkP5(a587)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_711, c_0_653]), c_0_353]), ['final']).
cnf(c_0_1346,negated_conjecture,(ssSkC8|c1_2(a587,a557)|c2_2(a587,a557)|c5_1(a565)|c7_1(a587)|ssSkC20|ssSkC23|~c6_2(a587,a516)|~c4_2(a587,a557)|~ssSkP5(a587)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_711, c_0_656]), c_0_353]), ['final']).
cnf(c_0_1347,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,a456)|c2_2(X1,X2)|ssSkP0(X1)|ndr1_1(a559)|~c2_2(X1,a456)|~c8_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1348,negated_conjecture,(ssSkC22|c8_2(a532,a576)|c8_1(a532)|ndr1_1(X1)|c7_1(a532)|ssSkC13|~c7_2(a532,a576)|~c8_1(X1)|~c2_1(a532)|~c6_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_712, c_0_706]), c_0_707]), ['final']).
cnf(c_0_1349,negated_conjecture,(c8_2(a565,a487)|c8_1(a565)|ssSkC5|ndr1_1(X1)|c5_1(a565)|ssSkC20|~c7_2(a565,a487)|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_712, c_0_708]), c_0_534]), c_0_16]), ['final']).
cnf(c_0_1350,negated_conjecture,(c8_2(a542,a487)|c8_1(a542)|ssSkC5|ndr1_1(X1)|c5_1(a542)|ssSkC16|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_712, c_0_386]), c_0_534]), c_0_34]), c_0_409]), ['final']).
cnf(c_0_1351,negated_conjecture,(c8_2(a517,a487)|c8_1(a517)|ssSkC5|ndr1_1(X1)|c5_1(a517)|ssSkC11|~c7_2(a517,a487)|~c8_1(X1)|~c6_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_712, c_0_709]), c_0_534]), c_0_22]), ['final']).
cnf(c_0_1352,negated_conjecture,(ssSkC12|ssSkC8|c4_2(a587,a527)|c3_2(a587,a557)|c2_2(a587,a588)|c2_2(a587,a557)|ssSkP8(a587)|c5_1(a565)|c7_1(a587)|ssSkC20|ssSkC23|~ssSkP6(a587)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_713]), c_0_353]), ['final']).
cnf(c_0_1353,negated_conjecture,(ssSkC8|c4_2(a592,a526)|c2_2(a592,a588)|c8_1(a592)|ssSkP6(a592)|c1_1(a573)|ssSkP8(a592)|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_714]), c_0_21]), ['final']).
cnf(c_0_1354,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,a456)|c2_2(X1,X2)|ssSkP0(X1)|c7_1(a559)|~c2_2(X1,a456)|~c8_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1355,negated_conjecture,(ssSkC8|c7_2(a592,a503)|c7_2(X1,a583)|ssSkP4(a592)|c4_1(a592)|c5_1(X1)|c7_1(a592)|c7_1(X1)|ssSkC24|~c2_2(a592,a503)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_691]), c_0_353]), ['final']).
cnf(c_0_1356,negated_conjecture,(ssSkC8|c7_2(a592,a503)|c2_2(X1,a583)|ssSkP4(a592)|c4_1(a592)|c5_1(X1)|c7_1(a592)|c7_1(X1)|ssSkC24|~c2_2(a592,a503)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_691]), c_0_353]), ['final']).
cnf(c_0_1357,negated_conjecture,(ssSkC8|c7_2(a491,a503)|c7_2(X1,a583)|ssSkP4(a491)|c5_1(X1)|c7_1(a491)|c7_1(X1)|ssSkC6|~c2_2(a491,a503)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_693]), c_0_353]), c_0_291]), ['final']).
cnf(c_0_1358,negated_conjecture,(ssSkC8|c7_2(a491,a503)|c2_2(X1,a583)|ssSkP4(a491)|c5_1(X1)|c7_1(a491)|c7_1(X1)|ssSkC6|~c2_2(a491,a503)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_693]), c_0_353]), c_0_291]), ['final']).
cnf(c_0_1359,negated_conjecture,(ssSkC8|c7_2(a568,a503)|c7_2(X1,a583)|ssSkP4(a568)|c5_1(X1)|c7_1(a568)|c7_1(X1)|ssSkC21|~c2_2(a568,a503)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_692]), c_0_353]), c_0_290]), ['final']).
cnf(c_0_1360,negated_conjecture,(ssSkC8|c7_2(a568,a503)|c2_2(X1,a583)|ssSkP4(a568)|c5_1(X1)|c7_1(a568)|c7_1(X1)|ssSkC21|~c2_2(a568,a503)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_692]), c_0_353]), c_0_290]), ['final']).
cnf(c_0_1361,negated_conjecture,(c3_2(a570,X1)|c1_2(a570,a557)|c1_2(a570,a493)|c2_2(a570,X1)|ssSkC7|c5_1(a565)|c7_1(a570)|ssSkC20|~c5_2(a570,X1)|~ssSkP2(a570)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_715]), c_0_421])]), ['final']).
cnf(c_0_1362,negated_conjecture,(c4_2(a498,a526)|c1_2(a498,a499)|c1_2(a498,a493)|c2_2(a498,a499)|c8_1(a498)|ssSkP6(a498)|c1_1(a573)|ssSkC7|c2_0|~ssSkP2(a498)), inference(spm,[status(thm)],[c_0_562, c_0_716]), ['final']).
cnf(c_0_1363,negated_conjecture,(c4_2(a509,a526)|c5_2(a509,a526)|c1_2(a509,a511)|c2_2(a509,a511)|c8_1(a509)|ssSkP6(a509)|c1_1(a573)|~c4_0), inference(spm,[status(thm)],[c_0_450, c_0_717]), ['final']).
cnf(c_0_1364,negated_conjecture,(c4_2(a496,a526)|c5_2(a496,a526)|c1_2(a496,a497)|c2_2(a496,a497)|c8_1(a496)|ssSkP6(a496)|c1_1(a573)|~c4_0), inference(spm,[status(thm)],[c_0_450, c_0_663]), ['final']).
cnf(c_0_1365,negated_conjecture,(c4_2(a512,a526)|c1_2(a512,a493)|c8_1(a512)|ssSkP6(a512)|c1_1(a573)|ssSkC7|ssSkC9|~ssSkP2(a512)), inference(spm,[status(thm)],[c_0_562, c_0_718]), ['final']).
cnf(c_0_1366,negated_conjecture,(ssSkC8|c6_2(a552,a515)|ssSkP5(a552)|c4_1(a531)|c5_1(a552)|ssSkC18|~c1_2(a552,a515)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_719, c_0_23]), c_0_277]), c_0_24]), ['final']).
cnf(c_0_1367,negated_conjecture,(c3_2(a520,X1)|c1_2(a520,a557)|c1_2(a520,a493)|c2_2(a520,X1)|ssSkC7|c5_1(a565)|c7_1(a520)|ssSkC20|~c5_2(a520,X1)|~ssSkP2(a520)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_401]), c_0_380])]), ['final']).
cnf(c_0_1368,negated_conjecture,(c4_2(a498,a526)|c1_2(a498,a499)|c1_2(a498,a493)|c2_2(a498,a499)|c8_1(a498)|ssSkP6(a498)|ssSkC7|c2_0|c3_1(a573)|~ssSkP2(a498)), inference(spm,[status(thm)],[c_0_562, c_0_720]), ['final']).
cnf(c_0_1369,negated_conjecture,(c6_2(a498,X1)|c3_2(a498,a500)|c3_2(a498,X1)|c4_1(a531)|c2_0|c5_1(a498)|~c7_2(a498,X1)|~c1_2(a498,a500)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_533]), c_0_525]), ['final']).
cnf(c_0_1370,negated_conjecture,(c4_2(a559,a560)|c8_2(a566,X1)|c7_2(a566,X1)|c3_2(a566,a567)|c2_2(a566,X1)|~c2_2(a566,a567)|~c8_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1371,negated_conjecture,(c4_2(a559,a560)|c8_2(a505,X1)|c7_2(a505,X1)|c3_2(a505,a506)|c2_2(a505,X1)|~ssSkC8|~c2_2(a505,a506)|~c8_1(a505)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1372,negated_conjecture,(c4_2(a509,a515)|c5_2(a509,a515)|c1_2(a509,a511)|c2_2(a509,a511)|c8_1(a509)|ssSkP5(a509)|c1_1(a573)|~c4_0), inference(spm,[status(thm)],[c_0_417, c_0_717]), ['final']).
cnf(c_0_1373,negated_conjecture,(c4_2(a496,a515)|c5_2(a496,a515)|c1_2(a496,a497)|c2_2(a496,a497)|c8_1(a496)|ssSkP5(a496)|c1_1(a573)|~c4_0), inference(spm,[status(thm)],[c_0_417, c_0_663]), ['final']).
cnf(c_0_1374,negated_conjecture,(c4_2(a498,a515)|c5_2(a498,a515)|c1_2(a498,a499)|c2_2(a498,a499)|c8_1(a498)|ssSkP5(a498)|c1_1(a573)|c2_0), inference(spm,[status(thm)],[c_0_417, c_0_523]), ['final']).
cnf(c_0_1375,negated_conjecture,(ssSkC8|c4_2(a592,a526)|c2_2(a592,a588)|c8_1(a592)|ssSkP6(a592)|ssSkP8(a592)|c3_1(a573)|ssSkC24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_721]), c_0_21]), ['final']).
cnf(c_0_1376,negated_conjecture,(c8_2(a566,X1)|c7_2(a559,a560)|c7_2(a566,X1)|c3_2(a566,a567)|c2_2(a566,X1)|~c2_2(a566,a567)|~c8_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1377,negated_conjecture,(c8_2(a505,X1)|c7_2(a559,a560)|c7_2(a505,X1)|c3_2(a505,a506)|c2_2(a505,X1)|~ssSkC8|~c2_2(a505,a506)|~c8_1(a505)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1378,negated_conjecture,(c7_2(a505,X1)|c1_2(a505,X1)|c6_1(a505)|ndr1_1(a534)|~c3_2(a505,a506)|~c1_2(a505,a506)|~c2_2(a505,X1)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_705, c_0_285]), c_0_353]), c_0_660]), ['final']).
cnf(c_0_1379,negated_conjecture,(ssSkC3|c6_2(a570,a479)|c3_2(a570,X1)|c1_2(a570,a557)|c1_2(a570,X1)|c2_2(a570,X1)|c5_1(a565)|c7_1(a570)|ssSkC20|~c8_1(a570)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_715]), c_0_421])]), ['final']).
cnf(c_0_1380,negated_conjecture,(c4_2(a509,a526)|c5_2(a509,a526)|c1_2(a509,a511)|c2_2(a509,a511)|c8_1(a509)|ssSkP6(a509)|c3_1(a573)|~c4_0), inference(spm,[status(thm)],[c_0_449, c_0_717]), ['final']).
cnf(c_0_1381,negated_conjecture,(c4_2(a496,a526)|c5_2(a496,a526)|c1_2(a496,a497)|c2_2(a496,a497)|c8_1(a496)|ssSkP6(a496)|c3_1(a573)|~c4_0), inference(spm,[status(thm)],[c_0_449, c_0_663]), ['final']).
cnf(c_0_1382,negated_conjecture,(c4_2(a512,a526)|c1_2(a512,a493)|c8_1(a512)|ssSkP6(a512)|ssSkC7|c3_1(a573)|ssSkC9|~ssSkP2(a512)), inference(spm,[status(thm)],[c_0_562, c_0_722]), ['final']).
cnf(c_0_1383,negated_conjecture,(ssSkC3|c6_2(a520,a479)|c3_2(a520,X1)|c1_2(a520,a557)|c1_2(a520,X1)|c2_2(a520,X1)|c5_1(a565)|c7_1(a520)|ssSkC20|~c8_1(a520)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_401]), c_0_380])]), ['final']).
cnf(c_0_1384,negated_conjecture,(c4_2(a534,a535)|c7_2(a496,X1)|c1_2(a496,X1)|c6_1(a496)|~c1_2(a496,a497)|~c2_2(a496,X1)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_700, c_0_661]), c_0_662]), c_0_663]), ['final']).
cnf(c_0_1385,negated_conjecture,(c8_2(a534,a535)|c7_2(a496,X1)|c1_2(a496,X1)|c6_1(a496)|~c1_2(a496,a497)|~c2_2(a496,X1)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_701, c_0_661]), c_0_662]), c_0_663]), ['final']).
cnf(c_0_1386,negated_conjecture,(ssSkC8|c6_2(a552,X1)|c3_2(a552,X1)|c1_2(a552,a493)|c2_2(a552,X1)|ssSkC7|ssSkC18|~c5_2(a552,X1)|~ssSkP2(a552)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_363]), c_0_353]), ['final']).
cnf(c_0_1387,negated_conjecture,(c4_2(a546,a547)|c5_2(a546,a547)|c1_2(a546,X1)|c2_2(a546,X1)|c8_1(a546)|c1_1(a573)|ssSkC17|~c3_2(a546,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_126]), c_0_45]), ['final']).
cnf(c_0_1388,negated_conjecture,(c6_2(a494,X1)|c3_2(a494,X1)|c4_1(a531)|c5_1(a494)|~c7_2(a494,X1)|~c1_2(a494,a495)|~ssSkC7|~ssSkC6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_293]), c_0_294]), c_0_295]), ['final']).
cnf(c_0_1389,negated_conjecture,(c4_2(a509,a515)|c5_2(a509,a515)|c1_2(a509,a511)|c2_2(a509,a511)|c8_1(a509)|ssSkP5(a509)|c3_1(a573)|~c4_0), inference(spm,[status(thm)],[c_0_419, c_0_717]), ['final']).
cnf(c_0_1390,negated_conjecture,(c4_2(a496,a515)|c5_2(a496,a515)|c1_2(a496,a497)|c2_2(a496,a497)|c8_1(a496)|ssSkP5(a496)|c3_1(a573)|~c4_0), inference(spm,[status(thm)],[c_0_419, c_0_663]), ['final']).
cnf(c_0_1391,negated_conjecture,(c4_2(a498,a515)|c5_2(a498,a515)|c1_2(a498,a499)|c2_2(a498,a499)|c8_1(a498)|ssSkP5(a498)|c2_0|c3_1(a573)), inference(spm,[status(thm)],[c_0_419, c_0_523]), ['final']).
cnf(c_0_1392,negated_conjecture,(c6_2(X1,a526)|ssSkP6(X1)|c4_1(a531)|c5_1(X1)|~c1_2(X1,a526)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_98]), c_0_31]), c_0_20]), c_0_21]), ['final']).
cnf(c_0_1393,negated_conjecture,(c8_2(a566,X1)|c7_2(a566,X1)|c3_2(a566,a567)|c2_2(a566,X1)|ndr1_1(a559)|~c2_2(a566,a567)|~c8_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1394,negated_conjecture,(c8_2(a505,X1)|c7_2(a505,X1)|c3_2(a505,a506)|c2_2(a505,X1)|ndr1_1(a559)|~c2_2(a505,a506)|~c8_1(a505)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_285]), c_0_353]), c_0_660]), ['final']).
cnf(c_0_1395,negated_conjecture,(c8_2(a566,X1)|c7_2(a566,X1)|c3_2(a566,a567)|c2_2(a566,X1)|c7_1(a559)|~c2_2(a566,a567)|~c8_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1396,negated_conjecture,(c8_2(a505,X1)|c7_2(a505,X1)|c3_2(a505,a506)|c2_2(a505,X1)|c7_1(a559)|~ssSkC8|~c2_2(a505,a506)|~c8_1(a505)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1397,negated_conjecture,(c3_2(a570,X1)|c1_2(a570,a557)|c2_2(a570,X1)|c5_1(a565)|c7_1(a570)|ssSkC20|~c6_2(a570,X1)|~c4_2(a570,a548)|~c1_1(a570)|~ssSkC17), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_715]), c_0_421])]), ['final']).
cnf(c_0_1398,negated_conjecture,(c3_2(a570,X1)|c1_2(a570,a557)|c2_2(a570,X1)|c5_1(a565)|c7_1(a570)|ssSkC20|~c6_2(a570,X1)|~c7_2(a570,a548)|~c1_1(a570)|~ssSkC17), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_715]), c_0_421])]), ['final']).
cnf(c_0_1399,negated_conjecture,(c3_2(a520,X1)|c1_2(a520,a557)|c2_2(a520,X1)|c5_1(a565)|c7_1(a520)|ssSkC20|~c6_2(a520,X1)|~c4_2(a520,a548)|~c1_1(a520)|~ssSkC17), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_401]), c_0_380])]), ['final']).
cnf(c_0_1400,negated_conjecture,(c3_2(a520,X1)|c1_2(a520,a557)|c2_2(a520,X1)|c5_1(a565)|c7_1(a520)|ssSkC20|~c6_2(a520,X1)|~c7_2(a520,a548)|~c1_1(a520)|~ssSkC17), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_401]), c_0_380])]), ['final']).
cnf(c_0_1401,negated_conjecture,(c4_2(a498,a499)|c5_2(a498,a499)|c1_2(a498,a499)|c2_2(a498,a499)|c8_1(a498)|c1_1(a573)|c2_0), inference(spm,[status(thm)],[c_0_723, c_0_523]), ['final']).
cnf(c_0_1402,negated_conjecture,(c4_2(a546,a547)|c5_2(a546,a547)|c1_2(a546,X1)|c2_2(a546,X1)|c8_1(a546)|c3_1(a573)|ssSkC17|~c3_2(a546,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_126]), c_0_45]), ['final']).
cnf(c_0_1403,negated_conjecture,(c4_2(a498,a499)|c5_2(a498,a499)|c1_2(a498,a499)|c2_2(a498,a499)|c8_1(a498)|c2_0|c3_1(a573)), inference(spm,[status(thm)],[c_0_724, c_0_523]), ['final']).
cnf(c_0_1404,negated_conjecture,(c4_2(a559,a560)|c8_2(a549,X1)|c7_2(a549,X1)|c3_2(a549,a550)|c2_2(a549,X1)|~c2_2(a549,a550)|~c8_1(a549)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_694]), c_0_725]), ['final']).
cnf(c_0_1405,negated_conjecture,(c8_2(a549,X1)|c7_2(a559,a560)|c7_2(a549,X1)|c3_2(a549,a550)|c2_2(a549,X1)|~c2_2(a549,a550)|~c8_1(a549)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_694]), c_0_725]), ['final']).
cnf(c_0_1406,negated_conjecture,(c7_2(a496,X1)|c1_2(a496,X1)|c8_1(a534)|c6_1(a496)|~c1_2(a496,a497)|~c2_2(a496,X1)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_704, c_0_661]), c_0_662]), c_0_663]), ['final']).
cnf(c_0_1407,negated_conjecture,(ssSkC12|c2_2(a552,a503)|ssSkP4(a552)|c1_1(a508)|c2_1(a552)|c4_1(a458)|ssSkC18|~c7_2(a552,a527)|~c3_2(a552,a527)|~c4_1(a552)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_726]), c_0_9]), ['final']).
cnf(c_0_1408,negated_conjecture,(ssSkC12|c2_2(a512,a503)|ssSkP4(a512)|c1_1(a508)|c2_1(a512)|c4_1(a458)|ssSkC9|~c7_2(a512,a527)|~c3_2(a512,a527)|~c4_1(a512)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_727]), c_0_32]), ['final']).
cnf(c_0_1409,negated_conjecture,(ssSkC12|c4_2(a570,a527)|c8_2(a570,X1)|c3_2(a570,a557)|c2_2(a570,a557)|c2_2(a570,X1)|c5_1(a565)|c7_1(a570)|ssSkC20|~ssSkP6(a570)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_715]), c_0_421])]), ['final']).
cnf(c_0_1410,negated_conjecture,(c7_2(a496,X1)|c1_2(a496,X1)|c6_1(a496)|ndr1_1(a534)|~c1_2(a496,a497)|~c2_2(a496,X1)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_705, c_0_661]), c_0_662]), c_0_663]), ['final']).
cnf(c_0_1411,negated_conjecture,(c6_2(a494,a495)|c4_1(a531)|c5_1(a494)|~c7_2(a494,a495)|~c1_2(a494,a495)|~ssSkC7|~ssSkC6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_293]), c_0_294]), c_0_295]), ['final']).
cnf(c_0_1412,negated_conjecture,(c6_2(a549,a551)|c3_2(a549,a551)|c4_1(a531)|c5_1(a549)|~c7_2(a549,a551)|~c1_2(a549,a551)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_728]), c_0_725]), ['final']).
cnf(c_0_1413,negated_conjecture,(c6_2(X1,a460)|c3_2(X1,a460)|c4_1(a531)|c5_1(X1)|~ssSkC0|~c7_2(X1,a460)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_729]), c_0_541]), c_0_581]), ['final']).
cnf(c_0_1414,negated_conjecture,(ssSkC8|c6_2(a491,a548)|c3_2(a491,a548)|c3_2(a491,X1)|ssSkC6|~c6_2(a491,X1)|~c8_2(a491,X1)|~c1_1(a491)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_511]), c_0_353]), ['final']).
cnf(c_0_1415,negated_conjecture,(ssSkC12|c4_2(a520,a527)|c8_2(a520,X1)|c3_2(a520,a557)|c2_2(a520,a557)|c2_2(a520,X1)|c5_1(a565)|c7_1(a520)|ssSkC20|~ssSkP6(a520)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_401]), c_0_380])]), ['final']).
cnf(c_0_1416,negated_conjecture,(ssSkC8|c4_2(a565,a502)|c8_2(a565,a481)|ssSkC4|c8_1(a565)|ssSkP3(a565)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20|~c3_2(a565,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_634]), c_0_353]), ['final']).
cnf(c_0_1417,negated_conjecture,(ssSkC8|c4_2(a565,a502)|ssSkC4|c2_2(a565,a481)|c8_1(a565)|ssSkP3(a565)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20|~c3_2(a565,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_634]), c_0_353]), ['final']).
cnf(c_0_1418,negated_conjecture,(ssSkC8|c4_2(a517,a502)|c8_2(a517,a481)|ssSkC4|c8_1(a517)|ssSkP3(a517)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11|~c3_2(a517,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_635]), c_0_353]), ['final']).
cnf(c_0_1419,negated_conjecture,(ssSkC8|c4_2(a517,a502)|ssSkC4|c2_2(a517,a481)|c8_1(a517)|ssSkP3(a517)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11|~c3_2(a517,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_635]), c_0_353]), ['final']).
cnf(c_0_1420,negated_conjecture,(ssSkC8|c3_2(a600,a557)|c2_2(a600,a502)|ssSkP3(a600)|c2_1(a600)|c4_1(a458)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c6_1(a600)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_477]), c_0_353]), ['final']).
cnf(c_0_1421,negated_conjecture,(c7_2(X1,a456)|c7_2(X2,a583)|ssSkP0(X1)|c4_1(X1)|c5_1(X2)|c7_1(X1)|c7_1(X2)|~c2_2(X1,a456)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1422,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(a565)|c7_1(X1)|ssSkC20|~c6_2(X1,X2)|~c4_2(X1,a548)|~c1_1(X1)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_357]), c_0_353]), ['final']).
cnf(c_0_1423,negated_conjecture,(ssSkC8|c3_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(a565)|c7_1(X1)|ssSkC20|~c6_2(X1,X2)|~c7_2(X1,a548)|~c1_1(X1)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_357]), c_0_353]), ['final']).
cnf(c_0_1424,negated_conjecture,(c6_2(a549,X1)|c3_2(a549,a551)|c3_2(a549,X1)|c4_1(a531)|c5_1(a549)|~c7_2(a549,X1)|~c1_2(a549,a551)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_728]), c_0_725]), ['final']).
cnf(c_0_1425,negated_conjecture,(c8_2(a549,X1)|c7_2(a549,X1)|c3_2(a549,a550)|c2_2(a549,X1)|ndr1_1(a559)|~c2_2(a549,a550)|~c8_1(a549)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_694]), c_0_725]), ['final']).
cnf(c_0_1426,negated_conjecture,(c8_2(a549,X1)|c7_2(a549,X1)|c3_2(a549,a550)|c2_2(a549,X1)|c7_1(a559)|~c2_2(a549,a550)|~c8_1(a549)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_694]), c_0_725]), ['final']).
cnf(c_0_1427,negated_conjecture,(c4_2(a542,a502)|c8_2(a542,a481)|ssSkC4|c8_1(a542)|ssSkP3(a542)|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c3_2(a542,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1428,negated_conjecture,(c4_2(a542,a502)|ssSkC4|c2_2(a542,a481)|c8_1(a542)|ssSkP3(a542)|c1_1(a542)|c7_0|c5_1(a501)|ssSkC16|~c3_2(a542,a502)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_628]), c_0_8]), ['final']).
cnf(c_0_1429,negated_conjecture,(c7_2(X1,a456)|c2_2(X2,a583)|ssSkP0(X1)|c4_1(X1)|c5_1(X2)|c7_1(X1)|c7_1(X2)|~c2_2(X1,a456)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1430,negated_conjecture,(c7_2(X1,a583)|c4_1(X1)|c5_1(X1)|c7_1(X1)|~c6_2(X1,a583)|~c2_2(X1,a583)|~ndr1_1(X1)), inference(ef,[status(thm)],[c_0_695]), ['final']).
cnf(c_0_1431,negated_conjecture,(ssSkC8|c1_2(a587,a526)|ssSkP6(a587)|ssSkC23|~c4_2(a587,a526)|~c2_2(a587,a516)|~ssSkP5(a587)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_732, c_0_653]), c_0_353]), ['final']).
cnf(c_0_1432,negated_conjecture,(ssSkC8|c1_2(a587,a526)|ssSkP6(a587)|ssSkC23|~c6_2(a587,a516)|~c4_2(a587,a526)|~ssSkP5(a587)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_732, c_0_656]), c_0_353]), ['final']).
cnf(c_0_1433,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c4_2(a574,X3)|~c3_2(a574,X3)|~c2_2(a574,X3)|~c6_1(X1)|~c7_0|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_733, c_0_596]), ['final']).
cnf(c_0_1434,negated_conjecture,(c8_2(X1,a467)|c8_1(a509)|~c7_2(a509,a510)|~c8_1(X1)|~c2_1(a509)|~c6_1(X1)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_734]), c_0_735]), c_0_736]), ['final']).
cnf(c_0_1435,negated_conjecture,(c3_2(a566,a456)|c1_2(a566,a493)|ssSkP0(a566)|ssSkC7|~c1_2(a566,a456)|~c2_2(a566,a456)|~ssSkP2(a566)|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_737, c_0_738]), ['final']).
cnf(c_0_1436,negated_conjecture,(ssSkC8|c1_2(a564,a493)|ssSkP4(a564)|ssSkC7|ssSkC19|~c5_2(a564,a503)|~ssSkP2(a564)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_739]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1437,negated_conjecture,(c7_2(X1,a456)|ssSkP0(X1)|c4_1(X1)|ndr1_1(X2)|c5_1(X2)|c7_1(X1)|c7_1(X2)|~c2_2(X1,a456)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_121]), c_0_100]), ['final']).
cnf(c_0_1438,negated_conjecture,(c3_2(X1,a467)|c8_1(a509)|~c7_2(a509,a510)|~c8_1(X1)|~c2_1(a509)|~c6_1(X1)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_734]), c_0_735]), c_0_736]), ['final']).
cnf(c_0_1439,negated_conjecture,(c2_2(X1,a467)|c8_1(a509)|~c7_2(a509,a510)|~c8_1(X1)|~c2_1(a509)|~c6_1(X1)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_734]), c_0_735]), c_0_736]), ['final']).
cnf(c_0_1440,negated_conjecture,(c7_2(a566,a567)|c7_2(X1,a583)|c4_1(a566)|c5_1(X1)|c7_1(a566)|c7_1(X1)|~c2_2(a566,a567)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1441,negated_conjecture,(c7_2(a505,a506)|c7_2(X1,a583)|c4_1(a505)|c5_1(X1)|c7_1(a505)|c7_1(X1)|~ssSkC8|~c2_2(a505,a506)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1442,negated_conjecture,(c7_2(a566,a567)|c2_2(X1,a583)|c4_1(a566)|c5_1(X1)|c7_1(a566)|c7_1(X1)|~c2_2(a566,a567)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1443,negated_conjecture,(c1_2(X1,a526)|c8_1(a566)|c8_1(X1)|ssSkP6(X1)|c4_1(X1)|c3_1(a566)|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_530, c_0_292]), ['final']).
cnf(c_0_1444,negated_conjecture,(c1_2(X1,a526)|c8_1(a538)|c8_1(X1)|ssSkP6(X1)|c4_1(X1)|c3_1(a538)|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_530, c_0_280]), ['final']).
cnf(c_0_1445,negated_conjecture,(c1_2(X1,a526)|c8_1(a494)|c8_1(X1)|ssSkP6(X1)|c4_1(X1)|c3_1(a494)|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_530, c_0_281]), ['final']).
cnf(c_0_1446,negated_conjecture,(ssSkC22|ssSkC8|c6_2(a592,a557)|c3_2(a592,a557)|c2_2(a592,a576)|c5_1(a565)|c7_1(a592)|ssSkC20|ssSkC24|~c1_1(a592)), inference(spm,[status(thm)],[c_0_665, c_0_392]), ['final']).
cnf(c_0_1447,negated_conjecture,(c1_2(X1,a526)|c8_1(a558)|c8_1(X1)|ssSkP6(X1)|c4_1(X1)|c3_1(a558)|~c7_0), inference(spm,[status(thm)],[c_0_530, c_0_282]), ['final']).
cnf(c_0_1448,negated_conjecture,(c1_2(X1,a526)|c8_1(X1)|ssSkP6(X1)|c4_1(X1)|c3_1(a466)|~c7_0), inference(spm,[status(thm)],[c_0_740, c_0_596]), ['final']).
cnf(c_0_1449,negated_conjecture,(c7_2(a505,a506)|c2_2(X1,a583)|c4_1(a505)|c5_1(X1)|c7_1(a505)|c7_1(X1)|~ssSkC8|~c2_2(a505,a506)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1450,negated_conjecture,(c3_2(X1,a456)|c1_2(X1,a456)|c8_1(a566)|c8_1(X1)|ssSkP0(X1)|c4_1(X1)|c3_1(a566)|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_531, c_0_292]), ['final']).
cnf(c_0_1451,negated_conjecture,(c3_2(X1,a456)|c1_2(X1,a456)|c8_1(a538)|c8_1(X1)|ssSkP0(X1)|c4_1(X1)|c3_1(a538)|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_531, c_0_280]), ['final']).
cnf(c_0_1452,negated_conjecture,(c3_2(X1,a456)|c1_2(X1,a456)|c8_1(a494)|c8_1(X1)|ssSkP0(X1)|c4_1(X1)|c3_1(a494)|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_531, c_0_281]), ['final']).
cnf(c_0_1453,negated_conjecture,(c3_2(X1,a456)|c1_2(X1,a456)|c8_1(a558)|c8_1(X1)|ssSkP0(X1)|c4_1(X1)|c3_1(a558)|~c7_0), inference(spm,[status(thm)],[c_0_531, c_0_282]), ['final']).
cnf(c_0_1454,negated_conjecture,(c3_2(X1,a456)|c1_2(X1,a456)|c8_1(X1)|ssSkP0(X1)|c4_1(X1)|c3_1(a466)|~c7_0), inference(spm,[status(thm)],[c_0_741, c_0_596]), ['final']).
cnf(c_0_1455,negated_conjecture,(c3_2(a542,a543)|c1_2(a542,a543)|c8_1(a566)|c8_1(a542)|c4_1(a542)|c3_1(a566)|ssSkC16|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_532, c_0_292]), ['final']).
cnf(c_0_1456,negated_conjecture,(c3_2(a542,a543)|c1_2(a542,a543)|c8_1(a538)|c8_1(a542)|c4_1(a542)|c3_1(a538)|ssSkC16|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_532, c_0_280]), ['final']).
cnf(c_0_1457,negated_conjecture,(c3_2(a542,a543)|c1_2(a542,a543)|c8_1(a494)|c8_1(a542)|c4_1(a542)|c3_1(a494)|ssSkC16|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_532, c_0_281]), ['final']).
cnf(c_0_1458,negated_conjecture,(c3_2(a542,a543)|c1_2(a542,a543)|c8_1(a558)|c8_1(a542)|c4_1(a542)|c3_1(a558)|ssSkC16|~c7_0), inference(spm,[status(thm)],[c_0_532, c_0_282]), ['final']).
cnf(c_0_1459,negated_conjecture,(c3_2(a542,a543)|c1_2(a542,a543)|c8_1(a542)|c4_1(a542)|c3_1(a466)|ssSkC16|~c7_0), inference(spm,[status(thm)],[c_0_742, c_0_596]), ['final']).
cnf(c_0_1460,negated_conjecture,(ssSkC8|ssSkP4(a564)|ssSkC19|~c6_2(a564,a503)|~c7_2(a564,a548)|~c1_1(a564)|~ssSkC17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_739]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1461,negated_conjecture,(c8_1(a509)|ndr1_1(X1)|~c7_2(a509,a510)|~c8_1(X1)|~c2_1(a509)|~c6_1(X1)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_712, c_0_734]), c_0_735]), c_0_736]), ['final']).
cnf(c_0_1462,negated_conjecture,(c7_2(a566,a567)|c4_1(a566)|ndr1_1(X1)|c5_1(X1)|c7_1(a566)|c7_1(X1)|~c2_2(a566,a567)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1463,negated_conjecture,(c7_2(a505,a506)|c4_1(a505)|ndr1_1(X1)|c5_1(X1)|c7_1(a505)|c7_1(X1)|~c2_2(a505,a506)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_285]), c_0_353]), c_0_660]), ['final']).
cnf(c_0_1464,negated_conjecture,(ssSkC8|c3_2(a537,a456)|c1_2(a537,a493)|c2_2(a537,a456)|c2_2(a537,a572)|ssSkP0(a537)|ssSkC7|ssSkC15|~ssSkP2(a537)), inference(spm,[status(thm)],[c_0_737, c_0_471]), ['final']).
cnf(c_0_1465,negated_conjecture,(ssSkC8|c3_2(a517,a456)|c1_2(a517,a493)|c2_2(a517,a456)|c2_2(a517,a572)|ssSkP0(a517)|ssSkC7|ssSkC11|~ssSkP2(a517)), inference(spm,[status(thm)],[c_0_737, c_0_472]), ['final']).
cnf(c_0_1466,negated_conjecture,(c3_2(a498,a500)|c1_2(a498,a500)|c8_1(a566)|c8_1(a498)|c4_1(a498)|c2_0|c3_1(a566)|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_743, c_0_292]), ['final']).
cnf(c_0_1467,negated_conjecture,(c3_2(a498,a500)|c1_2(a498,a500)|c8_1(a538)|c8_1(a498)|c4_1(a498)|c2_0|c3_1(a538)|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_743, c_0_280]), ['final']).
cnf(c_0_1468,negated_conjecture,(c3_2(a498,a500)|c1_2(a498,a500)|c8_1(a494)|c8_1(a498)|c4_1(a498)|c2_0|c3_1(a494)|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_743, c_0_281]), ['final']).
cnf(c_0_1469,negated_conjecture,(ssSkC8|c6_2(a568,a548)|c3_2(a568,X1)|ssSkC21|~c6_2(a568,X1)|~c8_2(a568,X1)|~c1_1(a568)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_513]), c_0_353]), ['final']).
cnf(c_0_1470,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|~c6_2(X3,a525)|~c6_2(X1,X2)|~c4_2(X1,X2)|~c7_0|~c4_1(X1)|~ndr1_1(X1)|~c3_1(X3)), inference(spm,[status(thm)],[c_0_744, c_0_596]), ['final']).
cnf(c_0_1471,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~c2_2(X3,a525)|~c7_0|~c4_1(X1)|~ndr1_1(X1)|~c3_1(X3)), inference(spm,[status(thm)],[c_0_745, c_0_596]), ['final']).
cnf(c_0_1472,negated_conjecture,(c7_2(a549,a550)|c7_2(X1,a583)|c4_1(a549)|c5_1(X1)|c7_1(a549)|c7_1(X1)|~c2_2(a549,a550)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_694]), c_0_725]), ['final']).
cnf(c_0_1473,negated_conjecture,(c7_2(a496,a497)|c7_2(X1,a583)|c4_1(a496)|c5_1(X1)|c7_1(a496)|c7_1(X1)|~c2_2(a496,a497)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_661]), c_0_662]), ['final']).
cnf(c_0_1474,negated_conjecture,(c1_2(X1,a493)|ssSkP6(X1)|ssSkC7|~c4_2(X1,a526)|~c2_2(X1,a516)|~ssSkP2(X1)|~ssSkP5(X1)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_562, c_0_653]), c_0_31]), ['final']).
cnf(c_0_1475,negated_conjecture,(c1_2(X1,a493)|ssSkP6(X1)|ssSkC7|~c6_2(X1,a516)|~c4_2(X1,a526)|~ssSkP2(X1)|~ssSkP5(X1)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_562, c_0_656]), c_0_31]), ['final']).
cnf(c_0_1476,negated_conjecture,(c1_2(a566,a493)|ssSkP6(a566)|ssSkC7|~c1_2(a566,a526)|~c2_2(a566,a526)|~ssSkP2(a566)|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_562, c_0_738]), ['final']).
cnf(c_0_1477,negated_conjecture,(c3_2(a570,a456)|c1_2(a570,a493)|c2_2(a570,a456)|c2_2(a570,a572)|ssSkP0(a570)|ssSkC7|~ssSkP2(a570)), inference(spm,[status(thm)],[c_0_737, c_0_463]), ['final']).
cnf(c_0_1478,negated_conjecture,(ssSkC8|c1_2(a537,a493)|c2_2(a537,a526)|c2_2(a537,a572)|ssSkP6(a537)|ssSkC7|ssSkC15|~ssSkP2(a537)), inference(spm,[status(thm)],[c_0_562, c_0_471]), ['final']).
cnf(c_0_1479,negated_conjecture,(ssSkC8|c1_2(a517,a493)|c2_2(a517,a526)|c2_2(a517,a572)|ssSkP6(a517)|ssSkC7|ssSkC11|~ssSkP2(a517)), inference(spm,[status(thm)],[c_0_562, c_0_472]), ['final']).
cnf(c_0_1480,negated_conjecture,(c1_2(a570,a493)|c2_2(a570,a526)|c2_2(a570,a572)|ssSkP6(a570)|ssSkC7|~ssSkP2(a570)), inference(spm,[status(thm)],[c_0_562, c_0_463]), ['final']).
cnf(c_0_1481,negated_conjecture,(ssSkC8|c1_2(a536,a493)|ssSkP6(a536)|ssSkC7|ssSkC14|~ssSkP2(a536)), inference(spm,[status(thm)],[c_0_562, c_0_746]), ['final']).
cnf(c_0_1482,negated_conjecture,(c7_2(a549,a550)|c2_2(X1,a583)|c4_1(a549)|c5_1(X1)|c7_1(a549)|c7_1(X1)|~c2_2(a549,a550)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_694]), c_0_725]), ['final']).
cnf(c_0_1483,negated_conjecture,(c7_2(a496,a497)|c2_2(X1,a583)|c4_1(a496)|c5_1(X1)|c7_1(a496)|c7_1(X1)|~c2_2(a496,a497)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_696, c_0_661]), c_0_662]), ['final']).
cnf(c_0_1484,negated_conjecture,(c7_2(X1,a525)|c3_2(X2,X3)|c8_1(X2)|c6_1(X1)|~c6_2(X2,X3)|~c4_2(X2,X3)|~c7_0|~c4_1(X2)|~ndr1_1(X2)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_747, c_0_596]), ['final']).
cnf(c_0_1485,negated_conjecture,(~c6_2(a494,a495)|~c7_2(a494,a548)|~ssSkC7|~ssSkC6|~ssSkC17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_293]), c_0_294]), c_0_281]), c_0_295]), ['final']).
cnf(c_0_1486,negated_conjecture,(c7_2(a549,a550)|c4_1(a549)|ndr1_1(X1)|c5_1(X1)|c7_1(a549)|c7_1(X1)|~c2_2(a549,a550)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_694]), c_0_725]), ['final']).
cnf(c_0_1487,negated_conjecture,(c7_2(a496,a497)|c4_1(a496)|ndr1_1(X1)|c5_1(X1)|c7_1(a496)|c7_1(X1)|~c2_2(a496,a497)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_661]), c_0_662]), ['final']).
cnf(c_0_1488,negated_conjecture,(ssSkC8|c3_2(a592,X1)|ssSkC24|~c4_2(a592,X1)|~c2_2(a592,a516)|~ssSkP5(a592)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_653]), c_0_353]), ['final']).
cnf(c_0_1489,negated_conjecture,(ssSkC8|c3_2(a592,X1)|ssSkC24|~c6_2(a592,a516)|~c4_2(a592,X1)|~ssSkP5(a592)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_656]), c_0_353]), ['final']).
cnf(c_0_1490,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,a460)|c3_2(X1,X2)|c4_1(a531)|c5_1(X1)|~ssSkC0|~c7_2(X1,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_729]), c_0_541]), c_0_581]), ['final']).
cnf(c_0_1491,negated_conjecture,(c3_2(X1,a460)|~ssSkC0|~c6_2(X1,a460)|~c7_2(X1,a548)|~c1_1(X1)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_729]), c_0_541]), ['final']).
cnf(c_0_1492,negated_conjecture,(c3_2(a549,a551)|c1_2(a549,a551)|c8_1(a549)|c8_1(X1)|c4_1(a549)|c3_1(X1)|~c1_1(X1)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_728]), c_0_725]), ['final']).
cnf(c_0_1493,negated_conjecture,(ssSkC12|c4_2(a565,a527)|c3_2(a565,a487)|c2_2(a565,a487)|ssSkC5|c5_1(a565)|ssSkC20|~ssSkP6(a565)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_613]), c_0_748]), ['final']).
cnf(c_0_1494,negated_conjecture,(ssSkC12|c4_2(a542,a527)|c3_2(a542,a487)|c2_2(a542,a487)|ssSkC5|c5_1(a542)|ssSkC16|~ssSkP6(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_605]), c_0_27]), ['final']).
cnf(c_0_1495,negated_conjecture,(ssSkC12|c4_2(a517,a527)|c3_2(a517,a487)|c2_2(a517,a487)|ssSkC5|c5_1(a517)|ssSkC11|~ssSkP6(a517)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_599]), c_0_749]), ['final']).
cnf(c_0_1496,negated_conjecture,(ssSkC22|ssSkC8|c7_2(a600,a557)|c2_2(a600,a576)|c5_1(a565)|c7_1(a600)|ssSkC20|ssSkC26|~c1_1(a600)), inference(spm,[status(thm)],[c_0_665, c_0_394]), ['final']).
cnf(c_0_1497,negated_conjecture,(c8_2(X1,a481)|ssSkC4|c8_1(X1)|ssSkP0(X1)|~c3_2(X1,a456)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1498,negated_conjecture,(ssSkC8|c1_2(a482,a557)|c2_2(a482,X1)|c5_1(a565)|c7_1(a482)|ssSkC20|~c6_2(a482,X1)|~ssSkC4), inference(spm,[status(thm)],[c_0_750, c_0_357]), ['final']).
cnf(c_0_1499,negated_conjecture,(c6_2(a565,X1)|c7_2(a565,X1)|c1_2(a565,X1)|c1_1(a565)|c7_0|c5_1(a501)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_362, c_0_541]), ['final']).
cnf(c_0_1500,negated_conjecture,(ssSkC22|c1_1(a508)|c7_1(a532)|ssSkC13|~c7_2(a532,a576)|~c3_2(a532,a576)|~c4_1(a532)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_706]), c_0_707]), ['final']).
cnf(c_0_1501,negated_conjecture,(c1_1(a508)|ssSkC5|c5_1(a565)|ssSkC20|~c7_2(a565,a487)|~c3_2(a565,a487)|~c4_1(a565)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_708]), c_0_748]), ['final']).
cnf(c_0_1502,negated_conjecture,(c1_1(a508)|ssSkC5|c5_1(a542)|ssSkC16|~c3_2(a542,a487)|~c4_1(a542)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_386]), c_0_27]), c_0_409]), ['final']).
cnf(c_0_1503,negated_conjecture,(ssSkC4|c2_2(X1,a481)|c8_1(X1)|ssSkP0(X1)|~c3_2(X1,a456)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_121]), c_0_100]), c_0_122]), ['final']).
cnf(c_0_1504,negated_conjecture,(c6_2(a517,X1)|c7_2(a517,X1)|c1_2(a517,X1)|c1_1(a517)|c7_0|c5_1(a501)|ssSkC11|~ssSkC0), inference(spm,[status(thm)],[c_0_369, c_0_541]), ['final']).
cnf(c_0_1505,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|ndr1_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~c7_0|~c4_1(X1)|~ndr1_1(X1)|~c3_1(X3)), inference(spm,[status(thm)],[c_0_751, c_0_596]), ['final']).
cnf(c_0_1506,negated_conjecture,(c5_2(X1,X2)|c2_2(X1,X2)|~c6_2(X1,a516)|~c4_2(X1,a572)|~ssSkP5(X1)|~c4_1(X1)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), inference(spm,[status(thm)],[c_0_752, c_0_656]), ['final']).
cnf(c_0_1507,negated_conjecture,(c8_1(X1)|c1_1(X1)|~c6_2(a530,X2)|~c8_2(a530,X2)|~c3_2(a530,X2)|~c6_1(X1)|~c7_0|~ndr1_1(a530)), inference(spm,[status(thm)],[c_0_753, c_0_596]), ['final']).
cnf(c_0_1508,negated_conjecture,(c8_2(a494,X1)|c1_2(a494,a557)|c2_2(a494,X1)|c5_1(a565)|c7_1(a494)|ssSkC20|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_352, c_0_294]), ['final']).
cnf(c_0_1509,negated_conjecture,(c1_1(a508)|ssSkC5|c5_1(a517)|ssSkC11|~c7_2(a517,a487)|~c3_2(a517,a487)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_709]), c_0_749]), c_0_53]), ['final']).
cnf(c_0_1510,negated_conjecture,(c8_2(a482,X1)|c1_2(a482,a557)|c2_2(a482,X1)|c5_1(a565)|c7_1(a482)|ssSkC20|~ssSkC4), inference(spm,[status(thm)],[c_0_352, c_0_535]), ['final']).
cnf(c_0_1511,negated_conjecture,(c8_2(a509,X1)|c1_2(a509,a557)|c2_2(a509,X1)|c5_1(a565)|c7_1(a509)|ssSkC20|~c4_0), inference(spm,[status(thm)],[c_0_352, c_0_735]), ['final']).
cnf(c_0_1512,negated_conjecture,(c5_2(X1,X2)|c2_2(X1,X2)|~c4_2(X1,a572)|~c2_2(X1,a516)|~ssSkP5(X1)|~c4_1(X1)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), inference(spm,[status(thm)],[c_0_752, c_0_653]), ['final']).
cnf(c_0_1513,negated_conjecture,(c8_2(a505,X1)|c1_2(a505,a557)|c2_2(a505,X1)|c5_1(a565)|c7_1(a505)|ssSkC20|~c4_0), inference(spm,[status(thm)],[c_0_352, c_0_660]), ['final']).
cnf(c_0_1514,negated_conjecture,(c8_2(a549,X1)|c1_2(a549,a557)|c2_2(a549,X1)|c5_1(a565)|c7_1(a549)|ssSkC20|~c4_0), inference(spm,[status(thm)],[c_0_352, c_0_725]), ['final']).
cnf(c_0_1515,negated_conjecture,(c8_2(a496,X1)|c1_2(a496,a557)|c2_2(a496,X1)|c5_1(a565)|c7_1(a496)|ssSkC20|~c4_0), inference(spm,[status(thm)],[c_0_352, c_0_662]), ['final']).
cnf(c_0_1516,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(a565)|c7_1(X1)|ssSkC20|~ssSkC0), inference(spm,[status(thm)],[c_0_352, c_0_541]), ['final']).
cnf(c_0_1517,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(a552)|c7_1(X1)|ssSkC18|~ndr1_1(X1)|~c3_1(a552)), inference(spm,[status(thm)],[c_0_351, c_0_41]), ['final']).
cnf(c_0_1518,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(a517)|c7_1(X1)|ssSkC11|~ndr1_1(X1)|~c3_1(a517)), inference(spm,[status(thm)],[c_0_351, c_0_296]), ['final']).
cnf(c_0_1519,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(a512)|c7_1(X1)|ssSkC9|~ndr1_1(X1)|~c3_1(a512)), inference(spm,[status(thm)],[c_0_351, c_0_42]), ['final']).
cnf(c_0_1520,negated_conjecture,(c3_2(a549,a551)|~c6_2(a549,a551)|~c7_2(a549,a548)|~c1_1(a549)|~c4_0|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_728]), c_0_725]), ['final']).
cnf(c_0_1521,negated_conjecture,(c3_2(a542,a543)|ssSkC16|~c6_2(a542,a543)|~c7_2(a542,a548)|~c1_1(a542)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_54]), c_0_27]), ['final']).
cnf(c_0_1522,negated_conjecture,(c3_2(a498,a500)|c2_0|~c6_2(a498,a500)|~c7_2(a498,a548)|~c1_1(a498)|~ssSkC17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_533]), c_0_525]), ['final']).
cnf(c_0_1523,negated_conjecture,(c4_2(a566,a567)|c8_2(a566,a481)|ssSkC4|c8_1(a566)|~c3_2(a566,a567)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1524,negated_conjecture,(c4_2(a505,a506)|c8_2(a505,a481)|ssSkC4|c8_1(a505)|~ssSkC8|~c3_2(a505,a506)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1525,negated_conjecture,(c4_2(a566,a567)|ssSkC4|c2_2(a566,a481)|c8_1(a566)|~c3_2(a566,a567)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_284]), c_0_160]), ['final']).
cnf(c_0_1526,negated_conjecture,(c4_2(a505,a506)|ssSkC4|c2_2(a505,a481)|c8_1(a505)|~ssSkC8|~c3_2(a505,a506)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_285]), c_0_660]), ['final']).
cnf(c_0_1527,negated_conjecture,(c5_2(a566,X1)|c2_2(a566,X1)|~c1_2(a566,a572)|~c4_1(a566)|~ssSkC19|~ssSkC20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_752, c_0_738]), c_0_395]), c_0_160]), ['final']).
cnf(c_0_1528,negated_conjecture,(ssSkP6(a558)|c2_1(a558)|c4_1(a458)|~c7_0), inference(spm,[status(thm)],[c_0_378, c_0_297]), ['final']).
cnf(c_0_1529,negated_conjecture,(ssSkP6(a468)|c2_1(a468)|c4_1(a458)|~c4_0), inference(spm,[status(thm)],[c_0_378, c_0_298]), ['final']).
cnf(c_0_1530,negated_conjecture,(c3_2(a542,a543)|c1_2(a542,a493)|ssSkC7|ssSkC16|~c5_2(a542,a543)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_54]), c_0_27]), c_0_388]), ['final']).
cnf(c_0_1531,negated_conjecture,(c3_2(a498,a500)|c1_2(a498,a493)|ssSkC7|c2_0|~c5_2(a498,a500)|~ssSkP2(a498)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_533]), c_0_525]), ['final']).
cnf(c_0_1532,negated_conjecture,(ssSkC8|c3_2(a470,a456)|ssSkP0(a470)|c2_1(a470)|c4_1(a458)|ssSkC2|~c6_1(a470)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_668]), c_0_353]), ['final']).
cnf(c_0_1533,negated_conjecture,(ssSkC8|ssSkP4(a600)|c2_1(a600)|c4_1(a458)|ssSkC26|~c6_1(a600)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_479]), c_0_353]), c_0_63]), ['final']).
cnf(c_0_1534,negated_conjecture,(ssSkC3|c6_2(a581,a479)|c1_2(a581,a526)|ssSkP6(a581)|~c7_0), inference(spm,[status(thm)],[c_0_754, c_0_596]), ['final']).
cnf(c_0_1535,negated_conjecture,(ssSkP5(a552)|c2_1(a552)|c4_1(a458)|ssSkC18), inference(spm,[status(thm)],[c_0_755, c_0_41]), ['final']).
cnf(c_0_1536,negated_conjecture,(ssSkP5(a512)|c2_1(a512)|c4_1(a458)|ssSkC9), inference(spm,[status(thm)],[c_0_755, c_0_42]), ['final']).
cnf(c_0_1537,negated_conjecture,(ssSkP5(a558)|c2_1(a558)|c4_1(a458)|~c7_0), inference(spm,[status(thm)],[c_0_755, c_0_297]), ['final']).
cnf(c_0_1538,negated_conjecture,(ssSkP5(a468)|c2_1(a468)|c4_1(a458)|~c4_0), inference(spm,[status(thm)],[c_0_755, c_0_298]), ['final']).
cnf(c_0_1539,negated_conjecture,(c7_2(a530,X1)|c8_1(X2)|c1_1(X2)|~c6_2(a530,X1)|~c4_2(a530,X1)|~c6_1(X2)|~c7_0|~ndr1_1(a530)), inference(spm,[status(thm)],[c_0_756, c_0_596]), ['final']).
cnf(c_0_1540,negated_conjecture,(ssSkC8|~c6_2(X1,a516)|~c4_2(X1,a504)|~ssSkP3(X1)|~ssSkP4(X1)|~ssSkP5(X1)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_757, c_0_656]), c_0_353]), ['final']).
cnf(c_0_1541,negated_conjecture,(ssSkC8|~c4_2(X1,a504)|~c2_2(X1,a516)|~ssSkP3(X1)|~ssSkP4(X1)|~ssSkP5(X1)|~ssSkC9|~ssSkC10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_757, c_0_653]), c_0_353]), ['final']).
cnf(c_0_1542,negated_conjecture,(ssSkC8|ssSkP2(a470)|ssSkC2|~c1_2(a470,a492)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_758, c_0_36]), c_0_37]), ['final']).
cnf(c_0_1543,negated_conjecture,(ssSkC8|ssSkP2(a514)|ssSkC10|~c8_2(a514,a492)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_759, c_0_36]), c_0_299]), ['final']).
cnf(c_0_1544,negated_conjecture,(ssSkC3|c6_2(a581,a479)|c3_2(a581,a456)|c1_2(a581,a456)|ssSkP0(a581)|~c7_0), inference(spm,[status(thm)],[c_0_760, c_0_596]), ['final']).
cnf(c_0_1545,negated_conjecture,(ssSkC8|ssSkC15|~c7_2(a537,X1)|~c3_2(a537,X1)|~c2_2(a537,X1)), inference(spm,[status(thm)],[c_0_300, c_0_353]), ['final']).
cnf(c_0_1546,negated_conjecture,(c3_2(a549,a551)|c1_2(a549,a493)|ssSkC7|~ssSkP2(a549)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_728]), c_0_725]), c_0_761]), ['final']).
cnf(c_0_1547,negated_conjecture,(c8_2(a496,a481)|ssSkC4|c8_1(a496)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1548,negated_conjecture,(ssSkC4|c2_2(a496,a481)|c8_1(a496)|~c4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_661]), c_0_662]), c_0_663]), c_0_664]), ['final']).
cnf(c_0_1549,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c6_2(a574,a575)|~c6_1(X1)|~c7_0|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_762, c_0_596]), ['final']).
cnf(c_0_1550,negated_conjecture,(ssSkC22|ssSkC8|c7_2(a532,a576)|c7_1(a532)|ssSkC13|~c8_2(a532,a576)), inference(spm,[status(thm)],[c_0_763, c_0_706]), ['final']).
cnf(c_0_1551,negated_conjecture,(ssSkC8|ssSkP2(a532)|ssSkC13|~c8_2(a532,a492)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_763, c_0_36]), c_0_37]), ['final']).
cnf(c_0_1552,negated_conjecture,(ssSkC3|c6_2(a542,a479)|c3_2(a542,a543)|c1_2(a542,a543)|ssSkC16|~c8_1(a542)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_54]), c_0_27]), ['final']).
cnf(c_0_1553,negated_conjecture,(ssSkC3|c6_2(a498,a479)|c3_2(a498,a500)|c1_2(a498,a500)|c2_0|~c8_1(a498)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_533]), c_0_525]), ['final']).
cnf(c_0_1554,negated_conjecture,(ssSkC8|c1_2(a587,a456)|ssSkP0(a587)|ssSkC23|~c5_2(a587,a456)), inference(spm,[status(thm)],[c_0_439, c_0_99]), ['final']).
cnf(c_0_1555,negated_conjecture,(ssSkC3|c6_2(a494,a479)|c1_2(a494,a495)|~c8_1(a494)|~ssSkC7|~ssSkC6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_293]), c_0_294]), c_0_295]), ['final']).
cnf(c_0_1556,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c1_2(a574,a575)|~c6_1(X1)|~c7_0|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_764, c_0_596]), ['final']).
cnf(c_0_1557,negated_conjecture,(ssSkC3|c6_2(a549,a479)|c3_2(a549,a551)|c1_2(a549,a551)|~c8_1(a549)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_728]), c_0_725]), ['final']).
cnf(c_0_1558,negated_conjecture,(c4_2(a530,X1)|c8_2(a530,X1)|c7_2(a530,X1)|c8_1(X2)|c1_1(X2)|~c6_1(X2)|~c7_0|~ndr1_1(a530)), inference(spm,[status(thm)],[c_0_765, c_0_596]), ['final']).
cnf(c_0_1559,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(a558)|c7_1(X1)|~c7_0|~ndr1_1(X1)|~c3_1(a558)), inference(spm,[status(thm)],[c_0_351, c_0_297]), ['final']).
cnf(c_0_1560,negated_conjecture,(c6_2(X1,X2)|c4_2(a574,a575)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c6_1(X1)|~c7_0|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_766, c_0_596]), ['final']).
cnf(c_0_1561,negated_conjecture,(c1_1(a508)|~c7_2(a509,a510)|~c3_2(a509,a510)|~c4_0|~c4_1(a509)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_734]), c_0_735]), ['final']).
cnf(c_0_1562,negated_conjecture,(c1_1(a508)|~c7_2(a549,a550)|~c3_2(a549,a550)|~c4_0|~c4_1(a549)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_767]), c_0_725]), ['final']).
cnf(c_0_1563,negated_conjecture,(ssSkC8|~c1_2(a566,a504)|~c2_2(a566,a504)|~ssSkP3(a566)|~ssSkP4(a566)|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_757, c_0_738]), ['final']).
cnf(c_0_1564,negated_conjecture,(c3_2(X1,a456)|ssSkP0(X1)|~c7_2(X1,a548)|~c1_1(X1)|~ssSkC17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_99]), c_0_100]), c_0_121]), ['final']).
cnf(c_0_1565,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c6_1(X1)|~c7_0|~ndr1_1(X1)|~c3_1(a574)), inference(spm,[status(thm)],[c_0_768, c_0_596]), ['final']).
cnf(c_0_1566,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|ndr1_1(a574)|c5_1(X1)|~c6_1(X1)|~c7_0|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_528, c_0_596]), ['final']).
cnf(c_0_1567,negated_conjecture,(c3_2(X1,a460)|c1_2(X1,a493)|ssSkC7|~ssSkC0|~ssSkP2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_729]), c_0_541]), c_0_699]), ['final']).
cnf(c_0_1568,negated_conjecture,(c5_2(a537,X1)|c2_2(a537,a572)|c2_2(a537,X1)|ssSkC15|~ssSkC0), inference(spm,[status(thm)],[c_0_423, c_0_541]), ['final']).
cnf(c_0_1569,negated_conjecture,(c5_2(a517,X1)|c2_2(a517,a572)|c2_2(a517,X1)|ssSkC11|~ssSkC0), inference(spm,[status(thm)],[c_0_424, c_0_541]), ['final']).
cnf(c_0_1570,negated_conjecture,(ssSkC8|c2_2(a537,a504)|c2_2(a537,a572)|ssSkC15|~ssSkP3(a537)|~ssSkP4(a537)), inference(spm,[status(thm)],[c_0_757, c_0_471]), ['final']).
cnf(c_0_1571,negated_conjecture,(ssSkC8|c2_2(a517,a504)|c2_2(a517,a572)|ssSkC11|~ssSkP3(a517)|~ssSkP4(a517)), inference(spm,[status(thm)],[c_0_757, c_0_472]), ['final']).
cnf(c_0_1572,negated_conjecture,(ssSkC12|c4_2(X1,a527)|c3_2(X1,a460)|c2_2(X1,a460)|~ssSkC0|~ssSkP6(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_581]), c_0_541]), ['final']).
cnf(c_0_1573,negated_conjecture,(c5_2(a494,X1)|c2_2(a494,a572)|c2_2(a494,X1)|~ssSkC7|~ssSkC6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_301]), c_0_294]), ['final']).
cnf(c_0_1574,negated_conjecture,(c5_2(a488,X1)|c2_2(a488,a572)|c2_2(a488,X1)|~ssSkC5|~ndr1_1(a488)), inference(spm,[status(thm)],[c_0_395, c_0_769]), ['final']).
cnf(c_0_1575,negated_conjecture,(ssSkC15|~ssSkC0|~c7_2(a537,X1)|~c3_2(a537,X1)|~c2_2(a537,X1)), inference(spm,[status(thm)],[c_0_300, c_0_541]), ['final']).
cnf(c_0_1576,negated_conjecture,(ssSkC22|ssSkC8|c6_2(a552,a576)|c2_2(a552,a576)|c7_1(a552)|ssSkC18|~c1_1(a552)), inference(spm,[status(thm)],[c_0_665, c_0_363]), ['final']).
cnf(c_0_1577,negated_conjecture,(ssSkC8|c6_2(a552,a502)|c2_2(a552,a502)|ssSkP3(a552)|ssSkC18), inference(spm,[status(thm)],[c_0_49, c_0_363]), ['final']).
cnf(c_0_1578,negated_conjecture,(c3_2(a546,a547)|c2_1(a546)|c4_1(a458)|ssSkC17|~c6_1(a546)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_126]), c_0_45]), ['final']).
cnf(c_0_1579,negated_conjecture,(ssSkC8|c2_2(a570,a504)|c2_2(a570,a572)|~ssSkP3(a570)|~ssSkP4(a570)), inference(spm,[status(thm)],[c_0_757, c_0_463]), ['final']).
cnf(c_0_1580,negated_conjecture,(ssSkC8|ssSkP2(a596)|ssSkC25|~c1_2(a596,a492)), inference(spm,[status(thm)],[c_0_770, c_0_36]), ['final']).
cnf(c_0_1581,negated_conjecture,(c5_2(a569,X1)|c2_2(a569,a572)|c2_2(a569,X1)|~c7_0|~ndr1_1(a569)|~ssSkC21), inference(spm,[status(thm)],[c_0_395, c_0_302]), ['final']).
cnf(c_0_1582,negated_conjecture,(c8_2(a581,X1)|c7_2(a581,X1)|c2_2(a581,X1)|~c7_0|~ndr1_1(a581)), inference(spm,[status(thm)],[c_0_303, c_0_596]), ['final']).
cnf(c_0_1583,negated_conjecture,(c7_2(a470,X1)|ssSkC2|~ssSkC0|~c4_2(a470,X1)|~c1_2(a470,X1)), inference(spm,[status(thm)],[c_0_170, c_0_541]), ['final']).
cnf(c_0_1584,negated_conjecture,(c6_2(a514,X1)|ssSkC10|~ssSkC0|~c4_2(a514,X1)|~c8_2(a514,X1)), inference(spm,[status(thm)],[c_0_171, c_0_541]), ['final']).
cnf(c_0_1585,negated_conjecture,(c7_2(a532,X1)|ssSkC13|~ssSkC0|~c4_2(a532,X1)|~c8_2(a532,X1)), inference(spm,[status(thm)],[c_0_174, c_0_541]), ['final']).
cnf(c_0_1586,negated_conjecture,(c1_2(a587,X1)|ssSkC23|~ssSkC0|~c8_2(a587,X1)|~c5_2(a587,X1)), inference(spm,[status(thm)],[c_0_59, c_0_541]), ['final']).
cnf(c_0_1587,negated_conjecture,(ssSkC8|c7_2(a470,a460)|ssSkC2|~ssSkC0), inference(spm,[status(thm)],[c_0_501, c_0_729]), ['final']).
cnf(c_0_1588,negated_conjecture,(ssSkC8|c3_2(a592,a460)|ssSkC24|~ssSkC0), inference(spm,[status(thm)],[c_0_374, c_0_699]), ['final']).
cnf(c_0_1589,negated_conjecture,(ssSkC8|c7_2(a600,a460)|ssSkC26|~ssSkC0), inference(spm,[status(thm)],[c_0_376, c_0_581]), ['final']).
cnf(c_0_1590,negated_conjecture,(c8_2(a564,X1)|c3_2(a564,X1)|ssSkC19|~ssSkC0|~c1_2(a564,X1)), inference(spm,[status(thm)],[c_0_97, c_0_541]), ['final']).
cnf(c_0_1591,negated_conjecture,(c5_2(a536,X1)|c3_2(a536,X1)|ssSkC14|~ssSkC0|~c7_2(a536,X1)), inference(spm,[status(thm)],[c_0_80, c_0_541]), ['final']).
cnf(c_0_1592,negated_conjecture,(c6_2(a592,X1)|c5_2(a592,X1)|ssSkC24|~ssSkC0|~c1_2(a592,X1)), inference(spm,[status(thm)],[c_0_18, c_0_541]), ['final']).
cnf(c_0_1593,negated_conjecture,(ssSkP6(a496)|~c2_2(a496,a526)|~c4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_771, c_0_98]), c_0_21]), ['final']).
cnf(c_0_1594,negated_conjecture,(c3_2(a496,a456)|ssSkP0(a496)|~c2_2(a496,a456)|~c4_0), inference(spm,[status(thm)],[c_0_771, c_0_99]), ['final']).
cnf(c_0_1595,negated_conjecture,(c6_2(a491,X1)|c4_2(a491,X1)|c3_2(a491,X1)|ssSkC6|~ssSkC0), inference(spm,[status(thm)],[c_0_83, c_0_541]), ['final']).
cnf(c_0_1596,negated_conjecture,(c6_2(a552,X1)|c8_2(a552,X1)|c2_2(a552,X1)|ssSkC18|~ssSkC0), inference(spm,[status(thm)],[c_0_17, c_0_541]), ['final']).
cnf(c_0_1597,negated_conjecture,(ssSkC25|~ssSkC0|~c4_2(a596,X1)|~c1_2(a596,X1)), inference(spm,[status(thm)],[c_0_181, c_0_541]), ['final']).
cnf(c_0_1598,negated_conjecture,(c2_2(a482,a526)|ssSkP6(a482)|~c6_2(a482,a526)|~ssSkC4), inference(spm,[status(thm)],[c_0_750, c_0_98]), ['final']).
cnf(c_0_1599,negated_conjecture,(c2_2(a482,a456)|ssSkP0(a482)|~ssSkC4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_750, c_0_99]), c_0_121]), ['final']).
cnf(c_0_1600,negated_conjecture,(c8_2(a546,a492)|ssSkP2(a546)|ssSkC17|~c1_2(a546,a492)), inference(spm,[status(thm)],[c_0_772, c_0_36]), ['final']).
cnf(c_0_1601,negated_conjecture,(c3_2(a496,a460)|~ssSkC0|~c2_2(a496,a460)|~c4_0), inference(spm,[status(thm)],[c_0_771, c_0_729]), ['final']).
cnf(c_0_1602,negated_conjecture,(c2_2(a482,a460)|~ssSkC0|~c6_2(a482,a460)|~ssSkC4), inference(spm,[status(thm)],[c_0_750, c_0_729]), ['final']).
cnf(c_0_1603,negated_conjecture,(ssSkC22|c4_2(a566,a576)|c7_1(a566)|~ssSkC19|~ssSkC20), inference(spm,[status(thm)],[c_0_519, c_0_292]), ['final']).
cnf(c_0_1604,negated_conjecture,(ssSkC22|c4_2(a538,a576)|c7_1(a538)|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_519, c_0_280]), ['final']).
cnf(c_0_1605,negated_conjecture,(ssSkC22|c4_2(a494,a576)|c7_1(a494)|~ssSkC7|~ssSkC6), inference(spm,[status(thm)],[c_0_519, c_0_281]), ['final']).
cnf(c_0_1606,negated_conjecture,(c7_2(a470,X1)|ssSkC2|~ssSkC0|~c8_2(a470,X1)), inference(spm,[status(thm)],[c_0_79, c_0_541]), ['final']).
cnf(c_0_1607,negated_conjecture,(c3_2(a592,X1)|ssSkC24|~ssSkC0|~c5_2(a592,X1)), inference(spm,[status(thm)],[c_0_28, c_0_541]), ['final']).
cnf(c_0_1608,negated_conjecture,(c7_2(a600,X1)|ssSkC26|~ssSkC0|~c1_2(a600,X1)), inference(spm,[status(thm)],[c_0_29, c_0_541]), ['final']).
cnf(c_0_1609,negated_conjecture,(ssSkC0|c7_2(a581,a459)|~c7_0|~c3_1(a581)), inference(spm,[status(thm)],[c_0_773, c_0_596]), ['final']).
cnf(c_0_1610,negated_conjecture,(ndr1_1(a569)|~c8_1(a569)|~c7_0|~ssSkC21), inference(spm,[status(thm)],[c_0_774, c_0_304]), ['final']).
cnf(c_0_1611,negated_conjecture,(ssSkC22|ndr1_1(a538)|c7_1(a538)|~ssSkC14|~ssSkC15), inference(spm,[status(thm)],[c_0_520, c_0_280]), ['final']).
cnf(c_0_1612,negated_conjecture,(ssSkC22|c4_2(a466,a576)|c7_1(a466)|~c7_0), inference(spm,[status(thm)],[c_0_775, c_0_596]), ['final']).
cnf(c_0_1613,negated_conjecture,(ssSkC22|c4_2(a558,a576)|c7_1(a558)|~c7_0), inference(spm,[status(thm)],[c_0_519, c_0_282]), ['final']).
cnf(c_0_1614,negated_conjecture,(c6_2(a568,X1)|c4_2(a568,X1)|ssSkC21|~ssSkC0), inference(spm,[status(thm)],[c_0_84, c_0_541]), ['final']).
cnf(c_0_1615,negated_conjecture,(ssSkC22|ndr1_1(a466)|c7_1(a466)|~c7_0), inference(spm,[status(thm)],[c_0_776, c_0_596]), ['final']).
cnf(c_0_1616,negated_conjecture,(ssSkC22|ndr1_1(a558)|c7_1(a558)|~c7_0), inference(spm,[status(thm)],[c_0_520, c_0_282]), ['final']).
cnf(c_0_1617,negated_conjecture,(ndr1_1(a600)|ssSkC26|~c8_1(a600)), inference(spm,[status(thm)],[c_0_774, c_0_305]), ['final']).
cnf(c_0_1618,negated_conjecture,(ndr1_1(a587)|ssSkC23|~c8_1(a587)), inference(spm,[status(thm)],[c_0_774, c_0_306]), ['final']).
cnf(c_0_1619,negated_conjecture,(ndr1_1(a565)|ssSkC20|~c8_1(a565)), inference(spm,[status(thm)],[c_0_774, c_0_6]), ['final']).
cnf(c_0_1620,negated_conjecture,(ndr1_1(a532)|ssSkC13|~c8_1(a532)), inference(spm,[status(thm)],[c_0_774, c_0_307]), ['final']).
cnf(c_0_1621,negated_conjecture,(~c8_1(a466)|~c7_0), inference(spm,[status(thm)],[c_0_118, c_0_596]), ['final']).
cnf(c_0_1622,negated_conjecture,(~c6_1(a581)|~c7_0), inference(spm,[status(thm)],[c_0_308, c_0_596]), ['final']).
cnf(c_0_1623,negated_conjecture,(c7_2(X1,X2)|c5_2(X3,X4)|c3_2(X1,X2)|c1_2(X5,X6)|c1_2(X1,X2)|c8_1(X3)|c8_1(X5)|c1_1(X1)|c2_1(X1)|~c6_2(X3,X4)|~c6_2(X5,X6)|~c7_2(X3,X4)|~c5_2(X3,a465)|~c2_2(X5,X6)|~c1_1(X5)|~ndr1_1(X3)|~ndr1_1(X5)|~ndr1_1(X1)), c_0_462, ['final']).
cnf(c_0_1624,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X4,X5)|c2_2(X1,X2)|c6_1(X1)|~c4_2(X1,X3)|~c8_2(X1,X2)|~c5_2(X4,a589)|~c2_2(X1,X3)|~c2_2(X4,X5)|~ssSkP8(X4)|~ndr1_1(X1)|~ndr1_1(X4)|~ssSkC23), c_0_777, ['final']).
cnf(c_0_1625,negated_conjecture,(c7_2(X1,X2)|c5_2(X3,X4)|c3_2(X1,X2)|c1_2(X5,X6)|c1_2(X1,X2)|c8_1(X3)|c8_1(X5)|c1_1(X1)|c2_1(X1)|~c6_2(X3,X4)|~c6_2(X5,X6)|~c7_2(X3,a465)|~c7_2(X3,X4)|~c2_2(X5,X6)|~c1_1(X5)|~ndr1_1(X3)|~ndr1_1(X5)|~ndr1_1(X1)), c_0_778, ['final']).
cnf(c_0_1626,negated_conjecture,(c7_2(X1,X2)|c5_2(X3,X4)|c3_2(X1,X2)|c1_2(X5,X6)|c1_2(X1,X2)|c8_1(X3)|c8_1(X5)|c1_1(X1)|c2_1(X1)|~c6_2(X3,X4)|~c6_2(X5,X6)|~c7_2(X3,X4)|~c1_2(X3,a465)|~c2_2(X5,X6)|~c1_1(X5)|~ndr1_1(X3)|~ndr1_1(X5)|~ndr1_1(X1)), c_0_458, ['final']).
cnf(c_0_1627,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X4,X5)|c2_2(X1,X2)|c6_1(X1)|~c4_2(X1,X3)|~c8_2(X1,X2)|~c7_2(X4,a589)|~c2_2(X1,X3)|~c2_2(X4,X5)|~ssSkP8(X4)|~ndr1_1(X1)|~ndr1_1(X4)|~ssSkC23), c_0_779, ['final']).
cnf(c_0_1628,negated_conjecture,(c6_2(X1,X2)|c8_2(X3,a589)|c8_2(X1,X4)|c5_2(X3,X5)|c2_2(X1,X2)|c6_1(X1)|~c4_2(X1,X4)|~c8_2(X1,X2)|~c2_2(X1,X4)|~c2_2(X3,X5)|~ssSkP8(X3)|~ndr1_1(X1)|~ndr1_1(X3)|~ssSkC23), c_0_780, ['final']).
cnf(c_0_1629,negated_conjecture,(c6_2(X1,a474)|c4_2(X2,X3)|c4_2(X4,X5)|c1_2(X2,X3)|c8_1(X4)|c6_1(X2)|~c6_2(X2,X3)|~c6_2(X4,X5)|~c3_2(X4,X5)|~ndr1_1(X2)|~ndr1_1(X4)|~c3_1(X4)|~c5_1(X1)|~c7_1(X2)), c_0_482, ['final']).
cnf(c_0_1630,negated_conjecture,(c4_2(X1,X2)|c4_2(X3,X4)|c7_2(X5,a474)|c1_2(X1,X2)|c8_1(X3)|c6_1(X1)|~c6_2(X1,X2)|~c6_2(X3,X4)|~c3_2(X3,X4)|~ndr1_1(X1)|~ndr1_1(X3)|~c3_1(X3)|~c5_1(X5)|~c7_1(X1)), c_0_485, ['final']).
cnf(c_0_1631,negated_conjecture,(c4_2(X1,X2)|c4_2(X3,X4)|c1_2(X5,a474)|c1_2(X1,X2)|c8_1(X3)|c6_1(X1)|~c6_2(X1,X2)|~c6_2(X3,X4)|~c3_2(X3,X4)|~ndr1_1(X1)|~ndr1_1(X3)|~c3_1(X3)|~c5_1(X5)|~c7_1(X1)), c_0_487, ['final']).
cnf(c_0_1632,negated_conjecture,(c4_2(X1,X2)|c4_2(X3,X4)|c1_2(X1,X2)|c8_1(X3)|c6_1(X1)|ndr1_1(X5)|~c6_2(X1,X2)|~c6_2(X3,X4)|~c3_2(X3,X4)|~ndr1_1(X1)|~ndr1_1(X3)|~c3_1(X3)|~c5_1(X5)|~c7_1(X1)), c_0_489, ['final']).
cnf(c_0_1633,negated_conjecture,(c4_2(X1,X2)|c5_2(a573,X3)|c5_2(X1,X2)|c1_2(X1,X4)|c2_2(X1,X4)|c8_1(X1)|~c6_2(a573,X3)|~c7_2(X1,X2)|~c3_2(X1,X4)|~c2_2(a573,X3)|~ndr1_1(a573)|~ndr1_1(X1)), c_0_430, ['final']).
cnf(c_0_1634,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X3)|c3_2(X1,X4)|c2_2(X1,X2)|~ssSkC22|~c6_2(a577,a579)|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), c_0_781, ['final']).
cnf(c_0_1635,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X3)|c3_2(X1,X4)|c2_2(X1,X2)|~ssSkC22|~c8_2(a577,a579)|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), c_0_782, ['final']).
cnf(c_0_1636,negated_conjecture,(c6_2(a577,a578)|c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X3)|c3_2(X1,X4)|c2_2(X1,X2)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), c_0_783, ['final']).
cnf(c_0_1637,negated_conjecture,(c6_2(X1,X2)|c8_2(a577,a578)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X3)|c3_2(X1,X4)|c2_2(X1,X2)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), c_0_784, ['final']).
cnf(c_0_1638,negated_conjecture,(c1_2(X1,X2)|c2_2(X1,X3)|c2_2(X1,X2)|c1_1(X4)|c2_1(X4)|c2_0|~c7_2(X1,X3)|~c5_2(X1,X2)|~c1_2(X1,X3)|~c4_1(X4)|~ndr1_1(X1)|~c5_1(X1)), c_0_624, ['final']).
cnf(c_0_1639,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X3)|c3_2(X1,X4)|c1_2(a577,a578)|c2_2(X1,X2)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), c_0_785, ['final']).
cnf(c_0_1640,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X3)|c3_2(X1,X4)|c2_2(X1,X2)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)|~c7_1(a577)), c_0_786, ['final']).
cnf(c_0_1641,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c5_2(X1,X4)|c3_2(X1,X3)|c3_2(X1,X4)|c2_2(X1,X2)|ndr1_1(a577)|~ssSkC22|~c1_2(X1,X4)|~c2_2(X1,X3)|~ndr1_1(X1)), c_0_787, ['final']).
cnf(c_0_1642,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|~c6_2(a534,a535)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_788, c_0_360])]), ['final']).
cnf(c_0_1643,negated_conjecture,(c7_2(X1,X2)|c1_2(X1,X2)|c6_1(X1)|~c6_2(X1,X3)|~c3_2(X1,X3)|~c1_2(X1,X3)|~c2_2(X1,X2)|~c2_1(a534)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_789, c_0_360])]), ['final']).
cnf(c_0_1644,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c3_2(X1,X3)|c2_2(X1,X2)|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_1(X1)|~c6_1(a559)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_790, c_0_360])]), ['final']).
cnf(c_0_1645,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X3)|c3_2(X1,X2)|c5_1(X1)|~c8_2(X1,X3)|~c7_2(X1,X2)|~c1_2(X1,X3)|~c6_1(a531)|~ndr1_1(X1)), c_0_791, ['final']).
cnf(c_0_1646,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X3)|c3_2(X1,X2)|c5_1(X1)|~c8_2(X1,X3)|~c7_2(X1,X2)|~c1_2(X1,X3)|~ndr1_1(X1)|~c5_1(a531)), c_0_792, ['final']).
cnf(c_0_1647,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X3)|c3_2(X1,X2)|c4_1(a531)|c5_1(X1)|~c8_2(X1,X3)|~c7_2(X1,X2)|~c1_2(X1,X3)|~ndr1_1(X1)), c_0_356, ['final']).
cnf(c_0_1648,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c4_2(a574,X3)|~c3_2(a574,X3)|~c2_2(a574,X3)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), c_0_733, ['final']).
cnf(c_0_1649,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|~c6_2(X3,a525)|~c6_2(X1,X2)|~c4_2(X1,X2)|~c4_1(X1)|~ndr1_1(X1)|~c2_0|~c3_1(X3)), c_0_744, ['final']).
cnf(c_0_1650,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~c2_2(X3,a525)|~c4_1(X1)|~ndr1_1(X1)|~c2_0|~c3_1(X3)), c_0_745, ['final']).
cnf(c_0_1651,negated_conjecture,(c7_2(X1,a525)|c3_2(X2,X3)|c8_1(X2)|c6_1(X1)|~c6_2(X2,X3)|~c4_2(X2,X3)|~c4_1(X2)|~ndr1_1(X2)|~c2_0|~c3_1(X1)), c_0_747, ['final']).
cnf(c_0_1652,negated_conjecture,(c8_2(X1,a467)|c8_2(X2,X3)|c8_1(X2)|~c4_2(X2,X3)|~c7_2(X2,X3)|~c8_1(X1)|~c2_1(X2)|~c6_1(X1)|~ndr1_1(X2)), c_0_583, ['final']).
cnf(c_0_1653,negated_conjecture,(c8_2(X1,X2)|c3_2(X3,a467)|c8_1(X1)|~c4_2(X1,X2)|~c7_2(X1,X2)|~c8_1(X3)|~c2_1(X1)|~c6_1(X3)|~ndr1_1(X1)), c_0_585, ['final']).
cnf(c_0_1654,negated_conjecture,(c8_2(X1,X2)|c2_2(X3,a467)|c8_1(X1)|~c4_2(X1,X2)|~c7_2(X1,X2)|~c8_1(X3)|~c2_1(X1)|~c6_1(X3)|~ndr1_1(X1)), c_0_586, ['final']).
cnf(c_0_1655,negated_conjecture,(c3_2(X1,X2)|c8_1(X1)|c6_1(X3)|ndr1_1(X3)|~c6_2(X1,X2)|~c4_2(X1,X2)|~c4_1(X1)|~ndr1_1(X1)|~c2_0|~c3_1(X3)), c_0_751, ['final']).
cnf(c_0_1656,negated_conjecture,(c4_2(X1,X2)|c5_2(X1,X2)|c1_2(X1,X3)|c2_2(X1,X3)|c8_1(X1)|c1_1(a573)|~c7_2(X1,X2)|~c3_2(X1,X3)|~ndr1_1(X1)), c_0_391, ['final']).
cnf(c_0_1657,negated_conjecture,(c4_2(X1,X2)|c5_2(X1,X2)|c1_2(X1,X3)|c2_2(X1,X3)|c8_1(X1)|c3_1(a573)|~c7_2(X1,X2)|~c3_2(X1,X3)|~ndr1_1(X1)), c_0_393, ['final']).
cnf(c_0_1658,negated_conjecture,(c8_2(X1,X2)|c8_1(X1)|ndr1_1(X3)|~c4_2(X1,X2)|~c7_2(X1,X2)|~c8_1(X3)|~c2_1(X1)|~c6_1(X3)|~ndr1_1(X1)), c_0_712, ['final']).
cnf(c_0_1659,negated_conjecture,(c7_2(X1,X2)|c4_1(X1)|c5_1(X3)|c7_1(X3)|c7_1(X1)|~c6_2(X1,X2)|~c4_2(X3,a583)|~c2_2(X1,X2)|~ndr1_1(X1)), c_0_510, ['final']).
cnf(c_0_1660,negated_conjecture,(c7_2(X1,a583)|c7_2(X2,X3)|c4_1(X2)|c5_1(X1)|c7_1(X1)|c7_1(X2)|~c6_2(X2,X3)|~c2_2(X2,X3)|~ndr1_1(X2)), c_0_695, ['final']).
cnf(c_0_1661,negated_conjecture,(c7_2(X1,X2)|c2_2(X3,a583)|c4_1(X1)|c5_1(X3)|c7_1(X3)|c7_1(X1)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)), c_0_696, ['final']).
cnf(c_0_1662,negated_conjecture,(c4_2(X1,a473)|~c4_2(X2,X3)|~c8_2(X2,X3)|~c1_2(X2,X3)|~c7_0|~ndr1_1(X2)|~c3_1(X1)), c_0_793, ['final']).
cnf(c_0_1663,negated_conjecture,(c3_2(X1,X2)|~c6_2(X1,X2)|~c4_2(X1,a548)|~c8_2(X1,X2)|~c1_1(X1)|~ndr1_1(X1)|~ssSkC17), c_0_603, ['final']).
cnf(c_0_1664,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,a557)|c2_2(X1,X2)|c5_1(X3)|c7_1(X1)|~c6_1(X3)|~ndr1_1(X1)|~c3_1(X3)), c_0_351, ['final']).
cnf(c_0_1665,negated_conjecture,(c8_2(X1,a473)|~c4_2(X2,X3)|~c8_2(X2,X3)|~c1_2(X2,X3)|~c7_0|~ndr1_1(X2)|~c3_1(X1)), c_0_794, ['final']).
cnf(c_0_1666,negated_conjecture,(c8_2(X1,X2)|c2_2(X1,X2)|c5_1(X3)|c7_1(X1)|~c6_2(X1,a557)|~c6_1(X3)|~ndr1_1(X1)|~c3_1(X3)), c_0_795, ['final']).
cnf(c_0_1667,negated_conjecture,(c8_2(X1,X2)|c2_2(X1,X2)|c5_1(X3)|c7_1(X1)|~c2_2(X1,a557)|~c6_1(X3)|~ndr1_1(X1)|~c3_1(X3)), c_0_796, ['final']).
cnf(c_0_1668,negated_conjecture,(c7_2(X1,X2)|c4_1(X1)|ndr1_1(X3)|c5_1(X3)|c7_1(X3)|c7_1(X1)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)), c_0_702, ['final']).
cnf(c_0_1669,negated_conjecture,(c3_2(X1,X2)|c1_2(X1,X2)|c8_1(X1)|c8_1(X3)|c4_1(X1)|c3_1(X3)|~c8_2(X1,X2)|~c1_1(X3)|~ndr1_1(X1)), c_0_400, ['final']).
cnf(c_0_1670,negated_conjecture,(c3_2(X1,X2)|~c6_2(X1,X2)|~c8_2(X1,X2)|~c7_2(X1,a548)|~c1_1(X1)|~ndr1_1(X1)|~ssSkC17), c_0_473, ['final']).
cnf(c_0_1671,negated_conjecture,(c6_2(X1,X2)|ssSkC1|~c6_2(X1,a462)|~c4_2(X1,X2)|~c2_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), c_0_797, ['final']).
cnf(c_0_1672,negated_conjecture,(c6_2(X1,X2)|ssSkC1|~c4_2(X1,X2)|~c7_2(X1,a462)|~c2_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), c_0_798, ['final']).
cnf(c_0_1673,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X2)|c7_0|~c4_2(X1,a593)|~c7_2(X1,X2)|~c6_1(X1)|~ndr1_1(X1)|~ssSkC24), c_0_799, ['final']).
cnf(c_0_1674,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X2)|c7_0|~c7_2(X1,a593)|~c7_2(X1,X2)|~c6_1(X1)|~ndr1_1(X1)|~ssSkC24), c_0_800, ['final']).
cnf(c_0_1675,negated_conjecture,(ndr1_1(X1)|~c4_2(X2,X3)|~c8_2(X2,X3)|~c1_2(X2,X3)|~c7_0|~ndr1_1(X2)|~c3_1(X1)), c_0_801, ['final']).
cnf(c_0_1676,negated_conjecture,(c8_2(X1,X2)|c3_2(X1,X2)|c2_1(X1)|~ssSkC12|~c7_2(X1,X2)|~c5_2(X1,a528)|~c4_0|~ndr1_1(X1)), c_0_802, ['final']).
cnf(c_0_1677,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c6_2(a574,a575)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), c_0_762, ['final']).
cnf(c_0_1678,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c1_2(a574,a575)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), c_0_764, ['final']).
cnf(c_0_1679,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,a593)|c2_2(X1,X2)|c7_0|~c7_2(X1,X2)|~c6_1(X1)|~ndr1_1(X1)|~ssSkC24), c_0_803, ['final']).
cnf(c_0_1680,negated_conjecture,(ssSkC12|c3_2(X1,X2)|c2_2(X1,X2)|~c1_2(X1,a527)|~c1_2(X1,X2)|~ssSkP6(X1)|~ndr1_1(X1)), c_0_698, ['final']).
cnf(c_0_1681,negated_conjecture,(c6_2(X1,X2)|c4_2(a574,a575)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), c_0_766, ['final']).
cnf(c_0_1682,negated_conjecture,(c6_2(X1,X2)|ssSkC1|~c4_2(X1,X2)|~c1_2(X1,a462)|~c2_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), c_0_697, ['final']).
cnf(c_0_1683,negated_conjecture,(c3_2(X1,X2)|ssSkC7|~c6_2(X1,a493)|~c8_2(X1,X2)|~c5_2(X1,X2)|~ssSkP2(X1)|~ndr1_1(X1)), c_0_804, ['final']).
cnf(c_0_1684,negated_conjecture,(c3_2(X1,X2)|ssSkC7|~c8_2(X1,X2)|~c5_2(X1,X2)|~c2_2(X1,a493)|~ssSkP2(X1)|~ndr1_1(X1)), c_0_805, ['final']).
cnf(c_0_1685,negated_conjecture,(c7_2(X1,X2)|c2_1(X1)|c4_1(X1)|~c6_2(X1,X2)|~c2_2(X1,X2)|~ndr1_1(X1)|~c3_1(a533)|~ssSkC13), c_0_806, ['final']).
cnf(c_0_1686,negated_conjecture,(c3_2(X1,X2)|c2_2(X1,a548)|~c6_2(X1,X2)|~c8_2(X1,X2)|~c1_1(X1)|~ndr1_1(X1)|~ssSkC17), c_0_807, ['final']).
cnf(c_0_1687,negated_conjecture,(ssSkC12|c3_2(X1,X2)|c2_2(X1,X2)|~c1_2(X1,X2)|~c2_2(X1,a527)|~ssSkP6(X1)|~ndr1_1(X1)), c_0_808, ['final']).
cnf(c_0_1688,negated_conjecture,(c3_2(X1,X2)|c1_2(X1,a493)|ssSkC7|~c8_2(X1,X2)|~c5_2(X1,X2)|~ssSkP2(X1)|~ndr1_1(X1)), c_0_387, ['final']).
cnf(c_0_1689,negated_conjecture,(c6_2(X1,a528)|c8_2(X1,X2)|c3_2(X1,X2)|c2_1(X1)|~ssSkC12|~c7_2(X1,X2)|~c4_0|~ndr1_1(X1)), c_0_809, ['final']).
cnf(c_0_1690,negated_conjecture,(c8_2(X1,X2)|c3_2(X1,X2)|c2_2(X1,a528)|c2_1(X1)|~ssSkC12|~c7_2(X1,X2)|~c4_0|~ndr1_1(X1)), c_0_810, ['final']).
cnf(c_0_1691,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|c5_1(X1)|~c6_1(X1)|~ndr1_1(X1)|~c2_0|~c3_1(a574)), c_0_768, ['final']).
cnf(c_0_1692,negated_conjecture,(c6_2(X1,X2)|c7_2(X1,X2)|c1_2(X1,X2)|c1_1(X1)|c7_0|~c8_1(a501)|~c2_1(X1)|~ndr1_1(X1)), c_0_811, ['final']).
cnf(c_0_1693,negated_conjecture,(c6_2(X1,X2)|c7_2(X1,X2)|c1_2(X1,X2)|c1_1(X1)|c7_0|~c2_1(X1)|~c6_1(a501)|~ndr1_1(X1)), c_0_812, ['final']).
cnf(c_0_1694,negated_conjecture,(ssSkC12|c4_2(X1,a527)|c3_2(X1,X2)|c2_2(X1,X2)|~c1_2(X1,X2)|~ssSkP6(X1)|~ndr1_1(X1)), c_0_377, ['final']).
cnf(c_0_1695,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X2)|c7_2(X1,X2)|ndr1_1(a574)|c5_1(X1)|~c6_1(X1)|~ndr1_1(X1)|~c2_0), c_0_528, ['final']).
cnf(c_0_1696,negated_conjecture,(c8_1(X1)|c1_1(X1)|~c6_2(a530,X2)|~c8_2(a530,X2)|~c3_2(a530,X2)|~c6_1(X1)|~ndr1_1(a530)|~c2_0), c_0_753, ['final']).
cnf(c_0_1697,negated_conjecture,(c7_2(a530,X1)|c8_1(X2)|c1_1(X2)|~c6_2(a530,X1)|~c4_2(a530,X1)|~c6_1(X2)|~ndr1_1(a530)|~c2_0), c_0_756, ['final']).
cnf(c_0_1698,negated_conjecture,(c4_2(a530,X1)|c8_2(a530,X1)|c7_2(a530,X1)|c8_1(X2)|c1_1(X2)|~c6_1(X2)|~ndr1_1(a530)|~c2_0), c_0_765, ['final']).
cnf(c_0_1699,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|~c7_2(X1,X2)|~c2_1(a458)|~c6_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_813, c_0_360])]), ['final']).
cnf(c_0_1700,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|~c7_2(X1,X2)|~c6_1(a458)|~c6_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_814, c_0_360])]), ['final']).
cnf(c_0_1701,negated_conjecture,(ssSkC3|c3_2(X1,X2)|c1_2(X1,X2)|~c8_2(X1,X2)|~c7_2(X1,a479)|~c8_1(X1)|~ndr1_1(X1)), c_0_815, ['final']).
cnf(c_0_1702,negated_conjecture,(c1_1(a508)|~c4_2(X1,X2)|~c7_2(X1,X2)|~c3_2(X1,X2)|~c4_1(X1)|~ndr1_1(X1)), c_0_629, ['final']).
cnf(c_0_1703,negated_conjecture,(c6_2(X1,X2)|c7_2(X1,X2)|c1_2(X1,X2)|c1_1(X1)|c7_0|c5_1(a501)|~c2_1(X1)|~ndr1_1(X1)), c_0_358, ['final']).
cnf(c_0_1704,negated_conjecture,(c4_2(X1,X2)|c8_2(X1,a481)|ssSkC4|c8_1(X1)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)), c_0_730, ['final']).
cnf(c_0_1705,negated_conjecture,(c4_2(X1,X2)|ssSkC4|c2_2(X1,a481)|c8_1(X1)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)), c_0_731, ['final']).
cnf(c_0_1706,negated_conjecture,(c5_2(X1,X2)|~c6_2(X1,a516)|~c4_2(X1,X2)|~ssSkP5(X1)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), c_0_656, ['final']).
cnf(c_0_1707,negated_conjecture,(c5_2(X1,X2)|c2_2(X1,X2)|~c5_2(X1,a572)|~c4_1(X1)|~ndr1_1(X1)), c_0_752, ['final']).
cnf(c_0_1708,negated_conjecture,(ssSkC15|~c7_2(a537,X1)|~c3_2(a537,X1)|~c2_2(a537,X1)|~ndr1_1(a537)), c_0_300, ['final']).
cnf(c_0_1709,negated_conjecture,(c8_2(a538,X1)|c7_2(a538,X1)|~c1_2(a538,X1)|~ndr1_1(a538)|~ssSkC14|~ssSkC15), c_0_309, ['final']).
cnf(c_0_1710,negated_conjecture,(c6_2(a569,X1)|c3_2(a569,X1)|~c4_2(a569,X1)|~c7_0|~ndr1_1(a569)|~ssSkC21), c_0_310, ['final']).
cnf(c_0_1711,negated_conjecture,(c3_2(a538,X1)|c2_2(a538,X1)|~c7_2(a538,X1)|~ndr1_1(a538)|~ssSkC14|~ssSkC15), c_0_311, ['final']).
cnf(c_0_1712,negated_conjecture,(c7_2(a470,X1)|ssSkC2|~c4_2(a470,X1)|~c1_2(a470,X1)|~ndr1_1(a470)), c_0_170, ['final']).
cnf(c_0_1713,negated_conjecture,(c6_2(a514,X1)|ssSkC10|~c4_2(a514,X1)|~c8_2(a514,X1)|~ndr1_1(a514)), c_0_171, ['final']).
cnf(c_0_1714,negated_conjecture,(c7_2(a532,X1)|ssSkC13|~c4_2(a532,X1)|~c8_2(a532,X1)|~ndr1_1(a532)), c_0_174, ['final']).
cnf(c_0_1715,negated_conjecture,(c1_2(a587,X1)|ssSkC23|~c8_2(a587,X1)|~c5_2(a587,X1)|~ndr1_1(a587)), c_0_59, ['final']).
cnf(c_0_1716,negated_conjecture,(c4_2(a544,X1)|c1_2(a544,X1)|~c8_1(X2)|~c2_1(X2)|~ssSkC16), c_0_816, ['final']).
cnf(c_0_1717,negated_conjecture,(c5_2(X1,X2)|~c4_2(X1,X2)|~c2_2(X1,a516)|~ssSkP5(X1)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), c_0_653, ['final']).
cnf(c_0_1718,negated_conjecture,(c8_2(a564,X1)|c3_2(a564,X1)|ssSkC19|~c1_2(a564,X1)|~ndr1_1(a564)), c_0_97, ['final']).
cnf(c_0_1719,negated_conjecture,(ssSkC3|c6_2(X1,a479)|c3_2(X1,X2)|c1_2(X1,X2)|~c8_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), c_0_402, ['final']).
cnf(c_0_1720,negated_conjecture,(c5_2(a488,X1)|c2_2(a488,X1)|~c6_2(a488,X1)|~ssSkC5|~ndr1_1(a488)), c_0_817, ['final']).
cnf(c_0_1721,negated_conjecture,(epred1_0|~c4_2(X1,X2)|~c8_2(X1,X2)|~c1_2(X1,X2)|~c7_0|~ndr1_1(X1)), c_0_818, ['final']).
cnf(c_0_1722,negated_conjecture,(c5_2(a536,X1)|c3_2(a536,X1)|ssSkC14|~c7_2(a536,X1)|~ndr1_1(a536)), c_0_80, ['final']).
cnf(c_0_1723,negated_conjecture,(c6_2(a592,X1)|c5_2(a592,X1)|ssSkC24|~c1_2(a592,X1)|~ndr1_1(a592)), c_0_18, ['final']).
cnf(c_0_1724,negated_conjecture,(c5_2(a566,X1)|~c1_2(a566,X1)|~c2_2(a566,X1)|~ssSkC19|~ssSkC20), c_0_738, ['final']).
cnf(c_0_1725,negated_conjecture,(c8_2(a581,X1)|c7_2(a581,X1)|c2_2(a581,X1)|~ndr1_1(a581)|~c2_0), c_0_303, ['final']).
cnf(c_0_1726,negated_conjecture,(c6_2(a491,X1)|c4_2(a491,X1)|c3_2(a491,X1)|ssSkC6|~ndr1_1(a491)), c_0_83, ['final']).
cnf(c_0_1727,negated_conjecture,(c3_2(a496,X1)|~c8_2(a496,X1)|~c2_2(a496,X1)|~c4_0), c_0_771, ['final']).
cnf(c_0_1728,negated_conjecture,(c2_2(a482,X1)|~c6_2(a482,X1)|~c8_2(a482,X1)|~ssSkC4), c_0_750, ['final']).
cnf(c_0_1729,negated_conjecture,(c4_2(a475,X1)|c5_2(a475,X1)|c7_0|~c8_2(a475,X1)|~c2_0), c_0_819, ['final']).
cnf(c_0_1730,negated_conjecture,(c8_2(a546,X1)|ssSkC17|~c4_2(a546,X1)|~c1_2(a546,X1)), c_0_772, ['final']).
cnf(c_0_1731,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|~c5_2(a539,a540)|~c7_0), c_0_820, ['final']).
cnf(c_0_1732,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|~c5_2(a539,a541)|~c7_0), c_0_821, ['final']).
cnf(c_0_1733,negated_conjecture,(c6_2(a552,X1)|c8_2(a552,X1)|c2_2(a552,X1)|ssSkC18|~ndr1_1(a552)), c_0_17, ['final']).
cnf(c_0_1734,negated_conjecture,(c1_1(X1)|c6_1(X1)|c5_1(X1)|~c2_2(a539,a541)|~c7_0), c_0_822, ['final']).
cnf(c_0_1735,negated_conjecture,(c7_0|~c3_2(X1,a553)|~c4_1(X1)|~c5_1(X1)|~ssSkC18), c_0_823, ['final']).
cnf(c_0_1736,negated_conjecture,(c7_0|~c2_2(X1,a553)|~c4_1(X1)|~c5_1(X1)|~ssSkC18), c_0_824, ['final']).
cnf(c_0_1737,negated_conjecture,(c5_2(X1,X2)|c2_2(X1,X2)|~c6_2(X1,a572)|~c4_1(X1)|~ndr1_1(X1)), c_0_825, ['final']).
cnf(c_0_1738,negated_conjecture,(c5_2(X1,X2)|c2_2(X1,a572)|c2_2(X1,X2)|~c4_1(X1)|~ndr1_1(X1)), c_0_395, ['final']).
cnf(c_0_1739,negated_conjecture,(ssSkC0|~c6_2(X1,a459)|~c8_1(X1)|~c3_1(X1)), c_0_826, ['final']).
cnf(c_0_1740,negated_conjecture,(ssSkC0|~c3_2(X1,a459)|~c8_1(X1)|~c3_1(X1)), c_0_827, ['final']).
cnf(c_0_1741,negated_conjecture,(ssSkC8|~c5_2(X1,a504)|~ssSkP3(X1)|~ssSkP4(X1)), c_0_757, ['final']).
cnf(c_0_1742,negated_conjecture,(ssSkC8|~c3_2(X1,a504)|~ssSkP3(X1)|~ssSkP4(X1)), c_0_828, ['final']).
cnf(c_0_1743,negated_conjecture,(c1_1(X1)|c3_1(X1)|~c5_2(X1,a582)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_829, c_0_360])]), ['final']).
cnf(c_0_1744,negated_conjecture,(ssSkC25|~c4_2(a596,X1)|~c1_2(a596,X1)|~ndr1_1(a596)), c_0_181, ['final']).
cnf(c_0_1745,negated_conjecture,(c8_2(a539,a540)|c1_1(X1)|c6_1(X1)|c5_1(X1)|~c7_0), c_0_830, ['final']).
cnf(c_0_1746,negated_conjecture,(c8_2(a539,a541)|c1_1(X1)|c6_1(X1)|c5_1(X1)|~c7_0), c_0_831, ['final']).
cnf(c_0_1747,negated_conjecture,(c3_1(X1)|~ssSkC3|~c6_2(X1,a480)|~c2_1(X1)), c_0_832, ['final']).
cnf(c_0_1748,negated_conjecture,(c7_2(a470,X1)|ssSkC2|~c8_2(a470,X1)|~ndr1_1(a470)), c_0_79, ['final']).
cnf(c_0_1749,negated_conjecture,(c3_2(a592,X1)|ssSkC24|~c5_2(a592,X1)|~ndr1_1(a592)), c_0_28, ['final']).
cnf(c_0_1750,negated_conjecture,(c7_2(a600,X1)|ssSkC26|~c1_2(a600,X1)|~ndr1_1(a600)), c_0_29, ['final']).
cnf(c_0_1751,negated_conjecture,(c3_2(a539,a540)|c1_1(X1)|c6_1(X1)|c5_1(X1)|~c7_0), c_0_833, ['final']).
cnf(c_0_1752,negated_conjecture,(ssSkC22|c7_1(X1)|~c8_2(X1,a576)|~c1_1(X1)), c_0_665, ['final']).
cnf(c_0_1753,negated_conjecture,(c2_1(X1)|c7_1(X1)|~c7_2(X1,a602)|~c7_0), c_0_834, ['final']).
cnf(c_0_1754,negated_conjecture,(c7_2(X1,a553)|c7_0|~c4_1(X1)|~c5_1(X1)|~ssSkC18), c_0_835, ['final']).
cnf(c_0_1755,negated_conjecture,(c3_2(X1,a582)|c1_1(X1)|c3_1(X1)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_836, c_0_360])]), ['final']).
cnf(c_0_1756,negated_conjecture,(c2_2(X1,a582)|c1_1(X1)|c3_1(X1)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_837, c_0_360])]), ['final']).
cnf(c_0_1757,negated_conjecture,(ssSkC0|c7_2(X1,a459)|~c8_1(X1)|~c3_1(X1)), c_0_539, ['final']).
cnf(c_0_1758,negated_conjecture,(ssSkC8|c8_2(X1,a504)|~ssSkP3(X1)|~ssSkP4(X1)), c_0_838, ['final']).
cnf(c_0_1759,negated_conjecture,(c3_2(X1,a480)|c3_1(X1)|~ssSkC3|~c2_1(X1)), c_0_839, ['final']).
cnf(c_0_1760,negated_conjecture,(c4_2(X1,a480)|c3_1(X1)|~ssSkC3|~c2_1(X1)), c_0_840, ['final']).
cnf(c_0_1761,negated_conjecture,(c4_2(a561,X1)|c8_2(a561,X1)|~c4_0|~ndr1_1(a561)), c_0_312, ['final']).
cnf(c_0_1762,negated_conjecture,(ssSkC22|c4_2(X1,a576)|c7_1(X1)|~c1_1(X1)), c_0_519, ['final']).
cnf(c_0_1763,negated_conjecture,(c4_2(X1,a487)|ssSkC5|c5_1(X1)|~c2_1(X1)), c_0_371, ['final']).
cnf(c_0_1764,negated_conjecture,(c6_2(a568,X1)|c4_2(a568,X1)|ssSkC21|~ndr1_1(a568)), c_0_84, ['final']).
cnf(c_0_1765,negated_conjecture,(c4_2(X1,a602)|c2_1(X1)|c7_1(X1)|~c7_0), c_0_841, ['final']).
cnf(c_0_1766,negated_conjecture,(c3_2(X1,a602)|c2_1(X1)|c7_1(X1)|~c7_0), c_0_842, ['final']).
cnf(c_0_1767,negated_conjecture,(c7_2(a542,X1)|ssSkC16|~c4_2(a542,X1)), c_0_372, ['final']).
cnf(c_0_1768,negated_conjecture,(c1_2(X1,a487)|ssSkC5|c5_1(X1)|~c2_1(X1)), c_0_469, ['final']).
cnf(c_0_1769,negated_conjecture,(epred2_0|~c3_2(X1,a473)|~c3_1(X1)), c_0_843, ['final']).
cnf(c_0_1770,negated_conjecture,(c7_0|ndr1_1(X1)|~c4_1(X1)|~c5_1(X1)|~ssSkC18), c_0_844, ['final']).
cnf(c_0_1771,negated_conjecture,(ssSkP0(X1)|~c4_2(X1,a456)), c_0_122, ['final']).
cnf(c_0_1772,negated_conjecture,(ssSkP1(X1)|~c3_2(X1,a485)), c_0_64, ['final']).
cnf(c_0_1773,negated_conjecture,(ssSkP1(X1)|~c1_2(X1,a485)), c_0_47, ['final']).
cnf(c_0_1774,negated_conjecture,(ssSkP1(X1)|~c2_2(X1,a485)), c_0_313, ['final']).
cnf(c_0_1775,negated_conjecture,(ssSkP2(X1)|~c6_2(X1,a492)), c_0_299, ['final']).
cnf(c_0_1776,negated_conjecture,(ssSkP2(X1)|~c7_2(X1,a492)), c_0_37, ['final']).
cnf(c_0_1777,negated_conjecture,(ssSkP3(X1)|~c8_2(X1,a502)), c_0_49, ['final']).
cnf(c_0_1778,negated_conjecture,(ssSkP3(X1)|~c7_2(X1,a502)), c_0_125, ['final']).
cnf(c_0_1779,negated_conjecture,(ssSkP3(X1)|~c1_2(X1,a502)), c_0_124, ['final']).
cnf(c_0_1780,negated_conjecture,(c7_0|~c3_2(a475,a476)|~c2_0), c_0_314, ['final']).
cnf(c_0_1781,negated_conjecture,(c6_2(a505,a506)|~ssSkC8|~c4_0), c_0_285, ['final']).
cnf(c_0_1782,negated_conjecture,(c6_2(a566,a567)|~ssSkC19|~ssSkC20), c_0_284, ['final']).
cnf(c_0_1783,negated_conjecture,(ssSkP4(X1)|~c4_2(X1,a503)), c_0_140, ['final']).
cnf(c_0_1784,negated_conjecture,(ssSkP4(X1)|~c3_2(X1,a503)), c_0_63, ['final']).
cnf(c_0_1785,negated_conjecture,(ssSkP5(X1)|~c3_2(X1,a515)), c_0_24, ['final']).
cnf(c_0_1786,negated_conjecture,(ssSkP5(X1)|~c2_2(X1,a515)), c_0_277, ['final']).
cnf(c_0_1787,negated_conjecture,(ssSkP6(X1)|~c3_2(X1,a526)), c_0_21, ['final']).
cnf(c_0_1788,negated_conjecture,(ssSkP7(X1)|~c6_2(X1,a554)), c_0_315, ['final']).
cnf(c_0_1789,negated_conjecture,(ssSkP7(X1)|~c3_2(X1,a554)), c_0_316, ['final']).
cnf(c_0_1790,negated_conjecture,(ndr1_1(a544)|~c8_1(X1)|~c2_1(X1)|~ssSkC16), c_0_542, ['final']).
cnf(c_0_1791,negated_conjecture,(c1_1(X1)|c6_1(a539)|c6_1(X1)|c5_1(X1)|~c7_0), c_0_845, ['final']).
cnf(c_0_1792,negated_conjecture,(c8_2(a494,a495)|~ssSkC7|~ssSkC6), c_0_293, ['final']).
cnf(c_0_1793,negated_conjecture,(c2_2(a549,a551)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_317, c_0_360])]), ['final']).
cnf(c_0_1794,negated_conjecture,(c6_2(a475,a477)|c7_0|~c2_0), c_0_318, ['final']).
cnf(c_0_1795,negated_conjecture,(c4_2(a475,a477)|c7_0|~c2_0), c_0_319, ['final']).
cnf(c_0_1796,negated_conjecture,(c7_2(a475,a476)|c7_0|~c2_0), c_0_320, ['final']).
cnf(c_0_1797,negated_conjecture,(c7_2(a475,a477)|c7_0|~c2_0), c_0_321, ['final']).
cnf(c_0_1798,negated_conjecture,(ssSkP8(X1)|~c1_2(X1,a588)), c_0_46, ['final']).
cnf(c_0_1799,negated_conjecture,(c5_2(a475,a476)|c7_0|~c2_0), c_0_322, ['final']).
cnf(c_0_1800,negated_conjecture,(c1_1(X1)|c6_1(X1)|ndr1_1(a539)|c5_1(X1)|~c7_0), c_0_846, ['final']).
cnf(c_0_1801,negated_conjecture,(c1_1(X1)|ndr1_1(X1)|c3_1(X1)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_847, c_0_360])]), ['final']).
cnf(c_0_1802,negated_conjecture,(ssSkC5|ndr1_1(X1)|c5_1(X1)|~c2_1(X1)), c_0_534, ['final']).
cnf(c_0_1803,negated_conjecture,(ndr1_1(X1)|c3_1(X1)|~ssSkC3|~c2_1(X1)), c_0_848, ['final']).
cnf(c_0_1804,negated_conjecture,(c6_2(X1,a456)|ssSkP0(X1)), c_0_121, ['final']).
cnf(c_0_1805,negated_conjecture,(c8_2(X1,a456)|ssSkP0(X1)), c_0_99, ['final']).
cnf(c_0_1806,negated_conjecture,(c4_2(X1,a492)|ssSkP2(X1)), c_0_36, ['final']).
cnf(c_0_1807,negated_conjecture,(c2_0|~c8_2(a498,a499)), c_0_849, ['final']).
cnf(c_0_1808,negated_conjecture,(c2_0|~c7_2(a498,a500)), c_0_850, ['final']).
cnf(c_0_1809,negated_conjecture,(ssSkC9|~c1_2(a512,a513)), c_0_136, ['final']).
cnf(c_0_1810,negated_conjecture,(ssSkC9|~c2_2(a512,a513)), c_0_135, ['final']).
cnf(c_0_1811,negated_conjecture,(ssSkC16|~c4_2(a542,a543)), c_0_283, ['final']).
cnf(c_0_1812,negated_conjecture,(c1_2(X1,a503)|ssSkP4(X1)), c_0_62, ['final']).
cnf(c_0_1813,negated_conjecture,(c7_2(X1,a515)|ssSkP5(X1)), c_0_23, ['final']).
cnf(c_0_1814,negated_conjecture,(c8_2(X1,a526)|ssSkP6(X1)), c_0_98, ['final']).
cnf(c_0_1815,negated_conjecture,(c7_2(X1,a526)|ssSkP6(X1)), c_0_20, ['final']).
cnf(c_0_1816,negated_conjecture,(c5_2(X1,a588)|ssSkP8(X1)), c_0_56, ['final']).
cnf(c_0_1817,negated_conjecture,(ssSkC22|ndr1_1(X1)|c7_1(X1)|~c1_1(X1)), c_0_520, ['final']).
cnf(c_0_1818,negated_conjecture,(c1_2(X1,a460)|~ssSkC0), c_0_581, ['final']).
cnf(c_0_1819,negated_conjecture,(c5_2(X1,a460)|~ssSkC0), c_0_699, ['final']).
cnf(c_0_1820,negated_conjecture,(c8_2(X1,a460)|~ssSkC0), c_0_729, ['final']).
cnf(c_0_1821,negated_conjecture,(c2_1(X1)|ndr1_1(X1)|c7_1(X1)|~c7_0), c_0_851, ['final']).
cnf(c_0_1822,negated_conjecture,(ssSkC17|~c8_2(a546,a547)), c_0_323, ['final']).
cnf(c_0_1823,negated_conjecture,(ssSkC17|~c2_2(a546,a547)), c_0_289, ['final']).
cnf(c_0_1824,negated_conjecture,(c1_1(a494)|~ssSkC7|~ssSkC6), c_0_281, ['final']).
cnf(c_0_1825,negated_conjecture,(c3_2(a496,a497)|~c4_0), c_0_663, ['final']).
cnf(c_0_1826,negated_conjecture,(c6_2(a496,a497)|~c4_0), c_0_661, ['final']).
cnf(c_0_1827,negated_conjecture,(c4_2(a509,a510)|~c4_0), c_0_734, ['final']).
cnf(c_0_1828,negated_conjecture,(c3_2(a509,a511)|~c4_0), c_0_717, ['final']).
cnf(c_0_1829,negated_conjecture,(c8_2(a482,a483)|~ssSkC4), c_0_852, ['final']).
cnf(c_0_1830,negated_conjecture,(c1_1(a538)|~ssSkC14|~ssSkC15), c_0_280, ['final']).
cnf(c_0_1831,negated_conjecture,(c5_2(a482,a484)|~ssSkC4), c_0_672, ['final']).
cnf(c_0_1832,negated_conjecture,(c1_1(a566)|~ssSkC19|~ssSkC20), c_0_292, ['final']).
cnf(c_0_1833,negated_conjecture,(c4_1(a569)|~c7_0|~ssSkC21), c_0_302, ['final']).
cnf(c_0_1834,negated_conjecture,(c4_1(a494)|~ssSkC7|~ssSkC6), c_0_301, ['final']).
cnf(c_0_1835,negated_conjecture,(ndr1_1(a494)|~ssSkC7|~ssSkC6), c_0_294, ['final']).
cnf(c_0_1836,negated_conjecture,(ndr1_1(a566)|~ssSkC19|~ssSkC20), c_0_160, ['final']).
cnf(c_0_1837,negated_conjecture,(c3_1(a569)|~c7_0|~ssSkC21), c_0_304, ['final']).
cnf(c_0_1838,negated_conjecture,(c3_2(a482,a483)|~ssSkC4), c_0_853, ['final']).
cnf(c_0_1839,negated_conjecture,(c2_2(a482,a483)|~ssSkC4), c_0_854, ['final']).
cnf(c_0_1840,negated_conjecture,(c4_2(a498,a500)|c2_0), c_0_855, ['final']).
cnf(c_0_1841,negated_conjecture,(c8_2(a498,a500)|c2_0), c_0_533, ['final']).
cnf(c_0_1842,negated_conjecture,(c7_2(a498,a499)|c2_0), c_0_524, ['final']).
cnf(c_0_1843,negated_conjecture,(c3_2(a498,a499)|c2_0), c_0_523, ['final']).
cnf(c_0_1844,negated_conjecture,(c3_2(a512,a513)|ssSkC9), c_0_134, ['final']).
cnf(c_0_1845,negated_conjecture,(c8_2(a542,a543)|ssSkC16), c_0_54, ['final']).
cnf(c_0_1846,negated_conjecture,(c7_2(a546,a547)|ssSkC17), c_0_126, ['final']).
cnf(c_0_1847,negated_conjecture,(ssSkC6|~c4_1(a491)), c_0_291, ['final']).
cnf(c_0_1848,negated_conjecture,(ssSkC6|~c5_1(a491)), c_0_139, ['final']).
cnf(c_0_1849,negated_conjecture,(ssSkC10|~c1_1(a514)), c_0_324, ['final']).
cnf(c_0_1850,negated_conjecture,(ssSkC14|~c7_1(a536)), c_0_288, ['final']).
cnf(c_0_1851,negated_conjecture,(c7_0|ndr1_1(a475)|~c2_0), c_0_225, ['final']).
cnf(c_0_1852,negated_conjecture,(ssSkC15|~c5_1(a537)), c_0_325, ['final']).
cnf(c_0_1853,negated_conjecture,(ssSkC17|~c4_1(a546)), c_0_58, ['final']).
cnf(c_0_1854,negated_conjecture,(ssSkC18|~c8_1(a552)), c_0_68, ['final']).
cnf(c_0_1855,negated_conjecture,(ssSkC21|~c4_1(a568)), c_0_290, ['final']).
cnf(c_0_1856,negated_conjecture,(ssSkC24|~c6_1(a592)), c_0_326, ['final']).
cnf(c_0_1857,negated_conjecture,(ssSkC25|~c6_1(a596)), c_0_327, ['final']).
cnf(c_0_1858,negated_conjecture,(ndr1_1(a496)|~c4_0), c_0_662, ['final']).
cnf(c_0_1859,negated_conjecture,(ndr1_1(a509)|~c4_0), c_0_735, ['final']).
cnf(c_0_1860,negated_conjecture,(c4_1(a488)|~ssSkC5), c_0_769, ['final']).
cnf(c_0_1861,negated_conjecture,(ndr1_1(a482)|~ssSkC4), c_0_535, ['final']).
cnf(c_0_1862,negated_conjecture,(c8_1(a581)|~c2_0), c_0_138, ['final']).
cnf(c_0_1863,negated_conjecture,(c1_1(a466)|~c2_0), c_0_117, ['final']).
cnf(c_0_1864,negated_conjecture,(c1_1(a558)|~c7_0), c_0_282, ['final']).
cnf(c_0_1865,negated_conjecture,(c6_1(a468)|~c4_0), c_0_298, ['final']).
cnf(c_0_1866,negated_conjecture,(c6_1(a558)|~c7_0), c_0_297, ['final']).
cnf(c_0_1867,negated_conjecture,(c5_1(a601)|~ssSkC26), c_0_328, ['final']).
cnf(c_0_1868,negated_conjecture,(c7_1(a601)|~ssSkC26), c_0_329, ['final']).
cnf(c_0_1869,negated_conjecture,(ndr1_1(a498)|c2_0), c_0_525, ['final']).
cnf(c_0_1870,negated_conjecture,(c2_0|c7_1(a498)), c_0_856, ['final']).
cnf(c_0_1871,negated_conjecture,(c7_1(a470)|ssSkC2), c_0_330, ['final']).
cnf(c_0_1872,negated_conjecture,(c6_1(a512)|ssSkC9), c_0_42, ['final']).
cnf(c_0_1873,negated_conjecture,(ndr1_1(a512)|ssSkC9), c_0_32, ['final']).
cnf(c_0_1874,negated_conjecture,(c5_1(a514)|ssSkC10), c_0_331, ['final']).
cnf(c_0_1875,negated_conjecture,(c2_1(a517)|ssSkC11), c_0_22, ['final']).
cnf(c_0_1876,negated_conjecture,(~c4_2(X1,X2)|~c7_2(X1,X2)|~c3_2(X1,X2)|~c2_1(a508)|~c4_1(X1)|~ndr1_1(X1)), c_0_857, ['final']).
cnf(c_0_1877,negated_conjecture,(c6_1(a517)|ssSkC11), c_0_296, ['final']).
cnf(c_0_1878,negated_conjecture,(c4_1(a517)|ssSkC11), c_0_53, ['final']).
cnf(c_0_1879,negated_conjecture,(c1_1(a532)|ssSkC13), c_0_147, ['final']).
cnf(c_0_1880,negated_conjecture,(c3_1(a532)|ssSkC13), c_0_307, ['final']).
cnf(c_0_1881,negated_conjecture,(c5_1(a536)|ssSkC14), c_0_332, ['final']).
cnf(c_0_1882,negated_conjecture,(c4_1(a537)|ssSkC15), c_0_52, ['final']).
cnf(c_0_1883,negated_conjecture,(c2_1(a542)|ssSkC16), c_0_34, ['final']).
cnf(c_0_1884,negated_conjecture,(ndr1_1(a542)|ssSkC16), c_0_27, ['final']).
cnf(c_0_1885,negated_conjecture,(~c4_2(a544,a545)|~c8_1(X1)|~c2_1(X1)|~ssSkC16), c_0_858, ['final']).
cnf(c_0_1886,negated_conjecture,(~c7_2(a544,a545)|~c8_1(X1)|~c2_1(X1)|~ssSkC16), c_0_859, ['final']).
cnf(c_0_1887,negated_conjecture,(~c3_2(a544,a545)|~c8_1(X1)|~c2_1(X1)|~ssSkC16), c_0_860, ['final']).
cnf(c_0_1888,negated_conjecture,(~ssSkC8|~c8_2(a505,a506)|~c4_0), c_0_333, ['final']).
cnf(c_0_1889,negated_conjecture,(~ssSkC8|~c1_2(a505,a506)|~c4_0), c_0_287, ['final']).
cnf(c_0_1890,negated_conjecture,(~c3_2(a494,a495)|~ssSkC7|~ssSkC6), c_0_295, ['final']).
cnf(c_0_1891,negated_conjecture,(~c8_2(a566,a567)|~ssSkC19|~ssSkC20), c_0_334, ['final']).
cnf(c_0_1892,negated_conjecture,(~c1_2(a566,a567)|~ssSkC19|~ssSkC20), c_0_286, ['final']).
cnf(c_0_1893,negated_conjecture,(~c1_2(a549,a550)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_335, c_0_360])]), ['final']).
cnf(c_0_1894,negated_conjecture,(ndr1_1(a546)|ssSkC17), c_0_45, ['final']).
cnf(c_0_1895,negated_conjecture,(c6_1(a552)|ssSkC18), c_0_41, ['final']).
cnf(c_0_1896,negated_conjecture,(c5_1(a564)|ssSkC19), c_0_336, ['final']).
cnf(c_0_1897,negated_conjecture,(c7_1(a564)|ssSkC19), c_0_337, ['final']).
cnf(c_0_1898,negated_conjecture,(c2_1(a565)|ssSkC20), c_0_16, ['final']).
cnf(c_0_1899,negated_conjecture,(c6_1(a565)|ssSkC20), c_0_5, ['final']).
cnf(c_0_1900,negated_conjecture,(c3_1(a565)|ssSkC20), c_0_6, ['final']).
cnf(c_0_1901,negated_conjecture,(c3_1(a587)|ssSkC23), c_0_306, ['final']).
cnf(c_0_1902,negated_conjecture,(~c4_2(a496,a497)|~c4_0), c_0_664, ['final']).
cnf(c_0_1903,negated_conjecture,(~c6_2(a509,a511)|~c4_0), c_0_861, ['final']).
cnf(c_0_1904,negated_conjecture,(~c8_2(a509,a510)|~c4_0), c_0_736, ['final']).
cnf(c_0_1905,negated_conjecture,(~c5_2(a509,a511)|~c4_0), c_0_862, ['final']).
cnf(c_0_1906,negated_conjecture,(~c1_2(a509,a510)|~c4_0), c_0_863, ['final']).
cnf(c_0_1907,negated_conjecture,(~c6_2(a482,a484)|~ssSkC4), c_0_864, ['final']).
cnf(c_0_1908,negated_conjecture,(~c4_2(a482,a484)|~ssSkC4), c_0_865, ['final']).
cnf(c_0_1909,negated_conjecture,(c7_1(a596)|ssSkC25), c_0_338, ['final']).
cnf(c_0_1910,negated_conjecture,(c3_1(a600)|ssSkC26), c_0_305, ['final']).
cnf(c_0_1911,negated_conjecture,(ssSkP0(X1)|ndr1_1(X1)), c_0_100, ['final']).
cnf(c_0_1912,negated_conjecture,(ssSkP1(X1)|ndr1_1(X1)), c_0_279, ['final']).
cnf(c_0_1913,negated_conjecture,(ssSkP2(X1)|ndr1_1(X1)), c_0_339, ['final']).
cnf(c_0_1914,negated_conjecture,(ssSkP3(X1)|ndr1_1(X1)), c_0_8, ['final']).
cnf(c_0_1915,negated_conjecture,(ssSkP4(X1)|ndr1_1(X1)), c_0_9, ['final']).
cnf(c_0_1916,negated_conjecture,(~ssSkC8|~c2_1(a505)|~c4_0), c_0_340, ['final']).
cnf(c_0_1917,negated_conjecture,(~c2_1(a549)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_341, c_0_360])]), ['final']).
cnf(c_0_1918,negated_conjecture,(~c2_1(a496)|~c4_0), c_0_866, ['final']).
cnf(c_0_1919,negated_conjecture,(~ssSkC5|~c3_1(a488)), c_0_867, ['final']).
cnf(c_0_1920,negated_conjecture,(~c2_1(a561)|~c4_0), c_0_342, ['final']).
cnf(c_0_1921,negated_conjecture,(~c4_0|~c4_1(a468)), c_0_343, ['final']).
cnf(c_0_1922,negated_conjecture,(~c4_0|~c3_1(a468)), c_0_344, ['final']).
cnf(c_0_1923,negated_conjecture,(ssSkP5(X1)|ndr1_1(X1)), c_0_48, ['final']).
cnf(c_0_1924,negated_conjecture,(ssSkP6(X1)|ndr1_1(X1)), c_0_31, ['final']).
cnf(c_0_1925,negated_conjecture,(ssSkP7(X1)|ndr1_1(X1)), c_0_345, ['final']).
cnf(c_0_1926,negated_conjecture,(ssSkP8(X1)|ndr1_1(X1)), c_0_278, ['final']).
cnf(c_0_1927,negated_conjecture,(ndr1_1(X1)|~ssSkC0), c_0_541, ['final']).
cnf(c_0_1928,negated_conjecture,(~c4_0|~c5_1(a561)), c_0_346, ['final']).
cnf(c_0_1929,negated_conjecture,(~c8_1(a466)|~c2_0), c_0_118, ['final']).
cnf(c_0_1930,negated_conjecture,(~c6_1(a581)|~c2_0), c_0_308, ['final']).
cnf(c_0_1931,negated_conjecture,(~c1_1(a601)|~ssSkC26), c_0_347, ['final']).
cnf(c_0_1932,negated_conjecture,(~epred2_0|~epred1_0), c_0_868, ['final']).
cnf(c_0_1933,negated_conjecture,(ssSkC8|ndr1_1(X1)), c_0_353, ['final']).
cnf(c_0_1934,negated_conjecture,(c2_0|~c7_0), c_0_596, ['final']).
cnf(c_0_1935,negated_conjecture,(~c7_0|~c4_0), c_0_150, ['final']).
cnf(c_0_1936,negated_conjecture,(~c4_2(a520,a521)), c_0_869, ['final']).
cnf(c_0_1937,negated_conjecture,(~c1_2(a520,a521)), c_0_870, ['final']).
cnf(c_0_1938,negated_conjecture,(~c2_1(a520)), c_0_871, ['final']).
cnf(c_0_1939,negated_conjecture,(~c5_1(a520)), c_0_433, ['final']).
cnf(c_0_1940,negated_conjecture,(~c1_0), c_0_366, ['final']).
cnf(c_0_1941,negated_conjecture,(~c5_0), c_0_350, ['final']).
cnf(c_0_1942,negated_conjecture,(c2_2(a520,a521)), c_0_872, ['final']).
cnf(c_0_1943,negated_conjecture,(c3_2(a570,a571)), c_0_873, ['final']).
cnf(c_0_1944,negated_conjecture,(c1_2(a570,a571)), c_0_874, ['final']).
cnf(c_0_1945,negated_conjecture,(c2_2(a570,a571)), c_0_875, ['final']).
cnf(c_0_1946,negated_conjecture,(ndr1_1(a520)), c_0_380, ['final']).
cnf(c_0_1947,negated_conjecture,(c4_1(a570)), c_0_420, ['final']).
cnf(c_0_1948,negated_conjecture,(ndr1_1(a570)), c_0_421, ['final']).
cnf(c_0_1949,negated_conjecture,(c5_1(a570)), c_0_876, ['final']).
cnf(c_0_1950,negated_conjecture,(c6_0), c_0_441, ['final']).
cnf(c_0_1951,negated_conjecture,(ndr1_0), c_0_349, ['final']).
cnf(c_0_1952,negated_conjecture,(c3_0), c_0_348, ['final']).
# SZS output end Saturation.
EOF
