% Problem    : PUZ056-2.010 : TPTP v6.1.0. Released v3.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n179.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:30:56 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Presaturation interreduction done
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___008_C45_F1_PI_AE_Q4_CS_SP_S4c
# and selection function SelectCQPrecWNTNp.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(p(X1,X11,X3,X4,X5,X6,X7,X8,X9,X10)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~neq(X1,X2)|~neq(X1,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule2)).
cnf(c_0_1,axiom,(neq(s1,s2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', neq6)).
cnf(c_0_2,axiom,(neq(s1,s0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', neq4)).
cnf(c_0_3,axiom,(p(X11,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule1)).
cnf(c_0_4,axiom,(p(s0,s0,s0,s0,s0,s0,s0,s0,s0,s0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', init)).
cnf(c_0_5,axiom,(p(X1,X2,X11,X4,X5,X6,X7,X8,X9,X10)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~neq(X1,X3)|~neq(X1,X11)|~neq(X2,X3)|~neq(X2,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule3)).
cnf(c_0_6,axiom,(neq(s2,s1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', neq8)).
cnf(c_0_7,axiom,(neq(s2,s0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', neq7)).
cnf(c_0_8,axiom,(neq(s0,s1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', neq2)).
cnf(c_0_9,axiom,(p(X1,X2,X3,X11,X5,X6,X7,X8,X9,X10)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~neq(X1,X4)|~neq(X1,X11)|~neq(X2,X4)|~neq(X2,X11)|~neq(X3,X4)|~neq(X3,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule4)).
cnf(c_0_10,axiom,(neq(s0,s2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', neq3)).
cnf(c_0_11,axiom,(p(X1,X2,X3,X4,X11,X6,X7,X8,X9,X10)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~neq(X1,X5)|~neq(X1,X11)|~neq(X2,X5)|~neq(X2,X11)|~neq(X3,X5)|~neq(X3,X11)|~neq(X4,X5)|~neq(X4,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule5)).
cnf(c_0_12,axiom,(p(X1,X2,X3,X4,X5,X11,X7,X8,X9,X10)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~neq(X1,X6)|~neq(X1,X11)|~neq(X2,X6)|~neq(X2,X11)|~neq(X3,X6)|~neq(X3,X11)|~neq(X4,X6)|~neq(X4,X11)|~neq(X5,X6)|~neq(X5,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule6)).
cnf(c_0_13,axiom,(p(X1,X2,X3,X4,X5,X6,X11,X8,X9,X10)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~neq(X1,X7)|~neq(X1,X11)|~neq(X2,X7)|~neq(X2,X11)|~neq(X3,X7)|~neq(X3,X11)|~neq(X4,X7)|~neq(X4,X11)|~neq(X5,X7)|~neq(X5,X11)|~neq(X6,X7)|~neq(X6,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule7)).
cnf(c_0_14,axiom,(p(X1,X2,X3,X4,X5,X6,X7,X11,X9,X10)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~neq(X1,X8)|~neq(X1,X11)|~neq(X2,X8)|~neq(X2,X11)|~neq(X3,X8)|~neq(X3,X11)|~neq(X4,X8)|~neq(X4,X11)|~neq(X5,X8)|~neq(X5,X11)|~neq(X6,X8)|~neq(X6,X11)|~neq(X7,X8)|~neq(X7,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule8)).
cnf(c_0_15,axiom,(p(X1,X2,X3,X4,X5,X6,X7,X8,X11,X10)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~neq(X1,X9)|~neq(X1,X11)|~neq(X2,X9)|~neq(X2,X11)|~neq(X3,X9)|~neq(X3,X11)|~neq(X4,X9)|~neq(X4,X11)|~neq(X5,X9)|~neq(X5,X11)|~neq(X6,X9)|~neq(X6,X11)|~neq(X7,X9)|~neq(X7,X11)|~neq(X8,X9)|~neq(X8,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule9)).
cnf(c_0_16,axiom,(p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X11)|~p(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~neq(X1,X10)|~neq(X1,X11)|~neq(X2,X10)|~neq(X2,X11)|~neq(X3,X10)|~neq(X3,X11)|~neq(X4,X10)|~neq(X4,X11)|~neq(X5,X10)|~neq(X5,X11)|~neq(X6,X10)|~neq(X6,X11)|~neq(X7,X10)|~neq(X7,X11)|~neq(X8,X10)|~neq(X8,X11)|~neq(X9,X10)|~neq(X9,X11)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rule10)).
cnf(c_0_17,negated_conjecture,(~p(s2,s2,s2,s2,s2,s2,s2,s2,s2,s2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goal)).
cnf(c_0_18,plain,(p(s1,s2,X1,X2,X3,X4,X5,X6,X7,X8)|~neq(s1,X9)|~p(s1,X9,X1,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_19,plain,(p(s1,s2,X1,X2,X3,X4,X5,X6,X7,X8)|~p(s1,s0,X1,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_18, c_0_2])).
cnf(c_0_20,plain,(p(X1,s0,s0,s0,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_21,plain,(p(X1,s2,s1,X2,X3,X4,X5,X6,X7,X8)|~neq(s2,X9)|~neq(X1,s1)|~neq(X1,X9)|~p(X1,s2,X9,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_22,plain,(p(s1,s2,s0,s0,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_23,plain,(p(s2,s2,s1,X1,X2,X3,X4,X5,X6,X7)|~p(s2,s2,s0,X1,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_7]), c_0_7]), c_0_6])])).
cnf(c_0_24,plain,(p(X1,s2,s0,s0,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_22])).
cnf(c_0_25,plain,(p(s0,s1,X1,X2,X3,X4,X5,X6,X7,X8)|~neq(s0,X9)|~p(s0,X9,X1,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_0, c_0_8])).
cnf(c_0_26,plain,(p(s2,s2,s1,s0,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27,plain,(p(X1,X2,s1,s2,X3,X4,X5,X6,X7,X8)|~neq(s1,X9)|~neq(X2,s2)|~neq(X1,s2)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,s1,X9,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_9, c_0_1])).
cnf(c_0_28,plain,(p(s0,s1,X1,X2,X3,X4,X5,X6,X7,X8)|~p(s0,s2,X1,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_25, c_0_10])).
cnf(c_0_29,plain,(p(X1,s2,s1,s0,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_26])).
cnf(c_0_30,plain,(p(X1,s1,s1,s2,X2,X3,X4,X5,X6,X7)|~neq(X1,s2)|~neq(X1,s0)|~p(X1,s1,s1,s0,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_2]), c_0_2]), c_0_1])])).
cnf(c_0_31,plain,(p(s0,s1,s1,s0,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_32,plain,(p(s1,s1,s1,s2,X1,X2,X3,X4,X5,X6)|~p(s1,s1,s1,s0,X1,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_1]), c_0_2])])).
cnf(c_0_33,plain,(p(X1,s1,s1,s0,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_31])).
cnf(c_0_34,plain,(p(s2,s0,X1,X2,X3,X4,X5,X6,X7,X8)|~neq(s2,X9)|~p(s2,X9,X1,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_0, c_0_7])).
cnf(c_0_35,plain,(p(s1,s1,s1,s2,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36,plain,(p(s2,s0,X1,X2,X3,X4,X5,X6,X7,X8)|~p(s2,s1,X1,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_34, c_0_6])).
cnf(c_0_37,plain,(p(X1,s1,s1,s2,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_35])).
cnf(c_0_38,plain,(p(X1,s0,s2,X2,X3,X4,X5,X6,X7,X8)|~neq(s0,X9)|~neq(X1,s2)|~neq(X1,X9)|~p(X1,s0,X9,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_5, c_0_10])).
cnf(c_0_39,plain,(p(s2,s0,s1,s2,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_40,plain,(p(s0,s0,s2,X1,X2,X3,X4,X5,X6,X7)|~p(s0,s0,s1,X1,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_8]), c_0_8]), c_0_10])])).
cnf(c_0_41,plain,(p(X1,s0,s1,s2,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_39])).
cnf(c_0_42,plain,(p(X1,X2,X3,s2,s1,X4,X5,X6,X7,X8)|~neq(s2,X9)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,s2,X9,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_11, c_0_6])).
cnf(c_0_43,plain,(p(s0,s0,s2,s2,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_44,plain,(p(X1,X2,s2,s2,s1,X3,X4,X5,X6,X7)|~neq(X2,s1)|~neq(X1,s1)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,s2,s2,s0,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_7]), c_0_7]), c_0_6])])).
cnf(c_0_45,plain,(p(X1,s0,s2,s2,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_43])).
cnf(c_0_46,plain,(p(X1,s2,s2,s2,s1,X2,X3,X4,X5,X6)|~neq(X1,s1)|~neq(X1,s0)|~p(X1,s2,s2,s2,s0,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_6]), c_0_7])])).
cnf(c_0_47,plain,(p(s1,s2,s2,s2,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_45])).
cnf(c_0_48,plain,(p(s2,s2,s2,s2,s1,X1,X2,X3,X4,X5)|~p(s2,s2,s2,s2,s0,X1,X2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_6]), c_0_7])])).
cnf(c_0_49,plain,(p(X1,s2,s2,s2,s0,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_47])).
cnf(c_0_50,plain,(p(s2,s2,s2,s2,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_51,plain,(p(X1,s2,s2,s2,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_50])).
cnf(c_0_52,plain,(p(X1,s1,s0,X2,X3,X4,X5,X6,X7,X8)|~neq(s1,X9)|~neq(X1,s0)|~neq(X1,X9)|~p(X1,s1,X9,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_5, c_0_2])).
cnf(c_0_53,plain,(p(s0,s1,s2,s2,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_51])).
cnf(c_0_54,plain,(p(s1,s1,s0,X1,X2,X3,X4,X5,X6,X7)|~p(s1,s1,s2,X1,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1]), c_0_1]), c_0_2])])).
cnf(c_0_55,plain,(p(X1,s1,s2,s2,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_53])).
cnf(c_0_56,plain,(p(s1,s1,s0,s2,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_57,plain,(p(X1,X2,s0,s1,X3,X4,X5,X6,X7,X8)|~neq(s0,X9)|~neq(X2,s1)|~neq(X1,s1)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,s0,X9,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_9, c_0_8])).
cnf(c_0_58,plain,(p(X1,s1,s0,s2,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_56])).
cnf(c_0_59,plain,(p(X1,s0,s0,s1,X2,X3,X4,X5,X6,X7)|~neq(X1,s1)|~neq(X1,s2)|~p(X1,s0,s0,s2,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_10]), c_0_10]), c_0_8])])).
cnf(c_0_60,plain,(p(s2,s0,s0,s2,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_58])).
cnf(c_0_61,plain,(p(s0,s0,s0,s1,X1,X2,X3,X4,X5,X6)|~p(s0,s0,s0,s2,X1,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_8]), c_0_10])])).
cnf(c_0_62,plain,(p(X1,s0,s0,s2,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_60])).
cnf(c_0_63,plain,(p(s0,s0,s0,s1,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_64,plain,(p(X1,s0,s0,s1,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_63])).
cnf(c_0_65,plain,(p(s1,s2,s0,s1,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_64])).
cnf(c_0_66,plain,(p(X1,X2,X3,X4,s1,s2,X5,X6,X7,X8)|~neq(s1,X9)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,s1,X9,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_12, c_0_1])).
cnf(c_0_67,plain,(p(X1,s2,s0,s1,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_65])).
cnf(c_0_68,plain,(p(X1,X2,X3,s1,s1,s2,X4,X5,X6,X7)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,s1,s1,s0,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_2]), c_0_2]), c_0_1])])).
cnf(c_0_69,plain,(p(s2,s2,s1,s1,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_67])).
cnf(c_0_70,plain,(p(X1,X2,s1,s1,s1,s2,X3,X4,X5,X6)|~neq(X2,s2)|~neq(X1,s2)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,s1,s1,s1,s0,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_1]), c_0_2])])).
cnf(c_0_71,plain,(p(X1,s2,s1,s1,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_69])).
cnf(c_0_72,plain,(p(X1,s1,s1,s1,s1,s2,X2,X3,X4,X5)|~neq(X1,s2)|~neq(X1,s0)|~p(X1,s1,s1,s1,s1,s0,X2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_1]), c_0_2])])).
cnf(c_0_73,plain,(p(s0,s1,s1,s1,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_71])).
cnf(c_0_74,plain,(p(s1,s1,s1,s1,s1,s2,X1,X2,X3,X4)|~p(s1,s1,s1,s1,s1,s0,X1,X2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_1]), c_0_2])])).
cnf(c_0_75,plain,(p(X1,s1,s1,s1,s1,s0,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_73])).
cnf(c_0_76,plain,(p(s1,s1,s1,s1,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_77,plain,(p(X1,s1,s1,s1,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_76])).
cnf(c_0_78,plain,(p(s2,s0,s1,s1,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_77])).
cnf(c_0_79,plain,(p(X1,s0,s1,s1,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_78])).
cnf(c_0_80,plain,(p(s0,s0,s2,s1,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_79])).
cnf(c_0_81,plain,(p(X1,X2,s2,s0,X3,X4,X5,X6,X7,X8)|~neq(s2,X9)|~neq(X2,s0)|~neq(X1,s0)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,s2,X9,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_9, c_0_7])).
cnf(c_0_82,plain,(p(X1,s0,s2,s1,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_80])).
cnf(c_0_83,plain,(p(X1,s2,s2,s0,X2,X3,X4,X5,X6,X7)|~neq(X1,s0)|~neq(X1,s1)|~p(X1,s2,s2,s1,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_6]), c_0_6]), c_0_7])])).
cnf(c_0_84,plain,(p(s1,s2,s2,s1,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_82])).
cnf(c_0_85,plain,(p(s2,s2,s2,s0,X1,X2,X3,X4,X5,X6)|~p(s2,s2,s2,s1,X1,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_7]), c_0_6])])).
cnf(c_0_86,plain,(p(X1,s2,s2,s1,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_84])).
cnf(c_0_87,plain,(p(s2,s2,s2,s0,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_88,plain,(p(X1,s2,s2,s0,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_87])).
cnf(c_0_89,plain,(p(s0,s1,s2,s0,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_88])).
cnf(c_0_90,plain,(p(X1,s1,s2,s0,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_89])).
cnf(c_0_91,plain,(p(X1,X2,X3,s0,s2,X4,X5,X6,X7,X8)|~neq(s0,X9)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,s0,X9,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_11, c_0_10])).
cnf(c_0_92,plain,(p(s1,s1,s0,s0,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_90])).
cnf(c_0_93,plain,(p(X1,X2,s0,s0,s2,X3,X4,X5,X6,X7)|~neq(X2,s2)|~neq(X1,s2)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,s0,s0,s1,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_8]), c_0_8]), c_0_10])])).
cnf(c_0_94,plain,(p(X1,s1,s0,s0,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_92])).
cnf(c_0_95,plain,(p(X1,s0,s0,s0,s2,X2,X3,X4,X5,X6)|~neq(X1,s2)|~neq(X1,s1)|~p(X1,s0,s0,s0,s1,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_10]), c_0_8])])).
cnf(c_0_96,plain,(p(s2,s0,s0,s0,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_94])).
cnf(c_0_97,plain,(p(s0,s0,s0,s0,s2,X1,X2,X3,X4,X5)|~p(s0,s0,s0,s0,s1,X1,X2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_10]), c_0_8])])).
cnf(c_0_98,plain,(p(X1,s0,s0,s0,s1,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_96])).
cnf(c_0_99,plain,(p(s0,s0,s0,s0,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_100,plain,(p(X1,s0,s0,s0,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_99])).
cnf(c_0_101,plain,(p(s1,s2,s0,s0,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_100])).
cnf(c_0_102,plain,(p(X1,s2,s0,s0,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_101])).
cnf(c_0_103,plain,(p(s2,s2,s1,s0,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_102])).
cnf(c_0_104,plain,(p(X1,s2,s1,s0,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_103])).
cnf(c_0_105,plain,(p(s0,s1,s1,s0,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_104])).
cnf(c_0_106,plain,(p(X1,s1,s1,s0,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_105])).
cnf(c_0_107,plain,(p(s1,s1,s1,s2,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_32, c_0_106])).
cnf(c_0_108,plain,(p(X1,s1,s1,s2,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_107])).
cnf(c_0_109,plain,(p(X1,X2,X3,X4,X5,s2,s1,X6,X7,X8)|~neq(s2,X9)|~neq(X5,s1)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X5,X9)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,X5,s2,X9,X6,X7,X8)), inference(spm,[status(thm)],[c_0_13, c_0_6])).
cnf(c_0_110,plain,(p(s2,s0,s1,s2,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_108])).
cnf(c_0_111,plain,(p(X1,X2,X3,X4,s2,s2,s1,X5,X6,X7)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,s2,s2,s0,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_7]), c_0_7]), c_0_6])])).
cnf(c_0_112,plain,(p(X1,s0,s1,s2,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_110])).
cnf(c_0_113,plain,(p(X1,X2,X3,s2,s2,s2,s1,X4,X5,X6)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,s2,s2,s2,s0,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_6]), c_0_7])])).
cnf(c_0_114,plain,(p(s0,s0,s2,s2,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_112])).
cnf(c_0_115,plain,(p(X1,X2,s2,s2,s2,s2,s1,X3,X4,X5)|~neq(X2,s1)|~neq(X1,s1)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,s2,s2,s2,s2,s0,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_6]), c_0_7])])).
cnf(c_0_116,plain,(p(X1,s0,s2,s2,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_114])).
cnf(c_0_117,plain,(p(X1,s2,s2,s2,s2,s2,s1,X2,X3,X4)|~neq(X1,s1)|~neq(X1,s0)|~p(X1,s2,s2,s2,s2,s2,s0,X2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_6]), c_0_7])])).
cnf(c_0_118,plain,(p(s1,s2,s2,s2,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_116])).
cnf(c_0_119,plain,(p(s2,s2,s2,s2,s2,s2,s1,X1,X2,X3)|~p(s2,s2,s2,s2,s2,s2,s0,X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_6]), c_0_7])])).
cnf(c_0_120,plain,(p(X1,s2,s2,s2,s2,s2,s0,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_118])).
cnf(c_0_121,plain,(p(s2,s2,s2,s2,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_122,plain,(p(X1,s2,s2,s2,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_121])).
cnf(c_0_123,plain,(p(s0,s1,s2,s2,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_122])).
cnf(c_0_124,plain,(p(X1,s1,s2,s2,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_123])).
cnf(c_0_125,plain,(p(s1,s1,s0,s2,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_124])).
cnf(c_0_126,plain,(p(X1,s1,s0,s2,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_125])).
cnf(c_0_127,plain,(p(s2,s0,s0,s2,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_126])).
cnf(c_0_128,plain,(p(X1,s0,s0,s2,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_127])).
cnf(c_0_129,plain,(p(s0,s0,s0,s1,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_61, c_0_128])).
cnf(c_0_130,plain,(p(X1,s0,s0,s1,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_129])).
cnf(c_0_131,plain,(p(s1,s2,s0,s1,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_130])).
cnf(c_0_132,plain,(p(X1,s2,s0,s1,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_131])).
cnf(c_0_133,plain,(p(X1,X2,X3,s1,s0,X4,X5,X6,X7,X8)|~neq(s1,X9)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,s1,X9,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_11, c_0_2])).
cnf(c_0_134,plain,(p(s2,s2,s1,s1,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_132])).
cnf(c_0_135,plain,(p(X1,X2,s1,s1,s0,X3,X4,X5,X6,X7)|~neq(X2,s0)|~neq(X1,s0)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,s1,s1,s2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_1]), c_0_1]), c_0_2])])).
cnf(c_0_136,plain,(p(X1,s2,s1,s1,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_134])).
cnf(c_0_137,plain,(p(X1,s1,s1,s1,s0,X2,X3,X4,X5,X6)|~neq(X1,s0)|~neq(X1,s2)|~p(X1,s1,s1,s1,s2,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_2]), c_0_1])])).
cnf(c_0_138,plain,(p(s0,s1,s1,s1,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_136])).
cnf(c_0_139,plain,(p(s1,s1,s1,s1,s0,X1,X2,X3,X4,X5)|~p(s1,s1,s1,s1,s2,X1,X2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_2]), c_0_1])])).
cnf(c_0_140,plain,(p(X1,s1,s1,s1,s2,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_138])).
cnf(c_0_141,plain,(p(s1,s1,s1,s1,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_142,plain,(p(X1,s1,s1,s1,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_141])).
cnf(c_0_143,plain,(p(s2,s0,s1,s1,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_142])).
cnf(c_0_144,plain,(p(X1,s0,s1,s1,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_143])).
cnf(c_0_145,plain,(p(s0,s0,s2,s1,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_144])).
cnf(c_0_146,plain,(p(X1,s0,s2,s1,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_145])).
cnf(c_0_147,plain,(p(s1,s2,s2,s1,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_146])).
cnf(c_0_148,plain,(p(X1,s2,s2,s1,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_147])).
cnf(c_0_149,plain,(p(s2,s2,s2,s0,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_85, c_0_148])).
cnf(c_0_150,plain,(p(X1,s2,s2,s0,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_149])).
cnf(c_0_151,plain,(p(s0,s1,s2,s0,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_150])).
cnf(c_0_152,plain,(p(X1,X2,X3,X4,s0,s1,X5,X6,X7,X8)|~neq(s0,X9)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,s0,X9,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_12, c_0_8])).
cnf(c_0_153,plain,(p(X1,s1,s2,s0,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_151])).
cnf(c_0_154,plain,(p(X1,X2,X3,s0,s0,s1,X4,X5,X6,X7)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,X3,s0,s0,s2,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_10]), c_0_10]), c_0_8])])).
cnf(c_0_155,plain,(p(s1,s1,s0,s0,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_153])).
cnf(c_0_156,plain,(p(X1,X2,s0,s0,s0,s1,X3,X4,X5,X6)|~neq(X2,s1)|~neq(X1,s1)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,s0,s0,s0,s2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_8]), c_0_10])])).
cnf(c_0_157,plain,(p(X1,s1,s0,s0,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_155])).
cnf(c_0_158,plain,(p(X1,s0,s0,s0,s0,s1,X2,X3,X4,X5)|~neq(X1,s1)|~neq(X1,s2)|~p(X1,s0,s0,s0,s0,s2,X2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_8]), c_0_10])])).
cnf(c_0_159,plain,(p(s2,s0,s0,s0,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_157])).
cnf(c_0_160,plain,(p(s0,s0,s0,s0,s0,s1,X1,X2,X3,X4)|~p(s0,s0,s0,s0,s0,s2,X1,X2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_8]), c_0_10])])).
cnf(c_0_161,plain,(p(X1,s0,s0,s0,s0,s2,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_159])).
cnf(c_0_162,plain,(p(s0,s0,s0,s0,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_163,plain,(p(X1,s0,s0,s0,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_162])).
cnf(c_0_164,plain,(p(s1,s2,s0,s0,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_163])).
cnf(c_0_165,plain,(p(X1,s2,s0,s0,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_164])).
cnf(c_0_166,plain,(p(s2,s2,s1,s0,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_165])).
cnf(c_0_167,plain,(p(X1,s2,s1,s0,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_166])).
cnf(c_0_168,plain,(p(s0,s1,s1,s0,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_167])).
cnf(c_0_169,plain,(p(X1,s1,s1,s0,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_168])).
cnf(c_0_170,plain,(p(s1,s1,s1,s2,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_32, c_0_169])).
cnf(c_0_171,plain,(p(X1,s1,s1,s2,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_170])).
cnf(c_0_172,plain,(p(s2,s0,s1,s2,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_171])).
cnf(c_0_173,plain,(p(X1,s0,s1,s2,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_172])).
cnf(c_0_174,plain,(p(s0,s0,s2,s2,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_173])).
cnf(c_0_175,plain,(p(X1,s0,s2,s2,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_174])).
cnf(c_0_176,plain,(p(s1,s2,s2,s2,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_175])).
cnf(c_0_177,plain,(p(X1,s2,s2,s2,s0,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_176])).
cnf(c_0_178,plain,(p(s2,s2,s2,s2,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_48, c_0_177])).
cnf(c_0_179,plain,(p(X1,s2,s2,s2,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_178])).
cnf(c_0_180,plain,(p(s0,s1,s2,s2,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_179])).
cnf(c_0_181,plain,(p(X1,s1,s2,s2,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_180])).
cnf(c_0_182,plain,(p(s1,s1,s0,s2,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_181])).
cnf(c_0_183,plain,(p(X1,s1,s0,s2,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_182])).
cnf(c_0_184,plain,(p(s2,s0,s0,s2,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_183])).
cnf(c_0_185,plain,(p(X1,s0,s0,s2,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_184])).
cnf(c_0_186,plain,(p(s0,s0,s0,s1,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_61, c_0_185])).
cnf(c_0_187,plain,(p(X1,X2,X3,X4,X5,X6,s1,s2,X7,X8)|~neq(s1,X9)|~neq(X6,s2)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X6,X9)|~neq(X5,X9)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,X5,X6,s1,X9,X7,X8)), inference(spm,[status(thm)],[c_0_14, c_0_1])).
cnf(c_0_188,plain,(p(X1,s0,s0,s1,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_186])).
cnf(c_0_189,plain,(p(X1,X2,X3,X4,X5,s1,s1,s2,X6,X7)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X5,s0)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,X5,s1,s1,s0,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_2]), c_0_2]), c_0_1])])).
cnf(c_0_190,plain,(p(s1,s2,s0,s1,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_188])).
cnf(c_0_191,plain,(p(X1,X2,X3,X4,s1,s1,s1,s2,X5,X6)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,s1,s1,s1,s0,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_1]), c_0_2])])).
cnf(c_0_192,plain,(p(X1,s2,s0,s1,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_190])).
cnf(c_0_193,plain,(p(X1,X2,X3,s1,s1,s1,s1,s2,X4,X5)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,s1,s1,s1,s1,s0,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_1]), c_0_2])])).
cnf(c_0_194,plain,(p(s2,s2,s1,s1,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_192])).
cnf(c_0_195,plain,(p(X1,X2,s1,s1,s1,s1,s1,s2,X3,X4)|~neq(X2,s2)|~neq(X1,s2)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,s1,s1,s1,s1,s1,s0,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_1]), c_0_2])])).
cnf(c_0_196,plain,(p(X1,s2,s1,s1,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_194])).
cnf(c_0_197,plain,(p(X1,s1,s1,s1,s1,s1,s1,s2,X2,X3)|~neq(X1,s2)|~neq(X1,s0)|~p(X1,s1,s1,s1,s1,s1,s1,s0,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_1]), c_0_2])])).
cnf(c_0_198,plain,(p(s0,s1,s1,s1,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_196])).
cnf(c_0_199,plain,(p(s1,s1,s1,s1,s1,s1,s1,s2,X1,X2)|~p(s1,s1,s1,s1,s1,s1,s1,s0,X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_1]), c_0_2])])).
cnf(c_0_200,plain,(p(X1,s1,s1,s1,s1,s1,s1,s0,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_198])).
cnf(c_0_201,plain,(p(s1,s1,s1,s1,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_199, c_0_200])).
cnf(c_0_202,plain,(p(X1,s1,s1,s1,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_201])).
cnf(c_0_203,plain,(p(s2,s0,s1,s1,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_202])).
cnf(c_0_204,plain,(p(X1,s0,s1,s1,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_203])).
cnf(c_0_205,plain,(p(s0,s0,s2,s1,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_204])).
cnf(c_0_206,plain,(p(X1,s0,s2,s1,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_205])).
cnf(c_0_207,plain,(p(s1,s2,s2,s1,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_206])).
cnf(c_0_208,plain,(p(X1,s2,s2,s1,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_207])).
cnf(c_0_209,plain,(p(s2,s2,s2,s0,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_85, c_0_208])).
cnf(c_0_210,plain,(p(X1,s2,s2,s0,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_209])).
cnf(c_0_211,plain,(p(s0,s1,s2,s0,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_210])).
cnf(c_0_212,plain,(p(X1,s1,s2,s0,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_211])).
cnf(c_0_213,plain,(p(s1,s1,s0,s0,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_212])).
cnf(c_0_214,plain,(p(X1,s1,s0,s0,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_213])).
cnf(c_0_215,plain,(p(s2,s0,s0,s0,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_214])).
cnf(c_0_216,plain,(p(X1,s0,s0,s0,s1,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_215])).
cnf(c_0_217,plain,(p(s0,s0,s0,s0,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_97, c_0_216])).
cnf(c_0_218,plain,(p(X1,s0,s0,s0,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_217])).
cnf(c_0_219,plain,(p(s1,s2,s0,s0,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_218])).
cnf(c_0_220,plain,(p(X1,s2,s0,s0,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_219])).
cnf(c_0_221,plain,(p(s2,s2,s1,s0,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_220])).
cnf(c_0_222,plain,(p(X1,s2,s1,s0,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_221])).
cnf(c_0_223,plain,(p(s0,s1,s1,s0,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_222])).
cnf(c_0_224,plain,(p(X1,s1,s1,s0,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_223])).
cnf(c_0_225,plain,(p(s1,s1,s1,s2,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_32, c_0_224])).
cnf(c_0_226,plain,(p(X1,s1,s1,s2,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_225])).
cnf(c_0_227,plain,(p(s2,s0,s1,s2,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_226])).
cnf(c_0_228,plain,(p(X1,X2,X3,X4,s2,s0,X5,X6,X7,X8)|~neq(s2,X9)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,s2,X9,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_12, c_0_7])).
cnf(c_0_229,plain,(p(X1,s0,s1,s2,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_227])).
cnf(c_0_230,plain,(p(X1,X2,X3,s2,s2,s0,X4,X5,X6,X7)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,s2,s2,s1,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_6]), c_0_6]), c_0_7])])).
cnf(c_0_231,plain,(p(s0,s0,s2,s2,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_229])).
cnf(c_0_232,plain,(p(X1,X2,s2,s2,s2,s0,X3,X4,X5,X6)|~neq(X2,s0)|~neq(X1,s0)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,s2,s2,s2,s1,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_7]), c_0_6])])).
cnf(c_0_233,plain,(p(X1,s0,s2,s2,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_231])).
cnf(c_0_234,plain,(p(X1,s2,s2,s2,s2,s0,X2,X3,X4,X5)|~neq(X1,s0)|~neq(X1,s1)|~p(X1,s2,s2,s2,s2,s1,X2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_7]), c_0_6])])).
cnf(c_0_235,plain,(p(s1,s2,s2,s2,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_233])).
cnf(c_0_236,plain,(p(s2,s2,s2,s2,s2,s0,X1,X2,X3,X4)|~p(s2,s2,s2,s2,s2,s1,X1,X2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_7]), c_0_6])])).
cnf(c_0_237,plain,(p(X1,s2,s2,s2,s2,s1,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_235])).
cnf(c_0_238,plain,(p(s2,s2,s2,s2,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_236, c_0_237])).
cnf(c_0_239,plain,(p(X1,s2,s2,s2,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_238])).
cnf(c_0_240,plain,(p(s0,s1,s2,s2,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_239])).
cnf(c_0_241,plain,(p(X1,s1,s2,s2,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_240])).
cnf(c_0_242,plain,(p(s1,s1,s0,s2,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_241])).
cnf(c_0_243,plain,(p(X1,s1,s0,s2,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_242])).
cnf(c_0_244,plain,(p(s2,s0,s0,s2,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_243])).
cnf(c_0_245,plain,(p(X1,s0,s0,s2,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_244])).
cnf(c_0_246,plain,(p(s0,s0,s0,s1,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_61, c_0_245])).
cnf(c_0_247,plain,(p(X1,s0,s0,s1,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_246])).
cnf(c_0_248,plain,(p(s1,s2,s0,s1,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_247])).
cnf(c_0_249,plain,(p(X1,s2,s0,s1,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_248])).
cnf(c_0_250,plain,(p(s2,s2,s1,s1,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_249])).
cnf(c_0_251,plain,(p(X1,s2,s1,s1,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_250])).
cnf(c_0_252,plain,(p(s0,s1,s1,s1,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_251])).
cnf(c_0_253,plain,(p(X1,s1,s1,s1,s2,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_252])).
cnf(c_0_254,plain,(p(s1,s1,s1,s1,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_139, c_0_253])).
cnf(c_0_255,plain,(p(X1,s1,s1,s1,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_254])).
cnf(c_0_256,plain,(p(s2,s0,s1,s1,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_255])).
cnf(c_0_257,plain,(p(X1,s0,s1,s1,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_256])).
cnf(c_0_258,plain,(p(s0,s0,s2,s1,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_257])).
cnf(c_0_259,plain,(p(X1,s0,s2,s1,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_258])).
cnf(c_0_260,plain,(p(s1,s2,s2,s1,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_259])).
cnf(c_0_261,plain,(p(X1,s2,s2,s1,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_260])).
cnf(c_0_262,plain,(p(s2,s2,s2,s0,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_85, c_0_261])).
cnf(c_0_263,plain,(p(X1,s2,s2,s0,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_262])).
cnf(c_0_264,plain,(p(X1,X2,X3,X4,X5,s0,s2,X6,X7,X8)|~neq(s0,X9)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X5,X9)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,X5,s0,X9,X6,X7,X8)), inference(spm,[status(thm)],[c_0_13, c_0_10])).
cnf(c_0_265,plain,(p(s0,s1,s2,s0,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_263])).
cnf(c_0_266,plain,(p(X1,X2,X3,X4,s0,s0,s2,X5,X6,X7)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,X4,s0,s0,s1,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_8]), c_0_8]), c_0_10])])).
cnf(c_0_267,plain,(p(X1,s1,s2,s0,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_265])).
cnf(c_0_268,plain,(p(X1,X2,X3,s0,s0,s0,s2,X4,X5,X6)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,s0,s0,s0,s1,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_10]), c_0_8])])).
cnf(c_0_269,plain,(p(s1,s1,s0,s0,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_267])).
cnf(c_0_270,plain,(p(X1,X2,s0,s0,s0,s0,s2,X3,X4,X5)|~neq(X2,s2)|~neq(X1,s2)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,s0,s0,s0,s0,s1,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_10]), c_0_8])])).
cnf(c_0_271,plain,(p(X1,s1,s0,s0,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_269])).
cnf(c_0_272,plain,(p(X1,s0,s0,s0,s0,s0,s2,X2,X3,X4)|~neq(X1,s2)|~neq(X1,s1)|~p(X1,s0,s0,s0,s0,s0,s1,X2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_10]), c_0_8])])).
cnf(c_0_273,plain,(p(s2,s0,s0,s0,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_271])).
cnf(c_0_274,plain,(p(s0,s0,s0,s0,s0,s0,s2,X1,X2,X3)|~p(s0,s0,s0,s0,s0,s0,s1,X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_10]), c_0_8])])).
cnf(c_0_275,plain,(p(X1,s0,s0,s0,s0,s0,s1,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_273])).
cnf(c_0_276,plain,(p(s0,s0,s0,s0,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_274, c_0_275])).
cnf(c_0_277,plain,(p(X1,s0,s0,s0,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_276])).
cnf(c_0_278,plain,(p(s1,s2,s0,s0,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_277])).
cnf(c_0_279,plain,(p(X1,s2,s0,s0,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_278])).
cnf(c_0_280,plain,(p(s2,s2,s1,s0,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_279])).
cnf(c_0_281,plain,(p(X1,s2,s1,s0,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_280])).
cnf(c_0_282,plain,(p(s0,s1,s1,s0,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_281])).
cnf(c_0_283,plain,(p(X1,s1,s1,s0,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_282])).
cnf(c_0_284,plain,(p(s1,s1,s1,s2,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_32, c_0_283])).
cnf(c_0_285,plain,(p(X1,s1,s1,s2,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_284])).
cnf(c_0_286,plain,(p(s2,s0,s1,s2,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_285])).
cnf(c_0_287,plain,(p(X1,s0,s1,s2,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_286])).
cnf(c_0_288,plain,(p(s0,s0,s2,s2,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_287])).
cnf(c_0_289,plain,(p(X1,s0,s2,s2,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_288])).
cnf(c_0_290,plain,(p(s1,s2,s2,s2,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_289])).
cnf(c_0_291,plain,(p(X1,s2,s2,s2,s0,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_290])).
cnf(c_0_292,plain,(p(s2,s2,s2,s2,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_48, c_0_291])).
cnf(c_0_293,plain,(p(X1,s2,s2,s2,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_292])).
cnf(c_0_294,plain,(p(s0,s1,s2,s2,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_293])).
cnf(c_0_295,plain,(p(X1,s1,s2,s2,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_294])).
cnf(c_0_296,plain,(p(s1,s1,s0,s2,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_295])).
cnf(c_0_297,plain,(p(X1,s1,s0,s2,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_296])).
cnf(c_0_298,plain,(p(s2,s0,s0,s2,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_297])).
cnf(c_0_299,plain,(p(X1,s0,s0,s2,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_298])).
cnf(c_0_300,plain,(p(s0,s0,s0,s1,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_61, c_0_299])).
cnf(c_0_301,plain,(p(X1,s0,s0,s1,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_300])).
cnf(c_0_302,plain,(p(s1,s2,s0,s1,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_301])).
cnf(c_0_303,plain,(p(X1,s2,s0,s1,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_302])).
cnf(c_0_304,plain,(p(s2,s2,s1,s1,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_303])).
cnf(c_0_305,plain,(p(X1,s2,s1,s1,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_304])).
cnf(c_0_306,plain,(p(s0,s1,s1,s1,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_305])).
cnf(c_0_307,plain,(p(X1,s1,s1,s1,s1,s0,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_306])).
cnf(c_0_308,plain,(p(s1,s1,s1,s1,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_74, c_0_307])).
cnf(c_0_309,plain,(p(X1,s1,s1,s1,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_308])).
cnf(c_0_310,plain,(p(s2,s0,s1,s1,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_309])).
cnf(c_0_311,plain,(p(X1,s0,s1,s1,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_310])).
cnf(c_0_312,plain,(p(s0,s0,s2,s1,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_311])).
cnf(c_0_313,plain,(p(X1,s0,s2,s1,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_312])).
cnf(c_0_314,plain,(p(s1,s2,s2,s1,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_313])).
cnf(c_0_315,plain,(p(X1,s2,s2,s1,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_314])).
cnf(c_0_316,plain,(p(s2,s2,s2,s0,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_85, c_0_315])).
cnf(c_0_317,plain,(p(X1,s2,s2,s0,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_316])).
cnf(c_0_318,plain,(p(s0,s1,s2,s0,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_317])).
cnf(c_0_319,plain,(p(X1,s1,s2,s0,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_318])).
cnf(c_0_320,plain,(p(s1,s1,s0,s0,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_54, c_0_319])).
cnf(c_0_321,plain,(p(X1,s1,s0,s0,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_320])).
cnf(c_0_322,plain,(p(s2,s0,s0,s0,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_321])).
cnf(c_0_323,plain,(p(X1,s0,s0,s0,s1,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_322])).
cnf(c_0_324,plain,(p(s0,s0,s0,s0,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_97, c_0_323])).
cnf(c_0_325,plain,(p(X1,s0,s0,s0,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_324])).
cnf(c_0_326,plain,(p(s1,s2,s0,s0,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_325])).
cnf(c_0_327,plain,(p(X1,s2,s0,s0,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_326])).
cnf(c_0_328,plain,(p(s2,s2,s1,s0,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_23, c_0_327])).
cnf(c_0_329,plain,(p(X1,s2,s1,s0,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_328])).
cnf(c_0_330,plain,(p(s0,s1,s1,s0,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_28, c_0_329])).
cnf(c_0_331,plain,(p(X1,s1,s1,s0,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_330])).
cnf(c_0_332,plain,(p(X1,X2,X3,X4,X5,X6,X7,s2,s1,X8)|~neq(s2,X9)|~neq(X7,s1)|~neq(X6,s1)|~neq(X5,s1)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X7,X9)|~neq(X6,X9)|~neq(X5,X9)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,X5,X6,X7,s2,X9,X8)), inference(spm,[status(thm)],[c_0_15, c_0_6])).
cnf(c_0_333,plain,(p(s1,s1,s1,s2,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_32, c_0_331])).
cnf(c_0_334,plain,(p(X1,X2,X3,X4,X5,X6,s2,s2,s1,X7)|~neq(X6,s1)|~neq(X5,s1)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X6,s0)|~neq(X5,s0)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,X5,X6,s2,s2,s0,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_7]), c_0_7]), c_0_6])])).
cnf(c_0_335,plain,(p(X1,s1,s1,s2,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_333])).
cnf(c_0_336,plain,(p(X1,X2,X3,X4,X5,s2,s2,s2,s1,X6)|~neq(X5,s1)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X5,s0)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,X5,s2,s2,s2,s0,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_6]), c_0_7])])).
cnf(c_0_337,plain,(p(s2,s0,s1,s2,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_36, c_0_335])).
cnf(c_0_338,plain,(p(X1,X2,X3,X4,s2,s2,s2,s2,s1,X5)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,s2,s2,s2,s2,s0,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_6]), c_0_7])])).
cnf(c_0_339,plain,(p(X1,s0,s1,s2,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_337])).
cnf(c_0_340,plain,(p(X1,X2,X3,s2,s2,s2,s2,s2,s1,X4)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,s2,s2,s2,s2,s2,s0,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_6]), c_0_7])])).
cnf(c_0_341,plain,(p(s0,s0,s2,s2,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_40, c_0_339])).
cnf(c_0_342,plain,(p(X1,X2,s2,s2,s2,s2,s2,s2,s1,X3)|~neq(X2,s1)|~neq(X1,s1)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,s2,s2,s2,s2,s2,s2,s0,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_6]), c_0_7])])).
cnf(c_0_343,plain,(p(X1,s0,s2,s2,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_341])).
cnf(c_0_344,plain,(p(X1,s2,s2,s2,s2,s2,s2,s2,s1,X2)|~neq(X1,s1)|~neq(X1,s0)|~p(X1,s2,s2,s2,s2,s2,s2,s2,s0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_6]), c_0_7])])).
cnf(c_0_345,plain,(p(s1,s2,s2,s2,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_19, c_0_343])).
cnf(c_0_346,plain,(p(s2,s2,s2,s2,s2,s2,s2,s2,s1,X1)|~p(s2,s2,s2,s2,s2,s2,s2,s2,s0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_6]), c_0_7])])).
cnf(c_0_347,plain,(p(X1,s2,s2,s2,s2,s2,s2,s2,s0,s0)), inference(spm,[status(thm)],[c_0_3, c_0_345])).
cnf(c_0_348,plain,(p(s2,s2,s2,s2,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_346, c_0_347])).
cnf(c_0_349,plain,(p(X1,s2,s2,s2,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_348])).
cnf(c_0_350,plain,(p(s0,s1,s2,s2,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_349])).
cnf(c_0_351,plain,(p(X1,s1,s2,s2,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_350])).
cnf(c_0_352,plain,(p(s1,s1,s0,s2,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_351])).
cnf(c_0_353,plain,(p(X1,s1,s0,s2,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_352])).
cnf(c_0_354,plain,(p(s2,s0,s0,s2,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_353])).
cnf(c_0_355,plain,(p(X1,s0,s0,s2,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_354])).
cnf(c_0_356,plain,(p(s0,s0,s0,s1,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_61, c_0_355])).
cnf(c_0_357,plain,(p(X1,s0,s0,s1,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_356])).
cnf(c_0_358,plain,(p(s1,s2,s0,s1,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_357])).
cnf(c_0_359,plain,(p(X1,s2,s0,s1,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_358])).
cnf(c_0_360,plain,(p(s2,s2,s1,s1,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_359])).
cnf(c_0_361,plain,(p(X1,s2,s1,s1,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_360])).
cnf(c_0_362,plain,(p(s0,s1,s1,s1,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_361])).
cnf(c_0_363,plain,(p(X1,s1,s1,s1,s2,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_362])).
cnf(c_0_364,plain,(p(s1,s1,s1,s1,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_139, c_0_363])).
cnf(c_0_365,plain,(p(X1,s1,s1,s1,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_364])).
cnf(c_0_366,plain,(p(s2,s0,s1,s1,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_365])).
cnf(c_0_367,plain,(p(X1,s0,s1,s1,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_366])).
cnf(c_0_368,plain,(p(s0,s0,s2,s1,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_367])).
cnf(c_0_369,plain,(p(X1,s0,s2,s1,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_368])).
cnf(c_0_370,plain,(p(s1,s2,s2,s1,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_369])).
cnf(c_0_371,plain,(p(X1,s2,s2,s1,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_370])).
cnf(c_0_372,plain,(p(s2,s2,s2,s0,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_85, c_0_371])).
cnf(c_0_373,plain,(p(X1,s2,s2,s0,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_372])).
cnf(c_0_374,plain,(p(s0,s1,s2,s0,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_373])).
cnf(c_0_375,plain,(p(X1,s1,s2,s0,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_374])).
cnf(c_0_376,plain,(p(s1,s1,s0,s0,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_375])).
cnf(c_0_377,plain,(p(X1,s1,s0,s0,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_376])).
cnf(c_0_378,plain,(p(s2,s0,s0,s0,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_377])).
cnf(c_0_379,plain,(p(X1,s0,s0,s0,s0,s2,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_378])).
cnf(c_0_380,plain,(p(s0,s0,s0,s0,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_160, c_0_379])).
cnf(c_0_381,plain,(p(X1,s0,s0,s0,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_380])).
cnf(c_0_382,plain,(p(s1,s2,s0,s0,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_381])).
cnf(c_0_383,plain,(p(X1,s2,s0,s0,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_382])).
cnf(c_0_384,plain,(p(s2,s2,s1,s0,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_383])).
cnf(c_0_385,plain,(p(X1,s2,s1,s0,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_384])).
cnf(c_0_386,plain,(p(s0,s1,s1,s0,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_385])).
cnf(c_0_387,plain,(p(X1,s1,s1,s0,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_386])).
cnf(c_0_388,plain,(p(s1,s1,s1,s2,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_32, c_0_387])).
cnf(c_0_389,plain,(p(X1,s1,s1,s2,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_388])).
cnf(c_0_390,plain,(p(s2,s0,s1,s2,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_389])).
cnf(c_0_391,plain,(p(X1,s0,s1,s2,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_390])).
cnf(c_0_392,plain,(p(s0,s0,s2,s2,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_391])).
cnf(c_0_393,plain,(p(X1,s0,s2,s2,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_392])).
cnf(c_0_394,plain,(p(s1,s2,s2,s2,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_393])).
cnf(c_0_395,plain,(p(X1,s2,s2,s2,s0,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_394])).
cnf(c_0_396,plain,(p(s2,s2,s2,s2,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_48, c_0_395])).
cnf(c_0_397,plain,(p(X1,s2,s2,s2,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_396])).
cnf(c_0_398,plain,(p(s0,s1,s2,s2,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_397])).
cnf(c_0_399,plain,(p(X1,s1,s2,s2,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_398])).
cnf(c_0_400,plain,(p(s1,s1,s0,s2,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_399])).
cnf(c_0_401,plain,(p(X1,s1,s0,s2,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_400])).
cnf(c_0_402,plain,(p(s2,s0,s0,s2,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_401])).
cnf(c_0_403,plain,(p(X1,s0,s0,s2,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_402])).
cnf(c_0_404,plain,(p(s0,s0,s0,s1,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_61, c_0_403])).
cnf(c_0_405,plain,(p(X1,s0,s0,s1,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_404])).
cnf(c_0_406,plain,(p(X1,X2,X3,X4,X5,s1,s0,X6,X7,X8)|~neq(s1,X9)|~neq(X5,s0)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X5,X9)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,X5,s1,X9,X6,X7,X8)), inference(spm,[status(thm)],[c_0_13, c_0_2])).
cnf(c_0_407,plain,(p(s1,s2,s0,s1,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_405])).
cnf(c_0_408,plain,(p(X1,X2,X3,X4,s1,s1,s0,X5,X6,X7)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,X3,X4,s1,s1,s2,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_1]), c_0_1]), c_0_2])])).
cnf(c_0_409,plain,(p(X1,s2,s0,s1,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_407])).
cnf(c_0_410,plain,(p(X1,X2,X3,s1,s1,s1,s0,X4,X5,X6)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,X3,s1,s1,s1,s2,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_2]), c_0_1])])).
cnf(c_0_411,plain,(p(s2,s2,s1,s1,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_409])).
cnf(c_0_412,plain,(p(X1,X2,s1,s1,s1,s1,s0,X3,X4,X5)|~neq(X2,s0)|~neq(X1,s0)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,s1,s1,s1,s1,s2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_2]), c_0_1])])).
cnf(c_0_413,plain,(p(X1,s2,s1,s1,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_411])).
cnf(c_0_414,plain,(p(X1,s1,s1,s1,s1,s1,s0,X2,X3,X4)|~neq(X1,s0)|~neq(X1,s2)|~p(X1,s1,s1,s1,s1,s1,s2,X2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_2]), c_0_1])])).
cnf(c_0_415,plain,(p(s0,s1,s1,s1,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_413])).
cnf(c_0_416,plain,(p(s1,s1,s1,s1,s1,s1,s0,X1,X2,X3)|~p(s1,s1,s1,s1,s1,s1,s2,X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_2]), c_0_1])])).
cnf(c_0_417,plain,(p(X1,s1,s1,s1,s1,s1,s2,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_415])).
cnf(c_0_418,plain,(p(s1,s1,s1,s1,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_416, c_0_417])).
cnf(c_0_419,plain,(p(X1,s1,s1,s1,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_418])).
cnf(c_0_420,plain,(p(s2,s0,s1,s1,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_419])).
cnf(c_0_421,plain,(p(X1,s0,s1,s1,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_420])).
cnf(c_0_422,plain,(p(s0,s0,s2,s1,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_421])).
cnf(c_0_423,plain,(p(X1,s0,s2,s1,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_422])).
cnf(c_0_424,plain,(p(s1,s2,s2,s1,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_423])).
cnf(c_0_425,plain,(p(X1,s2,s2,s1,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_424])).
cnf(c_0_426,plain,(p(s2,s2,s2,s0,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_85, c_0_425])).
cnf(c_0_427,plain,(p(X1,s2,s2,s0,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_426])).
cnf(c_0_428,plain,(p(s0,s1,s2,s0,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_427])).
cnf(c_0_429,plain,(p(X1,s1,s2,s0,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_428])).
cnf(c_0_430,plain,(p(s1,s1,s0,s0,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_429])).
cnf(c_0_431,plain,(p(X1,s1,s0,s0,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_430])).
cnf(c_0_432,plain,(p(s2,s0,s0,s0,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_431])).
cnf(c_0_433,plain,(p(X1,s0,s0,s0,s1,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_432])).
cnf(c_0_434,plain,(p(s0,s0,s0,s0,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_97, c_0_433])).
cnf(c_0_435,plain,(p(X1,s0,s0,s0,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_434])).
cnf(c_0_436,plain,(p(s1,s2,s0,s0,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_435])).
cnf(c_0_437,plain,(p(X1,s2,s0,s0,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_436])).
cnf(c_0_438,plain,(p(s2,s2,s1,s0,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_437])).
cnf(c_0_439,plain,(p(X1,s2,s1,s0,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_438])).
cnf(c_0_440,plain,(p(s0,s1,s1,s0,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_439])).
cnf(c_0_441,plain,(p(X1,s1,s1,s0,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_440])).
cnf(c_0_442,plain,(p(s1,s1,s1,s2,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_32, c_0_441])).
cnf(c_0_443,plain,(p(X1,s1,s1,s2,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_442])).
cnf(c_0_444,plain,(p(s2,s0,s1,s2,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_443])).
cnf(c_0_445,plain,(p(X1,s0,s1,s2,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_444])).
cnf(c_0_446,plain,(p(s0,s0,s2,s2,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_445])).
cnf(c_0_447,plain,(p(X1,s0,s2,s2,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_446])).
cnf(c_0_448,plain,(p(s1,s2,s2,s2,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_447])).
cnf(c_0_449,plain,(p(X1,s2,s2,s2,s2,s1,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_448])).
cnf(c_0_450,plain,(p(s2,s2,s2,s2,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_236, c_0_449])).
cnf(c_0_451,plain,(p(X1,s2,s2,s2,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_450])).
cnf(c_0_452,plain,(p(s0,s1,s2,s2,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_451])).
cnf(c_0_453,plain,(p(X1,s1,s2,s2,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_452])).
cnf(c_0_454,plain,(p(s1,s1,s0,s2,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_453])).
cnf(c_0_455,plain,(p(X1,s1,s0,s2,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_454])).
cnf(c_0_456,plain,(p(s2,s0,s0,s2,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_455])).
cnf(c_0_457,plain,(p(X1,s0,s0,s2,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_456])).
cnf(c_0_458,plain,(p(s0,s0,s0,s1,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_61, c_0_457])).
cnf(c_0_459,plain,(p(X1,s0,s0,s1,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_458])).
cnf(c_0_460,plain,(p(s1,s2,s0,s1,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_459])).
cnf(c_0_461,plain,(p(X1,s2,s0,s1,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_460])).
cnf(c_0_462,plain,(p(s2,s2,s1,s1,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_461])).
cnf(c_0_463,plain,(p(X1,s2,s1,s1,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_462])).
cnf(c_0_464,plain,(p(s0,s1,s1,s1,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_463])).
cnf(c_0_465,plain,(p(X1,s1,s1,s1,s2,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_464])).
cnf(c_0_466,plain,(p(s1,s1,s1,s1,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_139, c_0_465])).
cnf(c_0_467,plain,(p(X1,s1,s1,s1,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_466])).
cnf(c_0_468,plain,(p(s2,s0,s1,s1,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_467])).
cnf(c_0_469,plain,(p(X1,s0,s1,s1,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_468])).
cnf(c_0_470,plain,(p(s0,s0,s2,s1,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_469])).
cnf(c_0_471,plain,(p(X1,s0,s2,s1,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_470])).
cnf(c_0_472,plain,(p(s1,s2,s2,s1,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_471])).
cnf(c_0_473,plain,(p(X1,s2,s2,s1,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_472])).
cnf(c_0_474,plain,(p(s2,s2,s2,s0,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_85, c_0_473])).
cnf(c_0_475,plain,(p(X1,X2,X3,X4,X5,X6,s0,s1,X7,X8)|~neq(s0,X9)|~neq(X6,s1)|~neq(X5,s1)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X6,X9)|~neq(X5,X9)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,X5,X6,s0,X9,X7,X8)), inference(spm,[status(thm)],[c_0_14, c_0_8])).
cnf(c_0_476,plain,(p(X1,s2,s2,s0,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_474])).
cnf(c_0_477,plain,(p(X1,X2,X3,X4,X5,s0,s0,s1,X6,X7)|~neq(X5,s1)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,X3,X4,X5,s0,s0,s2,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_475, c_0_10]), c_0_10]), c_0_8])])).
cnf(c_0_478,plain,(p(s0,s1,s2,s0,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_476])).
cnf(c_0_479,plain,(p(X1,X2,X3,X4,s0,s0,s0,s1,X5,X6)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,X3,X4,s0,s0,s0,s2,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_477, c_0_8]), c_0_10])])).
cnf(c_0_480,plain,(p(X1,s1,s2,s0,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_478])).
cnf(c_0_481,plain,(p(X1,X2,X3,s0,s0,s0,s0,s1,X4,X5)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,X3,s0,s0,s0,s0,s2,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_479, c_0_8]), c_0_10])])).
cnf(c_0_482,plain,(p(s1,s1,s0,s0,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_480])).
cnf(c_0_483,plain,(p(X1,X2,s0,s0,s0,s0,s0,s1,X3,X4)|~neq(X2,s1)|~neq(X1,s1)|~neq(X2,s2)|~neq(X1,s2)|~p(X1,X2,s0,s0,s0,s0,s0,s2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_8]), c_0_10])])).
cnf(c_0_484,plain,(p(X1,s1,s0,s0,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_482])).
cnf(c_0_485,plain,(p(X1,s0,s0,s0,s0,s0,s0,s1,X2,X3)|~neq(X1,s1)|~neq(X1,s2)|~p(X1,s0,s0,s0,s0,s0,s0,s2,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_483, c_0_8]), c_0_10])])).
cnf(c_0_486,plain,(p(s2,s0,s0,s0,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_484])).
cnf(c_0_487,plain,(p(s0,s0,s0,s0,s0,s0,s0,s1,X1,X2)|~p(s0,s0,s0,s0,s0,s0,s0,s2,X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_8]), c_0_10])])).
cnf(c_0_488,plain,(p(X1,s0,s0,s0,s0,s0,s0,s2,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_486])).
cnf(c_0_489,plain,(p(s0,s0,s0,s0,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_487, c_0_488])).
cnf(c_0_490,plain,(p(X1,s0,s0,s0,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_489])).
cnf(c_0_491,plain,(p(s1,s2,s0,s0,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_490])).
cnf(c_0_492,plain,(p(X1,s2,s0,s0,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_491])).
cnf(c_0_493,plain,(p(s2,s2,s1,s0,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_492])).
cnf(c_0_494,plain,(p(X1,s2,s1,s0,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_493])).
cnf(c_0_495,plain,(p(s0,s1,s1,s0,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_494])).
cnf(c_0_496,plain,(p(X1,s1,s1,s0,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_495])).
cnf(c_0_497,plain,(p(s1,s1,s1,s2,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_32, c_0_496])).
cnf(c_0_498,plain,(p(X1,s1,s1,s2,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_497])).
cnf(c_0_499,plain,(p(s2,s0,s1,s2,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_498])).
cnf(c_0_500,plain,(p(X1,s0,s1,s2,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_499])).
cnf(c_0_501,plain,(p(s0,s0,s2,s2,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_500])).
cnf(c_0_502,plain,(p(X1,s0,s2,s2,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_501])).
cnf(c_0_503,plain,(p(s1,s2,s2,s2,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_502])).
cnf(c_0_504,plain,(p(X1,s2,s2,s2,s0,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_503])).
cnf(c_0_505,plain,(p(s2,s2,s2,s2,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_48, c_0_504])).
cnf(c_0_506,plain,(p(X1,s2,s2,s2,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_505])).
cnf(c_0_507,plain,(p(s0,s1,s2,s2,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_506])).
cnf(c_0_508,plain,(p(X1,s1,s2,s2,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_507])).
cnf(c_0_509,plain,(p(s1,s1,s0,s2,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_508])).
cnf(c_0_510,plain,(p(X1,s1,s0,s2,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_509])).
cnf(c_0_511,plain,(p(s2,s0,s0,s2,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_510])).
cnf(c_0_512,plain,(p(X1,s0,s0,s2,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_511])).
cnf(c_0_513,plain,(p(s0,s0,s0,s1,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_61, c_0_512])).
cnf(c_0_514,plain,(p(X1,s0,s0,s1,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_513])).
cnf(c_0_515,plain,(p(s1,s2,s0,s1,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_514])).
cnf(c_0_516,plain,(p(X1,s2,s0,s1,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_515])).
cnf(c_0_517,plain,(p(s2,s2,s1,s1,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_516])).
cnf(c_0_518,plain,(p(X1,s2,s1,s1,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_517])).
cnf(c_0_519,plain,(p(s0,s1,s1,s1,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_518])).
cnf(c_0_520,plain,(p(X1,s1,s1,s1,s1,s0,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_519])).
cnf(c_0_521,plain,(p(s1,s1,s1,s1,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_74, c_0_520])).
cnf(c_0_522,plain,(p(X1,s1,s1,s1,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_521])).
cnf(c_0_523,plain,(p(s2,s0,s1,s1,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_522])).
cnf(c_0_524,plain,(p(X1,s0,s1,s1,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_523])).
cnf(c_0_525,plain,(p(s0,s0,s2,s1,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_524])).
cnf(c_0_526,plain,(p(X1,s0,s2,s1,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_525])).
cnf(c_0_527,plain,(p(s1,s2,s2,s1,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_526])).
cnf(c_0_528,plain,(p(X1,s2,s2,s1,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_527])).
cnf(c_0_529,plain,(p(s2,s2,s2,s0,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_85, c_0_528])).
cnf(c_0_530,plain,(p(X1,s2,s2,s0,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_529])).
cnf(c_0_531,plain,(p(s0,s1,s2,s0,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_530])).
cnf(c_0_532,plain,(p(X1,s1,s2,s0,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_531])).
cnf(c_0_533,plain,(p(s1,s1,s0,s0,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_532])).
cnf(c_0_534,plain,(p(X1,s1,s0,s0,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_533])).
cnf(c_0_535,plain,(p(s2,s0,s0,s0,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_534])).
cnf(c_0_536,plain,(p(X1,s0,s0,s0,s1,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_535])).
cnf(c_0_537,plain,(p(s0,s0,s0,s0,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_97, c_0_536])).
cnf(c_0_538,plain,(p(X1,s0,s0,s0,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_537])).
cnf(c_0_539,plain,(p(s1,s2,s0,s0,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_538])).
cnf(c_0_540,plain,(p(X1,s2,s0,s0,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_539])).
cnf(c_0_541,plain,(p(s2,s2,s1,s0,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_540])).
cnf(c_0_542,plain,(p(X1,s2,s1,s0,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_541])).
cnf(c_0_543,plain,(p(s0,s1,s1,s0,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_542])).
cnf(c_0_544,plain,(p(X1,s1,s1,s0,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_543])).
cnf(c_0_545,plain,(p(s1,s1,s1,s2,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_32, c_0_544])).
cnf(c_0_546,plain,(p(X1,s1,s1,s2,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_545])).
cnf(c_0_547,plain,(p(s2,s0,s1,s2,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_546])).
cnf(c_0_548,plain,(p(X1,s0,s1,s2,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_547])).
cnf(c_0_549,plain,(p(s0,s0,s2,s2,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_548])).
cnf(c_0_550,plain,(p(X1,s0,s2,s2,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_549])).
cnf(c_0_551,plain,(p(s1,s2,s2,s2,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_550])).
cnf(c_0_552,plain,(p(X1,s2,s2,s2,s2,s2,s0,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_551])).
cnf(c_0_553,plain,(p(s2,s2,s2,s2,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_119, c_0_552])).
cnf(c_0_554,plain,(p(X1,s2,s2,s2,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_553])).
cnf(c_0_555,plain,(p(s0,s1,s2,s2,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_554])).
cnf(c_0_556,plain,(p(X1,s1,s2,s2,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_555])).
cnf(c_0_557,plain,(p(s1,s1,s0,s2,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_556])).
cnf(c_0_558,plain,(p(X1,s1,s0,s2,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_557])).
cnf(c_0_559,plain,(p(s2,s0,s0,s2,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_558])).
cnf(c_0_560,plain,(p(X1,s0,s0,s2,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_559])).
cnf(c_0_561,plain,(p(s0,s0,s0,s1,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_61, c_0_560])).
cnf(c_0_562,plain,(p(X1,s0,s0,s1,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_561])).
cnf(c_0_563,plain,(p(s1,s2,s0,s1,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_562])).
cnf(c_0_564,plain,(p(X1,s2,s0,s1,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_563])).
cnf(c_0_565,plain,(p(s2,s2,s1,s1,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_564])).
cnf(c_0_566,plain,(p(X1,s2,s1,s1,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_565])).
cnf(c_0_567,plain,(p(s0,s1,s1,s1,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_566])).
cnf(c_0_568,plain,(p(X1,s1,s1,s1,s2,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_567])).
cnf(c_0_569,plain,(p(s1,s1,s1,s1,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_139, c_0_568])).
cnf(c_0_570,plain,(p(X1,s1,s1,s1,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_569])).
cnf(c_0_571,plain,(p(s2,s0,s1,s1,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_570])).
cnf(c_0_572,plain,(p(X1,s0,s1,s1,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_571])).
cnf(c_0_573,plain,(p(s0,s0,s2,s1,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_572])).
cnf(c_0_574,plain,(p(X1,s0,s2,s1,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_573])).
cnf(c_0_575,plain,(p(s1,s2,s2,s1,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_574])).
cnf(c_0_576,plain,(p(X1,s2,s2,s1,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_575])).
cnf(c_0_577,plain,(p(s2,s2,s2,s0,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_85, c_0_576])).
cnf(c_0_578,plain,(p(X1,s2,s2,s0,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_577])).
cnf(c_0_579,plain,(p(s0,s1,s2,s0,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_578])).
cnf(c_0_580,plain,(p(X1,s1,s2,s0,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_579])).
cnf(c_0_581,plain,(p(s1,s1,s0,s0,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_580])).
cnf(c_0_582,plain,(p(X1,s1,s0,s0,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_581])).
cnf(c_0_583,plain,(p(s2,s0,s0,s0,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_582])).
cnf(c_0_584,plain,(p(X1,s0,s0,s0,s0,s2,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_583])).
cnf(c_0_585,plain,(p(s0,s0,s0,s0,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_160, c_0_584])).
cnf(c_0_586,plain,(p(X1,s0,s0,s0,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_585])).
cnf(c_0_587,plain,(p(s1,s2,s0,s0,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_586])).
cnf(c_0_588,plain,(p(X1,s2,s0,s0,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_587])).
cnf(c_0_589,plain,(p(s2,s2,s1,s0,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_588])).
cnf(c_0_590,plain,(p(X1,s2,s1,s0,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_589])).
cnf(c_0_591,plain,(p(s0,s1,s1,s0,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_590])).
cnf(c_0_592,plain,(p(X1,s1,s1,s0,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_591])).
cnf(c_0_593,plain,(p(s1,s1,s1,s2,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_32, c_0_592])).
cnf(c_0_594,plain,(p(X1,s1,s1,s2,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_593])).
cnf(c_0_595,plain,(p(s2,s0,s1,s2,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_594])).
cnf(c_0_596,plain,(p(X1,s0,s1,s2,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_595])).
cnf(c_0_597,plain,(p(s0,s0,s2,s2,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_40, c_0_596])).
cnf(c_0_598,plain,(p(X1,s0,s2,s2,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_597])).
cnf(c_0_599,plain,(p(s1,s2,s2,s2,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_598])).
cnf(c_0_600,plain,(p(X1,s2,s2,s2,s0,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_599])).
cnf(c_0_601,plain,(p(s2,s2,s2,s2,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_48, c_0_600])).
cnf(c_0_602,plain,(p(X1,s2,s2,s2,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_601])).
cnf(c_0_603,plain,(p(s0,s1,s2,s2,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_602])).
cnf(c_0_604,plain,(p(X1,s1,s2,s2,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_603])).
cnf(c_0_605,plain,(p(s1,s1,s0,s2,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_54, c_0_604])).
cnf(c_0_606,plain,(p(X1,s1,s0,s2,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_605])).
cnf(c_0_607,plain,(p(s2,s0,s0,s2,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_36, c_0_606])).
cnf(c_0_608,plain,(p(X1,X2,X3,X4,X5,X6,X7,X8,s1,s2)|~neq(s1,X9)|~neq(X8,s2)|~neq(X7,s2)|~neq(X6,s2)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X8,X9)|~neq(X7,X9)|~neq(X6,X9)|~neq(X5,X9)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,X5,X6,X7,X8,s1,X9)), inference(spm,[status(thm)],[c_0_16, c_0_1])).
cnf(c_0_609,plain,(p(X1,s0,s0,s2,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_607])).
cnf(c_0_610,plain,(p(X1,X2,X3,X4,X5,X6,X7,s1,s1,s2)|~neq(X7,s2)|~neq(X6,s2)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X7,s0)|~neq(X6,s0)|~neq(X5,s0)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,X5,X6,X7,s1,s1,s0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_608, c_0_2]), c_0_2]), c_0_1])])).
cnf(c_0_611,plain,(p(s0,s0,s0,s1,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_61, c_0_609])).
cnf(c_0_612,plain,(p(X1,X2,X3,X4,X5,X6,s1,s1,s1,s2)|~neq(X6,s2)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X6,s0)|~neq(X5,s0)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,X5,X6,s1,s1,s1,s0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_610, c_0_1]), c_0_2])])).
cnf(c_0_613,plain,(p(X1,s0,s0,s1,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_611])).
cnf(c_0_614,plain,(p(X1,X2,X3,X4,X5,s1,s1,s1,s1,s2)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X5,s0)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,X5,s1,s1,s1,s1,s0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_612, c_0_1]), c_0_2])])).
cnf(c_0_615,plain,(p(s1,s2,s0,s1,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_19, c_0_613])).
cnf(c_0_616,plain,(p(X1,X2,X3,X4,s1,s1,s1,s1,s1,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,X4,s1,s1,s1,s1,s1,s0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_614, c_0_1]), c_0_2])])).
cnf(c_0_617,plain,(p(X1,s2,s0,s1,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_615])).
cnf(c_0_618,plain,(p(X1,X2,X3,s1,s1,s1,s1,s1,s1,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,X3,s1,s1,s1,s1,s1,s1,s0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_616, c_0_1]), c_0_2])])).
cnf(c_0_619,plain,(p(s2,s2,s1,s1,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_23, c_0_617])).
cnf(c_0_620,plain,(p(X1,X2,s1,s1,s1,s1,s1,s1,s1,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X2,s0)|~neq(X1,s0)|~p(X1,X2,s1,s1,s1,s1,s1,s1,s1,s0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_618, c_0_1]), c_0_2])])).
cnf(c_0_621,plain,(p(X1,s2,s1,s1,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_619])).
cnf(c_0_622,plain,(p(X1,s1,s1,s1,s1,s1,s1,s1,s1,s2)|~neq(X1,s2)|~neq(X1,s0)|~p(X1,s1,s1,s1,s1,s1,s1,s1,s1,s0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_620, c_0_1]), c_0_2])])).
cnf(c_0_623,plain,(p(s0,s1,s1,s1,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_28, c_0_621])).
cnf(c_0_624,plain,(p(s1,s1,s1,s1,s1,s1,s1,s1,s1,s2)|~p(s1,s1,s1,s1,s1,s1,s1,s1,s1,s0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_622, c_0_1]), c_0_2])])).
cnf(c_0_625,plain,(p(X1,s1,s1,s1,s1,s1,s1,s1,s1,s0)), inference(spm,[status(thm)],[c_0_3, c_0_623])).
cnf(c_0_626,plain,(p(s1,s1,s1,s1,s1,s1,s1,s1,s1,s2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_624, c_0_625])])).
cnf(c_0_627,plain,(p(X1,s1,s1,s1,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_626])).
cnf(c_0_628,plain,(p(s2,s0,s1,s1,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_627])).
cnf(c_0_629,plain,(p(X1,s0,s1,s1,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_628])).
cnf(c_0_630,plain,(p(s0,s0,s2,s1,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_629])).
cnf(c_0_631,plain,(p(X1,s0,s2,s1,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_630])).
cnf(c_0_632,plain,(p(s1,s2,s2,s1,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_631])).
cnf(c_0_633,plain,(p(X1,s2,s2,s1,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_632])).
cnf(c_0_634,plain,(p(s2,s2,s2,s0,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_85, c_0_633])).
cnf(c_0_635,plain,(p(X1,s2,s2,s0,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_634])).
cnf(c_0_636,plain,(p(s0,s1,s2,s0,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_635])).
cnf(c_0_637,plain,(p(X1,s1,s2,s0,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_636])).
cnf(c_0_638,plain,(p(s1,s1,s0,s0,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_637])).
cnf(c_0_639,plain,(p(X1,s1,s0,s0,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_638])).
cnf(c_0_640,plain,(p(s2,s0,s0,s0,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_639])).
cnf(c_0_641,plain,(p(X1,s0,s0,s0,s1,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_640])).
cnf(c_0_642,plain,(p(s0,s0,s0,s0,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_97, c_0_641])).
cnf(c_0_643,plain,(p(X1,s0,s0,s0,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_642])).
cnf(c_0_644,plain,(p(s1,s2,s0,s0,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_643])).
cnf(c_0_645,plain,(p(X1,s2,s0,s0,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_644])).
cnf(c_0_646,plain,(p(s2,s2,s1,s0,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_645])).
cnf(c_0_647,plain,(p(X1,s2,s1,s0,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_646])).
cnf(c_0_648,plain,(p(s0,s1,s1,s0,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_647])).
cnf(c_0_649,plain,(p(X1,s1,s1,s0,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_648])).
cnf(c_0_650,plain,(p(s1,s1,s1,s2,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_32, c_0_649])).
cnf(c_0_651,plain,(p(X1,s1,s1,s2,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_650])).
cnf(c_0_652,plain,(p(s2,s0,s1,s2,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_651])).
cnf(c_0_653,plain,(p(X1,s0,s1,s2,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_652])).
cnf(c_0_654,plain,(p(s0,s0,s2,s2,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_653])).
cnf(c_0_655,plain,(p(X1,s0,s2,s2,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_654])).
cnf(c_0_656,plain,(p(s1,s2,s2,s2,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_655])).
cnf(c_0_657,plain,(p(X1,s2,s2,s2,s2,s1,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_656])).
cnf(c_0_658,plain,(p(s2,s2,s2,s2,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_236, c_0_657])).
cnf(c_0_659,plain,(p(X1,s2,s2,s2,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_658])).
cnf(c_0_660,plain,(p(s0,s1,s2,s2,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_659])).
cnf(c_0_661,plain,(p(X1,s1,s2,s2,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_660])).
cnf(c_0_662,plain,(p(s1,s1,s0,s2,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_661])).
cnf(c_0_663,plain,(p(X1,s1,s0,s2,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_662])).
cnf(c_0_664,plain,(p(s2,s0,s0,s2,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_663])).
cnf(c_0_665,plain,(p(X1,s0,s0,s2,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_664])).
cnf(c_0_666,plain,(p(s0,s0,s0,s1,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_61, c_0_665])).
cnf(c_0_667,plain,(p(X1,s0,s0,s1,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_666])).
cnf(c_0_668,plain,(p(s1,s2,s0,s1,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_667])).
cnf(c_0_669,plain,(p(X1,s2,s0,s1,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_668])).
cnf(c_0_670,plain,(p(s2,s2,s1,s1,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_669])).
cnf(c_0_671,plain,(p(X1,s2,s1,s1,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_670])).
cnf(c_0_672,plain,(p(s0,s1,s1,s1,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_671])).
cnf(c_0_673,plain,(p(X1,s1,s1,s1,s2,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_672])).
cnf(c_0_674,plain,(p(s1,s1,s1,s1,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_139, c_0_673])).
cnf(c_0_675,plain,(p(X1,s1,s1,s1,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_674])).
cnf(c_0_676,plain,(p(s2,s0,s1,s1,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_675])).
cnf(c_0_677,plain,(p(X1,s0,s1,s1,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_676])).
cnf(c_0_678,plain,(p(s0,s0,s2,s1,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_677])).
cnf(c_0_679,plain,(p(X1,s0,s2,s1,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_678])).
cnf(c_0_680,plain,(p(s1,s2,s2,s1,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_679])).
cnf(c_0_681,plain,(p(X1,s2,s2,s1,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_680])).
cnf(c_0_682,plain,(p(s2,s2,s2,s0,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_85, c_0_681])).
cnf(c_0_683,plain,(p(X1,s2,s2,s0,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_682])).
cnf(c_0_684,plain,(p(s0,s1,s2,s0,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_683])).
cnf(c_0_685,plain,(p(X1,s1,s2,s0,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_684])).
cnf(c_0_686,plain,(p(s1,s1,s0,s0,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_685])).
cnf(c_0_687,plain,(p(X1,s1,s0,s0,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_686])).
cnf(c_0_688,plain,(p(s2,s0,s0,s0,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_687])).
cnf(c_0_689,plain,(p(X1,s0,s0,s0,s0,s0,s1,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_688])).
cnf(c_0_690,plain,(p(s0,s0,s0,s0,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_274, c_0_689])).
cnf(c_0_691,plain,(p(X1,s0,s0,s0,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_690])).
cnf(c_0_692,plain,(p(s1,s2,s0,s0,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_691])).
cnf(c_0_693,plain,(p(X1,s2,s0,s0,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_692])).
cnf(c_0_694,plain,(p(s2,s2,s1,s0,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_693])).
cnf(c_0_695,plain,(p(X1,s2,s1,s0,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_694])).
cnf(c_0_696,plain,(p(s0,s1,s1,s0,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_695])).
cnf(c_0_697,plain,(p(X1,s1,s1,s0,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_696])).
cnf(c_0_698,plain,(p(s1,s1,s1,s2,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_32, c_0_697])).
cnf(c_0_699,plain,(p(X1,s1,s1,s2,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_698])).
cnf(c_0_700,plain,(p(s2,s0,s1,s2,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_699])).
cnf(c_0_701,plain,(p(X1,s0,s1,s2,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_700])).
cnf(c_0_702,plain,(p(s0,s0,s2,s2,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_701])).
cnf(c_0_703,plain,(p(X1,s0,s2,s2,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_702])).
cnf(c_0_704,plain,(p(s1,s2,s2,s2,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_703])).
cnf(c_0_705,plain,(p(X1,s2,s2,s2,s0,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_704])).
cnf(c_0_706,plain,(p(s2,s2,s2,s2,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_48, c_0_705])).
cnf(c_0_707,plain,(p(X1,s2,s2,s2,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_706])).
cnf(c_0_708,plain,(p(s0,s1,s2,s2,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_707])).
cnf(c_0_709,plain,(p(X1,s1,s2,s2,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_708])).
cnf(c_0_710,plain,(p(s1,s1,s0,s2,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_709])).
cnf(c_0_711,plain,(p(X1,s1,s0,s2,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_710])).
cnf(c_0_712,plain,(p(s2,s0,s0,s2,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_711])).
cnf(c_0_713,plain,(p(X1,s0,s0,s2,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_712])).
cnf(c_0_714,plain,(p(s0,s0,s0,s1,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_61, c_0_713])).
cnf(c_0_715,plain,(p(X1,s0,s0,s1,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_714])).
cnf(c_0_716,plain,(p(s1,s2,s0,s1,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_715])).
cnf(c_0_717,plain,(p(X1,s2,s0,s1,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_716])).
cnf(c_0_718,plain,(p(s2,s2,s1,s1,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_717])).
cnf(c_0_719,plain,(p(X1,s2,s1,s1,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_718])).
cnf(c_0_720,plain,(p(s0,s1,s1,s1,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_719])).
cnf(c_0_721,plain,(p(X1,s1,s1,s1,s1,s0,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_720])).
cnf(c_0_722,plain,(p(s1,s1,s1,s1,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_74, c_0_721])).
cnf(c_0_723,plain,(p(X1,s1,s1,s1,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_722])).
cnf(c_0_724,plain,(p(s2,s0,s1,s1,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_723])).
cnf(c_0_725,plain,(p(X1,s0,s1,s1,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_724])).
cnf(c_0_726,plain,(p(s0,s0,s2,s1,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_725])).
cnf(c_0_727,plain,(p(X1,s0,s2,s1,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_726])).
cnf(c_0_728,plain,(p(s1,s2,s2,s1,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_727])).
cnf(c_0_729,plain,(p(X1,s2,s2,s1,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_728])).
cnf(c_0_730,plain,(p(s2,s2,s2,s0,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_85, c_0_729])).
cnf(c_0_731,plain,(p(X1,s2,s2,s0,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_730])).
cnf(c_0_732,plain,(p(s0,s1,s2,s0,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_731])).
cnf(c_0_733,plain,(p(X1,s1,s2,s0,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_732])).
cnf(c_0_734,plain,(p(s1,s1,s0,s0,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_733])).
cnf(c_0_735,plain,(p(X1,s1,s0,s0,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_734])).
cnf(c_0_736,plain,(p(s2,s0,s0,s0,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_735])).
cnf(c_0_737,plain,(p(X1,s0,s0,s0,s1,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_736])).
cnf(c_0_738,plain,(p(s0,s0,s0,s0,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_97, c_0_737])).
cnf(c_0_739,plain,(p(X1,s0,s0,s0,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_738])).
cnf(c_0_740,plain,(p(s1,s2,s0,s0,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_739])).
cnf(c_0_741,plain,(p(X1,s2,s0,s0,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_740])).
cnf(c_0_742,plain,(p(s2,s2,s1,s0,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_741])).
cnf(c_0_743,plain,(p(X1,s2,s1,s0,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_742])).
cnf(c_0_744,plain,(p(s0,s1,s1,s0,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_743])).
cnf(c_0_745,plain,(p(X1,s1,s1,s0,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_744])).
cnf(c_0_746,plain,(p(s1,s1,s1,s2,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_32, c_0_745])).
cnf(c_0_747,plain,(p(X1,X2,X3,X4,X5,X6,s2,s0,X7,X8)|~neq(s2,X9)|~neq(X6,s0)|~neq(X5,s0)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X6,X9)|~neq(X5,X9)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,X5,X6,s2,X9,X7,X8)), inference(spm,[status(thm)],[c_0_14, c_0_7])).
cnf(c_0_748,plain,(p(X1,s1,s1,s2,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_746])).
cnf(c_0_749,plain,(p(X1,X2,X3,X4,X5,s2,s2,s0,X6,X7)|~neq(X5,s0)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X5,s1)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,X4,X5,s2,s2,s1,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_747, c_0_6]), c_0_6]), c_0_7])])).
cnf(c_0_750,plain,(p(s2,s0,s1,s2,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_748])).
cnf(c_0_751,plain,(p(X1,X2,X3,X4,s2,s2,s2,s0,X5,X6)|~neq(X4,s0)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,X4,s2,s2,s2,s1,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_749, c_0_7]), c_0_6])])).
cnf(c_0_752,plain,(p(X1,s0,s1,s2,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_750])).
cnf(c_0_753,plain,(p(X1,X2,X3,s2,s2,s2,s2,s0,X4,X5)|~neq(X3,s0)|~neq(X2,s0)|~neq(X1,s0)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,s2,s2,s2,s2,s1,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_751, c_0_7]), c_0_6])])).
cnf(c_0_754,plain,(p(s0,s0,s2,s2,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_752])).
cnf(c_0_755,plain,(p(X1,X2,s2,s2,s2,s2,s2,s0,X3,X4)|~neq(X2,s0)|~neq(X1,s0)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,s2,s2,s2,s2,s2,s1,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_753, c_0_7]), c_0_6])])).
cnf(c_0_756,plain,(p(X1,s0,s2,s2,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_754])).
cnf(c_0_757,plain,(p(X1,s2,s2,s2,s2,s2,s2,s0,X2,X3)|~neq(X1,s0)|~neq(X1,s1)|~p(X1,s2,s2,s2,s2,s2,s2,s1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_755, c_0_7]), c_0_6])])).
cnf(c_0_758,plain,(p(s1,s2,s2,s2,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_756])).
cnf(c_0_759,plain,(p(s2,s2,s2,s2,s2,s2,s2,s0,X1,X2)|~p(s2,s2,s2,s2,s2,s2,s2,s1,X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_757, c_0_7]), c_0_6])])).
cnf(c_0_760,plain,(p(X1,s2,s2,s2,s2,s2,s2,s1,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_758])).
cnf(c_0_761,plain,(p(s2,s2,s2,s2,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_759, c_0_760])).
cnf(c_0_762,plain,(p(X1,s2,s2,s2,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_761])).
cnf(c_0_763,plain,(p(s0,s1,s2,s2,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_762])).
cnf(c_0_764,plain,(p(X1,s1,s2,s2,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_763])).
cnf(c_0_765,plain,(p(s1,s1,s0,s2,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_764])).
cnf(c_0_766,plain,(p(X1,s1,s0,s2,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_765])).
cnf(c_0_767,plain,(p(s2,s0,s0,s2,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_766])).
cnf(c_0_768,plain,(p(X1,s0,s0,s2,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_767])).
cnf(c_0_769,plain,(p(s0,s0,s0,s1,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_61, c_0_768])).
cnf(c_0_770,plain,(p(X1,s0,s0,s1,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_769])).
cnf(c_0_771,plain,(p(s1,s2,s0,s1,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_770])).
cnf(c_0_772,plain,(p(X1,s2,s0,s1,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_771])).
cnf(c_0_773,plain,(p(s2,s2,s1,s1,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_772])).
cnf(c_0_774,plain,(p(X1,s2,s1,s1,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_773])).
cnf(c_0_775,plain,(p(s0,s1,s1,s1,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_774])).
cnf(c_0_776,plain,(p(X1,s1,s1,s1,s2,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_775])).
cnf(c_0_777,plain,(p(s1,s1,s1,s1,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_139, c_0_776])).
cnf(c_0_778,plain,(p(X1,s1,s1,s1,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_777])).
cnf(c_0_779,plain,(p(s2,s0,s1,s1,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_778])).
cnf(c_0_780,plain,(p(X1,s0,s1,s1,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_779])).
cnf(c_0_781,plain,(p(s0,s0,s2,s1,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_780])).
cnf(c_0_782,plain,(p(X1,s0,s2,s1,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_781])).
cnf(c_0_783,plain,(p(s1,s2,s2,s1,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_782])).
cnf(c_0_784,plain,(p(X1,s2,s2,s1,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_783])).
cnf(c_0_785,plain,(p(s2,s2,s2,s0,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_85, c_0_784])).
cnf(c_0_786,plain,(p(X1,s2,s2,s0,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_785])).
cnf(c_0_787,plain,(p(s0,s1,s2,s0,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_786])).
cnf(c_0_788,plain,(p(X1,s1,s2,s0,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_787])).
cnf(c_0_789,plain,(p(s1,s1,s0,s0,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_788])).
cnf(c_0_790,plain,(p(X1,s1,s0,s0,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_789])).
cnf(c_0_791,plain,(p(s2,s0,s0,s0,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_790])).
cnf(c_0_792,plain,(p(X1,s0,s0,s0,s0,s2,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_791])).
cnf(c_0_793,plain,(p(s0,s0,s0,s0,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_160, c_0_792])).
cnf(c_0_794,plain,(p(X1,s0,s0,s0,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_793])).
cnf(c_0_795,plain,(p(s1,s2,s0,s0,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_794])).
cnf(c_0_796,plain,(p(X1,s2,s0,s0,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_795])).
cnf(c_0_797,plain,(p(s2,s2,s1,s0,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_796])).
cnf(c_0_798,plain,(p(X1,s2,s1,s0,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_797])).
cnf(c_0_799,plain,(p(s0,s1,s1,s0,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_798])).
cnf(c_0_800,plain,(p(X1,s1,s1,s0,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_799])).
cnf(c_0_801,plain,(p(s1,s1,s1,s2,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_32, c_0_800])).
cnf(c_0_802,plain,(p(X1,s1,s1,s2,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_801])).
cnf(c_0_803,plain,(p(s2,s0,s1,s2,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_802])).
cnf(c_0_804,plain,(p(X1,s0,s1,s2,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_803])).
cnf(c_0_805,plain,(p(s0,s0,s2,s2,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_804])).
cnf(c_0_806,plain,(p(X1,s0,s2,s2,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_805])).
cnf(c_0_807,plain,(p(s1,s2,s2,s2,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_806])).
cnf(c_0_808,plain,(p(X1,s2,s2,s2,s0,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_807])).
cnf(c_0_809,plain,(p(s2,s2,s2,s2,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_48, c_0_808])).
cnf(c_0_810,plain,(p(X1,s2,s2,s2,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_809])).
cnf(c_0_811,plain,(p(s0,s1,s2,s2,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_810])).
cnf(c_0_812,plain,(p(X1,s1,s2,s2,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_811])).
cnf(c_0_813,plain,(p(s1,s1,s0,s2,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_812])).
cnf(c_0_814,plain,(p(X1,s1,s0,s2,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_813])).
cnf(c_0_815,plain,(p(s2,s0,s0,s2,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_814])).
cnf(c_0_816,plain,(p(X1,s0,s0,s2,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_815])).
cnf(c_0_817,plain,(p(s0,s0,s0,s1,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_61, c_0_816])).
cnf(c_0_818,plain,(p(X1,s0,s0,s1,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_817])).
cnf(c_0_819,plain,(p(s1,s2,s0,s1,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_818])).
cnf(c_0_820,plain,(p(X1,s2,s0,s1,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_819])).
cnf(c_0_821,plain,(p(s2,s2,s1,s1,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_820])).
cnf(c_0_822,plain,(p(X1,s2,s1,s1,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_821])).
cnf(c_0_823,plain,(p(s0,s1,s1,s1,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_822])).
cnf(c_0_824,plain,(p(X1,s1,s1,s1,s1,s1,s2,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_823])).
cnf(c_0_825,plain,(p(s1,s1,s1,s1,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_416, c_0_824])).
cnf(c_0_826,plain,(p(X1,s1,s1,s1,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_825])).
cnf(c_0_827,plain,(p(s2,s0,s1,s1,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_826])).
cnf(c_0_828,plain,(p(X1,s0,s1,s1,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_827])).
cnf(c_0_829,plain,(p(s0,s0,s2,s1,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_828])).
cnf(c_0_830,plain,(p(X1,s0,s2,s1,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_829])).
cnf(c_0_831,plain,(p(s1,s2,s2,s1,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_830])).
cnf(c_0_832,plain,(p(X1,s2,s2,s1,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_831])).
cnf(c_0_833,plain,(p(s2,s2,s2,s0,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_85, c_0_832])).
cnf(c_0_834,plain,(p(X1,s2,s2,s0,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_833])).
cnf(c_0_835,plain,(p(s0,s1,s2,s0,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_834])).
cnf(c_0_836,plain,(p(X1,s1,s2,s0,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_835])).
cnf(c_0_837,plain,(p(s1,s1,s0,s0,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_836])).
cnf(c_0_838,plain,(p(X1,s1,s0,s0,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_837])).
cnf(c_0_839,plain,(p(s2,s0,s0,s0,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_838])).
cnf(c_0_840,plain,(p(X1,s0,s0,s0,s1,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_839])).
cnf(c_0_841,plain,(p(s0,s0,s0,s0,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_97, c_0_840])).
cnf(c_0_842,plain,(p(X1,s0,s0,s0,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_841])).
cnf(c_0_843,plain,(p(s1,s2,s0,s0,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_842])).
cnf(c_0_844,plain,(p(X1,s2,s0,s0,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_843])).
cnf(c_0_845,plain,(p(s2,s2,s1,s0,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_844])).
cnf(c_0_846,plain,(p(X1,s2,s1,s0,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_845])).
cnf(c_0_847,plain,(p(s0,s1,s1,s0,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_846])).
cnf(c_0_848,plain,(p(X1,s1,s1,s0,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_847])).
cnf(c_0_849,plain,(p(s1,s1,s1,s2,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_32, c_0_848])).
cnf(c_0_850,plain,(p(X1,s1,s1,s2,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_849])).
cnf(c_0_851,plain,(p(s2,s0,s1,s2,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_850])).
cnf(c_0_852,plain,(p(X1,s0,s1,s2,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_851])).
cnf(c_0_853,plain,(p(s0,s0,s2,s2,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_852])).
cnf(c_0_854,plain,(p(X1,s0,s2,s2,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_853])).
cnf(c_0_855,plain,(p(s1,s2,s2,s2,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_854])).
cnf(c_0_856,plain,(p(X1,s2,s2,s2,s2,s1,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_855])).
cnf(c_0_857,plain,(p(s2,s2,s2,s2,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_236, c_0_856])).
cnf(c_0_858,plain,(p(X1,s2,s2,s2,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_857])).
cnf(c_0_859,plain,(p(s0,s1,s2,s2,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_858])).
cnf(c_0_860,plain,(p(X1,s1,s2,s2,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_859])).
cnf(c_0_861,plain,(p(s1,s1,s0,s2,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_860])).
cnf(c_0_862,plain,(p(X1,s1,s0,s2,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_861])).
cnf(c_0_863,plain,(p(s2,s0,s0,s2,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_862])).
cnf(c_0_864,plain,(p(X1,s0,s0,s2,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_863])).
cnf(c_0_865,plain,(p(s0,s0,s0,s1,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_61, c_0_864])).
cnf(c_0_866,plain,(p(X1,s0,s0,s1,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_865])).
cnf(c_0_867,plain,(p(s1,s2,s0,s1,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_866])).
cnf(c_0_868,plain,(p(X1,s2,s0,s1,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_867])).
cnf(c_0_869,plain,(p(s2,s2,s1,s1,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_23, c_0_868])).
cnf(c_0_870,plain,(p(X1,s2,s1,s1,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_869])).
cnf(c_0_871,plain,(p(s0,s1,s1,s1,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_870])).
cnf(c_0_872,plain,(p(X1,s1,s1,s1,s2,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_871])).
cnf(c_0_873,plain,(p(s1,s1,s1,s1,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_139, c_0_872])).
cnf(c_0_874,plain,(p(X1,s1,s1,s1,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_873])).
cnf(c_0_875,plain,(p(s2,s0,s1,s1,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_874])).
cnf(c_0_876,plain,(p(X1,s0,s1,s1,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_875])).
cnf(c_0_877,plain,(p(s0,s0,s2,s1,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_40, c_0_876])).
cnf(c_0_878,plain,(p(X1,s0,s2,s1,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_877])).
cnf(c_0_879,plain,(p(s1,s2,s2,s1,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_19, c_0_878])).
cnf(c_0_880,plain,(p(X1,s2,s2,s1,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_879])).
cnf(c_0_881,plain,(p(X1,X2,X3,X4,X5,X6,X7,s0,s2,X8)|~neq(s0,X9)|~neq(X7,s2)|~neq(X6,s2)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X7,X9)|~neq(X6,X9)|~neq(X5,X9)|~neq(X4,X9)|~neq(X3,X9)|~neq(X2,X9)|~neq(X1,X9)|~p(X1,X2,X3,X4,X5,X6,X7,s0,X9,X8)), inference(spm,[status(thm)],[c_0_15, c_0_10])).
cnf(c_0_882,plain,(p(s2,s2,s2,s0,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_85, c_0_880])).
cnf(c_0_883,plain,(p(X1,X2,X3,X4,X5,X6,s0,s0,s2,X7)|~neq(X6,s2)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X6,s1)|~neq(X5,s1)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,X4,X5,X6,s0,s0,s1,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_881, c_0_8]), c_0_8]), c_0_10])])).
cnf(c_0_884,plain,(p(X1,s2,s2,s0,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_882])).
cnf(c_0_885,plain,(p(X1,X2,X3,X4,X5,s0,s0,s0,s2,X6)|~neq(X5,s2)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X5,s1)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,X4,X5,s0,s0,s0,s1,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_883, c_0_10]), c_0_8])])).
cnf(c_0_886,plain,(p(s0,s1,s2,s0,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_28, c_0_884])).
cnf(c_0_887,plain,(p(X1,X2,X3,X4,s0,s0,s0,s0,s2,X5)|~neq(X4,s2)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X4,s1)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,X4,s0,s0,s0,s0,s1,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_885, c_0_10]), c_0_8])])).
cnf(c_0_888,plain,(p(X1,s1,s2,s0,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_886])).
cnf(c_0_889,plain,(p(X1,X2,X3,s0,s0,s0,s0,s0,s2,X4)|~neq(X3,s2)|~neq(X2,s2)|~neq(X1,s2)|~neq(X3,s1)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,X3,s0,s0,s0,s0,s0,s1,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_887, c_0_10]), c_0_8])])).
cnf(c_0_890,plain,(p(s1,s1,s0,s0,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_54, c_0_888])).
cnf(c_0_891,plain,(p(X1,X2,s0,s0,s0,s0,s0,s0,s2,X3)|~neq(X2,s2)|~neq(X1,s2)|~neq(X2,s1)|~neq(X1,s1)|~p(X1,X2,s0,s0,s0,s0,s0,s0,s1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_889, c_0_10]), c_0_8])])).
cnf(c_0_892,plain,(p(X1,s1,s0,s0,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_890])).
cnf(c_0_893,plain,(p(X1,s0,s0,s0,s0,s0,s0,s0,s2,X2)|~neq(X1,s2)|~neq(X1,s1)|~p(X1,s0,s0,s0,s0,s0,s0,s0,s1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_891, c_0_10]), c_0_8])])).
cnf(c_0_894,plain,(p(s2,s0,s0,s0,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_36, c_0_892])).
cnf(c_0_895,plain,(p(s0,s0,s0,s0,s0,s0,s0,s0,s2,X1)|~p(s0,s0,s0,s0,s0,s0,s0,s0,s1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_893, c_0_10]), c_0_8])])).
cnf(c_0_896,plain,(p(X1,s0,s0,s0,s0,s0,s0,s0,s1,s2)), inference(spm,[status(thm)],[c_0_3, c_0_894])).
cnf(c_0_897,plain,(p(s0,s0,s0,s0,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_895, c_0_896])).
cnf(c_0_898,plain,(p(X1,s0,s0,s0,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_897])).
cnf(c_0_899,plain,(p(s1,s2,s0,s0,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_898])).
cnf(c_0_900,plain,(p(X1,s2,s0,s0,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_899])).
cnf(c_0_901,plain,(p(s2,s2,s1,s0,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_900])).
cnf(c_0_902,plain,(p(X1,s2,s1,s0,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_901])).
cnf(c_0_903,plain,(p(s0,s1,s1,s0,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_902])).
cnf(c_0_904,plain,(p(X1,s1,s1,s0,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_903])).
cnf(c_0_905,plain,(p(s1,s1,s1,s2,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_32, c_0_904])).
cnf(c_0_906,plain,(p(X1,s1,s1,s2,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_905])).
cnf(c_0_907,plain,(p(s2,s0,s1,s2,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_906])).
cnf(c_0_908,plain,(p(X1,s0,s1,s2,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_907])).
cnf(c_0_909,plain,(p(s0,s0,s2,s2,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_908])).
cnf(c_0_910,plain,(p(X1,s0,s2,s2,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_909])).
cnf(c_0_911,plain,(p(s1,s2,s2,s2,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_910])).
cnf(c_0_912,plain,(p(X1,s2,s2,s2,s0,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_911])).
cnf(c_0_913,plain,(p(s2,s2,s2,s2,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_48, c_0_912])).
cnf(c_0_914,plain,(p(X1,s2,s2,s2,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_913])).
cnf(c_0_915,plain,(p(s0,s1,s2,s2,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_914])).
cnf(c_0_916,plain,(p(X1,s1,s2,s2,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_915])).
cnf(c_0_917,plain,(p(s1,s1,s0,s2,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_916])).
cnf(c_0_918,plain,(p(X1,s1,s0,s2,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_917])).
cnf(c_0_919,plain,(p(s2,s0,s0,s2,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_918])).
cnf(c_0_920,plain,(p(X1,s0,s0,s2,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_919])).
cnf(c_0_921,plain,(p(s0,s0,s0,s1,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_61, c_0_920])).
cnf(c_0_922,plain,(p(X1,s0,s0,s1,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_921])).
cnf(c_0_923,plain,(p(s1,s2,s0,s1,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_922])).
cnf(c_0_924,plain,(p(X1,s2,s0,s1,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_923])).
cnf(c_0_925,plain,(p(s2,s2,s1,s1,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_924])).
cnf(c_0_926,plain,(p(X1,s2,s1,s1,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_925])).
cnf(c_0_927,plain,(p(s0,s1,s1,s1,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_926])).
cnf(c_0_928,plain,(p(X1,s1,s1,s1,s1,s0,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_927])).
cnf(c_0_929,plain,(p(s1,s1,s1,s1,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_74, c_0_928])).
cnf(c_0_930,plain,(p(X1,s1,s1,s1,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_929])).
cnf(c_0_931,plain,(p(s2,s0,s1,s1,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_930])).
cnf(c_0_932,plain,(p(X1,s0,s1,s1,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_931])).
cnf(c_0_933,plain,(p(s0,s0,s2,s1,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_932])).
cnf(c_0_934,plain,(p(X1,s0,s2,s1,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_933])).
cnf(c_0_935,plain,(p(s1,s2,s2,s1,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_934])).
cnf(c_0_936,plain,(p(X1,s2,s2,s1,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_935])).
cnf(c_0_937,plain,(p(s2,s2,s2,s0,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_85, c_0_936])).
cnf(c_0_938,plain,(p(X1,s2,s2,s0,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_937])).
cnf(c_0_939,plain,(p(s0,s1,s2,s0,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_938])).
cnf(c_0_940,plain,(p(X1,s1,s2,s0,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_939])).
cnf(c_0_941,plain,(p(s1,s1,s0,s0,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_940])).
cnf(c_0_942,plain,(p(X1,s1,s0,s0,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_941])).
cnf(c_0_943,plain,(p(s2,s0,s0,s0,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_942])).
cnf(c_0_944,plain,(p(X1,s0,s0,s0,s1,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_943])).
cnf(c_0_945,plain,(p(s0,s0,s0,s0,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_97, c_0_944])).
cnf(c_0_946,plain,(p(X1,s0,s0,s0,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_945])).
cnf(c_0_947,plain,(p(s1,s2,s0,s0,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_946])).
cnf(c_0_948,plain,(p(X1,s2,s0,s0,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_947])).
cnf(c_0_949,plain,(p(s2,s2,s1,s0,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_948])).
cnf(c_0_950,plain,(p(X1,s2,s1,s0,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_949])).
cnf(c_0_951,plain,(p(s0,s1,s1,s0,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_950])).
cnf(c_0_952,plain,(p(X1,s1,s1,s0,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_951])).
cnf(c_0_953,plain,(p(s1,s1,s1,s2,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_32, c_0_952])).
cnf(c_0_954,plain,(p(X1,s1,s1,s2,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_953])).
cnf(c_0_955,plain,(p(s2,s0,s1,s2,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_954])).
cnf(c_0_956,plain,(p(X1,s0,s1,s2,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_955])).
cnf(c_0_957,plain,(p(s0,s0,s2,s2,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_956])).
cnf(c_0_958,plain,(p(X1,s0,s2,s2,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_957])).
cnf(c_0_959,plain,(p(s1,s2,s2,s2,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_958])).
cnf(c_0_960,plain,(p(X1,s2,s2,s2,s2,s2,s0,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_959])).
cnf(c_0_961,plain,(p(s2,s2,s2,s2,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_119, c_0_960])).
cnf(c_0_962,plain,(p(X1,s2,s2,s2,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_961])).
cnf(c_0_963,plain,(p(s0,s1,s2,s2,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_962])).
cnf(c_0_964,plain,(p(X1,s1,s2,s2,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_963])).
cnf(c_0_965,plain,(p(s1,s1,s0,s2,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_964])).
cnf(c_0_966,plain,(p(X1,s1,s0,s2,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_965])).
cnf(c_0_967,plain,(p(s2,s0,s0,s2,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_966])).
cnf(c_0_968,plain,(p(X1,s0,s0,s2,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_967])).
cnf(c_0_969,plain,(p(s0,s0,s0,s1,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_61, c_0_968])).
cnf(c_0_970,plain,(p(X1,s0,s0,s1,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_969])).
cnf(c_0_971,plain,(p(s1,s2,s0,s1,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_970])).
cnf(c_0_972,plain,(p(X1,s2,s0,s1,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_971])).
cnf(c_0_973,plain,(p(s2,s2,s1,s1,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_972])).
cnf(c_0_974,plain,(p(X1,s2,s1,s1,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_973])).
cnf(c_0_975,plain,(p(s0,s1,s1,s1,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_974])).
cnf(c_0_976,plain,(p(X1,s1,s1,s1,s2,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_975])).
cnf(c_0_977,plain,(p(s1,s1,s1,s1,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_139, c_0_976])).
cnf(c_0_978,plain,(p(X1,s1,s1,s1,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_977])).
cnf(c_0_979,plain,(p(s2,s0,s1,s1,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_978])).
cnf(c_0_980,plain,(p(X1,s0,s1,s1,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_979])).
cnf(c_0_981,plain,(p(s0,s0,s2,s1,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_980])).
cnf(c_0_982,plain,(p(X1,s0,s2,s1,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_981])).
cnf(c_0_983,plain,(p(s1,s2,s2,s1,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_982])).
cnf(c_0_984,plain,(p(X1,s2,s2,s1,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_983])).
cnf(c_0_985,plain,(p(s2,s2,s2,s0,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_85, c_0_984])).
cnf(c_0_986,plain,(p(X1,s2,s2,s0,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_985])).
cnf(c_0_987,plain,(p(s0,s1,s2,s0,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_986])).
cnf(c_0_988,plain,(p(X1,s1,s2,s0,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_987])).
cnf(c_0_989,plain,(p(s1,s1,s0,s0,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_988])).
cnf(c_0_990,plain,(p(X1,s1,s0,s0,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_989])).
cnf(c_0_991,plain,(p(s2,s0,s0,s0,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_990])).
cnf(c_0_992,plain,(p(X1,s0,s0,s0,s0,s2,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_991])).
cnf(c_0_993,plain,(p(s0,s0,s0,s0,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_160, c_0_992])).
cnf(c_0_994,plain,(p(X1,s0,s0,s0,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_993])).
cnf(c_0_995,plain,(p(s1,s2,s0,s0,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_994])).
cnf(c_0_996,plain,(p(X1,s2,s0,s0,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_995])).
cnf(c_0_997,plain,(p(s2,s2,s1,s0,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_996])).
cnf(c_0_998,plain,(p(X1,s2,s1,s0,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_997])).
cnf(c_0_999,plain,(p(s0,s1,s1,s0,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_998])).
cnf(c_0_1000,plain,(p(X1,s1,s1,s0,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_999])).
cnf(c_0_1001,plain,(p(s1,s1,s1,s2,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_32, c_0_1000])).
cnf(c_0_1002,plain,(p(X1,s1,s1,s2,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1001])).
cnf(c_0_1003,plain,(p(s2,s0,s1,s2,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1002])).
cnf(c_0_1004,plain,(p(X1,s0,s1,s2,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1003])).
cnf(c_0_1005,plain,(p(s0,s0,s2,s2,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_1004])).
cnf(c_0_1006,plain,(p(X1,s0,s2,s2,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1005])).
cnf(c_0_1007,plain,(p(s1,s2,s2,s2,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1006])).
cnf(c_0_1008,plain,(p(X1,s2,s2,s2,s0,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1007])).
cnf(c_0_1009,plain,(p(s2,s2,s2,s2,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_48, c_0_1008])).
cnf(c_0_1010,plain,(p(X1,s2,s2,s2,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1009])).
cnf(c_0_1011,plain,(p(s0,s1,s2,s2,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1010])).
cnf(c_0_1012,plain,(p(X1,s1,s2,s2,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1011])).
cnf(c_0_1013,plain,(p(s1,s1,s0,s2,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_1012])).
cnf(c_0_1014,plain,(p(X1,s1,s0,s2,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1013])).
cnf(c_0_1015,plain,(p(s2,s0,s0,s2,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1014])).
cnf(c_0_1016,plain,(p(X1,s0,s0,s2,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1015])).
cnf(c_0_1017,plain,(p(s0,s0,s0,s1,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_61, c_0_1016])).
cnf(c_0_1018,plain,(p(X1,s0,s0,s1,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1017])).
cnf(c_0_1019,plain,(p(s1,s2,s0,s1,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1018])).
cnf(c_0_1020,plain,(p(X1,s2,s0,s1,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1019])).
cnf(c_0_1021,plain,(p(s2,s2,s1,s1,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_1020])).
cnf(c_0_1022,plain,(p(X1,s2,s1,s1,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1021])).
cnf(c_0_1023,plain,(p(s0,s1,s1,s1,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1022])).
cnf(c_0_1024,plain,(p(X1,s1,s1,s1,s1,s1,s1,s0,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1023])).
cnf(c_0_1025,plain,(p(s1,s1,s1,s1,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_199, c_0_1024])).
cnf(c_0_1026,plain,(p(X1,s1,s1,s1,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1025])).
cnf(c_0_1027,plain,(p(s2,s0,s1,s1,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1026])).
cnf(c_0_1028,plain,(p(X1,s0,s1,s1,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1027])).
cnf(c_0_1029,plain,(p(s0,s0,s2,s1,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_1028])).
cnf(c_0_1030,plain,(p(X1,s0,s2,s1,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1029])).
cnf(c_0_1031,plain,(p(s1,s2,s2,s1,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1030])).
cnf(c_0_1032,plain,(p(X1,s2,s2,s1,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1031])).
cnf(c_0_1033,plain,(p(s2,s2,s2,s0,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_85, c_0_1032])).
cnf(c_0_1034,plain,(p(X1,s2,s2,s0,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1033])).
cnf(c_0_1035,plain,(p(s0,s1,s2,s0,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1034])).
cnf(c_0_1036,plain,(p(X1,s1,s2,s0,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1035])).
cnf(c_0_1037,plain,(p(s1,s1,s0,s0,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_1036])).
cnf(c_0_1038,plain,(p(X1,s1,s0,s0,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1037])).
cnf(c_0_1039,plain,(p(s2,s0,s0,s0,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1038])).
cnf(c_0_1040,plain,(p(X1,s0,s0,s0,s1,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1039])).
cnf(c_0_1041,plain,(p(s0,s0,s0,s0,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_97, c_0_1040])).
cnf(c_0_1042,plain,(p(X1,s0,s0,s0,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1041])).
cnf(c_0_1043,plain,(p(s1,s2,s0,s0,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1042])).
cnf(c_0_1044,plain,(p(X1,s2,s0,s0,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1043])).
cnf(c_0_1045,plain,(p(s2,s2,s1,s0,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_1044])).
cnf(c_0_1046,plain,(p(X1,s2,s1,s0,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1045])).
cnf(c_0_1047,plain,(p(s0,s1,s1,s0,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1046])).
cnf(c_0_1048,plain,(p(X1,s1,s1,s0,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1047])).
cnf(c_0_1049,plain,(p(s1,s1,s1,s2,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_32, c_0_1048])).
cnf(c_0_1050,plain,(p(X1,s1,s1,s2,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1049])).
cnf(c_0_1051,plain,(p(s2,s0,s1,s2,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1050])).
cnf(c_0_1052,plain,(p(X1,s0,s1,s2,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1051])).
cnf(c_0_1053,plain,(p(s0,s0,s2,s2,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_1052])).
cnf(c_0_1054,plain,(p(X1,s0,s2,s2,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1053])).
cnf(c_0_1055,plain,(p(s1,s2,s2,s2,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1054])).
cnf(c_0_1056,plain,(p(X1,s2,s2,s2,s2,s1,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1055])).
cnf(c_0_1057,plain,(p(s2,s2,s2,s2,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_236, c_0_1056])).
cnf(c_0_1058,plain,(p(X1,s2,s2,s2,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1057])).
cnf(c_0_1059,plain,(p(s0,s1,s2,s2,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1058])).
cnf(c_0_1060,plain,(p(X1,s1,s2,s2,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1059])).
cnf(c_0_1061,plain,(p(s1,s1,s0,s2,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_1060])).
cnf(c_0_1062,plain,(p(X1,s1,s0,s2,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1061])).
cnf(c_0_1063,plain,(p(s2,s0,s0,s2,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1062])).
cnf(c_0_1064,plain,(p(X1,s0,s0,s2,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1063])).
cnf(c_0_1065,plain,(p(s0,s0,s0,s1,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_61, c_0_1064])).
cnf(c_0_1066,plain,(p(X1,s0,s0,s1,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1065])).
cnf(c_0_1067,plain,(p(s1,s2,s0,s1,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1066])).
cnf(c_0_1068,plain,(p(X1,s2,s0,s1,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1067])).
cnf(c_0_1069,plain,(p(s2,s2,s1,s1,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_1068])).
cnf(c_0_1070,plain,(p(X1,s2,s1,s1,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1069])).
cnf(c_0_1071,plain,(p(s0,s1,s1,s1,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1070])).
cnf(c_0_1072,plain,(p(X1,s1,s1,s1,s2,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1071])).
cnf(c_0_1073,plain,(p(s1,s1,s1,s1,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_139, c_0_1072])).
cnf(c_0_1074,plain,(p(X1,s1,s1,s1,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1073])).
cnf(c_0_1075,plain,(p(s2,s0,s1,s1,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1074])).
cnf(c_0_1076,plain,(p(X1,s0,s1,s1,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1075])).
cnf(c_0_1077,plain,(p(s0,s0,s2,s1,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_1076])).
cnf(c_0_1078,plain,(p(X1,s0,s2,s1,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1077])).
cnf(c_0_1079,plain,(p(s1,s2,s2,s1,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1078])).
cnf(c_0_1080,plain,(p(X1,s2,s2,s1,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1079])).
cnf(c_0_1081,plain,(p(s2,s2,s2,s0,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_85, c_0_1080])).
cnf(c_0_1082,plain,(p(X1,s2,s2,s0,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1081])).
cnf(c_0_1083,plain,(p(s0,s1,s2,s0,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1082])).
cnf(c_0_1084,plain,(p(X1,s1,s2,s0,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1083])).
cnf(c_0_1085,plain,(p(s1,s1,s0,s0,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_1084])).
cnf(c_0_1086,plain,(p(X1,s1,s0,s0,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1085])).
cnf(c_0_1087,plain,(p(s2,s0,s0,s0,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1086])).
cnf(c_0_1088,plain,(p(X1,s0,s0,s0,s0,s0,s1,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1087])).
cnf(c_0_1089,plain,(p(s0,s0,s0,s0,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_274, c_0_1088])).
cnf(c_0_1090,plain,(p(X1,s0,s0,s0,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1089])).
cnf(c_0_1091,plain,(p(s1,s2,s0,s0,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1090])).
cnf(c_0_1092,plain,(p(X1,s2,s0,s0,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1091])).
cnf(c_0_1093,plain,(p(s2,s2,s1,s0,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_1092])).
cnf(c_0_1094,plain,(p(X1,s2,s1,s0,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1093])).
cnf(c_0_1095,plain,(p(s0,s1,s1,s0,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1094])).
cnf(c_0_1096,plain,(p(X1,s1,s1,s0,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1095])).
cnf(c_0_1097,plain,(p(s1,s1,s1,s2,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_32, c_0_1096])).
cnf(c_0_1098,plain,(p(X1,s1,s1,s2,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1097])).
cnf(c_0_1099,plain,(p(s2,s0,s1,s2,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1098])).
cnf(c_0_1100,plain,(p(X1,s0,s1,s2,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1099])).
cnf(c_0_1101,plain,(p(s0,s0,s2,s2,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_1100])).
cnf(c_0_1102,plain,(p(X1,s0,s2,s2,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1101])).
cnf(c_0_1103,plain,(p(s1,s2,s2,s2,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1102])).
cnf(c_0_1104,plain,(p(X1,s2,s2,s2,s0,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1103])).
cnf(c_0_1105,plain,(p(s2,s2,s2,s2,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_48, c_0_1104])).
cnf(c_0_1106,plain,(p(X1,s2,s2,s2,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1105])).
cnf(c_0_1107,plain,(p(s0,s1,s2,s2,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1106])).
cnf(c_0_1108,plain,(p(X1,s1,s2,s2,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1107])).
cnf(c_0_1109,plain,(p(s1,s1,s0,s2,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_1108])).
cnf(c_0_1110,plain,(p(X1,s1,s0,s2,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1109])).
cnf(c_0_1111,plain,(p(s2,s0,s0,s2,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1110])).
cnf(c_0_1112,plain,(p(X1,s0,s0,s2,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1111])).
cnf(c_0_1113,plain,(p(s0,s0,s0,s1,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_61, c_0_1112])).
cnf(c_0_1114,plain,(p(X1,s0,s0,s1,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1113])).
cnf(c_0_1115,plain,(p(s1,s2,s0,s1,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1114])).
cnf(c_0_1116,plain,(p(X1,s2,s0,s1,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1115])).
cnf(c_0_1117,plain,(p(s2,s2,s1,s1,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_1116])).
cnf(c_0_1118,plain,(p(X1,s2,s1,s1,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1117])).
cnf(c_0_1119,plain,(p(s0,s1,s1,s1,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1118])).
cnf(c_0_1120,plain,(p(X1,s1,s1,s1,s1,s0,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1119])).
cnf(c_0_1121,plain,(p(s1,s1,s1,s1,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_74, c_0_1120])).
cnf(c_0_1122,plain,(p(X1,s1,s1,s1,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1121])).
cnf(c_0_1123,plain,(p(s2,s0,s1,s1,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1122])).
cnf(c_0_1124,plain,(p(X1,s0,s1,s1,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1123])).
cnf(c_0_1125,plain,(p(s0,s0,s2,s1,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_1124])).
cnf(c_0_1126,plain,(p(X1,s0,s2,s1,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1125])).
cnf(c_0_1127,plain,(p(s1,s2,s2,s1,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1126])).
cnf(c_0_1128,plain,(p(X1,s2,s2,s1,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1127])).
cnf(c_0_1129,plain,(p(s2,s2,s2,s0,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_85, c_0_1128])).
cnf(c_0_1130,plain,(p(X1,s2,s2,s0,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1129])).
cnf(c_0_1131,plain,(p(s0,s1,s2,s0,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1130])).
cnf(c_0_1132,plain,(p(X1,s1,s2,s0,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1131])).
cnf(c_0_1133,plain,(p(s1,s1,s0,s0,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_54, c_0_1132])).
cnf(c_0_1134,plain,(p(X1,s1,s0,s0,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1133])).
cnf(c_0_1135,plain,(p(s2,s0,s0,s0,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1134])).
cnf(c_0_1136,plain,(p(X1,s0,s0,s0,s1,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1135])).
cnf(c_0_1137,plain,(p(s0,s0,s0,s0,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_97, c_0_1136])).
cnf(c_0_1138,plain,(p(X1,s0,s0,s0,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1137])).
cnf(c_0_1139,plain,(p(s1,s2,s0,s0,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1138])).
cnf(c_0_1140,plain,(p(X1,s2,s0,s0,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1139])).
cnf(c_0_1141,plain,(p(s2,s2,s1,s0,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_23, c_0_1140])).
cnf(c_0_1142,plain,(p(X1,s2,s1,s0,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1141])).
cnf(c_0_1143,plain,(p(s0,s1,s1,s0,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_28, c_0_1142])).
cnf(c_0_1144,plain,(p(X1,s1,s1,s0,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1143])).
cnf(c_0_1145,plain,(p(s1,s1,s1,s2,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_32, c_0_1144])).
cnf(c_0_1146,plain,(p(X1,s1,s1,s2,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1145])).
cnf(c_0_1147,plain,(p(s2,s0,s1,s2,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_36, c_0_1146])).
cnf(c_0_1148,plain,(p(X1,s0,s1,s2,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1147])).
cnf(c_0_1149,plain,(p(s0,s0,s2,s2,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_40, c_0_1148])).
cnf(c_0_1150,plain,(p(X1,s0,s2,s2,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1149])).
cnf(c_0_1151,plain,(p(s1,s2,s2,s2,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_19, c_0_1150])).
cnf(c_0_1152,plain,(p(X1,s2,s2,s2,s2,s2,s2,s2,s2,s2)), inference(spm,[status(thm)],[c_0_3, c_0_1151])).
cnf(c_0_1153,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_1152])]), ['proof']).
# SZS output end CNFRefutation.
EOF
