% Problem    : GRP128-2.006 : TPTP v6.1.0. Released v1.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n119.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:44:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____011_C07_F1_PI_AE_OS_S1U
# and selection function SelectComplexAHPExceptRRHorn.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,axiom,(product(X1,X2,e_1)|product(X1,X2,e_2)|product(X1,X2,e_3)|product(X1,X2,e_4)|product(X1,X2,e_5)|product(X1,X2,e_6)|~group_element(X1)|~group_element(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', product_total_function1)).
cnf(c_0_1,axiom,(group_element(e_1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', element_1)).
cnf(c_0_2,axiom,(group_element(e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', element_2)).
cnf(c_0_3,axiom,(~product(X1,e_1,X2)|~next(X1,X3)|~greater(X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', no_redundancy)).
cnf(c_0_4,axiom,(group_element(e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', element_3)).
cnf(c_0_5,axiom,(greater(e_6,e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_6_greater_e_3)).
cnf(c_0_6,axiom,(next(e_2,e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_2_then_e_3)).
cnf(c_0_7,axiom,(equalish(X1,X4)|~product(X1,X2,X3)|~product(X4,X2,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', product_left_cancellation)).
cnf(c_0_8,axiom,(greater(e_6,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_6_greater_e_2)).
cnf(c_0_9,axiom,(next(e_1,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_1_then_e_2)).
cnf(c_0_10,axiom,(greater(e_6,e_4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_6_greater_e_4)).
cnf(c_0_11,axiom,(next(e_3,e_4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_3_then_e_4)).
cnf(c_0_12,axiom,(~equalish(e_1,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_1_is_not_e_2)).
cnf(c_0_13,axiom,(group_element(e_4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', element_4)).
cnf(c_0_14,axiom,(~equalish(e_1,e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_1_is_not_e_3)).
cnf(c_0_15,axiom,(greater(e_5,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_5_greater_e_2)).
cnf(c_0_16,axiom,(greater(e_6,e_5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_6_greater_e_5)).
cnf(c_0_17,axiom,(next(e_4,e_5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_4_then_e_5)).
cnf(c_0_18,axiom,(greater(e_4,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_4_greater_e_2)).
cnf(c_0_19,axiom,(~equalish(e_1,e_4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_1_is_not_e_4)).
cnf(c_0_20,axiom,(greater(e_3,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_3_greater_e_2)).
cnf(c_0_21,negated_conjecture,(product(X1,X3,X4)|~product(X1,X2,X3)|~product(X3,X2,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', qg3)).
cnf(c_0_22,axiom,(greater(e_5,e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_5_greater_e_3)).
cnf(c_0_23,axiom,(greater(e_5,e_4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_5_greater_e_4)).
cnf(c_0_24,axiom,(group_element(e_5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', element_5)).
cnf(c_0_25,axiom,(greater(e_4,e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_4_greater_e_3)).
cnf(c_0_26,axiom,(~equalish(e_3,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_3_is_not_e_2)).
cnf(c_0_27,axiom,(~equalish(e_5,e_1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_5_is_not_e_1)).
cnf(c_0_28,axiom,(equalish(X2,X4)|~product(X1,X2,X3)|~product(X1,X4,X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', product_right_cancellation)).
cnf(c_0_29,axiom,(group_element(e_6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', element_6)).
cnf(c_0_30,axiom,(~equalish(e_6,e_1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_6_is_not_e_1)).
cnf(c_0_31,axiom,(~equalish(e_5,e_4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_5_is_not_e_4)).
cnf(c_0_32,axiom,(~equalish(e_5,e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_5_is_not_e_3)).
cnf(c_0_33,axiom,(~equalish(e_6,e_4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_6_is_not_e_4)).
cnf(c_0_34,axiom,(~equalish(e_6,e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_6_is_not_e_3)).
cnf(c_0_35,axiom,(~equalish(e_5,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_5_is_not_e_2)).
cnf(c_0_36,axiom,(~equalish(e_6,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_6_is_not_e_2)).
cnf(c_0_37,axiom,(~equalish(e_6,e_5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_6_is_not_e_5)).
cnf(c_0_38,axiom,(~equalish(e_1,e_5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_1_is_not_e_5)).
cnf(c_0_39,axiom,(~equalish(e_4,e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_4_is_not_e_3)).
cnf(c_0_40,axiom,(~equalish(e_4,e_2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_4_is_not_e_2)).
cnf(c_0_41,axiom,(~equalish(e_3,e_1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_3_is_not_e_1)).
cnf(c_0_42,axiom,(equalish(X3,X4)|~product(X1,X2,X3)|~product(X1,X2,X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', product_total_function2)).
cnf(c_0_43,axiom,(~equalish(e_2,e_1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_2_is_not_e_1)).
cnf(c_0_44,axiom,(~equalish(e_2,e_5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_2_is_not_e_5)).
cnf(c_0_45,axiom,(~equalish(e_2,e_4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_2_is_not_e_4)).
cnf(c_0_46,axiom,(~equalish(e_4,e_6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_4_is_not_e_6)).
cnf(c_0_47,axiom,(~equalish(e_4,e_5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_4_is_not_e_5)).
cnf(c_0_48,axiom,(~equalish(e_3,e_6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_3_is_not_e_6)).
cnf(c_0_49,axiom,(~equalish(e_5,e_6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_5_is_not_e_6)).
cnf(c_0_50,axiom,(~equalish(e_1,e_6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_1_is_not_e_6)).
cnf(c_0_51,axiom,(~equalish(e_4,e_1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_4_is_not_e_1)).
cnf(c_0_52,axiom,(~equalish(e_2,e_3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_2_is_not_e_3)).
cnf(c_0_53,axiom,(~equalish(e_3,e_4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_3_is_not_e_4)).
cnf(c_0_54,axiom,(~equalish(e_3,e_5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_3_is_not_e_5)).
cnf(c_0_55,axiom,(~equalish(e_2,e_6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', e_2_is_not_e_6)).
cnf(c_0_56,plain,(product(X1,e_1,e_1)|product(X1,e_1,e_2)|product(X1,e_1,e_3)|product(X1,e_1,e_4)|product(X1,e_1,e_5)|product(X1,e_1,e_6)|~group_element(X1)), inference(spm,[status(thm)],[c_0_0, c_0_1])).
cnf(c_0_57,plain,(product(e_2,e_1,e_1)|product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(e_2,e_1,e_6)), inference(spm,[status(thm)],[c_0_56, c_0_2])).
cnf(c_0_58,plain,(product(e_2,e_1,e_1)|product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|~greater(e_6,X1)|~next(e_2,X1)), inference(spm,[status(thm)],[c_0_3, c_0_57])).
cnf(c_0_59,plain,(product(e_1,e_1,e_1)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|product(e_1,e_1,e_6)), inference(spm,[status(thm)],[c_0_56, c_0_1])).
cnf(c_0_60,plain,(product(e_3,e_1,e_1)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|product(e_3,e_1,e_6)), inference(spm,[status(thm)],[c_0_56, c_0_4])).
cnf(c_0_61,plain,(product(e_2,e_1,e_1)|product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_5]), c_0_6])])).
cnf(c_0_62,plain,(product(e_1,e_1,e_1)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|~greater(e_6,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_59])).
cnf(c_0_63,plain,(product(e_3,e_1,e_1)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|~greater(e_6,X1)|~next(e_3,X1)), inference(spm,[status(thm)],[c_0_3, c_0_60])).
cnf(c_0_64,plain,(equalish(X1,e_2)|product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|~product(X1,e_1,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_61])).
cnf(c_0_65,plain,(product(e_1,e_1,e_1)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_8]), c_0_9])])).
cnf(c_0_66,plain,(product(e_3,e_1,e_1)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_10]), c_0_11])])).
cnf(c_0_67,plain,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_12])).
cnf(c_0_68,plain,(equalish(X1,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|~product(X1,e_1,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_66])).
cnf(c_0_69,plain,(product(e_4,e_1,e_1)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|product(e_4,e_1,e_6)), inference(spm,[status(thm)],[c_0_56, c_0_13])).
cnf(c_0_70,plain,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|~greater(e_5,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_67])).
cnf(c_0_71,plain,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_65]), c_0_14])).
cnf(c_0_72,plain,(product(e_4,e_1,e_1)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|~greater(e_6,X1)|~next(e_4,X1)), inference(spm,[status(thm)],[c_0_3, c_0_69])).
cnf(c_0_73,plain,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_15]), c_0_9])])).
cnf(c_0_74,plain,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|~greater(e_5,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_71])).
cnf(c_0_75,plain,(product(e_4,e_1,e_1)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_16]), c_0_17])])).
cnf(c_0_76,plain,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|~greater(e_4,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_73])).
cnf(c_0_77,plain,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_15]), c_0_9])])).
cnf(c_0_78,plain,(equalish(X1,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|~product(X1,e_1,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_75])).
cnf(c_0_79,plain,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_18]), c_0_9])])).
cnf(c_0_80,plain,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|~greater(e_4,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_77])).
cnf(c_0_81,plain,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_65]), c_0_19])).
cnf(c_0_82,plain,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(e_1,e_1,e_2)|~greater(e_3,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_79])).
cnf(c_0_83,plain,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_18]), c_0_9])])).
cnf(c_0_84,plain,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|~greater(e_5,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_81])).
cnf(c_0_85,plain,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(e_1,e_1,e_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_20]), c_0_9])])).
cnf(c_0_86,plain,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|product(e_1,e_1,e_2)|~greater(e_3,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_83])).
cnf(c_0_87,plain,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_15]), c_0_9])])).
cnf(c_0_88,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(X1,e_1,e_2)|~product(X1,e_1,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_85])).
cnf(c_0_89,plain,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|product(e_1,e_1,e_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_20]), c_0_9])])).
cnf(c_0_90,plain,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|~greater(e_4,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_87])).
cnf(c_0_91,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)), inference(spm,[status(thm)],[c_0_88, c_0_61])).
cnf(c_0_92,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)|product(X1,e_1,e_2)|~product(X1,e_1,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_89])).
cnf(c_0_93,plain,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_18]), c_0_9])])).
cnf(c_0_94,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)|~greater(e_5,X1)|~next(e_2,X1)), inference(spm,[status(thm)],[c_0_3, c_0_91])).
cnf(c_0_95,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|product(e_3,e_1,e_5)), inference(spm,[status(thm)],[c_0_92, c_0_66])).
cnf(c_0_96,plain,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|product(e_1,e_1,e_2)|~greater(e_3,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_93])).
cnf(c_0_97,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|product(e_2,e_1,e_4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_22]), c_0_6])])).
cnf(c_0_98,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)|~greater(e_5,X1)|~next(e_3,X1)), inference(spm,[status(thm)],[c_0_3, c_0_95])).
cnf(c_0_99,plain,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|product(e_1,e_1,e_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_20]), c_0_9])])).
cnf(c_0_100,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)|~greater(e_4,X1)|~next(e_2,X1)), inference(spm,[status(thm)],[c_0_3, c_0_97])).
cnf(c_0_101,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_1,e_4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_23]), c_0_11])])).
cnf(c_0_102,plain,(product(e_5,e_1,e_1)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_5,e_1,e_6)), inference(spm,[status(thm)],[c_0_56, c_0_24])).
cnf(c_0_103,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)|product(X1,e_1,e_2)|~product(X1,e_1,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_99])).
cnf(c_0_104,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_1,e_3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_25]), c_0_6])])).
cnf(c_0_105,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(X1,e_3,e_4)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_101])).
cnf(c_0_106,negated_conjecture,(product(e_5,e_1,e_1)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(X1,e_5,e_6)|~product(X1,e_1,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_102])).
cnf(c_0_107,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_1,e_5)), inference(spm,[status(thm)],[c_0_103, c_0_75])).
cnf(c_0_108,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_1,e_2)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_104])).
cnf(c_0_109,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_1,e_2)), inference(spm,[status(thm)],[c_0_105, c_0_104])).
cnf(c_0_110,plain,(equalish(X1,e_1)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|~product(X1,e_1,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_65])).
cnf(c_0_111,negated_conjecture,(product(e_5,e_1,e_1)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_112,negated_conjecture,(product(e_2,e_1,e_2)|product(X1,e_2,e_3)|~product(X1,e_1,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_104])).
cnf(c_0_113,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_2,e_3,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_26])).
cnf(c_0_114,negated_conjecture,(equalish(X1,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|~product(X1,e_1,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_107])).
cnf(c_0_115,negated_conjecture,(product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_27])).
cnf(c_0_116,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_4)), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_117,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_19])).
cnf(c_0_118,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_3,e_4)|product(X1,e_3,e_3)|~product(X1,e_2,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_116])).
cnf(c_0_119,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)|~greater(e_4,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_117])).
cnf(c_0_120,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_1,e_2)|product(e_2,e_3,e_4)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_116])).
cnf(c_0_121,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_118, c_0_116])).
cnf(c_0_122,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_18]), c_0_9])])).
cnf(c_0_123,plain,(product(e_6,e_1,e_1)|product(e_6,e_1,e_2)|product(e_6,e_1,e_3)|product(e_6,e_1,e_4)|product(e_6,e_1,e_5)|product(e_6,e_1,e_6)), inference(spm,[status(thm)],[c_0_56, c_0_29])).
cnf(c_0_124,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_3,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_26])).
cnf(c_0_125,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)|~greater(e_3,X1)|~next(e_1,X1)), inference(spm,[status(thm)],[c_0_3, c_0_122])).
cnf(c_0_126,negated_conjecture,(product(e_6,e_1,e_1)|product(e_6,e_1,e_2)|product(e_6,e_1,e_3)|product(e_6,e_1,e_4)|product(e_6,e_1,e_5)|product(X1,e_6,e_6)|~product(X1,e_1,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_123])).
cnf(c_0_127,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|~product(X1,e_1,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_124])).
cnf(c_0_128,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_20]), c_0_9])])).
cnf(c_0_129,plain,(equalish(X1,e_1)|product(e_6,e_1,e_1)|product(e_6,e_1,e_2)|product(e_6,e_1,e_3)|product(e_6,e_1,e_4)|product(e_6,e_1,e_5)|~product(e_6,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_123])).
cnf(c_0_130,negated_conjecture,(product(e_6,e_1,e_1)|product(e_6,e_1,e_2)|product(e_6,e_1,e_3)|product(e_6,e_1,e_4)|product(e_6,e_1,e_5)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_126, c_0_123])).
cnf(c_0_131,negated_conjecture,(product(e_2,e_3,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_12])).
cnf(c_0_132,negated_conjecture,(product(e_6,e_1,e_1)|product(e_6,e_1,e_2)|product(e_6,e_1,e_3)|product(e_6,e_1,e_4)|product(e_6,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_30])).
cnf(c_0_133,negated_conjecture,(equalish(X1,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|~product(X1,e_1,e_4)), inference(spm,[status(thm)],[c_0_7, c_0_101])).
cnf(c_0_134,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_2,e_3,e_4)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_131]), c_0_31])).
cnf(c_0_135,negated_conjecture,(product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|product(e_6,e_1,e_2)|product(e_6,e_1,e_3)|product(e_6,e_1,e_4)|product(e_6,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_132]), c_0_30])).
cnf(c_0_136,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_2,e_3,e_4)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_32])).
cnf(c_0_137,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|product(e_6,e_1,e_2)|product(e_6,e_1,e_3)|product(e_6,e_1,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_135]), c_0_33])).
cnf(c_0_138,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_5,e_3,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_2,e_3,e_4)|product(e_5,e_1,e_2)|product(e_4,e_5,e_6)), inference(spm,[status(thm)],[c_0_105, c_0_136])).
cnf(c_0_139,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|product(e_6,e_1,e_2)|product(e_6,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_137]), c_0_34])).
cnf(c_0_140,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_5,e_3,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_138]), c_0_35])).
cnf(c_0_141,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_6,e_3,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|product(e_6,e_1,e_2)), inference(spm,[status(thm)],[c_0_105, c_0_139])).
cnf(c_0_142,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_7, c_0_140])).
cnf(c_0_143,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_6,e_3,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_141]), c_0_36])).
cnf(c_0_144,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_37])).
cnf(c_0_145,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_5,e_6)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_144]), c_0_19])).
cnf(c_0_146,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_2,e_3,e_4)|product(e_4,e_5,e_6)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_145]), c_0_14])).
cnf(c_0_147,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_1,e_3,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_2,e_3,e_4)|product(e_4,e_5,e_6)|product(e_1,e_1,e_2)), inference(spm,[status(thm)],[c_0_105, c_0_146])).
cnf(c_0_148,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_1,e_3,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_147]), c_0_12])).
cnf(c_0_149,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_148]), c_0_38])).
cnf(c_0_150,plain,(product(X1,e_5,e_1)|product(X1,e_5,e_2)|product(X1,e_5,e_3)|product(X1,e_5,e_4)|product(X1,e_5,e_5)|product(X1,e_5,e_6)|~group_element(X1)), inference(spm,[status(thm)],[c_0_0, c_0_24])).
cnf(c_0_151,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_149]), c_0_39])).
cnf(c_0_152,plain,(product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_5,e_6)), inference(spm,[status(thm)],[c_0_150, c_0_2])).
cnf(c_0_153,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)|product(e_2,e_3,e_4)|product(e_4,e_1,e_2)|product(e_4,e_5,e_6)), inference(spm,[status(thm)],[c_0_105, c_0_151])).
cnf(c_0_154,plain,(equalish(X1,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|~product(X1,e_5,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_152])).
cnf(c_0_155,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_153]), c_0_40])).
cnf(c_0_156,negated_conjecture,(product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_40])).
cnf(c_0_157,negated_conjecture,(product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(X1,e_4,e_4)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_156])).
cnf(c_0_158,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_156])).
cnf(c_0_159,negated_conjecture,(product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_157, c_0_156])).
cnf(c_0_160,negated_conjecture,(product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_39])).
cnf(c_0_161,negated_conjecture,(product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(X1,e_3,e_3)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_160])).
cnf(c_0_162,negated_conjecture,(product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_161, c_0_160])).
cnf(c_0_163,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_160])).
cnf(c_0_164,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_162]), c_0_26])).
cnf(c_0_165,negated_conjecture,(product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_41])).
cnf(c_0_166,plain,(product(e_1,e_5,e_1)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_1,e_5,e_5)|product(e_1,e_5,e_6)), inference(spm,[status(thm)],[c_0_150, c_0_1])).
cnf(c_0_167,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_165]), c_0_26])).
cnf(c_0_168,negated_conjecture,(product(e_1,e_5,e_1)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_1,e_5,e_5)|product(X1,e_1,e_6)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_166])).
cnf(c_0_169,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_167])).
cnf(c_0_170,negated_conjecture,(product(e_1,e_5,e_1)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_1,e_5,e_5)|product(e_2,e_1,e_6)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)), inference(spm,[status(thm)],[c_0_168, c_0_167])).
cnf(c_0_171,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_1,e_5,e_5)|product(e_2,e_1,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170]), c_0_12])).
cnf(c_0_172,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_2,e_1,e_6)|product(X1,e_1,e_5)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_171])).
cnf(c_0_173,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)), inference(spm,[status(thm)],[c_0_172, c_0_167])).
cnf(c_0_174,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(X1,e_1,e_4)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_173])).
cnf(c_0_175,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(X1,e_4,e_5)|~product(X1,e_1,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_107])).
cnf(c_0_176,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)), inference(spm,[status(thm)],[c_0_174, c_0_167])).
cnf(c_0_177,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_3,e_4,e_5)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(spm,[status(thm)],[c_0_175, c_0_101])).
cnf(c_0_178,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_1,e_5,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(X1,e_1,e_3)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_176])).
cnf(c_0_179,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_3,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_177]), c_0_39])).
cnf(c_0_180,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(X1,e_2,e_1)|~product(X1,e_5,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_167])).
cnf(c_0_181,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_1,e_5,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(spm,[status(thm)],[c_0_178, c_0_167])).
cnf(c_0_182,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)|product(e_4,e_1,e_2)|product(e_3,e_4,e_5)), inference(spm,[status(thm)],[c_0_105, c_0_179])).
cnf(c_0_183,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_1,e_2,e_1)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(spm,[status(thm)],[c_0_180, c_0_181])).
cnf(c_0_184,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)|product(e_3,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_182]), c_0_40])).
cnf(c_0_185,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|product(X1,e_1,e_1)|~product(X1,e_2,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_183])).
cnf(c_0_186,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|product(X1,e_4,e_4)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_184])).
cnf(c_0_187,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|~product(e_1,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_183])).
cnf(c_0_188,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|product(e_1,e_1,e_1)), inference(spm,[status(thm)],[c_0_185, c_0_183])).
cnf(c_0_189,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_184])).
cnf(c_0_190,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_186, c_0_184])).
cnf(c_0_191,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|~product(e_2,e_1,X1)), inference(spm,[status(thm)],[c_0_42, c_0_124])).
cnf(c_0_192,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_12])).
cnf(c_0_193,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_39])).
cnf(c_0_194,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192]), c_0_36])).
cnf(c_0_195,plain,(product(X1,e_4,e_1)|product(X1,e_4,e_2)|product(X1,e_4,e_3)|product(X1,e_4,e_4)|product(X1,e_4,e_5)|product(X1,e_4,e_6)|~group_element(X1)), inference(spm,[status(thm)],[c_0_0, c_0_13])).
cnf(c_0_196,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(X1,e_3,e_3)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_193])).
cnf(c_0_197,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_194]), c_0_35])).
cnf(c_0_198,plain,(product(e_5,e_4,e_1)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_5,e_4,e_6)), inference(spm,[status(thm)],[c_0_195, c_0_24])).
cnf(c_0_199,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_196, c_0_193])).
cnf(c_0_200,plain,(product(e_5,e_5,e_1)|product(e_5,e_5,e_2)|product(e_5,e_5,e_3)|product(e_5,e_5,e_4)|product(e_5,e_5,e_5)|product(e_5,e_5,e_6)), inference(spm,[status(thm)],[c_0_150, c_0_24])).
cnf(c_0_201,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_197]), c_0_40])).
cnf(c_0_202,negated_conjecture,(product(e_5,e_4,e_1)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(X1,e_5,e_6)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_198])).
cnf(c_0_203,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_4,e_5)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_199]), c_0_26])).
cnf(c_0_204,negated_conjecture,(product(e_5,e_5,e_1)|product(e_5,e_5,e_2)|product(e_5,e_5,e_3)|product(e_5,e_5,e_4)|product(e_5,e_5,e_5)|product(X1,e_5,e_6)|~product(X1,e_5,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_200])).
cnf(c_0_205,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_201]), c_0_26])).
cnf(c_0_206,negated_conjecture,(product(e_5,e_4,e_1)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_3,e_5,e_6)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_202, c_0_203])).
cnf(c_0_207,negated_conjecture,(product(e_5,e_5,e_1)|product(e_5,e_5,e_2)|product(e_5,e_5,e_3)|product(e_5,e_5,e_4)|product(e_5,e_5,e_5)|product(e_2,e_5,e_6)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)), inference(spm,[status(thm)],[c_0_204, c_0_205])).
cnf(c_0_208,negated_conjecture,(product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_3,e_5,e_6)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(X1,e_5,e_1)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_206])).
cnf(c_0_209,negated_conjecture,(product(e_5,e_5,e_2)|product(e_5,e_5,e_3)|product(e_5,e_5,e_4)|product(e_5,e_5,e_5)|product(e_2,e_5,e_6)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(X1,e_5,e_1)|~product(X1,e_5,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_207])).
cnf(c_0_210,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_203])).
cnf(c_0_211,negated_conjecture,(product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_3,e_5,e_6)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)), inference(spm,[status(thm)],[c_0_208, c_0_203])).
cnf(c_0_212,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|~product(X1,e_5,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_205])).
cnf(c_0_213,negated_conjecture,(product(e_5,e_5,e_2)|product(e_5,e_5,e_3)|product(e_5,e_5,e_4)|product(e_5,e_5,e_5)|product(e_2,e_5,e_6)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_1)), inference(spm,[status(thm)],[c_0_209, c_0_205])).
cnf(c_0_214,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_32])).
cnf(c_0_215,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_5,e_5,e_2)|product(e_5,e_5,e_3)|product(e_5,e_5,e_4)|product(e_2,e_5,e_6)|product(e_2,e_5,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213]), c_0_35])).
cnf(c_0_216,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(X1,e_5,e_4)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_214])).
cnf(c_0_217,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_5,e_5,e_2)|product(e_5,e_5,e_3)|product(e_2,e_5,e_6)|product(e_2,e_5,e_1)|product(X1,e_5,e_4)|~product(X1,e_5,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_215])).
cnf(c_0_218,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(X1,e_3,e_5)|~product(X1,e_4,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_203])).
cnf(c_0_219,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_216, c_0_203])).
cnf(c_0_220,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_5,e_5,e_2)|product(e_5,e_5,e_3)|product(e_2,e_5,e_6)|product(e_2,e_5,e_1)), inference(spm,[status(thm)],[c_0_217, c_0_205])).
cnf(c_0_221,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_5,e_4,e_2)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_218, c_0_219])).
cnf(c_0_222,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_5,e_5,e_2)|product(e_2,e_5,e_6)|product(e_2,e_5,e_1)|product(X1,e_5,e_3)|~product(X1,e_5,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_220])).
cnf(c_0_223,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(X1,e_5,e_2)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_221])).
cnf(c_0_224,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_5,e_5,e_2)|product(e_2,e_5,e_6)|product(e_2,e_5,e_1)), inference(spm,[status(thm)],[c_0_222, c_0_205])).
cnf(c_0_225,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)), inference(spm,[status(thm)],[c_0_223, c_0_203])).
cnf(c_0_226,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_6)|product(e_2,e_5,e_1)|product(X1,e_5,e_2)|~product(X1,e_5,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_224])).
cnf(c_0_227,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(X1,e_5,e_5)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_225])).
cnf(c_0_228,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|~product(e_2,e_5,X1)), inference(spm,[status(thm)],[c_0_42, c_0_205])).
cnf(c_0_229,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_6)|product(e_2,e_5,e_1)), inference(spm,[status(thm)],[c_0_226, c_0_205])).
cnf(c_0_230,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|~product(e_5,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_225])).
cnf(c_0_231,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_5,e_5,e_5)), inference(spm,[status(thm)],[c_0_227, c_0_225])).
cnf(c_0_232,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_37])).
cnf(c_0_233,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231]), c_0_32])).
cnf(c_0_234,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)|product(X1,e_4,e_6)|~product(X1,e_5,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_155])).
cnf(c_0_235,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_232]), c_0_38])).
cnf(c_0_236,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|~product(X1,e_5,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_233])).
cnf(c_0_237,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)), inference(spm,[status(thm)],[c_0_234, c_0_235])).
cnf(c_0_238,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_155]), c_0_39])).
cnf(c_0_239,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(X1,e_4,e_4)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_237])).
cnf(c_0_240,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(X1,e_4,e_4)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_238])).
cnf(c_0_241,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_237])).
cnf(c_0_242,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_239, c_0_237])).
cnf(c_0_243,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_238])).
cnf(c_0_244,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_240, c_0_238])).
cnf(c_0_245,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_242]), c_0_39])).
cnf(c_0_246,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244]), c_0_39])).
cnf(c_0_247,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(X1,e_3,e_3)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_245])).
cnf(c_0_248,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(X1,e_3,e_3)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_246])).
cnf(c_0_249,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_247, c_0_245])).
cnf(c_0_250,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)), inference(spm,[status(thm)],[c_0_248, c_0_246])).
cnf(c_0_251,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_245])).
cnf(c_0_252,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_249]), c_0_26])).
cnf(c_0_253,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_250]), c_0_26])).
cnf(c_0_254,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_252]), c_0_41])).
cnf(c_0_255,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(X1,e_3,e_1)|~product(X1,e_5,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_253])).
cnf(c_0_256,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_254]), c_0_26])).
cnf(c_0_257,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_2,e_3,e_1)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)), inference(spm,[status(thm)],[c_0_255, c_0_256])).
cnf(c_0_258,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_2)|product(e_2,e_3,e_1)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(X1,e_3,e_4)|~product(X1,e_5,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_257])).
cnf(c_0_259,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(X1,e_2,e_3)|~product(X1,e_5,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_256])).
cnf(c_0_260,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_5,e_2)|product(e_2,e_3,e_1)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)), inference(spm,[status(thm)],[c_0_258, c_0_256])).
cnf(c_0_261,plain,(product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_3,e_5,e_6)), inference(spm,[status(thm)],[c_0_150, c_0_4])).
cnf(c_0_262,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_193])).
cnf(c_0_263,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_3,e_3,e_3)|product(e_2,e_3,e_1)), inference(spm,[status(thm)],[c_0_259, c_0_260])).
cnf(c_0_264,negated_conjecture,(product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(X1,e_3,e_6)|~product(X1,e_5,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_261])).
cnf(c_0_265,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_263]), c_0_41])).
cnf(c_0_266,negated_conjecture,(product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_6)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)), inference(spm,[status(thm)],[c_0_264, c_0_256])).
cnf(c_0_267,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_4,e_5)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_265]), c_0_26])).
cnf(c_0_268,negated_conjecture,(product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_6)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(X1,e_3,e_1)|~product(X1,e_5,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_266])).
cnf(c_0_269,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|~product(e_3,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_267])).
cnf(c_0_270,negated_conjecture,(product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_6)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_3,e_1)), inference(spm,[status(thm)],[c_0_268, c_0_256])).
cnf(c_0_271,negated_conjecture,(equalish(X1,e_1)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|~product(e_3,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_101])).
cnf(c_0_272,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_270]), c_0_31])).
cnf(c_0_273,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_263])).
cnf(c_0_274,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272]), c_0_27])).
cnf(c_0_275,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_1,e_2)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_274]), c_0_14])).
cnf(c_0_276,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|~product(X1,e_5,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_256])).
cnf(c_0_277,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_275]), c_0_26])).
cnf(c_0_278,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_5,e_2)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_277]), c_0_26])).
cnf(c_0_279,plain,(product(e_6,e_4,e_1)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|product(e_6,e_4,e_6)), inference(spm,[status(thm)],[c_0_195, c_0_29])).
cnf(c_0_280,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(spm,[status(thm)],[c_0_259, c_0_278])).
cnf(c_0_281,negated_conjecture,(product(e_6,e_4,e_1)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|product(X1,e_6,e_6)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_279])).
cnf(c_0_282,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(X1,e_3,e_3)|~product(X1,e_2,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_280])).
cnf(c_0_283,plain,(equalish(X1,e_4)|product(e_6,e_4,e_1)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|~product(e_6,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_279])).
cnf(c_0_284,negated_conjecture,(product(e_6,e_4,e_1)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_281, c_0_279])).
cnf(c_0_285,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_280])).
cnf(c_0_286,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_282, c_0_280])).
cnf(c_0_287,negated_conjecture,(product(e_6,e_4,e_1)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_284]), c_0_33])).
cnf(c_0_288,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_3,e_4)|~product(e_2,X1,e_2)), inference(spm,[status(thm)],[c_0_28, c_0_124])).
cnf(c_0_289,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286]), c_0_26])).
cnf(c_0_290,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|product(X1,e_6,e_1)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_287])).
cnf(c_0_291,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_289]), c_0_27])).
cnf(c_0_292,plain,(product(e_1,e_4,e_1)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_1,e_4,e_6)), inference(spm,[status(thm)],[c_0_195, c_0_1])).
cnf(c_0_293,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(spm,[status(thm)],[c_0_290, c_0_291])).
cnf(c_0_294,plain,(equalish(X1,e_1)|product(e_1,e_4,e_1)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_292])).
cnf(c_0_295,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_2,e_6,e_1)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_293]), c_0_34])).
cnf(c_0_296,negated_conjecture,(product(e_1,e_4,e_1)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_291]), c_0_43])).
cnf(c_0_297,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(X1,e_6,e_4)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_295])).
cnf(c_0_298,plain,(product(e_4,e_4,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_4,e_4,e_6)), inference(spm,[status(thm)],[c_0_195, c_0_13])).
cnf(c_0_299,negated_conjecture,(product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(X1,e_1,e_1)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_296])).
cnf(c_0_300,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)), inference(spm,[status(thm)],[c_0_297, c_0_291])).
cnf(c_0_301,plain,(equalish(X1,e_5)|product(e_5,e_4,e_1)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_198])).
cnf(c_0_302,plain,(equalish(X1,e_4)|product(e_4,e_4,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_298])).
cnf(c_0_303,negated_conjecture,(equalish(X1,e_4)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(e_1,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_296])).
cnf(c_0_304,negated_conjecture,(product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_1,e_1,e_1)), inference(spm,[status(thm)],[c_0_299, c_0_296])).
cnf(c_0_305,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_6,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(X1,e_6,e_3)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_300])).
cnf(c_0_306,negated_conjecture,(product(e_5,e_4,e_1)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_291]), c_0_44])).
cnf(c_0_307,negated_conjecture,(product(e_4,e_4,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_291]), c_0_45])).
cnf(c_0_308,negated_conjecture,(product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_304]), c_0_19])).
cnf(c_0_309,negated_conjecture,(equalish(X1,e_6)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_287])).
cnf(c_0_310,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(X1,e_2,e_6)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_291])).
cnf(c_0_311,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_6,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_305, c_0_291])).
cnf(c_0_312,negated_conjecture,(equalish(X1,e_5)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_306])).
cnf(c_0_313,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(X1,e_4,e_1)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_307])).
cnf(c_0_314,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_308]), c_0_14])).
cnf(c_0_315,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_307]), c_0_46])).
cnf(c_0_316,plain,(product(X1,e_6,e_1)|product(X1,e_6,e_2)|product(X1,e_6,e_3)|product(X1,e_6,e_4)|product(X1,e_6,e_5)|product(X1,e_6,e_6)|~group_element(X1)), inference(spm,[status(thm)],[c_0_0, c_0_29])).
cnf(c_0_317,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_6,e_2,e_6)|product(e_3,e_3,e_3)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_310, c_0_311])).
cnf(c_0_318,negated_conjecture,(product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_307]), c_0_47])).
cnf(c_0_319,negated_conjecture,(equalish(X1,e_4)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_307])).
cnf(c_0_320,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_1,e_4,e_1)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)), inference(spm,[status(thm)],[c_0_313, c_0_314])).
cnf(c_0_321,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_315]), c_0_34])).
cnf(c_0_322,plain,(product(e_6,e_6,e_1)|product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_316, c_0_29])).
cnf(c_0_323,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(X1,e_6,e_6)|~product(X1,e_2,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_317])).
cnf(c_0_324,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_318]), c_0_32])).
cnf(c_0_325,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_320]), c_0_19])).
cnf(c_0_326,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_6,e_4,e_1)|product(e_3,e_3,e_3)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)), inference(spm,[status(thm)],[c_0_313, c_0_321])).
cnf(c_0_327,plain,(equalish(X1,e_6)|product(e_6,e_6,e_1)|product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_322])).
cnf(c_0_328,plain,(product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(e_3,e_6,e_6)), inference(spm,[status(thm)],[c_0_316, c_0_4])).
cnf(c_0_329,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(e_6,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_317])).
cnf(c_0_330,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_323, c_0_317])).
cnf(c_0_331,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_5,e_4,e_1)|product(e_3,e_3,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)), inference(spm,[status(thm)],[c_0_313, c_0_324])).
cnf(c_0_332,negated_conjecture,(equalish(X1,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|~product(X1,e_4,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_325])).
cnf(c_0_333,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_326]), c_0_33])).
cnf(c_0_334,plain,(product(e_6,e_6,e_1)|product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_328]), c_0_48])).
cnf(c_0_335,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_330]), c_0_36])).
cnf(c_0_336,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_331]), c_0_31])).
cnf(c_0_337,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|product(e_6,e_4,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_333]), c_0_30])).
cnf(c_0_338,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(X1,e_6,e_1)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_334])).
cnf(c_0_339,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_335]), c_0_41])).
cnf(c_0_340,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_1,e_3,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_1,e_4,e_2)), inference(spm,[status(thm)],[c_0_218, c_0_325])).
cnf(c_0_341,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_5,e_4,e_2)), inference(spm,[status(thm)],[c_0_218, c_0_336])).
cnf(c_0_342,negated_conjecture,(equalish(X1,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_337])).
cnf(c_0_343,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)), inference(spm,[status(thm)],[c_0_338, c_0_328])).
cnf(c_0_344,negated_conjecture,(equalish(X1,e_6)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_2,e_6,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_293])).
cnf(c_0_345,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_339]), c_0_26])).
cnf(c_0_346,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_1,e_2,e_6)|product(e_3,e_3,e_3)|product(e_1,e_3,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)), inference(spm,[status(thm)],[c_0_310, c_0_340])).
cnf(c_0_347,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_341])).
cnf(c_0_348,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|product(e_5,e_3,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_341]), c_0_49])).
cnf(c_0_349,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(X1,e_6,e_5)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_343])).
cnf(c_0_350,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_2,e_6,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_345]), c_0_48])).
cnf(c_0_351,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_6,e_2,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)), inference(spm,[status(thm)],[c_0_310, c_0_337])).
cnf(c_0_352,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_1,e_2,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_346])).
cnf(c_0_353,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_348]), c_0_38])).
cnf(c_0_354,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)), inference(spm,[status(thm)],[c_0_349, c_0_328])).
cnf(c_0_355,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(X1,e_6,e_4)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_350])).
cnf(c_0_356,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_6,e_3,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_6,e_4,e_2)), inference(spm,[status(thm)],[c_0_218, c_0_333])).
cnf(c_0_357,negated_conjecture,(equalish(X1,e_6)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|~product(X1,e_2,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_351])).
cnf(c_0_358,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_1,e_2,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_353]), c_0_27])).
cnf(c_0_359,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(X1,e_6,e_4)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_354])).
cnf(c_0_360,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_355, c_0_291])).
cnf(c_0_361,negated_conjecture,(equalish(X1,e_6)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_6,e_3,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_356])).
cnf(c_0_362,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_358]), c_0_50])).
cnf(c_0_363,negated_conjecture,(product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(X1,e_5,e_1)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_306])).
cnf(c_0_364,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)), inference(spm,[status(thm)],[c_0_359, c_0_328])).
cnf(c_0_365,negated_conjecture,(product(e_6,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(X1,e_6,e_3)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_360])).
cnf(c_0_366,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_353])).
cnf(c_0_367,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_6,e_3,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_362]), c_0_50])).
cnf(c_0_368,negated_conjecture,(product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_363, c_0_203])).
cnf(c_0_369,negated_conjecture,(product(e_6,e_6,e_2)|product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(X1,e_6,e_3)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_364])).
cnf(c_0_370,negated_conjecture,(product(e_6,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_365, c_0_291])).
cnf(c_0_371,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_367]), c_0_37])).
cnf(c_0_372,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_368]), c_0_32])).
cnf(c_0_373,negated_conjecture,(product(e_6,e_6,e_2)|product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)), inference(spm,[status(thm)],[c_0_369, c_0_328])).
cnf(c_0_374,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_6,e_2,e_6)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_310, c_0_370])).
cnf(c_0_375,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_7, c_0_314])).
cnf(c_0_376,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_371]), c_0_39])).
cnf(c_0_377,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_7, c_0_372])).
cnf(c_0_378,negated_conjecture,(product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(X1,e_6,e_2)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_373])).
cnf(c_0_379,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(X1,e_6,e_6)|~product(X1,e_2,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_374])).
cnf(c_0_380,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_375, c_0_376]), c_0_51])).
cnf(c_0_381,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_376]), c_0_47])).
cnf(c_0_382,negated_conjecture,(product(e_3,e_6,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)), inference(spm,[status(thm)],[c_0_378, c_0_328])).
cnf(c_0_383,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(e_6,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_374])).
cnf(c_0_384,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_379, c_0_374])).
cnf(c_0_385,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_1,e_3,e_5)|product(e_1,e_4,e_2)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)), inference(spm,[status(thm)],[c_0_218, c_0_380])).
cnf(c_0_386,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_5,e_4,e_2)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)), inference(spm,[status(thm)],[c_0_218, c_0_381])).
cnf(c_0_387,negated_conjecture,(equalish(X1,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_382])).
cnf(c_0_388,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_384]), c_0_36])).
cnf(c_0_389,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_1,e_2,e_6)|product(e_3,e_3,e_3)|product(e_1,e_3,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)), inference(spm,[status(thm)],[c_0_310, c_0_385])).
cnf(c_0_390,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(X1,e_5,e_2)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_386])).
cnf(c_0_391,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_5,e_2,e_6)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_3,e_5,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)), inference(spm,[status(thm)],[c_0_310, c_0_386])).
cnf(c_0_392,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|~product(e_3,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_203])).
cnf(c_0_393,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_388]), c_0_52])).
cnf(c_0_394,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_1,e_2,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_389])).
cnf(c_0_395,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_3,e_5,e_2)), inference(spm,[status(thm)],[c_0_390, c_0_203])).
cnf(c_0_396,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_5,e_2,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(X1,e_5,e_5)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_391])).
cnf(c_0_397,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_393]), c_0_33])).
cnf(c_0_398,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_1,e_2,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_394, c_0_395]), c_0_27])).
cnf(c_0_399,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_5,e_2,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|~product(e_5,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_391])).
cnf(c_0_400,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_5,e_2,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_5,e_5,e_5)), inference(spm,[status(thm)],[c_0_396, c_0_391])).
cnf(c_0_401,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_397]), c_0_30])).
cnf(c_0_402,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|~product(X1,e_2,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_398])).
cnf(c_0_403,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_5,e_2,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_400]), c_0_32])).
cnf(c_0_404,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_6,e_2)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(X1,e_3,e_3)|~product(X1,e_6,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_401])).
cnf(c_0_405,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_403]), c_0_27])).
cnf(c_0_406,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(X1,e_2,e_1)|~product(X1,e_6,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_388])).
cnf(c_0_407,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_6,e_2)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_404, c_0_401])).
cnf(c_0_408,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(X1,e_4,e_3)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_405])).
cnf(c_0_409,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_406, c_0_407])).
cnf(c_0_410,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|~product(X1,e_4,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_405])).
cnf(c_0_411,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_2)), inference(spm,[status(thm)],[c_0_408, c_0_372])).
cnf(c_0_412,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)|product(e_3,e_1,e_2)|product(e_3,e_3,e_3)|product(X1,e_3,e_3)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_409])).
cnf(c_0_413,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_5,e_4,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_411]), c_0_31])).
cnf(c_0_414,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_2)), inference(spm,[status(thm)],[c_0_408, c_0_314])).
cnf(c_0_415,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)|product(e_3,e_1,e_2)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_412, c_0_409])).
cnf(c_0_416,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_413])).
cnf(c_0_417,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_1,e_4,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_414]), c_0_19])).
cnf(c_0_418,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_415]), c_0_26])).
cnf(c_0_419,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_4,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_417]), c_0_38])).
cnf(c_0_420,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_418]), c_0_41])).
cnf(c_0_421,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(X1,e_4,e_2)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_419])).
cnf(c_0_422,negated_conjecture,(equalish(X1,e_6)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(e_3,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_393])).
cnf(c_0_423,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_420]), c_0_26])).
cnf(c_0_424,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)), inference(spm,[status(thm)],[c_0_421, c_0_372])).
cnf(c_0_425,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_423]), c_0_46])).
cnf(c_0_426,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_419])).
cnf(c_0_427,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_5,e_3,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_5,e_4,e_2)), inference(spm,[status(thm)],[c_0_218, c_0_424])).
cnf(c_0_428,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)), inference(spm,[status(thm)],[c_0_421, c_0_314])).
cnf(c_0_429,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_418])).
cnf(c_0_430,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_425]), c_0_30])).
cnf(c_0_431,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_5,e_3,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_427]), c_0_31])).
cnf(c_0_432,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_1,e_3,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_1,e_4,e_2)), inference(spm,[status(thm)],[c_0_218, c_0_428])).
cnf(c_0_433,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_6,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_430]), c_0_34])).
cnf(c_0_434,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_431])).
cnf(c_0_435,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_1,e_3,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_432]), c_0_19])).
cnf(c_0_436,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(spm,[status(thm)],[c_0_406, c_0_433])).
cnf(c_0_437,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_434, c_0_435]), c_0_38])).
cnf(c_0_438,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_436]), c_0_14])).
cnf(c_0_439,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_2)|~product(e_3,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_437])).
cnf(c_0_440,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_2,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_438]), c_0_26])).
cnf(c_0_441,negated_conjecture,(equalish(X1,e_6)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(e_3,X1,e_2)), inference(spm,[status(thm)],[c_0_28, c_0_407])).
cnf(c_0_442,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_5,e_2)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_439, c_0_440]), c_0_44])).
cnf(c_0_443,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_441, c_0_442]), c_0_49])).
cnf(c_0_444,negated_conjecture,(product(e_3,e_1,e_2)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(X1,e_3,e_3)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_443])).
cnf(c_0_445,negated_conjecture,(product(e_3,e_1,e_2)|product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_444, c_0_443])).
cnf(c_0_446,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_445]), c_0_26])).
cnf(c_0_447,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_446]), c_0_41])).
cnf(c_0_448,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_4,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_447]), c_0_26])).
cnf(c_0_449,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_448])).
cnf(c_0_450,negated_conjecture,(product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_5,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_363, c_0_448])).
cnf(c_0_451,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_3,e_5,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_450]), c_0_32])).
cnf(c_0_452,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_3,e_5,e_1)|product(X1,e_5,e_4)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_451])).
cnf(c_0_453,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_452, c_0_448])).
cnf(c_0_454,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_5,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(X1,e_5,e_3)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_453])).
cnf(c_0_455,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_5,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_3)), inference(spm,[status(thm)],[c_0_454, c_0_448])).
cnf(c_0_456,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_3)|product(X1,e_5,e_2)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_455])).
cnf(c_0_457,plain,(product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(e_4,e_6,e_6)), inference(spm,[status(thm)],[c_0_316, c_0_13])).
cnf(c_0_458,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(e_3,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_440])).
cnf(c_0_459,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_3)|product(e_3,e_5,e_2)), inference(spm,[status(thm)],[c_0_456, c_0_448])).
cnf(c_0_460,plain,(product(e_6,e_6,e_1)|product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_457]), c_0_46])).
cnf(c_0_461,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|~product(X1,e_5,e_4)), inference(spm,[status(thm)],[c_0_7, c_0_235])).
cnf(c_0_462,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_3)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_458, c_0_459]), c_0_35])).
cnf(c_0_463,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(X1,e_6,e_1)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_460])).
cnf(c_0_464,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_446])).
cnf(c_0_465,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_5,e_3)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_462]), c_0_26])).
cnf(c_0_466,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)), inference(spm,[status(thm)],[c_0_463, c_0_457])).
cnf(c_0_467,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(X1,e_2,e_4)|~product(X1,e_5,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_235])).
cnf(c_0_468,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_465]), c_0_32])).
cnf(c_0_469,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(X1,e_6,e_5)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_466])).
cnf(c_0_470,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(e_3,e_2,X1)), inference(spm,[status(thm)],[c_0_42, c_0_440])).
cnf(c_0_471,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_3,e_2,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_467, c_0_468])).
cnf(c_0_472,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)), inference(spm,[status(thm)],[c_0_469, c_0_457])).
cnf(c_0_473,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|~product(e_2,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_235])).
cnf(c_0_474,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_470, c_0_471]), c_0_51])).
cnf(c_0_475,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(X1,e_6,e_4)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_472])).
cnf(c_0_476,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_448]), c_0_46])).
cnf(c_0_477,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_5,e_3)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_462]), c_0_27])).
cnf(c_0_478,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(X1,e_3,e_1)|~product(X1,e_6,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_382])).
cnf(c_0_479,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_474]), c_0_37])).
cnf(c_0_480,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)), inference(spm,[status(thm)],[c_0_475, c_0_457])).
cnf(c_0_481,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_476]), c_0_30])).
cnf(c_0_482,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_477]), c_0_14])).
cnf(c_0_483,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(e_2,e_3,e_1)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_6)), inference(spm,[status(thm)],[c_0_478, c_0_479])).
cnf(c_0_484,negated_conjecture,(product(e_6,e_6,e_2)|product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(X1,e_6,e_3)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_480])).
cnf(c_0_485,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_6,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_481]), c_0_34])).
cnf(c_0_486,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_5,e_3)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_482]), c_0_26])).
cnf(c_0_487,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_3,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_483]), c_0_33])).
cnf(c_0_488,negated_conjecture,(product(e_6,e_6,e_2)|product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)), inference(spm,[status(thm)],[c_0_484, c_0_457])).
cnf(c_0_489,negated_conjecture,(equalish(X1,e_6)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|~product(e_3,X1,e_2)), inference(spm,[status(thm)],[c_0_28, c_0_485])).
cnf(c_0_490,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_486]), c_0_32])).
cnf(c_0_491,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_2,e_3,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_6)|product(X1,e_3,e_4)|~product(X1,e_6,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_487])).
cnf(c_0_492,negated_conjecture,(product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(X1,e_6,e_2)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_488])).
cnf(c_0_493,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_490]), c_0_49])).
cnf(c_0_494,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(X1,e_6,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_479])).
cnf(c_0_495,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_3,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_2,e_3,e_1)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_6)), inference(spm,[status(thm)],[c_0_491, c_0_479])).
cnf(c_0_496,plain,(equalish(X1,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|~product(X1,e_5,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_261])).
cnf(c_0_497,negated_conjecture,(product(e_4,e_6,e_1)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)), inference(spm,[status(thm)],[c_0_492, c_0_457])).
cnf(c_0_498,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_493]), c_0_14])).
cnf(c_0_499,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(X1,e_2,e_3)|~product(X1,e_6,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_479])).
cnf(c_0_500,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)|product(e_3,e_6,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_495]), c_0_26])).
cnf(c_0_501,negated_conjecture,(product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_496, c_0_155]), c_0_39])).
cnf(c_0_502,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(X1,e_4,e_1)|~product(X1,e_6,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_497])).
cnf(c_0_503,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_498]), c_0_26])).
cnf(c_0_504,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_2,e_3)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_499, c_0_500])).
cnf(c_0_505,negated_conjecture,(product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(X1,e_4,e_4)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_501])).
cnf(c_0_506,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(e_2,e_4,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_502, c_0_503])).
cnf(c_0_507,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_504]), c_0_41])).
cnf(c_0_508,negated_conjecture,(equalish(X1,e_3)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_501])).
cnf(c_0_509,negated_conjecture,(product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_505, c_0_501])).
cnf(c_0_510,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(X1,e_4,e_5)|~product(X1,e_6,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_506])).
cnf(c_0_511,negated_conjecture,(equalish(X1,e_6)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_6)|~product(e_3,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_483])).
cnf(c_0_512,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_4,e_5)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_507]), c_0_26])).
cnf(c_0_513,negated_conjecture,(product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_508, c_0_509]), c_0_39])).
cnf(c_0_514,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|~product(X1,e_6,e_4)), inference(spm,[status(thm)],[c_0_7, c_0_503])).
cnf(c_0_515,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_510, c_0_503])).
cnf(c_0_516,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_6)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_511, c_0_512]), c_0_46])).
cnf(c_0_517,negated_conjecture,(equalish(X1,e_1)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_513])).
cnf(c_0_518,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_514, c_0_515]), c_0_40])).
cnf(c_0_519,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_6)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_516]), c_0_30])).
cnf(c_0_520,negated_conjecture,(product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_3,e_4)|product(e_3,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_263]), c_0_41])).
cnf(c_0_521,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(e_4,e_6,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_4,e_3)|~product(X1,e_6,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_518])).
cnf(c_0_522,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_6,e_2)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_519]), c_0_26])).
cnf(c_0_523,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_5,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_520]), c_0_26])).
cnf(c_0_524,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(X1,e_2,e_4)|~product(X1,e_6,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_503])).
cnf(c_0_525,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(e_4,e_6,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)), inference(spm,[status(thm)],[c_0_521, c_0_503])).
cnf(c_0_526,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_2,e_3)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_504])).
cnf(c_0_527,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_2,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(spm,[status(thm)],[c_0_499, c_0_522])).
cnf(c_0_528,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(X1,e_3,e_1)|~product(X1,e_5,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_523])).
cnf(c_0_529,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(e_4,e_2,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)), inference(spm,[status(thm)],[c_0_524, c_0_525])).
cnf(c_0_530,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_2,e_3)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_526, c_0_527]), c_0_14])).
cnf(c_0_531,negated_conjecture,(product(e_2,e_3,e_4)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)|product(e_3,e_5,e_5)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)), inference(spm,[status(thm)],[c_0_528, c_0_256])).
cnf(c_0_532,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)|product(X1,e_4,e_4)|~product(X1,e_2,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_529])).
cnf(c_0_533,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_530]), c_0_26])).
cnf(c_0_534,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)|product(e_3,e_5,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_531]), c_0_31])).
cnf(c_0_535,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_529])).
cnf(c_0_536,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_532, c_0_529])).
cnf(c_0_537,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(X1,e_3,e_3)|~product(X1,e_2,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_533])).
cnf(c_0_538,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_534]), c_0_27])).
cnf(c_0_539,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_535, c_0_536]), c_0_40])).
cnf(c_0_540,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_533])).
cnf(c_0_541,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_537, c_0_533])).
cnf(c_0_542,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_1,e_2)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_538]), c_0_14])).
cnf(c_0_543,negated_conjecture,(equalish(X1,e_6)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)|~product(e_2,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_539])).
cnf(c_0_544,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_5,e_2)|product(e_2,e_5,e_3)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_540, c_0_541]), c_0_26])).
cnf(c_0_545,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_5,e_2)|product(e_3,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_542]), c_0_26])).
cnf(c_0_546,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_543, c_0_544]), c_0_49])).
cnf(c_0_547,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_545]), c_0_26])).
cnf(c_0_548,negated_conjecture,(equalish(X1,e_6)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|~product(e_2,e_4,X1)), inference(spm,[status(thm)],[c_0_42, c_0_291])).
cnf(c_0_549,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_546]), c_0_27])).
cnf(c_0_550,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_3,e_2,e_3)|product(e_2,e_3,e_1)), inference(spm,[status(thm)],[c_0_259, c_0_547])).
cnf(c_0_551,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_4,e_5)|product(e_2,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_549]), c_0_50])).
cnf(c_0_552,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_3,e_1)|product(X1,e_3,e_3)|~product(X1,e_2,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_550])).
cnf(c_0_553,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)|product(e_2,e_4,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_551]), c_0_49])).
cnf(c_0_554,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_3,e_1)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_550])).
cnf(c_0_555,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_3,e_1)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_552, c_0_550])).
cnf(c_0_556,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|product(e_2,e_3,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_553]), c_0_48])).
cnf(c_0_557,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_5,e_2)|product(e_2,e_3,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_554, c_0_555]), c_0_26])).
cnf(c_0_558,plain,(product(X1,e_3,e_1)|product(X1,e_3,e_2)|product(X1,e_3,e_3)|product(X1,e_3,e_4)|product(X1,e_3,e_5)|product(X1,e_3,e_6)|~group_element(X1)), inference(spm,[status(thm)],[c_0_0, c_0_4])).
cnf(c_0_559,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_3,e_4)|product(e_2,e_3,e_1)|~product(e_2,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_556])).
cnf(c_0_560,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_3,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_557]), c_0_27])).
cnf(c_0_561,plain,(product(e_1,e_3,e_1)|product(e_1,e_3,e_2)|product(e_1,e_3,e_3)|product(e_1,e_3,e_4)|product(e_1,e_3,e_5)|product(e_1,e_3,e_6)), inference(spm,[status(thm)],[c_0_558, c_0_1])).
cnf(c_0_562,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_3,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_559, c_0_560]), c_0_39])).
cnf(c_0_563,negated_conjecture,(product(e_1,e_3,e_1)|product(e_1,e_3,e_2)|product(e_1,e_3,e_3)|product(e_1,e_3,e_4)|product(e_1,e_3,e_5)|product(X1,e_1,e_6)|~product(X1,e_3,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_561])).
cnf(c_0_564,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|~product(X1,e_3,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_562])).
cnf(c_0_565,negated_conjecture,(product(e_1,e_3,e_1)|product(e_1,e_3,e_2)|product(e_1,e_3,e_3)|product(e_1,e_3,e_4)|product(e_1,e_3,e_5)|product(e_2,e_1,e_6)|product(e_2,e_3,e_4)), inference(spm,[status(thm)],[c_0_563, c_0_562])).
cnf(c_0_566,negated_conjecture,(product(e_2,e_3,e_4)|product(e_1,e_3,e_2)|product(e_1,e_3,e_3)|product(e_1,e_3,e_4)|product(e_1,e_3,e_5)|product(e_2,e_1,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_564, c_0_565]), c_0_12])).
cnf(c_0_567,negated_conjecture,(product(e_2,e_3,e_4)|product(e_1,e_3,e_2)|product(e_1,e_3,e_3)|product(e_1,e_3,e_4)|product(e_2,e_1,e_6)|product(X1,e_1,e_5)|~product(X1,e_3,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_566])).
cnf(c_0_568,negated_conjecture,(product(e_2,e_3,e_4)|product(e_1,e_3,e_2)|product(e_1,e_3,e_3)|product(e_1,e_3,e_4)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)), inference(spm,[status(thm)],[c_0_567, c_0_562])).
cnf(c_0_569,negated_conjecture,(product(e_2,e_3,e_4)|product(e_1,e_3,e_2)|product(e_1,e_3,e_3)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(X1,e_1,e_4)|~product(X1,e_3,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_568])).
cnf(c_0_570,negated_conjecture,(product(e_2,e_3,e_4)|product(e_1,e_3,e_2)|product(e_1,e_3,e_3)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)), inference(spm,[status(thm)],[c_0_569, c_0_562])).
cnf(c_0_571,negated_conjecture,(product(e_2,e_3,e_4)|product(e_1,e_3,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(X1,e_1,e_3)|~product(X1,e_3,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_570])).
cnf(c_0_572,negated_conjecture,(product(e_2,e_3,e_4)|product(X1,e_2,e_1)|~product(X1,e_3,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_562])).
cnf(c_0_573,negated_conjecture,(product(e_2,e_3,e_4)|product(e_1,e_3,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(spm,[status(thm)],[c_0_571, c_0_562])).
cnf(c_0_574,negated_conjecture,(product(e_2,e_3,e_4)|product(e_1,e_2,e_1)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(spm,[status(thm)],[c_0_572, c_0_573])).
cnf(c_0_575,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|product(X1,e_1,e_1)|~product(X1,e_2,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_574])).
cnf(c_0_576,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_3,e_4)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|~product(e_1,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_574])).
cnf(c_0_577,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|product(e_1,e_1,e_1)), inference(spm,[status(thm)],[c_0_575, c_0_574])).
cnf(c_0_578,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_576, c_0_577]), c_0_12])).
cnf(c_0_579,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_578]), c_0_36])).
cnf(c_0_580,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_579]), c_0_35])).
cnf(c_0_581,plain,(product(e_4,e_3,e_1)|product(e_4,e_3,e_2)|product(e_4,e_3,e_3)|product(e_4,e_3,e_4)|product(e_4,e_3,e_5)|product(e_4,e_3,e_6)), inference(spm,[status(thm)],[c_0_558, c_0_13])).
cnf(c_0_582,negated_conjecture,(product(e_2,e_3,e_4)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_580]), c_0_40])).
cnf(c_0_583,negated_conjecture,(product(e_4,e_3,e_1)|product(e_4,e_3,e_2)|product(e_4,e_3,e_3)|product(e_4,e_3,e_4)|product(e_4,e_3,e_5)|product(X1,e_4,e_6)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_581])).
cnf(c_0_584,negated_conjecture,(product(e_2,e_3,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_582]), c_0_26])).
cnf(c_0_585,negated_conjecture,(product(e_4,e_3,e_1)|product(e_4,e_3,e_2)|product(e_4,e_3,e_3)|product(e_4,e_3,e_4)|product(e_4,e_3,e_5)|product(e_2,e_4,e_6)), inference(spm,[status(thm)],[c_0_583, c_0_584])).
cnf(c_0_586,negated_conjecture,(product(e_4,e_3,e_2)|product(e_4,e_3,e_3)|product(e_4,e_3,e_4)|product(e_4,e_3,e_5)|product(e_2,e_4,e_6)|product(X1,e_4,e_1)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_585])).
cnf(c_0_587,negated_conjecture,(product(e_4,e_3,e_2)|product(e_4,e_3,e_3)|product(e_4,e_3,e_4)|product(e_4,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)), inference(spm,[status(thm)],[c_0_586, c_0_584])).
cnf(c_0_588,negated_conjecture,(product(e_4,e_3,e_2)|product(e_4,e_3,e_3)|product(e_4,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(X1,e_4,e_5)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_587])).
cnf(c_0_589,plain,(product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_3,e_3,e_6)), inference(spm,[status(thm)],[c_0_558, c_0_4])).
cnf(c_0_590,negated_conjecture,(equalish(X1,e_2)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_7, c_0_584])).
cnf(c_0_591,negated_conjecture,(product(e_4,e_3,e_2)|product(e_4,e_3,e_3)|product(e_4,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_588, c_0_584])).
cnf(c_0_592,negated_conjecture,(product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(X1,e_3,e_6)|~product(X1,e_3,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_589])).
cnf(c_0_593,negated_conjecture,(product(e_4,e_3,e_2)|product(e_4,e_3,e_3)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_590, c_0_591]), c_0_40])).
cnf(c_0_594,plain,(equalish(X1,e_3)|product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|~product(X1,e_3,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_589])).
cnf(c_0_595,negated_conjecture,(product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_4,e_3,e_6)|product(e_4,e_3,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_592, c_0_593])).
cnf(c_0_596,negated_conjecture,(product(X1,e_2,e_4)|~product(X1,e_3,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_584])).
cnf(c_0_597,negated_conjecture,(product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_4,e_3,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_594, c_0_595]), c_0_39])).
cnf(c_0_598,negated_conjecture,(product(e_4,e_2,e_4)|product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_596, c_0_597])).
cnf(c_0_599,negated_conjecture,(product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_4,e_4)|~product(X1,e_2,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_598])).
cnf(c_0_600,negated_conjecture,(equalish(X1,e_2)|product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_598])).
cnf(c_0_601,negated_conjecture,(product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_599, c_0_598])).
cnf(c_0_602,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_1,e_2)|product(e_3,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_179]), c_0_40])).
cnf(c_0_603,negated_conjecture,(product(e_3,e_3,e_1)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_600, c_0_601]), c_0_40])).
cnf(c_0_604,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_2,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)), inference(spm,[status(thm)],[c_0_112, c_0_602])).
cnf(c_0_605,negated_conjecture,(product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_3,e_1)|~product(X1,e_3,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_603])).
cnf(c_0_606,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|~product(e_4,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_604])).
cnf(c_0_607,negated_conjecture,(equalish(X1,e_3)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(X1,e_3,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_603])).
cnf(c_0_608,negated_conjecture,(product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_3,e_1)|product(e_4,e_3,e_2)), inference(spm,[status(thm)],[c_0_605, c_0_593])).
cnf(c_0_609,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|~product(e_4,X1,e_2)), inference(spm,[status(thm)],[c_0_28, c_0_602])).
cnf(c_0_610,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|product(e_4,e_3,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_606, c_0_593]), c_0_26])).
cnf(c_0_611,negated_conjecture,(product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_607, c_0_608]), c_0_39])).
cnf(c_0_612,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_609, c_0_610]), c_0_41])).
cnf(c_0_613,negated_conjecture,(product(e_4,e_2,e_4)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_596, c_0_611])).
cnf(c_0_614,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_612]), c_0_26])).
cnf(c_0_615,negated_conjecture,(product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_4,e_4)|~product(X1,e_2,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_613])).
cnf(c_0_616,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_112, c_0_614])).
cnf(c_0_617,negated_conjecture,(equalish(X1,e_2)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_613])).
cnf(c_0_618,negated_conjecture,(product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_615, c_0_613])).
cnf(c_0_619,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(e_3,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_616])).
cnf(c_0_620,negated_conjecture,(product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_3,e_3,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_617, c_0_618]), c_0_40])).
cnf(c_0_621,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_619, c_0_620]), c_0_53])).
cnf(c_0_622,negated_conjecture,(equalish(X1,e_4)|product(e_4,e_3,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(X1,e_3,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_593])).
cnf(c_0_623,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_590, c_0_621]), c_0_26])).
cnf(c_0_624,negated_conjecture,(product(e_4,e_3,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_622, c_0_623]), c_0_53])).
cnf(c_0_625,negated_conjecture,(product(e_4,e_2,e_4)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_3,e_2)), inference(spm,[status(thm)],[c_0_596, c_0_624])).
cnf(c_0_626,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_3,e_2)|product(X1,e_4,e_4)|~product(X1,e_2,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_625])).
cnf(c_0_627,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_3,e_2)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_625])).
cnf(c_0_628,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_3,e_2)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_626, c_0_625])).
cnf(c_0_629,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_628]), c_0_40])).
cnf(c_0_630,negated_conjecture,(product(e_3,e_2,e_4)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)), inference(spm,[status(thm)],[c_0_596, c_0_629])).
cnf(c_0_631,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_630]), c_0_43])).
cnf(c_0_632,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_631]), c_0_26])).
cnf(c_0_633,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_112, c_0_632])).
cnf(c_0_634,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_3,e_3)|~product(X1,e_2,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_633])).
cnf(c_0_635,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_633])).
cnf(c_0_636,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_634, c_0_633])).
cnf(c_0_637,negated_conjecture,(equalish(X1,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)|product(e_3,e_4,e_5)|~product(X1,e_1,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_182])).
cnf(c_0_638,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_635, c_0_636]), c_0_26])).
cnf(c_0_639,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_3,e_4)|product(e_3,e_4,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_637, c_0_638]), c_0_45])).
cnf(c_0_640,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_590, c_0_639]), c_0_40])).
cnf(c_0_641,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_3,e_3)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_640])).
cnf(c_0_642,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(X1,e_1,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_638])).
cnf(c_0_643,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_641, c_0_640])).
cnf(c_0_644,negated_conjecture,(equalish(X1,e_3)|product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(e_3,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_620])).
cnf(c_0_645,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_4,e_5)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_642, c_0_643]), c_0_26])).
cnf(c_0_646,negated_conjecture,(product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_3,e_3,e_4)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_644, c_0_645]), c_0_39])).
cnf(c_0_647,negated_conjecture,(product(e_3,e_3,e_2)|product(e_3,e_3,e_3)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_590, c_0_646]), c_0_26])).
cnf(c_0_648,negated_conjecture,(product(e_4,e_3,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_622, c_0_647]), c_0_53])).
cnf(c_0_649,negated_conjecture,(product(e_4,e_2,e_4)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_2)), inference(spm,[status(thm)],[c_0_596, c_0_648])).
cnf(c_0_650,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_2)|product(X1,e_4,e_4)|~product(X1,e_2,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_649])).
cnf(c_0_651,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_2)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_649])).
cnf(c_0_652,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_2)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_650, c_0_649])).
cnf(c_0_653,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_651, c_0_652]), c_0_40])).
cnf(c_0_654,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)|product(X1,e_1,e_2)|~product(X1,e_1,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_128])).
cnf(c_0_655,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_3,e_2)|~product(X1,e_3,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_653])).
cnf(c_0_656,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_5)|product(e_4,e_5,e_6)), inference(spm,[status(thm)],[c_0_654, c_0_111])).
cnf(c_0_657,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(X1,e_3,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_653])).
cnf(c_0_658,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_3,e_2)), inference(spm,[status(thm)],[c_0_655, c_0_593])).
cnf(c_0_659,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_5,e_1,e_4)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_656]), c_0_31])).
cnf(c_0_660,negated_conjecture,(product(e_2,e_4,e_6)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_658]), c_0_39])).
cnf(c_0_661,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_5,e_4,e_5)|product(e_5,e_1,e_2)|product(e_5,e_1,e_3)|product(e_4,e_5,e_6)), inference(spm,[status(thm)],[c_0_175, c_0_659])).
cnf(c_0_662,negated_conjecture,(product(e_4,e_4,e_1)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_660]), c_0_45])).
cnf(c_0_663,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|product(X1,e_4,e_2)|~product(X1,e_1,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_602])).
cnf(c_0_664,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_5,e_4,e_5)|product(e_5,e_1,e_2)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_661]), c_0_35])).
cnf(c_0_665,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_662]), c_0_46])).
cnf(c_0_666,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|product(e_3,e_4,e_2)), inference(spm,[status(thm)],[c_0_663, c_0_101])).
cnf(c_0_667,negated_conjecture,(product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_5,e_4,e_5)|product(e_4,e_5,e_6)), inference(spm,[status(thm)],[c_0_112, c_0_664])).
cnf(c_0_668,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_6,e_5)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_665])).
cnf(c_0_669,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_3,e_4,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_666]), c_0_26])).
cnf(c_0_670,negated_conjecture,(product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)|product(X1,e_5,e_5)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_667])).
cnf(c_0_671,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_6,e_4,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|product(e_6,e_1,e_2)|product(e_6,e_1,e_3)), inference(spm,[status(thm)],[c_0_175, c_0_137])).
cnf(c_0_672,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_4,e_1)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_662])).
cnf(c_0_673,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)), inference(spm,[status(thm)],[c_0_668, c_0_660])).
cnf(c_0_674,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_5)|product(e_3,e_4,e_2)), inference(spm,[status(thm)],[c_0_112, c_0_669])).
cnf(c_0_675,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)|~product(e_5,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_667])).
cnf(c_0_676,negated_conjecture,(product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)|product(e_5,e_5,e_5)), inference(spm,[status(thm)],[c_0_670, c_0_667])).
cnf(c_0_677,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_6,e_4,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)|product(e_6,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_671]), c_0_36])).
cnf(c_0_678,negated_conjecture,(equalish(X1,e_4)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_662])).
cnf(c_0_679,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_6,e_4,e_1)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_5)), inference(spm,[status(thm)],[c_0_672, c_0_673])).
cnf(c_0_680,negated_conjecture,(product(e_5,e_4,e_1)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_3,e_5,e_6)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)), inference(spm,[status(thm)],[c_0_202, c_0_674])).
cnf(c_0_681,negated_conjecture,(product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_675, c_0_676]), c_0_31])).
cnf(c_0_682,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_674])).
cnf(c_0_683,negated_conjecture,(product(e_2,e_1,e_2)|product(e_6,e_2,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_6,e_4,e_5)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)), inference(spm,[status(thm)],[c_0_112, c_0_677])).
cnf(c_0_684,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_679]), c_0_33])).
cnf(c_0_685,negated_conjecture,(product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_3,e_5,e_6)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(X1,e_5,e_1)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_680])).
cnf(c_0_686,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)|~product(X1,e_2,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_681])).
cnf(c_0_687,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_6,e_2,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_682, c_0_683]), c_0_34])).
cnf(c_0_688,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_6,e_4,e_2)|product(e_2,e_6,e_5)|product(X1,e_6,e_3)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_684])).
cnf(c_0_689,negated_conjecture,(product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_5,e_4,e_5)|product(e_3,e_5,e_6)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_1)), inference(spm,[status(thm)],[c_0_685, c_0_674])).
cnf(c_0_690,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)|product(e_1,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_686, c_0_687]), c_0_37])).
cnf(c_0_691,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_2,e_6)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_660])).
cnf(c_0_692,negated_conjecture,(product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_6,e_4,e_2)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_688, c_0_660])).
cnf(c_0_693,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_5,e_4,e_4)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_682, c_0_689]), c_0_32])).
cnf(c_0_694,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_2,e_1,e_2)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)|product(e_1,e_1,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_690]), c_0_19])).
cnf(c_0_695,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_6,e_2,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_691, c_0_692])).
cnf(c_0_696,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(X1,e_5,e_4)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_693])).
cnf(c_0_697,negated_conjecture,(product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_4,e_5)|product(e_2,e_1,e_2)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_1,e_1,e_2)|product(e_1,e_1,e_3)), inference(spm,[status(thm)],[c_0_175, c_0_694])).
cnf(c_0_698,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(X1,e_6,e_6)|~product(X1,e_2,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_695])).
cnf(c_0_699,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(X1,e_3,e_5)|~product(X1,e_4,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_674])).
cnf(c_0_700,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_5,e_4,e_2)|product(e_5,e_4,e_3)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_696, c_0_674])).
cnf(c_0_701,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_4,e_5)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_1,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_697]), c_0_12])).
cnf(c_0_702,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|~product(e_6,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_695])).
cnf(c_0_703,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_698, c_0_695])).
cnf(c_0_704,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_6,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_290, c_0_660])).
cnf(c_0_705,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_5,e_3,e_5)|product(e_5,e_4,e_2)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_699, c_0_700])).
cnf(c_0_706,negated_conjecture,(product(e_2,e_1,e_2)|product(e_1,e_2,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_1,e_4,e_5)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)), inference(spm,[status(thm)],[c_0_112, c_0_701])).
cnf(c_0_707,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_4,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_703]), c_0_36])).
cnf(c_0_708,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_6,e_5)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_704])).
cnf(c_0_709,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_5,e_3,e_5)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(X1,e_5,e_2)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_705])).
cnf(c_0_710,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_1,e_2,e_3)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_682, c_0_706]), c_0_14])).
cnf(c_0_711,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(X1,e_4,e_5)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_707])).
cnf(c_0_712,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)), inference(spm,[status(thm)],[c_0_708, c_0_660])).
cnf(c_0_713,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_5,e_3,e_5)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)), inference(spm,[status(thm)],[c_0_709, c_0_674])).
cnf(c_0_714,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_1,e_4)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_686, c_0_710]), c_0_38])).
cnf(c_0_715,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_707])).
cnf(c_0_716,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_6,e_4,e_5)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)), inference(spm,[status(thm)],[c_0_711, c_0_712])).
cnf(c_0_717,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(X1,e_5,e_5)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_713])).
cnf(c_0_718,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_714]), c_0_39])).
cnf(c_0_719,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_715, c_0_716]), c_0_33])).
cnf(c_0_720,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|~product(e_5,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_713])).
cnf(c_0_721,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_5,e_5,e_5)), inference(spm,[status(thm)],[c_0_717, c_0_713])).
cnf(c_0_722,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_1,e_2)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_718]), c_0_40])).
cnf(c_0_723,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_6,e_4,e_2)|product(e_2,e_6,e_1)|product(X1,e_6,e_3)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_719])).
cnf(c_0_724,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_6)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_720, c_0_721]), c_0_32])).
cnf(c_0_725,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(X1,e_4,e_2)|~product(X1,e_1,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_722])).
cnf(c_0_726,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_6,e_4,e_2)|product(e_2,e_6,e_1)), inference(spm,[status(thm)],[c_0_723, c_0_660])).
cnf(c_0_727,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|~product(X1,e_5,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_724])).
cnf(c_0_728,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)), inference(spm,[status(thm)],[c_0_725, c_0_101])).
cnf(c_0_729,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_6,e_2,e_6)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_2,e_6,e_1)), inference(spm,[status(thm)],[c_0_691, c_0_726])).
cnf(c_0_730,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_727, c_0_728]), c_0_39])).
cnf(c_0_731,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_2,e_6,e_1)|product(X1,e_6,e_6)|~product(X1,e_2,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_729])).
cnf(c_0_732,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_730]), c_0_26])).
cnf(c_0_733,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_2,e_6,e_1)|~product(e_6,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_729])).
cnf(c_0_734,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_2,e_6,e_1)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_731, c_0_729])).
cnf(c_0_735,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_1)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)), inference(spm,[status(thm)],[c_0_112, c_0_732])).
cnf(c_0_736,negated_conjecture,(equalish(X1,e_6)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_7, c_0_712])).
cnf(c_0_737,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_4,e_4,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_2,e_6,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_733, c_0_734]), c_0_36])).
cnf(c_0_738,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_735])).
cnf(c_0_739,negated_conjecture,(product(e_1,e_5,e_1)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_1,e_5,e_5)|product(e_3,e_1,e_6)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)), inference(spm,[status(thm)],[c_0_168, c_0_735])).
cnf(c_0_740,negated_conjecture,(product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_736, c_0_737]), c_0_46])).
cnf(c_0_741,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_1,e_5,e_5)|product(e_3,e_1,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_738, c_0_739]), c_0_14])).
cnf(c_0_742,negated_conjecture,(product(e_6,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_2,e_6,e_3)|product(X1,e_6,e_3)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_740])).
cnf(c_0_743,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_3,e_1,e_6)|product(X1,e_1,e_5)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_741])).
cnf(c_0_744,negated_conjecture,(product(e_6,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_742, c_0_660])).
cnf(c_0_745,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)), inference(spm,[status(thm)],[c_0_743, c_0_735])).
cnf(c_0_746,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_6,e_2,e_6)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_691, c_0_744])).
cnf(c_0_747,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)|product(X1,e_1,e_4)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_745])).
cnf(c_0_748,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_2,e_6,e_3)|product(X1,e_6,e_6)|~product(X1,e_2,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_746])).
cnf(c_0_749,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(X1,e_3,e_1)|~product(X1,e_5,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_735])).
cnf(c_0_750,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)), inference(spm,[status(thm)],[c_0_747, c_0_735])).
cnf(c_0_751,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_2,e_6,e_3)|~product(e_6,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_746])).
cnf(c_0_752,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_2,e_6,e_3)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_748, c_0_746])).
cnf(c_0_753,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_3,e_1)|product(e_1,e_5,e_2)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)), inference(spm,[status(thm)],[c_0_749, c_0_750])).
cnf(c_0_754,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_4,e_4,e_3)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_751, c_0_752]), c_0_36])).
cnf(c_0_755,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_3,e_1)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)|product(X1,e_1,e_2)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_753])).
cnf(c_0_756,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_2,e_6,e_3)|product(X1,e_4,e_3)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_754])).
cnf(c_0_757,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_3,e_1)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)|product(e_3,e_1,e_2)), inference(spm,[status(thm)],[c_0_755, c_0_735])).
cnf(c_0_758,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_2,e_6,e_3)|~product(X1,e_4,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_754])).
cnf(c_0_759,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_2,e_6,e_3)|product(e_6,e_4,e_3)|product(e_6,e_4,e_2)), inference(spm,[status(thm)],[c_0_756, c_0_712])).
cnf(c_0_760,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)|product(e_3,e_1,e_2)|product(X1,e_1,e_1)|~product(X1,e_3,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_757])).
cnf(c_0_761,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_2,e_6,e_3)|product(e_6,e_4,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_758, c_0_759]), c_0_33])).
cnf(c_0_762,plain,(product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_5,e_6,e_6)), inference(spm,[status(thm)],[c_0_316, c_0_24])).
cnf(c_0_763,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)|product(e_3,e_1,e_2)|~product(e_1,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_757])).
cnf(c_0_764,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)|product(e_3,e_1,e_2)|product(e_1,e_1,e_1)), inference(spm,[status(thm)],[c_0_760, c_0_757])).
cnf(c_0_765,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_6,e_2,e_6)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_691, c_0_761])).
cnf(c_0_766,plain,(product(e_6,e_6,e_1)|product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_762]), c_0_49])).
cnf(c_0_767,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_763, c_0_764]), c_0_14])).
cnf(c_0_768,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_2,e_6,e_3)|product(X1,e_6,e_6)|~product(X1,e_2,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_765])).
cnf(c_0_769,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(X1,e_6,e_1)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_766])).
cnf(c_0_770,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)|product(e_3,e_1,e_2)|~greater(e_6,X1)|~next(e_3,X1)), inference(spm,[status(thm)],[c_0_3, c_0_767])).
cnf(c_0_771,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_2,e_6,e_3)|~product(e_6,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_765])).
cnf(c_0_772,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_2,e_6,e_3)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_768, c_0_765])).
cnf(c_0_773,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)), inference(spm,[status(thm)],[c_0_769, c_0_762])).
cnf(c_0_774,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|~product(e_3,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_674])).
cnf(c_0_775,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_5)|product(e_3,e_1,e_4)|product(e_3,e_1,e_2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_770, c_0_10]), c_0_11])])).
cnf(c_0_776,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_4,e_4,e_2)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_771, c_0_772]), c_0_36])).
cnf(c_0_777,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(X1,e_6,e_5)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_773])).
cnf(c_0_778,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|~product(X1,e_1,e_4)), inference(spm,[status(thm)],[c_0_7, c_0_714])).
cnf(c_0_779,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_4)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_774, c_0_775]), c_0_19])).
cnf(c_0_780,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(X1,e_4,e_2)|~product(X1,e_4,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_776])).
cnf(c_0_781,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)), inference(spm,[status(thm)],[c_0_777, c_0_762])).
cnf(c_0_782,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_1,e_3)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_778, c_0_779]), c_0_53])).
cnf(c_0_783,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)), inference(spm,[status(thm)],[c_0_780, c_0_712])).
cnf(c_0_784,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(X1,e_6,e_4)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_781])).
cnf(c_0_785,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_1,e_2)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_782]), c_0_40])).
cnf(c_0_786,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_6,e_4,e_2)|product(X1,e_6,e_3)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_783])).
cnf(c_0_787,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)), inference(spm,[status(thm)],[c_0_784, c_0_762])).
cnf(c_0_788,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(X1,e_4,e_2)|~product(X1,e_1,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_785])).
cnf(c_0_789,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_7, c_0_776])).
cnf(c_0_790,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_6,e_4,e_2)), inference(spm,[status(thm)],[c_0_786, c_0_660])).
cnf(c_0_791,negated_conjecture,(product(e_6,e_6,e_2)|product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(X1,e_6,e_3)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_787])).
cnf(c_0_792,plain,(equalish(X1,e_1)|product(e_1,e_5,e_1)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_1,e_5,e_5)|~product(X1,e_5,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_166])).
cnf(c_0_793,negated_conjecture,(product(e_2,e_1,e_2)|product(e_4,e_5,e_6)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)), inference(spm,[status(thm)],[c_0_788, c_0_779])).
cnf(c_0_794,negated_conjecture,(equalish(X1,e_4)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_497])).
cnf(c_0_795,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_789, c_0_790]), c_0_33])).
cnf(c_0_796,negated_conjecture,(product(e_6,e_6,e_2)|product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)), inference(spm,[status(thm)],[c_0_791, c_0_762])).
cnf(c_0_797,negated_conjecture,(product(e_1,e_5,e_1)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_1,e_5,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_792, c_0_793]), c_0_51])).
cnf(c_0_798,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_795]), c_0_45])).
cnf(c_0_799,negated_conjecture,(product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(X1,e_6,e_2)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_796])).
cnf(c_0_800,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_1,e_5,e_5)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_738, c_0_797]), c_0_14])).
cnf(c_0_801,negated_conjecture,(equalish(X1,e_4)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_798])).
cnf(c_0_802,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_795]), c_0_52])).
cnf(c_0_803,negated_conjecture,(product(e_5,e_6,e_1)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)), inference(spm,[status(thm)],[c_0_799, c_0_762])).
cnf(c_0_804,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_3,e_1,e_2)|product(X1,e_1,e_5)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_800])).
cnf(c_0_805,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_801, c_0_802]), c_0_53])).
cnf(c_0_806,negated_conjecture,(equalish(X1,e_5)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_803])).
cnf(c_0_807,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(X1,e_4,e_6)|~product(X1,e_5,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_793])).
cnf(c_0_808,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_1,e_5,e_4)|product(e_3,e_1,e_2)|product(e_3,e_1,e_5)), inference(spm,[status(thm)],[c_0_804, c_0_735])).
cnf(c_0_809,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(X1,e_4,e_4)|~product(X1,e_6,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_805])).
cnf(c_0_810,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_806, c_0_795]), c_0_44])).
cnf(c_0_811,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_1,e_4,e_6)|product(e_1,e_5,e_2)|product(e_1,e_5,e_3)|product(e_3,e_1,e_5)), inference(spm,[status(thm)],[c_0_807, c_0_808])).
cnf(c_0_812,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(X1,e_3,e_5)|~product(X1,e_6,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_802])).
cnf(c_0_813,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_809, c_0_805])).
cnf(c_0_814,negated_conjecture,(equalish(X1,e_5)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_810])).
cnf(c_0_815,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_1,e_4,e_6)|product(e_1,e_5,e_2)|product(e_3,e_1,e_5)|product(X1,e_1,e_3)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_811])).
cnf(c_0_816,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(X1,e_2,e_1)|~product(X1,e_6,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_795])).
cnf(c_0_817,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_4,e_3,e_5)|product(e_4,e_6,e_2)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_812, c_0_813])).
cnf(c_0_818,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_814, c_0_802]), c_0_54])).
cnf(c_0_819,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_1,e_4,e_6)|product(e_1,e_5,e_2)|product(e_3,e_1,e_5)|product(e_3,e_1,e_3)), inference(spm,[status(thm)],[c_0_815, c_0_735])).
cnf(c_0_820,negated_conjecture,(equalish(X1,e_6)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_805])).
cnf(c_0_821,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_4,e_3,e_5)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_816, c_0_817])).
cnf(c_0_822,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(X1,e_5,e_4)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_818])).
cnf(c_0_823,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_1,e_4,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_3)|product(X1,e_1,e_2)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_819])).
cnf(c_0_824,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_4,e_2,e_1)|product(e_4,e_3,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_820, c_0_821]), c_0_46])).
cnf(c_0_825,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_822, c_0_802])).
cnf(c_0_826,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_7, c_0_660])).
cnf(c_0_827,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_1,e_4,e_6)|product(e_3,e_1,e_5)|product(e_3,e_1,e_3)), inference(spm,[status(thm)],[c_0_823, c_0_735])).
cnf(c_0_828,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_4,e_3,e_5)|product(e_4,e_6,e_2)|product(e_4,e_2,e_1)), inference(spm,[status(thm)],[c_0_812, c_0_824])).
cnf(c_0_829,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_5,e_3,e_5)|product(e_5,e_6,e_2)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_812, c_0_825])).
cnf(c_0_830,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_5)|product(e_3,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_826, c_0_827]), c_0_12])).
cnf(c_0_831,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_4,e_3,e_5)), inference(spm,[status(thm)],[c_0_816, c_0_828])).
cnf(c_0_832,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_5,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_5,e_3,e_5)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_816, c_0_829])).
cnf(c_0_833,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_774, c_0_830]), c_0_19])).
cnf(c_0_834,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(X1,e_4,e_5)|~product(X1,e_3,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_831])).
cnf(c_0_835,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_5,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|product(X1,e_5,e_5)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_832])).
cnf(c_0_836,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_833]), c_0_26])).
cnf(c_0_837,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)), inference(spm,[status(thm)],[c_0_834, c_0_584])).
cnf(c_0_838,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_5,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|~product(e_5,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_832])).
cnf(c_0_839,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_5,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|product(e_5,e_5,e_5)), inference(spm,[status(thm)],[c_0_835, c_0_832])).
cnf(c_0_840,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_5,e_4)|product(e_3,e_5,e_2)), inference(spm,[status(thm)],[c_0_112, c_0_836])).
cnf(c_0_841,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|~product(X1,e_2,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_837])).
cnf(c_0_842,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_5,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_838, c_0_839]), c_0_32])).
cnf(c_0_843,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_840]), c_0_27])).
cnf(c_0_844,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_841, c_0_842]), c_0_31])).
cnf(c_0_845,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_843]), c_0_26])).
cnf(c_0_846,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_5,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_844]), c_0_30])).
cnf(c_0_847,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_5,e_2)), inference(spm,[status(thm)],[c_0_112, c_0_845])).
cnf(c_0_848,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_5,e_4)|product(X1,e_3,e_3)|~product(X1,e_6,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_846])).
cnf(c_0_849,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(e_3,X1,e_2)), inference(spm,[status(thm)],[c_0_28, c_0_847])).
cnf(c_0_850,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_5,e_4)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_848, c_0_846])).
cnf(c_0_851,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_5,e_4)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_849, c_0_850]), c_0_37])).
cnf(c_0_852,plain,(product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)|product(e_1,e_6,e_6)), inference(spm,[status(thm)],[c_0_316, c_0_1])).
cnf(c_0_853,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_3,e_1,e_2)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_5,e_4)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_851]), c_0_26])).
cnf(c_0_854,plain,(product(e_6,e_6,e_1)|product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_852]), c_0_50])).
cnf(c_0_855,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|~product(e_3,e_5,X1)), inference(spm,[status(thm)],[c_0_42, c_0_847])).
cnf(c_0_856,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_5,e_4)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_112, c_0_853])).
cnf(c_0_857,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)|product(X1,e_6,e_1)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_854])).
cnf(c_0_858,negated_conjecture,(equalish(X1,e_6)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_5,e_4)|product(e_3,e_3,e_3)|~product(e_3,X1,e_2)), inference(spm,[status(thm)],[c_0_28, c_0_850])).
cnf(c_0_859,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_855, c_0_856]), c_0_40])).
cnf(c_0_860,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)), inference(spm,[status(thm)],[c_0_857, c_0_852])).
cnf(c_0_861,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_5,e_4)|product(e_3,e_3,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_858, c_0_859]), c_0_46])).
cnf(c_0_862,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)|product(X1,e_6,e_5)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_860])).
cnf(c_0_863,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_5,e_4)|product(e_3,e_3,e_3)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_861]), c_0_26])).
cnf(c_0_864,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)), inference(spm,[status(thm)],[c_0_862, c_0_852])).
cnf(c_0_865,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_5,e_4)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_112, c_0_863])).
cnf(c_0_866,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)|product(X1,e_6,e_4)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_864])).
cnf(c_0_867,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_865]), c_0_27])).
cnf(c_0_868,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)), inference(spm,[status(thm)],[c_0_866, c_0_852])).
cnf(c_0_869,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_3,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_867]), c_0_26])).
cnf(c_0_870,negated_conjecture,(product(e_6,e_6,e_2)|product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)|product(X1,e_6,e_3)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_868])).
cnf(c_0_871,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_112, c_0_869])).
cnf(c_0_872,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_774, c_0_802]), c_0_33])).
cnf(c_0_873,negated_conjecture,(product(e_6,e_6,e_2)|product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)), inference(spm,[status(thm)],[c_0_870, c_0_852])).
cnf(c_0_874,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_871])).
cnf(c_0_875,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_872]), c_0_30])).
cnf(c_0_876,negated_conjecture,(product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)|product(X1,e_6,e_2)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_873])).
cnf(c_0_877,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_2)|product(e_3,e_6,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_874, c_0_875]), c_0_34])).
cnf(c_0_878,negated_conjecture,(product(e_1,e_6,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)), inference(spm,[status(thm)],[c_0_876, c_0_852])).
cnf(c_0_879,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_849, c_0_877]), c_0_37])).
cnf(c_0_880,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(X1,e_5,e_1)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_803])).
cnf(c_0_881,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_806, c_0_878]), c_0_38])).
cnf(c_0_882,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_879]), c_0_26])).
cnf(c_0_883,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_1,e_5,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)), inference(spm,[status(thm)],[c_0_880, c_0_881])).
cnf(c_0_884,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_112, c_0_882])).
cnf(c_0_885,negated_conjecture,(equalish(X1,e_6)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_5,e_4)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_846])).
cnf(c_0_886,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_1,e_5,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(X1,e_1,e_4)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_883])).
cnf(c_0_887,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|~product(e_3,X1,e_2)), inference(spm,[status(thm)],[c_0_28, c_0_884])).
cnf(c_0_888,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_6,e_2)|product(e_3,e_5,e_4)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_885, c_0_871]), c_0_48])).
cnf(c_0_889,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_1,e_5,e_1)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_5,e_1,e_4)), inference(spm,[status(thm)],[c_0_886, c_0_803])).
cnf(c_0_890,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_5,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_887, c_0_888]), c_0_33])).
cnf(c_0_891,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_1,e_5,e_1)|product(e_1,e_6,e_2)|product(e_5,e_1,e_4)|product(X1,e_1,e_3)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_889])).
cnf(c_0_892,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_1,e_2)|product(e_3,e_5,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_890]), c_0_26])).
cnf(c_0_893,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_1,e_5,e_1)|product(e_1,e_6,e_2)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)), inference(spm,[status(thm)],[c_0_891, c_0_803])).
cnf(c_0_894,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_112, c_0_892])).
cnf(c_0_895,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_1,e_5,e_1)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(X1,e_1,e_2)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_893])).
cnf(c_0_896,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_894]), c_0_27])).
cnf(c_0_897,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_1,e_5,e_1)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)), inference(spm,[status(thm)],[c_0_895, c_0_803])).
cnf(c_0_898,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_896]), c_0_26])).
cnf(c_0_899,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|product(X1,e_1,e_1)|~product(X1,e_5,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_897])).
cnf(c_0_900,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_112, c_0_898])).
cnf(c_0_901,negated_conjecture,(equalish(X1,e_5)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|~product(e_1,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_897])).
cnf(c_0_902,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|product(e_1,e_1,e_1)), inference(spm,[status(thm)],[c_0_899, c_0_897])).
cnf(c_0_903,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(X1,e_3,e_3)|~product(X1,e_2,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_900])).
cnf(c_0_904,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_901, c_0_902]), c_0_38])).
cnf(c_0_905,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_900])).
cnf(c_0_906,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_903, c_0_900])).
cnf(c_0_907,negated_conjecture,(equalish(X1,e_5)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_904])).
cnf(c_0_908,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_905, c_0_906]), c_0_26])).
cnf(c_0_909,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_907, c_0_908]), c_0_44])).
cnf(c_0_910,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(X1,e_5,e_4)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_909])).
cnf(c_0_911,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)), inference(spm,[status(thm)],[c_0_910, c_0_908])).
cnf(c_0_912,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(X1,e_5,e_3)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_911])).
cnf(c_0_913,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(X1,e_2,e_5)|~product(X1,e_6,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_908])).
cnf(c_0_914,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(spm,[status(thm)],[c_0_912, c_0_908])).
cnf(c_0_915,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_6,e_4,e_4)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_682, c_0_704]), c_0_34])).
cnf(c_0_916,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_5,e_2,e_5)|product(e_5,e_1,e_4)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(spm,[status(thm)],[c_0_913, c_0_914])).
cnf(c_0_917,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(X1,e_6,e_4)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_915])).
cnf(c_0_918,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_5,e_2,e_5)|product(e_5,e_1,e_3)|product(e_5,e_1,e_2)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_916]), c_0_32])).
cnf(c_0_919,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_6,e_4,e_2)|product(e_6,e_4,e_3)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_4)), inference(spm,[status(thm)],[c_0_917, c_0_660])).
cnf(c_0_920,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_5,e_2,e_5)|product(e_5,e_1,e_2)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_918]), c_0_35])).
cnf(c_0_921,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_6,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_4)|product(X1,e_6,e_3)|~product(X1,e_4,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_919])).
cnf(c_0_922,negated_conjecture,(product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_5,e_2,e_5)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(spm,[status(thm)],[c_0_112, c_0_920])).
cnf(c_0_923,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_6,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_921, c_0_660])).
cnf(c_0_924,negated_conjecture,(product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(X1,e_5,e_5)|~product(X1,e_2,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_922])).
cnf(c_0_925,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_6,e_2,e_6)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(spm,[status(thm)],[c_0_691, c_0_923])).
cnf(c_0_926,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|~product(e_5,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_922])).
cnf(c_0_927,negated_conjecture,(product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_5,e_5,e_5)), inference(spm,[status(thm)],[c_0_924, c_0_922])).
cnf(c_0_928,plain,(product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_2,e_6,e_6)), inference(spm,[status(thm)],[c_0_316, c_0_2])).
cnf(c_0_929,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(X1,e_6,e_6)|~product(X1,e_2,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_925])).
cnf(c_0_930,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|~product(X1,e_2,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_604])).
cnf(c_0_931,negated_conjecture,(product(e_2,e_1,e_2)|product(e_5,e_2,e_3)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_926, c_0_927]), c_0_35])).
cnf(c_0_932,plain,(product(e_6,e_6,e_1)|product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_928]), c_0_55])).
cnf(c_0_933,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(e_6,X1,e_6)), inference(spm,[status(thm)],[c_0_28, c_0_925])).
cnf(c_0_934,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_6,e_6,e_6)), inference(spm,[status(thm)],[c_0_929, c_0_925])).
cnf(c_0_935,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_3,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_930, c_0_931]), c_0_31])).
cnf(c_0_936,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(X1,e_6,e_1)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_932])).
cnf(c_0_937,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_933, c_0_934]), c_0_36])).
cnf(c_0_938,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_3,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_935]), c_0_26])).
cnf(c_0_939,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_6,e_6,e_5)|product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)), inference(spm,[status(thm)],[c_0_936, c_0_928])).
cnf(c_0_940,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|~product(e_3,e_4,X1)), inference(spm,[status(thm)],[c_0_42, c_0_937])).
cnf(c_0_941,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_3,e_4,e_5)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(spm,[status(thm)],[c_0_112, c_0_938])).
cnf(c_0_942,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(X1,e_6,e_5)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_939])).
cnf(c_0_943,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_3,e_2,e_3)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_940, c_0_941]), c_0_35])).
cnf(c_0_944,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_6,e_6,e_4)|product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)), inference(spm,[status(thm)],[c_0_942, c_0_928])).
cnf(c_0_945,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(X1,e_3,e_3)|~product(X1,e_2,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_943])).
cnf(c_0_946,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(X1,e_6,e_4)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_944])).
cnf(c_0_947,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|~product(e_3,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_943])).
cnf(c_0_948,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_3,e_3,e_3)), inference(spm,[status(thm)],[c_0_945, c_0_943])).
cnf(c_0_949,negated_conjecture,(product(e_6,e_6,e_2)|product(e_6,e_6,e_3)|product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)), inference(spm,[status(thm)],[c_0_946, c_0_928])).
cnf(c_0_950,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|~product(e_2,e_6,X1)), inference(spm,[status(thm)],[c_0_42, c_0_908])).
cnf(c_0_951,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_1,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_4)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_947, c_0_948]), c_0_26])).
cnf(c_0_952,negated_conjecture,(product(e_6,e_6,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(X1,e_6,e_3)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_949])).
cnf(c_0_953,negated_conjecture,(equalish(X1,e_3)|~product(e_2,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_584])).
cnf(c_0_954,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_950, c_0_951]), c_0_38])).
cnf(c_0_955,negated_conjecture,(product(e_6,e_6,e_2)|product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)), inference(spm,[status(thm)],[c_0_952, c_0_928])).
cnf(c_0_956,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_1,e_2)|~product(e_2,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_104])).
cnf(c_0_957,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_6,e_3)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_954]), c_0_34])).
cnf(c_0_958,negated_conjecture,(product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(X1,e_6,e_2)|~product(X1,e_6,e_6)), inference(spm,[status(thm)],[c_0_21, c_0_955])).
cnf(c_0_959,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_956, c_0_957]), c_0_30])).
cnf(c_0_960,negated_conjecture,(product(e_2,e_6,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)), inference(spm,[status(thm)],[c_0_958, c_0_928])).
cnf(c_0_961,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(X1,e_2,e_2)|~product(X1,e_1,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_959])).
cnf(c_0_962,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_4,e_6,e_5)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_960]), c_0_45])).
cnf(c_0_963,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_2,e_2)), inference(spm,[status(thm)],[c_0_961, c_0_959])).
cnf(c_0_964,negated_conjecture,(equalish(X1,e_4)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_962])).
cnf(c_0_965,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_6,e_5)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_960]), c_0_52])).
cnf(c_0_966,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_5,e_3)|product(e_2,e_2,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_963]), c_0_32])).
cnf(c_0_967,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_4,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_964, c_0_965]), c_0_53])).
cnf(c_0_968,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_956, c_0_966]), c_0_27])).
cnf(c_0_969,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(X1,e_4,e_4)|~product(X1,e_6,e_4)), inference(spm,[status(thm)],[c_0_21, c_0_967])).
cnf(c_0_970,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(X1,e_2,e_2)|~product(X1,e_1,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_968])).
cnf(c_0_971,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(X1,e_3,e_5)|~product(X1,e_6,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_965])).
cnf(c_0_972,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_969, c_0_967])).
cnf(c_0_973,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_806, c_0_960]), c_0_44])).
cnf(c_0_974,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)), inference(spm,[status(thm)],[c_0_970, c_0_968])).
cnf(c_0_975,negated_conjecture,(product(e_1,e_4,e_1)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(X1,e_1,e_6)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_292])).
cnf(c_0_976,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(X1,e_2,e_1)|~product(X1,e_6,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_960])).
cnf(c_0_977,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_4,e_3,e_5)|product(e_4,e_6,e_2)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_971, c_0_972])).
cnf(c_0_978,negated_conjecture,(equalish(X1,e_5)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_973])).
cnf(c_0_979,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_960])).
cnf(c_0_980,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_974])).
cnf(c_0_981,negated_conjecture,(product(e_1,e_4,e_1)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_2,e_1,e_6)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)), inference(spm,[status(thm)],[c_0_975, c_0_974])).
cnf(c_0_982,negated_conjecture,(equalish(X1,e_6)|product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|~product(e_4,X1,e_4)), inference(spm,[status(thm)],[c_0_28, c_0_967])).
cnf(c_0_983,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_4,e_3,e_5)|product(e_4,e_4,e_4)), inference(spm,[status(thm)],[c_0_976, c_0_977])).
cnf(c_0_984,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_978, c_0_965]), c_0_54])).
cnf(c_0_985,negated_conjecture,(equalish(X1,e_3)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_965])).
cnf(c_0_986,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_1,e_6,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_979, c_0_878]), c_0_12])).
cnf(c_0_987,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_2,e_1,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_980, c_0_981]), c_0_12])).
cnf(c_0_988,negated_conjecture,(product(e_4,e_6,e_2)|product(e_4,e_6,e_3)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_4,e_2,e_1)|product(e_4,e_3,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_982, c_0_983]), c_0_46])).
cnf(c_0_989,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(X1,e_5,e_4)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_984])).
cnf(c_0_990,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_985, c_0_986]), c_0_14])).
cnf(c_0_991,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_2,e_1,e_6)|product(X1,e_1,e_5)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_987])).
cnf(c_0_992,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_4,e_3,e_5)|product(e_4,e_6,e_2)|product(e_4,e_2,e_1)), inference(spm,[status(thm)],[c_0_971, c_0_988])).
cnf(c_0_993,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_989, c_0_965])).
cnf(c_0_994,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(X1,e_1,e_4)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_990])).
cnf(c_0_995,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)), inference(spm,[status(thm)],[c_0_991, c_0_974])).
cnf(c_0_996,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_4,e_3,e_5)), inference(spm,[status(thm)],[c_0_976, c_0_992])).
cnf(c_0_997,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_5,e_3,e_5)|product(e_5,e_6,e_2)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_971, c_0_993])).
cnf(c_0_998,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_2,e_1,e_4)), inference(spm,[status(thm)],[c_0_994, c_0_960])).
cnf(c_0_999,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(X1,e_1,e_4)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_995])).
cnf(c_0_1000,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_996])).
cnf(c_0_1001,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_5,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_5,e_3,e_5)|product(e_3,e_5,e_4)), inference(spm,[status(thm)],[c_0_976, c_0_997])).
cnf(c_0_1002,negated_conjecture,(product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_3,e_5)|product(e_1,e_6,e_2)|product(e_2,e_1,e_4)), inference(spm,[status(thm)],[c_0_971, c_0_998])).
cnf(c_0_1003,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)), inference(spm,[status(thm)],[c_0_999, c_0_974])).
cnf(c_0_1004,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_5,e_2,e_1)|product(e_3,e_5,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1000, c_0_1001]), c_0_31])).
cnf(c_0_1005,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_1,e_3,e_5)|product(e_2,e_1,e_4)), inference(spm,[status(thm)],[c_0_976, c_0_1002])).
cnf(c_0_1006,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_1,e_4,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(X1,e_1,e_3)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_1003])).
cnf(c_0_1007,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|~product(X1,e_2,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_1004])).
cnf(c_0_1008,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_1,e_2,e_1)|product(e_2,e_1,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1000, c_0_1005]), c_0_19])).
cnf(c_0_1009,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(X1,e_2,e_1)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_974])).
cnf(c_0_1010,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_1,e_4,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(spm,[status(thm)],[c_0_1006, c_0_974])).
cnf(c_0_1011,negated_conjecture,(equalish(X1,e_6)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|~product(e_2,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_960])).
cnf(c_0_1012,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_4,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1007, c_0_1008]), c_0_38])).
cnf(c_0_1013,negated_conjecture,(equalish(X1,e_5)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_5,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|~product(X1,e_3,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_1001])).
cnf(c_0_1014,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_1,e_2,e_1)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(spm,[status(thm)],[c_0_1009, c_0_1010])).
cnf(c_0_1015,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1011, c_0_974]), c_0_46])).
cnf(c_0_1016,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|~product(X1,e_2,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_1012])).
cnf(c_0_1017,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_5,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|product(e_1,e_2,e_1)|product(e_2,e_1,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1013, c_0_1005]), c_0_38])).
cnf(c_0_1018,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(X1,e_1,e_5)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_986])).
cnf(c_0_1019,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|product(X1,e_1,e_1)|~product(X1,e_2,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_1014])).
cnf(c_0_1020,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_1015])).
cnf(c_0_1021,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_2,e_5,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)), inference(spm,[status(thm)],[c_0_880, c_0_1015])).
cnf(c_0_1022,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|product(e_5,e_2,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1016, c_0_1017]), c_0_19])).
cnf(c_0_1023,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_1,e_6,e_4)|product(e_2,e_1,e_5)), inference(spm,[status(thm)],[c_0_1018, c_0_960])).
cnf(c_0_1024,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|~product(e_1,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_1014])).
cnf(c_0_1025,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|product(e_1,e_1,e_1)), inference(spm,[status(thm)],[c_0_1019, c_0_1014])).
cnf(c_0_1026,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_2,e_5,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1020, c_0_1021]), c_0_35])).
cnf(c_0_1027,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_6,e_4)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1016, c_0_1022]), c_0_31])).
cnf(c_0_1028,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_2,e_1,e_5)|product(X1,e_1,e_4)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_1023])).
cnf(c_0_1029,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1024, c_0_1025]), c_0_12])).
cnf(c_0_1030,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_2,e_5,e_1)|product(X1,e_5,e_4)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_1026])).
cnf(c_0_1031,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|product(X1,e_3,e_4)|~product(X1,e_6,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_1027])).
cnf(c_0_1032,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_1,e_6,e_3)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)), inference(spm,[status(thm)],[c_0_1028, c_0_960])).
cnf(c_0_1033,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|~greater(e_6,X1)|~next(e_2,X1)), inference(spm,[status(thm)],[c_0_3, c_0_1029])).
cnf(c_0_1034,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)), inference(spm,[status(thm)],[c_0_1030, c_0_1015])).
cnf(c_0_1035,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|product(e_1,e_3,e_4)|product(e_1,e_6,e_2)|product(e_2,e_1,e_5)), inference(spm,[status(thm)],[c_0_1031, c_0_1032])).
cnf(c_0_1036,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1033, c_0_5]), c_0_6])])).
cnf(c_0_1037,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_5,e_6,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(X1,e_5,e_3)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_1034])).
cnf(c_0_1038,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|product(e_1,e_3,e_4)|product(e_2,e_1,e_5)), inference(spm,[status(thm)],[c_0_976, c_0_1035])).
cnf(c_0_1039,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|~greater(e_5,X1)|~next(e_2,X1)), inference(spm,[status(thm)],[c_0_3, c_0_1036])).
cnf(c_0_1040,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_5,e_6,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(spm,[status(thm)],[c_0_1037, c_0_1015])).
cnf(c_0_1041,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_2,e_1)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_590, c_0_1038]), c_0_12])).
cnf(c_0_1042,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1039, c_0_22]), c_0_6])])).
cnf(c_0_1043,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(X1,e_5,e_2)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_1040])).
cnf(c_0_1044,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(X1,e_1,e_1)|~product(X1,e_2,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_1041])).
cnf(c_0_1045,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1042]), c_0_14])).
cnf(c_0_1046,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(spm,[status(thm)],[c_0_1043, c_0_1015])).
cnf(c_0_1047,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|~product(e_1,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_1041])).
cnf(c_0_1048,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)|product(e_1,e_1,e_1)), inference(spm,[status(thm)],[c_0_1044, c_0_1041])).
cnf(c_0_1049,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|~product(e_2,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_1045])).
cnf(c_0_1050,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1046]), c_0_34])).
cnf(c_0_1051,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|~product(X1,e_6,e_3)), inference(spm,[status(thm)],[c_0_7, c_0_1032])).
cnf(c_0_1052,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_6,e_2)|product(e_3,e_6,e_3)|product(e_3,e_5,e_4)|product(e_2,e_1,e_4)|product(e_2,e_1,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1047, c_0_1048]), c_0_12])).
cnf(c_0_1053,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_6,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1049, c_0_1050]), c_0_30])).
cnf(c_0_1054,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_1,e_6,e_2)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_3,e_6,e_2)|product(e_3,e_5,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1051, c_0_1052]), c_0_41])).
cnf(c_0_1055,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)|product(X1,e_2,e_2)|~product(X1,e_6,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_1053])).
cnf(c_0_1056,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_3,e_6,e_2)|product(e_3,e_5,e_4)|product(X1,e_1,e_2)|~product(X1,e_6,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_1054])).
cnf(c_0_1057,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|~product(e_2,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_974])).
cnf(c_0_1058,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(spm,[status(thm)],[c_0_1055, c_0_1053])).
cnf(c_0_1059,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_3,e_6,e_2)|product(e_3,e_5,e_4)|product(e_2,e_1,e_2)), inference(spm,[status(thm)],[c_0_1056, c_0_960])).
cnf(c_0_1060,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1057, c_0_1058]), c_0_31])).
cnf(c_0_1061,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_2,e_1)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_3,e_5,e_4)|product(e_2,e_1,e_2)), inference(spm,[status(thm)],[c_0_976, c_0_1059])).
cnf(c_0_1062,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1060]), c_0_32])).
cnf(c_0_1063,negated_conjecture,(product(e_3,e_1,e_2)|product(e_3,e_1,e_3)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_2,e_1)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_1061]), c_0_27])).
cnf(c_0_1064,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1049, c_0_1062]), c_0_27])).
cnf(c_0_1065,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_4,e_5)|product(e_3,e_4,e_2)|product(X1,e_3,e_2)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_669])).
cnf(c_0_1066,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_2,e_1)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_1063]), c_0_26])).
cnf(c_0_1067,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)|product(X1,e_2,e_2)|~product(X1,e_5,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_1064])).
cnf(c_0_1068,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_4,e_5)|product(e_3,e_4,e_2)|product(e_2,e_3,e_2)), inference(spm,[status(thm)],[c_0_1065, c_0_104])).
cnf(c_0_1069,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_2,e_1)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(X1,e_3,e_2)|~product(X1,e_1,e_3)), inference(spm,[status(thm)],[c_0_21, c_0_1066])).
cnf(c_0_1070,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_2,e_2)), inference(spm,[status(thm)],[c_0_1067, c_0_1064])).
cnf(c_0_1071,negated_conjecture,(equalish(X1,e_3)|product(e_2,e_1,e_2)|product(e_3,e_4,e_2)|product(e_2,e_3,e_2)|~product(X1,e_4,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_1068])).
cnf(c_0_1072,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_3,e_2,e_1)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(spm,[status(thm)],[c_0_1069, c_0_104])).
cnf(c_0_1073,negated_conjecture,(product(e_2,e_2,e_2)|product(X1,e_2,e_5)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_1070])).
cnf(c_0_1074,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_4,e_2)|product(e_2,e_3,e_2)|product(e_2,e_2,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1071, c_0_1070]), c_0_52])).
cnf(c_0_1075,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)|~product(e_3,e_2,X1)), inference(spm,[status(thm)],[c_0_42, c_0_1072])).
cnf(c_0_1076,negated_conjecture,(product(e_2,e_2,e_2)|product(e_3,e_2,e_5)|product(e_2,e_1,e_2)|product(e_2,e_3,e_2)), inference(spm,[status(thm)],[c_0_1073, c_0_1074])).
cnf(c_0_1077,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_2,e_2)|~product(e_2,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_1070])).
cnf(c_0_1078,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)|product(e_2,e_2,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1075, c_0_1076]), c_0_27])).
cnf(c_0_1079,negated_conjecture,(product(e_2,e_2,e_2)|product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1077, c_0_1078]), c_0_33])).
cnf(c_0_1080,negated_conjecture,(product(e_2,e_2,e_2)|product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1079]), c_0_34])).
cnf(c_0_1081,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_2,e_2)|product(e_2,e_6,e_2)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_956, c_0_1080]), c_0_30])).
cnf(c_0_1082,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_2,e_2)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)|product(X1,e_2,e_2)|~product(X1,e_6,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_1081])).
cnf(c_0_1083,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_2,e_2)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(spm,[status(thm)],[c_0_1082, c_0_1081])).
cnf(c_0_1084,negated_conjecture,(product(e_2,e_2,e_2)|product(e_2,e_1,e_2)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1077, c_0_1083]), c_0_19])).
cnf(c_0_1085,negated_conjecture,(product(e_2,e_2,e_2)|product(e_2,e_1,e_2)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1084]), c_0_14])).
cnf(c_0_1086,negated_conjecture,(product(e_2,e_2,e_2)|product(e_2,e_3,e_2)|product(X1,e_2,e_2)|~product(X1,e_1,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_1085])).
cnf(c_0_1087,negated_conjecture,(equalish(X1,e_4)|~product(e_2,e_3,X1)), inference(spm,[status(thm)],[c_0_42, c_0_584])).
cnf(c_0_1088,negated_conjecture,(product(e_2,e_2,e_2)|product(e_2,e_3,e_2)), inference(spm,[status(thm)],[c_0_1086, c_0_1085])).
cnf(c_0_1089,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|~product(e_2,X1,e_2)), inference(spm,[status(thm)],[c_0_28, c_0_959])).
cnf(c_0_1090,negated_conjecture,(product(e_2,e_2,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1087, c_0_1088]), c_0_45])).
cnf(c_0_1091,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1089, c_0_1090]), c_0_43])).
cnf(c_0_1092,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)|product(e_2,e_5,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1091]), c_0_32])).
cnf(c_0_1093,negated_conjecture,(equalish(X1,e_2)|~product(e_2,X1,e_2)), inference(spm,[status(thm)],[c_0_28, c_0_1090])).
cnf(c_0_1094,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_956, c_0_1092]), c_0_27])).
cnf(c_0_1095,negated_conjecture,(product(e_2,e_4,e_1)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1093, c_0_1094]), c_0_12])).
cnf(c_0_1096,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_5)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_7, c_0_1095])).
cnf(c_0_1097,negated_conjecture,(product(e_1,e_4,e_1)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_2,e_1,e_6)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_975, c_0_1095])).
cnf(c_0_1098,negated_conjecture,(product(e_2,e_4,e_5)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_1,e_4,e_5)|product(e_2,e_1,e_6)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1096, c_0_1097]), c_0_12])).
cnf(c_0_1099,negated_conjecture,(product(e_2,e_4,e_5)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_2,e_1,e_6)|product(X1,e_1,e_5)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_1098])).
cnf(c_0_1100,negated_conjecture,(product(e_2,e_4,e_5)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_1,e_4,e_4)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)), inference(spm,[status(thm)],[c_0_1099, c_0_1095])).
cnf(c_0_1101,negated_conjecture,(product(e_2,e_4,e_5)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(X1,e_1,e_4)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_1100])).
cnf(c_0_1102,negated_conjecture,(product(e_2,e_4,e_5)|product(e_1,e_4,e_2)|product(e_1,e_4,e_3)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)), inference(spm,[status(thm)],[c_0_1101, c_0_1095])).
cnf(c_0_1103,negated_conjecture,(product(e_2,e_4,e_5)|product(e_1,e_4,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(X1,e_1,e_3)|~product(X1,e_4,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_1102])).
cnf(c_0_1104,negated_conjecture,(product(e_2,e_4,e_5)|product(X1,e_2,e_1)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_1095])).
cnf(c_0_1105,negated_conjecture,(product(e_2,e_4,e_5)|product(e_1,e_4,e_2)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(spm,[status(thm)],[c_0_1103, c_0_1095])).
cnf(c_0_1106,negated_conjecture,(product(e_2,e_4,e_5)|product(e_1,e_2,e_1)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(spm,[status(thm)],[c_0_1104, c_0_1105])).
cnf(c_0_1107,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1011, c_0_1095]), c_0_46])).
cnf(c_0_1108,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|product(X1,e_1,e_1)|~product(X1,e_2,e_1)), inference(spm,[status(thm)],[c_0_21, c_0_1106])).
cnf(c_0_1109,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_7, c_0_1107])).
cnf(c_0_1110,negated_conjecture,(product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_5,e_6,e_5)|product(e_2,e_5,e_1)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)), inference(spm,[status(thm)],[c_0_880, c_0_1107])).
cnf(c_0_1111,negated_conjecture,(equalish(X1,e_2)|product(e_2,e_4,e_5)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|~product(e_1,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_1106])).
cnf(c_0_1112,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|product(e_1,e_1,e_1)), inference(spm,[status(thm)],[c_0_1108, c_0_1106])).
cnf(c_0_1113,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_5,e_6,e_4)|product(e_2,e_5,e_1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1109, c_0_1110]), c_0_35])).
cnf(c_0_1114,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_1,e_6)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1111, c_0_1112]), c_0_12])).
cnf(c_0_1115,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_2,e_5,e_1)|product(X1,e_5,e_4)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_1113])).
cnf(c_0_1116,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|~greater(e_6,X1)|~next(e_2,X1)), inference(spm,[status(thm)],[c_0_3, c_0_1114])).
cnf(c_0_1117,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_5,e_6,e_2)|product(e_5,e_6,e_3)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)), inference(spm,[status(thm)],[c_0_1115, c_0_1107])).
cnf(c_0_1118,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1116, c_0_5]), c_0_6])])).
cnf(c_0_1119,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_5,e_6,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(X1,e_5,e_3)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_1117])).
cnf(c_0_1120,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)|~greater(e_5,X1)|~next(e_2,X1)), inference(spm,[status(thm)],[c_0_3, c_0_1118])).
cnf(c_0_1121,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_5,e_6,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)), inference(spm,[status(thm)],[c_0_1119, c_0_1107])).
cnf(c_0_1122,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_1,e_4)|product(e_2,e_1,e_3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1120, c_0_22]), c_0_6])])).
cnf(c_0_1123,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(X1,e_5,e_2)|~product(X1,e_6,e_5)), inference(spm,[status(thm)],[c_0_21, c_0_1121])).
cnf(c_0_1124,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_1,e_3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1122]), c_0_14])).
cnf(c_0_1125,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_4,e_5)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(spm,[status(thm)],[c_0_1123, c_0_1107])).
cnf(c_0_1126,negated_conjecture,(equalish(X1,e_1)|product(e_2,e_4,e_5)|~product(e_2,X1,e_3)), inference(spm,[status(thm)],[c_0_28, c_0_1124])).
cnf(c_0_1127,negated_conjecture,(product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_4,e_5)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1125]), c_0_34])).
cnf(c_0_1128,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_6,e_2)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1126, c_0_1127]), c_0_30])).
cnf(c_0_1129,negated_conjecture,(equalish(X1,e_4)|product(e_2,e_4,e_5)|~product(e_2,X1,e_1)), inference(spm,[status(thm)],[c_0_28, c_0_1095])).
cnf(c_0_1130,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_5,e_1)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1093, c_0_1128]), c_0_36])).
cnf(c_0_1131,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_5,e_4)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1129, c_0_1130]), c_0_31])).
cnf(c_0_1132,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_5,e_3)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1131]), c_0_32])).
cnf(c_0_1133,negated_conjecture,(product(e_2,e_4,e_5)|product(e_2,e_5,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1126, c_0_1132]), c_0_27])).
cnf(c_0_1134,negated_conjecture,(product(e_2,e_4,e_5)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1093, c_0_1133]), c_0_35])).
cnf(c_0_1135,negated_conjecture,(product(X1,e_2,e_5)|~product(X1,e_4,e_2)), inference(spm,[status(thm)],[c_0_21, c_0_1134])).
cnf(c_0_1136,negated_conjecture,(product(e_2,e_1,e_2)|product(e_3,e_4,e_2)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1071, c_0_1134]), c_0_52])).
cnf(c_0_1137,negated_conjecture,(product(e_3,e_2,e_5)|product(e_2,e_1,e_2)|product(e_2,e_3,e_2)), inference(spm,[status(thm)],[c_0_1135, c_0_1136])).
cnf(c_0_1138,negated_conjecture,(equalish(X1,e_4)|~product(e_2,X1,e_5)), inference(spm,[status(thm)],[c_0_28, c_0_1134])).
cnf(c_0_1139,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_6,e_5)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1075, c_0_1137]), c_0_27])).
cnf(c_0_1140,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_6,e_4)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1138, c_0_1139]), c_0_33])).
cnf(c_0_1141,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_6,e_3)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1140]), c_0_34])).
cnf(c_0_1142,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_6,e_2)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_956, c_0_1141]), c_0_30])).
cnf(c_0_1143,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_1,e_5)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1093, c_0_1142]), c_0_36])).
cnf(c_0_1144,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_1,e_4)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1138, c_0_1143]), c_0_19])).
cnf(c_0_1145,negated_conjecture,(product(e_2,e_1,e_2)|product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_1144]), c_0_14])).
cnf(c_0_1146,negated_conjecture,(product(e_2,e_3,e_2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1093, c_0_1145]), c_0_12])).
cnf(c_0_1147,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1093, c_0_1146]), c_0_26]), ['proof']).
# SZS output end CNFRefutation.
EOF
