% Problem    : SYN861-1 : TPTP v6.1.0. Released v2.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n099.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:27:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,negated_conjecture,(ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause2)).
cnf(c_0_1,negated_conjecture,(ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1)).
cnf(c_0_2,negated_conjecture,(ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause3)).
cnf(c_0_3,negated_conjecture,(ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause8)).
cnf(c_0_4,negated_conjecture,(ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause17)).
cnf(c_0_5,negated_conjecture,(ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause30)).
cnf(c_0_6,negated_conjecture,(ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause45)).
cnf(c_0_7,negated_conjecture,(ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause58)).
cnf(c_0_8,negated_conjecture,(ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause73)).
cnf(c_0_9,negated_conjecture,(ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause97)).
cnf(c_0_10,negated_conjecture,(ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause120)).
cnf(c_0_11,negated_conjecture,(ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause147)).
cnf(c_0_12,negated_conjecture,(ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause175)).
cnf(c_0_13,negated_conjecture,(ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause206)).
cnf(c_0_14,negated_conjecture,(ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause236)).
cnf(c_0_15,negated_conjecture,(ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause268)).
cnf(c_0_16,negated_conjecture,(ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause304)).
cnf(c_0_17,negated_conjecture,(ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause343)).
cnf(c_0_18,negated_conjecture,(ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause382)).
cnf(c_0_19,negated_conjecture,(ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause423)).
cnf(c_0_20,negated_conjecture,(ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause465)).
cnf(c_0_21,negated_conjecture,(ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause509)).
cnf(c_0_22,negated_conjecture,(ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause553)).
cnf(c_0_23,negated_conjecture,(ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause601)).
cnf(c_0_24,negated_conjecture,(ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause651)).
cnf(c_0_25,negated_conjecture,(ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause705)).
cnf(c_0_26,negated_conjecture,(ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause760)).
cnf(c_0_27,negated_conjecture,(ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause817)).
cnf(c_0_28,negated_conjecture,(ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause878)).
cnf(c_0_29,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv24_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause796)).
cnf(c_0_30,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv24_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv35_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause924)).
cnf(c_0_31,negated_conjecture,(ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause944)).
cnf(c_0_32,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause788)).
cnf(c_0_33,negated_conjecture,(ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1012)).
cnf(c_0_34,negated_conjecture,(ssPv35_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause774)).
cnf(c_0_35,negated_conjecture,(ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1077)).
cnf(c_0_36,negated_conjecture,(~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,skc45)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause707)).
cnf(c_0_37,negated_conjecture,(ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1147)).
cnf(c_0_38,negated_conjecture,(ssPv26_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv31_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv34_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause454)).
cnf(c_0_39,negated_conjecture,(ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1223)).
cnf(c_0_40,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssPv26_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause379)).
cnf(c_0_41,negated_conjecture,(ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1293)).
cnf(c_0_42,negated_conjecture,(ssPv31_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause365)).
cnf(c_0_43,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv31_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause781)).
cnf(c_0_44,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv31_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv34_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv38_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause922)).
cnf(c_0_45,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv35_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv38_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv32_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause324)).
cnf(c_0_46,negated_conjecture,(ssPv32_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause259)).
cnf(c_0_47,negated_conjecture,(~ssPv6_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv16_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv32_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv33_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1582)).
cnf(c_0_48,negated_conjecture,(ssPv16_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1414)).
cnf(c_0_49,negated_conjecture,(ssPv34_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause356)).
cnf(c_0_50,negated_conjecture,(ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv34_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause776)).
cnf(c_0_51,negated_conjecture,(ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssPv35_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause249)).
cnf(c_0_52,negated_conjecture,(ssPv32_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1382)).
cnf(c_0_53,negated_conjecture,(~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,skc51)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause238)).
cnf(c_0_54,negated_conjecture,(ssPv38_3r1r1r1(X1,X2,X3)|~ssPv38_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause768)).
cnf(c_0_55,negated_conjecture,(ssPv38_3r1r1r1(X1,X2,X3)|~ssPv38_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause244)).
cnf(c_0_56,negated_conjecture,(ssPv33_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1381)).
cnf(c_0_57,negated_conjecture,(~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,skc55)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause177)).
cnf(c_0_58,negated_conjecture,(~ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,skc53)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause208)).
cnf(c_0_59,negated_conjecture,(ssPv6_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv15_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv19_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv27_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1574)).
cnf(c_0_60,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,skc60)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause98)).
cnf(c_0_61,negated_conjecture,(~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,skc61)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause99)).
cnf(c_0_62,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc62)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause74)).
cnf(c_0_63,negated_conjecture,(~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv16_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv19_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv38_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1285)).
cnf(c_0_64,negated_conjecture,(ssPv16_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1131)).
cnf(c_0_65,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv15_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1416)).
cnf(c_0_66,negated_conjecture,(ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1365)).
cnf(c_0_67,negated_conjecture,(ssPv19_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1123)).
cnf(c_0_68,negated_conjecture,(ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssPv19_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1408)).
cnf(c_0_69,negated_conjecture,(ssPv38_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssPv38_3r1r1r1(X1,X2,X3)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1445)).
cnf(c_0_70,negated_conjecture,(ssPv38_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssPv38_3r1r1r1(X1,X2,X3)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1084)).
cnf(c_0_71,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv27_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1392)).
cnf(c_0_72,negated_conjecture,(ssPv32_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv5_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv10_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv38_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1658)).
cnf(c_0_73,negated_conjecture,(ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,skc34)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1013)).
cnf(c_0_74,negated_conjecture,(ssPv5_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv28_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv33_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv14_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1654)).
cnf(c_0_75,negated_conjecture,(ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,skc46)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause652)).
cnf(c_0_76,negated_conjecture,(ssPv10_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1506)).
cnf(c_0_77,negated_conjecture,(ssPv14_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv14_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1493)).
cnf(c_0_78,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv28_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1464)).
cnf(c_0_79,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv32_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1456)).
cnf(c_0_80,negated_conjecture,(ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv33_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1455)).
cnf(c_0_81,negated_conjecture,(ssPv14_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,skc42)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause761)).
cnf(c_0_82,negated_conjecture,(~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc63)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause75)).
cnf(c_0_83,negated_conjecture,(ssNder1_1r1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1])])).
cnf(c_0_84,negated_conjecture,(ssNder1_2r1r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_1]), c_0_83])])).
cnf(c_0_85,negated_conjecture,(ssNder1_3r1r1r1(X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_1]), c_0_83]), c_0_84])])).
cnf(c_0_86,negated_conjecture,(ssNder1_4r1r1r1r1(X1,X2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_1]), c_0_83]), c_0_84]), c_0_85])])).
cnf(c_0_87,negated_conjecture,(ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_88,negated_conjecture,(ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87])])).
cnf(c_0_89,negated_conjecture,(ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88])])).
cnf(c_0_90,negated_conjecture,(ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89])])).
cnf(c_0_91,negated_conjecture,(ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90])])).
cnf(c_0_92,negated_conjecture,(ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91])])).
cnf(c_0_93,negated_conjecture,(ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92])])).
cnf(c_0_94,negated_conjecture,(ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93])])).
cnf(c_0_95,negated_conjecture,(ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94])])).
cnf(c_0_96,negated_conjecture,(ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_97,negated_conjecture,(ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96])])).
cnf(c_0_98,negated_conjecture,(ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97])])).
cnf(c_0_99,negated_conjecture,(ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_100,negated_conjecture,(ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99])])).
cnf(c_0_101,negated_conjecture,(ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100])])).
cnf(c_0_102,negated_conjecture,(ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_103,negated_conjecture,(ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102])])).
cnf(c_0_104,negated_conjecture,(ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_105,negated_conjecture,(ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_106,negated_conjecture,(ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105])])).
cnf(c_0_107,negated_conjecture,(ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_108,negated_conjecture,(ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_109,negated_conjecture,(ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108])])).
cnf(c_0_110,negated_conjecture,(ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109])])).
cnf(c_0_111,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv24_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_112,negated_conjecture,(ssPv24_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv35_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_113,negated_conjecture,(ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110])])).
cnf(c_0_114,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_115,negated_conjecture,(ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv35_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_116,negated_conjecture,(ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113])])).
cnf(c_0_117,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv35_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_118,negated_conjecture,(ssPv35_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_119,negated_conjecture,(ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116])])).
cnf(c_0_120,negated_conjecture,(~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,skc45)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_121,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_122,negated_conjecture,(ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119])])).
cnf(c_0_123,negated_conjecture,(ssPv26_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv31_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv34_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_124,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_125,negated_conjecture,(ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122])])).
cnf(c_0_126,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssPv26_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_127,negated_conjecture,(ssPv26_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv34_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_128,negated_conjecture,(ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125])])).
cnf(c_0_129,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv34_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_130,negated_conjecture,(ssPv31_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_131,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv31_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_132,negated_conjecture,(ssPv31_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv34_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv38_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_133,negated_conjecture,(ssPv35_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv38_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv32_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_134,negated_conjecture,(ssPv32_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_135,negated_conjecture,(~ssPv6_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv16_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv32_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv33_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128])])).
cnf(c_0_136,negated_conjecture,(ssPv16_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128])])).
cnf(c_0_137,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv34_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_138,negated_conjecture,(ssPv34_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_139,negated_conjecture,(ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv34_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_140,negated_conjecture,(ssPv34_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv38_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_141,negated_conjecture,(ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssPv35_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_142,negated_conjecture,(ssPv35_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv38_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_143,negated_conjecture,(~ssPv6_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv32_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv33_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_144,negated_conjecture,(ssPv32_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128])])).
cnf(c_0_145,negated_conjecture,(~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,skc51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_146,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_147,negated_conjecture,(ssPv38_3r1r1r1(X1,X2,X3)|~ssPv38_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_148,negated_conjecture,(ssPv38_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_149,negated_conjecture,(ssPv38_3r1r1r1(X1,X2,X3)|~ssPv38_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_150,negated_conjecture,(ssPv38_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_151,negated_conjecture,(~ssPv6_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv33_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_152,negated_conjecture,(ssPv33_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128])])).
cnf(c_0_153,negated_conjecture,(~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,skc55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93])])).
cnf(c_0_154,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_145, c_0_146])).
cnf(c_0_155,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_147, c_0_148])).
cnf(c_0_156,negated_conjecture,(~ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,skc53)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94])])).
cnf(c_0_157,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|ssPv38_3r1r1r1(X1,X2,X3)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_149, c_0_150])).
cnf(c_0_158,negated_conjecture,(~ssPv6_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_159,negated_conjecture,(ssPv6_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv15_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv19_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv27_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128])])).
cnf(c_0_160,negated_conjecture,(~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_153, c_0_154])).
cnf(c_0_161,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,skc60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90])])).
cnf(c_0_162,negated_conjecture,(~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,skc61)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90])])).
cnf(c_0_163,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_120, c_0_155])).
cnf(c_0_164,negated_conjecture,(ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|ssPv38_3r1r1r1(X1,X2,X3)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_165,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89])])).
cnf(c_0_166,negated_conjecture,(~ssPv16_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv19_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv38_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116])])).
cnf(c_0_167,negated_conjecture,(ssPv16_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116])])).
cnf(c_0_168,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv15_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128])])).
cnf(c_0_169,negated_conjecture,(ssPv15_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv19_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv27_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_158, c_0_159])).
cnf(c_0_170,negated_conjecture,(ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128])])).
cnf(c_0_171,negated_conjecture,(~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_172,negated_conjecture,(ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_162, c_0_163])).
cnf(c_0_173,negated_conjecture,(ssPv35_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_174,negated_conjecture,(~ssPv19_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv38_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(spm,[status(thm)],[c_0_166, c_0_167])).
cnf(c_0_175,negated_conjecture,(ssPv19_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116])])).
cnf(c_0_176,negated_conjecture,(ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssPv19_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128])])).
cnf(c_0_177,negated_conjecture,(ssPv19_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv27_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_168, c_0_169])).
cnf(c_0_178,negated_conjecture,(ssPv38_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128]), c_0_170])])).
cnf(c_0_179,negated_conjecture,(ssPv38_3r1r1r1(X1,X2,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_173])).
cnf(c_0_180,negated_conjecture,(~ssPv38_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)), inference(spm,[status(thm)],[c_0_174, c_0_175])).
cnf(c_0_181,negated_conjecture,(ssPv38_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116])])).
cnf(c_0_182,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv27_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128])])).
cnf(c_0_183,negated_conjecture,(ssPv27_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_176, c_0_177])).
cnf(c_0_184,negated_conjecture,(ssPv32_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv5_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv10_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv38_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128]), c_0_170])])).
cnf(c_0_185,negated_conjecture,(ssPv38_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_178, c_0_179])])).
cnf(c_0_186,negated_conjecture,(~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_180, c_0_181])).
cnf(c_0_187,negated_conjecture,(ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,skc34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113])])).
cnf(c_0_188,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_182, c_0_183])).
cnf(c_0_189,negated_conjecture,(ssPv32_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv5_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv10_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_184, c_0_185])])).
cnf(c_0_190,negated_conjecture,(ssPv5_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv28_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv33_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv14_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128]), c_0_170])])).
cnf(c_0_191,negated_conjecture,(~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_186, c_0_187])).
cnf(c_0_192,negated_conjecture,(ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,skc46)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105])])).
cnf(c_0_193,negated_conjecture,(ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_120, c_0_188])).
cnf(c_0_194,negated_conjecture,(ssPv28_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv32_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv33_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv10_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv14_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)), inference(spm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_195,negated_conjecture,(ssPv10_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128]), c_0_170])])).
cnf(c_0_196,negated_conjecture,(~ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_191, c_0_192])).
cnf(c_0_197,negated_conjecture,(ssPv19_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_193, c_0_192])).
cnf(c_0_198,negated_conjecture,(ssPv28_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv32_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv33_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv14_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)), inference(spm,[status(thm)],[c_0_194, c_0_195])).
cnf(c_0_199,negated_conjecture,(ssPv14_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv14_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128]), c_0_170])])).
cnf(c_0_200,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_196, c_0_197])).
cnf(c_0_201,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv28_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128]), c_0_170])])).
cnf(c_0_202,negated_conjecture,(ssPv28_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv32_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv33_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv14_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(spm,[status(thm)],[c_0_198, c_0_199])).
cnf(c_0_203,negated_conjecture,(~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_156, c_0_200])).
cnf(c_0_204,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv32_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128]), c_0_170])])).
cnf(c_0_205,negated_conjecture,(ssPv32_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv33_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv14_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(spm,[status(thm)],[c_0_201, c_0_202])).
cnf(c_0_206,negated_conjecture,(~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_203, c_0_165])).
cnf(c_0_207,negated_conjecture,(ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv33_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107]), c_0_108]), c_0_109]), c_0_110]), c_0_113]), c_0_116]), c_0_119]), c_0_122]), c_0_125]), c_0_128]), c_0_170])])).
cnf(c_0_208,negated_conjecture,(ssPv33_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv14_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(spm,[status(thm)],[c_0_204, c_0_205])).
cnf(c_0_209,negated_conjecture,(~ssPv33_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_206, c_0_179])])).
cnf(c_0_210,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv10_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssPv14_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208]), c_0_209])).
cnf(c_0_211,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv14_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(spm,[status(thm)],[c_0_210, c_0_187])).
cnf(c_0_212,negated_conjecture,(ssPv14_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,skc42)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_213,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_211, c_0_212])).
cnf(c_0_214,negated_conjecture,(~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_1]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89])])).
cnf(c_0_215,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_153, c_0_213])).
cnf(c_0_216,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_214, c_0_215])]), ['proof']).
# SZS output end CNFRefutation.
EOF
