% Problem    : SYN850-1 : TPTP v6.1.0. Released v2.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n089.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:46:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,negated_conjecture,(ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause2)).
cnf(c_0_1,negated_conjecture,(ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1)).
cnf(c_0_2,negated_conjecture,(ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause3)).
cnf(c_0_3,negated_conjecture,(ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause8)).
cnf(c_0_4,negated_conjecture,(ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause14)).
cnf(c_0_5,negated_conjecture,(ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause26)).
cnf(c_0_6,negated_conjecture,(ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause40)).
cnf(c_0_7,negated_conjecture,(ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause56)).
cnf(c_0_8,negated_conjecture,(ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause74)).
cnf(c_0_9,negated_conjecture,(ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause93)).
cnf(c_0_10,negated_conjecture,(ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause117)).
cnf(c_0_11,negated_conjecture,(ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause139)).
cnf(c_0_12,negated_conjecture,(ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause166)).
cnf(c_0_13,negated_conjecture,(ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause197)).
cnf(c_0_14,negated_conjecture,(ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause228)).
cnf(c_0_15,negated_conjecture,(ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause260)).
cnf(c_0_16,negated_conjecture,(ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause296)).
cnf(c_0_17,negated_conjecture,(ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause330)).
cnf(c_0_18,negated_conjecture,(ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause370)).
cnf(c_0_19,negated_conjecture,(ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause409)).
cnf(c_0_20,negated_conjecture,(ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause450)).
cnf(c_0_21,negated_conjecture,(ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause492)).
cnf(c_0_22,negated_conjecture,(ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause536)).
cnf(c_0_23,negated_conjecture,(ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause585)).
cnf(c_0_24,negated_conjecture,(ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause634)).
cnf(c_0_25,negated_conjecture,(ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause684)).
cnf(c_0_26,negated_conjecture,(ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause738)).
cnf(c_0_27,negated_conjecture,(ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause796)).
cnf(c_0_28,negated_conjecture,(~ssPv1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv15_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv21_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1046)).
cnf(c_0_29,negated_conjecture,(ssPv1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|ssPv24_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv25_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1036)).
cnf(c_0_30,negated_conjecture,(ssPv30_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv17_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv20_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause846)).
cnf(c_0_31,negated_conjecture,(ssPv17_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause736)).
cnf(c_0_32,negated_conjecture,(ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssPv17_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause792)).
cnf(c_0_33,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv17_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv26_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv21_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause901)).
cnf(c_0_34,negated_conjecture,(ssPv15_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause921)).
cnf(c_0_35,negated_conjecture,(ssPv20_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause731)).
cnf(c_0_36,negated_conjecture,(ssPv21_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause782)).
cnf(c_0_37,negated_conjecture,(ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv20_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause785)).
cnf(c_0_38,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv20_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv24_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause908)).
cnf(c_0_39,negated_conjecture,(ssPv21_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause897)).
cnf(c_0_40,negated_conjecture,(ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv30_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause708)).
cnf(c_0_41,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssPv26_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause769)).
cnf(c_0_42,negated_conjecture,(ssPv24_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause774)).
cnf(c_0_43,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv24_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause891)).
cnf(c_0_44,negated_conjecture,(ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,skc46)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause635)).
cnf(c_0_45,negated_conjecture,(~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,skc45)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause686)).
cnf(c_0_46,negated_conjecture,(ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause766)).
cnf(c_0_47,negated_conjecture,(ssPv25_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause889)).
cnf(c_0_48,negated_conjecture,(ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,skc40)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause797)).
cnf(c_0_49,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,skc44)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause685)).
cnf(c_0_50,negated_conjecture,(ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|ssPv32_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv31_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause570)).
cnf(c_0_51,negated_conjecture,(ssPv31_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause469)).
cnf(c_0_52,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv32_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause467)).
cnf(c_0_53,negated_conjecture,(ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause464)).
cnf(c_0_54,negated_conjecture,(ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,skc48)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause261)).
cnf(c_0_55,negated_conjecture,(ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause856)).
cnf(c_0_56,negated_conjecture,(~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,skc51)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause230)).
cnf(c_0_57,negated_conjecture,(ssPv34_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv25_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv38_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1110)).
cnf(c_0_58,negated_conjecture,(ssPv25_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause958)).
cnf(c_0_59,negated_conjecture,(ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause924)).
cnf(c_0_60,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,skc54)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause167)).
cnf(c_0_61,negated_conjecture,(ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv34_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause940)).
cnf(c_0_62,negated_conjecture,(ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause984)).
cnf(c_0_63,negated_conjecture,(~ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,skc59)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause119)).
cnf(c_0_64,negated_conjecture,(ssPv38_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssPv38_3r1r1r1(X1,X2,X3)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause931)).
cnf(c_0_65,negated_conjecture,(ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1053)).
cnf(c_0_66,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,skc60)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause94)).
cnf(c_0_67,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,skc38)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause857)).
cnf(c_0_68,negated_conjecture,(ssPv38_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv23_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1320)).
cnf(c_0_69,negated_conjecture,(ssPv15_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1177)).
cnf(c_0_70,negated_conjecture,(~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc63)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause76)).
cnf(c_0_71,negated_conjecture,(ssPv23_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1157)).
cnf(c_0_72,negated_conjecture,(ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssPv23_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause722)).
cnf(c_0_73,negated_conjecture,(ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|ssPv23_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|ssPv38_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv32_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause837)).
cnf(c_0_74,negated_conjecture,(ssPv38_3r1r1r1(X1,X2,X3)|~ssPv38_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1127)).
cnf(c_0_75,negated_conjecture,(ssPv32_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause703)).
cnf(c_0_76,negated_conjecture,(ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,skc32)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1054)).
cnf(c_0_77,negated_conjecture,(ssPv38_3r1r1r1(X1,X2,X3)|~ssPv38_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause692)).
cnf(c_0_78,negated_conjecture,(~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,skc47)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause636)).
cnf(c_0_79,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc62)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause75)).
cnf(c_0_80,negated_conjecture,(ssNder1_1r1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1])])).
cnf(c_0_81,negated_conjecture,(ssNder1_2r1r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_1]), c_0_80])])).
cnf(c_0_82,negated_conjecture,(ssNder1_3r1r1r1(X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_1]), c_0_80]), c_0_81])])).
cnf(c_0_83,negated_conjecture,(ssNder1_4r1r1r1r1(X1,X2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_1]), c_0_80]), c_0_81]), c_0_82])])).
cnf(c_0_84,negated_conjecture,(ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83])])).
cnf(c_0_85,negated_conjecture,(ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84])])).
cnf(c_0_86,negated_conjecture,(ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85])])).
cnf(c_0_87,negated_conjecture,(ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_88,negated_conjecture,(ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87])])).
cnf(c_0_89,negated_conjecture,(ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88])])).
cnf(c_0_90,negated_conjecture,(ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89])])).
cnf(c_0_91,negated_conjecture,(ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90])])).
cnf(c_0_92,negated_conjecture,(ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91])])).
cnf(c_0_93,negated_conjecture,(ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92])])).
cnf(c_0_94,negated_conjecture,(ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93])])).
cnf(c_0_95,negated_conjecture,(ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94])])).
cnf(c_0_96,negated_conjecture,(ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_97,negated_conjecture,(ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96])])).
cnf(c_0_98,negated_conjecture,(ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97])])).
cnf(c_0_99,negated_conjecture,(ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_100,negated_conjecture,(ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99])])).
cnf(c_0_101,negated_conjecture,(ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100])])).
cnf(c_0_102,negated_conjecture,(ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_103,negated_conjecture,(ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102])])).
cnf(c_0_104,negated_conjecture,(ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_105,negated_conjecture,(ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_106,negated_conjecture,(ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105])])).
cnf(c_0_107,negated_conjecture,(~ssPv1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv15_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv21_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_108,negated_conjecture,(ssPv1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|ssPv24_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv25_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_109,negated_conjecture,(ssPv30_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv17_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv20_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_110,negated_conjecture,(ssPv17_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_111,negated_conjecture,(ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssPv17_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_112,negated_conjecture,(ssPv17_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv26_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv21_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_113,negated_conjecture,(ssPv24_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv15_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv21_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv25_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_114,negated_conjecture,(ssPv15_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_115,negated_conjecture,(ssPv30_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssPv20_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_116,negated_conjecture,(ssPv20_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_117,negated_conjecture,(ssPv26_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv21_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_118,negated_conjecture,(ssPv21_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_119,negated_conjecture,(ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv20_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_120,negated_conjecture,(ssPv20_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv24_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_121,negated_conjecture,(ssPv24_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv21_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv25_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_122,negated_conjecture,(ssPv21_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_123,negated_conjecture,(ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv30_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_124,negated_conjecture,(ssPv30_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_125,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssPv26_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_126,negated_conjecture,(ssPv26_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_127,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv24_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_128,negated_conjecture,(ssPv24_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_129,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv24_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_130,negated_conjecture,(ssPv24_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv25_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_131,negated_conjecture,(ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)), inference(spm,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_132,negated_conjecture,(ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,skc46)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102])])).
cnf(c_0_133,negated_conjecture,(~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,skc45)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_134,negated_conjecture,(ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_135,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(spm,[status(thm)],[c_0_127, c_0_128])).
cnf(c_0_136,negated_conjecture,(ssPv28_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_137,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv25_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_138,negated_conjecture,(ssPv25_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_139,negated_conjecture,(ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_140,negated_conjecture,(ssPv17_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_141,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_142,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_143,negated_conjecture,(ssPv13_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,skc40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105])])).
cnf(c_0_144,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_145,negated_conjecture,(ssPv20_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)), inference(spm,[status(thm)],[c_0_133, c_0_141])).
cnf(c_0_146,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_147,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,skc44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_148,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)), inference(spm,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_149,negated_conjecture,(ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|ssPv32_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv31_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_150,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_151,negated_conjecture,(ssPv32_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv31_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)), inference(spm,[status(thm)],[c_0_148, c_0_149])).
cnf(c_0_152,negated_conjecture,(ssPv31_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_153,negated_conjecture,(ssPv32_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv31_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_150, c_0_149])).
cnf(c_0_154,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv32_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_155,negated_conjecture,(ssPv32_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_156,negated_conjecture,(ssPv32_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(spm,[status(thm)],[c_0_153, c_0_152])).
cnf(c_0_157,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(spm,[status(thm)],[c_0_154, c_0_155])).
cnf(c_0_158,negated_conjecture,(ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_159,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv34_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(spm,[status(thm)],[c_0_154, c_0_156])).
cnf(c_0_160,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_157, c_0_158])).
cnf(c_0_161,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_159, c_0_158])).
cnf(c_0_162,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_163,negated_conjecture,(ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,skc48)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93])])).
cnf(c_0_164,negated_conjecture,(ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_165,negated_conjecture,(~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,skc51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92])])).
cnf(c_0_166,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_162, c_0_163])).
cnf(c_0_167,negated_conjecture,(ssPv34_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv25_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv38_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164])])).
cnf(c_0_168,negated_conjecture,(ssPv25_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164])])).
cnf(c_0_169,negated_conjecture,(ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164])])).
cnf(c_0_170,negated_conjecture,(ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_171,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,skc54)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90])])).
cnf(c_0_172,negated_conjecture,(ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv34_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164])])).
cnf(c_0_173,negated_conjecture,(ssPv34_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv38_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_167, c_0_168])).
cnf(c_0_174,negated_conjecture,(ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164]), c_0_169])])).
cnf(c_0_175,negated_conjecture,(~ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,skc59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88])])).
cnf(c_0_176,negated_conjecture,(ssPv30_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_170, c_0_171])).
cnf(c_0_177,negated_conjecture,(ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv38_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(spm,[status(thm)],[c_0_172, c_0_173])).
cnf(c_0_178,negated_conjecture,(ssPv38_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164])])).
cnf(c_0_179,negated_conjecture,(ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164]), c_0_169]), c_0_174])])).
cnf(c_0_180,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_175, c_0_176])).
cnf(c_0_181,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,skc60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87])])).
cnf(c_0_182,negated_conjecture,(ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_177, c_0_178])).
cnf(c_0_183,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,skc38)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_184,negated_conjecture,(ssPv38_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv23_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164]), c_0_169]), c_0_174]), c_0_179])])).
cnf(c_0_185,negated_conjecture,(ssPv15_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164]), c_0_169]), c_0_174]), c_0_179])])).
cnf(c_0_186,negated_conjecture,(~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc63)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_187,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_180, c_0_181])).
cnf(c_0_188,negated_conjecture,(ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv25_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_182, c_0_183])).
cnf(c_0_189,negated_conjecture,(ssPv38_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv23_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_184, c_0_185])).
cnf(c_0_190,negated_conjecture,(ssPv23_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164]), c_0_169]), c_0_174]), c_0_179])])).
cnf(c_0_191,negated_conjecture,(~ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_186, c_0_187])).
cnf(c_0_192,negated_conjecture,(ssPv34_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_188, c_0_163])).
cnf(c_0_193,negated_conjecture,(ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssPv23_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_194,negated_conjecture,(ssPv23_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|ssPv38_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv32_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_195,negated_conjecture,(ssPv38_3r1r1r1(X1,X2,X3)|~ssPv38_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164]), c_0_169]), c_0_174]), c_0_179])])).
cnf(c_0_196,negated_conjecture,(ssPv38_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_197,negated_conjecture,(~ssPv38_3r1r1r1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_191, c_0_192])).
cnf(c_0_198,negated_conjecture,(ssPv38_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssPv32_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(spm,[status(thm)],[c_0_193, c_0_194])).
cnf(c_0_199,negated_conjecture,(ssPv32_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_200,negated_conjecture,(~ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196]), c_0_197])).
cnf(c_0_201,negated_conjecture,(ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,skc32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_164]), c_0_169]), c_0_174])])).
cnf(c_0_202,negated_conjecture,(ssPv38_3r1r1r1(X1,X2,X3)|~ssPv38_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103])])).
cnf(c_0_203,negated_conjecture,(ssPv38_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_198, c_0_199])).
cnf(c_0_204,negated_conjecture,(~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_200, c_0_201])).
cnf(c_0_205,negated_conjecture,(~ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,skc47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102])])).
cnf(c_0_206,negated_conjecture,(ssPv16_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|ssPv38_3r1r1r1(X1,X2,X3)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_202, c_0_203])).
cnf(c_0_207,negated_conjecture,(~ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(spm,[status(thm)],[c_0_204, c_0_147])).
cnf(c_0_208,negated_conjecture,(ssPv23_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|ssPv38_3r1r1r1(X1,X2,X3)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_205, c_0_206])).
cnf(c_0_209,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_1]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_210,negated_conjecture,(~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208]), c_0_197])).
cnf(c_0_211,negated_conjecture,($false), inference(sr,[status(thm)],[c_0_209, c_0_210]), ['proof']).
# SZS output end CNFRefutation.
EOF
