% Problem    : SYN448-1 : TPTP v6.1.0. Released v2.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n114.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:37:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_N___023_B07_F1_SP_PI_Q7_CS_SP_CO_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,negated_conjecture,(ndr1_0|~hskp5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause25)).
cnf(c_0_1,negated_conjecture,(hskp5|hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1)).
cnf(c_0_2,negated_conjecture,(ndr1_0|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause19)).
cnf(c_0_3,negated_conjecture,(c1_1(X1)|hskp3|hskp12|~c0_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause137)).
cnf(c_0_4,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp19|hskp20|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause133)).
cnf(c_0_5,negated_conjecture,(c0_1(X1)|c3_1(X1)|hskp0|hskp12|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause128)).
cnf(c_0_6,negated_conjecture,(c0_1(X1)|c1_1(X1)|c2_1(X1)|hskp0|hskp1|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause125)).
cnf(c_0_7,negated_conjecture,(c3_1(a466)|~hskp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause80)).
cnf(c_0_8,negated_conjecture,(~c0_1(a466)|~hskp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause117)).
cnf(c_0_9,negated_conjecture,(c1_1(a466)|~hskp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause79)).
cnf(c_0_10,negated_conjecture,(c3_1(a503)|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause57)).
cnf(c_0_11,negated_conjecture,(~c1_1(a503)|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause89)).
cnf(c_0_12,negated_conjecture,(c2_1(a503)|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause56)).
cnf(c_0_13,negated_conjecture,(c3_1(a502)|~hskp19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause59)).
cnf(c_0_14,negated_conjecture,(c2_1(a502)|~hskp19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause58)).
cnf(c_0_15,negated_conjecture,(~c0_1(a502)|~hskp19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause90)).
cnf(c_0_16,negated_conjecture,(c2_1(a471)|~hskp6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause75)).
cnf(c_0_17,negated_conjecture,(~c1_1(a471)|~hskp6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause112)).
cnf(c_0_18,negated_conjecture,(~c3_1(a471)|~hskp6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause113)).
cnf(c_0_19,negated_conjecture,(c1_1(X1)|c2_1(X1)|hskp16|hskp17|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause129)).
cnf(c_0_20,negated_conjecture,(~c1_1(a463)|~hskp1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause122)).
cnf(c_0_21,negated_conjecture,(c2_1(a463)|~hskp1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause81)).
cnf(c_0_22,negated_conjecture,(~c0_1(a463)|~hskp1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause121)).
cnf(c_0_23,negated_conjecture,(~c1_1(a480)|~hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause103)).
cnf(c_0_24,negated_conjecture,(~c0_1(a480)|~hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause102)).
cnf(c_0_25,negated_conjecture,(~c2_1(a480)|~hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause104)).
cnf(c_0_26,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp12|hskp14|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause134)).
cnf(c_0_27,negated_conjecture,(c0_1(a540)|~hskp24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause50)).
cnf(c_0_28,negated_conjecture,(~c1_1(a540)|~hskp24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause83)).
cnf(c_0_29,negated_conjecture,(~c2_1(a540)|~hskp24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause84)).
cnf(c_0_30,negated_conjecture,(c0_1(a460)|~hskp0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause82)).
cnf(c_0_31,negated_conjecture,(~c2_1(a460)|~hskp0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause123)).
cnf(c_0_32,negated_conjecture,(~c3_1(a460)|~hskp0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause124)).
cnf(c_0_33,negated_conjecture,(hskp20|hskp6|hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause37)).
cnf(c_0_34,negated_conjecture,(~c1_1(a494)|~hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause93)).
cnf(c_0_35,negated_conjecture,(~c0_1(a494)|~hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause92)).
cnf(c_0_36,negated_conjecture,(~c3_1(a488)|~hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause100)).
cnf(c_0_37,negated_conjecture,(~c2_1(a488)|~hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause99)).
cnf(c_0_38,negated_conjecture,(~c1_1(a488)|~hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause98)).
cnf(c_0_39,negated_conjecture,(~c3_1(a475)|~hskp8), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause109)).
cnf(c_0_40,negated_conjecture,(c0_1(a475)|~hskp8), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause72)).
cnf(c_0_41,negated_conjecture,(~c3_1(a494)|~hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause94)).
cnf(c_0_42,negated_conjecture,(hskp24|hskp15|hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause36)).
cnf(c_0_43,negated_conjecture,(c1_1(a475)|~hskp8), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause73)).
cnf(c_0_44,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp27|hskp16|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause131)).
cnf(c_0_45,negated_conjecture,(c1_1(a493)|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause62)).
cnf(c_0_46,negated_conjecture,(~c0_1(a493)|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause95)).
cnf(c_0_47,negated_conjecture,(~c2_1(a493)|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause96)).
cnf(c_0_48,negated_conjecture,(c1_1(a492)|~hskp15), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause63)).
cnf(c_0_49,negated_conjecture,(~c2_1(a492)|~hskp15), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause97)).
cnf(c_0_50,negated_conjecture,(c3_1(a468)|~hskp5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause77)).
cnf(c_0_51,negated_conjecture,(c0_1(a468)|~hskp5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause76)).
cnf(c_0_52,negated_conjecture,(~c2_1(a468)|~hskp5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause114)).
cnf(c_0_53,negated_conjecture,(c0_1(X1)|c1_1(X1)|c0_1(X2)|c1_1(X2)|c2_1(X2)|hskp0|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause144)).
cnf(c_0_54,negated_conjecture,(c3_1(a492)|~hskp15), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause64)).
cnf(c_0_55,negated_conjecture,(hskp8|hskp15|hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause32)).
cnf(c_0_56,negated_conjecture,(c2_1(a478)|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause67)).
cnf(c_0_57,negated_conjecture,(~c0_1(a478)|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause105)).
cnf(c_0_58,negated_conjecture,(~c3_1(a478)|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause106)).
cnf(c_0_59,negated_conjecture,(c0_1(X1)|c1_1(X1)|hskp6|hskp2|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause127)).
cnf(c_0_60,negated_conjecture,(c2_1(X1)|c3_1(X1)|c0_1(X2)|c0_1(X3)|c1_1(X3)|c2_1(X3)|~c0_1(X1)|~ndr1_0|~c1_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause172)).
cnf(c_0_61,negated_conjecture,(c3_1(X1)|c1_1(X2)|c2_1(X2)|c3_1(X2)|c0_1(X3)|c1_1(X3)|c3_1(X3)|~c1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause171)).
cnf(c_0_62,negated_conjecture,(c3_1(a473)|~hskp27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause43)).
cnf(c_0_63,negated_conjecture,(c0_1(a473)|~hskp27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause41)).
cnf(c_0_64,negated_conjecture,(c1_1(a473)|~hskp27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause42)).
cnf(c_0_65,negated_conjecture,(c0_1(X1)|c2_1(X1)|c3_1(X1)|c0_1(X2)|c1_1(X2)|c3_1(X2)|hskp0|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause142)).
cnf(c_0_66,negated_conjecture,(c1_1(X1)|c0_1(X2)|c2_1(X2)|hskp8|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c1_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause154)).
cnf(c_0_67,negated_conjecture,(c3_1(a512)|~hskp21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause55)).
cnf(c_0_68,negated_conjecture,(c0_1(a512)|~hskp21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause54)).
cnf(c_0_69,negated_conjecture,(~c1_1(a512)|~hskp21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause88)).
cnf(c_0_70,negated_conjecture,(~c3_1(a477)|~hskp10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause107)).
cnf(c_0_71,negated_conjecture,(c2_1(a477)|~hskp10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause69)).
cnf(c_0_72,negated_conjecture,(c2_1(X1)|c1_1(X2)|c3_1(X2)|c0_1(X3)|~c0_1(X1)|~c1_1(X1)|~ndr1_0|~c2_1(X2)|~c2_1(X3)|~c3_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause176)).
cnf(c_0_73,negated_conjecture,(c3_1(X1)|c1_1(X2)|c0_1(X3)|c1_1(X3)|c2_1(X3)|~c0_1(X1)|~c2_1(X1)|~ndr1_0|~c2_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause173)).
cnf(c_0_74,negated_conjecture,(hskp21|hskp10|hskp6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause35)).
cnf(c_0_75,negated_conjecture,(~c3_1(a465)|~hskp2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause120)).
cnf(c_0_76,negated_conjecture,(~c0_1(a465)|~hskp2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause118)).
cnf(c_0_77,negated_conjecture,(~c2_1(a465)|~hskp2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause119)).
cnf(c_0_78,negated_conjecture,(c2_1(X2)|c3_1(X2)|hskp3|~c0_1(X1)|~c1_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause166)).
cnf(c_0_79,negated_conjecture,(c2_1(X1)|c2_1(X2)|c3_1(X2)|hskp21|~c0_1(X1)|~c3_1(X1)|~ndr1_0|~c1_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause159)).
cnf(c_0_80,negated_conjecture,(c2_1(X1)|c3_1(X1)|c1_1(X2)|c3_1(X2)|hskp12|~c0_1(X1)|~ndr1_0|~c0_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause151)).
cnf(c_0_81,negated_conjecture,(c1_1(X1)|c2_1(X1)|hskp7|hskp12|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause130)).
cnf(c_0_82,negated_conjecture,(c3_1(a461)|~hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause49)).
cnf(c_0_83,negated_conjecture,(c0_1(a461)|~hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause47)).
cnf(c_0_84,negated_conjecture,(c2_1(a461)|~hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause48)).
cnf(c_0_85,negated_conjecture,(c0_1(a533)|~hskp23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause51)).
cnf(c_0_86,negated_conjecture,(~c1_1(a533)|~hskp23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause85)).
cnf(c_0_87,negated_conjecture,(~c3_1(a533)|~hskp23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause86)).
cnf(c_0_88,negated_conjecture,(c0_1(X2)|hskp14|~c0_1(X1)|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c1_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause168)).
cnf(c_0_89,negated_conjecture,(c2_1(X1)|hskp15|hskp2|~c0_1(X1)|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause139)).
cnf(c_0_90,negated_conjecture,(c0_1(a476)|~hskp9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause70)).
cnf(c_0_91,negated_conjecture,(~c1_1(a476)|~hskp9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause108)).
cnf(c_0_92,negated_conjecture,(hskp9|hskp23|hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause34)).
cnf(c_0_93,negated_conjecture,(c1_1(X3)|c2_1(X3)|c3_1(X3)|c0_1(X2)|~c0_1(X1)|~c1_1(X1)|~c3_1(X1)|~ndr1_0|~c2_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause175)).
cnf(c_0_94,negated_conjecture,(c0_1(X1)|c0_1(X2)|c0_1(X3)|c2_1(X3)|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c1_1(X2)|~c3_1(X2)|~c1_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause174)).
cnf(c_0_95,negated_conjecture,(c1_1(X1)|c2_1(X1)|c3_1(X1)|hskp3|hskp15|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause126)).
cnf(c_0_96,negated_conjecture,(c2_1(a476)|~hskp9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause71)).
cnf(c_0_97,negated_conjecture,(hskp25|hskp5|hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause33)).
cnf(c_0_98,negated_conjecture,(c1_1(X1)|c3_1(X1)|c0_1(X2)|c1_1(X2)|c2_1(X2)|hskp25|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause145)).
cnf(c_0_99,negated_conjecture,(c2_1(X1)|hskp15|hskp12|~c1_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause140)).
cnf(c_0_100,negated_conjecture,(c1_1(a524)|~hskp22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause53)).
cnf(c_0_101,negated_conjecture,(c0_1(a524)|~hskp22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause52)).
cnf(c_0_102,negated_conjecture,(~c2_1(a524)|~hskp22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause87)).
cnf(c_0_103,negated_conjecture,(c3_1(a467)|~hskp4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause78)).
cnf(c_0_104,negated_conjecture,(~c0_1(a467)|~hskp4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause115)).
cnf(c_0_105,negated_conjecture,(~c1_1(a467)|~hskp4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause116)).
cnf(c_0_106,negated_conjecture,(c1_1(X1)|c2_1(X1)|c3_1(X1)|c0_1(X2)|c2_1(X2)|c3_1(X2)|hskp27|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause143)).
cnf(c_0_107,negated_conjecture,(hskp27|hskp22|hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause31)).
cnf(c_0_108,negated_conjecture,(c2_1(X1)|hskp18|hskp20|~c0_1(X1)|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause138)).
cnf(c_0_109,negated_conjecture,(c2_1(X1)|c0_1(X2)|c1_1(X2)|c3_1(X2)|hskp4|~c1_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause149)).
cnf(c_0_110,negated_conjecture,(c1_1(X2)|hskp17|~c0_1(X1)|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause170)).
cnf(c_0_111,negated_conjecture,(c0_1(X1)|hskp27|hskp10|~c1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause136)).
cnf(c_0_112,negated_conjecture,(c1_1(a500)|~hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause60)).
cnf(c_0_113,negated_conjecture,(c2_1(a500)|~hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause61)).
cnf(c_0_114,negated_conjecture,(~c0_1(a500)|~hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause91)).
cnf(c_0_115,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp18|hskp2|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause132)).
cnf(c_0_116,negated_conjecture,(c1_1(a477)|~hskp10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause68)).
cnf(c_0_117,negated_conjecture,(~c1_1(a474)|~hskp7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause110)).
cnf(c_0_118,negated_conjecture,(~c2_1(a474)|~hskp7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause111)).
cnf(c_0_119,negated_conjecture,(c3_1(a474)|~hskp7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause74)).
cnf(c_0_120,negated_conjecture,(c3_1(a470)|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause46)).
cnf(c_0_121,negated_conjecture,(c2_1(a470)|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause45)).
cnf(c_0_122,negated_conjecture,(c1_1(a470)|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause44)).
cnf(c_0_123,negated_conjecture,(c2_1(a490)|~hskp28), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause40)).
cnf(c_0_124,negated_conjecture,(c0_1(a490)|~hskp28), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause38)).
cnf(c_0_125,negated_conjecture,(c0_1(X1)|c0_1(X2)|c1_1(X2)|c3_1(X2)|hskp3|~c1_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause148)).
cnf(c_0_126,negated_conjecture,(c1_1(X1)|c2_1(X1)|c0_1(X2)|hskp9|~c0_1(X1)|~ndr1_0|~c1_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause156)).
cnf(c_0_127,negated_conjecture,(c0_1(X1)|c0_1(X2)|c1_1(X2)|c3_1(X2)|hskp2|~c1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause147)).
cnf(c_0_128,negated_conjecture,(c2_1(X1)|c1_1(X2)|c2_1(X2)|c3_1(X2)|hskp28|~c0_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause150)).
cnf(c_0_129,negated_conjecture,(c2_1(X1)|c3_1(X1)|c0_1(X2)|c2_1(X2)|c3_1(X2)|hskp7|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause146)).
cnf(c_0_130,negated_conjecture,(c1_1(a490)|~hskp28), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause39)).
cnf(c_0_131,negated_conjecture,(~c3_1(a483)|~hskp13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause101)).
cnf(c_0_132,negated_conjecture,(c0_1(a483)|~hskp13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause65)).
cnf(c_0_133,negated_conjecture,(c2_1(a483)|~hskp13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause66)).
cnf(c_0_134,negated_conjecture,(c0_1(X2)|c1_1(X2)|hskp26|~c0_1(X1)|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause160)).
cnf(c_0_135,negated_conjecture,(c1_1(X1)|c1_1(X2)|hskp9|~c0_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause164)).
cnf(c_0_136,negated_conjecture,(c1_1(X1)|c3_1(X1)|c0_1(X2)|hskp13|~c0_1(X1)|~ndr1_0|~c1_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause158)).
cnf(c_0_137,negated_conjecture,(c0_1(X1)|c0_1(X2)|c2_1(X2)|hskp9|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause155)).
fof(c_0_138, plain, ((~epred25_0<=>![X1]:((c0_1(X1)|~c2_1(X1))|~c3_1(X1)))), introduced(definition)).
fof(c_0_139, plain, ((~epred49_0<=>![X3]:(((c0_1(X3)|c2_1(X3))|~ndr1_0)|~c1_1(X3)))), introduced(definition)).
cnf(c_0_140,negated_conjecture,(ndr1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
fof(c_0_141, plain, ((~epred48_0<=>![X1]:(((c3_1(X1)|~ndr1_0)|~c0_1(X1))|~c2_1(X1)))), introduced(definition)).
fof(c_0_142, plain, ((~epred9_0<=>![X1]:((c0_1(X1)|c1_1(X1))|~c2_1(X1)))), introduced(definition)).
cnf(c_0_143,negated_conjecture,(epred25_0|c0_1(X1)|~c3_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_138])).
cnf(c_0_144,negated_conjecture,(epred49_0|c2_1(X1)|c0_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_139]), c_0_140])])).
cnf(c_0_145,negated_conjecture,(c1_1(X1)|hskp3|hskp12|~c3_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_140])])).
cnf(c_0_146,negated_conjecture,(c3_1(X1)|c1_1(X1)|hskp19|hskp20|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_140])])).
cnf(c_0_147,negated_conjecture,(epred48_0|c3_1(X1)|~c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_141]), c_0_140])])).
cnf(c_0_148,negated_conjecture,(c3_1(X1)|c0_1(X1)|hskp0|hskp12|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_140])])).
fof(c_0_149, plain, ((~epred38_0<=>![X2]:((((hskp3|c2_1(X2))|c3_1(X2))|~ndr1_0)|~c0_1(X2)))), introduced(definition)).
fof(c_0_150, plain, ((~epred8_0<=>![X2]:((((hskp0|c0_1(X2))|c1_1(X2))|c2_1(X2))|~ndr1_0))), introduced(definition)).
fof(c_0_151, plain, ((~epred17_0<=>![X1]:((c0_1(X1)|~c1_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_152,negated_conjecture,(epred9_0|c1_1(X1)|c0_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_142])).
cnf(c_0_153,negated_conjecture,(c2_1(X1)|c1_1(X1)|c0_1(X1)|hskp0|hskp1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_140])])).
fof(c_0_154, plain, ((~epred47_0<=>![X1]:(((c2_1(X1)|c3_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
cnf(c_0_155,negated_conjecture,(epred25_0|~c2_1(a466)|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_7]), c_0_8])).
cnf(c_0_156,negated_conjecture,(epred49_0|c2_1(a466)|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_9]), c_0_8])).
cnf(c_0_157,negated_conjecture,(hskp3|hskp12|~c0_1(a503)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_10]), c_0_11])).
cnf(c_0_158,negated_conjecture,(epred25_0|c0_1(a503)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_10]), c_0_12])).
cnf(c_0_159,negated_conjecture,(epred25_0|~hskp19), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_13]), c_0_14]), c_0_15])).
cnf(c_0_160,negated_conjecture,(hskp19|hskp20|~hskp6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_16]), c_0_17]), c_0_18])).
fof(c_0_161, plain, ((~epred6_0<=>![X2]:((c0_1(X2)|c1_1(X2))|c2_1(X2)))), introduced(definition)).
fof(c_0_162, plain, ((~epred3_0<=>![X1]:((((hskp0|c0_1(X1))|c2_1(X1))|c3_1(X1))|~ndr1_0))), introduced(definition)).
fof(c_0_163, plain, ((~epred4_0<=>![X2]:((c0_1(X2)|c1_1(X2))|c3_1(X2)))), introduced(definition)).
cnf(c_0_164,negated_conjecture,(epred48_0|~c0_1(a471)|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_147]), c_0_16])).
cnf(c_0_165,negated_conjecture,(c0_1(a471)|hskp0|hskp12|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_16]), c_0_18])).
cnf(c_0_166,negated_conjecture,(c2_1(X1)|c1_1(X1)|hskp16|hskp17|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_140])])).
fof(c_0_167, plain, ((~epred45_0<=>![X1]:(((c3_1(X1)|~ndr1_0)|~c1_1(X1))|~c2_1(X1)))), introduced(definition)).
fof(c_0_168, plain, ((~epred51_0<=>![X1]:((c2_1(X1)|~c0_1(X1))|~c1_1(X1)))), introduced(definition)).
cnf(c_0_169,negated_conjecture,(epred38_0|c3_1(X1)|c2_1(X1)|hskp3|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_149]), c_0_140])])).
cnf(c_0_170,negated_conjecture,(epred8_0|c2_1(X1)|c1_1(X1)|c0_1(X1)|hskp0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_150]), c_0_140])])).
cnf(c_0_171,negated_conjecture,(epred17_0|c0_1(X1)|~c3_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_151])).
cnf(c_0_172,negated_conjecture,(epred9_0|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_152]), c_0_21]), c_0_22])).
cnf(c_0_173,negated_conjecture,(hskp0|hskp1|~hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_153]), c_0_24]), c_0_25])).
cnf(c_0_174,negated_conjecture,(epred47_0|c3_1(X1)|c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_154]), c_0_140])])).
cnf(c_0_175,negated_conjecture,(epred49_0|epred25_0|~hskp3), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_176,negated_conjecture,(epred25_0|hskp3|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_157, c_0_158])).
cnf(c_0_177,negated_conjecture,(epred25_0|hskp20|~hskp6), inference(spm,[status(thm)],[c_0_159, c_0_160])).
cnf(c_0_178,negated_conjecture,(epred6_0|c2_1(X1)|c1_1(X1)|c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_161])).
cnf(c_0_179,negated_conjecture,(epred3_0|c3_1(X1)|c2_1(X1)|c0_1(X1)|hskp0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_162]), c_0_140])])).
cnf(c_0_180,negated_conjecture,(c3_1(X1)|c1_1(X1)|hskp12|hskp14|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_140])])).
cnf(c_0_181,negated_conjecture,(epred4_0|c3_1(X1)|c1_1(X1)|c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_163])).
fof(c_0_182, plain, ((~epred28_0<=>![X2]:((((hskp8|c0_1(X2))|c2_1(X2))|~ndr1_0)|~c1_1(X2)))), introduced(definition)).
fof(c_0_183, plain, ((~epred29_0<=>![X1]:((c1_1(X1)|~c2_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_184,negated_conjecture,(epred48_0|hskp0|hskp12|~hskp6), inference(spm,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_185,negated_conjecture,(hskp16|hskp17|~hskp24), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_27]), c_0_28]), c_0_29])).
cnf(c_0_186,negated_conjecture,(epred45_0|c3_1(X1)|~c2_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_167]), c_0_140])])).
cnf(c_0_187,negated_conjecture,(epred51_0|c2_1(X1)|~c1_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_168])).
fof(c_0_188, plain, ((~epred11_0<=>![X1]:((c2_1(X1)|~c0_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_189,negated_conjecture,(epred38_0|hskp3|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_190,negated_conjecture,(epred8_0|hskp0|~hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_170]), c_0_24]), c_0_25])).
cnf(c_0_191,negated_conjecture,(epred17_0|~hskp3), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_7]), c_0_9]), c_0_8])).
cnf(c_0_192,negated_conjecture,(epred9_0|hskp0|~hskp12), inference(spm,[status(thm)],[c_0_172, c_0_173])).
cnf(c_0_193,negated_conjecture,(epred47_0|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_174]), c_0_30]), c_0_31])).
cnf(c_0_194,negated_conjecture,(epred49_0|epred25_0|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_175, c_0_176])).
cnf(c_0_195,negated_conjecture,(epred25_0|hskp12|hskp20), inference(spm,[status(thm)],[c_0_177, c_0_33])).
cnf(c_0_196,negated_conjecture,(epred6_0|c2_1(a494)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_178]), c_0_35])).
cnf(c_0_197,negated_conjecture,(epred6_0|~hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_178]), c_0_24]), c_0_25])).
cnf(c_0_198,negated_conjecture,(epred47_0|~c0_1(a488)|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_174]), c_0_37])).
cnf(c_0_199,negated_conjecture,(epred3_0|c0_1(a488)|hskp0|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_179]), c_0_37])).
cnf(c_0_200,negated_conjecture,(hskp12|hskp14|~hskp6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_16]), c_0_17]), c_0_18])).
cnf(c_0_201,negated_conjecture,(epred4_0|c0_1(a488)|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_181]), c_0_38])).
fof(c_0_202, plain, ((~epred39_0<=>![X1]:((~c0_1(X1)|~c1_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_203,negated_conjecture,(epred48_0|~c2_1(a475)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_147]), c_0_40])).
cnf(c_0_204,negated_conjecture,(epred47_0|c2_1(a475)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_174]), c_0_40])).
cnf(c_0_205,negated_conjecture,(epred28_0|c2_1(X1)|c0_1(X1)|hskp8|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_182]), c_0_140])])).
cnf(c_0_206,negated_conjecture,(epred29_0|c1_1(X1)|~c3_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_183])).
cnf(c_0_207,negated_conjecture,(epred48_0|hskp0|hskp12|hskp20), inference(spm,[status(thm)],[c_0_184, c_0_33])).
fof(c_0_208, plain, ((~epred46_0<=>![X2]:((c1_1(X2)|c2_1(X2))|c3_1(X2)))), introduced(definition)).
cnf(c_0_209,negated_conjecture,(epred4_0|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_181]), c_0_35]), c_0_34])).
cnf(c_0_210,negated_conjecture,(hskp15|hskp16|hskp17), inference(spm,[status(thm)],[c_0_185, c_0_42])).
cnf(c_0_211,negated_conjecture,(epred45_0|~c2_1(a475)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_186]), c_0_43])).
cnf(c_0_212,negated_conjecture,(epred51_0|c2_1(a475)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_43]), c_0_40])).
cnf(c_0_213,negated_conjecture,(epred11_0|c2_1(X1)|~c3_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_188])).
cnf(c_0_214,negated_conjecture,(epred38_0|epred8_0|hskp3|~hskp12), inference(spm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_215,negated_conjecture,(epred25_0|epred17_0|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_191, c_0_176])).
cnf(c_0_216,negated_conjecture,(epred38_0|epred9_0|hskp3|~hskp12), inference(spm,[status(thm)],[c_0_189, c_0_192])).
cnf(c_0_217,negated_conjecture,(epred47_0|epred8_0|~hskp12), inference(spm,[status(thm)],[c_0_193, c_0_190])).
cnf(c_0_218,negated_conjecture,(epred49_0|epred25_0|hskp12), inference(spm,[status(thm)],[c_0_194, c_0_195])).
cnf(c_0_219,negated_conjecture,(c3_1(X1)|c1_1(X1)|hskp16|hskp27|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_140])])).
cnf(c_0_220,negated_conjecture,(epred6_0|c0_1(a488)|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_178]), c_0_37])).
cnf(c_0_221,negated_conjecture,(epred6_0|hskp14|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_196]), c_0_197]), c_0_34]), c_0_41])).
cnf(c_0_222,negated_conjecture,(epred47_0|epred3_0|hskp0|~hskp14), inference(spm,[status(thm)],[c_0_198, c_0_199])).
cnf(c_0_223,negated_conjecture,(hskp12|hskp14|hskp20), inference(spm,[status(thm)],[c_0_200, c_0_33])).
cnf(c_0_224,negated_conjecture,(epred47_0|epred4_0|~hskp14), inference(spm,[status(thm)],[c_0_198, c_0_201])).
cnf(c_0_225,negated_conjecture,(c0_1(a488)|hskp0|hskp1|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_153]), c_0_37])).
cnf(c_0_226,negated_conjecture,(epred39_0|~c3_1(X1)|~c1_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_202])).
cnf(c_0_227,negated_conjecture,(epred48_0|epred47_0|~hskp8), inference(spm,[status(thm)],[c_0_203, c_0_204])).
cnf(c_0_228,negated_conjecture,(epred28_0|hskp8|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_229,negated_conjecture,(epred9_0|c0_1(a471)|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_152]), c_0_16])).
cnf(c_0_230,negated_conjecture,(epred29_0|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_10]), c_0_12]), c_0_11])).
cnf(c_0_231,negated_conjecture,(epred48_0|epred47_0|hskp12|hskp20), inference(spm,[status(thm)],[c_0_193, c_0_207])).
cnf(c_0_232,negated_conjecture,(epred46_0|c3_1(X1)|c2_1(X1)|c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_208])).
cnf(c_0_233,negated_conjecture,(epred51_0|~c0_1(a492)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_48]), c_0_49])).
cnf(c_0_234,negated_conjecture,(epred49_0|c0_1(a492)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_48]), c_0_49])).
cnf(c_0_235,negated_conjecture,(epred49_0|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_236,negated_conjecture,(epred4_0|hskp15|hskp16), inference(spm,[status(thm)],[c_0_209, c_0_210])).
cnf(c_0_237,negated_conjecture,(epred51_0|epred45_0|~hskp8), inference(spm,[status(thm)],[c_0_211, c_0_212])).
cnf(c_0_238,negated_conjecture,(epred11_0|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_239,negated_conjecture,(epred38_0|epred17_0|epred8_0|~hskp12), inference(spm,[status(thm)],[c_0_191, c_0_214])).
cnf(c_0_240,negated_conjecture,(epred25_0|epred17_0|hskp12), inference(spm,[status(thm)],[c_0_215, c_0_195])).
cnf(c_0_241,negated_conjecture,(epred38_0|epred17_0|epred9_0|~hskp12), inference(spm,[status(thm)],[c_0_191, c_0_216])).
cnf(c_0_242,negated_conjecture,(~epred9_0|~epred8_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_53, c_0_150]), c_0_142])).
cnf(c_0_243,negated_conjecture,(epred49_0|epred47_0|epred25_0|epred8_0), inference(spm,[status(thm)],[c_0_217, c_0_218])).
cnf(c_0_244,negated_conjecture,(epred6_0|hskp16|hskp27|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220]), c_0_38]), c_0_36])).
cnf(c_0_245,negated_conjecture,(epred6_0|hskp14|hskp15|hskp16), inference(spm,[status(thm)],[c_0_221, c_0_210])).
cnf(c_0_246,negated_conjecture,(epred47_0|epred3_0|hskp12|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_223]), c_0_193])).
cnf(c_0_247,negated_conjecture,(epred47_0|epred4_0|hskp12|hskp20), inference(spm,[status(thm)],[c_0_224, c_0_223])).
cnf(c_0_248,negated_conjecture,(epred47_0|hskp1|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_225]), c_0_193])).
cnf(c_0_249,negated_conjecture,(epred39_0|~c0_1(a492)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_54]), c_0_48])).
cnf(c_0_250,negated_conjecture,(epred28_0|c0_1(a492)|hskp8|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_48]), c_0_49])).
cnf(c_0_251,negated_conjecture,(epred48_0|epred47_0|epred28_0|~hskp16), inference(spm,[status(thm)],[c_0_227, c_0_228])).
cnf(c_0_252,negated_conjecture,(epred48_0|epred47_0|hskp15|hskp16), inference(spm,[status(thm)],[c_0_227, c_0_55])).
cnf(c_0_253,negated_conjecture,(epred48_0|epred9_0|~hskp6), inference(spm,[status(thm)],[c_0_164, c_0_229])).
cnf(c_0_254,negated_conjecture,(epred48_0|epred47_0|epred29_0|hskp12), inference(spm,[status(thm)],[c_0_230, c_0_231])).
cnf(c_0_255,negated_conjecture,(epred46_0|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_232]), c_0_38]), c_0_37])).
cnf(c_0_256,negated_conjecture,(epred51_0|epred49_0|~hskp15), inference(spm,[status(thm)],[c_0_233, c_0_234])).
cnf(c_0_257,negated_conjecture,(epred49_0|epred4_0|hskp15), inference(spm,[status(thm)],[c_0_235, c_0_236])).
cnf(c_0_258,negated_conjecture,(epred51_0|epred45_0|hskp15|hskp16), inference(spm,[status(thm)],[c_0_237, c_0_55])).
cnf(c_0_259,negated_conjecture,(hskp0|hskp12|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_56]), c_0_57]), c_0_58])).
cnf(c_0_260,negated_conjecture,(epred11_0|hskp11), inference(spm,[status(thm)],[c_0_238, c_0_1])).
fof(c_0_261, plain, ((~epred50_0<=>![X2]:(((c1_1(X2)|c3_1(X2))|~ndr1_0)|~c2_1(X2)))), introduced(definition)).
cnf(c_0_262,negated_conjecture,(c1_1(X1)|c0_1(X1)|hskp2|hskp6|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_140])])).
cnf(c_0_263,negated_conjecture,(~epred47_0|~epred17_0|~epred6_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_60, c_0_154]), c_0_161]), c_0_151])).
cnf(c_0_264,negated_conjecture,(epred49_0|epred25_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_218])).
cnf(c_0_265,negated_conjecture,(epred38_0|epred25_0|epred17_0|epred8_0), inference(spm,[status(thm)],[c_0_239, c_0_240])).
cnf(c_0_266,negated_conjecture,(epred38_0|epred25_0|epred17_0|epred9_0), inference(spm,[status(thm)],[c_0_241, c_0_240])).
cnf(c_0_267,negated_conjecture,(epred47_0|epred9_0|~hskp12), inference(spm,[status(thm)],[c_0_193, c_0_192])).
cnf(c_0_268,negated_conjecture,(epred49_0|epred47_0|epred25_0|~epred9_0), inference(spm,[status(thm)],[c_0_242, c_0_243])).
cnf(c_0_269,negated_conjecture,(epred6_0|hskp15|hskp16|hskp27), inference(spm,[status(thm)],[c_0_244, c_0_245])).
cnf(c_0_270,negated_conjecture,(epred47_0|epred29_0|epred3_0|hskp12), inference(spm,[status(thm)],[c_0_230, c_0_246])).
cnf(c_0_271,negated_conjecture,(epred47_0|epred29_0|epred4_0|hskp12), inference(spm,[status(thm)],[c_0_230, c_0_247])).
cnf(c_0_272,negated_conjecture,(epred47_0|epred9_0|~hskp14), inference(spm,[status(thm)],[c_0_172, c_0_248])).
cnf(c_0_273,negated_conjecture,(epred39_0|epred28_0|hskp8|~hskp15), inference(spm,[status(thm)],[c_0_249, c_0_250])).
cnf(c_0_274,negated_conjecture,(epred48_0|epred47_0|epred28_0|hskp15), inference(spm,[status(thm)],[c_0_251, c_0_252])).
cnf(c_0_275,negated_conjecture,(epred48_0|epred9_0|hskp12|hskp20), inference(spm,[status(thm)],[c_0_253, c_0_33])).
cnf(c_0_276,negated_conjecture,(epred48_0|epred47_0|epred29_0|epred8_0), inference(spm,[status(thm)],[c_0_217, c_0_254])).
cnf(c_0_277,negated_conjecture,(epred46_0|epred6_0|hskp15|hskp16), inference(spm,[status(thm)],[c_0_255, c_0_245])).
cnf(c_0_278,negated_conjecture,(~epred46_0|~epred45_0|~epred4_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_61, c_0_167]), c_0_208]), c_0_163])).
cnf(c_0_279,negated_conjecture,(epred51_0|epred49_0|epred4_0), inference(spm,[status(thm)],[c_0_256, c_0_257])).
cnf(c_0_280,negated_conjecture,(epred51_0|epred49_0|epred45_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_258]), c_0_256])).
cnf(c_0_281,negated_conjecture,(epred51_0|epred48_0|~hskp8), inference(spm,[status(thm)],[c_0_203, c_0_212])).
cnf(c_0_282,negated_conjecture,(epred11_0|hskp0|hskp12), inference(spm,[status(thm)],[c_0_259, c_0_260])).
fof(c_0_283, plain, ((~epred32_0<=>![X1]:((~c0_1(X1)|~c2_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_284,negated_conjecture,(epred50_0|c3_1(X1)|c1_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_261]), c_0_140])])).
cnf(c_0_285,negated_conjecture,(hskp2|hskp6|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_21]), c_0_22]), c_0_20])).
cnf(c_0_286,negated_conjecture,(epred49_0|epred25_0|~epred47_0|~epred17_0), inference(spm,[status(thm)],[c_0_263, c_0_264])).
cnf(c_0_287,negated_conjecture,(epred38_0|epred25_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_265]), c_0_266])).
cnf(c_0_288,negated_conjecture,(epred49_0|epred47_0|epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_218]), c_0_268])).
cnf(c_0_289,negated_conjecture,(epred39_0|~hskp27), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_62]), c_0_63]), c_0_64])).
cnf(c_0_290,negated_conjecture,(epred49_0|epred6_0|hskp15|hskp27), inference(spm,[status(thm)],[c_0_235, c_0_269])).
cnf(c_0_291,negated_conjecture,(epred49_0|epred39_0|~hskp15), inference(spm,[status(thm)],[c_0_249, c_0_234])).
cnf(c_0_292,negated_conjecture,(~epred4_0|~epred3_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_65, c_0_162]), c_0_163])).
cnf(c_0_293,negated_conjecture,(epred47_0|epred29_0|epred8_0|epred3_0), inference(spm,[status(thm)],[c_0_217, c_0_270])).
cnf(c_0_294,negated_conjecture,(epred47_0|epred29_0|epred8_0|epred4_0), inference(spm,[status(thm)],[c_0_217, c_0_271])).
cnf(c_0_295,negated_conjecture,(epred47_0|epred9_0|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_223]), c_0_267])).
cnf(c_0_296,negated_conjecture,(~epred29_0|~epred28_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_66, c_0_182]), c_0_183])).
cnf(c_0_297,negated_conjecture,(epred48_0|epred47_0|epred39_0|epred28_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_273]), c_0_274])).
cnf(c_0_298,negated_conjecture,(epred48_0|epred29_0|epred9_0|hskp12), inference(spm,[status(thm)],[c_0_230, c_0_275])).
cnf(c_0_299,negated_conjecture,(epred48_0|epred47_0|epred29_0|~epred9_0), inference(spm,[status(thm)],[c_0_242, c_0_276])).
cnf(c_0_300,negated_conjecture,(epred49_0|epred46_0|epred6_0|hskp15), inference(spm,[status(thm)],[c_0_235, c_0_277])).
cnf(c_0_301,negated_conjecture,(epred51_0|epred49_0|~epred46_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_279]), c_0_280])).
cnf(c_0_302,negated_conjecture,(epred51_0|epred48_0|hskp15|hskp16), inference(spm,[status(thm)],[c_0_281, c_0_55])).
fof(c_0_303, plain, ((~epred12_0<=>![X1]:((((hskp12|c2_1(X1))|c3_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
fof(c_0_304, plain, ((~epred13_0<=>![X2]:((c1_1(X2)|c3_1(X2))|~c0_1(X2)))), introduced(definition)).
cnf(c_0_305,negated_conjecture,(epred47_0|epred11_0|hskp12), inference(spm,[status(thm)],[c_0_193, c_0_282])).
cnf(c_0_306,negated_conjecture,(c0_1(a503)|hskp2|hskp6|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_12]), c_0_11])).
cnf(c_0_307,negated_conjecture,(hskp3|hskp12|~hskp21), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_67]), c_0_68]), c_0_69])).
cnf(c_0_308,negated_conjecture,(epred48_0|~c0_1(a477)|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_147]), c_0_71])).
cnf(c_0_309,negated_conjecture,(c0_1(a477)|hskp0|hskp12|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_71]), c_0_70])).
cnf(c_0_310,negated_conjecture,(epred32_0|~c3_1(X1)|~c2_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_283])).
cnf(c_0_311,negated_conjecture,(epred50_0|~hskp6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_284]), c_0_16]), c_0_17])).
cnf(c_0_312,negated_conjecture,(hskp0|hskp2|hskp6|~hskp12), inference(spm,[status(thm)],[c_0_285, c_0_173])).
cnf(c_0_313,negated_conjecture,(~epred51_0|~epred50_0|~epred25_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_72, c_0_261]), c_0_168]), c_0_138])).
cnf(c_0_314,negated_conjecture,(epred49_0|epred38_0|epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287]), c_0_288])).
cnf(c_0_315,negated_conjecture,(~epred48_0|~epred29_0|~epred6_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_73, c_0_141]), c_0_161]), c_0_183])).
cnf(c_0_316,negated_conjecture,(epred49_0|epred39_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_290]), c_0_291])).
cnf(c_0_317,negated_conjecture,(epred47_0|epred29_0|epred8_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293]), c_0_294])).
cnf(c_0_318,negated_conjecture,(epred47_0|epred29_0|epred9_0), inference(spm,[status(thm)],[c_0_230, c_0_295])).
cnf(c_0_319,negated_conjecture,(epred48_0|epred47_0|epred39_0|~epred29_0), inference(spm,[status(thm)],[c_0_296, c_0_297])).
cnf(c_0_320,negated_conjecture,(epred48_0|epred47_0|epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_298]), c_0_299])).
cnf(c_0_321,negated_conjecture,(epred51_0|epred49_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_300]), c_0_301])).
cnf(c_0_322,negated_conjecture,(epred51_0|epred49_0|epred48_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_302]), c_0_256])).
cnf(c_0_323,negated_conjecture,(epred12_0|c3_1(X1)|c2_1(X1)|hskp12|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_303]), c_0_140])])).
cnf(c_0_324,negated_conjecture,(epred13_0|c3_1(X1)|c1_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_304])).
fof(c_0_325, plain, ((~epred52_0<=>![X3]:(((c1_1(X3)|c2_1(X3))|c3_1(X3))|~ndr1_0))), introduced(definition)).
fof(c_0_326, plain, ((~epred22_0<=>![X2]:((((hskp21|c2_1(X2))|c3_1(X2))|~ndr1_0)|~c1_1(X2)))), introduced(definition)).
cnf(c_0_327,negated_conjecture,(epred47_0|epred11_0|epred8_0), inference(spm,[status(thm)],[c_0_217, c_0_305])).
cnf(c_0_328,negated_conjecture,(hskp2|hskp3|hskp6|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_306]), c_0_33])).
cnf(c_0_329,negated_conjecture,(hskp3|hskp6|hskp10|hskp12), inference(spm,[status(thm)],[c_0_307, c_0_74])).
cnf(c_0_330,negated_conjecture,(epred48_0|hskp0|hskp12|~hskp10), inference(spm,[status(thm)],[c_0_308, c_0_309])).
cnf(c_0_331,negated_conjecture,(epred32_0|~c0_1(a503)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_10]), c_0_12])).
fof(c_0_332, plain, ((~epred42_0<=>![X2]:((((hskp14|c0_1(X2))|~ndr1_0)|~c1_1(X2))|~c3_1(X2)))), introduced(definition)).
cnf(c_0_333,negated_conjecture,(epred3_0|hskp0|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_179]), c_0_76]), c_0_77])).
cnf(c_0_334,negated_conjecture,(epred50_0|hskp0|hskp2|~hskp12), inference(spm,[status(thm)],[c_0_311, c_0_312])).
cnf(c_0_335,negated_conjecture,(~epred39_0|~epred38_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_78, c_0_149]), c_0_202])).
cnf(c_0_336,negated_conjecture,(epred49_0|epred38_0|~epred51_0|~epred50_0), inference(spm,[status(thm)],[c_0_313, c_0_314])).
cnf(c_0_337,negated_conjecture,(epred49_0|epred39_0|~epred48_0|~epred29_0), inference(spm,[status(thm)],[c_0_315, c_0_316])).
cnf(c_0_338,negated_conjecture,(epred47_0|epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_317]), c_0_318])).
cnf(c_0_339,negated_conjecture,(epred48_0|epred47_0|epred39_0), inference(spm,[status(thm)],[c_0_319, c_0_320])).
cnf(c_0_340,negated_conjecture,(epred51_0|epred49_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_321]), c_0_322])).
cnf(c_0_341,negated_conjecture,(epred12_0|epred6_0|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_220]), c_0_197]), c_0_37]), c_0_36])).
cnf(c_0_342,negated_conjecture,(epred13_0|~c0_1(a488)|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_324]), c_0_38])).
cnf(c_0_343,negated_conjecture,(epred17_0|~c1_1(a502)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_13]), c_0_15])).
cnf(c_0_344,negated_conjecture,(epred29_0|c1_1(a502)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_13]), c_0_14])).
cnf(c_0_345,negated_conjecture,(epred52_0|c3_1(X1)|c2_1(X1)|c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_325]), c_0_140])])).
cnf(c_0_346,negated_conjecture,(epred22_0|c3_1(X1)|c2_1(X1)|hskp21|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_326]), c_0_140])])).
cnf(c_0_347,negated_conjecture,(epred4_0|c1_1(a465)|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_181]), c_0_76])).
cnf(c_0_348,negated_conjecture,(epred47_0|epred11_0|~epred9_0), inference(spm,[status(thm)],[c_0_242, c_0_327])).
cnf(c_0_349,negated_conjecture,(epred47_0|epred6_0|~hskp14), inference(spm,[status(thm)],[c_0_198, c_0_220])).
cnf(c_0_350,negated_conjecture,(hskp2|hskp3|hskp12|hskp14), inference(spm,[status(thm)],[c_0_200, c_0_328])).
cnf(c_0_351,negated_conjecture,(epred48_0|hskp0|hskp3|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_329]), c_0_330])).
cnf(c_0_352,negated_conjecture,(epred32_0|hskp2|hskp6|~hskp20), inference(spm,[status(thm)],[c_0_331, c_0_306])).
cnf(c_0_353,negated_conjecture,(epred42_0|c0_1(X1)|hskp14|~c3_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_332]), c_0_140])])).
cnf(c_0_354,negated_conjecture,(epred47_0|epred3_0|~hskp2), inference(spm,[status(thm)],[c_0_193, c_0_333])).
cnf(c_0_355,negated_conjecture,(epred50_0|epred47_0|hskp2|~hskp12), inference(spm,[status(thm)],[c_0_193, c_0_334])).
cnf(c_0_356,negated_conjecture,(epred49_0|~epred51_0|~epred50_0|~epred39_0), inference(spm,[status(thm)],[c_0_335, c_0_336])).
cnf(c_0_357,negated_conjecture,(epred49_0|epred47_0|epred39_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_338]), c_0_339])).
cnf(c_0_358,negated_conjecture,(epred51_0|epred49_0|epred47_0), inference(spm,[status(thm)],[c_0_340, c_0_338])).
cnf(c_0_359,negated_conjecture,(epred12_0|epred6_0|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_223]), c_0_197])).
cnf(c_0_360,negated_conjecture,(epred13_0|epred6_0|~hskp14), inference(spm,[status(thm)],[c_0_342, c_0_220])).
cnf(c_0_361,negated_conjecture,(epred29_0|epred17_0|~hskp19), inference(spm,[status(thm)],[c_0_343, c_0_344])).
cnf(c_0_362,negated_conjecture,(epred52_0|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_345]), c_0_38]), c_0_37])).
cnf(c_0_363,negated_conjecture,(epred22_0|epred4_0|hskp21|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_347]), c_0_77]), c_0_75])).
cnf(c_0_364,negated_conjecture,(~epred22_0|~epred11_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_79, c_0_326]), c_0_188])).
cnf(c_0_365,negated_conjecture,(epred47_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_305]), c_0_348])).
cnf(c_0_366,negated_conjecture,(epred47_0|epred6_0|hskp2|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_350]), c_0_197])).
cnf(c_0_367,negated_conjecture,(epred9_0|c0_1(a503)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_152]), c_0_12])).
cnf(c_0_368,negated_conjecture,(epred48_0|epred47_0|hskp3|hskp12), inference(spm,[status(thm)],[c_0_193, c_0_351])).
cnf(c_0_369,negated_conjecture,(epred50_0|epred32_0|hskp2|~hskp20), inference(spm,[status(thm)],[c_0_311, c_0_352])).
cnf(c_0_370,negated_conjecture,(epred42_0|hskp14|~hskp3), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_7]), c_0_9]), c_0_8])).
cnf(c_0_371,negated_conjecture,(epred50_0|hskp2|hskp3|hskp12), inference(spm,[status(thm)],[c_0_311, c_0_328])).
cnf(c_0_372,negated_conjecture,(epred50_0|epred47_0|epred3_0|~hskp12), inference(spm,[status(thm)],[c_0_354, c_0_355])).
cnf(c_0_373,negated_conjecture,(epred49_0|epred4_0|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_347]), c_0_76]), c_0_77])).
cnf(c_0_374,negated_conjecture,(epred49_0|epred47_0|~epred50_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_357]), c_0_358])).
cnf(c_0_375,negated_conjecture,(epred29_0|epred12_0|epred6_0), inference(spm,[status(thm)],[c_0_230, c_0_359])).
cnf(c_0_376,negated_conjecture,(epred13_0|epred6_0|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_360, c_0_223]), c_0_197])).
cnf(c_0_377,negated_conjecture,(epred29_0|epred17_0|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_160]), c_0_230])).
cnf(c_0_378,negated_conjecture,(epred52_0|hskp12|hskp20), inference(spm,[status(thm)],[c_0_362, c_0_223])).
cnf(c_0_379,negated_conjecture,(epred22_0|epred4_0|hskp3|hskp12|~hskp2), inference(spm,[status(thm)],[c_0_307, c_0_363])).
cnf(c_0_380,negated_conjecture,(epred47_0|epred4_0|hskp2|hskp3|hskp12), inference(spm,[status(thm)],[c_0_224, c_0_350])).
cnf(c_0_381,negated_conjecture,(epred47_0|~epred22_0), inference(spm,[status(thm)],[c_0_364, c_0_365])).
cnf(c_0_382,negated_conjecture,(epred47_0|epred6_0|epred3_0|hskp3), inference(spm,[status(thm)],[c_0_354, c_0_366])).
cnf(c_0_383,negated_conjecture,(epred9_0|hskp3|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_157, c_0_367])).
cnf(c_0_384,negated_conjecture,(epred48_0|epred47_0|epred17_0|hskp12), inference(spm,[status(thm)],[c_0_191, c_0_368])).
cnf(c_0_385,negated_conjecture,(epred50_0|epred47_0|epred32_0|epred3_0|~hskp20), inference(spm,[status(thm)],[c_0_354, c_0_369])).
cnf(c_0_386,negated_conjecture,(epred50_0|hskp12|hskp20), inference(spm,[status(thm)],[c_0_311, c_0_33])).
cnf(c_0_387,negated_conjecture,(epred32_0|hskp2|hskp12|hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_352]), c_0_223])).
cnf(c_0_388,negated_conjecture,(epred47_0|epred42_0|epred3_0|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_370]), c_0_193])).
cnf(c_0_389,negated_conjecture,(epred50_0|epred47_0|epred3_0|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_371]), c_0_372])).
cnf(c_0_390,negated_conjecture,(epred49_0|epred47_0|epred4_0|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_355]), c_0_374])).
cnf(c_0_391,negated_conjecture,(~epred13_0|~epred12_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_80, c_0_303]), c_0_304])).
cnf(c_0_392,negated_conjecture,(epred29_0|epred12_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_375]), c_0_338])).
cnf(c_0_393,negated_conjecture,(epred29_0|epred13_0|epred6_0), inference(spm,[status(thm)],[c_0_230, c_0_376])).
cnf(c_0_394,negated_conjecture,(epred29_0|epred17_0|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_33]), c_0_230])).
fof(c_0_395, plain, ((~epred5_0<=>![X1]:((((hskp25|c1_1(X1))|c3_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
cnf(c_0_396,negated_conjecture,(epred29_0|epred25_0|hskp12), inference(spm,[status(thm)],[c_0_230, c_0_195])).
cnf(c_0_397,negated_conjecture,(epred52_0|epred29_0|hskp12), inference(spm,[status(thm)],[c_0_230, c_0_378])).
cnf(c_0_398,negated_conjecture,(c2_1(X1)|c1_1(X1)|hskp7|hskp12|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_140])])).
cnf(c_0_399,negated_conjecture,(epred47_0|epred4_0|hskp3|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_379, c_0_380]), c_0_381])).
cnf(c_0_400,negated_conjecture,(epred47_0|epred17_0|epred6_0|epred3_0), inference(spm,[status(thm)],[c_0_191, c_0_382])).
cnf(c_0_401,negated_conjecture,(epred17_0|epred9_0|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_191, c_0_383])).
cnf(c_0_402,negated_conjecture,(epred48_0|epred47_0|epred17_0|epred8_0), inference(spm,[status(thm)],[c_0_217, c_0_384])).
cnf(c_0_403,negated_conjecture,(epred50_0|epred47_0|epred32_0|epred3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_385, c_0_386]), c_0_372])).
cnf(c_0_404,negated_conjecture,(epred47_0|epred32_0|epred4_0|hskp2|hskp12), inference(spm,[status(thm)],[c_0_224, c_0_387])).
cnf(c_0_405,negated_conjecture,(epred50_0|epred47_0|epred42_0|epred3_0), inference(spm,[status(thm)],[c_0_388, c_0_389])).
cnf(c_0_406,negated_conjecture,(epred47_0|epred42_0|epred4_0|~hskp3), inference(spm,[status(thm)],[c_0_224, c_0_370])).
cnf(c_0_407,negated_conjecture,(epred49_0|epred47_0|epred4_0|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_380]), c_0_390])).
cnf(c_0_408,negated_conjecture,(epred32_0|epred25_0|~hskp20), inference(spm,[status(thm)],[c_0_331, c_0_158])).
cnf(c_0_409,negated_conjecture,(epred29_0|~epred17_0|~epred13_0), inference(spm,[status(thm)],[c_0_391, c_0_392])).
cnf(c_0_410,negated_conjecture,(epred29_0|epred13_0|~epred47_0|~epred17_0), inference(spm,[status(thm)],[c_0_263, c_0_393])).
cnf(c_0_411,negated_conjecture,(epred38_0|epred29_0|epred17_0|epred8_0), inference(spm,[status(thm)],[c_0_239, c_0_394])).
cnf(c_0_412,negated_conjecture,(epred38_0|epred29_0|epred17_0|epred9_0), inference(spm,[status(thm)],[c_0_241, c_0_394])).
cnf(c_0_413,negated_conjecture,(epred5_0|c3_1(X1)|c1_1(X1)|hskp25|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_395]), c_0_140])])).
cnf(c_0_414,negated_conjecture,(epred29_0|epred25_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_396])).
cnf(c_0_415,negated_conjecture,(epred52_0|epred29_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_397])).
cnf(c_0_416,negated_conjecture,(epred29_0|~c2_1(a512)|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_67]), c_0_69])).
cnf(c_0_417,negated_conjecture,(c2_1(a512)|hskp7|hskp12|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_67]), c_0_69])).
cnf(c_0_418,negated_conjecture,(epred47_0|epred17_0|epred4_0|hskp12), inference(spm,[status(thm)],[c_0_191, c_0_399])).
cnf(c_0_419,negated_conjecture,(epred47_0|epred17_0|epred6_0|~epred4_0), inference(spm,[status(thm)],[c_0_292, c_0_400])).
cnf(c_0_420,negated_conjecture,(epred48_0|epred17_0|epred9_0|hskp12), inference(spm,[status(thm)],[c_0_401, c_0_275])).
cnf(c_0_421,negated_conjecture,(epred48_0|epred47_0|epred17_0|~epred9_0), inference(spm,[status(thm)],[c_0_242, c_0_402])).
cnf(c_0_422,negated_conjecture,(epred50_0|epred47_0|epred32_0|~epred4_0), inference(spm,[status(thm)],[c_0_292, c_0_403])).
cnf(c_0_423,negated_conjecture,(epred49_0|epred47_0|epred32_0|epred4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_404]), c_0_390])).
cnf(c_0_424,negated_conjecture,(epred50_0|epred47_0|epred42_0|~epred4_0), inference(spm,[status(thm)],[c_0_292, c_0_405])).
cnf(c_0_425,negated_conjecture,(epred49_0|epred47_0|epred42_0|epred4_0), inference(spm,[status(thm)],[c_0_406, c_0_407])).
cnf(c_0_426,negated_conjecture,(epred32_0|epred25_0|hskp12), inference(spm,[status(thm)],[c_0_408, c_0_195])).
cnf(c_0_427,negated_conjecture,(epred29_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_410]), c_0_338])).
cnf(c_0_428,negated_conjecture,(epred38_0|epred29_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_411]), c_0_412])).
cnf(c_0_429,negated_conjecture,(epred39_0|~c1_1(a461)|~hskp25), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_82]), c_0_83])).
cnf(c_0_430,negated_conjecture,(epred29_0|c1_1(a461)|~hskp25), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_82]), c_0_84])).
cnf(c_0_431,negated_conjecture,(epred5_0|hskp25|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_85]), c_0_86]), c_0_87])).
cnf(c_0_432,negated_conjecture,(epred29_0|epred25_0|~epred47_0|~epred17_0), inference(spm,[status(thm)],[c_0_263, c_0_414])).
cnf(c_0_433,negated_conjecture,(epred52_0|epred29_0|~epred47_0|~epred17_0), inference(spm,[status(thm)],[c_0_263, c_0_415])).
cnf(c_0_434,negated_conjecture,(epred29_0|hskp7|hskp12|~hskp21), inference(spm,[status(thm)],[c_0_416, c_0_417])).
cnf(c_0_435,negated_conjecture,(epred47_0|epred17_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_418]), c_0_419])).
cnf(c_0_436,negated_conjecture,(epred48_0|epred47_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_420]), c_0_421])).
cnf(c_0_437,negated_conjecture,(~epred42_0|~epred32_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_88, c_0_332]), c_0_283])).
cnf(c_0_438,negated_conjecture,(epred49_0|epred47_0|epred32_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_423]), c_0_374])).
cnf(c_0_439,negated_conjecture,(epred49_0|epred47_0|epred42_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_424, c_0_425]), c_0_374])).
cnf(c_0_440,negated_conjecture,(epred47_0|epred32_0|epred25_0|epred8_0), inference(spm,[status(thm)],[c_0_217, c_0_426])).
cnf(c_0_441,negated_conjecture,(c2_1(X1)|hskp2|hskp15|~c1_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_140])])).
cnf(c_0_442,negated_conjecture,(c1_1(a461)|hskp3|hskp12|~hskp25), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_82]), c_0_83])).
cnf(c_0_443,negated_conjecture,(epred39_0|~c1_1(a468)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_50]), c_0_51])).
cnf(c_0_444,negated_conjecture,(c1_1(a468)|hskp3|hskp12|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_50]), c_0_51])).
cnf(c_0_445,negated_conjecture,(epred49_0|epred25_0|~epred48_0|~epred29_0), inference(spm,[status(thm)],[c_0_315, c_0_264])).
cnf(c_0_446,negated_conjecture,(epred38_0|epred29_0), inference(spm,[status(thm)],[c_0_427, c_0_428])).
fof(c_0_447, plain, ((~epred1_0<=>![X1]:((((hskp27|c1_1(X1))|c2_1(X1))|c3_1(X1))|~ndr1_0))), introduced(definition)).
cnf(c_0_448,negated_conjecture,(epred5_0|c3_1(a476)|hskp25|~hskp9), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_90]), c_0_91])).
cnf(c_0_449,negated_conjecture,(epred39_0|epred29_0|~hskp25), inference(spm,[status(thm)],[c_0_429, c_0_430])).
cnf(c_0_450,negated_conjecture,(epred5_0|hskp9|hskp20|hskp25), inference(spm,[status(thm)],[c_0_431, c_0_92])).
cnf(c_0_451,negated_conjecture,(~epred52_0|~epred39_0|~epred25_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_93, c_0_325]), c_0_138]), c_0_202])).
cnf(c_0_452,negated_conjecture,(epred38_0|epred29_0|epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_432, c_0_428]), c_0_338])).
cnf(c_0_453,negated_conjecture,(epred52_0|epred38_0|epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_428]), c_0_338])).
cnf(c_0_454,negated_conjecture,(epred29_0|hskp6|hskp7|hskp10|hskp12), inference(spm,[status(thm)],[c_0_434, c_0_74])).
cnf(c_0_455,negated_conjecture,(epred47_0|epred42_0|epred6_0|~hskp3), inference(spm,[status(thm)],[c_0_349, c_0_370])).
cnf(c_0_456,negated_conjecture,(~epred49_0|~epred25_0|~epred17_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_94, c_0_139]), c_0_138]), c_0_151])).
cnf(c_0_457,negated_conjecture,(epred47_0|epred17_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_435]), c_0_338]), c_0_436])).
cnf(c_0_458,negated_conjecture,(epred49_0|epred47_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_437, c_0_438]), c_0_439])).
cnf(c_0_459,negated_conjecture,(epred47_0|epred32_0|epred25_0|~epred9_0), inference(spm,[status(thm)],[c_0_242, c_0_440])).
cnf(c_0_460,negated_conjecture,(c2_1(a475)|hskp2|hskp15|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_441, c_0_43]), c_0_40])).
cnf(c_0_461,negated_conjecture,(c3_1(X1)|c2_1(X1)|c1_1(X1)|hskp3|hskp15), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_140])])).
cnf(c_0_462,negated_conjecture,(epred39_0|hskp3|hskp12|~hskp25), inference(spm,[status(thm)],[c_0_429, c_0_442])).
cnf(c_0_463,negated_conjecture,(epred39_0|hskp3|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_443, c_0_444])).
cnf(c_0_464,negated_conjecture,(epred38_0|epred3_0|hskp3|~hskp2), inference(spm,[status(thm)],[c_0_189, c_0_333])).
cnf(c_0_465,negated_conjecture,(epred50_0|epred38_0|hskp2|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_334]), c_0_371])).
cnf(c_0_466,negated_conjecture,(epred49_0|~epred51_0|~epred50_0|~epred48_0|~epred29_0), inference(spm,[status(thm)],[c_0_313, c_0_445])).
cnf(c_0_467,negated_conjecture,(epred51_0|epred49_0|epred38_0), inference(spm,[status(thm)],[c_0_340, c_0_446])).
cnf(c_0_468,negated_conjecture,(epred49_0|epred6_0|c2_1(X1)|c0_1(X1)), inference(spm,[status(thm)],[c_0_144, c_0_178])).
fof(c_0_469, plain, ((~epred2_0<=>![X2]:((c0_1(X2)|c2_1(X2))|c3_1(X2)))), introduced(definition)).
cnf(c_0_470,negated_conjecture,(epred1_0|c3_1(X1)|c2_1(X1)|c1_1(X1)|hskp27), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_447]), c_0_140])])).
cnf(c_0_471,negated_conjecture,(epred28_0|c2_1(a466)|hskp8|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_9]), c_0_8])).
cnf(c_0_472,negated_conjecture,(epred29_0|epred5_0|hskp25|~hskp9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_448]), c_0_96]), c_0_91])).
cnf(c_0_473,negated_conjecture,(epred39_0|epred29_0|epred5_0|hskp9), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_450]), c_0_230])).
cnf(c_0_474,negated_conjecture,(epred29_0|~epred39_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_452]), c_0_453]), c_0_335])).
cnf(c_0_475,negated_conjecture,(epred29_0|hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_454]), c_0_394]), c_0_394]), c_0_427])).
cnf(c_0_476,negated_conjecture,(epred47_0|epred42_0|epred6_0|epred3_0), inference(spm,[status(thm)],[c_0_455, c_0_382])).
cnf(c_0_477,negated_conjecture,(epred47_0|~epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_457]), c_0_458])).
cnf(c_0_478,negated_conjecture,(epred47_0|epred32_0|epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_426]), c_0_459])).
fof(c_0_479, plain, ((~epred15_0<=>![X1]:((c2_1(X1)|~c1_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_480,negated_conjecture,(epred48_0|epred32_0|~c2_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_310, c_0_147])).
cnf(c_0_481,negated_conjecture,(epred48_0|hskp2|hskp15|~hskp8), inference(spm,[status(thm)],[c_0_203, c_0_460])).
cnf(c_0_482,negated_conjecture,(hskp3|hskp15|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_461]), c_0_38]), c_0_37])).
cnf(c_0_483,negated_conjecture,(epred39_0|hskp3|hskp12|hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_462, c_0_97]), c_0_463])).
cnf(c_0_484,negated_conjecture,(epred17_0|c0_1(a492)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_54]), c_0_48])).
cnf(c_0_485,negated_conjecture,(epred50_0|epred38_0|epred3_0|hskp3), inference(spm,[status(thm)],[c_0_464, c_0_465])).
cnf(c_0_486,negated_conjecture,(epred49_0|epred38_0|~epred50_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_466, c_0_446]), c_0_336]), c_0_467])).
cnf(c_0_487,negated_conjecture,(epred49_0|epred6_0|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_468]), c_0_76])).
cnf(c_0_488,negated_conjecture,(epred52_0|hskp2|hskp3|hskp12), inference(spm,[status(thm)],[c_0_362, c_0_350])).
cnf(c_0_489,negated_conjecture,(epred2_0|c3_1(X1)|c2_1(X1)|c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_469])).
cnf(c_0_490,negated_conjecture,(epred1_0|hskp27|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_470]), c_0_38]), c_0_37])).
cnf(c_0_491,negated_conjecture,(epred28_0|epred25_0|hskp8|~hskp3), inference(spm,[status(thm)],[c_0_155, c_0_471])).
cnf(c_0_492,negated_conjecture,(~epred6_0|~epred5_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_98, c_0_395]), c_0_161])).
cnf(c_0_493,negated_conjecture,(epred29_0|epred5_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_472]), c_0_473]), c_0_474])).
cnf(c_0_494,negated_conjecture,(epred29_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_475])).
cnf(c_0_495,negated_conjecture,(epred47_0|epred42_0|epred4_0|hskp12), inference(spm,[status(thm)],[c_0_406, c_0_399])).
cnf(c_0_496,negated_conjecture,(epred47_0|epred42_0|epred6_0|~epred4_0), inference(spm,[status(thm)],[c_0_292, c_0_476])).
cnf(c_0_497,negated_conjecture,(epred47_0|epred32_0), inference(spm,[status(thm)],[c_0_477, c_0_478])).
cnf(c_0_498,negated_conjecture,(epred15_0|c2_1(X1)|~c3_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_479])).
cnf(c_0_499,negated_conjecture,(epred48_0|epred32_0|~hskp25), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_84]), c_0_83])).
cnf(c_0_500,negated_conjecture,(epred48_0|epred32_0|~hskp9), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_96]), c_0_90])).
fof(c_0_501, plain, ((~epred30_0<=>![X1]:((~c1_1(X1)|~c2_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_502,negated_conjecture,(epred48_0|hskp2|hskp15|hskp16), inference(spm,[status(thm)],[c_0_481, c_0_55])).
cnf(c_0_503,negated_conjecture,(epred39_0|hskp3|hskp12|hskp15), inference(spm,[status(thm)],[c_0_482, c_0_483])).
cnf(c_0_504,negated_conjecture,(epred39_0|epred17_0|~hskp15), inference(spm,[status(thm)],[c_0_249, c_0_484])).
cnf(c_0_505,negated_conjecture,(epred52_0|epred32_0|hskp2|hskp12), inference(spm,[status(thm)],[c_0_362, c_0_387])).
cnf(c_0_506,negated_conjecture,(epred50_0|epred38_0|epred17_0|epred3_0), inference(spm,[status(thm)],[c_0_191, c_0_485])).
cnf(c_0_507,negated_conjecture,(epred49_0|epred38_0|epred4_0|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_465]), c_0_486])).
cnf(c_0_508,negated_conjecture,(epred52_0|epred42_0|~hskp3), inference(spm,[status(thm)],[c_0_362, c_0_370])).
cnf(c_0_509,negated_conjecture,(epred52_0|epred49_0|epred6_0|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_488]), c_0_197])).
cnf(c_0_510,negated_conjecture,(epred2_0|c0_1(a488)|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_489]), c_0_37])).
cnf(c_0_511,negated_conjecture,(epred39_0|epred1_0|hskp3|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_490, c_0_483]), c_0_289])).
cnf(c_0_512,negated_conjecture,(epred48_0|epred47_0|epred28_0|epred25_0|~hskp3), inference(spm,[status(thm)],[c_0_227, c_0_491])).
cnf(c_0_513,negated_conjecture,(epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_493]), c_0_494])).
cnf(c_0_514,negated_conjecture,(epred47_0|epred42_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_495]), c_0_496])).
cnf(c_0_515,negated_conjecture,(epred47_0|~epred42_0), inference(spm,[status(thm)],[c_0_437, c_0_497])).
fof(c_0_516, plain, ((~epred14_0<=>![X2]:((((hskp4|c0_1(X2))|c1_1(X2))|c3_1(X2))|~ndr1_0))), introduced(definition)).
fof(c_0_517, plain, ((~epred41_0<=>![X2]:((((hskp17|c1_1(X2))|~ndr1_0)|~c0_1(X2))|~c2_1(X2)))), introduced(definition)).
cnf(c_0_518,negated_conjecture,(epred15_0|c2_1(a466)|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_498, c_0_7]), c_0_9])).
cnf(c_0_519,negated_conjecture,(epred48_0|epred32_0|epred5_0|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_450]), c_0_500])).
cnf(c_0_520,negated_conjecture,(epred30_0|~c3_1(X1)|~c2_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_501])).
cnf(c_0_521,negated_conjecture,(c2_1(X1)|hskp12|hskp15|~c3_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_140])])).
cnf(c_0_522,negated_conjecture,(epred49_0|epred48_0|epred6_0|hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_502]), c_0_235])).
cnf(c_0_523,negated_conjecture,(epred49_0|epred48_0|epred47_0|hskp15), inference(spm,[status(thm)],[c_0_235, c_0_252])).
cnf(c_0_524,negated_conjecture,(epred39_0|epred17_0|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_503]), c_0_504])).
cnf(c_0_525,negated_conjecture,(epred52_0|epred49_0|epred32_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_505]), c_0_197])).
cnf(c_0_526,negated_conjecture,(epred50_0|epred38_0|epred17_0|~epred4_0), inference(spm,[status(thm)],[c_0_292, c_0_506])).
cnf(c_0_527,negated_conjecture,(epred49_0|epred38_0|epred17_0|epred4_0), inference(spm,[status(thm)],[c_0_191, c_0_507])).
cnf(c_0_528,negated_conjecture,(epred52_0|epred49_0|epred42_0|epred6_0), inference(spm,[status(thm)],[c_0_508, c_0_509])).
cnf(c_0_529,negated_conjecture,(epred47_0|epred2_0|~hskp14), inference(spm,[status(thm)],[c_0_198, c_0_510])).
cnf(c_0_530,negated_conjecture,(epred39_0|epred17_0|epred1_0|hskp12), inference(spm,[status(thm)],[c_0_191, c_0_511])).
cnf(c_0_531,negated_conjecture,(hskp2|hskp15|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_441, c_0_100]), c_0_101]), c_0_102])).
cnf(c_0_532,negated_conjecture,(hskp3|hskp12|hskp15|hskp20), inference(spm,[status(thm)],[c_0_482, c_0_223])).
cnf(c_0_533,negated_conjecture,(epred48_0|epred47_0|epred28_0|epred25_0|hskp12), inference(spm,[status(thm)],[c_0_512, c_0_368])).
cnf(c_0_534,negated_conjecture,(~epred28_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_296, c_0_513])])).
cnf(c_0_535,negated_conjecture,(epred47_0|~epred48_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_514]), c_0_338]), c_0_515])).
cnf(c_0_536,negated_conjecture,(epred25_0|~c2_1(a467)|~hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_103]), c_0_104])).
cnf(c_0_537,negated_conjecture,(epred6_0|c2_1(a467)|~hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_178]), c_0_104])).
cnf(c_0_538,negated_conjecture,(epred14_0|c3_1(X1)|c1_1(X1)|c0_1(X1)|hskp4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_516]), c_0_140])])).
cnf(c_0_539,negated_conjecture,(epred41_0|c1_1(X1)|hskp17|~c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_517]), c_0_140])])).
cnf(c_0_540,negated_conjecture,(epred6_0|hskp19|hskp20|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_196]), c_0_34]), c_0_41])).
cnf(c_0_541,negated_conjecture,(epred25_0|epred15_0|~hskp3), inference(spm,[status(thm)],[c_0_155, c_0_518])).
cnf(c_0_542,negated_conjecture,(epred48_0|epred38_0|hskp3|hskp12), inference(spm,[status(thm)],[c_0_189, c_0_351])).
cnf(c_0_543,negated_conjecture,(epred48_0|epred32_0|epred25_0|epred5_0), inference(spm,[status(thm)],[c_0_408, c_0_519])).
cnf(c_0_544,negated_conjecture,(epred30_0|~c2_1(a466)|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_520, c_0_7]), c_0_9])).
cnf(c_0_545,negated_conjecture,(c2_1(a466)|hskp12|hskp15|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_521, c_0_7]), c_0_9])).
cnf(c_0_546,negated_conjecture,(epred49_0|epred48_0|epred39_0|epred6_0), inference(spm,[status(thm)],[c_0_291, c_0_522])).
cnf(c_0_547,negated_conjecture,(epred49_0|epred48_0|epred47_0|epred39_0), inference(spm,[status(thm)],[c_0_291, c_0_523])).
cnf(c_0_548,negated_conjecture,(epred39_0|epred38_0|epred17_0|epred8_0), inference(spm,[status(thm)],[c_0_239, c_0_524])).
cnf(c_0_549,negated_conjecture,(epred39_0|epred38_0|epred17_0|epred9_0), inference(spm,[status(thm)],[c_0_241, c_0_524])).
cnf(c_0_550,negated_conjecture,(epred52_0|epred49_0|epred32_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_525]), c_0_458])).
cnf(c_0_551,negated_conjecture,(epred49_0|epred38_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_526, c_0_527]), c_0_486])).
cnf(c_0_552,negated_conjecture,(epred52_0|epred49_0|epred42_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_528]), c_0_439])).
cnf(c_0_553,negated_conjecture,(epred39_0|hskp3|hskp12|hskp11), inference(spm,[status(thm)],[c_0_463, c_0_1])).
cnf(c_0_554,negated_conjecture,(epred47_0|epred39_0|epred2_0|hskp3|hskp12), inference(spm,[status(thm)],[c_0_529, c_0_483])).
cnf(c_0_555,negated_conjecture,(~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_106, c_0_447]), c_0_469])).
cnf(c_0_556,negated_conjecture,(epred47_0|epred39_0|epred17_0|epred8_0|epred1_0), inference(spm,[status(thm)],[c_0_217, c_0_530])).
cnf(c_0_557,negated_conjecture,(hskp2|hskp15|hskp17|hskp27), inference(spm,[status(thm)],[c_0_531, c_0_107])).
cnf(c_0_558,negated_conjecture,(epred11_0|~c0_1(a492)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_54]), c_0_49])).
cnf(c_0_559,negated_conjecture,(epred17_0|hskp12|hskp15|hskp20), inference(spm,[status(thm)],[c_0_191, c_0_532])).
cnf(c_0_560,negated_conjecture,(epred47_0|hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_533, c_0_534]), c_0_477]), c_0_535])).
cnf(c_0_561,negated_conjecture,(epred25_0|epred6_0|~hskp4), inference(spm,[status(thm)],[c_0_536, c_0_537])).
cnf(c_0_562,negated_conjecture,(epred14_0|hskp4|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_538]), c_0_35]), c_0_34])).
cnf(c_0_563,negated_conjecture,(epred41_0|hskp17|~c0_1(a503)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_539, c_0_12]), c_0_11])).
cnf(c_0_564,negated_conjecture,(epred25_0|epred6_0|hskp20|~hskp17), inference(spm,[status(thm)],[c_0_159, c_0_540])).
cnf(c_0_565,negated_conjecture,(epred6_0|hskp16|hskp17|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_220]), c_0_38]), c_0_37])).
cnf(c_0_566,negated_conjecture,(epred25_0|epred15_0|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_541, c_0_176])).
cnf(c_0_567,negated_conjecture,(epred48_0|epred38_0|epred17_0|hskp12), inference(spm,[status(thm)],[c_0_191, c_0_542])).
cnf(c_0_568,negated_conjecture,(epred48_0|epred32_0|epred25_0|~epred6_0), inference(spm,[status(thm)],[c_0_492, c_0_543])).
cnf(c_0_569,negated_conjecture,(epred32_0|epred25_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_426])).
cnf(c_0_570,negated_conjecture,(epred30_0|hskp12|hskp15|~hskp3), inference(spm,[status(thm)],[c_0_544, c_0_545])).
cnf(c_0_571,negated_conjecture,(epred49_0|epred48_0|epred39_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_546]), c_0_547])).
cnf(c_0_572,negated_conjecture,(epred39_0|epred38_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_548]), c_0_549])).
cnf(c_0_573,negated_conjecture,(epred52_0|epred49_0|epred38_0|epred32_0), inference(spm,[status(thm)],[c_0_550, c_0_551])).
cnf(c_0_574,negated_conjecture,(epred52_0|epred49_0|epred42_0|epred38_0), inference(spm,[status(thm)],[c_0_552, c_0_551])).
cnf(c_0_575,negated_conjecture,(epred39_0|hskp0|hskp3|hskp12), inference(spm,[status(thm)],[c_0_259, c_0_553])).
cnf(c_0_576,negated_conjecture,(epred47_0|epred39_0|epred17_0|epred2_0|hskp12), inference(spm,[status(thm)],[c_0_191, c_0_554])).
cnf(c_0_577,negated_conjecture,(epred47_0|epred39_0|epred17_0|epred8_0|~epred2_0), inference(spm,[status(thm)],[c_0_555, c_0_556])).
cnf(c_0_578,negated_conjecture,(epred2_0|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_489]), c_0_76]), c_0_77])).
cnf(c_0_579,negated_conjecture,(epred4_0|hskp2|hskp15|hskp27), inference(spm,[status(thm)],[c_0_209, c_0_557])).
cnf(c_0_580,negated_conjecture,(epred38_0|epred11_0|hskp3|hskp12), inference(spm,[status(thm)],[c_0_189, c_0_282])).
cnf(c_0_581,negated_conjecture,(epred17_0|epred11_0|~hskp15), inference(spm,[status(thm)],[c_0_558, c_0_484])).
cnf(c_0_582,negated_conjecture,(epred17_0|epred9_0|hskp12|hskp15), inference(spm,[status(thm)],[c_0_401, c_0_559])).
cnf(c_0_583,negated_conjecture,(epred47_0|epred8_0), inference(spm,[status(thm)],[c_0_217, c_0_560])).
cnf(c_0_584,negated_conjecture,(epred47_0|epred9_0), inference(spm,[status(thm)],[c_0_267, c_0_560])).
cnf(c_0_585,negated_conjecture,(epred25_0|epred14_0|epred6_0|~hskp17), inference(spm,[status(thm)],[c_0_561, c_0_562])).
cnf(c_0_586,negated_conjecture,(epred41_0|epred25_0|hskp17|~hskp20), inference(spm,[status(thm)],[c_0_563, c_0_158])).
cnf(c_0_587,negated_conjecture,(epred25_0|epred6_0|hskp16|hskp20|~hskp14), inference(spm,[status(thm)],[c_0_564, c_0_565])).
cnf(c_0_588,negated_conjecture,(epred25_0|epred15_0|hskp12), inference(spm,[status(thm)],[c_0_566, c_0_195])).
cnf(c_0_589,negated_conjecture,(epred48_0|epred38_0|epred17_0|epred8_0), inference(spm,[status(thm)],[c_0_239, c_0_567])).
cnf(c_0_590,negated_conjecture,(epred48_0|epred38_0|epred17_0|epred9_0), inference(spm,[status(thm)],[c_0_241, c_0_420])).
cnf(c_0_591,negated_conjecture,(epred48_0|epred32_0|epred25_0), inference(spm,[status(thm)],[c_0_568, c_0_569])).
cnf(c_0_592,negated_conjecture,(epred49_0|epred30_0|~hskp3), inference(spm,[status(thm)],[c_0_544, c_0_156])).
cnf(c_0_593,negated_conjecture,(epred39_0|epred30_0|hskp12|hskp15), inference(spm,[status(thm)],[c_0_570, c_0_503])).
cnf(c_0_594,negated_conjecture,(epred49_0|epred48_0|epred39_0|epred38_0), inference(spm,[status(thm)],[c_0_571, c_0_572])).
cnf(c_0_595,negated_conjecture,(epred52_0|epred49_0|epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_437, c_0_573]), c_0_574])).
cnf(c_0_596,negated_conjecture,(epred49_0|~epred52_0|~epred39_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_314]), c_0_335])).
cnf(c_0_597,negated_conjecture,(epred47_0|epred39_0|hskp3|hskp12), inference(spm,[status(thm)],[c_0_193, c_0_575])).
cnf(c_0_598,negated_conjecture,(epred47_0|epred39_0|epred17_0|epred8_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_576]), c_0_577])).
cnf(c_0_599,negated_conjecture,(epred45_0|hskp2|hskp15|~hskp8), inference(spm,[status(thm)],[c_0_211, c_0_460])).
cnf(c_0_600,negated_conjecture,(epred4_0|epred2_0|hskp15|hskp27), inference(spm,[status(thm)],[c_0_578, c_0_579])).
cnf(c_0_601,negated_conjecture,(epred38_0|epred17_0|epred11_0|hskp12), inference(spm,[status(thm)],[c_0_191, c_0_580])).
cnf(c_0_602,negated_conjecture,(epred17_0|epred11_0|epred9_0|hskp12), inference(spm,[status(thm)],[c_0_581, c_0_582])).
cnf(c_0_603,negated_conjecture,(epred38_0|epred6_0|hskp3|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_220]), c_0_37]), c_0_36])).
cnf(c_0_604,negated_conjecture,(epred47_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_583]), c_0_584])).
cnf(c_0_605,negated_conjecture,(epred41_0|epred25_0|epred14_0|epred6_0|~hskp20), inference(spm,[status(thm)],[c_0_585, c_0_586])).
cnf(c_0_606,negated_conjecture,(epred25_0|epred6_0|hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_587, c_0_223]), c_0_195]), c_0_197])).
cnf(c_0_607,negated_conjecture,(epred25_0|epred15_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_588])).
cnf(c_0_608,negated_conjecture,(epred48_0|epred38_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_589]), c_0_590])).
cnf(c_0_609,negated_conjecture,(epred32_0|epred25_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_569]), c_0_591])).
cnf(c_0_610,negated_conjecture,(epred49_0|epred38_0|epred30_0|epred8_0|~hskp12), inference(spm,[status(thm)],[c_0_592, c_0_214])).
cnf(c_0_611,negated_conjecture,(epred49_0|epred39_0|epred30_0|hskp12), inference(spm,[status(thm)],[c_0_291, c_0_593])).
cnf(c_0_612,negated_conjecture,(epred49_0|epred39_0|epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_446]), c_0_594])).
cnf(c_0_613,negated_conjecture,(epred49_0|~epred39_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_595]), c_0_596])).
cnf(c_0_614,negated_conjecture,(epred47_0|epred39_0|epred17_0|hskp12), inference(spm,[status(thm)],[c_0_191, c_0_597])).
cnf(c_0_615,negated_conjecture,(epred47_0|epred39_0|epred17_0|~epred9_0), inference(spm,[status(thm)],[c_0_242, c_0_598])).
cnf(c_0_616,negated_conjecture,(epred45_0|hskp2|hskp15|hskp16), inference(spm,[status(thm)],[c_0_599, c_0_55])).
cnf(c_0_617,negated_conjecture,(epred39_0|epred17_0|epred6_0|epred1_0), inference(spm,[status(thm)],[c_0_197, c_0_530])).
cnf(c_0_618,negated_conjecture,(epred39_0|epred4_0|epred2_0|hskp15), inference(spm,[status(thm)],[c_0_289, c_0_600])).
cnf(c_0_619,negated_conjecture,(epred46_0|epred39_0|hskp3|hskp12), inference(spm,[status(thm)],[c_0_255, c_0_483])).
cnf(c_0_620,negated_conjecture,(epred38_0|epred4_0|hskp3|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_201]), c_0_37]), c_0_36])).
cnf(c_0_621,negated_conjecture,(epred38_0|epred17_0|epred11_0|epred8_0), inference(spm,[status(thm)],[c_0_239, c_0_601])).
cnf(c_0_622,negated_conjecture,(epred38_0|epred17_0|epred11_0|epred9_0), inference(spm,[status(thm)],[c_0_241, c_0_602])).
cnf(c_0_623,negated_conjecture,(epred38_0|epred6_0|hskp2|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_350]), c_0_197])).
cnf(c_0_624,negated_conjecture,(~epred17_0|~epred6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_263, c_0_604])])).
cnf(c_0_625,negated_conjecture,(epred41_0|epred25_0|epred14_0|epred6_0), inference(spm,[status(thm)],[c_0_605, c_0_606])).
cnf(c_0_626,negated_conjecture,(epred25_0|epred15_0|~epred47_0|~epred17_0), inference(spm,[status(thm)],[c_0_263, c_0_607])).
cnf(c_0_627,negated_conjecture,(epred48_0|epred38_0|~epred49_0|~epred25_0), inference(spm,[status(thm)],[c_0_456, c_0_608])).
cnf(c_0_628,negated_conjecture,(epred32_0|epred25_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_609, c_0_513])])).
cnf(c_0_629,negated_conjecture,(epred49_0|epred38_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_610, c_0_611]), c_0_612]), c_0_612]), c_0_613])).
cnf(c_0_630,negated_conjecture,(epred47_0|epred39_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_614]), c_0_615])).
cnf(c_0_631,negated_conjecture,(epred45_0|epred28_0|hskp2|hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_599, c_0_228]), c_0_616])).
cnf(c_0_632,negated_conjecture,(epred39_0|epred17_0|epred6_0|~epred2_0), inference(spm,[status(thm)],[c_0_555, c_0_617])).
cnf(c_0_633,negated_conjecture,(epred39_0|epred17_0|epred4_0|epred2_0), inference(spm,[status(thm)],[c_0_504, c_0_618])).
cnf(c_0_634,negated_conjecture,(epred46_0|epred39_0|epred17_0|hskp12), inference(spm,[status(thm)],[c_0_191, c_0_619])).
cnf(c_0_635,negated_conjecture,(c2_1(X1)|hskp18|hskp20|~c1_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_140])])).
cnf(c_0_636,negated_conjecture,(epred38_0|epred4_0|hskp2|hskp3|hskp12), inference(spm,[status(thm)],[c_0_620, c_0_350])).
cnf(c_0_637,negated_conjecture,(epred38_0|epred17_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_621]), c_0_622])).
cnf(c_0_638,negated_conjecture,(epred38_0|epred6_0|epred3_0|hskp3), inference(spm,[status(thm)],[c_0_464, c_0_623])).
cnf(c_0_639,negated_conjecture,(~epred15_0|~epred14_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_109, c_0_516]), c_0_479])).
cnf(c_0_640,negated_conjecture,(epred41_0|epred25_0|epred14_0|~epred17_0), inference(spm,[status(thm)],[c_0_624, c_0_625])).
cnf(c_0_641,negated_conjecture,(epred25_0|epred15_0|~epred17_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_626, c_0_604])])).
cnf(c_0_642,negated_conjecture,(~epred41_0|~epred32_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_110, c_0_517]), c_0_283])).
cnf(c_0_643,negated_conjecture,(epred48_0|epred38_0|epred32_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_628]), c_0_629])).
cnf(c_0_644,negated_conjecture,(epred47_0|epred39_0|~epred49_0|~epred25_0), inference(spm,[status(thm)],[c_0_456, c_0_630])).
cnf(c_0_645,negated_conjecture,(epred45_0|epred28_0|epred2_0|hskp15), inference(spm,[status(thm)],[c_0_578, c_0_631])).
cnf(c_0_646,negated_conjecture,(epred39_0|epred17_0|epred6_0|epred4_0), inference(spm,[status(thm)],[c_0_632, c_0_633])).
cnf(c_0_647,negated_conjecture,(epred46_0|epred39_0|epred17_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_634])).
cnf(c_0_648,negated_conjecture,(c0_1(X1)|hskp10|hskp27|~c2_1(X1)|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_140])])).
cnf(c_0_649,negated_conjecture,(hskp18|hskp20|~c0_1(a492)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_635, c_0_48]), c_0_49])).
cnf(c_0_650,negated_conjecture,(hskp12|hskp15|~c1_1(a468)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_521, c_0_50]), c_0_52])).
cnf(c_0_651,negated_conjecture,(epred48_0|~c2_1(a533)|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_147]), c_0_85])).
cnf(c_0_652,negated_conjecture,(epred12_0|c2_1(a533)|hskp12|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_85]), c_0_87])).
cnf(c_0_653,negated_conjecture,(epred38_0|epred22_0|epred4_0|hskp3|hskp12), inference(spm,[status(thm)],[c_0_379, c_0_636])).
cnf(c_0_654,negated_conjecture,(epred38_0|epred17_0|~epred22_0), inference(spm,[status(thm)],[c_0_364, c_0_637])).
cnf(c_0_655,negated_conjecture,(epred38_0|epred17_0|epred6_0|epred3_0), inference(spm,[status(thm)],[c_0_191, c_0_638])).
cnf(c_0_656,negated_conjecture,(epred41_0|epred25_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_639, c_0_640]), c_0_641])).
cnf(c_0_657,negated_conjecture,(epred48_0|epred38_0|~epred41_0), inference(spm,[status(thm)],[c_0_642, c_0_643])).
cnf(c_0_658,negated_conjecture,(epred47_0|epred39_0|epred32_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_644, c_0_609]), c_0_338]), c_0_357])).
cnf(c_0_659,negated_conjecture,(epred45_0|epred39_0|epred28_0|epred17_0|epred2_0), inference(spm,[status(thm)],[c_0_504, c_0_645])).
cnf(c_0_660,negated_conjecture,(epred39_0|epred17_0|epred6_0|~epred45_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_646]), c_0_647])).
cnf(c_0_661,negated_conjecture,(epred32_0|~hskp25), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_82]), c_0_83]), c_0_84])).
cnf(c_0_662,negated_conjecture,(hskp10|hskp27|~hskp18), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_648, c_0_112]), c_0_113]), c_0_114])).
cnf(c_0_663,negated_conjecture,(epred17_0|hskp18|hskp20|~hskp15), inference(spm,[status(thm)],[c_0_649, c_0_484])).
cnf(c_0_664,negated_conjecture,(hskp3|hskp12|hskp15|~hskp5), inference(spm,[status(thm)],[c_0_650, c_0_444])).
cnf(c_0_665,negated_conjecture,(epred48_0|epred12_0|hskp12|~hskp23), inference(spm,[status(thm)],[c_0_651, c_0_652])).
cnf(c_0_666,negated_conjecture,(epred38_0|epred17_0|epred4_0|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_653]), c_0_654])).
cnf(c_0_667,negated_conjecture,(epred38_0|epred17_0|epred6_0|~epred4_0), inference(spm,[status(thm)],[c_0_292, c_0_655])).
cnf(c_0_668,negated_conjecture,(epred48_0|epred38_0|epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_656, c_0_608]), c_0_657])).
cnf(c_0_669,negated_conjecture,(epred47_0|epred39_0|epred6_0|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_483]), c_0_197])).
cnf(c_0_670,negated_conjecture,(epred47_0|epred39_0|~epred42_0), inference(spm,[status(thm)],[c_0_437, c_0_658])).
cnf(c_0_671,negated_conjecture,(epred39_0|epred28_0|epred17_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_632, c_0_659]), c_0_660])).
cnf(c_0_672,negated_conjecture,(epred32_0|epred5_0|hskp9|hskp20), inference(spm,[status(thm)],[c_0_661, c_0_450])).
cnf(c_0_673,negated_conjecture,(epred50_0|~c2_1(a494)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_284]), c_0_34])).
cnf(c_0_674,negated_conjecture,(epred2_0|c2_1(a494)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_489]), c_0_35])).
cnf(c_0_675,negated_conjecture,(epred50_0|epred29_0|c1_1(X1)|~c2_1(X1)), inference(spm,[status(thm)],[c_0_206, c_0_284])).
cnf(c_0_676,negated_conjecture,(epred17_0|hskp10|hskp20|hskp27|~hskp15), inference(spm,[status(thm)],[c_0_662, c_0_663])).
cnf(c_0_677,negated_conjecture,(hskp3|hskp12|hskp15|hskp11), inference(spm,[status(thm)],[c_0_664, c_0_1])).
cnf(c_0_678,negated_conjecture,(epred48_0|epred12_0|hskp9|hskp12|hskp20), inference(spm,[status(thm)],[c_0_665, c_0_92])).
cnf(c_0_679,negated_conjecture,(epred13_0|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_324]), c_0_85]), c_0_86])).
cnf(c_0_680,negated_conjecture,(~epred48_0|~epred6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_315, c_0_513])])).
cnf(c_0_681,negated_conjecture,(epred38_0|epred17_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_666]), c_0_667])).
cnf(c_0_682,negated_conjecture,(epred48_0|epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_668]), c_0_629])).
cnf(c_0_683,negated_conjecture,(epred42_0|epred12_0|epred6_0|~hskp3), inference(spm,[status(thm)],[c_0_341, c_0_370])).
cnf(c_0_684,negated_conjecture,(epred39_0|epred12_0|epred6_0|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_483]), c_0_197])).
cnf(c_0_685,negated_conjecture,(epred47_0|epred39_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_455, c_0_669]), c_0_670])).
cnf(c_0_686,negated_conjecture,(epred42_0|epred13_0|epred6_0|~hskp3), inference(spm,[status(thm)],[c_0_360, c_0_370])).
cnf(c_0_687,negated_conjecture,(epred39_0|epred13_0|epred6_0|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_360, c_0_483]), c_0_197])).
cnf(c_0_688,negated_conjecture,(epred39_0|epred17_0|~epred48_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_671]), c_0_296])).
cnf(c_0_689,negated_conjecture,(epred32_0|epred29_0|epred5_0|hskp9), inference(spm,[status(thm)],[c_0_230, c_0_672])).
cnf(c_0_690,negated_conjecture,(epred50_0|epred2_0|~hskp17), inference(spm,[status(thm)],[c_0_673, c_0_674])).
cnf(c_0_691,negated_conjecture,(epred50_0|epred29_0|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_675]), c_0_16])).
cnf(c_0_692,negated_conjecture,(epred1_0|c1_1(a460)|hskp27|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_470]), c_0_31])).
cnf(c_0_693,negated_conjecture,(epred29_0|epred17_0|hskp10|hskp27|~hskp15), inference(spm,[status(thm)],[c_0_230, c_0_676])).
cnf(c_0_694,negated_conjecture,(epred2_0|hskp6|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_674]), c_0_578]), c_0_35]), c_0_34])).
cnf(c_0_695,negated_conjecture,(hskp0|hskp3|hskp12|hskp15), inference(spm,[status(thm)],[c_0_259, c_0_677])).
fof(c_0_696, plain, ((~epred35_0<=>![X1]:((((hskp9|c1_1(X1))|~ndr1_0)|~c0_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_697,negated_conjecture,(epred51_0|epred28_0|hskp8|~hskp15), inference(spm,[status(thm)],[c_0_233, c_0_250])).
cnf(c_0_698,negated_conjecture,(epred48_0|epred29_0|epred12_0|hskp9|hskp12), inference(spm,[status(thm)],[c_0_230, c_0_678])).
cnf(c_0_699,negated_conjecture,(epred48_0|epred32_0|epred29_0|epred5_0), inference(spm,[status(thm)],[c_0_230, c_0_519])).
cnf(c_0_700,negated_conjecture,(epred13_0|hskp9|hskp20), inference(spm,[status(thm)],[c_0_679, c_0_92])).
cnf(c_0_701,negated_conjecture,(epred38_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_680, c_0_681]), c_0_682])).
cnf(c_0_702,negated_conjecture,(epred42_0|epred39_0|epred12_0|epred6_0), inference(spm,[status(thm)],[c_0_683, c_0_684])).
cnf(c_0_703,negated_conjecture,(epred47_0|epred39_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_685]), c_0_338]), c_0_339])).
cnf(c_0_704,negated_conjecture,(epred42_0|epred39_0|epred13_0|epred6_0), inference(spm,[status(thm)],[c_0_686, c_0_687])).
cnf(c_0_705,negated_conjecture,(epred39_0|~epred48_0|~epred29_0|~epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_688]), c_0_337])).
cnf(c_0_706,negated_conjecture,(epred32_0|epred29_0|epred5_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_661, c_0_472]), c_0_689])).
cnf(c_0_707,negated_conjecture,(epred50_0|epred2_0|hskp15|hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_557]), c_0_578])).
cnf(c_0_708,negated_conjecture,(epred50_0|epred29_0|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_675]), c_0_12])).
cnf(c_0_709,negated_conjecture,(epred50_0|epred29_0|hskp12|hskp20), inference(spm,[status(thm)],[c_0_691, c_0_33])).
cnf(c_0_710,negated_conjecture,(epred1_0|c1_1(a465)|hskp27|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_470]), c_0_77])).
cnf(c_0_711,negated_conjecture,(epred1_0|hskp2|hskp15|hskp27|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_441, c_0_692]), c_0_30]), c_0_31])).
cnf(c_0_712,negated_conjecture,(epred29_0|hskp0|hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_312]), c_0_394]), c_0_427])).
cnf(c_0_713,negated_conjecture,(epred29_0|~hskp15), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_693]), c_0_504]), c_0_427]), c_0_474])).
cnf(c_0_714,negated_conjecture,(epred2_0|hskp19|hskp20|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_674]), c_0_34]), c_0_41])).
cnf(c_0_715,negated_conjecture,(epred29_0|epred17_0|epred2_0|~hskp17), inference(spm,[status(thm)],[c_0_377, c_0_694])).
fof(c_0_716, plain, ((~epred18_0<=>![X2]:((((hskp2|c0_1(X2))|c1_1(X2))|c3_1(X2))|~ndr1_0))), introduced(definition)).
fof(c_0_717, plain, ((~epred19_0<=>![X1]:((c0_1(X1)|~c1_1(X1))|~c2_1(X1)))), introduced(definition)).
fof(c_0_718, plain, ((~epred16_0<=>![X2]:((((hskp3|c0_1(X2))|c1_1(X2))|c3_1(X2))|~ndr1_0))), introduced(definition)).
cnf(c_0_719,negated_conjecture,(epred47_0|hskp3|hskp12|hskp15), inference(spm,[status(thm)],[c_0_193, c_0_695])).
fof(c_0_720, plain, ((~epred10_0<=>![X2]:((((hskp28|c1_1(X2))|c2_1(X2))|c3_1(X2))|~ndr1_0))), introduced(definition)).
cnf(c_0_721,negated_conjecture,(epred35_0|c1_1(X1)|hskp9|~c3_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_696]), c_0_140])])).
fof(c_0_722, plain, ((~epred26_0<=>![X1]:((((hskp9|c1_1(X1))|c2_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
cnf(c_0_723,negated_conjecture,(epred51_0|epred48_0|epred28_0|~hskp16), inference(spm,[status(thm)],[c_0_281, c_0_228])).
cnf(c_0_724,negated_conjecture,(epred51_0|epred48_0|epred28_0|~hskp15), inference(spm,[status(thm)],[c_0_281, c_0_697])).
cnf(c_0_725,negated_conjecture,(epred48_0|epred32_0|epred29_0|epred12_0|hskp12), inference(spm,[status(thm)],[c_0_500, c_0_698])).
cnf(c_0_726,negated_conjecture,(epred48_0|epred32_0|epred29_0|~epred6_0), inference(spm,[status(thm)],[c_0_492, c_0_699])).
cnf(c_0_727,negated_conjecture,(epred29_0|epred13_0|hskp9), inference(spm,[status(thm)],[c_0_230, c_0_700])).
cnf(c_0_728,negated_conjecture,(c1_1(a468)|hskp7|hskp12|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_50]), c_0_52])).
cnf(c_0_729,negated_conjecture,(epred38_0|~epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_701]), c_0_629])).
cnf(c_0_730,negated_conjecture,(epred41_0|epred38_0|epred25_0), inference(spm,[status(thm)],[c_0_656, c_0_287])).
cnf(c_0_731,negated_conjecture,(epred42_0|epred39_0|epred12_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_702]), c_0_703])).
cnf(c_0_732,negated_conjecture,(epred42_0|epred39_0|epred13_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_704]), c_0_703])).
cnf(c_0_733,negated_conjecture,(epred39_0|epred32_0|~epred48_0|~epred29_0), inference(spm,[status(thm)],[c_0_705, c_0_609])).
cnf(c_0_734,negated_conjecture,(epred32_0|epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_706]), c_0_494])).
cnf(c_0_735,negated_conjecture,(epred50_0|epred39_0|epred2_0|hskp15), inference(spm,[status(thm)],[c_0_289, c_0_707])).
cnf(c_0_736,negated_conjecture,(epred50_0|epred29_0|~hskp9), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_675]), c_0_96])).
cnf(c_0_737,negated_conjecture,(epred50_0|epred29_0|hskp12), inference(spm,[status(thm)],[c_0_708, c_0_709])).
cnf(c_0_738,negated_conjecture,(epred49_0|epred1_0|hskp27|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_710]), c_0_76]), c_0_77])).
cnf(c_0_739,negated_conjecture,(epred29_0|epred1_0|hskp2|hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_711, c_0_712]), c_0_713])).
cnf(c_0_740,negated_conjecture,(epred17_0|epred2_0|c2_1(X1)|c0_1(X1)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_489])).
cnf(c_0_741,negated_conjecture,(epred29_0|epred2_0|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_714]), c_0_715]), c_0_427])).
fof(c_0_742, plain, ((~epred36_0<=>![X2]:((c1_1(X2)|~c0_1(X2))|~c2_1(X2)))), introduced(definition)).
cnf(c_0_743,negated_conjecture,(epred18_0|c3_1(X1)|c1_1(X1)|c0_1(X1)|hskp2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_716]), c_0_140])])).
cnf(c_0_744,negated_conjecture,(epred51_0|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_100]), c_0_101]), c_0_102])).
cnf(c_0_745,negated_conjecture,(epred19_0|c0_1(X1)|~c2_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_717])).
cnf(c_0_746,negated_conjecture,(epred16_0|c3_1(X1)|c1_1(X1)|c0_1(X1)|hskp3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_718]), c_0_140])])).
cnf(c_0_747,negated_conjecture,(epred51_0|epred17_0|~hskp15), inference(spm,[status(thm)],[c_0_233, c_0_484])).
cnf(c_0_748,negated_conjecture,(epred47_0|epred17_0|hskp12|hskp15), inference(spm,[status(thm)],[c_0_191, c_0_719])).
fof(c_0_749, plain, ((~epred33_0<=>![X2]:((((hskp26|c0_1(X2))|c1_1(X2))|~ndr1_0)|~c2_1(X2)))), introduced(definition)).
cnf(c_0_750,negated_conjecture,(epred10_0|c3_1(X1)|c2_1(X1)|c1_1(X1)|hskp28), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_720]), c_0_140])])).
cnf(c_0_751,negated_conjecture,(epred35_0|c1_1(a461)|hskp9|~hskp25), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_721, c_0_82]), c_0_83])).
cnf(c_0_752,negated_conjecture,(epred35_0|c1_1(a468)|hskp9|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_721, c_0_50]), c_0_51])).
cnf(c_0_753,negated_conjecture,(epred26_0|c2_1(X1)|c1_1(X1)|hskp9|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_722]), c_0_140])])).
cnf(c_0_754,negated_conjecture,(epred35_0|hskp9|~hskp21), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_721, c_0_67]), c_0_68]), c_0_69])).
cnf(c_0_755,negated_conjecture,(epred51_0|epred48_0|epred28_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_723, c_0_302]), c_0_724])).
cnf(c_0_756,negated_conjecture,(epred48_0|epred32_0|epred29_0|epred12_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_725]), c_0_726])).
cnf(c_0_757,negated_conjecture,(epred48_0|epred32_0|epred29_0|epred13_0), inference(spm,[status(thm)],[c_0_500, c_0_727])).
fof(c_0_758, plain, ((~epred7_0<=>![X1]:((((hskp7|c2_1(X1))|c3_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
cnf(c_0_759,negated_conjecture,(hskp7|hskp12|hskp15|~hskp5), inference(spm,[status(thm)],[c_0_650, c_0_728])).
cnf(c_0_760,negated_conjecture,(epred38_0|epred32_0), inference(spm,[status(thm)],[c_0_729, c_0_628])).
cnf(c_0_761,negated_conjecture,(epred41_0|epred38_0), inference(spm,[status(thm)],[c_0_729, c_0_730])).
cnf(c_0_762,negated_conjecture,(epred42_0|epred39_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_731]), c_0_732])).
cnf(c_0_763,negated_conjecture,(epred39_0|epred32_0|~epred48_0), inference(spm,[status(thm)],[c_0_733, c_0_734])).
cnf(c_0_764,negated_conjecture,(epred50_0|epred39_0|epred17_0|epred2_0), inference(spm,[status(thm)],[c_0_504, c_0_735])).
cnf(c_0_765,negated_conjecture,(epred50_0|epred39_0|epred29_0|epred5_0), inference(spm,[status(thm)],[c_0_736, c_0_473])).
cnf(c_0_766,negated_conjecture,(epred50_0|epred29_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_737])).
cnf(c_0_767,negated_conjecture,(epred49_0|epred29_0|epred1_0|hskp27), inference(spm,[status(thm)],[c_0_738, c_0_739])).
cnf(c_0_768,negated_conjecture,(epred17_0|epred2_0|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_740, c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_769,negated_conjecture,(epred29_0|epred2_0|hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_741, c_0_210]), c_0_713])).
cnf(c_0_770,negated_conjecture,(epred36_0|c1_1(X1)|~c2_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_742])).
cnf(c_0_771,negated_conjecture,(epred18_0|hskp2|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_743]), c_0_35]), c_0_34])).
cnf(c_0_772,negated_conjecture,(epred51_0|hskp17|hskp27), inference(spm,[status(thm)],[c_0_744, c_0_107])).
cnf(c_0_773,negated_conjecture,(epred19_0|~c2_1(a466)|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_745, c_0_9]), c_0_8])).
cnf(c_0_774,negated_conjecture,(epred16_0|hskp3|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_746]), c_0_35]), c_0_34])).
cnf(c_0_775,negated_conjecture,(epred51_0|epred47_0|epred17_0|hskp12), inference(spm,[status(thm)],[c_0_747, c_0_748])).
cnf(c_0_776,negated_conjecture,(epred33_0|c1_1(X1)|c0_1(X1)|hskp26|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_749]), c_0_140])])).
cnf(c_0_777,negated_conjecture,(epred10_0|c2_1(a533)|hskp28|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_750]), c_0_86])).
cnf(c_0_778,negated_conjecture,(c3_1(X1)|c1_1(X1)|hskp2|hskp18|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_140])])).
cnf(c_0_779,negated_conjecture,(epred39_0|epred35_0|hskp9|~hskp25), inference(spm,[status(thm)],[c_0_429, c_0_751])).
cnf(c_0_780,negated_conjecture,(epred39_0|epred35_0|hskp9|~hskp5), inference(spm,[status(thm)],[c_0_443, c_0_752])).
cnf(c_0_781,negated_conjecture,(epred50_0|~c2_1(a533)|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_284]), c_0_86])).
cnf(c_0_782,negated_conjecture,(epred26_0|c2_1(a533)|hskp9|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_753, c_0_85]), c_0_86])).
cnf(c_0_783,negated_conjecture,(epred19_0|c0_1(a477)|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_745, c_0_116]), c_0_71])).
cnf(c_0_784,negated_conjecture,(epred35_0|hskp6|hskp9|hskp10), inference(spm,[status(thm)],[c_0_754, c_0_74])).
cnf(c_0_785,negated_conjecture,(epred51_0|epred48_0|~epred29_0), inference(spm,[status(thm)],[c_0_296, c_0_755])).
cnf(c_0_786,negated_conjecture,(epred48_0|epred32_0|epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_756]), c_0_757])).
cnf(c_0_787,negated_conjecture,(epred6_0|c0_1(a474)|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_178]), c_0_118])).
cnf(c_0_788,negated_conjecture,(epred39_0|hskp7|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_443, c_0_728])).
cnf(c_0_789,negated_conjecture,(hskp2|hskp3|hskp12|hskp15), inference(spm,[status(thm)],[c_0_482, c_0_350])).
cnf(c_0_790,negated_conjecture,(hskp3|hskp12|~c0_1(a474)|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_119]), c_0_117])).
cnf(c_0_791,negated_conjecture,(epred7_0|c3_1(X1)|c2_1(X1)|hskp7|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_758]), c_0_140])])).
cnf(c_0_792,negated_conjecture,(hskp7|hskp12|hskp15|hskp11), inference(spm,[status(thm)],[c_0_759, c_0_1])).
cnf(c_0_793,negated_conjecture,(epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_642, c_0_760]), c_0_761])).
cnf(c_0_794,negated_conjecture,(epred42_0|epred39_0|~epred48_0|~epred29_0), inference(spm,[status(thm)],[c_0_762, c_0_688])).
cnf(c_0_795,negated_conjecture,(epred39_0|~epred48_0|~epred42_0), inference(spm,[status(thm)],[c_0_437, c_0_763])).
cnf(c_0_796,negated_conjecture,(epred50_0|epred39_0|epred17_0|epred6_0), inference(spm,[status(thm)],[c_0_632, c_0_764])).
cnf(c_0_797,negated_conjecture,(epred50_0|epred39_0|epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_765]), c_0_766])).
cnf(c_0_798,negated_conjecture,(epred49_0|epred29_0|epred1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_767]), c_0_474])).
cnf(c_0_799,negated_conjecture,(epred29_0|epred2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_768, c_0_769]), c_0_427])).
cnf(c_0_800,negated_conjecture,(epred36_0|~c2_1(a533)|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_770]), c_0_85])).
cnf(c_0_801,negated_conjecture,(epred32_0|~c2_1(a473)|~hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_62]), c_0_63])).
cnf(c_0_802,negated_conjecture,(epred51_0|c2_1(a473)|~hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_64]), c_0_63])).
cnf(c_0_803,negated_conjecture,(epred51_0|epred18_0|hskp2|hskp27), inference(spm,[status(thm)],[c_0_771, c_0_772])).
cnf(c_0_804,negated_conjecture,(epred49_0|epred19_0|~hskp3), inference(spm,[status(thm)],[c_0_773, c_0_156])).
cnf(c_0_805,negated_conjecture,(epred16_0|hskp3|hskp15|hskp16), inference(spm,[status(thm)],[c_0_774, c_0_210])).
cnf(c_0_806,negated_conjecture,(epred51_0|epred47_0|epred17_0|epred8_0), inference(spm,[status(thm)],[c_0_217, c_0_775])).
cnf(c_0_807,negated_conjecture,(epred33_0|c1_1(a502)|hskp26|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_776, c_0_14]), c_0_15])).
cnf(c_0_808,negated_conjecture,(epred32_0|~c0_1(a470)|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_120]), c_0_121])).
cnf(c_0_809,negated_conjecture,(epred17_0|c0_1(a470)|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_120]), c_0_122])).
cnf(c_0_810,negated_conjecture,(epred33_0|c0_1(a503)|hskp26|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_776, c_0_12]), c_0_11])).
cnf(c_0_811,negated_conjecture,(epred48_0|epred32_0|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_123]), c_0_124])).
cnf(c_0_812,negated_conjecture,(epred48_0|epred10_0|hskp28|~hskp23), inference(spm,[status(thm)],[c_0_651, c_0_777])).
cnf(c_0_813,negated_conjecture,(epred19_0|~hskp18), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_745, c_0_112]), c_0_113]), c_0_114])).
cnf(c_0_814,negated_conjecture,(epred6_0|hskp2|hskp18|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_778, c_0_220]), c_0_38]), c_0_36])).
cnf(c_0_815,negated_conjecture,(epred39_0|epred35_0|hskp9|hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_779, c_0_97]), c_0_780])).
cnf(c_0_816,negated_conjecture,(epred50_0|epred26_0|hskp9|~hskp23), inference(spm,[status(thm)],[c_0_781, c_0_782])).
cnf(c_0_817,negated_conjecture,(epred48_0|epred19_0|~hskp10), inference(spm,[status(thm)],[c_0_308, c_0_783])).
cnf(c_0_818,negated_conjecture,(epred50_0|epred35_0|hskp9|hskp10), inference(spm,[status(thm)],[c_0_311, c_0_784])).
cnf(c_0_819,negated_conjecture,(epred51_0|epred48_0|epred32_0), inference(spm,[status(thm)],[c_0_785, c_0_786])).
cnf(c_0_820,negated_conjecture,(epred6_0|hskp16|hskp17|~hskp7), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_787]), c_0_117]), c_0_118])).
cnf(c_0_821,negated_conjecture,(epred39_0|hskp7|hskp12|hskp11), inference(spm,[status(thm)],[c_0_788, c_0_1])).
cnf(c_0_822,negated_conjecture,(epred51_0|epred6_0|hskp14|hskp27), inference(spm,[status(thm)],[c_0_221, c_0_772])).
cnf(c_0_823,negated_conjecture,(epred2_0|hskp3|hskp12|hskp15), inference(spm,[status(thm)],[c_0_578, c_0_789])).
cnf(c_0_824,negated_conjecture,(epred6_0|hskp3|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_790, c_0_787]), c_0_197])).
cnf(c_0_825,negated_conjecture,(epred7_0|hskp7|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_791, c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_826,negated_conjecture,(hskp0|hskp7|hskp12|hskp15), inference(spm,[status(thm)],[c_0_259, c_0_792])).
cnf(c_0_827,negated_conjecture,(epred48_0|epred28_0|hskp2|hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_228]), c_0_502])).
cnf(c_0_828,negated_conjecture,(~epred39_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_335, c_0_793])])).
cnf(c_0_829,negated_conjecture,(epred39_0|~epred48_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_794, c_0_513])]), c_0_795])).
cnf(c_0_830,negated_conjecture,(epred50_0|epred49_0|epred32_0|epred6_0|~hskp20), inference(spm,[status(thm)],[c_0_487, c_0_369])).
cnf(c_0_831,negated_conjecture,(epred50_0|epred39_0|epred17_0|~epred48_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_796]), c_0_797])).
cnf(c_0_832,negated_conjecture,(epred49_0|epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_798]), c_0_799])).
cnf(c_0_833,negated_conjecture,(epred36_0|~c0_1(a503)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_770]), c_0_12])).
cnf(c_0_834,negated_conjecture,(epred36_0|epred12_0|hskp12|~hskp23), inference(spm,[status(thm)],[c_0_800, c_0_652])).
cnf(c_0_835,negated_conjecture,(epred36_0|~hskp9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_770]), c_0_90]), c_0_96])).
cnf(c_0_836,negated_conjecture,(epred47_0|c2_1(a533)|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_174]), c_0_85])).
cnf(c_0_837,negated_conjecture,(epred51_0|epred32_0|~hskp27), inference(spm,[status(thm)],[c_0_801, c_0_802])).
cnf(c_0_838,negated_conjecture,(epred51_0|epred47_0|epred18_0|epred3_0|hskp27), inference(spm,[status(thm)],[c_0_354, c_0_803])).
cnf(c_0_839,negated_conjecture,(epred51_0|epred4_0|hskp27), inference(spm,[status(thm)],[c_0_209, c_0_772])).
cnf(c_0_840,negated_conjecture,(epred49_0|epred19_0|epred16_0|hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_804, c_0_805]), c_0_235])).
cnf(c_0_841,negated_conjecture,(epred51_0|epred47_0|epred17_0|~epred9_0), inference(spm,[status(thm)],[c_0_242, c_0_806])).
cnf(c_0_842,negated_conjecture,(epred33_0|epred17_0|hskp26|~hskp19), inference(spm,[status(thm)],[c_0_343, c_0_807])).
cnf(c_0_843,negated_conjecture,(epred32_0|epred17_0|~hskp26), inference(spm,[status(thm)],[c_0_808, c_0_809])).
cnf(c_0_844,negated_conjecture,(epred33_0|epred32_0|hskp26|~hskp20), inference(spm,[status(thm)],[c_0_331, c_0_810])).
cnf(c_0_845,negated_conjecture,(epred48_0|epred32_0|epred10_0|~hskp23), inference(spm,[status(thm)],[c_0_811, c_0_812])).
cnf(c_0_846,negated_conjecture,(epred32_0|hskp14|hskp5), inference(spm,[status(thm)],[c_0_661, c_0_97])).
cnf(c_0_847,negated_conjecture,(epred19_0|epred6_0|hskp2|~hskp14), inference(spm,[status(thm)],[c_0_813, c_0_814])).
cnf(c_0_848,negated_conjecture,(epred39_0|epred35_0|epred1_0|hskp9), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_490, c_0_815]), c_0_289])).
cnf(c_0_849,negated_conjecture,(epred48_0|~epred39_0), inference(spm,[status(thm)],[c_0_335, c_0_682])).
cnf(c_0_850,negated_conjecture,(epred50_0|epred32_0|epred2_0|~hskp20), inference(spm,[status(thm)],[c_0_578, c_0_369])).
cnf(c_0_851,negated_conjecture,(epred50_0|epred26_0|hskp9|hskp20), inference(spm,[status(thm)],[c_0_816, c_0_92])).
cnf(c_0_852,negated_conjecture,(epred50_0|epred48_0|epred35_0|epred19_0|hskp9), inference(spm,[status(thm)],[c_0_817, c_0_818])).
cnf(c_0_853,negated_conjecture,(epred48_0|epred32_0|~epred50_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_591]), c_0_819])).
cnf(c_0_854,negated_conjecture,(epred6_0|hskp14|hskp16|~hskp7), inference(spm,[status(thm)],[c_0_221, c_0_820])).
cnf(c_0_855,negated_conjecture,(epred12_0|hskp12|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_856,negated_conjecture,(epred39_0|hskp0|hskp7|hskp12), inference(spm,[status(thm)],[c_0_259, c_0_821])).
cnf(c_0_857,negated_conjecture,(epred1_0|c2_1(a494)|hskp27|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_470]), c_0_34])).
cnf(c_0_858,negated_conjecture,(epred51_0|epred12_0|epred6_0|hskp27), inference(spm,[status(thm)],[c_0_341, c_0_822])).
cnf(c_0_859,negated_conjecture,(epred17_0|epred2_0|hskp12|hskp15), inference(spm,[status(thm)],[c_0_191, c_0_823])).
cnf(c_0_860,negated_conjecture,(epred17_0|epred6_0|~hskp7), inference(spm,[status(thm)],[c_0_191, c_0_824])).
cnf(c_0_861,negated_conjecture,(epred7_0|hskp7|hskp12|hskp15), inference(spm,[status(thm)],[c_0_825, c_0_826])).
cnf(c_0_862,negated_conjecture,(epred48_0|hskp2|hskp15), inference(sr,[status(thm)],[c_0_827, c_0_534])).
cnf(c_0_863,negated_conjecture,(~epred48_0), inference(spm,[status(thm)],[c_0_828, c_0_829])).
cnf(c_0_864,negated_conjecture,(epred50_0|epred49_0|epred32_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_830, c_0_386]), c_0_197])).
cnf(c_0_865,negated_conjecture,(epred50_0|epred29_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_831]), c_0_797]), c_0_474])).
cnf(c_0_866,negated_conjecture,(epred49_0|epred39_0|~epred48_0), inference(spm,[status(thm)],[c_0_337, c_0_832])).
cnf(c_0_867,negated_conjecture,(epred51_0|epred49_0|~epred39_0), inference(spm,[status(thm)],[c_0_340, c_0_474])).
cnf(c_0_868,negated_conjecture,(epred36_0|epred25_0|~hskp20), inference(spm,[status(thm)],[c_0_833, c_0_158])).
cnf(c_0_869,negated_conjecture,(epred36_0|epred12_0|hskp12|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_834, c_0_92]), c_0_835])).
cnf(c_0_870,negated_conjecture,(epred47_0|epred36_0|~hskp23), inference(spm,[status(thm)],[c_0_800, c_0_836])).
cnf(c_0_871,negated_conjecture,(epred51_0|epred47_0|epred32_0|epred18_0|epred3_0), inference(spm,[status(thm)],[c_0_837, c_0_838])).
cnf(c_0_872,negated_conjecture,(epred51_0|epred32_0|epred4_0), inference(spm,[status(thm)],[c_0_837, c_0_839])).
cnf(c_0_873,negated_conjecture,(~epred17_0|~epred16_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_125, c_0_718]), c_0_151])).
cnf(c_0_874,negated_conjecture,(epred51_0|epred49_0|epred19_0|epred16_0), inference(spm,[status(thm)],[c_0_256, c_0_840])).
cnf(c_0_875,negated_conjecture,(epred51_0|epred47_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_775]), c_0_841])).
cnf(c_0_876,negated_conjecture,(epred33_0|epred17_0|hskp20|hskp26|~hskp6), inference(spm,[status(thm)],[c_0_842, c_0_160])).
cnf(c_0_877,negated_conjecture,(epred33_0|epred32_0|epred17_0|~hskp20), inference(spm,[status(thm)],[c_0_843, c_0_844])).
cnf(c_0_878,negated_conjecture,(epred48_0|epred32_0|epred10_0|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_845, c_0_92]), c_0_500])).
cnf(c_0_879,negated_conjecture,(epred32_0|hskp3|hskp15|hskp5), inference(spm,[status(thm)],[c_0_482, c_0_846])).
cnf(c_0_880,negated_conjecture,(epred19_0|epred6_0|hskp2|hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_847, c_0_350]), c_0_197])).
cnf(c_0_881,negated_conjecture,(epred48_0|epred35_0|epred32_0|epred1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_848]), c_0_849])).
cnf(c_0_882,negated_conjecture,(epred50_0|epred32_0|epred26_0|epred2_0|hskp9), inference(spm,[status(thm)],[c_0_850, c_0_851])).
cnf(c_0_883,negated_conjecture,(~epred26_0|~epred19_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_126, c_0_722]), c_0_717])).
cnf(c_0_884,negated_conjecture,(epred48_0|epred35_0|epred32_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_852]), c_0_853])).
cnf(c_0_885,negated_conjecture,(epred12_0|c2_1(a475)|hskp12|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_40]), c_0_39])).
cnf(c_0_886,negated_conjecture,(epred12_0|epred6_0|hskp16|~hskp7), inference(spm,[status(thm)],[c_0_341, c_0_854])).
cnf(c_0_887,negated_conjecture,(epred39_0|epred12_0|hskp7|hskp12), inference(spm,[status(thm)],[c_0_855, c_0_856])).
cnf(c_0_888,negated_conjecture,(epred1_0|hskp12|hskp27|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_857]), c_0_490]), c_0_34]), c_0_41])).
cnf(c_0_889,negated_conjecture,(epred51_0|epred39_0|epred12_0|epred6_0), inference(spm,[status(thm)],[c_0_289, c_0_858])).
cnf(c_0_890,negated_conjecture,(epred51_0|epred17_0|epred2_0|hskp12), inference(spm,[status(thm)],[c_0_747, c_0_859])).
cnf(c_0_891,negated_conjecture,(epred17_0|epred7_0|epred6_0|hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_860, c_0_861]), c_0_197])).
cnf(c_0_892,negated_conjecture,(hskp2|hskp15), inference(sr,[status(thm)],[c_0_862, c_0_863])).
cnf(c_0_893,negated_conjecture,(epred50_0|epred49_0|epred32_0|~epred48_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_864]), c_0_865])).
cnf(c_0_894,negated_conjecture,(epred52_0|epred49_0|epred42_0|~epred48_0|~epred29_0), inference(spm,[status(thm)],[c_0_315, c_0_528])).
cnf(c_0_895,negated_conjecture,(epred49_0|~epred52_0|~epred48_0), inference(spm,[status(thm)],[c_0_596, c_0_866])).
cnf(c_0_896,negated_conjecture,(epred49_0|~epred50_0|~epred39_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_466, c_0_474]), c_0_356]), c_0_867])).
cnf(c_0_897,negated_conjecture,(epred36_0|epred25_0|epred12_0|hskp12), inference(spm,[status(thm)],[c_0_868, c_0_869])).
cnf(c_0_898,negated_conjecture,(epred47_0|epred36_0|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_870, c_0_92]), c_0_835])).
cnf(c_0_899,negated_conjecture,(~epred19_0|~epred18_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_127, c_0_716]), c_0_717])).
cnf(c_0_900,negated_conjecture,(epred51_0|epred47_0|epred32_0|epred18_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_871]), c_0_872])).
cnf(c_0_901,negated_conjecture,(epred51_0|epred49_0|epred19_0|~epred17_0), inference(spm,[status(thm)],[c_0_873, c_0_874])).
cnf(c_0_902,negated_conjecture,(epred51_0|epred47_0|~epred49_0|~epred25_0), inference(spm,[status(thm)],[c_0_456, c_0_875])).
cnf(c_0_903,negated_conjecture,(epred33_0|epred32_0|epred17_0|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_843, c_0_876]), c_0_877])).
cnf(c_0_904,negated_conjecture,(~epred11_0|~epred10_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_128, c_0_720]), c_0_188])).
cnf(c_0_905,negated_conjecture,(epred48_0|epred33_0|epred32_0|epred17_0|epred10_0), inference(spm,[status(thm)],[c_0_877, c_0_878])).
cnf(c_0_906,negated_conjecture,(epred32_0|epred11_0|hskp3|hskp15), inference(spm,[status(thm)],[c_0_238, c_0_879])).
cnf(c_0_907,negated_conjecture,(epred49_0|epred19_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_880]), c_0_804])).
cnf(c_0_908,negated_conjecture,(epred41_0|hskp17|~hskp9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_539, c_0_96]), c_0_90]), c_0_91])).
cnf(c_0_909,negated_conjecture,(epred48_0|epred35_0|epred32_0|~epred2_0), inference(spm,[status(thm)],[c_0_555, c_0_881])).
cnf(c_0_910,negated_conjecture,(epred48_0|epred32_0|epred26_0|epred2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_882]), c_0_853])).
cnf(c_0_911,negated_conjecture,(epred48_0|epred35_0|epred32_0|~epred26_0), inference(spm,[status(thm)],[c_0_883, c_0_884])).
cnf(c_0_912,negated_conjecture,(epred48_0|epred12_0|hskp12|~hskp8), inference(spm,[status(thm)],[c_0_203, c_0_885])).
cnf(c_0_913,negated_conjecture,(epred39_0|epred12_0|epred6_0|hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_886, c_0_887]), c_0_197])).
cnf(c_0_914,negated_conjecture,(epred51_0|epred1_0|hskp27|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_692]), c_0_30]), c_0_31])).
cnf(c_0_915,negated_conjecture,(epred51_0|epred1_0|hskp12|hskp27), inference(spm,[status(thm)],[c_0_888, c_0_772])).
cnf(c_0_916,negated_conjecture,(epred51_0|epred13_0|epred6_0|hskp27), inference(spm,[status(thm)],[c_0_360, c_0_822])).
cnf(c_0_917,negated_conjecture,(epred51_0|epred39_0|epred12_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_889]), c_0_703])).
cnf(c_0_918,negated_conjecture,(~epred7_0|~epred2_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_129, c_0_758]), c_0_469])).
cnf(c_0_919,negated_conjecture,(epred51_0|epred17_0|epred6_0|epred2_0), inference(spm,[status(thm)],[c_0_197, c_0_890])).
cnf(c_0_920,negated_conjecture,(epred51_0|epred17_0|epred7_0|epred6_0), inference(spm,[status(thm)],[c_0_747, c_0_891])).
cnf(c_0_921,negated_conjecture,(epred9_0|~c2_1(a494)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_152]), c_0_35])).
cnf(c_0_922,negated_conjecture,(epred49_0|epred1_0|hskp15|hskp27), inference(spm,[status(thm)],[c_0_738, c_0_892])).
cnf(c_0_923,negated_conjecture,(epred48_0|epred28_0|epred2_0|hskp15), inference(spm,[status(thm)],[c_0_578, c_0_827])).
cnf(c_0_924,negated_conjecture,(epred50_0|epred49_0|~epred48_0|~epred42_0), inference(spm,[status(thm)],[c_0_437, c_0_893])).
cnf(c_0_925,negated_conjecture,(epred49_0|epred42_0|~epred48_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_894, c_0_832]), c_0_895])).
cnf(c_0_926,negated_conjecture,(epred49_0|~epred50_0|~epred48_0), inference(spm,[status(thm)],[c_0_896, c_0_866])).
cnf(c_0_927,negated_conjecture,(epred26_0|hskp9|~hskp24), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_753, c_0_27]), c_0_28]), c_0_29])).
cnf(c_0_928,negated_conjecture,(epred36_0|epred25_0|epred12_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_897])).
cnf(c_0_929,negated_conjecture,(epred47_0|epred36_0|epred25_0), inference(spm,[status(thm)],[c_0_868, c_0_898])).
cnf(c_0_930,negated_conjecture,(epred51_0|epred47_0|epred32_0|~epred19_0), inference(spm,[status(thm)],[c_0_899, c_0_900])).
cnf(c_0_931,negated_conjecture,(epred51_0|epred49_0|epred47_0|epred19_0), inference(spm,[status(thm)],[c_0_901, c_0_875])).
cnf(c_0_932,negated_conjecture,(epred51_0|epred47_0|epred32_0|~epred49_0), inference(spm,[status(thm)],[c_0_902, c_0_478])).
cnf(c_0_933,negated_conjecture,(epred33_0|epred32_0|epred17_0|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_903, c_0_33]), c_0_877])).
cnf(c_0_934,negated_conjecture,(epred48_0|epred33_0|epred32_0|epred17_0|~epred11_0), inference(spm,[status(thm)],[c_0_904, c_0_905])).
cnf(c_0_935,negated_conjecture,(epred32_0|epred17_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_906]), c_0_581])).
cnf(c_0_936,negated_conjecture,(epred49_0|epred6_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_804, c_0_684]), c_0_907]), c_0_316]), c_0_613])).
cnf(c_0_937,negated_conjecture,(epred41_0|epred6_0|hskp14|~hskp9), inference(spm,[status(thm)],[c_0_221, c_0_908])).
cnf(c_0_938,negated_conjecture,(epred48_0|epred35_0|epred32_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_909, c_0_910]), c_0_911])).
cnf(c_0_939,negated_conjecture,(epred48_0|epred28_0|epred12_0|hskp12|~hskp16), inference(spm,[status(thm)],[c_0_912, c_0_228])).
cnf(c_0_940,negated_conjecture,(epred12_0|epred6_0|hskp16), inference(sr,[status(thm)],[c_0_913, c_0_828])).
cnf(c_0_941,negated_conjecture,(epred51_0|epred8_0|epred1_0|hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_914, c_0_190]), c_0_915])).
cnf(c_0_942,negated_conjecture,(epred51_0|epred39_0|epred13_0|epred6_0), inference(spm,[status(thm)],[c_0_289, c_0_916])).
cnf(c_0_943,negated_conjecture,(epred51_0|epred39_0|~epred17_0|~epred13_0), inference(spm,[status(thm)],[c_0_391, c_0_917])).
cnf(c_0_944,negated_conjecture,(epred51_0|epred17_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_918, c_0_919]), c_0_920])).
cnf(c_0_945,negated_conjecture,(epred9_0|epred1_0|hskp27|~hskp17), inference(spm,[status(thm)],[c_0_921, c_0_857])).
cnf(c_0_946,negated_conjecture,(epred9_0|epred2_0|~hskp17), inference(spm,[status(thm)],[c_0_921, c_0_674])).
cnf(c_0_947,negated_conjecture,(hskp18|hskp20|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_635, c_0_100]), c_0_101]), c_0_102])).
fof(c_0_948, plain, ((~epred23_0<=>![X1]:((((hskp13|c1_1(X1))|c3_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
cnf(c_0_949,negated_conjecture,(epred41_0|epred25_0|epred4_0|~hskp20), inference(spm,[status(thm)],[c_0_209, c_0_586])).
cnf(c_0_950,negated_conjecture,(epred49_0|epred1_0|hskp15), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_922]), c_0_828])).
cnf(c_0_951,negated_conjecture,(epred48_0|epred2_0|hskp15), inference(sr,[status(thm)],[c_0_923, c_0_534])).
cnf(c_0_952,negated_conjecture,(epred49_0|~epred48_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_924, c_0_925]), c_0_926])).
cnf(c_0_953,negated_conjecture,(c2_1(a533)|hskp3|hskp15|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_461]), c_0_86])).
cnf(c_0_954,negated_conjecture,(epred26_0|hskp9|hskp15|hskp16), inference(spm,[status(thm)],[c_0_927, c_0_42])).
cnf(c_0_955,negated_conjecture,(epred36_0|epred25_0|epred12_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_928]), c_0_929])).
cnf(c_0_956,negated_conjecture,(epred36_0|epred25_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_868, c_0_700]), c_0_835])).
cnf(c_0_957,negated_conjecture,(epred51_0|epred32_0|epred12_0|epred6_0), inference(spm,[status(thm)],[c_0_837, c_0_858])).
cnf(c_0_958,negated_conjecture,(epred51_0|epred47_0|epred32_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_930, c_0_931]), c_0_932])).
cnf(c_0_959,negated_conjecture,(epred45_0|epred39_0|~c2_1(X1)|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_226, c_0_186])).
cnf(c_0_960,negated_conjecture,(epred10_0|c2_1(a494)|hskp28|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_750]), c_0_34])).
cnf(c_0_961,negated_conjecture,(epred33_0|epred32_0|epred17_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_933])).
cnf(c_0_962,negated_conjecture,(epred48_0|epred33_0|epred32_0|epred17_0), inference(spm,[status(thm)],[c_0_934, c_0_935])).
cnf(c_0_963,negated_conjecture,(epred49_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_936]), c_0_458])).
cnf(c_0_964,negated_conjecture,(epred41_0|epred13_0|epred6_0|~hskp9), inference(spm,[status(thm)],[c_0_360, c_0_937])).
cnf(c_0_965,negated_conjecture,(epred39_0|epred35_0|epred13_0|epred6_0|hskp9), inference(spm,[status(thm)],[c_0_360, c_0_815])).
cnf(c_0_966,negated_conjecture,(epred35_0|epred32_0), inference(sr,[status(thm)],[c_0_938, c_0_863])).
cnf(c_0_967,negated_conjecture,(epred12_0|epred6_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_939, c_0_940]), c_0_863]), c_0_534]), c_0_197])).
cnf(c_0_968,negated_conjecture,(epred51_0|epred39_0|epred8_0|epred1_0), inference(spm,[status(thm)],[c_0_289, c_0_941])).
cnf(c_0_969,negated_conjecture,(epred51_0|epred39_0|~epred17_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_942]), c_0_703]), c_0_943])).
cnf(c_0_970,negated_conjecture,(epred51_0|epred17_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_944]), c_0_785])).
cnf(c_0_971,negated_conjecture,(epred51_0|epred9_0|epred1_0|hskp27), inference(spm,[status(thm)],[c_0_945, c_0_772])).
cnf(c_0_972,negated_conjecture,(epred51_0|epred9_0|epred2_0|hskp27), inference(spm,[status(thm)],[c_0_946, c_0_772])).
cnf(c_0_973,negated_conjecture,(epred32_0|epred25_0|~epred47_0|~epred17_0), inference(spm,[status(thm)],[c_0_263, c_0_569])).
cnf(c_0_974,negated_conjecture,(hskp17|hskp18|hskp20|hskp27), inference(spm,[status(thm)],[c_0_947, c_0_107])).
cnf(c_0_975,negated_conjecture,(epred23_0|c3_1(X1)|c1_1(X1)|hskp13|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_948]), c_0_140])])).
cnf(c_0_976,negated_conjecture,(epred41_0|epred25_0|epred4_0|hskp12), inference(spm,[status(thm)],[c_0_949, c_0_195])).
cnf(c_0_977,negated_conjecture,(epred49_0|epred1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_950]), c_0_828])).
cnf(c_0_978,negated_conjecture,(epred49_0|epred2_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_951]), c_0_613]), c_0_952])).
cnf(c_0_979,negated_conjecture,(epred36_0|hskp3|hskp15|~hskp23), inference(spm,[status(thm)],[c_0_800, c_0_953])).
cnf(c_0_980,negated_conjecture,(epred49_0|epred26_0|hskp9|hskp15), inference(spm,[status(thm)],[c_0_235, c_0_954])).
cnf(c_0_981,negated_conjecture,(epred36_0|epred25_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_955]), c_0_956])).
cnf(c_0_982,negated_conjecture,(epred39_0|epred36_0|epred29_0|epred5_0), inference(spm,[status(thm)],[c_0_835, c_0_473])).
cnf(c_0_983,negated_conjecture,(epred36_0|epred29_0|epred12_0|hskp12), inference(spm,[status(thm)],[c_0_230, c_0_869])).
cnf(c_0_984,negated_conjecture,(epred51_0|epred32_0|epred12_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_957]), c_0_958])).
cnf(c_0_985,negated_conjecture,(epred39_0|~c0_1(a470)|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_120]), c_0_122])).
cnf(c_0_986,negated_conjecture,(c0_1(a470)|hskp10|hskp27|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_648, c_0_122]), c_0_121])).
cnf(c_0_987,negated_conjecture,(epred45_0|~hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_959, c_0_130]), c_0_828]), c_0_124]), c_0_123])).
cnf(c_0_988,negated_conjecture,(epred33_0|epred10_0|hskp26|hskp28|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_776, c_0_960]), c_0_35]), c_0_34])).
cnf(c_0_989,negated_conjecture,(epred33_0|epred32_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_680, c_0_961]), c_0_962])).
cnf(c_0_990,negated_conjecture,(epred49_0|epred33_0|epred32_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_963, c_0_962]), c_0_952])).
cnf(c_0_991,negated_conjecture,(epred39_0|epred35_0|hskp3|hskp9|hskp15), inference(spm,[status(thm)],[c_0_482, c_0_815])).
cnf(c_0_992,negated_conjecture,(epred41_0|epred39_0|epred35_0|epred13_0|epred6_0), inference(spm,[status(thm)],[c_0_964, c_0_965])).
cnf(c_0_993,negated_conjecture,(epred35_0|~epred41_0), inference(spm,[status(thm)],[c_0_642, c_0_966])).
cnf(c_0_994,negated_conjecture,(epred12_0|~epred17_0), inference(spm,[status(thm)],[c_0_624, c_0_967])).
cnf(c_0_995,negated_conjecture,(epred51_0|epred39_0|epred8_0|~epred2_0), inference(spm,[status(thm)],[c_0_555, c_0_968])).
cnf(c_0_996,negated_conjecture,(epred51_0|epred39_0|~epred29_0), inference(spm,[status(thm)],[c_0_969, c_0_970])).
cnf(c_0_997,negated_conjecture,(epred51_0|epred39_0|epred9_0|epred1_0), inference(spm,[status(thm)],[c_0_289, c_0_971])).
cnf(c_0_998,negated_conjecture,(epred51_0|epred39_0|epred9_0|epred2_0), inference(spm,[status(thm)],[c_0_289, c_0_972])).
cnf(c_0_999,negated_conjecture,(epred51_0|~epred49_0|~epred29_0|~epred25_0), inference(spm,[status(thm)],[c_0_456, c_0_970])).
cnf(c_0_1000,negated_conjecture,(epred38_0|epred32_0|epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_973, c_0_287]), c_0_478])).
cnf(c_0_1001,negated_conjecture,(epred19_0|hskp17|hskp20|hskp27), inference(spm,[status(thm)],[c_0_813, c_0_974])).
cnf(c_0_1002,negated_conjecture,(epred48_0|~hskp13), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_147]), c_0_132]), c_0_133])).
cnf(c_0_1003,negated_conjecture,(epred23_0|hskp13|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_975, c_0_85]), c_0_86]), c_0_87])).
cnf(c_0_1004,negated_conjecture,(epred46_0|c2_1(a494)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_232]), c_0_34])).
cnf(c_0_1005,negated_conjecture,(epred41_0|epred25_0|epred6_0|epred4_0), inference(spm,[status(thm)],[c_0_197, c_0_976])).
cnf(c_0_1006,negated_conjecture,(epred49_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_977]), c_0_978])).
cnf(c_0_1007,negated_conjecture,(epred36_0|hskp3|hskp15|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_979, c_0_92]), c_0_835])).
cnf(c_0_1008,negated_conjecture,(epred51_0|epred49_0|epred19_0|~epred29_0), inference(spm,[status(thm)],[c_0_901, c_0_970])).
cnf(c_0_1009,negated_conjecture,(epred49_0|epred36_0|epred26_0|hskp15), inference(spm,[status(thm)],[c_0_835, c_0_980])).
cnf(c_0_1010,negated_conjecture,(epred51_0|epred36_0|epred25_0|~epred29_0), inference(spm,[status(thm)],[c_0_981, c_0_970])).
cnf(c_0_1011,negated_conjecture,(epred39_0|epred36_0|epred29_0|~epred6_0), inference(spm,[status(thm)],[c_0_492, c_0_982])).
cnf(c_0_1012,negated_conjecture,(epred36_0|epred29_0|epred12_0|epred6_0), inference(spm,[status(thm)],[c_0_197, c_0_983])).
cnf(c_0_1013,negated_conjecture,(epred51_0|epred32_0|epred13_0|epred6_0), inference(spm,[status(thm)],[c_0_837, c_0_916])).
cnf(c_0_1014,negated_conjecture,(epred51_0|epred32_0|~epred17_0|~epred13_0), inference(spm,[status(thm)],[c_0_391, c_0_984])).
cnf(c_0_1015,negated_conjecture,(epred50_0|epred1_0|hskp27|~hskp17), inference(spm,[status(thm)],[c_0_673, c_0_857])).
cnf(c_0_1016,negated_conjecture,(epred42_0|c0_1(a492)|hskp14|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_54]), c_0_48])).
cnf(c_0_1017,negated_conjecture,(epred39_0|hskp10|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_985, c_0_986]), c_0_289])).
cnf(c_0_1018,negated_conjecture,(epred45_0|epred33_0|epred10_0|hskp26|~hskp17), inference(spm,[status(thm)],[c_0_987, c_0_988])).
cnf(c_0_1019,negated_conjecture,(epred45_0|~hskp10), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_186]), c_0_116]), c_0_71])).
cnf(c_0_1020,negated_conjecture,(epred33_0|epred32_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_989]), c_0_628]), c_0_990])).
cnf(c_0_1021,negated_conjecture,(epred39_0|epred35_0|epred17_0|hskp9), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_991]), c_0_504])).
cnf(c_0_1022,negated_conjecture,(epred35_0|epred13_0|epred6_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_992, c_0_828]), c_0_993])).
cnf(c_0_1023,negated_conjecture,(~epred17_0|~epred13_0), inference(spm,[status(thm)],[c_0_391, c_0_994])).
cnf(c_0_1024,negated_conjecture,(epred45_0|~c1_1(a478)|~hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_186]), c_0_56])).
cnf(c_0_1025,negated_conjecture,(epred33_0|c1_1(a478)|hskp26|~hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_776, c_0_56]), c_0_57])).
cnf(c_0_1026,negated_conjecture,(epred51_0|epred39_0|epred8_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_995, c_0_799]), c_0_996])).
cnf(c_0_1027,negated_conjecture,(epred51_0|epred39_0|epred9_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_997]), c_0_998])).
cnf(c_0_1028,negated_conjecture,(epred51_0|epred38_0|epred32_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1000]), c_0_446]), c_0_467])).
cnf(c_0_1029,negated_conjecture,(epred19_0|epred4_0|hskp20|hskp27), inference(spm,[status(thm)],[c_0_209, c_0_1001])).
cnf(c_0_1030,negated_conjecture,(epred48_0|epred23_0|~hskp23), inference(spm,[status(thm)],[c_0_1002, c_0_1003])).
cnf(c_0_1031,negated_conjecture,(epred39_0|epred17_0|~hskp26), inference(spm,[status(thm)],[c_0_985, c_0_809])).
cnf(c_0_1032,negated_conjecture,(epred46_0|epred33_0|hskp26|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_776, c_0_1004]), c_0_35]), c_0_34])).
cnf(c_0_1033,negated_conjecture,(epred41_0|epred25_0|epred6_0|~epred46_0|~epred45_0), inference(spm,[status(thm)],[c_0_278, c_0_1005])).
cnf(c_0_1034,negated_conjecture,(~epred25_0|~epred17_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_456, c_0_1006])])).
cnf(c_0_1035,negated_conjecture,(epred19_0|epred17_0|hskp20|~hskp15), inference(spm,[status(thm)],[c_0_813, c_0_663])).
cnf(c_0_1036,negated_conjecture,(epred36_0|epred17_0|hskp15|hskp20), inference(spm,[status(thm)],[c_0_191, c_0_1007])).
cnf(c_0_1037,negated_conjecture,(epred36_0|epred26_0|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_800, c_0_782]), c_0_835])).
cnf(c_0_1038,negated_conjecture,(epred51_0|epred49_0|~epred29_0|~epred26_0), inference(spm,[status(thm)],[c_0_883, c_0_1008])).
cnf(c_0_1039,negated_conjecture,(epred51_0|epred49_0|epred36_0|epred26_0), inference(spm,[status(thm)],[c_0_256, c_0_1009])).
cnf(c_0_1040,negated_conjecture,(epred51_0|epred36_0|~epred49_0|~epred29_0), inference(spm,[status(thm)],[c_0_999, c_0_1010])).
cnf(c_0_1041,negated_conjecture,(epred39_0|epred36_0|epred29_0|epred12_0), inference(spm,[status(thm)],[c_0_1011, c_0_1012])).
cnf(c_0_1042,negated_conjecture,(epred36_0|epred29_0|epred13_0), inference(spm,[status(thm)],[c_0_835, c_0_727])).
cnf(c_0_1043,negated_conjecture,(epred51_0|epred32_0|epred8_0|epred1_0), inference(spm,[status(thm)],[c_0_837, c_0_941])).
cnf(c_0_1044,negated_conjecture,(epred51_0|epred32_0|~epred17_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_1013]), c_0_958]), c_0_1014])).
cnf(c_0_1045,negated_conjecture,(epred50_0|epred41_0|epred1_0|hskp27|~hskp9), inference(spm,[status(thm)],[c_0_1015, c_0_908])).
cnf(c_0_1046,negated_conjecture,(epred4_0|hskp16|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_201]), c_0_209]), c_0_38]), c_0_37])).
cnf(c_0_1047,negated_conjecture,(epred42_0|epred39_0|hskp14|~hskp15), inference(spm,[status(thm)],[c_0_249, c_0_1016])).
cnf(c_0_1048,negated_conjecture,(epred45_0|epred33_0|epred10_0|~hskp17), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1017, c_0_1018]), c_0_828]), c_0_1019])).
cnf(c_0_1049,negated_conjecture,(epred33_0|~epred41_0), inference(spm,[status(thm)],[c_0_642, c_0_1020])).
cnf(c_0_1050,negated_conjecture,(epred35_0|hskp3|hskp9|hskp15), inference(sr,[status(thm)],[c_0_991, c_0_828])).
cnf(c_0_1051,negated_conjecture,(epred35_0|epred17_0|hskp9), inference(sr,[status(thm)],[c_0_1021, c_0_828])).
cnf(c_0_1052,negated_conjecture,(epred35_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_1022]), c_0_1023])).
cnf(c_0_1053,negated_conjecture,(~epred33_0|~epred32_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_134, c_0_749]), c_0_283])).
cnf(c_0_1054,negated_conjecture,(epred45_0|epred33_0|hskp26|~hskp11), inference(spm,[status(thm)],[c_0_1024, c_0_1025])).
cnf(c_0_1055,negated_conjecture,(epred51_0|epred39_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_1026]), c_0_1027])).
cnf(c_0_1056,negated_conjecture,(epred25_0|epred15_0|~epred48_0|~epred29_0), inference(spm,[status(thm)],[c_0_315, c_0_607])).
cnf(c_0_1057,negated_conjecture,(epred51_0|epred38_0|~epred41_0), inference(spm,[status(thm)],[c_0_642, c_0_1028])).
cnf(c_0_1058,negated_conjecture,(epred41_0|epred25_0|epred19_0|epred4_0|hskp27), inference(spm,[status(thm)],[c_0_949, c_0_1029])).
cnf(c_0_1059,negated_conjecture,(epred48_0|epred23_0|hskp9|hskp20), inference(spm,[status(thm)],[c_0_1030, c_0_92])).
cnf(c_0_1060,negated_conjecture,(epred46_0|epred33_0|epred17_0|~hskp17), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1031, c_0_1032]), c_0_828])).
cnf(c_0_1061,negated_conjecture,(epred41_0|~epred17_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_1033]), c_0_656]), c_0_656]), c_0_1034])).
cnf(c_0_1062,negated_conjecture,(epred36_0|epred25_0|epred19_0|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_868, c_0_1035]), c_0_981])).
cnf(c_0_1063,negated_conjecture,(epred36_0|epred25_0|hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_868, c_0_1036]), c_0_981])).
cnf(c_0_1064,negated_conjecture,(epred36_0|epred26_0|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1037, c_0_92]), c_0_835])).
cnf(c_0_1065,negated_conjecture,(epred51_0|epred36_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1038, c_0_1039]), c_0_1040])).
cnf(c_0_1066,negated_conjecture,(epred39_0|epred36_0|epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_1041]), c_0_1042])).
cnf(c_0_1067,negated_conjecture,(epred51_0|epred32_0|epred8_0|~epred2_0), inference(spm,[status(thm)],[c_0_555, c_0_1043])).
cnf(c_0_1068,negated_conjecture,(epred51_0|epred32_0|~epred29_0), inference(spm,[status(thm)],[c_0_1044, c_0_970])).
cnf(c_0_1069,negated_conjecture,(epred51_0|epred32_0|epred9_0|epred1_0), inference(spm,[status(thm)],[c_0_837, c_0_971])).
cnf(c_0_1070,negated_conjecture,(epred51_0|epred32_0|epred9_0|epred2_0), inference(spm,[status(thm)],[c_0_837, c_0_972])).
cnf(c_0_1071,negated_conjecture,(epred50_0|epred41_0|epred39_0|epred1_0|~hskp9), inference(spm,[status(thm)],[c_0_289, c_0_1045])).
cnf(c_0_1072,negated_conjecture,(epred35_0|epred1_0|hskp9), inference(sr,[status(thm)],[c_0_848, c_0_828])).
cnf(c_0_1073,negated_conjecture,(epred42_0|epred39_0|epred4_0|hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1046, c_0_1047]), c_0_236])).
cnf(c_0_1074,negated_conjecture,(epred39_0|epred17_0|epred2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_504, c_0_951]), c_0_829])).
cnf(c_0_1075,negated_conjecture,(epred45_0|epred33_0|epred10_0|~hskp9), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1048, c_0_908]), c_0_1049])).
cnf(c_0_1076,negated_conjecture,(epred35_0|hskp9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_1050]), c_0_1051]), c_0_1052])).
cnf(c_0_1077,negated_conjecture,(epred35_0|~epred33_0), inference(spm,[status(thm)],[c_0_1053, c_0_966])).
cnf(c_0_1078,negated_conjecture,(epred45_0|epred39_0|epred33_0|~hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1017, c_0_1054]), c_0_1019])).
cnf(c_0_1079,negated_conjecture,(epred51_0|epred49_0), inference(spm,[status(thm)],[c_0_867, c_0_1055])).
cnf(c_0_1080,negated_conjecture,(epred41_0|epred25_0|epred14_0|~epred48_0), inference(spm,[status(thm)],[c_0_680, c_0_625])).
cnf(c_0_1081,negated_conjecture,(epred25_0|epred15_0|~epred48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1056, c_0_513])])).
cnf(c_0_1082,negated_conjecture,(epred51_0|epred48_0|~epred39_0), inference(spm,[status(thm)],[c_0_785, c_0_474])).
cnf(c_0_1083,negated_conjecture,(epred51_0|~epred41_0|~epred39_0), inference(spm,[status(thm)],[c_0_335, c_0_1057])).
cnf(c_0_1084,negated_conjecture,(epred41_0|epred39_0|epred25_0|epred19_0|epred4_0), inference(spm,[status(thm)],[c_0_289, c_0_1058])).
cnf(c_0_1085,negated_conjecture,(epred23_0|hskp9|hskp20), inference(sr,[status(thm)],[c_0_1059, c_0_863])).
cnf(c_0_1086,negated_conjecture,(epred41_0|epred4_0|~hskp9), inference(spm,[status(thm)],[c_0_209, c_0_908])).
cnf(c_0_1087,negated_conjecture,(epred46_0|epred33_0|~hskp9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1060, c_0_908]), c_0_1061]), c_0_1049])).
cnf(c_0_1088,negated_conjecture,(epred36_0|epred25_0|epred19_0), inference(spm,[status(thm)],[c_0_1062, c_0_1063])).
cnf(c_0_1089,negated_conjecture,(epred36_0|epred26_0|epred25_0), inference(spm,[status(thm)],[c_0_868, c_0_1064])).
cnf(c_0_1090,negated_conjecture,(epred51_0|epred39_0|epred36_0), inference(spm,[status(thm)],[c_0_1065, c_0_1066])).
cnf(c_0_1091,negated_conjecture,(epred51_0|epred32_0|epred8_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1067, c_0_799]), c_0_1068])).
cnf(c_0_1092,negated_conjecture,(epred51_0|epred32_0|epred9_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_1069]), c_0_1070])).
cnf(c_0_1093,negated_conjecture,(epred50_0|epred35_0|epred1_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1071, c_0_1072]), c_0_828]), c_0_993])).
cnf(c_0_1094,negated_conjecture,(epred42_0|epred39_0|epred2_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_768, c_0_1073]), c_0_1074]), c_0_762])).
cnf(c_0_1095,negated_conjecture,(epred45_0|epred35_0|epred10_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1075, c_0_1076]), c_0_1077])).
cnf(c_0_1096,negated_conjecture,(epred45_0|epred39_0|epred33_0|epred11_0), inference(spm,[status(thm)],[c_0_1078, c_0_260])).
cnf(c_0_1097,negated_conjecture,(epred51_0|~epred25_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_999, c_0_513])]), c_0_1079])).
cnf(c_0_1098,negated_conjecture,(epred41_0|epred25_0|~epred48_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_639, c_0_1080]), c_0_1081])).
cnf(c_0_1099,negated_conjecture,(epred51_0|epred48_0), inference(spm,[status(thm)],[c_0_1082, c_0_1055])).
cnf(c_0_1100,negated_conjecture,(epred51_0|~epred41_0), inference(spm,[status(thm)],[c_0_1083, c_0_1055])).
cnf(c_0_1101,negated_conjecture,(epred41_0|epred25_0|epred19_0|epred4_0), inference(sr,[status(thm)],[c_0_1084, c_0_828])).
cnf(c_0_1102,negated_conjecture,(epred41_0|epred25_0|epred23_0|epred4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_949, c_0_1085]), c_0_1086])).
cnf(c_0_1103,negated_conjecture,(~epred36_0|~epred35_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_135, c_0_696]), c_0_742])).
cnf(c_0_1104,negated_conjecture,(epred46_0|epred35_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1087, c_0_1076]), c_0_1077])).
cnf(c_0_1105,negated_conjecture,(epred36_0|epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_883, c_0_1088]), c_0_1089])).
cnf(c_0_1106,negated_conjecture,(epred51_0|epred36_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1065, c_0_474]), c_0_1090])).
cnf(c_0_1107,negated_conjecture,(epred32_0|~epred51_0|~epred50_0|~epred29_0), inference(spm,[status(thm)],[c_0_313, c_0_609])).
cnf(c_0_1108,negated_conjecture,(epred51_0|epred32_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_1091]), c_0_1092])).
cnf(c_0_1109,negated_conjecture,(epred36_0|hskp2|hskp6|~hskp20), inference(spm,[status(thm)],[c_0_833, c_0_306])).
cnf(c_0_1110,negated_conjecture,(epred50_0|epred35_0|~epred2_0), inference(spm,[status(thm)],[c_0_555, c_0_1093])).
cnf(c_0_1111,negated_conjecture,(epred42_0|epred2_0), inference(sr,[status(thm)],[c_0_1094, c_0_828])).
cnf(c_0_1112,negated_conjecture,(epred35_0|~epred42_0), inference(spm,[status(thm)],[c_0_437, c_0_966])).
cnf(c_0_1113,negated_conjecture,(epred45_0|epred35_0|~epred11_0), inference(spm,[status(thm)],[c_0_904, c_0_1095])).
cnf(c_0_1114,negated_conjecture,(epred45_0|epred33_0|epred11_0), inference(sr,[status(thm)],[c_0_1096, c_0_828])).
cnf(c_0_1115,negated_conjecture,(epred51_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1097, c_0_1098]), c_0_1099]), c_0_1100])).
cnf(c_0_1116,negated_conjecture,(~epred23_0|~epred19_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_136, c_0_948]), c_0_717])).
cnf(c_0_1117,negated_conjecture,(epred41_0|epred25_0|epred19_0|~epred46_0|~epred45_0), inference(spm,[status(thm)],[c_0_278, c_0_1101])).
cnf(c_0_1118,negated_conjecture,(epred41_0|epred25_0|epred23_0|~epred46_0|~epred45_0), inference(spm,[status(thm)],[c_0_278, c_0_1102])).
cnf(c_0_1119,negated_conjecture,(epred46_0|~epred36_0), inference(spm,[status(thm)],[c_0_1103, c_0_1104])).
cnf(c_0_1120,negated_conjecture,(epred36_0|~epred50_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_1105]), c_0_1106])).
cnf(c_0_1121,negated_conjecture,(epred32_0|~epred50_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1107, c_0_734]), c_0_1108])).
cnf(c_0_1122,negated_conjecture,(epred50_0|epred36_0|hskp2|~hskp20), inference(spm,[status(thm)],[c_0_311, c_0_1109])).
cnf(c_0_1123,negated_conjecture,(epred50_0|epred35_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1110, c_0_1111]), c_0_1112])).
cnf(c_0_1124,negated_conjecture,(epred45_0|epred35_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1113, c_0_1114]), c_0_1077])).
cnf(c_0_1125,negated_conjecture,(~epred50_0|~epred25_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_313, c_0_1115])])).
cnf(c_0_1126,negated_conjecture,(epred41_0|epred25_0|~epred46_0|~epred45_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1116, c_0_1117]), c_0_1118])).
cnf(c_0_1127,negated_conjecture,(epred46_0|~epred50_0), inference(spm,[status(thm)],[c_0_1119, c_0_1120])).
cnf(c_0_1128,negated_conjecture,(~epred50_0|~epred41_0), inference(spm,[status(thm)],[c_0_642, c_0_1121])).
cnf(c_0_1129,negated_conjecture,(epred50_0|epred36_0|epred2_0|~hskp20), inference(spm,[status(thm)],[c_0_578, c_0_1122])).
cnf(c_0_1130,negated_conjecture,(epred50_0|~epred36_0), inference(spm,[status(thm)],[c_0_1103, c_0_1123])).
cnf(c_0_1131,negated_conjecture,(epred45_0|~epred36_0), inference(spm,[status(thm)],[c_0_1103, c_0_1124])).
cnf(c_0_1132,negated_conjecture,(~epred50_0|~epred45_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1125, c_0_1126]), c_0_1127]), c_0_1128])).
cnf(c_0_1133,negated_conjecture,(epred50_0|epred25_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1129, c_0_606]), c_0_1105]), c_0_1105]), c_0_1130])).
cnf(c_0_1134,negated_conjecture,(~epred50_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1131, c_0_1120]), c_0_1132])).
cnf(c_0_1135,negated_conjecture,(epred25_0), inference(sr,[status(thm)],[c_0_1133, c_0_1134])).
cnf(c_0_1136,negated_conjecture,(epred17_0|epred2_0), inference(sr,[status(thm)],[c_0_1074, c_0_828])).
cnf(c_0_1137,negated_conjecture,(~epred17_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1034, c_0_1135])])).
cnf(c_0_1138,negated_conjecture,(epred39_0|epred1_0|c3_1(X1)|c2_1(X1)|c1_1(X1)), inference(spm,[status(thm)],[c_0_289, c_0_470])).
cnf(c_0_1139,negated_conjecture,(epred2_0), inference(sr,[status(thm)],[c_0_1136, c_0_1137])).
fof(c_0_1140, plain, ((~epred24_0<=>![X2]:((((hskp9|c0_1(X2))|c2_1(X2))|~ndr1_0)|~c3_1(X2)))), introduced(definition)).
cnf(c_0_1141,negated_conjecture,(epred48_0|epred39_0|~c2_1(X1)|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_226, c_0_147])).
cnf(c_0_1142,negated_conjecture,(epred35_0), inference(sr,[status(thm)],[c_0_1123, c_0_1134])).
cnf(c_0_1143,negated_conjecture,(epred17_0|hskp12), inference(sr,[status(thm)],[c_0_524, c_0_828])).
cnf(c_0_1144,negated_conjecture,(epred1_0|c3_1(X1)|c2_1(X1)|c1_1(X1)), inference(sr,[status(thm)],[c_0_1138, c_0_828])).
cnf(c_0_1145,negated_conjecture,(~epred1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_555, c_0_1139])])).
cnf(c_0_1146,negated_conjecture,(~epred25_0|~epred24_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_137, c_0_1140]), c_0_138])).
cnf(c_0_1147,negated_conjecture,(epred48_0|epred39_0|epred36_0|~c2_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_1141, c_0_770])).
cnf(c_0_1148,negated_conjecture,(~epred36_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1103, c_0_1142])])).
cnf(c_0_1149,negated_conjecture,(hskp12), inference(sr,[status(thm)],[c_0_1143, c_0_1137])).
cnf(c_0_1150,negated_conjecture,(epred24_0|c2_1(X1)|c0_1(X1)|hskp9|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_1140]), c_0_140])])).
cnf(c_0_1151,negated_conjecture,(c3_1(X1)|c2_1(X1)|c1_1(X1)), inference(sr,[status(thm)],[c_0_1144, c_0_1145])).
cnf(c_0_1152,negated_conjecture,(~epred24_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1146, c_0_1135])])).
cnf(c_0_1153,negated_conjecture,(~hskp9), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1147, c_0_96]), c_0_863]), c_0_828]), c_0_1148]), c_0_90])).
cnf(c_0_1154,negated_conjecture,(~c1_1(a480)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_1149])])).
cnf(c_0_1155,negated_conjecture,(c2_1(X1)|c1_1(X1)|c0_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1150, c_0_1151]), c_0_1152]), c_0_1153])).
cnf(c_0_1156,negated_conjecture,(~c2_1(a480)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_1149])])).
cnf(c_0_1157,negated_conjecture,(~c0_1(a480)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_1149])])).
cnf(c_0_1158,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1154, c_0_1155]), c_0_1156]), c_0_1157]), ['proof']).
# SZS output end CNFRefutation.
EOF
