% Problem    : SYN862-1 : TPTP v6.1.0. Released v2.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n135.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:15:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation.
cnf(c_0_0,negated_conjecture,(ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause2)).
cnf(c_0_1,negated_conjecture,(ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1)).
cnf(c_0_2,negated_conjecture,(ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause3)).
cnf(c_0_3,negated_conjecture,(ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause12)).
cnf(c_0_4,negated_conjecture,(ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause20)).
cnf(c_0_5,negated_conjecture,(ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause32)).
cnf(c_0_6,negated_conjecture,(ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause48)).
cnf(c_0_7,negated_conjecture,(ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause63)).
cnf(c_0_8,negated_conjecture,(ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause80)).
cnf(c_0_9,negated_conjecture,(ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause100)).
cnf(c_0_10,negated_conjecture,(ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause126)).
cnf(c_0_11,negated_conjecture,(ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause151)).
cnf(c_0_12,negated_conjecture,(ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause179)).
cnf(c_0_13,negated_conjecture,(ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause215)).
cnf(c_0_14,negated_conjecture,(ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause244)).
cnf(c_0_15,negated_conjecture,(ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause279)).
cnf(c_0_16,negated_conjecture,(ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause317)).
cnf(c_0_17,negated_conjecture,(ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause355)).
cnf(c_0_18,negated_conjecture,(ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause391)).
cnf(c_0_19,negated_conjecture,(ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause429)).
cnf(c_0_20,negated_conjecture,(ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause473)).
cnf(c_0_21,negated_conjecture,(ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause514)).
cnf(c_0_22,negated_conjecture,(ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause558)).
cnf(c_0_23,negated_conjecture,(ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause606)).
cnf(c_0_24,negated_conjecture,(ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause653)).
cnf(c_0_25,negated_conjecture,(ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause705)).
cnf(c_0_26,negated_conjecture,(ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause760)).
cnf(c_0_27,negated_conjecture,(ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause817)).
cnf(c_0_28,negated_conjecture,(ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause882)).
cnf(c_0_29,negated_conjecture,(ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause946)).
cnf(c_0_30,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv24_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv18_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv26_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1131)).
cnf(c_0_31,negated_conjecture,(ssPv18_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause996)).
cnf(c_0_32,negated_conjecture,(ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1007)).
cnf(c_0_33,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv24_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause981)).
cnf(c_0_34,negated_conjecture,(ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1075)).
cnf(c_0_35,negated_conjecture,(~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,skc39)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause884)).
cnf(c_0_36,negated_conjecture,(ssPv26_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause978)).
cnf(c_0_37,negated_conjecture,(ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|ssPv18_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|ssPv21_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1347)).
cnf(c_0_38,negated_conjecture,(ssPv15_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1204)).
cnf(c_0_39,negated_conjecture,(ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv18_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1194)).
cnf(c_0_40,negated_conjecture,(ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv27_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv28_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv21_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause743)).
cnf(c_0_41,negated_conjecture,(ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv21_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1189)).
cnf(c_0_42,negated_conjecture,(ssPv29_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv40_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv24_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1127)).
cnf(c_0_43,negated_conjecture,(ssPv24_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause982)).
cnf(c_0_44,negated_conjecture,(ssPv21_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause647)).
cnf(c_0_45,negated_conjecture,(~ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,skc33)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1077)).
cnf(c_0_46,negated_conjecture,(ssPv40_1r1(X1)|~ssPv40_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause949)).
cnf(c_0_47,negated_conjecture,(ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssPv29_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause972)).
cnf(c_0_48,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv27_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause634)).
cnf(c_0_49,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,skc44)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause706)).
cnf(c_0_50,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv28_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause631)).
cnf(c_0_51,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,skc38)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause883)).
cnf(c_0_52,negated_conjecture,(ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssPv29_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause275)).
cnf(c_0_53,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv29_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv39_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv40_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause334)).
cnf(c_0_54,negated_conjecture,(ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1149)).
cnf(c_0_55,negated_conjecture,(ssPv39_2r1r1(X1,X2)|~ssPv39_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause250)).
cnf(c_0_56,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,skc50)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause245)).
cnf(c_0_57,negated_conjecture,(ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1219)).
cnf(c_0_58,negated_conjecture,(ssPv40_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssPv40_1r1(X1)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause247)).
cnf(c_0_59,negated_conjecture,(~ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,skc53)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause217)).
cnf(c_0_60,negated_conjecture,(ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1292)).
cnf(c_0_61,negated_conjecture,(~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,skc55)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause181)).
cnf(c_0_62,negated_conjecture,(ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1363)).
cnf(c_0_63,negated_conjecture,(~ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,skc57)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause153)).
cnf(c_0_64,negated_conjecture,(ssNder1_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1433)).
cnf(c_0_65,negated_conjecture,(ssPv40_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssNder1_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssPv40_1r1(X1)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1508)).
cnf(c_0_66,negated_conjecture,(ssPv39_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssNder1_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssPv39_2r1r1(X1,X2)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1510)).
cnf(c_0_67,negated_conjecture,(ssPv31_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv4_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv40_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssNder1_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1733)).
cnf(c_0_68,negated_conjecture,(ssPv4_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|ssPv12_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv39_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssNder1_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1729)).
cnf(c_0_69,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv12_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssNder1_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1564)).
cnf(c_0_70,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv31_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssNder1_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1526)).
cnf(c_0_71,negated_conjecture,(ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssNder1_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)|~ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)|~ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)|~ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)|~ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)|~ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)|~ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)|~ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)|~ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)|~ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)|~ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)|~ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)|~ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)|~ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)|~ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)|~ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)|~ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)|~ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause1524)).
cnf(c_0_72,negated_conjecture,(~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,skc61)|~ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause102)).
cnf(c_0_73,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc62)|~ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)|~ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)|~ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)|~ssNder1_4r1r1r1r1(X1,X2,X3,X4)|~ssNder1_3r1r1r1(X1,X2,X3)|~ssNder1_2r1r1(X1,X2)|~ssNder1_1r1(X1)|~ssNder1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause81)).
cnf(c_0_74,negated_conjecture,(ssNder1_1r1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1])])).
cnf(c_0_75,negated_conjecture,(ssNder1_2r1r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_1]), c_0_74])])).
cnf(c_0_76,negated_conjecture,(ssNder1_3r1r1r1(X1,X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_1]), c_0_74]), c_0_75])])).
cnf(c_0_77,negated_conjecture,(ssNder1_4r1r1r1r1(X1,X2,X3,X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_1]), c_0_74]), c_0_75]), c_0_76])])).
cnf(c_0_78,negated_conjecture,(ssNder1_5r1r1r1r1r1(X1,X2,X3,X4,X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77])])).
cnf(c_0_79,negated_conjecture,(ssNder1_6r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78])])).
cnf(c_0_80,negated_conjecture,(ssNder1_7r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79])])).
cnf(c_0_81,negated_conjecture,(ssNder1_8r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80])])).
cnf(c_0_82,negated_conjecture,(ssNder1_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81])])).
cnf(c_0_83,negated_conjecture,(ssNder1_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82])])).
cnf(c_0_84,negated_conjecture,(ssNder1_11r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83])])).
cnf(c_0_85,negated_conjecture,(ssNder1_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84])])).
cnf(c_0_86,negated_conjecture,(ssNder1_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85])])).
cnf(c_0_87,negated_conjecture,(ssNder1_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_88,negated_conjecture,(ssNder1_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87])])).
cnf(c_0_89,negated_conjecture,(ssNder1_16r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88])])).
cnf(c_0_90,negated_conjecture,(ssNder1_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89])])).
cnf(c_0_91,negated_conjecture,(ssNder1_18r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90])])).
cnf(c_0_92,negated_conjecture,(ssNder1_19r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91])])).
cnf(c_0_93,negated_conjecture,(ssNder1_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92])])).
cnf(c_0_94,negated_conjecture,(ssNder1_21r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93])])).
cnf(c_0_95,negated_conjecture,(ssNder1_22r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94])])).
cnf(c_0_96,negated_conjecture,(ssNder1_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_97,negated_conjecture,(ssNder1_24r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96])])).
cnf(c_0_98,negated_conjecture,(ssNder1_25r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97])])).
cnf(c_0_99,negated_conjecture,(ssNder1_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98])])).
cnf(c_0_100,negated_conjecture,(ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99])])).
cnf(c_0_101,negated_conjecture,(ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100])])).
cnf(c_0_102,negated_conjecture,(ssNder1_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_103,negated_conjecture,(ssPv24_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv18_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv26_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_104,negated_conjecture,(ssPv18_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_105,negated_conjecture,(ssNder1_30r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102])])).
cnf(c_0_106,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv24_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_107,negated_conjecture,(ssPv24_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv26_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_108,negated_conjecture,(ssNder1_31r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105])])).
cnf(c_0_109,negated_conjecture,(~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,skc39)|~ssNder1_27r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99])])).
cnf(c_0_110,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv26_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_111,negated_conjecture,(ssPv26_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_112,negated_conjecture,(ssPv18_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|ssPv21_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108])])).
cnf(c_0_113,negated_conjecture,(ssPv15_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108])])).
cnf(c_0_114,negated_conjecture,(~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,skc39)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_100])])).
cnf(c_0_115,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_116,negated_conjecture,(ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv18_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108])])).
cnf(c_0_117,negated_conjecture,(ssPv18_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|ssPv21_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_118,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_119,negated_conjecture,(ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv27_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv28_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv21_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_120,negated_conjecture,(ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv21_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108])])).
cnf(c_0_121,negated_conjecture,(ssPv21_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_122,negated_conjecture,(ssPv29_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv40_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv24_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_123,negated_conjecture,(ssPv24_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_124,negated_conjecture,(ssPv27_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv28_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv21_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(spm,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_125,negated_conjecture,(ssPv21_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_126,negated_conjecture,(~ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,skc33)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105])])).
cnf(c_0_127,negated_conjecture,(ssPv9_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)|ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_128,negated_conjecture,(ssPv40_1r1(X1)|~ssPv40_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssNder1_28r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100])])).
cnf(c_0_129,negated_conjecture,(ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssPv29_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101])])).
cnf(c_0_130,negated_conjecture,(ssPv29_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv40_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_131,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv27_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_132,negated_conjecture,(ssPv27_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv28_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(spm,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_133,negated_conjecture,(ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_134,negated_conjecture,(ssPv15_26r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,skc44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97])])).
cnf(c_0_135,negated_conjecture,(ssPv40_1r1(X1)|~ssPv40_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_101])])).
cnf(c_0_136,negated_conjecture,(ssPv40_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_137,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv28_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_138,negated_conjecture,(ssPv28_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_139,negated_conjecture,(ssPv18_23r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23)|ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_140,negated_conjecture,(ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|ssPv40_1r1(X1)|~ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_141,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,skc38)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100])])).
cnf(c_0_142,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_143,negated_conjecture,(ssPv21_20r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20)|ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(spm,[status(thm)],[c_0_118, c_0_139])).
cnf(c_0_144,negated_conjecture,(ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssPv29_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_145,negated_conjecture,(ssPv29_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv39_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv40_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_146,negated_conjecture,(ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|ssPv40_1r1(X1)|~ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)), inference(spm,[status(thm)],[c_0_140, c_0_141])).
cnf(c_0_147,negated_conjecture,(ssPv24_17r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(spm,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_148,negated_conjecture,(ssNder1_32r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108])])).
cnf(c_0_149,negated_conjecture,(ssPv39_2r1r1(X1,X2)|~ssPv39_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_150,negated_conjecture,(ssPv39_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|~ssPv40_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_151,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|ssPv40_1r1(X1)|~ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_152,negated_conjecture,(ssPv26_15r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,skc50)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_153,negated_conjecture,(ssNder1_33r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148])])).
cnf(c_0_154,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|ssPv39_2r1r1(X1,X2)|~ssPv40_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)), inference(spm,[status(thm)],[c_0_149, c_0_150])).
cnf(c_0_155,negated_conjecture,(ssPv40_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|~ssPv40_1r1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_156,negated_conjecture,(~ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,skc53)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85])])).
cnf(c_0_157,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|ssPv40_1r1(X1)), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_158,negated_conjecture,(ssNder1_34r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153])])).
cnf(c_0_159,negated_conjecture,(ssPv27_14r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14)|ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|ssPv39_2r1r1(X1,X2)|~ssPv40_1r1(X1)), inference(spm,[status(thm)],[c_0_154, c_0_155])).
cnf(c_0_160,negated_conjecture,(~ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,skc55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84])])).
cnf(c_0_161,negated_conjecture,(ssPv28_13r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13)|ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|ssPv40_1r1(X1)), inference(spm,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_162,negated_conjecture,(ssNder1_35r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153]), c_0_158])])).
cnf(c_0_163,negated_conjecture,(~ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,skc57)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83])])).
cnf(c_0_164,negated_conjecture,(ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|ssPv39_2r1r1(X1,X2)|~ssPv40_1r1(X1)), inference(spm,[status(thm)],[c_0_156, c_0_159])).
cnf(c_0_165,negated_conjecture,(ssPv29_12r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12)|ssPv40_1r1(X1)), inference(spm,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_166,negated_conjecture,(ssNder1_36r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153]), c_0_158]), c_0_162])])).
cnf(c_0_167,negated_conjecture,(ssPv39_2r1r1(X1,X2)|~ssPv40_1r1(X1)), inference(spm,[status(thm)],[c_0_163, c_0_164])).
cnf(c_0_168,negated_conjecture,(ssPv40_1r1(X1)), inference(spm,[status(thm)],[c_0_163, c_0_165])).
cnf(c_0_169,negated_conjecture,(ssPv40_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv40_1r1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153]), c_0_158]), c_0_162]), c_0_166])])).
cnf(c_0_170,negated_conjecture,(ssPv39_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv39_2r1r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153]), c_0_158]), c_0_162]), c_0_166])])).
cnf(c_0_171,negated_conjecture,(ssPv39_2r1r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_167, c_0_168])])).
cnf(c_0_172,negated_conjecture,(ssPv31_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv4_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv40_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153]), c_0_158]), c_0_162]), c_0_166])])).
cnf(c_0_173,negated_conjecture,(ssPv40_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_169, c_0_168])])).
cnf(c_0_174,negated_conjecture,(ssPv4_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|ssPv12_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv39_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153]), c_0_158]), c_0_162]), c_0_166])])).
cnf(c_0_175,negated_conjecture,(ssPv39_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_170, c_0_171])])).
cnf(c_0_176,negated_conjecture,(ssPv31_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv4_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_172, c_0_173])])).
cnf(c_0_177,negated_conjecture,(ssPv4_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|ssPv12_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_174, c_0_175])])).
cnf(c_0_178,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv12_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153]), c_0_158]), c_0_162]), c_0_166])])).
cnf(c_0_179,negated_conjecture,(ssPv12_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|ssPv31_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(spm,[status(thm)],[c_0_176, c_0_177])).
cnf(c_0_180,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv31_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153]), c_0_158]), c_0_162]), c_0_166])])).
cnf(c_0_181,negated_conjecture,(ssPv31_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|~ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(spm,[status(thm)],[c_0_178, c_0_179])).
cnf(c_0_182,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)), inference(spm,[status(thm)],[c_0_180, c_0_181])).
cnf(c_0_183,negated_conjecture,(ssPv32_37r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91]), c_0_92]), c_0_93]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_105]), c_0_108]), c_0_148]), c_0_153]), c_0_158]), c_0_162]), c_0_166])])).
cnf(c_0_184,negated_conjecture,(ssPv12_29r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29)|ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_182, c_0_183])).
cnf(c_0_185,negated_conjecture,(~ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,skc61)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80]), c_0_81])])).
cnf(c_0_186,negated_conjecture,(ssPv31_10r1r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9,X10)|~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_114, c_0_184])).
cnf(c_0_187,negated_conjecture,(ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,skc62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_1]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_80])])).
cnf(c_0_188,negated_conjecture,(~ssPv32_9r1r1r1r1r1r1r1r1r1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_185, c_0_186])).
cnf(c_0_189,negated_conjecture,($false), inference(sr,[status(thm)],[c_0_187, c_0_188]), ['proof']).
# SZS output end CNFRefutation.
EOF
