% Problem    : CSR115+53 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n122.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:46:36 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_news_1308)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((attr(c1341,c1342)&sub(c1341,firma_1_1))&sub(c1342,name_1_1))&val(c1342,rover_0))&attr(c1516,c1517))&sub(c1516,firma_1_1))&sub(c1517,name_1_1))&val(c1517,bmw_0))&pmod(c1524,zweit_1_1,engagement_1_1))&prop(c1532,britisch__1_1))&sub(c1532,gesellschaft_1_1))&bei(c1536,c1341))&equ(c3,c3))&loc(c3,c1536))&mode(c3,c1532))&prop(c3,gro__337_1_1))&subs(c3,c1524))&subs(c3,einstieg_1_1))&arg1(c855,c3))&arg2(c855,c3))&benf(c855,c1516))&subr(c855,equ_0))&sort(c1341,d))&sort(c1341,io))&card(c1341,int1))&etype(c1341,int0))&fact(c1341,real))&gener(c1341,sp))&quant(c1341,one))&refer(c1341,det))&varia(c1341,con))&sort(c1342,na))&card(c1342,int1))&etype(c1342,int0))&fact(c1342,real))&gener(c1342,sp))&quant(c1342,one))&refer(c1342,indet))&varia(c1342,varia_c))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(rover_0,fe))&sort(c1516,d))&sort(c1516,io))&card(c1516,int1))&etype(c1516,int0))&fact(c1516,real))&gener(c1516,sp))&quant(c1516,one))&refer(c1516,det))&varia(c1516,con))&sort(c1517,na))&card(c1517,int1))&etype(c1517,int0))&fact(c1517,real))&gener(c1517,sp))&quant(c1517,one))&refer(c1517,indet))&varia(c1517,varia_c))&sort(bmw_0,fe))&sort(c1524,ad))&card(c1524,int1))&etype(c1524,int0))&fact(c1524,real))&gener(c1524,ge))&quant(c1524,one))&refer(c1524,refer_c))&varia(c1524,varia_c))&sort(zweit_1_1,oq))&card(zweit_1_1,int2))&sort(engagement_1_1,ad))&card(engagement_1_1,int1))&etype(engagement_1_1,int0))&fact(engagement_1_1,real))&gener(engagement_1_1,ge))&quant(engagement_1_1,one))&refer(engagement_1_1,refer_c))&varia(engagement_1_1,varia_c))&sort(c1532,io))&card(c1532,int1))&etype(c1532,int1))&fact(c1532,real))&gener(c1532,sp))&quant(c1532,one))&refer(c1532,indet))&varia(c1532,varia_c))&sort(britisch__1_1,nq))&sort(gesellschaft_1_1,io))&card(gesellschaft_1_1,card_c))&etype(gesellschaft_1_1,int1))&fact(gesellschaft_1_1,real))&gener(gesellschaft_1_1,ge))&quant(gesellschaft_1_1,quant_c))&refer(gesellschaft_1_1,refer_c))&varia(gesellschaft_1_1,varia_c))&sort(c1536,l))&card(c1536,int1))&etype(c1536,int0))&fact(c1536,real))&gener(c1536,sp))&quant(c1536,one))&refer(c1536,det))&varia(c1536,con))&sort(c3,ad))&card(c3,int1))&etype(c3,int0))&fact(c3,real))&gener(c3,sp))&quant(c3,one))&refer(c3,det))&varia(c3,con))&sort(gro__337_1_1,mq))&sort(einstieg_1_1,ad))&card(einstieg_1_1,int1))&etype(einstieg_1_1,int0))&fact(einstieg_1_1,real))&gener(einstieg_1_1,ge))&quant(einstieg_1_1,one))&refer(einstieg_1_1,refer_c))&varia(einstieg_1_1,varia_c))&sort(c855,st))&fact(c855,real))&gener(c855,sp))&sort(equ_0,st))&fact(equ_0,real))&gener(equ_0,gener_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_news_1308)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_3, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_4, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, negated_conjecture, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((((~attr(X8,X9)|~attr(X11,X10))|~attr(X12,X13))|~obj(X14,X8))|~sub(X9,name_1_1))|~sub(X8,firma_1_1))|~sub(X10,name_1_1))|~val(X9,bmw_0))|~val(X10,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_6,negated_conjecture,(~val(X1,bmw_0)|~val(X2,bmw_0)|~sub(X1,name_1_1)|~sub(X3,firma_1_1)|~sub(X2,name_1_1)|~obj(X4,X3)|~attr(X5,X6)|~attr(X7,X1)|~attr(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,hypothesis,(val(c1517,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(sub(c1517,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(~val(X1,bmw_0)|~attr(X2,c1517)|~attr(X3,X1)|~attr(X4,X5)|~obj(X6,X2)|~sub(X2,firma_1_1)|~sub(X1,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_10,hypothesis,(~attr(X1,c1517)|~attr(X2,c1517)|~attr(X3,X4)|~obj(X5,X1)|~sub(X1,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_8])])).
cnf(c_0_11,hypothesis,(attr(c1516,c1517)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_12,hypothesis,(sub(c1516,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(~attr(X1,c1517)|~attr(X2,X3)|~obj(X4,c1516)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk15_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk15_3(X7,X8,X9),esk16_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk14_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk14_3(X7,X8,X9),esk15_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk14_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk16_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk15_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk14_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_15,hypothesis,(~attr(X1,X2)|~obj(X3,c1516)), inference(spm,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_16,plain,(obj(esk14_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, (![X4]:![X5]:(((arg1(esk17_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk17_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk17_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_18,hypothesis,(~subr(X1,sub_0)|~attr(X2,X3)|~arg2(X1,X4)|~arg1(X1,c1516)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(subr(esk17_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,hypothesis,(~attr(X1,X2)|~arg2(esk17_2(X3,X4),X5)|~arg1(esk17_2(X3,X4),c1516)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(arg2(esk17_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,hypothesis,(~attr(X1,X2)|~arg1(esk17_2(X3,X4),c1516)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,plain,(arg1(esk17_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,hypothesis,(~attr(X1,X2)|~sub(c1516,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,hypothesis,(attr(c1341,c1342)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_26,hypothesis,(~attr(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_12])).
cnf(c_0_27,hypothesis,($false), inference(sr,[status(thm)],[c_0_25, c_0_26]), ['proof']).
# SZS output end CNFRefutation.
EOF
