% Problem    : CSR116+24 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n129.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:25:01 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:?[X8]:?[X9]:?[X10]:((((((((((((((in(X6,X7)&arg1(X4,X1))&arg2(X4,X5))&attr(X1,X2))&attr(X1,X3))&attr(X7,X8))&obj(X9,X1))&sub(X2,familiename_1_1))&sub(X3,eigenname_1_1))&sub(X5,X10))&sub(X8,name_1_1))&subr(X4,rprs_0))&val(X2,mandela_0))&val(X3,nelson_0))&val(X8,s__374dafrika_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_010_mira_news_1788)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((origl(c15,c323)&sub(c15,an_f__374hrer_1_1))&sub(c15,c15))&sub(c15,hirte_1_1))&arg1(c19,c15))&arg2(c19,c15))&subr(c19,sub_0))&pred(c307,autobiographie_1_1))&attch(c316,c307))&attr(c316,c317))&attr(c316,c318))&prop(c316,s__374dafrikanisch_1_1))&sub(c316,pr__344sident_1_1))&sub(c317,eigenname_1_1))&val(c317,nelson_0))&sub(c318,familiename_1_1))&val(c318,mandela_0))&flp(c323,c307))&agt(c324,c15))&subs(c324,f__374hren_1_1))&sort(c15,d))&card(c15,int1))&etype(c15,int0))&fact(c15,real))&gener(c15,sp))&quant(c15,one))&refer(c15,indet))&varia(c15,varia_c))&sort(c323,l))&card(c323,cons(x_constant,cons(int1,nil))))&etype(c323,int1))&fact(c323,real))&gener(c323,sp))&quant(c323,mult))&refer(c323,det))&varia(c323,con))&sort(an_f__374hrer_1_1,d))&card(an_f__374hrer_1_1,int1))&etype(an_f__374hrer_1_1,int0))&fact(an_f__374hrer_1_1,real))&gener(an_f__374hrer_1_1,ge))&quant(an_f__374hrer_1_1,one))&refer(an_f__374hrer_1_1,refer_c))&varia(an_f__374hrer_1_1,varia_c))&sort(hirte_1_1,d))&card(hirte_1_1,int1))&etype(hirte_1_1,int0))&fact(hirte_1_1,real))&gener(hirte_1_1,ge))&quant(hirte_1_1,one))&refer(hirte_1_1,refer_c))&varia(hirte_1_1,varia_c))&sort(c19,st))&fact(c19,real))&gener(c19,sp))&sort(sub_0,st))&fact(sub_0,real))&gener(sub_0,gener_c))&sort(c307,d))&sort(c307,io))&card(c307,cons(x_constant,cons(int1,nil))))&etype(c307,int1))&fact(c307,real))&gener(c307,sp))&quant(c307,mult))&refer(c307,det))&varia(c307,con))&sort(autobiographie_1_1,d))&sort(autobiographie_1_1,io))&card(autobiographie_1_1,int1))&etype(autobiographie_1_1,int0))&fact(autobiographie_1_1,real))&gener(autobiographie_1_1,ge))&quant(autobiographie_1_1,one))&refer(autobiographie_1_1,refer_c))&varia(autobiographie_1_1,varia_c))&sort(c316,d))&card(c316,int1))&etype(c316,int0))&fact(c316,real))&gener(c316,sp))&quant(c316,one))&refer(c316,det))&varia(c316,con))&sort(c317,na))&card(c317,int1))&etype(c317,int0))&fact(c317,real))&gener(c317,sp))&quant(c317,one))&refer(c317,indet))&varia(c317,varia_c))&sort(c318,na))&card(c318,int1))&etype(c318,int0))&fact(c318,real))&gener(c318,sp))&quant(c318,one))&refer(c318,indet))&varia(c318,varia_c))&sort(s__374dafrikanisch_1_1,nq))&sort(pr__344sident_1_1,d))&card(pr__344sident_1_1,int1))&etype(pr__344sident_1_1,int0))&fact(pr__344sident_1_1,real))&gener(pr__344sident_1_1,ge))&quant(pr__344sident_1_1,one))&refer(pr__344sident_1_1,refer_c))&varia(pr__344sident_1_1,varia_c))&sort(eigenname_1_1,na))&card(eigenname_1_1,int1))&etype(eigenname_1_1,int0))&fact(eigenname_1_1,real))&gener(eigenname_1_1,ge))&quant(eigenname_1_1,one))&refer(eigenname_1_1,refer_c))&varia(eigenname_1_1,varia_c))&sort(nelson_0,fe))&sort(familiename_1_1,na))&card(familiename_1_1,int1))&etype(familiename_1_1,int0))&fact(familiename_1_1,real))&gener(familiename_1_1,ge))&quant(familiename_1_1,one))&refer(familiename_1_1,refer_c))&varia(familiename_1_1,varia_c))&sort(mandela_0,fe))&sort(c324,da))&fact(c324,real))&gener(c324,sp))&sort(f__374hren_1_1,da))&fact(f__374hren_1_1,real))&gener(f__374hren_1_1,ge))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_010_mira_news_1788)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:((prop(X1,X2)&state_adjective_state_binding(X2,X3))=>?[X4]:?[X5]:?[X6]:(((((in(X6,X4)&attr(X4,X5))&loc(X1,X6))&sub(X4,land_1_1))&sub(X5,name_1_1))&val(X5,X3)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', state_adjective__in_state)).
fof(c_0_3, axiom, (state_adjective_state_binding(s__374dafrikanisch_1_1,s__374dafrika_0)), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', fact_8980)).
fof(c_0_4, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_5, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_6, axiom, (![X1]:![X2]:((has_fact_leq(X2,real)&loc(X2,X1))=>?[X3]:((loc(X3,X1)&obj(X3,X2))&subs(X3,geben_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', loc__geben_1_1_loc)).
fof(c_0_7, axiom, (![X1]:![X2]:(fact(X1,X2)=>has_fact_leq(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', has_fact_eq)).
fof(c_0_8, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:?[X8]:?[X9]:?[X10]:((((((((((((((in(X6,X7)&arg1(X4,X1))&arg2(X4,X5))&attr(X1,X2))&attr(X1,X3))&attr(X7,X8))&obj(X9,X1))&sub(X2,familiename_1_1))&sub(X3,eigenname_1_1))&sub(X5,X10))&sub(X8,name_1_1))&subr(X4,rprs_0))&val(X2,mandela_0))&val(X3,nelson_0))&val(X8,s__374dafrika_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_9, negated_conjecture, (![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:((((((((((((((~in(X16,X17)|~arg1(X14,X11))|~arg2(X14,X15))|~attr(X11,X12))|~attr(X11,X13))|~attr(X17,X18))|~obj(X19,X11))|~sub(X12,familiename_1_1))|~sub(X13,eigenname_1_1))|~sub(X15,X20))|~sub(X18,name_1_1))|~subr(X14,rprs_0))|~val(X12,mandela_0))|~val(X13,nelson_0))|~val(X18,s__374dafrika_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
cnf(c_0_10,negated_conjecture,(~val(X1,s__374dafrika_0)|~val(X2,nelson_0)|~val(X3,mandela_0)|~subr(X4,rprs_0)|~sub(X1,name_1_1)|~sub(X5,X6)|~sub(X2,eigenname_1_1)|~sub(X3,familiename_1_1)|~obj(X7,X8)|~attr(X9,X1)|~attr(X8,X2)|~attr(X8,X3)|~arg2(X4,X5)|~arg1(X4,X8)|~in(X10,X9)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_11,hypothesis,(val(c318,mandela_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_12,hypothesis,(sub(c318,familiename_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(~val(X1,nelson_0)|~val(X2,s__374dafrika_0)|~subr(X3,rprs_0)|~in(X4,X5)|~attr(X6,c318)|~attr(X5,X2)|~attr(X6,X1)|~arg2(X3,X7)|~arg1(X3,X6)|~obj(X8,X6)|~sub(X1,eigenname_1_1)|~sub(X2,name_1_1)|~sub(X7,X9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
cnf(c_0_14,hypothesis,(val(c317,nelson_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_15,hypothesis,(sub(c317,eigenname_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_16, plain, (![X7]:![X8]:![X9]:((((((in(esk9_3(X7,X8,X9),esk7_3(X7,X8,X9))|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9)))&(attr(esk7_3(X7,X8,X9),esk8_3(X7,X8,X9))|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(loc(X7,esk9_3(X7,X8,X9))|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(sub(esk7_3(X7,X8,X9),land_1_1)|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(sub(esk8_3(X7,X8,X9),name_1_1)|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(val(esk8_3(X7,X8,X9),X9)|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_17,hypothesis,(~val(X1,s__374dafrika_0)|~subr(X2,rprs_0)|~in(X3,X4)|~attr(X5,c318)|~attr(X5,c317)|~attr(X4,X1)|~arg2(X2,X6)|~arg1(X2,X5)|~obj(X7,X5)|~sub(X1,name_1_1)|~sub(X6,X8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
cnf(c_0_18,plain,(val(esk8_3(X3,X1,X2),X2)|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_19,plain,(sub(esk8_3(X3,X1,X2),name_1_1)|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20,hypothesis,(~state_adjective_state_binding(X1,s__374dafrika_0)|~prop(X2,X1)|~subr(X3,rprs_0)|~in(X4,X5)|~attr(X5,esk8_3(X2,X1,s__374dafrika_0))|~attr(X6,c318)|~attr(X6,c317)|~arg2(X3,X7)|~arg1(X3,X6)|~obj(X8,X6)|~sub(X7,X9)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_21,plain,(attr(esk7_3(X3,X1,X2),esk8_3(X3,X1,X2))|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,hypothesis,(~state_adjective_state_binding(X1,s__374dafrika_0)|~prop(X2,X1)|~subr(X3,rprs_0)|~in(X4,esk7_3(X2,X1,s__374dafrika_0))|~attr(X5,c318)|~attr(X5,c317)|~arg2(X3,X6)|~arg1(X3,X5)|~obj(X7,X5)|~sub(X6,X8)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,plain,(in(esk9_3(X3,X1,X2),esk7_3(X3,X1,X2))|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24,hypothesis,(~state_adjective_state_binding(X1,s__374dafrika_0)|~prop(X2,X1)|~subr(X3,rprs_0)|~attr(X4,c318)|~attr(X4,c317)|~arg2(X3,X5)|~arg1(X3,X4)|~obj(X6,X4)|~sub(X5,X7)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,plain,(state_adjective_state_binding(s__374dafrikanisch_1_1,s__374dafrika_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_26,hypothesis,(~prop(X1,s__374dafrikanisch_1_1)|~subr(X2,rprs_0)|~attr(X3,c318)|~attr(X3,c317)|~arg2(X2,X4)|~arg1(X2,X3)|~obj(X5,X3)|~sub(X4,X6)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_27,hypothesis,(prop(c316,s__374dafrikanisch_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_28, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk13_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk13_3(X7,X8,X9),esk14_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk12_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk12_3(X7,X8,X9),esk13_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk12_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk14_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk13_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk12_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
cnf(c_0_29,hypothesis,(~subr(X1,rprs_0)|~attr(X2,c318)|~attr(X2,c317)|~arg2(X1,X3)|~arg1(X1,X2)|~obj(X4,X2)|~sub(X3,X5)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_30,plain,(subr(esk13_3(X1,X3,X2),rprs_0)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_31,hypothesis,(~subr(X1,sub_0)|~attr(X2,c318)|~attr(X2,c317)|~arg2(esk13_3(X1,X3,X4),X5)|~arg2(X1,X4)|~arg1(esk13_3(X1,X3,X4),X2)|~arg1(X1,X3)|~obj(X6,X2)|~sub(X5,X7)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_32,plain,(arg2(esk13_3(X1,X3,X2),esk14_3(X1,X3,X2))|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33,hypothesis,(~subr(X1,sub_0)|~attr(X2,c318)|~attr(X2,c317)|~arg2(X1,X3)|~arg1(esk13_3(X1,X4,X3),X2)|~arg1(X1,X4)|~obj(X5,X2)|~sub(esk14_3(X1,X4,X3),X6)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_34,plain,(arg1(esk13_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35,hypothesis,(~subr(X1,sub_0)|~attr(X2,c318)|~attr(X2,c317)|~arg2(X1,X3)|~arg1(X1,X2)|~obj(X4,X2)|~sub(esk14_3(X1,X2,X3),X5)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_36,plain,(sub(esk14_3(X1,X3,X2),X2)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_37, plain, (![X4]:![X5]:(((arg1(esk15_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk15_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk15_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
cnf(c_0_38,hypothesis,(~subr(X1,sub_0)|~attr(X2,c318)|~attr(X2,c317)|~arg2(X1,X3)|~arg1(X1,X2)|~obj(X4,X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39,plain,(subr(esk15_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_40,hypothesis,(~attr(X1,c318)|~attr(X1,c317)|~arg2(esk15_2(X2,X3),X4)|~arg1(esk15_2(X2,X3),X1)|~obj(X5,X1)|~sub(X2,X3)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_41,plain,(arg2(esk15_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_42,hypothesis,(~attr(X1,c318)|~attr(X1,c317)|~arg1(esk15_2(X2,X3),X1)|~obj(X4,X1)|~sub(X2,X3)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_43,plain,(arg1(esk15_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44,hypothesis,(~attr(X1,c318)|~attr(X1,c317)|~obj(X2,X1)|~sub(X1,X3)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_45,hypothesis,(attr(c316,c318)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_46,hypothesis,(attr(c316,c317)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_47, plain, (![X4]:![X5]:(((loc(esk23_2(X4,X5),X4)|(~has_fact_leq(X5,real)|~loc(X5,X4)))&(obj(esk23_2(X4,X5),X5)|(~has_fact_leq(X5,real)|~loc(X5,X4))))&(subs(esk23_2(X4,X5),geben_1_1)|(~has_fact_leq(X5,real)|~loc(X5,X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_48,hypothesis,(~obj(X1,c316)|~sub(c316,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])])).
cnf(c_0_49,plain,(obj(esk23_2(X2,X1),X1)|~loc(X1,X2)|~has_fact_leq(X1,real)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_50, plain, (![X3]:![X4]:(~fact(X3,X4)|has_fact_leq(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
cnf(c_0_51,hypothesis,(~loc(c316,X1)|~sub(c316,X2)|~has_fact_leq(c316,real)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_52,plain,(has_fact_leq(X1,X2)|~fact(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_53,hypothesis,(fact(c316,real)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_54,hypothesis,(~loc(c316,X1)|~sub(c316,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])])).
cnf(c_0_55,plain,(loc(X3,esk9_3(X3,X1,X2))|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_56,hypothesis,(~state_adjective_state_binding(X1,X2)|~prop(c316,X1)|~sub(c316,X3)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_57,hypothesis,(sub(c316,pr__344sident_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_58,hypothesis,(~state_adjective_state_binding(X1,X2)|~prop(c316,X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_59,hypothesis,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_25]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation.
EOF
