% Problem    : KRS267+1 : TPTP v6.1.0. Bugfixed v5.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n117.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:02:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C18_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X7]:![X8]:(![X1]:![X2]:(status(X1,X2,X7)=>~(status(X1,X2,X8)))<=>nevera(X7,X8))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', nevera)).
fof(c_0_1, axiom, (![X1]:![X2]:((~(?[X3]:model(X3,X1))&![X4]:model(X4,X2))<=>status(X1,X2,tca))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', tca)).
fof(c_0_2, axiom, (![X1]:![X2]:(![X3]:(model(X3,X1)=>model(X3,X2))<=>status(X1,X2,thm))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', thm)).
fof(c_0_3, axiom, (?[X10]:![X9]:model(X9,X10)), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', tautology)).
fof(c_0_4, axiom, (![X7]:![X8]:(?[X1]:?[X2]:(status(X1,X2,X7)&status(X1,X2,X8))<=>mighta(X7,X8))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', mighta)).
fof(c_0_5, conjecture, (mighta(tca,thm)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', mighta_tca_thm)).
fof(c_0_6, axiom, (?[X10]:![X9]:~(model(X9,X10))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', contradiction)).
fof(c_0_7, plain, (![X7]:![X8]:(![X1]:![X2]:(status(X1,X2,X7)=>~status(X1,X2,X8))<=>nevera(X7,X8))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_8, plain, (![X5]:![X6]:![X9]:![X10]:![X11]:![X12]:(((model(esk33_2(X5,X6),X5)|~model(esk34_2(X5,X6),X6))|status(X5,X6,tca))&((~model(X11,X9)|~status(X9,X10,tca))&(model(X12,X10)|~status(X9,X10,tca))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])])).
fof(c_0_9, plain, (![X9]:![X10]:![X13]:![X14]:![X15]:![X16]:(((status(esk49_2(X9,X10),esk50_2(X9,X10),X9)|nevera(X9,X10))&(status(esk49_2(X9,X10),esk50_2(X9,X10),X10)|nevera(X9,X10)))&(~nevera(X13,X14)|(~status(X15,X16,X13)|~status(X15,X16,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])])).
cnf(c_0_10,plain,(~status(X1,X2,tca)|~model(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11,plain,(nevera(X1,X2)|status(esk49_2(X1,X2),esk50_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_12, plain, (![X4]:![X5]:![X7]:![X8]:![X9]:(((model(esk10_2(X4,X5),X4)|status(X4,X5,thm))&(~model(esk10_2(X4,X5),X5)|status(X4,X5,thm)))&(~status(X7,X8,thm)|(~model(X9,X7)|model(X9,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
fof(c_0_13, plain, (![X12]:model(X12,esk53_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_3])])).
fof(c_0_14, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:(((~status(X11,X12,X9)|~status(X11,X12,X10))|mighta(X9,X10))&((status(esk43_2(X13,X14),esk44_2(X13,X14),X13)|~mighta(X13,X14))&(status(esk43_2(X13,X14),esk44_2(X13,X14),X14)|~mighta(X13,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
cnf(c_0_15,plain,(nevera(X1,tca)|~model(X2,esk49_2(X1,tca))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_16,plain,(status(X1,X2,thm)|model(esk10_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(status(X1,X2,thm)|~model(esk10_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(model(X1,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(mighta(X1,X2)|~status(X3,X4,X2)|~status(X3,X4,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(nevera(X1,tca)|status(esk49_2(X1,tca),X2,thm)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_21, negated_conjecture, (~mighta(tca,thm)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_5])])).
cnf(c_0_22,plain,(~status(X1,X2,X3)|~status(X1,X2,X4)|~nevera(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23,plain,(status(X1,esk53_0,thm)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24,plain,(nevera(X1,tca)|mighta(X2,thm)|~status(esk49_2(X1,tca),X3,X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25,negated_conjecture,(~mighta(tca,thm)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_26, plain, (?[X10]:![X9]:~model(X9,X10)), inference(fof_simplification,[status(thm)],[c_0_6])).
cnf(c_0_27,plain,(~nevera(thm,X1)|~status(X2,esk53_0,X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28,plain,(nevera(X1,tca)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_11]), c_0_25])).
fof(c_0_29, plain, (![X12]:~model(X12,esk57_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_26])])).
cnf(c_0_30,plain,(status(X1,X2,tca)|model(esk33_2(X1,X2),X1)|~model(esk34_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_31,plain,(~status(X1,esk53_0,tca)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32,plain,(~model(X1,esk57_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33,plain,(model(esk33_2(X1,esk53_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_18]), c_0_31])).
cnf(c_0_34,plain,($false), inference(spm,[status(thm)],[c_0_32, c_0_33]), ['proof']).
# SZS output end CNFRefutation.
EOF
