% Problem    : CSR057+2 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n106.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:39:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_C45_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X13]:(mtvisible(c_tptpgeo_member8_mt)=>inregion(X13,c_georegion_l4_x75_y75))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', query107)).
fof(c_0_1, axiom, (![X20]:![X21]:((mtvisible(X20)&genlmt(X20,X21))=>mtvisible(X21))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_1123)).
fof(c_0_2, axiom, (genlmt(c_tptpgeo_member8_mt,c_tptpgeo_spindleheadmt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_1)).
fof(c_0_3, axiom, ((mtvisible(c_worldgeographymt)=>geolevel_4(c_georegion_l4_x75_y75))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_464)).
fof(c_0_4, axiom, (genlmt(c_tptpgeo_spindleheadmt,c_worldgeographymt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_326)).
fof(c_0_5, axiom, (![X1]:(geolevel_4(X1)=>geographicalregion(X1))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_462)).
fof(c_0_6, axiom, (![X13]:(geographicalregion(X13)=>geographicalsubregions(X13,X13))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_1042)).
fof(c_0_7, axiom, (![X5]:![X6]:(geographicalsubregions(X5,X6)=>inregion(X6,X5))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_128)).
fof(c_0_8, negated_conjecture, (~(?[X13]:(mtvisible(c_tptpgeo_member8_mt)=>inregion(X13,c_georegion_l4_x75_y75)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_9, plain, (![X22]:![X23]:((~mtvisible(X22)|~genlmt(X22,X23))|mtvisible(X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
fof(c_0_10, negated_conjecture, (![X14]:(mtvisible(c_tptpgeo_member8_mt)&~inregion(X14,c_georegion_l4_x75_y75))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
cnf(c_0_11,plain,(mtvisible(X1)|~genlmt(X2,X1)|~mtvisible(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12,plain,(genlmt(c_tptpgeo_member8_mt,c_tptpgeo_spindleheadmt)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_13,negated_conjecture,(mtvisible(c_tptpgeo_member8_mt)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ((~mtvisible(c_worldgeographymt)|geolevel_4(c_georegion_l4_x75_y75))), inference(fof_nnf,[status(thm)],[c_0_3])).
cnf(c_0_15,plain,(genlmt(c_tptpgeo_spindleheadmt,c_worldgeographymt)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,plain,(mtvisible(c_tptpgeo_spindleheadmt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
fof(c_0_17, plain, (![X2]:(~geolevel_4(X2)|geographicalregion(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
cnf(c_0_18,plain,(geolevel_4(c_georegion_l4_x75_y75)|~mtvisible(c_worldgeographymt)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,plain,(mtvisible(c_worldgeographymt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_15]), c_0_16])])).
fof(c_0_20, plain, (![X14]:(~geographicalregion(X14)|geographicalsubregions(X14,X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
cnf(c_0_21,plain,(geographicalregion(X1)|~geolevel_4(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,plain,(geolevel_4(c_georegion_l4_x75_y75)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19])])).
fof(c_0_23, plain, (![X7]:![X8]:(~geographicalsubregions(X7,X8)|inregion(X8,X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
cnf(c_0_24,plain,(geographicalsubregions(X1,X1)|~geographicalregion(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25,plain,(geographicalregion(c_georegion_l4_x75_y75)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26,plain,(inregion(X1,X2)|~geographicalsubregions(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27,plain,(geographicalsubregions(c_georegion_l4_x75_y75,c_georegion_l4_x75_y75)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28,negated_conjecture,(~inregion(X1,c_georegion_l4_x75_y75)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_29,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), ['proof']).
# SZS output end CNFRefutation.
EOF
