% Problem    : LCL656+1.015 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n096.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:58:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_B02_F1_AE_CS_SP_PI_S0S
# and selection function SelectComplexG.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (~(?[X1]:~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p7(X2)))))))))))))))))|~(((![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|((((((((((((((((((((((((((((((((((((((((((((((((~(![X1]:(~(r1(X2,X1))|~(((~(p16(X1))&~(p116(X1)))&p115(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p16(X1)&~(p116(X1)))&p115(X1))))))|~((~(p115(X2))&p114(X2))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p15(X1))&~(p115(X1)))&p114(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p15(X1)&~(p115(X1)))&p114(X1))))))|~((~(p114(X2))&p113(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p14(X1))&~(p114(X1)))&p113(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p14(X1)&~(p114(X1)))&p113(X1))))))|~((~(p113(X2))&p112(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p13(X1))&~(p113(X1)))&p112(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p13(X1)&~(p113(X1)))&p112(X1))))))|~((~(p112(X2))&p111(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p12(X1))&~(p112(X1)))&p111(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p12(X1)&~(p112(X1)))&p111(X1))))))|~((~(p111(X2))&p110(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p11(X1))&~(p111(X1)))&p110(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p11(X1)&~(p111(X1)))&p110(X1))))))|~((~(p110(X2))&p109(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p10(X1))&~(p110(X1)))&p109(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p10(X1)&~(p110(X1)))&p109(X1))))))|~((~(p109(X2))&p108(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p9(X1))&~(p109(X1)))&p108(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p9(X1)&~(p109(X1)))&p108(X1))))))|~((~(p108(X2))&p107(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p8(X1))&~(p108(X1)))&p107(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p8(X1)&~(p108(X1)))&p107(X1))))))|~((~(p107(X2))&p106(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p7(X1))&~(p107(X1)))&p106(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p7(X1)&~(p107(X1)))&p106(X1))))))|~((~(p106(X2))&p105(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p6(X1))&~(p106(X1)))&p105(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p6(X1)&~(p106(X1)))&p105(X1))))))|~((~(p105(X2))&p104(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p5(X1))&~(p105(X1)))&p104(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p5(X1)&~(p105(X1)))&p104(X1))))))|~((~(p104(X2))&p103(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p4(X1))&~(p104(X1)))&p103(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p4(X1)&~(p104(X1)))&p103(X1))))))|~((~(p103(X2))&p102(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p3(X1))&~(p103(X1)))&p102(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p3(X1)&~(p103(X1)))&p102(X1))))))|~((~(p102(X2))&p101(X2)))))&((~(![X1]:(~(r1(X2,X1))|~(((~(p2(X1))&~(p102(X1)))&p101(X1)))))&~(![X1]:(~(r1(X2,X1))|~(((p2(X1)&~(p102(X1)))&p101(X1))))))|~((~(p101(X2))&p100(X2)))))&(((![X1]:((~(r1(X2,X1))|~(p16(X1)))|~(p115(X1)))|p16(X2))&(![X1]:((~(r1(X2,X1))|p16(X1))|~(p115(X1)))|~(p16(X2))))|~(p115(X2))))&(((![X1]:((~(r1(X2,X1))|~(p15(X1)))|~(p114(X1)))|p15(X2))&(![X1]:((~(r1(X2,X1))|p15(X1))|~(p114(X1)))|~(p15(X2))))|~(p114(X2))))&(((![X1]:((~(r1(X2,X1))|~(p14(X1)))|~(p113(X1)))|p14(X2))&(![X1]:((~(r1(X2,X1))|p14(X1))|~(p113(X1)))|~(p14(X2))))|~(p113(X2))))&(((![X1]:((~(r1(X2,X1))|~(p13(X1)))|~(p112(X1)))|p13(X2))&(![X1]:((~(r1(X2,X1))|p13(X1))|~(p112(X1)))|~(p13(X2))))|~(p112(X2))))&(((![X1]:((~(r1(X2,X1))|~(p12(X1)))|~(p111(X1)))|p12(X2))&(![X1]:((~(r1(X2,X1))|p12(X1))|~(p111(X1)))|~(p12(X2))))|~(p111(X2))))&(((![X1]:((~(r1(X2,X1))|~(p11(X1)))|~(p110(X1)))|p11(X2))&(![X1]:((~(r1(X2,X1))|p11(X1))|~(p110(X1)))|~(p11(X2))))|~(p110(X2))))&(((![X1]:((~(r1(X2,X1))|~(p10(X1)))|~(p109(X1)))|p10(X2))&(![X1]:((~(r1(X2,X1))|p10(X1))|~(p109(X1)))|~(p10(X2))))|~(p109(X2))))&(((![X1]:((~(r1(X2,X1))|~(p9(X1)))|~(p108(X1)))|p9(X2))&(![X1]:((~(r1(X2,X1))|p9(X1))|~(p108(X1)))|~(p9(X2))))|~(p108(X2))))&(((![X1]:((~(r1(X2,X1))|~(p8(X1)))|~(p107(X1)))|p8(X2))&(![X1]:((~(r1(X2,X1))|p8(X1))|~(p107(X1)))|~(p8(X2))))|~(p107(X2))))&(((![X1]:((~(r1(X2,X1))|~(p7(X1)))|~(p106(X1)))|p7(X2))&(![X1]:((~(r1(X2,X1))|p7(X1))|~(p106(X1)))|~(p7(X2))))|~(p106(X2))))&(((![X1]:((~(r1(X2,X1))|~(p6(X1)))|~(p105(X1)))|p6(X2))&(![X1]:((~(r1(X2,X1))|p6(X1))|~(p105(X1)))|~(p6(X2))))|~(p105(X2))))&(((![X1]:((~(r1(X2,X1))|~(p5(X1)))|~(p104(X1)))|p5(X2))&(![X1]:((~(r1(X2,X1))|p5(X1))|~(p104(X1)))|~(p5(X2))))|~(p104(X2))))&(((![X1]:((~(r1(X2,X1))|~(p4(X1)))|~(p103(X1)))|p4(X2))&(![X1]:((~(r1(X2,X1))|p4(X1))|~(p103(X1)))|~(p4(X2))))|~(p103(X2))))&(((![X1]:((~(r1(X2,X1))|~(p3(X1)))|~(p102(X1)))|p3(X2))&(![X1]:((~(r1(X2,X1))|p3(X1))|~(p102(X1)))|~(p3(X2))))|~(p102(X2))))&(((![X1]:((~(r1(X2,X1))|~(p2(X1)))|~(p101(X1)))|p2(X2))&(![X1]:((~(r1(X2,X1))|p2(X1))|~(p101(X1)))|~(p2(X2))))|~(p101(X2))))&(((![X1]:((~(r1(X2,X1))|~(p1(X1)))|~(p100(X1)))|p1(X2))&(![X1]:((~(r1(X2,X1))|p1(X1))|~(p100(X1)))|~(p1(X2))))|~(p100(X2))))&(p115(X2)|~(p116(X2))))&(p114(X2)|~(p115(X2))))&(p113(X2)|~(p114(X2))))&(p112(X2)|~(p113(X2))))&(p111(X2)|~(p112(X2))))&(p110(X2)|~(p111(X2))))&(p109(X2)|~(p110(X2))))&(p108(X2)|~(p109(X2))))&(p107(X2)|~(p108(X2))))&(p106(X2)|~(p107(X2))))&(p105(X2)|~(p106(X2))))&(p104(X2)|~(p105(X2))))&(p103(X2)|~(p104(X2))))&(p102(X2)|~(p103(X2))))&(p101(X2)|~(p102(X2))))&(p100(X2)|~(p101(X2)))))))))))))))))))&~(p101(X1)))&p100(X1))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', main)).
fof(c_0_1, axiom, (![X1]:r1(X1,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', reflexivity)).
fof(c_0_2, plain, (![X2]:(epred1_1(X2)<=>((((((((((((((((((((((((((((((((((((((((((((((((~(![X1]:(~r1(X2,X1)|~(((~p16(X1)&~p116(X1))&p115(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p16(X1)&~p116(X1))&p115(X1))))))|~((~p115(X2)&p114(X2))))&((~(![X1]:(~r1(X2,X1)|~(((~p15(X1)&~p115(X1))&p114(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p15(X1)&~p115(X1))&p114(X1))))))|~((~p114(X2)&p113(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p14(X1)&~p114(X1))&p113(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p14(X1)&~p114(X1))&p113(X1))))))|~((~p113(X2)&p112(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p13(X1)&~p113(X1))&p112(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p13(X1)&~p113(X1))&p112(X1))))))|~((~p112(X2)&p111(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p12(X1)&~p112(X1))&p111(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p12(X1)&~p112(X1))&p111(X1))))))|~((~p111(X2)&p110(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p11(X1)&~p111(X1))&p110(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p11(X1)&~p111(X1))&p110(X1))))))|~((~p110(X2)&p109(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p10(X1)&~p110(X1))&p109(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p10(X1)&~p110(X1))&p109(X1))))))|~((~p109(X2)&p108(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p9(X1)&~p109(X1))&p108(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p9(X1)&~p109(X1))&p108(X1))))))|~((~p108(X2)&p107(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p8(X1)&~p108(X1))&p107(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p8(X1)&~p108(X1))&p107(X1))))))|~((~p107(X2)&p106(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p7(X1)&~p107(X1))&p106(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p7(X1)&~p107(X1))&p106(X1))))))|~((~p106(X2)&p105(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p6(X1)&~p106(X1))&p105(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p6(X1)&~p106(X1))&p105(X1))))))|~((~p105(X2)&p104(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p5(X1)&~p105(X1))&p104(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p5(X1)&~p105(X1))&p104(X1))))))|~((~p104(X2)&p103(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p4(X1)&~p104(X1))&p103(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p4(X1)&~p104(X1))&p103(X1))))))|~((~p103(X2)&p102(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p3(X1)&~p103(X1))&p102(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p3(X1)&~p103(X1))&p102(X1))))))|~((~p102(X2)&p101(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p2(X1)&~p102(X1))&p101(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p2(X1)&~p102(X1))&p101(X1))))))|~((~p101(X2)&p100(X2)))))&(((![X1]:((~r1(X2,X1)|~p16(X1))|~p115(X1))|p16(X2))&(![X1]:((~r1(X2,X1)|p16(X1))|~p115(X1))|~p16(X2)))|~p115(X2)))&(((![X1]:((~r1(X2,X1)|~p15(X1))|~p114(X1))|p15(X2))&(![X1]:((~r1(X2,X1)|p15(X1))|~p114(X1))|~p15(X2)))|~p114(X2)))&(((![X1]:((~r1(X2,X1)|~p14(X1))|~p113(X1))|p14(X2))&(![X1]:((~r1(X2,X1)|p14(X1))|~p113(X1))|~p14(X2)))|~p113(X2)))&(((![X1]:((~r1(X2,X1)|~p13(X1))|~p112(X1))|p13(X2))&(![X1]:((~r1(X2,X1)|p13(X1))|~p112(X1))|~p13(X2)))|~p112(X2)))&(((![X1]:((~r1(X2,X1)|~p12(X1))|~p111(X1))|p12(X2))&(![X1]:((~r1(X2,X1)|p12(X1))|~p111(X1))|~p12(X2)))|~p111(X2)))&(((![X1]:((~r1(X2,X1)|~p11(X1))|~p110(X1))|p11(X2))&(![X1]:((~r1(X2,X1)|p11(X1))|~p110(X1))|~p11(X2)))|~p110(X2)))&(((![X1]:((~r1(X2,X1)|~p10(X1))|~p109(X1))|p10(X2))&(![X1]:((~r1(X2,X1)|p10(X1))|~p109(X1))|~p10(X2)))|~p109(X2)))&(((![X1]:((~r1(X2,X1)|~p9(X1))|~p108(X1))|p9(X2))&(![X1]:((~r1(X2,X1)|p9(X1))|~p108(X1))|~p9(X2)))|~p108(X2)))&(((![X1]:((~r1(X2,X1)|~p8(X1))|~p107(X1))|p8(X2))&(![X1]:((~r1(X2,X1)|p8(X1))|~p107(X1))|~p8(X2)))|~p107(X2)))&(((![X1]:((~r1(X2,X1)|~p7(X1))|~p106(X1))|p7(X2))&(![X1]:((~r1(X2,X1)|p7(X1))|~p106(X1))|~p7(X2)))|~p106(X2)))&(((![X1]:((~r1(X2,X1)|~p6(X1))|~p105(X1))|p6(X2))&(![X1]:((~r1(X2,X1)|p6(X1))|~p105(X1))|~p6(X2)))|~p105(X2)))&(((![X1]:((~r1(X2,X1)|~p5(X1))|~p104(X1))|p5(X2))&(![X1]:((~r1(X2,X1)|p5(X1))|~p104(X1))|~p5(X2)))|~p104(X2)))&(((![X1]:((~r1(X2,X1)|~p4(X1))|~p103(X1))|p4(X2))&(![X1]:((~r1(X2,X1)|p4(X1))|~p103(X1))|~p4(X2)))|~p103(X2)))&(((![X1]:((~r1(X2,X1)|~p3(X1))|~p102(X1))|p3(X2))&(![X1]:((~r1(X2,X1)|p3(X1))|~p102(X1))|~p3(X2)))|~p102(X2)))&(((![X1]:((~r1(X2,X1)|~p2(X1))|~p101(X1))|p2(X2))&(![X1]:((~r1(X2,X1)|p2(X1))|~p101(X1))|~p2(X2)))|~p101(X2)))&(((![X1]:((~r1(X2,X1)|~p1(X1))|~p100(X1))|p1(X2))&(![X1]:((~r1(X2,X1)|p1(X1))|~p100(X1))|~p1(X2)))|~p100(X2)))&(p115(X2)|~p116(X2)))&(p114(X2)|~p115(X2)))&(p113(X2)|~p114(X2)))&(p112(X2)|~p113(X2)))&(p111(X2)|~p112(X2)))&(p110(X2)|~p111(X2)))&(p109(X2)|~p110(X2)))&(p108(X2)|~p109(X2)))&(p107(X2)|~p108(X2)))&(p106(X2)|~p107(X2)))&(p105(X2)|~p106(X2)))&(p104(X2)|~p105(X2)))&(p103(X2)|~p104(X2)))&(p102(X2)|~p103(X2)))&(p101(X2)|~p102(X2)))&(p100(X2)|~p101(X2))))), introduced(definition)).
fof(c_0_3, negated_conjecture, (~(~(?[X1]:~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p7(X2)))))))))))))))))|~(((![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|epred1_1(X2))))))))))))))))&~p101(X1))&p100(X1)))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_2])).
fof(c_0_4, negated_conjecture, (![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:((~r1(esk1_0,X4)|(~r1(X4,X5)|(~r1(X5,X6)|(~r1(X6,X7)|(~r1(X7,X8)|(~r1(X8,X9)|(~r1(X9,X10)|(~r1(X10,X11)|(~r1(X11,X12)|(~r1(X12,X13)|(~r1(X13,X14)|(~r1(X14,X15)|(~r1(X15,X16)|(~r1(X16,X17)|(~r1(X17,X18)|p7(X18))))))))))))))))&(((~r1(esk1_0,X19)|(~r1(X19,X20)|(~r1(X20,X21)|(~r1(X21,X22)|(~r1(X22,X23)|(~r1(X23,X24)|(~r1(X24,X25)|(~r1(X25,X26)|(~r1(X26,X27)|(~r1(X27,X28)|(~r1(X28,X29)|(~r1(X29,X30)|(~r1(X30,X31)|(~r1(X31,X32)|(~r1(X32,X33)|epred1_1(X33))))))))))))))))&~p101(esk1_0))&p100(esk1_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
fof(c_0_5, plain, (![X2]:r1(X2,X2)), inference(variable_rename,[status(thm)],[c_0_1])).
cnf(c_0_6,negated_conjecture,(epred1_1(X1)|~r1(X2,X1)|~r1(X3,X2)|~r1(X4,X3)|~r1(X5,X4)|~r1(X6,X5)|~r1(X7,X6)|~r1(X8,X7)|~r1(X9,X8)|~r1(X10,X9)|~r1(X11,X10)|~r1(X12,X11)|~r1(X13,X12)|~r1(X14,X13)|~r1(X15,X14)|~r1(esk1_0,X15)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(r1(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X1)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_9,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X1)), inference(spm,[status(thm)],[c_0_8, c_0_7])).
cnf(c_0_10,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X1)), inference(spm,[status(thm)],[c_0_9, c_0_7])).
cnf(c_0_11,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X1)), inference(spm,[status(thm)],[c_0_10, c_0_7])).
cnf(c_0_12,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X1)), inference(spm,[status(thm)],[c_0_11, c_0_7])).
cnf(c_0_13,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X1)), inference(spm,[status(thm)],[c_0_12, c_0_7])).
cnf(c_0_14,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X1)), inference(spm,[status(thm)],[c_0_13, c_0_7])).
cnf(c_0_15,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X1)), inference(spm,[status(thm)],[c_0_14, c_0_7])).
cnf(c_0_16,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X1)), inference(spm,[status(thm)],[c_0_15, c_0_7])).
cnf(c_0_17,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X1)), inference(spm,[status(thm)],[c_0_16, c_0_7])).
cnf(c_0_18,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X1)), inference(spm,[status(thm)],[c_0_17, c_0_7])).
cnf(c_0_19,negated_conjecture,(p7(X1)|~r1(X2,X1)|~r1(X3,X2)|~r1(X4,X3)|~r1(X5,X4)|~r1(X6,X5)|~r1(X7,X6)|~r1(X8,X7)|~r1(X9,X8)|~r1(X10,X9)|~r1(X11,X10)|~r1(X12,X11)|~r1(X13,X12)|~r1(X14,X13)|~r1(X15,X14)|~r1(esk1_0,X15)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_20,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_18, c_0_7])).
cnf(c_0_21,negated_conjecture,(p7(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X1)), inference(spm,[status(thm)],[c_0_19, c_0_7])).
fof(c_0_22, plain, (![X2]:(epred1_1(X2)=>((((((((((((((((((((((((((((((((((((((((((((((((~(![X1]:(~r1(X2,X1)|~(((~p16(X1)&~p116(X1))&p115(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p16(X1)&~p116(X1))&p115(X1))))))|~((~p115(X2)&p114(X2))))&((~(![X1]:(~r1(X2,X1)|~(((~p15(X1)&~p115(X1))&p114(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p15(X1)&~p115(X1))&p114(X1))))))|~((~p114(X2)&p113(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p14(X1)&~p114(X1))&p113(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p14(X1)&~p114(X1))&p113(X1))))))|~((~p113(X2)&p112(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p13(X1)&~p113(X1))&p112(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p13(X1)&~p113(X1))&p112(X1))))))|~((~p112(X2)&p111(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p12(X1)&~p112(X1))&p111(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p12(X1)&~p112(X1))&p111(X1))))))|~((~p111(X2)&p110(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p11(X1)&~p111(X1))&p110(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p11(X1)&~p111(X1))&p110(X1))))))|~((~p110(X2)&p109(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p10(X1)&~p110(X1))&p109(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p10(X1)&~p110(X1))&p109(X1))))))|~((~p109(X2)&p108(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p9(X1)&~p109(X1))&p108(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p9(X1)&~p109(X1))&p108(X1))))))|~((~p108(X2)&p107(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p8(X1)&~p108(X1))&p107(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p8(X1)&~p108(X1))&p107(X1))))))|~((~p107(X2)&p106(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p7(X1)&~p107(X1))&p106(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p7(X1)&~p107(X1))&p106(X1))))))|~((~p106(X2)&p105(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p6(X1)&~p106(X1))&p105(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p6(X1)&~p106(X1))&p105(X1))))))|~((~p105(X2)&p104(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p5(X1)&~p105(X1))&p104(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p5(X1)&~p105(X1))&p104(X1))))))|~((~p104(X2)&p103(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p4(X1)&~p104(X1))&p103(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p4(X1)&~p104(X1))&p103(X1))))))|~((~p103(X2)&p102(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p3(X1)&~p103(X1))&p102(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p3(X1)&~p103(X1))&p102(X1))))))|~((~p102(X2)&p101(X2)))))&((~(![X1]:(~r1(X2,X1)|~(((~p2(X1)&~p102(X1))&p101(X1)))))&~(![X1]:(~r1(X2,X1)|~(((p2(X1)&~p102(X1))&p101(X1))))))|~((~p101(X2)&p100(X2)))))&(((![X1]:((~r1(X2,X1)|~p16(X1))|~p115(X1))|p16(X2))&(![X1]:((~r1(X2,X1)|p16(X1))|~p115(X1))|~p16(X2)))|~p115(X2)))&(((![X1]:((~r1(X2,X1)|~p15(X1))|~p114(X1))|p15(X2))&(![X1]:((~r1(X2,X1)|p15(X1))|~p114(X1))|~p15(X2)))|~p114(X2)))&(((![X1]:((~r1(X2,X1)|~p14(X1))|~p113(X1))|p14(X2))&(![X1]:((~r1(X2,X1)|p14(X1))|~p113(X1))|~p14(X2)))|~p113(X2)))&(((![X1]:((~r1(X2,X1)|~p13(X1))|~p112(X1))|p13(X2))&(![X1]:((~r1(X2,X1)|p13(X1))|~p112(X1))|~p13(X2)))|~p112(X2)))&(((![X1]:((~r1(X2,X1)|~p12(X1))|~p111(X1))|p12(X2))&(![X1]:((~r1(X2,X1)|p12(X1))|~p111(X1))|~p12(X2)))|~p111(X2)))&(((![X1]:((~r1(X2,X1)|~p11(X1))|~p110(X1))|p11(X2))&(![X1]:((~r1(X2,X1)|p11(X1))|~p110(X1))|~p11(X2)))|~p110(X2)))&(((![X1]:((~r1(X2,X1)|~p10(X1))|~p109(X1))|p10(X2))&(![X1]:((~r1(X2,X1)|p10(X1))|~p109(X1))|~p10(X2)))|~p109(X2)))&(((![X1]:((~r1(X2,X1)|~p9(X1))|~p108(X1))|p9(X2))&(![X1]:((~r1(X2,X1)|p9(X1))|~p108(X1))|~p9(X2)))|~p108(X2)))&(((![X1]:((~r1(X2,X1)|~p8(X1))|~p107(X1))|p8(X2))&(![X1]:((~r1(X2,X1)|p8(X1))|~p107(X1))|~p8(X2)))|~p107(X2)))&(((![X1]:((~r1(X2,X1)|~p7(X1))|~p106(X1))|p7(X2))&(![X1]:((~r1(X2,X1)|p7(X1))|~p106(X1))|~p7(X2)))|~p106(X2)))&(((![X1]:((~r1(X2,X1)|~p6(X1))|~p105(X1))|p6(X2))&(![X1]:((~r1(X2,X1)|p6(X1))|~p105(X1))|~p6(X2)))|~p105(X2)))&(((![X1]:((~r1(X2,X1)|~p5(X1))|~p104(X1))|p5(X2))&(![X1]:((~r1(X2,X1)|p5(X1))|~p104(X1))|~p5(X2)))|~p104(X2)))&(((![X1]:((~r1(X2,X1)|~p4(X1))|~p103(X1))|p4(X2))&(![X1]:((~r1(X2,X1)|p4(X1))|~p103(X1))|~p4(X2)))|~p103(X2)))&(((![X1]:((~r1(X2,X1)|~p3(X1))|~p102(X1))|p3(X2))&(![X1]:((~r1(X2,X1)|p3(X1))|~p102(X1))|~p3(X2)))|~p102(X2)))&(((![X1]:((~r1(X2,X1)|~p2(X1))|~p101(X1))|p2(X2))&(![X1]:((~r1(X2,X1)|p2(X1))|~p101(X1))|~p2(X2)))|~p101(X2)))&(((![X1]:((~r1(X2,X1)|~p1(X1))|~p100(X1))|p1(X2))&(![X1]:((~r1(X2,X1)|p1(X1))|~p100(X1))|~p1(X2)))|~p100(X2)))&(p115(X2)|~p116(X2)))&(p114(X2)|~p115(X2)))&(p113(X2)|~p114(X2)))&(p112(X2)|~p113(X2)))&(p111(X2)|~p112(X2)))&(p110(X2)|~p111(X2)))&(p109(X2)|~p110(X2)))&(p108(X2)|~p109(X2)))&(p107(X2)|~p108(X2)))&(p106(X2)|~p107(X2)))&(p105(X2)|~p106(X2)))&(p104(X2)|~p105(X2)))&(p103(X2)|~p104(X2)))&(p102(X2)|~p103(X2)))&(p101(X2)|~p102(X2)))&(p100(X2)|~p101(X2))))), inference(split_equiv,[status(thm)],[c_0_2])).
cnf(c_0_23,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_20, c_0_7])).
cnf(c_0_24,negated_conjecture,(p7(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X1)), inference(spm,[status(thm)],[c_0_21, c_0_7])).
fof(c_0_25, plain, (![X3]:![X34]:![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:![X43]:![X44]:![X45]:![X46]:![X47]:![X48]:![X49]:![X50]:![X51]:![X52]:![X53]:![X54]:![X55]:![X56]:![X57]:![X58]:![X59]:![X60]:![X61]:![X62]:![X63]:![X64]:![X65]:((((((((((((((((((((((((((((((((((((((((((((((((((r1(X3,esk2_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))&((((~p16(esk2_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))&((~p116(esk2_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3)))&((p115(esk2_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))))&(((r1(X3,esk3_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))&((((p16(esk3_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))&((~p116(esk3_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3)))&((p115(esk3_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3)))))&((((r1(X3,esk4_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))&((((~p15(esk4_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))&((~p115(esk4_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3)))&((p114(esk4_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))))&(((r1(X3,esk5_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))&((((p15(esk5_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))&((~p115(esk5_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3)))&((p114(esk5_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))))))&((((r1(X3,esk6_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))&((((~p14(esk6_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))&((~p114(esk6_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3)))&((p113(esk6_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))))&(((r1(X3,esk7_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))&((((p14(esk7_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))&((~p114(esk7_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3)))&((p113(esk7_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))))))&((((r1(X3,esk8_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))&((((~p13(esk8_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))&((~p113(esk8_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3)))&((p112(esk8_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))))&(((r1(X3,esk9_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))&((((p13(esk9_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))&((~p113(esk9_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3)))&((p112(esk9_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))))))&((((r1(X3,esk10_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))&((((~p12(esk10_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))&((~p112(esk10_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3)))&((p111(esk10_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))))&(((r1(X3,esk11_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))&((((p12(esk11_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))&((~p112(esk11_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3)))&((p111(esk11_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))))))&((((r1(X3,esk12_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))&((((~p11(esk12_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))&((~p111(esk12_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3)))&((p110(esk12_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))))&(((r1(X3,esk13_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))&((((p11(esk13_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))&((~p111(esk13_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3)))&((p110(esk13_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))))))&((((r1(X3,esk14_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))&((((~p10(esk14_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))&((~p110(esk14_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3)))&((p109(esk14_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))))&(((r1(X3,esk15_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))&((((p10(esk15_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))&((~p110(esk15_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3)))&((p109(esk15_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))))))&((((r1(X3,esk16_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))&((((~p9(esk16_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))&((~p109(esk16_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3)))&((p108(esk16_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))))&(((r1(X3,esk17_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))&((((p9(esk17_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))&((~p109(esk17_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3)))&((p108(esk17_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))))))&((((r1(X3,esk18_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))&((((~p8(esk18_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))&((~p108(esk18_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3)))&((p107(esk18_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))))&(((r1(X3,esk19_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))&((((p8(esk19_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))&((~p108(esk19_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3)))&((p107(esk19_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))))))&((((r1(X3,esk20_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))&((((~p7(esk20_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))&((~p107(esk20_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3)))&((p106(esk20_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))))&(((r1(X3,esk21_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))&((((p7(esk21_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))&((~p107(esk21_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3)))&((p106(esk21_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))))))&((((r1(X3,esk22_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))&((((~p6(esk22_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))&((~p106(esk22_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3)))&((p105(esk22_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))))&(((r1(X3,esk23_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))&((((p6(esk23_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))&((~p106(esk23_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3)))&((p105(esk23_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))))))&((((r1(X3,esk24_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))&((((~p5(esk24_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))&((~p105(esk24_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3)))&((p104(esk24_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))))&(((r1(X3,esk25_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))&((((p5(esk25_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))&((~p105(esk25_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3)))&((p104(esk25_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))))))&((((r1(X3,esk26_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))&((((~p4(esk26_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))&((~p104(esk26_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3)))&((p103(esk26_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))))&(((r1(X3,esk27_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))&((((p4(esk27_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))&((~p104(esk27_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3)))&((p103(esk27_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))))))&((((r1(X3,esk28_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))&((((~p3(esk28_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))&((~p103(esk28_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3)))&((p102(esk28_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))))&(((r1(X3,esk29_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))&((((p3(esk29_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))&((~p103(esk29_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3)))&((p102(esk29_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))))))&((((r1(X3,esk30_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))&((((~p2(esk30_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))&((~p102(esk30_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3)))&((p101(esk30_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))))&(((r1(X3,esk31_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))&((((p2(esk31_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))&((~p102(esk31_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3)))&((p101(esk31_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))))))&((((((~r1(X3,X34)|~p16(X34))|~p115(X34))|p16(X3))|~p115(X3))|~epred1_1(X3))&(((((~r1(X3,X35)|p16(X35))|~p115(X35))|~p16(X3))|~p115(X3))|~epred1_1(X3))))&((((((~r1(X3,X36)|~p15(X36))|~p114(X36))|p15(X3))|~p114(X3))|~epred1_1(X3))&(((((~r1(X3,X37)|p15(X37))|~p114(X37))|~p15(X3))|~p114(X3))|~epred1_1(X3))))&((((((~r1(X3,X38)|~p14(X38))|~p113(X38))|p14(X3))|~p113(X3))|~epred1_1(X3))&(((((~r1(X3,X39)|p14(X39))|~p113(X39))|~p14(X3))|~p113(X3))|~epred1_1(X3))))&((((((~r1(X3,X40)|~p13(X40))|~p112(X40))|p13(X3))|~p112(X3))|~epred1_1(X3))&(((((~r1(X3,X41)|p13(X41))|~p112(X41))|~p13(X3))|~p112(X3))|~epred1_1(X3))))&((((((~r1(X3,X42)|~p12(X42))|~p111(X42))|p12(X3))|~p111(X3))|~epred1_1(X3))&(((((~r1(X3,X43)|p12(X43))|~p111(X43))|~p12(X3))|~p111(X3))|~epred1_1(X3))))&((((((~r1(X3,X44)|~p11(X44))|~p110(X44))|p11(X3))|~p110(X3))|~epred1_1(X3))&(((((~r1(X3,X45)|p11(X45))|~p110(X45))|~p11(X3))|~p110(X3))|~epred1_1(X3))))&((((((~r1(X3,X46)|~p10(X46))|~p109(X46))|p10(X3))|~p109(X3))|~epred1_1(X3))&(((((~r1(X3,X47)|p10(X47))|~p109(X47))|~p10(X3))|~p109(X3))|~epred1_1(X3))))&((((((~r1(X3,X48)|~p9(X48))|~p108(X48))|p9(X3))|~p108(X3))|~epred1_1(X3))&(((((~r1(X3,X49)|p9(X49))|~p108(X49))|~p9(X3))|~p108(X3))|~epred1_1(X3))))&((((((~r1(X3,X50)|~p8(X50))|~p107(X50))|p8(X3))|~p107(X3))|~epred1_1(X3))&(((((~r1(X3,X51)|p8(X51))|~p107(X51))|~p8(X3))|~p107(X3))|~epred1_1(X3))))&((((((~r1(X3,X52)|~p7(X52))|~p106(X52))|p7(X3))|~p106(X3))|~epred1_1(X3))&(((((~r1(X3,X53)|p7(X53))|~p106(X53))|~p7(X3))|~p106(X3))|~epred1_1(X3))))&((((((~r1(X3,X54)|~p6(X54))|~p105(X54))|p6(X3))|~p105(X3))|~epred1_1(X3))&(((((~r1(X3,X55)|p6(X55))|~p105(X55))|~p6(X3))|~p105(X3))|~epred1_1(X3))))&((((((~r1(X3,X56)|~p5(X56))|~p104(X56))|p5(X3))|~p104(X3))|~epred1_1(X3))&(((((~r1(X3,X57)|p5(X57))|~p104(X57))|~p5(X3))|~p104(X3))|~epred1_1(X3))))&((((((~r1(X3,X58)|~p4(X58))|~p103(X58))|p4(X3))|~p103(X3))|~epred1_1(X3))&(((((~r1(X3,X59)|p4(X59))|~p103(X59))|~p4(X3))|~p103(X3))|~epred1_1(X3))))&((((((~r1(X3,X60)|~p3(X60))|~p102(X60))|p3(X3))|~p102(X3))|~epred1_1(X3))&(((((~r1(X3,X61)|p3(X61))|~p102(X61))|~p3(X3))|~p102(X3))|~epred1_1(X3))))&((((((~r1(X3,X62)|~p2(X62))|~p101(X62))|p2(X3))|~p101(X3))|~epred1_1(X3))&(((((~r1(X3,X63)|p2(X63))|~p101(X63))|~p2(X3))|~p101(X3))|~epred1_1(X3))))&((((((~r1(X3,X64)|~p1(X64))|~p100(X64))|p1(X3))|~p100(X3))|~epred1_1(X3))&(((((~r1(X3,X65)|p1(X65))|~p100(X65))|~p1(X3))|~p100(X3))|~epred1_1(X3))))&((p115(X3)|~p116(X3))|~epred1_1(X3)))&((p114(X3)|~p115(X3))|~epred1_1(X3)))&((p113(X3)|~p114(X3))|~epred1_1(X3)))&((p112(X3)|~p113(X3))|~epred1_1(X3)))&((p111(X3)|~p112(X3))|~epred1_1(X3)))&((p110(X3)|~p111(X3))|~epred1_1(X3)))&((p109(X3)|~p110(X3))|~epred1_1(X3)))&((p108(X3)|~p109(X3))|~epred1_1(X3)))&((p107(X3)|~p108(X3))|~epred1_1(X3)))&((p106(X3)|~p107(X3))|~epred1_1(X3)))&((p105(X3)|~p106(X3))|~epred1_1(X3)))&((p104(X3)|~p105(X3))|~epred1_1(X3)))&((p103(X3)|~p104(X3))|~epred1_1(X3)))&((p102(X3)|~p103(X3))|~epred1_1(X3)))&((p101(X3)|~p102(X3))|~epred1_1(X3)))&((p100(X3)|~p101(X3))|~epred1_1(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])).
cnf(c_0_26,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_7])).
cnf(c_0_27,negated_conjecture,(p7(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X1)), inference(spm,[status(thm)],[c_0_24, c_0_7])).
cnf(c_0_28,plain,(p101(X1)|r1(X1,esk30_1(X1))|~epred1_1(X1)|~p100(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29,negated_conjecture,(epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_26, c_0_7])).
cnf(c_0_30,negated_conjecture,(p100(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_31,negated_conjecture,(~p101(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_32,negated_conjecture,(p7(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X1)), inference(spm,[status(thm)],[c_0_27, c_0_7])).
cnf(c_0_33,plain,(r1(esk1_0,esk30_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_34,plain,(p101(X1)|p101(esk30_1(X1))|~epred1_1(X1)|~p100(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35,negated_conjecture,(p7(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X1)), inference(spm,[status(thm)],[c_0_32, c_0_7])).
cnf(c_0_36,plain,(p102(X1)|r1(X1,esk28_1(X1))|~epred1_1(X1)|~p101(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37,negated_conjecture,(epred1_1(esk30_1(esk1_0))), inference(spm,[status(thm)],[c_0_26, c_0_33])).
cnf(c_0_38,plain,(p101(esk30_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_39,negated_conjecture,(p7(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X1)), inference(spm,[status(thm)],[c_0_35, c_0_7])).
cnf(c_0_40,negated_conjecture,(epred1_1(X1)|~r1(esk30_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_23, c_0_33])).
cnf(c_0_41,plain,(p102(esk30_1(esk1_0))|r1(esk30_1(esk1_0),esk28_1(esk30_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])])).
cnf(c_0_42,plain,(p102(X1)|p102(esk28_1(X1))|~epred1_1(X1)|~p101(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_43,negated_conjecture,(p7(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X1)), inference(spm,[status(thm)],[c_0_39, c_0_7])).
cnf(c_0_44,negated_conjecture,(epred1_1(X1)|~r1(esk30_1(esk1_0),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_20, c_0_33])).
cnf(c_0_45,plain,(p103(X1)|r1(X1,esk26_1(X1))|~epred1_1(X1)|~p102(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_46,negated_conjecture,(epred1_1(esk28_1(esk30_1(esk1_0)))|p102(esk30_1(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47,plain,(p102(esk28_1(esk30_1(esk1_0)))|p102(esk30_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_37]), c_0_38])])).
cnf(c_0_48,negated_conjecture,(p7(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X1)), inference(spm,[status(thm)],[c_0_43, c_0_7])).
cnf(c_0_49,negated_conjecture,(epred1_1(X1)|~r1(esk30_1(esk1_0),X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_18, c_0_33])).
cnf(c_0_50,negated_conjecture,(epred1_1(X1)|p102(esk30_1(esk1_0))|~r1(esk28_1(esk30_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_44, c_0_41])).
cnf(c_0_51,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|r1(esk28_1(esk30_1(esk1_0)),esk26_1(esk28_1(esk30_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])).
cnf(c_0_52,plain,(p103(X1)|p103(esk26_1(X1))|~epred1_1(X1)|~p102(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_53,negated_conjecture,(p7(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X1)), inference(spm,[status(thm)],[c_0_48, c_0_7])).
cnf(c_0_54,negated_conjecture,(epred1_1(X1)|~r1(esk30_1(esk1_0),X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X1)), inference(spm,[status(thm)],[c_0_17, c_0_33])).
cnf(c_0_55,plain,(epred1_1(X1)|p102(esk30_1(esk1_0))|~r1(esk28_1(esk30_1(esk1_0)),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_49, c_0_41])).
cnf(c_0_56,plain,(p104(X1)|r1(X1,esk24_1(X1))|~epred1_1(X1)|~p103(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_57,negated_conjecture,(epred1_1(esk26_1(esk28_1(esk30_1(esk1_0))))|p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58,plain,(p102(esk30_1(esk1_0))|p103(esk26_1(esk28_1(esk30_1(esk1_0))))|p103(esk28_1(esk30_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_46]), c_0_47])).
cnf(c_0_59,negated_conjecture,(p7(X1)|~r1(esk30_1(esk1_0),X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X1)), inference(spm,[status(thm)],[c_0_53, c_0_33])).
cnf(c_0_60,plain,(epred1_1(X1)|p102(esk30_1(esk1_0))|~r1(esk28_1(esk30_1(esk1_0)),X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_54, c_0_41])).
cnf(c_0_61,plain,(epred1_1(X1)|p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|~r1(esk26_1(esk28_1(esk30_1(esk1_0))),X1)), inference(spm,[status(thm)],[c_0_55, c_0_51])).
cnf(c_0_62,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|r1(esk26_1(esk28_1(esk30_1(esk1_0))),esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])).
cnf(c_0_63,plain,(p104(X1)|p104(esk24_1(X1))|~epred1_1(X1)|~p103(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_64,plain,(p102(esk30_1(esk1_0))|p7(X1)|~r1(esk28_1(esk30_1(esk1_0)),X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X1)), inference(spm,[status(thm)],[c_0_59, c_0_41])).
cnf(c_0_65,plain,(epred1_1(X1)|p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|~r1(esk26_1(esk28_1(esk30_1(esk1_0))),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_60, c_0_51])).
cnf(c_0_66,plain,(p105(X1)|r1(X1,esk22_1(X1))|~epred1_1(X1)|~p104(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_67,plain,(epred1_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))|p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_57]), c_0_58])).
cnf(c_0_69,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p7(X1)|~r1(esk26_1(esk28_1(esk30_1(esk1_0))),X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_64, c_0_51])).
cnf(c_0_70,plain,(epred1_1(X1)|p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|~r1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))),X1)), inference(spm,[status(thm)],[c_0_65, c_0_62])).
cnf(c_0_71,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|p105(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))|r1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))),esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0))))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_72,plain,(p105(X1)|p105(esk22_1(X1))|~epred1_1(X1)|~p104(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_73,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|p7(X1)|~r1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_69, c_0_62])).
cnf(c_0_74,plain,(p106(X1)|r1(X1,esk20_1(X1))|~epred1_1(X1)|~p105(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_75,plain,(epred1_1(esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0))))))|p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|p105(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_76,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|p105(esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0))))))|p105(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_67]), c_0_68])).
cnf(c_0_77,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|p105(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))|p7(X1)|~r1(esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0))))),X1)), inference(spm,[status(thm)],[c_0_73, c_0_71])).
cnf(c_0_78,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|p105(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))|p106(esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0))))))|r1(esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0))))),esk20_1(esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
cnf(c_0_79,plain,(p106(X1)|~epred1_1(X1)|~p105(X1)|~p7(esk20_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_80,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|p105(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))|p106(esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0))))))|p7(esk20_1(esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))))), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_81,plain,(p105(X1)|~epred1_1(X1)|~p104(X1)|~p106(esk22_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_82,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|p105(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))|p106(esk22_1(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0))))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_76]), c_0_75])).
cnf(c_0_83,plain,(p104(X1)|~epred1_1(X1)|~p103(X1)|~p105(esk24_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_84,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))|p105(esk24_1(esk26_1(esk28_1(esk30_1(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_68]), c_0_67])).
cnf(c_0_85,plain,(p103(X1)|~epred1_1(X1)|~p102(X1)|~p104(esk26_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_86,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))|p104(esk26_1(esk28_1(esk30_1(esk1_0))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_58]), c_0_57])).
cnf(c_0_87,plain,(p102(X1)|~epred1_1(X1)|~p101(X1)|~p103(esk28_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_88,plain,(p102(esk30_1(esk1_0))|p103(esk28_1(esk30_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_47]), c_0_46])).
cnf(c_0_89,plain,(p101(X1)|~epred1_1(X1)|~p100(X1)|~p102(esk30_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_90,plain,(p102(esk30_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_37]), c_0_38])])).
cnf(c_0_91,plain,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_29]), c_0_30])]), c_0_31]), ['proof']).
# SZS output end CNFRefutation.
EOF
