% Problem    : CSR115+8 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n137.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:48:21 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:((((((((agt(X4,X3)&attr(X3,X2))&attr(X5,X6))&obj(X4,X1))&prop(X1,britisch__1_1))&sub(X1,firma_1_1))&sub(X2,name_1_1))&subs(X4,n374bernehmen_1_1))&val(X2,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_news_1099_a19984)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((assoc(autobauer_1_1,auto__1_1)&sub(autobauer_1_1,fabrikant_1_1))&assoc(autokonzern_1_1,auto__1_1))&sub(autokonzern_1_1,firmengruppe_1_1))&attr(c1724,c1725))&poss(c1724,c2008))&prop(c1724,japanisch__1_1))&sub(c1724,autokonzern_1_1))&sub(c1725,name_1_1))&val(c1725,honda_0))&agt(c1730,c2006))&ante(c1730,c2016))&obj(c1730,c1912))&subs(c1730,annahme_1_1))&attr(c18,c19))&sub(c18,stadt__1_1))&sub(c19,name_1_1))&val(c19,tokio_0))&attr(c1912,c1913))&prop(c1912,britisch__1_1))&sub(c1912,firma_1_1))&sub(c1913,name_1_1))&val(c1913,rover_0))&attr(c2006,c2007))&prop(c2006,bundesdeutsch_1_1))&sub(c2006,autobauer_1_1))&sub(c2007,name_1_1))&val(c2007,bmw_0))&attch(c2008,c2036))&sub(c2008,anteil_1_1))&agt(c2016,c1724))&modl(c2016,wollen_0))&obj(c2016,c2008))&subs(c2016,abziehen_1_2))&quant_p3(c2036,c2013,hundertstel__1_1))&sub(c21,feb_1_1))&tupl(c34,c18,c21))&sort(autobauer_1_1,d))&sort(autobauer_1_1,io))&card(autobauer_1_1,int1))&etype(autobauer_1_1,int0))&fact(autobauer_1_1,real))&gener(autobauer_1_1,ge))&quant(autobauer_1_1,one))&refer(autobauer_1_1,refer_c))&varia(autobauer_1_1,varia_c))&sort(auto__1_1,d))&card(auto__1_1,int1))&etype(auto__1_1,int0))&fact(auto__1_1,real))&gener(auto__1_1,ge))&quant(auto__1_1,one))&refer(auto__1_1,refer_c))&varia(auto__1_1,varia_c))&sort(fabrikant_1_1,d))&sort(fabrikant_1_1,io))&card(fabrikant_1_1,int1))&etype(fabrikant_1_1,int0))&fact(fabrikant_1_1,real))&gener(fabrikant_1_1,ge))&quant(fabrikant_1_1,one))&refer(fabrikant_1_1,refer_c))&varia(fabrikant_1_1,varia_c))&sort(autokonzern_1_1,d))&sort(autokonzern_1_1,io))&card(autokonzern_1_1,int1))&etype(autokonzern_1_1,int0))&fact(autokonzern_1_1,real))&gener(autokonzern_1_1,ge))&quant(autokonzern_1_1,one))&refer(autokonzern_1_1,refer_c))&varia(autokonzern_1_1,varia_c))&sort(firmengruppe_1_1,d))&sort(firmengruppe_1_1,io))&card(firmengruppe_1_1,int1))&etype(firmengruppe_1_1,int0))&fact(firmengruppe_1_1,real))&gener(firmengruppe_1_1,ge))&quant(firmengruppe_1_1,one))&refer(firmengruppe_1_1,refer_c))&varia(firmengruppe_1_1,varia_c))&sort(c1724,d))&sort(c1724,io))&card(c1724,int1))&etype(c1724,int0))&fact(c1724,real))&gener(c1724,sp))&quant(c1724,one))&refer(c1724,det))&varia(c1724,con))&sort(c1725,na))&card(c1725,int1))&etype(c1725,int0))&fact(c1725,real))&gener(c1725,sp))&quant(c1725,one))&refer(c1725,indet))&varia(c1725,varia_c))&sort(c2008,co))&card(c2008,card_c))&etype(c2008,etype_c))&fact(c2008,real))&gener(c2008,sp))&quant(c2008,quant_c))&refer(c2008,det))&varia(c2008,varia_c))&sort(japanisch__1_1,nq))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(honda_0,fe))&sort(c1730,ad))&card(c1730,int1))&etype(c1730,int0))&fact(c1730,real))&gener(c1730,sp))&quant(c1730,one))&refer(c1730,det))&varia(c1730,con))&sort(c2006,d))&sort(c2006,io))&card(c2006,int1))&etype(c2006,int0))&fact(c2006,real))&gener(c2006,sp))&quant(c2006,one))&refer(c2006,det))&varia(c2006,con))&sort(c2016,da))&fact(c2016,real))&gener(c2016,sp))&sort(c1912,d))&sort(c1912,io))&card(c1912,int1))&etype(c1912,int0))&fact(c1912,real))&gener(c1912,sp))&quant(c1912,one))&refer(c1912,det))&varia(c1912,con))&sort(annahme_1_1,ad))&card(annahme_1_1,int1))&etype(annahme_1_1,int0))&fact(annahme_1_1,real))&gener(annahme_1_1,ge))&quant(annahme_1_1,one))&refer(annahme_1_1,refer_c))&varia(annahme_1_1,varia_c))&sort(c18,d))&sort(c18,io))&card(c18,int1))&etype(c18,int0))&fact(c18,real))&gener(c18,sp))&quant(c18,one))&refer(c18,det))&varia(c18,con))&sort(c19,na))&card(c19,int1))&etype(c19,int0))&fact(c19,real))&gener(c19,sp))&quant(c19,one))&refer(c19,indet))&varia(c19,varia_c))&sort(stadt__1_1,d))&sort(stadt__1_1,io))&card(stadt__1_1,int1))&etype(stadt__1_1,int0))&fact(stadt__1_1,real))&gener(stadt__1_1,ge))&quant(stadt__1_1,one))&refer(stadt__1_1,refer_c))&varia(stadt__1_1,varia_c))&sort(tokio_0,fe))&sort(c1913,na))&card(c1913,int1))&etype(c1913,int0))&fact(c1913,real))&gener(c1913,sp))&quant(c1913,one))&refer(c1913,indet))&varia(c1913,varia_c))&sort(britisch__1_1,nq))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(rover_0,fe))&sort(c2007,na))&card(c2007,int1))&etype(c2007,int0))&fact(c2007,real))&gener(c2007,sp))&quant(c2007,one))&refer(c2007,indet))&varia(c2007,varia_c))&sort(bundesdeutsch_1_1,tq))&sort(bmw_0,fe))&sort(c2036,co))&card(c2036,card_c))&etype(c2036,etype_c))&fact(c2036,real))&gener(c2036,gener_c))&quant(c2036,quant_c))&refer(c2036,refer_c))&varia(c2036,con))&sort(anteil_1_1,co))&card(anteil_1_1,card_c))&etype(anteil_1_1,etype_c))&fact(anteil_1_1,real))&gener(anteil_1_1,ge))&quant(anteil_1_1,quant_c))&refer(anteil_1_1,refer_c))&varia(anteil_1_1,varia_c))&sort(wollen_0,md))&fact(wollen_0,real))&gener(wollen_0,gener_c))&sort(abziehen_1_2,da))&fact(abziehen_1_2,real))&gener(abziehen_1_2,ge))&sort(c2013,nu))&card(c2013,int20))&sort(hundertstel__1_1,me))&gener(hundertstel__1_1,ge))&sort(c21,o))&card(c21,int1))&etype(c21,int0))&fact(c21,real))&gener(c21,gener_c))&quant(c21,one))&refer(c21,refer_c))&varia(c21,varia_c))&sort(feb_1_1,o))&card(feb_1_1,int1))&etype(feb_1_1,int0))&fact(feb_1_1,real))&gener(feb_1_1,ge))&quant(feb_1_1,one))&refer(feb_1_1,refer_c))&varia(feb_1_1,varia_c))&sort(c34,ent))&card(c34,card_c))&etype(c34,etype_c))&fact(c34,real))&gener(c34,gener_c))&quant(c34,quant_c))&refer(c34,refer_c))&varia(c34,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_news_1099_a19984)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:((agt(X1,X3)&chea(X2,X1))=>agt(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', chea_agt_abs__event)).
fof(c_0_3, axiom, (![X1]:![X2]:![X3]:((chea(X3,X2)&subs(X1,X2))=>?[X4]:(chea(X4,X1)&subs(X4,X3)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', chea_subs_abs__event)).
fof(c_0_4, axiom, (![X1]:![X2]:![X3]:((chea(X2,X1)&obj(X1,X3))=>obj(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', chea_obj_abs__event)).
fof(c_0_5, axiom, (chea(n374bernehmen_1_1,annahme_1_1)), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', fact_8354)).
fof(c_0_6, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:((((((((agt(X4,X3)&attr(X3,X2))&attr(X5,X6))&obj(X4,X1))&prop(X1,britisch__1_1))&sub(X1,firma_1_1))&sub(X2,name_1_1))&subs(X4,n374bernehmen_1_1))&val(X2,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_7, negated_conjecture, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((((((((~agt(X10,X9)|~attr(X9,X8))|~attr(X11,X12))|~obj(X10,X7))|~prop(X7,britisch__1_1))|~sub(X7,firma_1_1))|~sub(X8,name_1_1))|~subs(X10,n374bernehmen_1_1))|~val(X8,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_8,negated_conjecture,(~val(X1,bmw_0)|~subs(X2,n374bernehmen_1_1)|~sub(X1,name_1_1)|~sub(X3,firma_1_1)|~prop(X3,britisch__1_1)|~obj(X2,X3)|~attr(X4,X5)|~attr(X6,X1)|~agt(X2,X6)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_9,hypothesis,(val(c2007,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_10,hypothesis,(sub(c2007,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_11,hypothesis,(~prop(X1,britisch__1_1)|~attr(X2,c2007)|~attr(X3,X4)|~agt(X5,X2)|~obj(X5,X1)|~subs(X5,n374bernehmen_1_1)|~sub(X1,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
cnf(c_0_12,hypothesis,(prop(c1912,britisch__1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(sub(c1912,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:((~agt(X4,X6)|~chea(X5,X4))|agt(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
fof(c_0_15, plain, (![X5]:![X6]:![X7]:((chea(esk11_3(X5,X6,X7),X5)|(~chea(X7,X6)|~subs(X5,X6)))&(subs(esk11_3(X5,X6,X7),X7)|(~chea(X7,X6)|~subs(X5,X6))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_16,hypothesis,(~attr(X1,c2007)|~attr(X2,X3)|~agt(X4,X1)|~obj(X4,c1912)|~subs(X4,n374bernehmen_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_17,hypothesis,(attr(c2006,c2007)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_18,plain,(agt(X1,X2)|~chea(X1,X3)|~agt(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,plain,(chea(esk11_3(X1,X2,X3),X1)|~subs(X1,X2)|~chea(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:((~chea(X5,X4)|~obj(X4,X6))|obj(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
cnf(c_0_21,hypothesis,(~attr(X1,X2)|~agt(X3,c2006)|~obj(X3,c1912)|~subs(X3,n374bernehmen_1_1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22,plain,(agt(esk11_3(X1,X2,X3),X4)|~agt(X1,X4)|~chea(X3,X2)|~subs(X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23,plain,(obj(X1,X2)|~obj(X3,X2)|~chea(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,hypothesis,(~attr(X1,X2)|~agt(X3,c2006)|~obj(esk11_3(X3,X4,X5),c1912)|~chea(X5,X4)|~subs(esk11_3(X3,X4,X5),n374bernehmen_1_1)|~subs(X3,X4)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25,plain,(obj(esk11_3(X1,X2,X3),X4)|~obj(X1,X4)|~chea(X3,X2)|~subs(X1,X2)), inference(spm,[status(thm)],[c_0_23, c_0_19])).
cnf(c_0_26,hypothesis,(~attr(X1,X2)|~agt(X3,c2006)|~obj(X3,c1912)|~chea(X4,X5)|~subs(esk11_3(X3,X5,X4),n374bernehmen_1_1)|~subs(X3,X5)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_27,plain,(subs(esk11_3(X1,X2,X3),X3)|~subs(X1,X2)|~chea(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28,hypothesis,(~attr(X1,X2)|~agt(X3,c2006)|~obj(X3,c1912)|~chea(n374bernehmen_1_1,X4)|~subs(X3,X4)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_29,hypothesis,(agt(c1730,c2006)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_30,hypothesis,(obj(c1730,c1912)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_31,hypothesis,(~attr(X1,X2)|~chea(n374bernehmen_1_1,X3)|~subs(c1730,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_32,plain,(chea(n374bernehmen_1_1,annahme_1_1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_33,hypothesis,(subs(c1730,annahme_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_34,hypothesis,(attr(c1724,c1725)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_35,hypothesis,(~attr(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])])).
cnf(c_0_36,hypothesis,($false), inference(sr,[status(thm)],[c_0_34, c_0_35]), ['proof']).
# SZS output end CNFRefutation.
EOF
