% Problem    : KRS261+1 : TPTP v6.1.0. Bugfixed v5.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n105.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:14:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C18_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X9]:![X10]:(model(X9,X10)<~>model(X9,not(X10)))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', completeness)).
fof(c_0_1, axiom, (![X1]:![X2]:(?[X3]:(model(X3,X1)&model(X3,not(X2)))<=>status(X1,X2,csa))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', csa)).
fof(c_0_2, axiom, (![X7]:![X8]:(![X1]:![X2]:(status(X1,X2,X7)=>~(status(X1,X2,X8)))<=>nevera(X7,X8))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', nevera)).
fof(c_0_3, axiom, (![X1]:![X2]:(((?[X3]:model(X3,X1)&?[X4]:~(model(X4,X1)))&![X5]:(model(X5,X1)<=>model(X5,X2)))<=>status(X1,X2,eth))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', eth)).
fof(c_0_4, axiom, (![X1]:![X2]:(![X3]:(model(X3,X1)=>model(X3,X2))<=>status(X1,X2,thm))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', thm)).
fof(c_0_5, axiom, (![X7]:![X8]:(?[X1]:?[X2]:(status(X1,X2,X7)&status(X1,X2,X8))<=>mighta(X7,X8))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', mighta)).
fof(c_0_6, axiom, (?[X3]:?[X1]:?[X2]:((model(X3,X1)&~(model(X3,X2)))&?[X4]:model(X4,X2))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', non_thm_spt)).
fof(c_0_7, axiom, (![X1]:![X2]:((~(?[X3]:model(X3,X1))=>~(?[X4]:model(X4,X2)))<=>status(X1,X2,unp))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', unp)).
fof(c_0_8, conjecture, (mighta(eth,thm)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', mighta_eth_thm)).
fof(c_0_9, plain, (![X9]:![X10]:~((model(X9,X10)<=>model(X9,not(X10))))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:![X7]:![X8]:(((~model(X6,X4)|~model(X6,not(X5)))|status(X4,X5,csa))&((model(esk38_2(X7,X8),X7)|~status(X7,X8,csa))&(model(esk38_2(X7,X8),not(X8))|~status(X7,X8,csa))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])])).
fof(c_0_11, plain, (![X11]:![X12]:![X13]:![X14]:((~model(X11,X12)|~model(X11,not(X12)))&(model(X13,X14)|model(X13,not(X14))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
fof(c_0_12, plain, (![X7]:![X8]:(![X1]:![X2]:(status(X1,X2,X7)=>~status(X1,X2,X8))<=>nevera(X7,X8))), inference(fof_simplification,[status(thm)],[c_0_2])).
fof(c_0_13, plain, (![X1]:![X2]:(((?[X3]:model(X3,X1)&?[X4]:~model(X4,X1))&![X5]:(model(X5,X1)<=>model(X5,X2)))<=>status(X1,X2,eth))), inference(fof_simplification,[status(thm)],[c_0_3])).
cnf(c_0_14,plain,(status(X1,X2,csa)|~model(X3,not(X2))|~model(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(model(X1,not(X2))|model(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, (![X4]:![X5]:![X7]:![X8]:![X9]:(((model(esk10_2(X4,X5),X4)|status(X4,X5,thm))&(~model(esk10_2(X4,X5),X5)|status(X4,X5,thm)))&(~status(X7,X8,thm)|(~model(X9,X7)|model(X9,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
fof(c_0_17, plain, (![X9]:![X10]:![X13]:![X14]:![X15]:![X16]:(((status(esk49_2(X9,X10),esk50_2(X9,X10),X9)|nevera(X9,X10))&(status(esk49_2(X9,X10),esk50_2(X9,X10),X10)|nevera(X9,X10)))&(~nevera(X13,X14)|(~status(X15,X16,X13)|~status(X15,X16,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])])).
fof(c_0_18, plain, (![X6]:![X7]:![X8]:![X9]:![X11]:![X12]:![X15]:![X16]:(((((~model(esk18_2(X6,X7),X6)|~model(esk18_2(X6,X7),X7))|(~model(X8,X6)|model(X9,X6)))|status(X6,X7,eth))&(((model(esk18_2(X6,X7),X6)|model(esk18_2(X6,X7),X7))|(~model(X8,X6)|model(X9,X6)))|status(X6,X7,eth)))&(((model(esk19_2(X11,X12),X11)|~status(X11,X12,eth))&(~model(esk20_2(X11,X12),X11)|~status(X11,X12,eth)))&(((~model(X15,X11)|model(X15,X12))|~status(X11,X12,eth))&((~model(X16,X12)|model(X16,X11))|~status(X11,X12,eth)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])])).
cnf(c_0_19,plain,(status(X1,X2,csa)|model(X3,X2)|~model(X3,X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20,plain,(status(X1,X2,thm)|model(esk10_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21,plain,(model(esk38_2(X1,X2),not(X2))|~status(X1,X2,csa)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22,plain,(nevera(X1,X2)|status(esk49_2(X1,X2),esk50_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,plain,(model(X3,X2)|~status(X1,X2,eth)|~model(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(nevera(X1,X2)|status(esk49_2(X1,X2),esk50_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(model(esk38_2(X1,X2),X1)|~status(X1,X2,csa)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_26, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:(((~status(X11,X12,X9)|~status(X11,X12,X10))|mighta(X9,X10))&((status(esk43_2(X13,X14),esk44_2(X13,X14),X13)|~mighta(X13,X14))&(status(esk43_2(X13,X14),esk44_2(X13,X14),X14)|~mighta(X13,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).
cnf(c_0_27,plain,(status(X1,X2,thm)|~model(esk10_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28,plain,(status(X1,X2,thm)|status(X1,X3,csa)|model(esk10_2(X1,X2),X3)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_29,plain,(~model(X1,not(X2))|~model(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_30,plain,(nevera(csa,X1)|model(esk38_2(esk49_2(csa,X1),esk50_2(csa,X1)),not(esk50_2(csa,X1)))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_31,plain,(nevera(X1,eth)|model(X2,esk50_2(X1,eth))|~model(X2,esk49_2(X1,eth))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32,plain,(nevera(csa,X1)|model(esk38_2(esk49_2(csa,X1),esk50_2(csa,X1)),esk49_2(csa,X1))), inference(spm,[status(thm)],[c_0_25, c_0_22])).
cnf(c_0_33,plain,(mighta(X1,X2)|~status(X3,X4,X2)|~status(X3,X4,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,plain,(~status(X1,X2,X3)|~status(X1,X2,X4)|~nevera(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_35,plain,(status(X1,X2,csa)|status(X1,X2,thm)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36,plain,(nevera(csa,X1)|~model(esk38_2(esk49_2(csa,X1),esk50_2(csa,X1)),esk50_2(csa,X1))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37,plain,(nevera(csa,eth)|model(esk38_2(esk49_2(csa,eth),esk50_2(csa,eth)),esk50_2(csa,eth))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38,plain,(nevera(X1,X2)|mighta(X3,X2)|~status(esk49_2(X1,X2),esk50_2(X1,X2),X3)), inference(spm,[status(thm)],[c_0_33, c_0_24])).
cnf(c_0_39,plain,(status(X1,X2,thm)|~nevera(csa,X3)|~status(X1,X2,X3)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40,plain,(nevera(csa,eth)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41,plain,(status(esk43_2(X1,X2),esk44_2(X1,X2),X1)|~mighta(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_42,plain,(nevera(X1,X2)|mighta(X2,X2)), inference(spm,[status(thm)],[c_0_38, c_0_24])).
fof(c_0_43, plain, (?[X3]:?[X1]:?[X2]:((model(X3,X1)&~model(X3,X2))&?[X4]:model(X4,X2))), inference(fof_simplification,[status(thm)],[c_0_6])).
cnf(c_0_44,plain,(status(X1,X2,thm)|~status(X1,X2,eth)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45,plain,(nevera(X1,X2)|status(esk43_2(X2,X2),esk44_2(X2,X2),X2)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_46, plain, (![X5]:![X6]:![X7]:![X9]:![X10]:![X12]:(((~model(X7,X5)|status(X5,X6,unp))&(model(esk1_2(X5,X6),X6)|status(X5,X6,unp)))&(~status(X9,X10,unp)|(model(esk2_2(X9,X10),X9)|~model(X12,X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])])).
fof(c_0_47, plain, (((model(esk68_0,esk69_0)&~model(esk68_0,esk70_0))&model(esk71_0,esk70_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_43])])).
cnf(c_0_48,plain,(nevera(X1,eth)|status(esk43_2(eth,eth),esk44_2(eth,eth),thm)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
fof(c_0_49, negated_conjecture, (~mighta(eth,thm)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_8])])).
cnf(c_0_50,plain,(status(X1,X2,unp)|~model(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_51,plain,(model(esk71_0,esk70_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_52,plain,(nevera(X1,eth)|mighta(X2,thm)|~status(esk43_2(eth,eth),esk44_2(eth,eth),X2)), inference(spm,[status(thm)],[c_0_33, c_0_48])).
cnf(c_0_53,negated_conjecture,(~mighta(eth,thm)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_54,plain,(status(esk70_0,X1,unp)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_55,plain,(nevera(X1,eth)|nevera(X2,eth)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_45]), c_0_53])).
cnf(c_0_56,plain,(status(X1,X2,eth)|model(X3,X1)|model(esk18_2(X1,X2),X2)|model(esk18_2(X1,X2),X1)|~model(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_57,plain,(~nevera(unp,X1)|~status(esk70_0,X2,X1)), inference(spm,[status(thm)],[c_0_34, c_0_54])).
cnf(c_0_58,plain,(nevera(X1,eth)), inference(ef,[status(thm)],[c_0_55])).
cnf(c_0_59,plain,(status(esk70_0,X1,eth)|model(esk18_2(esk70_0,X1),esk70_0)|model(esk18_2(esk70_0,X1),X1)|model(X2,esk70_0)), inference(spm,[status(thm)],[c_0_56, c_0_51])).
cnf(c_0_60,plain,(~status(esk70_0,X1,eth)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_61,plain,(model(esk18_2(esk70_0,X1),esk70_0)|model(esk18_2(esk70_0,X1),X1)|model(X2,esk70_0)), inference(sr,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_62,plain,(status(X1,X2,eth)|model(X3,X1)|~model(X4,X1)|~model(esk18_2(X1,X2),X2)|~model(esk18_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_63,plain,(model(esk18_2(esk70_0,esk70_0),esk70_0)), inference(ef,[status(thm)],[c_0_61])).
cnf(c_0_64,plain,(model(X1,esk70_0)|~model(X2,esk70_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_63])]), c_0_60])).
cnf(c_0_65,plain,(~model(esk68_0,esk70_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_66,plain,(model(X1,esk70_0)), inference(spm,[status(thm)],[c_0_64, c_0_51])).
cnf(c_0_67,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_66])]), ['proof']).
# SZS output end CNFRefutation.
EOF
