% Problem    : CSR116+36 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n122.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:27:26 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:?[X8]:?[X9]:?[X10]:((((((((((((((in(X6,X7)&arg1(X4,X1))&arg2(X4,X5))&attr(X1,X2))&attr(X1,X3))&attr(X7,X8))&obj(X9,X1))&sub(X2,familiename_1_1))&sub(X3,eigenname_1_1))&sub(X5,X10))&sub(X8,name_1_1))&subr(X4,rprs_0))&val(X2,mandela_0))&val(X3,nelson_0))&val(X8,s__374dafrika_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_010_mira_wp_716)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((sub(c11,hauptagentur_1_1)&attch(c15,c11))&loc(c15,c32))&prop(c15,s__374dafrikanisch_1_1))&sub(c15,pr__344sident_1_1))&attr(c21,c22))&sub(c21,stadt__1_1))&sub(c22,name_1_1))&val(c22,kapstadt_0))&agt(c26,c5))&init(c26,c34))&obj(c26,c11))&rslt(c26,c33))&subs(c26,umbenennen_1_1))&attr(c28,c29))&sub(c28,stadt__1_1))&sub(c29,name_1_1))&val(c29,genadendal_0))&in(c32,c21))&arg1(c33,c11))&arg2(c33,c28))&subr(c33,name_0))&arg1(c34,c11))&subr(c34,name_0))&attr(c5,c6))&attr(c5,c7))&sub(c5,mensch_1_1))&sub(c6,eigenname_1_1))&val(c6,nelson_0))&sub(c7,familiename_1_1))&val(c7,mandela_0))&assoc(hauptagentur_1_1,haupt_1_1))&sub(hauptagentur_1_1,b__374ro_1_1))&sort(c11,d))&sort(c11,io))&card(c11,int1))&etype(c11,int0))&fact(c11,real))&gener(c11,sp))&quant(c11,one))&refer(c11,det))&varia(c11,con))&sort(hauptagentur_1_1,d))&sort(hauptagentur_1_1,io))&card(hauptagentur_1_1,int1))&etype(hauptagentur_1_1,int0))&fact(hauptagentur_1_1,real))&gener(hauptagentur_1_1,ge))&quant(hauptagentur_1_1,one))&refer(hauptagentur_1_1,refer_c))&varia(hauptagentur_1_1,varia_c))&sort(c15,d))&card(c15,int1))&etype(c15,int0))&fact(c15,real))&gener(c15,sp))&quant(c15,one))&refer(c15,det))&varia(c15,con))&sort(c32,l))&card(c32,int1))&etype(c32,int0))&fact(c32,real))&gener(c32,sp))&quant(c32,one))&refer(c32,det))&varia(c32,con))&sort(s__374dafrikanisch_1_1,nq))&sort(pr__344sident_1_1,d))&card(pr__344sident_1_1,int1))&etype(pr__344sident_1_1,int0))&fact(pr__344sident_1_1,real))&gener(pr__344sident_1_1,ge))&quant(pr__344sident_1_1,one))&refer(pr__344sident_1_1,refer_c))&varia(pr__344sident_1_1,varia_c))&sort(c21,d))&sort(c21,io))&card(c21,int1))&etype(c21,int0))&fact(c21,real))&gener(c21,sp))&quant(c21,one))&refer(c21,det))&varia(c21,con))&sort(c22,na))&card(c22,int1))&etype(c22,int0))&fact(c22,real))&gener(c22,sp))&quant(c22,one))&refer(c22,indet))&varia(c22,varia_c))&sort(stadt__1_1,d))&sort(stadt__1_1,io))&card(stadt__1_1,int1))&etype(stadt__1_1,int0))&fact(stadt__1_1,real))&gener(stadt__1_1,ge))&quant(stadt__1_1,one))&refer(stadt__1_1,refer_c))&varia(stadt__1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(kapstadt_0,fe))&sort(c26,da))&fact(c26,real))&gener(c26,sp))&sort(c5,d))&card(c5,int1))&etype(c5,int0))&fact(c5,real))&gener(c5,sp))&quant(c5,one))&refer(c5,det))&varia(c5,con))&sort(c34,st))&fact(c34,real))&gener(c34,sp))&sort(c33,st))&fact(c33,real))&gener(c33,sp))&sort(umbenennen_1_1,da))&fact(umbenennen_1_1,real))&gener(umbenennen_1_1,ge))&sort(c28,d))&sort(c28,io))&card(c28,int1))&etype(c28,int0))&fact(c28,real))&gener(c28,sp))&quant(c28,one))&refer(c28,det))&varia(c28,con))&sort(c29,na))&card(c29,int1))&etype(c29,int0))&fact(c29,real))&gener(c29,sp))&quant(c29,one))&refer(c29,indet))&varia(c29,varia_c))&sort(genadendal_0,fe))&sort(name_0,st))&fact(name_0,real))&gener(name_0,gener_c))&sort(c6,na))&card(c6,int1))&etype(c6,int0))&fact(c6,real))&gener(c6,sp))&quant(c6,one))&refer(c6,indet))&varia(c6,varia_c))&sort(c7,na))&card(c7,int1))&etype(c7,int0))&fact(c7,real))&gener(c7,sp))&quant(c7,one))&refer(c7,indet))&varia(c7,varia_c))&sort(mensch_1_1,d))&card(mensch_1_1,int1))&etype(mensch_1_1,int0))&fact(mensch_1_1,real))&gener(mensch_1_1,ge))&quant(mensch_1_1,one))&refer(mensch_1_1,refer_c))&varia(mensch_1_1,varia_c))&sort(eigenname_1_1,na))&card(eigenname_1_1,int1))&etype(eigenname_1_1,int0))&fact(eigenname_1_1,real))&gener(eigenname_1_1,ge))&quant(eigenname_1_1,one))&refer(eigenname_1_1,refer_c))&varia(eigenname_1_1,varia_c))&sort(nelson_0,fe))&sort(familiename_1_1,na))&card(familiename_1_1,int1))&etype(familiename_1_1,int0))&fact(familiename_1_1,real))&gener(familiename_1_1,ge))&quant(familiename_1_1,one))&refer(familiename_1_1,refer_c))&varia(familiename_1_1,varia_c))&sort(mandela_0,fe))&sort(haupt_1_1,d))&card(haupt_1_1,int1))&etype(haupt_1_1,int0))&fact(haupt_1_1,real))&gener(haupt_1_1,ge))&quant(haupt_1_1,one))&refer(haupt_1_1,refer_c))&varia(haupt_1_1,varia_c))&sort(b__374ro_1_1,d))&sort(b__374ro_1_1,io))&card(b__374ro_1_1,int1))&etype(b__374ro_1_1,int0))&fact(b__374ro_1_1,real))&gener(b__374ro_1_1,ge))&quant(b__374ro_1_1,one))&refer(b__374ro_1_1,refer_c))&varia(b__374ro_1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_010_mira_wp_716)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:((prop(X1,X2)&state_adjective_state_binding(X2,X3))=>?[X4]:?[X5]:?[X6]:(((((in(X6,X4)&attr(X4,X5))&loc(X1,X6))&sub(X4,land_1_1))&sub(X5,name_1_1))&val(X5,X3)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', state_adjective__in_state)).
fof(c_0_3, axiom, (state_adjective_state_binding(s__374dafrikanisch_1_1,s__374dafrika_0)), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', fact_8980)).
fof(c_0_4, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_5, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_6, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:?[X8]:?[X9]:?[X10]:((((((((((((((in(X6,X7)&arg1(X4,X1))&arg2(X4,X5))&attr(X1,X2))&attr(X1,X3))&attr(X7,X8))&obj(X9,X1))&sub(X2,familiename_1_1))&sub(X3,eigenname_1_1))&sub(X5,X10))&sub(X8,name_1_1))&subr(X4,rprs_0))&val(X2,mandela_0))&val(X3,nelson_0))&val(X8,s__374dafrika_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_7, negated_conjecture, (![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:((((((((((((((~in(X16,X17)|~arg1(X14,X11))|~arg2(X14,X15))|~attr(X11,X12))|~attr(X11,X13))|~attr(X17,X18))|~obj(X19,X11))|~sub(X12,familiename_1_1))|~sub(X13,eigenname_1_1))|~sub(X15,X20))|~sub(X18,name_1_1))|~subr(X14,rprs_0))|~val(X12,mandela_0))|~val(X13,nelson_0))|~val(X18,s__374dafrika_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_8,negated_conjecture,(~val(X1,s__374dafrika_0)|~val(X2,nelson_0)|~val(X3,mandela_0)|~subr(X4,rprs_0)|~sub(X1,name_1_1)|~sub(X5,X6)|~sub(X2,eigenname_1_1)|~sub(X3,familiename_1_1)|~obj(X7,X8)|~attr(X9,X1)|~attr(X8,X2)|~attr(X8,X3)|~arg2(X4,X5)|~arg1(X4,X8)|~in(X10,X9)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_9,hypothesis,(val(c7,mandela_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_10,hypothesis,(sub(c7,familiename_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_11,hypothesis,(~val(X1,nelson_0)|~val(X2,s__374dafrika_0)|~subr(X3,rprs_0)|~in(X4,X5)|~attr(X6,c7)|~attr(X5,X2)|~attr(X6,X1)|~arg2(X3,X7)|~arg1(X3,X6)|~obj(X8,X6)|~sub(X1,eigenname_1_1)|~sub(X2,name_1_1)|~sub(X7,X9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
cnf(c_0_12,hypothesis,(val(c6,nelson_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(sub(c6,eigenname_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:((((((in(esk9_3(X7,X8,X9),esk7_3(X7,X8,X9))|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9)))&(attr(esk7_3(X7,X8,X9),esk8_3(X7,X8,X9))|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(loc(X7,esk9_3(X7,X8,X9))|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(sub(esk7_3(X7,X8,X9),land_1_1)|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(sub(esk8_3(X7,X8,X9),name_1_1)|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(val(esk8_3(X7,X8,X9),X9)|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_15,hypothesis,(~val(X1,s__374dafrika_0)|~subr(X2,rprs_0)|~in(X3,X4)|~attr(X5,c7)|~attr(X5,c6)|~attr(X4,X1)|~arg2(X2,X6)|~arg1(X2,X5)|~obj(X7,X5)|~sub(X1,name_1_1)|~sub(X6,X8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_16,plain,(val(esk8_3(X3,X1,X2),X2)|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17,plain,(sub(esk8_3(X3,X1,X2),name_1_1)|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18,hypothesis,(~state_adjective_state_binding(X1,s__374dafrika_0)|~prop(X2,X1)|~subr(X3,rprs_0)|~in(X4,X5)|~attr(X5,esk8_3(X2,X1,s__374dafrika_0))|~attr(X6,c7)|~attr(X6,c6)|~arg2(X3,X7)|~arg1(X3,X6)|~obj(X8,X6)|~sub(X7,X9)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
cnf(c_0_19,plain,(attr(esk7_3(X3,X1,X2),esk8_3(X3,X1,X2))|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,hypothesis,(~state_adjective_state_binding(X1,s__374dafrika_0)|~prop(X2,X1)|~subr(X3,rprs_0)|~in(X4,esk7_3(X2,X1,s__374dafrika_0))|~attr(X5,c7)|~attr(X5,c6)|~arg2(X3,X6)|~arg1(X3,X5)|~obj(X7,X5)|~sub(X6,X8)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(in(esk9_3(X3,X1,X2),esk7_3(X3,X1,X2))|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,hypothesis,(~state_adjective_state_binding(X1,s__374dafrika_0)|~prop(X2,X1)|~subr(X3,rprs_0)|~attr(X4,c7)|~attr(X4,c6)|~arg2(X3,X5)|~arg1(X3,X4)|~obj(X6,X4)|~sub(X5,X7)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,plain,(state_adjective_state_binding(s__374dafrikanisch_1_1,s__374dafrika_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_24,hypothesis,(~prop(X1,s__374dafrikanisch_1_1)|~subr(X2,rprs_0)|~attr(X3,c7)|~attr(X3,c6)|~arg2(X2,X4)|~arg1(X2,X3)|~obj(X5,X3)|~sub(X4,X6)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,hypothesis,(prop(c15,s__374dafrikanisch_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_26, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk13_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk13_3(X7,X8,X9),esk14_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk12_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk12_3(X7,X8,X9),esk13_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk12_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk14_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk13_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk12_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
cnf(c_0_27,hypothesis,(~subr(X1,rprs_0)|~attr(X2,c7)|~attr(X2,c6)|~arg2(X1,X3)|~arg1(X1,X2)|~obj(X4,X2)|~sub(X3,X5)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28,plain,(subr(esk13_3(X1,X3,X2),rprs_0)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_29,hypothesis,(~subr(X1,sub_0)|~attr(X2,c7)|~attr(X2,c6)|~arg2(esk13_3(X1,X3,X4),X5)|~arg2(X1,X4)|~arg1(esk13_3(X1,X3,X4),X2)|~arg1(X1,X3)|~obj(X6,X2)|~sub(X5,X7)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_30,plain,(arg2(esk13_3(X1,X3,X2),esk14_3(X1,X3,X2))|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31,hypothesis,(~subr(X1,sub_0)|~attr(X2,c7)|~attr(X2,c6)|~arg2(X1,X3)|~arg1(esk13_3(X1,X4,X3),X2)|~arg1(X1,X4)|~obj(X5,X2)|~sub(esk14_3(X1,X4,X3),X6)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_32,plain,(arg1(esk13_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33,hypothesis,(~subr(X1,sub_0)|~attr(X2,c7)|~attr(X2,c6)|~arg2(X1,X3)|~arg1(X1,X2)|~obj(X4,X2)|~sub(esk14_3(X1,X2,X3),X5)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_34,plain,(sub(esk14_3(X1,X3,X2),X2)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_35, plain, (![X4]:![X5]:(((arg1(esk15_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk15_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk15_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
cnf(c_0_36,hypothesis,(~subr(X1,sub_0)|~attr(X2,c7)|~attr(X2,c6)|~arg2(X1,X3)|~arg1(X1,X2)|~obj(X4,X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_37,plain,(subr(esk15_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_38,hypothesis,(~attr(X1,c7)|~attr(X1,c6)|~arg2(esk15_2(X2,X3),X4)|~arg1(esk15_2(X2,X3),X1)|~obj(X5,X1)|~sub(X2,X3)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_39,plain,(arg2(esk15_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40,hypothesis,(~attr(X1,c7)|~attr(X1,c6)|~arg1(esk15_2(X2,X3),X1)|~obj(X4,X1)|~sub(X2,X3)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_41,plain,(arg1(esk15_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42,hypothesis,(~attr(X1,c7)|~attr(X1,c6)|~obj(X2,X1)|~sub(X1,X3)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_43,hypothesis,(attr(c5,c7)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_44,hypothesis,(attr(c5,c6)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_45,hypothesis,(~obj(X1,c5)|~sub(c5,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])])).
cnf(c_0_46,plain,(obj(esk12_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_47,hypothesis,(~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,c5)|~sub(c5,X3)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_48,hypothesis,(~arg2(esk15_2(X1,X2),X3)|~arg1(esk15_2(X1,X2),c5)|~sub(c5,X4)|~sub(X1,X2)), inference(spm,[status(thm)],[c_0_47, c_0_37])).
cnf(c_0_49,hypothesis,(~arg1(esk15_2(X1,X2),c5)|~sub(c5,X3)|~sub(X1,X2)), inference(spm,[status(thm)],[c_0_48, c_0_39])).
cnf(c_0_50,hypothesis,(~sub(c5,X1)|~sub(c5,X2)), inference(spm,[status(thm)],[c_0_49, c_0_41])).
cnf(c_0_51,hypothesis,(sub(c5,mensch_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_52,hypothesis,(~sub(c5,X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_53,hypothesis,($false), inference(sr,[status(thm)],[c_0_51, c_0_52]), ['proof']).
# SZS output end CNFRefutation.
EOF
