% Problem    : CSR115+97 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n116.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:17:01 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:(((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&sub(X7,jahr__1_1))&val(X2,bmw_0))&val(X3,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mw3_195)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((sub(c631,bau_1_1)&attch(c634,c631))&sub(c634,goggomobil_1_1))&sub(c640,man_1_1))&attr(c649,c650))&attr(c649,c651))&attr(c649,c652))&sub(c650,tag_1_1))&val(c650,c646))&sub(c651,monat_1_1))&val(c651,c647))&sub(c652,jahr__1_1))&val(c652,c648))&quant_p3(c658,c654,jahr__1_1))&quant_p3(c663,c659,monat_1_1))&attr(c703,c704))&sub(c703,firma_1_1))&sub(c704,name_1_1))&val(c704,bmw_0))&sub(c705,firma_1_1))&sub(c711,glas_1_1))&attr(c713,c714))&sub(c713,mensch_1_1))&sub(c714,eigenname_1_1))&val(c714,hans_0))&tupl_p10(c745,c631,c640,c649,c658,c663,c703,c705,c713,c711))&sort(c631,d))&card(c631,int1))&etype(c631,int0))&fact(c631,real))&gener(c631,sp))&quant(c631,one))&refer(c631,det))&varia(c631,con))&sort(bau_1_1,d))&card(bau_1_1,int1))&etype(bau_1_1,int0))&fact(bau_1_1,real))&gener(bau_1_1,ge))&quant(bau_1_1,one))&refer(bau_1_1,refer_c))&varia(bau_1_1,varia_c))&sort(c634,o))&card(c634,int1))&etype(c634,int0))&fact(c634,real))&gener(c634,sp))&quant(c634,one))&refer(c634,det))&varia(c634,con))&sort(goggomobil_1_1,o))&card(goggomobil_1_1,int1))&etype(goggomobil_1_1,int0))&fact(goggomobil_1_1,real))&gener(goggomobil_1_1,ge))&quant(goggomobil_1_1,one))&refer(goggomobil_1_1,refer_c))&varia(goggomobil_1_1,varia_c))&sort(c640,d))&card(c640,int1))&etype(c640,int0))&fact(c640,real))&gener(c640,ge))&quant(c640,one))&refer(c640,refer_c))&varia(c640,varia_c))&sort(man_1_1,d))&card(man_1_1,int1))&etype(man_1_1,int0))&fact(man_1_1,real))&gener(man_1_1,ge))&quant(man_1_1,one))&refer(man_1_1,refer_c))&varia(man_1_1,varia_c))&sort(c649,t))&card(c649,int1))&etype(c649,int0))&fact(c649,real))&gener(c649,sp))&quant(c649,one))&refer(c649,det))&varia(c649,con))&sort(c650,me))&sort(c650,oa))&sort(c650,ta))&card(c650,card_c))&etype(c650,etype_c))&fact(c650,real))&gener(c650,sp))&quant(c650,quant_c))&refer(c650,refer_c))&varia(c650,varia_c))&sort(c651,me))&sort(c651,oa))&sort(c651,ta))&card(c651,card_c))&etype(c651,etype_c))&fact(c651,real))&gener(c651,sp))&quant(c651,quant_c))&refer(c651,refer_c))&varia(c651,varia_c))&sort(c652,me))&sort(c652,oa))&sort(c652,ta))&card(c652,card_c))&etype(c652,etype_c))&fact(c652,real))&gener(c652,sp))&quant(c652,quant_c))&refer(c652,refer_c))&varia(c652,varia_c))&sort(tag_1_1,me))&sort(tag_1_1,oa))&sort(tag_1_1,ta))&card(tag_1_1,card_c))&etype(tag_1_1,etype_c))&fact(tag_1_1,real))&gener(tag_1_1,ge))&quant(tag_1_1,quant_c))&refer(tag_1_1,refer_c))&varia(tag_1_1,varia_c))&sort(c646,nu))&card(c646,int30))&sort(monat_1_1,me))&sort(monat_1_1,oa))&sort(monat_1_1,ta))&card(monat_1_1,card_c))&etype(monat_1_1,etype_c))&fact(monat_1_1,real))&gener(monat_1_1,ge))&quant(monat_1_1,quant_c))&refer(monat_1_1,refer_c))&varia(monat_1_1,varia_c))&sort(c647,nu))&card(c647,int6))&sort(jahr__1_1,me))&sort(jahr__1_1,oa))&sort(jahr__1_1,ta))&card(jahr__1_1,card_c))&etype(jahr__1_1,etype_c))&fact(jahr__1_1,real))&gener(jahr__1_1,ge))&quant(jahr__1_1,quant_c))&refer(jahr__1_1,refer_c))&varia(jahr__1_1,varia_c))&sort(c648,nu))&card(c648,int1969))&sort(c658,m))&sort(c658,ta))&card(c658,card_c))&etype(c658,etype_c))&fact(c658,real))&gener(c658,gener_c))&quant(c658,quant_c))&refer(c658,refer_c))&varia(c658,varia_c))&sort(c654,nu))&card(c654,int2))&sort(c663,m))&sort(c663,ta))&card(c663,card_c))&etype(c663,etype_c))&fact(c663,real))&gener(c663,gener_c))&quant(c663,quant_c))&refer(c663,refer_c))&varia(c663,varia_c))&sort(c659,nu))&card(c659,int6))&sort(c703,d))&sort(c703,io))&card(c703,int1))&etype(c703,int0))&fact(c703,real))&gener(c703,sp))&quant(c703,one))&refer(c703,det))&varia(c703,con))&sort(c704,na))&card(c704,int1))&etype(c704,int0))&fact(c704,real))&gener(c704,sp))&quant(c704,one))&refer(c704,indet))&varia(c704,varia_c))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(bmw_0,fe))&sort(c705,d))&sort(c705,io))&card(c705,int1))&etype(c705,int0))&fact(c705,real))&gener(c705,sp))&quant(c705,one))&refer(c705,det))&varia(c705,con))&sort(c711,s))&card(c711,int1))&etype(c711,int0))&fact(c711,real))&gener(c711,gener_c))&quant(c711,one))&refer(c711,refer_c))&varia(c711,varia_c))&sort(glas_1_1,s))&card(glas_1_1,int1))&etype(glas_1_1,int0))&fact(glas_1_1,real))&gener(glas_1_1,ge))&quant(glas_1_1,one))&refer(glas_1_1,refer_c))&varia(glas_1_1,varia_c))&sort(c713,d))&card(c713,int1))&etype(c713,int0))&fact(c713,real))&gener(c713,sp))&quant(c713,one))&refer(c713,det))&varia(c713,con))&sort(c714,na))&card(c714,int1))&etype(c714,int0))&fact(c714,real))&gener(c714,sp))&quant(c714,one))&refer(c714,indet))&varia(c714,varia_c))&sort(mensch_1_1,d))&card(mensch_1_1,int1))&etype(mensch_1_1,int0))&fact(mensch_1_1,real))&gener(mensch_1_1,ge))&quant(mensch_1_1,one))&refer(mensch_1_1,refer_c))&varia(mensch_1_1,varia_c))&sort(eigenname_1_1,na))&card(eigenname_1_1,int1))&etype(eigenname_1_1,int0))&fact(eigenname_1_1,real))&gener(eigenname_1_1,ge))&quant(eigenname_1_1,one))&refer(eigenname_1_1,refer_c))&varia(eigenname_1_1,varia_c))&sort(hans_0,fe))&sort(c745,ent))&card(c745,card_c))&etype(c745,etype_c))&fact(c745,real))&gener(c745,gener_c))&quant(c745,quant_c))&refer(c745,refer_c))&varia(c745,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mw3_195)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_3, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_4, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:(((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&sub(X7,jahr__1_1))&val(X2,bmw_0))&val(X3,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, negated_conjecture, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:(((((((((~attr(X8,X9)|~attr(X11,X10))|~attr(X13,X14))|~obj(X12,X8))|~sub(X9,name_1_1))|~sub(X8,firma_1_1))|~sub(X10,name_1_1))|~sub(X14,jahr__1_1))|~val(X9,bmw_0))|~val(X10,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_6,negated_conjecture,(~val(X1,bmw_0)|~val(X2,bmw_0)|~sub(X3,jahr__1_1)|~sub(X1,name_1_1)|~sub(X4,firma_1_1)|~sub(X2,name_1_1)|~obj(X5,X4)|~attr(X6,X3)|~attr(X7,X1)|~attr(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,hypothesis,(val(c704,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(sub(c704,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(~val(X1,bmw_0)|~attr(X2,c704)|~attr(X3,X1)|~attr(X4,X5)|~obj(X6,X2)|~sub(X2,firma_1_1)|~sub(X5,jahr__1_1)|~sub(X1,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_10,hypothesis,(~attr(X1,c704)|~attr(X2,c704)|~attr(X3,X4)|~obj(X5,X1)|~sub(X1,firma_1_1)|~sub(X4,jahr__1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_8])])).
cnf(c_0_11,hypothesis,(attr(c703,c704)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_12,hypothesis,(sub(c703,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(~attr(X1,c704)|~attr(X2,X3)|~obj(X4,c703)|~sub(X3,jahr__1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk38_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk38_3(X7,X8,X9),esk39_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk37_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk37_3(X7,X8,X9),esk38_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk37_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk39_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk38_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk37_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_15,hypothesis,(~attr(X1,X2)|~obj(X3,c703)|~sub(X2,jahr__1_1)), inference(spm,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_16,plain,(obj(esk37_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, (![X4]:![X5]:(((arg1(esk40_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk40_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk40_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_18,hypothesis,(~subr(X1,sub_0)|~attr(X2,X3)|~arg2(X1,X4)|~arg1(X1,c703)|~sub(X3,jahr__1_1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(subr(esk40_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,hypothesis,(~attr(X1,X2)|~arg2(esk40_2(X3,X4),X5)|~arg1(esk40_2(X3,X4),c703)|~sub(X2,jahr__1_1)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(arg2(esk40_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,hypothesis,(~attr(X1,X2)|~arg1(esk40_2(X3,X4),c703)|~sub(X2,jahr__1_1)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,plain,(arg1(esk40_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,hypothesis,(~attr(X1,X2)|~sub(X2,jahr__1_1)|~sub(c703,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,hypothesis,(~attr(X1,X2)|~sub(X2,jahr__1_1)), inference(spm,[status(thm)],[c_0_24, c_0_12])).
cnf(c_0_26,hypothesis,(attr(c649,c652)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_27,hypothesis,(sub(c652,jahr__1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_28,hypothesis,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation.
EOF
