% Problem    : CSR061+2 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n142.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:28:56 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_C45_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X15]:![X6]:![X16]:((disjointwith(X15,X6)&genls(X16,X15))=>disjointwith(X16,X6))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_1122)).
fof(c_0_1, axiom, (![X13]:![X14]:(disjointwith(X13,X14)=>disjointwith(X14,X13))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_1120)).
fof(c_0_2, axiom, (genls(c_tptpcol_4_114689,c_tptpcol_3_114688)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_235)).
fof(c_0_3, axiom, (disjointwith(c_tptpcol_3_98305,c_tptpcol_3_114688)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_487)).
fof(c_0_4, axiom, (genls(c_tptpcol_5_114690,c_tptpcol_4_114689)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_70)).
fof(c_0_5, axiom, (genls(c_tptpcol_6_116738,c_tptpcol_5_114690)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_100)).
fof(c_0_6, axiom, (genls(c_tptpcol_7_117762,c_tptpcol_6_116738)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_361)).
fof(c_0_7, axiom, (genls(c_tptpcol_8_117763,c_tptpcol_7_117762)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_356)).
fof(c_0_8, axiom, (genls(c_tptpcol_9_118019,c_tptpcol_8_117763)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_387)).
fof(c_0_9, axiom, (genls(c_tptpcol_10_118020,c_tptpcol_9_118019)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_160)).
fof(c_0_10, axiom, (genls(c_tptpcol_11_118084,c_tptpcol_10_118020)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_129)).
fof(c_0_11, axiom, (genls(c_tptpcol_12_118116,c_tptpcol_11_118084)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_258)).
fof(c_0_12, axiom, (genls(c_tptpcol_13_118117,c_tptpcol_12_118116)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_422)).
fof(c_0_13, axiom, (genls(c_tptpcol_14_118118,c_tptpcol_13_118117)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_499)).
fof(c_0_14, axiom, (genls(c_tptpcol_4_106497,c_tptpcol_3_98305)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_68)).
fof(c_0_15, axiom, (genls(c_tptpcol_5_110593,c_tptpcol_4_106497)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_489)).
fof(c_0_16, axiom, (genls(c_tptpcol_6_112641,c_tptpcol_5_110593)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_393)).
fof(c_0_17, axiom, (genls(c_tptpcol_7_113665,c_tptpcol_6_112641)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_215)).
fof(c_0_18, conjecture, ((mtvisible(c_timehasnoendmt)=>disjointwith(c_tptpcol_8_114177,c_tptpcol_14_118118))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', query111)).
fof(c_0_19, axiom, (genls(c_tptpcol_8_114177,c_tptpcol_7_113665)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+1.ax', ax1_401)).
fof(c_0_20, plain, (![X17]:![X18]:![X19]:((~disjointwith(X17,X18)|~genls(X19,X17))|disjointwith(X19,X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])).
fof(c_0_21, plain, (![X15]:![X16]:(~disjointwith(X15,X16)|disjointwith(X16,X15))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
cnf(c_0_22,plain,(disjointwith(X1,X2)|~genls(X1,X3)|~disjointwith(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_23,plain,(genls(c_tptpcol_4_114689,c_tptpcol_3_114688)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_24,plain,(disjointwith(X1,X2)|~disjointwith(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25,plain,(disjointwith(c_tptpcol_3_98305,c_tptpcol_3_114688)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_26,plain,(genls(c_tptpcol_5_114690,c_tptpcol_4_114689)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_27,plain,(disjointwith(c_tptpcol_4_114689,X1)|~disjointwith(c_tptpcol_3_114688,X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28,plain,(disjointwith(c_tptpcol_3_114688,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29,plain,(genls(c_tptpcol_6_116738,c_tptpcol_5_114690)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_30,plain,(disjointwith(c_tptpcol_5_114690,X1)|~disjointwith(c_tptpcol_4_114689,X1)), inference(spm,[status(thm)],[c_0_22, c_0_26])).
cnf(c_0_31,plain,(disjointwith(c_tptpcol_4_114689,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32,plain,(genls(c_tptpcol_7_117762,c_tptpcol_6_116738)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_33,plain,(disjointwith(c_tptpcol_6_116738,X1)|~disjointwith(c_tptpcol_5_114690,X1)), inference(spm,[status(thm)],[c_0_22, c_0_29])).
cnf(c_0_34,plain,(disjointwith(c_tptpcol_5_114690,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35,plain,(genls(c_tptpcol_8_117763,c_tptpcol_7_117762)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_36,plain,(disjointwith(c_tptpcol_7_117762,X1)|~disjointwith(c_tptpcol_6_116738,X1)), inference(spm,[status(thm)],[c_0_22, c_0_32])).
cnf(c_0_37,plain,(disjointwith(c_tptpcol_6_116738,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38,plain,(genls(c_tptpcol_9_118019,c_tptpcol_8_117763)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_39,plain,(disjointwith(c_tptpcol_8_117763,X1)|~disjointwith(c_tptpcol_7_117762,X1)), inference(spm,[status(thm)],[c_0_22, c_0_35])).
cnf(c_0_40,plain,(disjointwith(c_tptpcol_7_117762,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41,plain,(genls(c_tptpcol_10_118020,c_tptpcol_9_118019)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_42,plain,(disjointwith(c_tptpcol_9_118019,X1)|~disjointwith(c_tptpcol_8_117763,X1)), inference(spm,[status(thm)],[c_0_22, c_0_38])).
cnf(c_0_43,plain,(disjointwith(c_tptpcol_8_117763,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44,plain,(genls(c_tptpcol_11_118084,c_tptpcol_10_118020)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_45,plain,(disjointwith(c_tptpcol_10_118020,X1)|~disjointwith(c_tptpcol_9_118019,X1)), inference(spm,[status(thm)],[c_0_22, c_0_41])).
cnf(c_0_46,plain,(disjointwith(c_tptpcol_9_118019,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47,plain,(genls(c_tptpcol_12_118116,c_tptpcol_11_118084)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_48,plain,(disjointwith(c_tptpcol_11_118084,X1)|~disjointwith(c_tptpcol_10_118020,X1)), inference(spm,[status(thm)],[c_0_22, c_0_44])).
cnf(c_0_49,plain,(disjointwith(c_tptpcol_10_118020,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50,plain,(genls(c_tptpcol_13_118117,c_tptpcol_12_118116)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_51,plain,(disjointwith(c_tptpcol_12_118116,X1)|~disjointwith(c_tptpcol_11_118084,X1)), inference(spm,[status(thm)],[c_0_22, c_0_47])).
cnf(c_0_52,plain,(disjointwith(c_tptpcol_11_118084,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53,plain,(genls(c_tptpcol_14_118118,c_tptpcol_13_118117)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54,plain,(disjointwith(c_tptpcol_13_118117,X1)|~disjointwith(c_tptpcol_12_118116,X1)), inference(spm,[status(thm)],[c_0_22, c_0_50])).
cnf(c_0_55,plain,(disjointwith(c_tptpcol_12_118116,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56,plain,(disjointwith(c_tptpcol_14_118118,X1)|~disjointwith(c_tptpcol_13_118117,X1)), inference(spm,[status(thm)],[c_0_22, c_0_53])).
cnf(c_0_57,plain,(disjointwith(c_tptpcol_13_118117,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_58,plain,(genls(c_tptpcol_4_106497,c_tptpcol_3_98305)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_59,plain,(disjointwith(c_tptpcol_14_118118,c_tptpcol_3_98305)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_60,plain,(genls(c_tptpcol_5_110593,c_tptpcol_4_106497)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_61,plain,(disjointwith(c_tptpcol_4_106497,X1)|~disjointwith(c_tptpcol_3_98305,X1)), inference(spm,[status(thm)],[c_0_22, c_0_58])).
cnf(c_0_62,plain,(disjointwith(c_tptpcol_3_98305,c_tptpcol_14_118118)), inference(spm,[status(thm)],[c_0_24, c_0_59])).
cnf(c_0_63,plain,(genls(c_tptpcol_6_112641,c_tptpcol_5_110593)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_64,plain,(disjointwith(c_tptpcol_5_110593,X1)|~disjointwith(c_tptpcol_4_106497,X1)), inference(spm,[status(thm)],[c_0_22, c_0_60])).
cnf(c_0_65,plain,(disjointwith(c_tptpcol_4_106497,c_tptpcol_14_118118)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_66,plain,(genls(c_tptpcol_7_113665,c_tptpcol_6_112641)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_67,plain,(disjointwith(c_tptpcol_6_112641,X1)|~disjointwith(c_tptpcol_5_110593,X1)), inference(spm,[status(thm)],[c_0_22, c_0_63])).
cnf(c_0_68,plain,(disjointwith(c_tptpcol_5_110593,c_tptpcol_14_118118)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
fof(c_0_69, negated_conjecture, (~((mtvisible(c_timehasnoendmt)=>disjointwith(c_tptpcol_8_114177,c_tptpcol_14_118118)))), inference(assume_negation,[status(cth)],[c_0_18])).
cnf(c_0_70,plain,(genls(c_tptpcol_8_114177,c_tptpcol_7_113665)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_71,plain,(disjointwith(c_tptpcol_7_113665,X1)|~disjointwith(c_tptpcol_6_112641,X1)), inference(spm,[status(thm)],[c_0_22, c_0_66])).
cnf(c_0_72,plain,(disjointwith(c_tptpcol_6_112641,c_tptpcol_14_118118)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
fof(c_0_73, negated_conjecture, ((mtvisible(c_timehasnoendmt)&~disjointwith(c_tptpcol_8_114177,c_tptpcol_14_118118))), inference(fof_nnf,[status(thm)],[c_0_69])).
cnf(c_0_74,plain,(disjointwith(c_tptpcol_8_114177,X1)|~disjointwith(c_tptpcol_7_113665,X1)), inference(spm,[status(thm)],[c_0_22, c_0_70])).
cnf(c_0_75,plain,(disjointwith(c_tptpcol_7_113665,c_tptpcol_14_118118)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_76,negated_conjecture,(~disjointwith(c_tptpcol_8_114177,c_tptpcol_14_118118)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_77,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76]), ['proof']).
# SZS output end CNFRefutation.
EOF
