% Problem    : KRS215+1 : TPTP v6.1.0. Bugfixed v5.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n117.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:28:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C18_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X9]:![X10]:(model(X9,X10)<~>model(X9,not(X10)))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', completeness)).
fof(c_0_1, axiom, (?[X10]:![X9]:~(model(X9,X10))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', contradiction)).
fof(c_0_2, axiom, (![X1]:![X2]:(![X3]:(model(X3,X1)=>model(X3,X2))<=>status(X1,X2,thm))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', thm)).
fof(c_0_3, axiom, (?[X10]:![X9]:model(X9,X10)), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', tautology)).
fof(c_0_4, axiom, (![X1]:![X2]:(?[X3]:(model(X3,X1)&model(X3,not(X2)))<=>status(X1,X2,csa))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', csa)).
fof(c_0_5, axiom, (?[X3]:?[X1]:?[X2]:((model(X3,X1)&~(model(X3,X2)))&?[X4]:model(X4,X2))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', non_thm_spt)).
fof(c_0_6, axiom, (![X7]:![X8]:(?[X1]:?[X2]:(status(X1,X2,X7)&~(status(X1,X2,X8)))<=>nota(X7,X8))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', nota)).
fof(c_0_7, conjecture, (nota(csa,thm)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', nota_csa_thm)).
fof(c_0_8, plain, (![X9]:![X10]:~((model(X9,X10)<=>model(X9,not(X10))))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_9, plain, (?[X10]:![X9]:~model(X9,X10)), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_10, plain, (![X4]:![X5]:![X7]:![X8]:![X9]:(((model(esk10_2(X4,X5),X4)|status(X4,X5,thm))&(~model(esk10_2(X4,X5),X5)|status(X4,X5,thm)))&(~status(X7,X8,thm)|(~model(X9,X7)|model(X9,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
fof(c_0_11, plain, (![X11]:![X12]:![X13]:![X14]:((~model(X11,X12)|~model(X11,not(X12)))&(model(X13,X14)|model(X13,not(X14))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
fof(c_0_12, plain, (![X12]:~model(X12,esk57_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_9])])).
cnf(c_0_13,plain,(status(X1,X2,thm)|~model(esk10_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,plain,(model(X1,not(X2))|model(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,plain,(~model(X1,esk57_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(status(X1,not(X2),thm)|model(esk10_2(X1,not(X2)),X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17,plain,(model(X1,X2)|~model(X1,X3)|~status(X3,X2,thm)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,plain,(status(X1,not(esk57_0),thm)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_19, plain, (![X12]:model(X12,esk53_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_3])])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:![X7]:![X8]:(((~model(X6,X4)|~model(X6,not(X5)))|status(X4,X5,csa))&((model(esk38_2(X7,X8),X7)|~status(X7,X8,csa))&(model(esk38_2(X7,X8),not(X8))|~status(X7,X8,csa))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
cnf(c_0_21,plain,(model(X1,not(esk57_0))|~model(X1,X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22,plain,(model(X1,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_23, plain, (?[X3]:?[X1]:?[X2]:((model(X3,X1)&~model(X3,X2))&?[X4]:model(X4,X2))), inference(fof_simplification,[status(thm)],[c_0_5])).
fof(c_0_24, plain, (![X7]:![X8]:(?[X1]:?[X2]:(status(X1,X2,X7)&~status(X1,X2,X8))<=>nota(X7,X8))), inference(fof_simplification,[status(thm)],[c_0_6])).
cnf(c_0_25,plain,(status(X1,X2,csa)|~model(X3,not(X2))|~model(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,plain,(model(X1,not(esk57_0))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_27, plain, (((model(esk68_0,esk69_0)&~model(esk68_0,esk70_0))&model(esk71_0,esk70_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_23])])).
fof(c_0_28, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:(((~status(X11,X12,X9)|status(X11,X12,X10))|nota(X9,X10))&((status(esk47_2(X13,X14),esk48_2(X13,X14),X13)|~nota(X13,X14))&(~status(esk47_2(X13,X14),esk48_2(X13,X14),X14)|~nota(X13,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])])])])).
cnf(c_0_29,plain,(status(X1,esk57_0,csa)|~model(X2,X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_30,plain,(model(esk71_0,esk70_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_31, negated_conjecture, (~nota(csa,thm)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_7])])).
cnf(c_0_32,plain,(nota(X1,X2)|status(X3,X4,X2)|~status(X3,X4,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33,plain,(status(esk70_0,esk57_0,csa)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34,negated_conjecture,(~nota(csa,thm)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_35,plain,(nota(csa,X1)|status(esk70_0,esk57_0,X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36,negated_conjecture,(status(esk70_0,esk57_0,thm)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_37,negated_conjecture,(~model(X1,esk70_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_36]), c_0_15])).
cnf(c_0_38,plain,($false), inference(sr,[status(thm)],[c_0_30, c_0_37]), ['proof']).
# SZS output end CNFRefutation.
EOF
