% Problem    : CSR050+3 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n135.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:08:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, ((mtvisible(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32))=>tptptypes_7_691(c_tptpcol_16_26939,f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', query150)).
fof(c_0_1, axiom, (![X3]:![X4]:(tptptypes_8_692(X3,X4)=>tptptypes_7_691(X4,X3))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1024)).
fof(c_0_2, axiom, (![X21]:![X22]:((mtvisible(X21)&genlmt(X21,X22))=>mtvisible(X22))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_7997)).
fof(c_0_3, axiom, (![X3]:![X4]:(tptptypes_9_693(X3,X4)=>tptptypes_8_692(X3,X4))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1858)).
fof(c_0_4, axiom, (genlmt(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32),c_machinelearningspindleheadmt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1040)).
fof(c_0_5, axiom, ((mtvisible(c_currentworlddatacollectormt_nonhomocentric)=>tptptypes_9_693(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1259)).
fof(c_0_6, axiom, (genlmt(c_analystdatabasemt,c_currentworlddatacollectormt_nonhomocentric)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_3332)).
fof(c_0_7, axiom, (genlmt(c_machinelearningspindleheadmt,c_analystdatabasemt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_943)).
fof(c_0_8, negated_conjecture, (~((mtvisible(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32))=>tptptypes_7_691(c_tptpcol_16_26939,f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_9, negated_conjecture, ((mtvisible(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32))&~tptptypes_7_691(c_tptpcol_16_26939,f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma)))), inference(fof_nnf,[status(thm)],[c_0_8])).
fof(c_0_10, plain, (![X5]:![X6]:(~tptptypes_8_692(X5,X6)|tptptypes_7_691(X6,X5))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
fof(c_0_11, plain, (![X23]:![X24]:((~mtvisible(X23)|~genlmt(X23,X24))|mtvisible(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
cnf(c_0_12,negated_conjecture,(~tptptypes_7_691(c_tptpcol_16_26939,f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(tptptypes_7_691(X1,X2)|~tptptypes_8_692(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, (![X5]:![X6]:(~tptptypes_9_693(X5,X6)|tptptypes_8_692(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
cnf(c_0_15,plain,(mtvisible(X1)|~genlmt(X2,X1)|~mtvisible(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(genlmt(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32),c_machinelearningspindleheadmt)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,negated_conjecture,(mtvisible(f_contentmtofcdafromeventfn(f_urlreferentfn(f_urlfn(s_http_memberstripodcomindygalfordtriviahtm)),c_translation_32))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18,negated_conjecture,(~tptptypes_8_692(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19,plain,(tptptypes_8_692(X1,X2)|~tptptypes_9_693(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_20, plain, ((~mtvisible(c_currentworlddatacollectormt_nonhomocentric)|tptptypes_9_693(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939))), inference(fof_nnf,[status(thm)],[c_0_5])).
cnf(c_0_21,plain,(genlmt(c_analystdatabasemt,c_currentworlddatacollectormt_nonhomocentric)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22,plain,(genlmt(c_machinelearningspindleheadmt,c_analystdatabasemt)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23,plain,(mtvisible(c_machinelearningspindleheadmt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
cnf(c_0_24,negated_conjecture,(~tptptypes_9_693(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25,plain,(tptptypes_9_693(f_subcollectionofwithrelationtofn(c_ship,c_objectfoundinlocation,c_cityofbostonma),c_tptpcol_16_26939)|~mtvisible(c_currentworlddatacollectormt_nonhomocentric)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,plain,(mtvisible(c_currentworlddatacollectormt_nonhomocentric)|~mtvisible(c_analystdatabasemt)), inference(spm,[status(thm)],[c_0_15, c_0_21])).
cnf(c_0_27,plain,(mtvisible(c_analystdatabasemt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_22]), c_0_23])])).
cnf(c_0_28,negated_conjecture,(~mtvisible(c_currentworlddatacollectormt_nonhomocentric)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29,plain,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27])]), c_0_28]), ['proof']).
# SZS output end CNFRefutation.
EOF
