% Problem    : CSR062+3 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n128.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:14:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X3]:(mtvisible(c_tptp_member3205_mt)=>tptptypes_5_387(X3,c_pushingwithfingers))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', query162)).
fof(c_0_1, axiom, (![X21]:![X22]:((mtvisible(X21)&genlmt(X21,X22))=>mtvisible(X22))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_7997)).
fof(c_0_2, axiom, (genlmt(c_tptp_member3205_mt,c_tptp_spindleheadmt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_2187)).
fof(c_0_3, axiom, (![X3]:![X4]:(tptptypes_6_388(X3,X4)=>tptptypes_5_387(X3,X4))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1143)).
fof(c_0_4, axiom, ((mtvisible(c_cyclistsmt)=>tptptypes_8_390(c_pushingwithfingers,c_tptpcol_15_4027))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_968)).
fof(c_0_5, axiom, (genlmt(c_tptp_spindleheadmt,c_cyclistsmt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_4288)).
fof(c_0_6, axiom, (![X3]:![X4]:(tptptypes_7_389(X3,X4)=>tptptypes_6_388(X4,X3))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1462)).
fof(c_0_7, axiom, (![X3]:![X4]:(tptptypes_8_390(X3,X4)=>tptptypes_7_389(X3,X4))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_2531)).
fof(c_0_8, negated_conjecture, (~(?[X3]:(mtvisible(c_tptp_member3205_mt)=>tptptypes_5_387(X3,c_pushingwithfingers)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_9, plain, (![X23]:![X24]:((~mtvisible(X23)|~genlmt(X23,X24))|mtvisible(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
fof(c_0_10, negated_conjecture, (![X4]:(mtvisible(c_tptp_member3205_mt)&~tptptypes_5_387(X4,c_pushingwithfingers))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
cnf(c_0_11,plain,(mtvisible(X1)|~genlmt(X2,X1)|~mtvisible(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12,plain,(genlmt(c_tptp_member3205_mt,c_tptp_spindleheadmt)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_13,negated_conjecture,(mtvisible(c_tptp_member3205_mt)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, (![X5]:![X6]:(~tptptypes_6_388(X5,X6)|tptptypes_5_387(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
fof(c_0_15, plain, ((~mtvisible(c_cyclistsmt)|tptptypes_8_390(c_pushingwithfingers,c_tptpcol_15_4027))), inference(fof_nnf,[status(thm)],[c_0_4])).
cnf(c_0_16,plain,(genlmt(c_tptp_spindleheadmt,c_cyclistsmt)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17,plain,(mtvisible(c_tptp_spindleheadmt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_18,negated_conjecture,(~tptptypes_5_387(X1,c_pushingwithfingers)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19,plain,(tptptypes_5_387(X1,X2)|~tptptypes_6_388(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_20, plain, (![X5]:![X6]:(~tptptypes_7_389(X5,X6)|tptptypes_6_388(X6,X5))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
fof(c_0_21, plain, (![X5]:![X6]:(~tptptypes_8_390(X5,X6)|tptptypes_7_389(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
cnf(c_0_22,plain,(tptptypes_8_390(c_pushingwithfingers,c_tptpcol_15_4027)|~mtvisible(c_cyclistsmt)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,(mtvisible(c_cyclistsmt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_16]), c_0_17])])).
cnf(c_0_24,negated_conjecture,(~tptptypes_6_388(X1,c_pushingwithfingers)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25,plain,(tptptypes_6_388(X1,X2)|~tptptypes_7_389(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,plain,(tptptypes_7_389(X1,X2)|~tptptypes_8_390(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27,plain,(tptptypes_8_390(c_pushingwithfingers,c_tptpcol_15_4027)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23])])).
cnf(c_0_28,negated_conjecture,(~tptptypes_7_389(c_pushingwithfingers,X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), ['proof']).
# SZS output end CNFRefutation.
EOF
