% Problem    : CSR036+3 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n129.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:18:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___008_C18_F1_AE_Q4_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X3]:![X18]:![X19]:((disjointwith(X3,X18)&genls(X19,X18))=>disjointwith(X3,X19))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_7585)).
fof(c_0_1, axiom, (![X15]:![X16]:(disjointwith(X15,X16)=>disjointwith(X16,X15))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_7584)).
fof(c_0_2, axiom, (genls(c_tptpcol_2_2,c_tptpcol_1_1)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_3253)).
fof(c_0_3, axiom, (disjointwith(c_tptpcol_1_1,c_tptpcol_1_65536)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1248)).
fof(c_0_4, axiom, (genls(c_tptpcol_3_16386,c_tptpcol_2_2)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_460)).
fof(c_0_5, axiom, (genls(c_tptpcol_4_16387,c_tptpcol_3_16386)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_783)).
fof(c_0_6, axiom, (genls(c_tptpcol_5_20483,c_tptpcol_4_16387)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1818)).
fof(c_0_7, axiom, (genls(c_tptpcol_6_20484,c_tptpcol_5_20483)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_793)).
fof(c_0_8, axiom, (genls(c_tptpcol_7_21508,c_tptpcol_6_20484)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_2436)).
fof(c_0_9, axiom, (genls(c_tptpcol_8_22020,c_tptpcol_7_21508)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_336)).
fof(c_0_10, axiom, (genls(c_tptpcol_9_22021,c_tptpcol_8_22020)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1731)).
fof(c_0_11, axiom, (genls(c_tptpcol_10_22022,c_tptpcol_9_22021)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_70)).
fof(c_0_12, axiom, (genls(c_tptpcol_11_22023,c_tptpcol_10_22022)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_487)).
fof(c_0_13, axiom, (genls(c_tptpcol_12_22055,c_tptpcol_11_22023)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_879)).
fof(c_0_14, axiom, (genls(c_tptpcol_13_22071,c_tptpcol_12_22055)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1366)).
fof(c_0_15, axiom, (genls(c_tptpcol_14_22072,c_tptpcol_13_22071)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1975)).
fof(c_0_16, axiom, (genls(c_tptpcol_15_22076,c_tptpcol_14_22072)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1300)).
fof(c_0_17, axiom, (genls(c_tptpcol_2_65537,c_tptpcol_1_65536)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_867)).
fof(c_0_18, axiom, (genls(c_tptpcol_3_65538,c_tptpcol_2_65537)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_587)).
fof(c_0_19, axiom, (genls(c_tptpcol_4_65539,c_tptpcol_3_65538)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_176)).
fof(c_0_20, axiom, (genls(c_tptpcol_5_69635,c_tptpcol_4_65539)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_814)).
fof(c_0_21, axiom, (genls(c_tptpcol_6_71683,c_tptpcol_5_69635)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_20)).
fof(c_0_22, axiom, (genls(c_tptpcol_7_72707,c_tptpcol_6_71683)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_2287)).
fof(c_0_23, axiom, (genls(c_tptpcol_8_72708,c_tptpcol_7_72707)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_940)).
fof(c_0_24, axiom, (genls(c_tptpcol_9_72709,c_tptpcol_8_72708)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_615)).
fof(c_0_25, axiom, (genls(c_tptpcol_10_72710,c_tptpcol_9_72709)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_227)).
fof(c_0_26, axiom, (genls(c_tptpcol_11_72774,c_tptpcol_10_72710)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_759)).
fof(c_0_27, axiom, (genls(c_tptpcol_12_72775,c_tptpcol_11_72774)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_853)).
fof(c_0_28, axiom, (genls(c_tptpcol_13_72791,c_tptpcol_12_72775)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_3764)).
fof(c_0_29, axiom, (genls(c_tptpcol_14_72792,c_tptpcol_13_72791)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1760)).
fof(c_0_30, axiom, (genls(c_tptpcol_15_72793,c_tptpcol_14_72792)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_2540)).
fof(c_0_31, conjecture, ((mtvisible(c_tptp_member974_mt)=>disjointwith(c_tptpcol_15_22076,c_tptpcol_16_72795))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', query136)).
fof(c_0_32, axiom, (genls(c_tptpcol_16_72795,c_tptpcol_15_72793)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_512)).
fof(c_0_33, plain, (![X20]:![X21]:![X22]:((~disjointwith(X20,X21)|~genls(X22,X21))|disjointwith(X20,X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])).
fof(c_0_34, plain, (![X17]:![X18]:(~disjointwith(X17,X18)|disjointwith(X18,X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
cnf(c_0_35,plain,(disjointwith(X1,X2)|~genls(X2,X3)|~disjointwith(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_36,plain,(genls(c_tptpcol_2_2,c_tptpcol_1_1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_37,plain,(disjointwith(X1,X2)|~disjointwith(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_38,plain,(disjointwith(c_tptpcol_1_1,c_tptpcol_1_65536)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_39,plain,(genls(c_tptpcol_3_16386,c_tptpcol_2_2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_40,plain,(disjointwith(X1,c_tptpcol_2_2)|~disjointwith(X1,c_tptpcol_1_1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_1_1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_42,plain,(genls(c_tptpcol_4_16387,c_tptpcol_3_16386)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_43,plain,(disjointwith(X1,c_tptpcol_3_16386)|~disjointwith(X1,c_tptpcol_2_2)), inference(spm,[status(thm)],[c_0_35, c_0_39])).
cnf(c_0_44,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_2_2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45,plain,(genls(c_tptpcol_5_20483,c_tptpcol_4_16387)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_46,plain,(disjointwith(X1,c_tptpcol_4_16387)|~disjointwith(X1,c_tptpcol_3_16386)), inference(spm,[status(thm)],[c_0_35, c_0_42])).
cnf(c_0_47,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_3_16386)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_48,plain,(genls(c_tptpcol_6_20484,c_tptpcol_5_20483)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_49,plain,(disjointwith(X1,c_tptpcol_5_20483)|~disjointwith(X1,c_tptpcol_4_16387)), inference(spm,[status(thm)],[c_0_35, c_0_45])).
cnf(c_0_50,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_4_16387)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51,plain,(genls(c_tptpcol_7_21508,c_tptpcol_6_20484)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_52,plain,(disjointwith(X1,c_tptpcol_6_20484)|~disjointwith(X1,c_tptpcol_5_20483)), inference(spm,[status(thm)],[c_0_35, c_0_48])).
cnf(c_0_53,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_5_20483)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54,plain,(genls(c_tptpcol_8_22020,c_tptpcol_7_21508)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_55,plain,(disjointwith(X1,c_tptpcol_7_21508)|~disjointwith(X1,c_tptpcol_6_20484)), inference(spm,[status(thm)],[c_0_35, c_0_51])).
cnf(c_0_56,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_6_20484)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_57,plain,(genls(c_tptpcol_9_22021,c_tptpcol_8_22020)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_58,plain,(disjointwith(X1,c_tptpcol_8_22020)|~disjointwith(X1,c_tptpcol_7_21508)), inference(spm,[status(thm)],[c_0_35, c_0_54])).
cnf(c_0_59,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_7_21508)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_60,plain,(genls(c_tptpcol_10_22022,c_tptpcol_9_22021)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_61,plain,(disjointwith(X1,c_tptpcol_9_22021)|~disjointwith(X1,c_tptpcol_8_22020)), inference(spm,[status(thm)],[c_0_35, c_0_57])).
cnf(c_0_62,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_8_22020)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63,plain,(genls(c_tptpcol_11_22023,c_tptpcol_10_22022)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_64,plain,(disjointwith(X1,c_tptpcol_10_22022)|~disjointwith(X1,c_tptpcol_9_22021)), inference(spm,[status(thm)],[c_0_35, c_0_60])).
cnf(c_0_65,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_9_22021)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_66,plain,(genls(c_tptpcol_12_22055,c_tptpcol_11_22023)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_67,plain,(disjointwith(X1,c_tptpcol_11_22023)|~disjointwith(X1,c_tptpcol_10_22022)), inference(spm,[status(thm)],[c_0_35, c_0_63])).
cnf(c_0_68,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_10_22022)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_69,plain,(genls(c_tptpcol_13_22071,c_tptpcol_12_22055)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_70,plain,(disjointwith(X1,c_tptpcol_12_22055)|~disjointwith(X1,c_tptpcol_11_22023)), inference(spm,[status(thm)],[c_0_35, c_0_66])).
cnf(c_0_71,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_11_22023)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_72,plain,(genls(c_tptpcol_14_22072,c_tptpcol_13_22071)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_73,plain,(disjointwith(X1,c_tptpcol_13_22071)|~disjointwith(X1,c_tptpcol_12_22055)), inference(spm,[status(thm)],[c_0_35, c_0_69])).
cnf(c_0_74,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_12_22055)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_75,plain,(genls(c_tptpcol_15_22076,c_tptpcol_14_22072)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_76,plain,(disjointwith(X1,c_tptpcol_14_22072)|~disjointwith(X1,c_tptpcol_13_22071)), inference(spm,[status(thm)],[c_0_35, c_0_72])).
cnf(c_0_77,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_13_22071)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_78,plain,(disjointwith(X1,c_tptpcol_15_22076)|~disjointwith(X1,c_tptpcol_14_22072)), inference(spm,[status(thm)],[c_0_35, c_0_75])).
cnf(c_0_79,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_14_22072)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_80,plain,(genls(c_tptpcol_2_65537,c_tptpcol_1_65536)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_81,plain,(disjointwith(c_tptpcol_1_65536,c_tptpcol_15_22076)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_82,plain,(genls(c_tptpcol_3_65538,c_tptpcol_2_65537)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_83,plain,(disjointwith(X1,c_tptpcol_2_65537)|~disjointwith(X1,c_tptpcol_1_65536)), inference(spm,[status(thm)],[c_0_35, c_0_80])).
cnf(c_0_84,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_1_65536)), inference(spm,[status(thm)],[c_0_37, c_0_81])).
cnf(c_0_85,plain,(genls(c_tptpcol_4_65539,c_tptpcol_3_65538)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_86,plain,(disjointwith(X1,c_tptpcol_3_65538)|~disjointwith(X1,c_tptpcol_2_65537)), inference(spm,[status(thm)],[c_0_35, c_0_82])).
cnf(c_0_87,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_2_65537)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_88,plain,(genls(c_tptpcol_5_69635,c_tptpcol_4_65539)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_89,plain,(disjointwith(X1,c_tptpcol_4_65539)|~disjointwith(X1,c_tptpcol_3_65538)), inference(spm,[status(thm)],[c_0_35, c_0_85])).
cnf(c_0_90,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_3_65538)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_91,plain,(genls(c_tptpcol_6_71683,c_tptpcol_5_69635)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_92,plain,(disjointwith(X1,c_tptpcol_5_69635)|~disjointwith(X1,c_tptpcol_4_65539)), inference(spm,[status(thm)],[c_0_35, c_0_88])).
cnf(c_0_93,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_4_65539)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_94,plain,(genls(c_tptpcol_7_72707,c_tptpcol_6_71683)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_95,plain,(disjointwith(X1,c_tptpcol_6_71683)|~disjointwith(X1,c_tptpcol_5_69635)), inference(spm,[status(thm)],[c_0_35, c_0_91])).
cnf(c_0_96,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_5_69635)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_97,plain,(genls(c_tptpcol_8_72708,c_tptpcol_7_72707)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_98,plain,(disjointwith(X1,c_tptpcol_7_72707)|~disjointwith(X1,c_tptpcol_6_71683)), inference(spm,[status(thm)],[c_0_35, c_0_94])).
cnf(c_0_99,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_6_71683)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_100,plain,(genls(c_tptpcol_9_72709,c_tptpcol_8_72708)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_101,plain,(disjointwith(X1,c_tptpcol_8_72708)|~disjointwith(X1,c_tptpcol_7_72707)), inference(spm,[status(thm)],[c_0_35, c_0_97])).
cnf(c_0_102,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_7_72707)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_103,plain,(genls(c_tptpcol_10_72710,c_tptpcol_9_72709)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_104,plain,(disjointwith(X1,c_tptpcol_9_72709)|~disjointwith(X1,c_tptpcol_8_72708)), inference(spm,[status(thm)],[c_0_35, c_0_100])).
cnf(c_0_105,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_8_72708)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_106,plain,(genls(c_tptpcol_11_72774,c_tptpcol_10_72710)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_107,plain,(disjointwith(X1,c_tptpcol_10_72710)|~disjointwith(X1,c_tptpcol_9_72709)), inference(spm,[status(thm)],[c_0_35, c_0_103])).
cnf(c_0_108,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_9_72709)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_109,plain,(genls(c_tptpcol_12_72775,c_tptpcol_11_72774)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_110,plain,(disjointwith(X1,c_tptpcol_11_72774)|~disjointwith(X1,c_tptpcol_10_72710)), inference(spm,[status(thm)],[c_0_35, c_0_106])).
cnf(c_0_111,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_10_72710)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_112,plain,(genls(c_tptpcol_13_72791,c_tptpcol_12_72775)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_113,plain,(disjointwith(X1,c_tptpcol_12_72775)|~disjointwith(X1,c_tptpcol_11_72774)), inference(spm,[status(thm)],[c_0_35, c_0_109])).
cnf(c_0_114,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_11_72774)), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_115,plain,(genls(c_tptpcol_14_72792,c_tptpcol_13_72791)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_116,plain,(disjointwith(X1,c_tptpcol_13_72791)|~disjointwith(X1,c_tptpcol_12_72775)), inference(spm,[status(thm)],[c_0_35, c_0_112])).
cnf(c_0_117,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_12_72775)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_118,plain,(genls(c_tptpcol_15_72793,c_tptpcol_14_72792)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_119,plain,(disjointwith(X1,c_tptpcol_14_72792)|~disjointwith(X1,c_tptpcol_13_72791)), inference(spm,[status(thm)],[c_0_35, c_0_115])).
cnf(c_0_120,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_13_72791)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
fof(c_0_121, negated_conjecture, (~((mtvisible(c_tptp_member974_mt)=>disjointwith(c_tptpcol_15_22076,c_tptpcol_16_72795)))), inference(assume_negation,[status(cth)],[c_0_31])).
cnf(c_0_122,plain,(genls(c_tptpcol_16_72795,c_tptpcol_15_72793)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_123,plain,(disjointwith(X1,c_tptpcol_15_72793)|~disjointwith(X1,c_tptpcol_14_72792)), inference(spm,[status(thm)],[c_0_35, c_0_118])).
cnf(c_0_124,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_14_72792)), inference(spm,[status(thm)],[c_0_119, c_0_120])).
fof(c_0_125, negated_conjecture, ((mtvisible(c_tptp_member974_mt)&~disjointwith(c_tptpcol_15_22076,c_tptpcol_16_72795))), inference(fof_nnf,[status(thm)],[c_0_121])).
cnf(c_0_126,plain,(disjointwith(X1,c_tptpcol_16_72795)|~disjointwith(X1,c_tptpcol_15_72793)), inference(spm,[status(thm)],[c_0_35, c_0_122])).
cnf(c_0_127,plain,(disjointwith(c_tptpcol_15_22076,c_tptpcol_15_72793)), inference(spm,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_128,negated_conjecture,(~disjointwith(c_tptpcol_15_22076,c_tptpcol_16_72795)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_129,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_128]), ['proof']).
# SZS output end CNFRefutation.
EOF
