% Problem    : CSR115+72 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n143.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:49:31 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_wp_482)).
fof(c_0_1, hypothesis, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((agt(c2,c1)&avrt(c2,c456))&benf(c2,c451))&subs(c2,verlassen_1_3))&attr(c451,c452))&attr(c451,c453))&sub(c451,hauptaktion__344r_1_1))&sub(c452,eigenname_1_1))&val(c452,camillo_0))&sub(c453,familiename_1_1))&val(c453,castiglioni_0))&sub(c456,firma_1_1))&sub(c528,namensrechte_1_1))&attr(c595,c596))&sub(c595,firma_1_1))&sub(c596,name_1_1))&val(c596,bmw_0))&agt(c598,c1))&dircl(c598,c601))&obj(c598,c528))&semrel(c598,c2))&subs(c598,mitnehmen_1_1))&flp(c601,c595))&assoc(hauptaktion__344r_1_1,haupt_1_1))&sub(hauptaktion__344r_1_1,aktion__344r_1_1))&assoc(namensrechte_1_1,name_1_1))&sub(namensrechte_1_1,rechte_1_1))&sort(c2,da))&fact(c2,real))&gener(c2,sp))&sort(c1,co))&card(c1,card_c))&etype(c1,etype_c))&fact(c1,real))&gener(c1,sp))&quant(c1,quant_c))&refer(c1,refer_c))&varia(c1,varia_c))&sort(c456,d))&sort(c456,io))&card(c456,int1))&etype(c456,int0))&fact(c456,real))&gener(c456,sp))&quant(c456,one))&refer(c456,det))&varia(c456,con))&sort(c451,d))&card(c451,int1))&etype(c451,int0))&fact(c451,real))&gener(c451,sp))&quant(c451,one))&refer(c451,det))&varia(c451,varia_c))&sort(verlassen_1_3,da))&fact(verlassen_1_3,real))&gener(verlassen_1_3,ge))&sort(c452,na))&card(c452,int1))&etype(c452,int0))&fact(c452,real))&gener(c452,sp))&quant(c452,one))&refer(c452,indet))&varia(c452,varia_c))&sort(c453,na))&card(c453,int1))&etype(c453,int0))&fact(c453,real))&gener(c453,sp))&quant(c453,one))&refer(c453,det))&varia(c453,varia_c))&sort(hauptaktion__344r_1_1,d))&sort(hauptaktion__344r_1_1,io))&card(hauptaktion__344r_1_1,int1))&etype(hauptaktion__344r_1_1,int0))&fact(hauptaktion__344r_1_1,real))&gener(hauptaktion__344r_1_1,ge))&quant(hauptaktion__344r_1_1,one))&refer(hauptaktion__344r_1_1,refer_c))&varia(hauptaktion__344r_1_1,varia_c))&sort(eigenname_1_1,na))&card(eigenname_1_1,int1))&etype(eigenname_1_1,int0))&fact(eigenname_1_1,real))&gener(eigenname_1_1,ge))&quant(eigenname_1_1,one))&refer(eigenname_1_1,refer_c))&varia(eigenname_1_1,varia_c))&sort(camillo_0,fe))&sort(familiename_1_1,na))&card(familiename_1_1,int1))&etype(familiename_1_1,int0))&fact(familiename_1_1,real))&gener(familiename_1_1,ge))&quant(familiename_1_1,one))&refer(familiename_1_1,refer_c))&varia(familiename_1_1,varia_c))&sort(castiglioni_0,fe))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(c528,d))&sort(c528,io))&card(c528,int1))&etype(c528,int1))&fact(c528,real))&gener(c528,sp))&quant(c528,one))&refer(c528,det))&varia(c528,con))&sort(namensrechte_1_1,d))&sort(namensrechte_1_1,io))&card(namensrechte_1_1,card_c))&etype(namensrechte_1_1,int1))&fact(namensrechte_1_1,real))&gener(namensrechte_1_1,ge))&quant(namensrechte_1_1,quant_c))&refer(namensrechte_1_1,refer_c))&varia(namensrechte_1_1,varia_c))&sort(c595,d))&sort(c595,io))&card(c595,int1))&etype(c595,int0))&fact(c595,real))&gener(c595,sp))&quant(c595,one))&refer(c595,det))&varia(c595,con))&sort(c596,na))&card(c596,int1))&etype(c596,int0))&fact(c596,real))&gener(c596,sp))&quant(c596,one))&refer(c596,indet))&varia(c596,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(bmw_0,fe))&sort(c598,da))&fact(c598,real))&gener(c598,sp))&sort(c601,l))&card(c601,int1))&etype(c601,int0))&fact(c601,real))&gener(c601,sp))&quant(c601,one))&refer(c601,det))&varia(c601,con))&sort(mitnehmen_1_1,da))&fact(mitnehmen_1_1,real))&gener(mitnehmen_1_1,ge))&sort(haupt_1_1,d))&card(haupt_1_1,int1))&etype(haupt_1_1,int0))&fact(haupt_1_1,real))&gener(haupt_1_1,ge))&quant(haupt_1_1,one))&refer(haupt_1_1,refer_c))&varia(haupt_1_1,varia_c))&sort(aktion__344r_1_1,d))&sort(aktion__344r_1_1,io))&card(aktion__344r_1_1,int1))&etype(aktion__344r_1_1,int0))&fact(aktion__344r_1_1,real))&gener(aktion__344r_1_1,ge))&quant(aktion__344r_1_1,one))&refer(aktion__344r_1_1,refer_c))&varia(aktion__344r_1_1,varia_c))&sort(rechte_1_1,d))&sort(rechte_1_1,io))&card(rechte_1_1,card_c))&etype(rechte_1_1,int1))&fact(rechte_1_1,real))&gener(rechte_1_1,ge))&quant(rechte_1_1,quant_c))&refer(rechte_1_1,refer_c))&varia(rechte_1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_wp_482)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_3, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_4, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, negated_conjecture, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((((~attr(X8,X9)|~attr(X11,X10))|~attr(X12,X13))|~obj(X14,X8))|~sub(X9,name_1_1))|~sub(X8,firma_1_1))|~sub(X10,name_1_1))|~val(X9,bmw_0))|~val(X10,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_6,negated_conjecture,(~val(X1,bmw_0)|~val(X2,bmw_0)|~sub(X1,name_1_1)|~sub(X3,firma_1_1)|~sub(X2,name_1_1)|~obj(X4,X3)|~attr(X5,X6)|~attr(X7,X1)|~attr(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,hypothesis,(val(c596,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(sub(c596,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(~val(X1,bmw_0)|~attr(X2,c596)|~attr(X3,X1)|~attr(X4,X5)|~obj(X6,X2)|~sub(X2,firma_1_1)|~sub(X1,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_10,hypothesis,(~attr(X1,c596)|~attr(X2,c596)|~attr(X3,X4)|~obj(X5,X1)|~sub(X1,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_8])])).
cnf(c_0_11,hypothesis,(attr(c595,c596)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_12,hypothesis,(sub(c595,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(~attr(X1,c596)|~attr(X2,X3)|~obj(X4,c595)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk39_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk39_3(X7,X8,X9),esk40_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk38_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk38_3(X7,X8,X9),esk39_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk38_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk40_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk39_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk38_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_15,hypothesis,(~attr(X1,X2)|~obj(X3,c595)), inference(spm,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_16,plain,(obj(esk38_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, (![X4]:![X5]:(((arg1(esk41_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk41_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk41_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_18,hypothesis,(~subr(X1,sub_0)|~attr(X2,X3)|~arg2(X1,X4)|~arg1(X1,c595)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(subr(esk41_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,hypothesis,(~attr(X1,X2)|~arg2(esk41_2(X3,X4),X5)|~arg1(esk41_2(X3,X4),c595)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(arg2(esk41_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,hypothesis,(~attr(X1,X2)|~arg1(esk41_2(X3,X4),c595)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,plain,(arg1(esk41_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,hypothesis,(~attr(X1,X2)|~sub(c595,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,hypothesis,(attr(c451,c452)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_26,hypothesis,(~attr(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_12])).
cnf(c_0_27,hypothesis,($false), inference(sr,[status(thm)],[c_0_25, c_0_26]), ['proof']).
# SZS output end CNFRefutation.
EOF
