% Problem    : CSR071+3 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n100.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:45:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, ((mtvisible(c_tptp_spindlecollectormt)=>tptpofobject(c_tptpnavypersonnel_3,f_tptpquantityfn_6(n_414)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', query171)).
fof(c_0_1, axiom, (![X2]:((mtvisible(c_tptp_member698_mt)&militaryperson(X2))=>tptpofobject(X2,f_tptpquantityfn_6(n_414)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_897)).
fof(c_0_2, axiom, (![X1]:(navypersonnel(X1)=>militaryperson(X1))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1623)).
fof(c_0_3, axiom, ((mtvisible(c_tptp_member1672_mt)=>navypersonnel(c_tptpnavypersonnel_3))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1657)).
fof(c_0_4, axiom, (![X21]:![X22]:((mtvisible(X21)&genlmt(X21,X22))=>mtvisible(X22))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_7997)).
fof(c_0_5, axiom, (genlmt(c_tptp_spindlecollectormt,c_tptp_member698_mt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1450)).
fof(c_0_6, axiom, (genlmt(c_tptp_spindlecollectormt,c_tptp_member1672_mt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_89)).
fof(c_0_7, negated_conjecture, (~((mtvisible(c_tptp_spindlecollectormt)=>tptpofobject(c_tptpnavypersonnel_3,f_tptpquantityfn_6(n_414))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_8, negated_conjecture, ((mtvisible(c_tptp_spindlecollectormt)&~tptpofobject(c_tptpnavypersonnel_3,f_tptpquantityfn_6(n_414)))), inference(fof_nnf,[status(thm)],[c_0_7])).
fof(c_0_9, plain, (![X3]:((~mtvisible(c_tptp_member698_mt)|~militaryperson(X3))|tptpofobject(X3,f_tptpquantityfn_6(n_414)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
fof(c_0_10, plain, (![X2]:(~navypersonnel(X2)|militaryperson(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
fof(c_0_11, plain, ((~mtvisible(c_tptp_member1672_mt)|navypersonnel(c_tptpnavypersonnel_3))), inference(fof_nnf,[status(thm)],[c_0_3])).
cnf(c_0_12,negated_conjecture,(~tptpofobject(c_tptpnavypersonnel_3,f_tptpquantityfn_6(n_414))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(tptpofobject(X1,f_tptpquantityfn_6(n_414))|~militaryperson(X1)|~mtvisible(c_tptp_member698_mt)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(militaryperson(X1)|~navypersonnel(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(navypersonnel(c_tptpnavypersonnel_3)|~mtvisible(c_tptp_member1672_mt)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, (![X23]:![X24]:((~mtvisible(X23)|~genlmt(X23,X24))|mtvisible(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
cnf(c_0_17,negated_conjecture,(~militaryperson(c_tptpnavypersonnel_3)|~mtvisible(c_tptp_member698_mt)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18,plain,(militaryperson(c_tptpnavypersonnel_3)|~mtvisible(c_tptp_member1672_mt)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19,plain,(mtvisible(X1)|~genlmt(X2,X1)|~mtvisible(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20,plain,(genlmt(c_tptp_spindlecollectormt,c_tptp_member698_mt)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_21,negated_conjecture,(mtvisible(c_tptp_spindlecollectormt)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_22,negated_conjecture,(~mtvisible(c_tptp_member698_mt)|~mtvisible(c_tptp_member1672_mt)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23,plain,(mtvisible(c_tptp_member698_mt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
cnf(c_0_24,plain,(genlmt(c_tptp_spindlecollectormt,c_tptp_member1672_mt)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25,negated_conjecture,(~mtvisible(c_tptp_member1672_mt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23])])).
cnf(c_0_26,plain,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_24]), c_0_21])]), c_0_25]), ['proof']).
# SZS output end CNFRefutation.
EOF
