% Problem    : LCL658+1.001 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n114.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:22:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S0S
# and selection function SelectComplexG.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (~(?[X1]:~((~(![X2]:(~(r1(X1,X2))|((((((![X1]:(~(r1(X2,X1))|p1(X1))|~(p1(X2)))|![X1]:(~(r1(X2,X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2))))))|~(![X1]:(((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1)))|~(![X2]:((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1))))|~((![X1]:(~(r1(X2,X1))|p1(X1))|~(p1(X2)))))))))&((p1(X2)|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|~(![X1]:((~(r1(X2,X1))|p1(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2))))))))&![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))&(![X1]:(~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p1(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1))))))|~(![X1]:((~(r1(X2,X1))|p1(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2))))))))))|~(~(![X2]:(((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|~(![X1]:((~(r1(X2,X1))|p1(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2))))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', main)).
fof(c_0_1, plain, (![X2]:(epred1_1(X2)<=>((((((![X1]:(~r1(X2,X1)|p1(X1))|~p1(X2))|![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2)))))|~(![X1]:(((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1)))|~((![X1]:(~r1(X2,X1)|p1(X1))|~p1(X2))))))))&((p1(X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2)))))))&![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))|~(![X2]:(~r1(X1,X2)|p1(X2)))))&(![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))))))))), introduced(definition)).
fof(c_0_2, negated_conjecture, (~(~(?[X1]:~((~(![X2]:(~r1(X1,X2)|epred1_1(X2)))|~(~(![X2]:(((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))))))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_1])).
fof(c_0_3, plain, (![X2]:(epred1_1(X2)=>((((((![X1]:(~r1(X2,X1)|p1(X1))|~p1(X2))|![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2)))))|~(![X1]:(((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1)))|~((![X1]:(~r1(X2,X1)|p1(X1))|~p1(X2))))))))&((p1(X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2)))))))&![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))|~(![X2]:(~r1(X1,X2)|p1(X2)))))&(![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))))))))), inference(split_equiv,[status(thm)],[c_0_1])).
fof(c_0_4, negated_conjecture, (![X4]:![X8]:![X9]:((~r1(esk1_0,X4)|epred1_1(X4))&(((r1(esk1_0,esk2_0)&(r1(esk2_0,esk3_0)&~p1(esk3_0)))&(r1(esk2_0,esk4_0)&(~r1(esk4_0,X8)|p1(X8))))&(((r1(X9,esk5_1(X9))|(~r1(esk2_0,X9)|p1(X9)))&((r1(esk5_1(X9),esk6_1(X9))|(~r1(esk2_0,X9)|p1(X9)))&(~p1(esk6_1(X9))|(~r1(esk2_0,X9)|p1(X9)))))&(p1(esk5_1(X9))|(~r1(esk2_0,X9)|p1(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
fof(c_0_5, plain, (![X3]:![X4]:![X5]:![X10]:![X11]:![X12]:![X14]:![X17]:![X18]:![X19]:![X20]:![X21]:![X23]:![X24]:![X28]:![X29]:((((((((((r1(X3,esk9_1(X3))|((r1(X5,esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((r1(esk9_1(X3),esk10_1(X3))|((r1(X5,esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&((~p1(esk10_1(X3))|((r1(X5,esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((p1(esk9_1(X3))|((r1(X5,esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((((r1(X10,esk11_2(X3,X10))|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(X5,esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((~p1(esk11_2(X3,X10))|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(X5,esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((p1(X10)|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(X5,esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((((((r1(X3,esk9_1(X3))|((r1(esk7_2(X3,X5),esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((r1(esk9_1(X3),esk10_1(X3))|((r1(esk7_2(X3,X5),esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&((~p1(esk10_1(X3))|((r1(esk7_2(X3,X5),esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((p1(esk9_1(X3))|((r1(esk7_2(X3,X5),esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((((r1(X10,esk11_2(X3,X10))|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(esk7_2(X3,X5),esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((~p1(esk11_2(X3,X10))|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(esk7_2(X3,X5),esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((p1(X10)|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(esk7_2(X3,X5),esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&(((((r1(X3,esk9_1(X3))|((~p1(esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((r1(esk9_1(X3),esk10_1(X3))|((~p1(esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&((~p1(esk10_1(X3))|((~p1(esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((p1(esk9_1(X3))|((~p1(esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((((r1(X10,esk11_2(X3,X10))|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((~p1(esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((~p1(esk11_2(X3,X10))|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((~p1(esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((p1(X10)|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((~p1(esk8_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))))&(((((r1(X3,esk9_1(X3))|((p1(esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((r1(esk9_1(X3),esk10_1(X3))|((p1(esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&((~p1(esk10_1(X3))|((p1(esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((p1(esk9_1(X3))|((p1(esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((((r1(X10,esk11_2(X3,X10))|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((p1(esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((~p1(esk11_2(X3,X10))|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((p1(esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((p1(X10)|(~r1(esk9_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((p1(esk7_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))))&(((((r1(X3,esk13_1(X3))|((r1(X14,esk12_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3))&((~p1(esk13_1(X3))|((r1(X14,esk12_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3)))&((((~r1(esk13_1(X3),X17)|(~r1(X17,X18)|p1(X18)))|~p1(X17))|((r1(X14,esk12_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3)))&((((r1(X3,esk13_1(X3))|((~p1(esk12_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3))&((~p1(esk13_1(X3))|((~p1(esk12_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3)))&((((~r1(esk13_1(X3),X17)|(~r1(X17,X18)|p1(X18)))|~p1(X17))|((~p1(esk12_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3)))))&(((r1(X19,esk14_2(X3,X19))|(~r1(X3,X19)|(~r1(X19,X20)|(~r1(X20,X21)|p1(X21)))))|~epred1_1(X3))&((~p1(esk14_2(X3,X19))|(~r1(X3,X19)|(~r1(X19,X20)|(~r1(X20,X21)|p1(X21)))))|~epred1_1(X3))))&((((((r1(X3,esk17_1(X3))|((r1(X24,esk15_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3))&((~p1(esk17_1(X3))|((r1(X24,esk15_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((~r1(esk17_1(X3),X28)|(~r1(X28,X29)|p1(X29)))|~p1(X28))|((r1(X24,esk15_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&(((((r1(X3,esk17_1(X3))|((r1(esk15_3(X3,X23,X24),esk16_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3))&((~p1(esk17_1(X3))|((r1(esk15_3(X3,X23,X24),esk16_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((~r1(esk17_1(X3),X28)|(~r1(X28,X29)|p1(X29)))|~p1(X28))|((r1(esk15_3(X3,X23,X24),esk16_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((r1(X3,esk17_1(X3))|((~p1(esk16_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3))&((~p1(esk17_1(X3))|((~p1(esk16_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((~r1(esk17_1(X3),X28)|(~r1(X28,X29)|p1(X29)))|~p1(X28))|((~p1(esk16_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))))&((((r1(X3,esk17_1(X3))|((p1(esk15_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3))&((~p1(esk17_1(X3))|((p1(esk15_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((~r1(esk17_1(X3),X28)|(~r1(X28,X29)|p1(X29)))|~p1(X28))|((p1(esk15_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_6,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,negated_conjecture,(r1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,plain,(p1(X1)|r1(X2,esk12_2(X1,X2))|r1(X1,esk13_1(X1))|~epred1_1(X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,negated_conjecture,(r1(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(epred1_1(esk2_0)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_11,negated_conjecture,(p1(X1)|~r1(esk4_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12,negated_conjecture,(p1(esk2_0)|r1(esk4_0,esk12_2(esk2_0,esk4_0))|r1(esk2_0,esk13_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
cnf(c_0_13,plain,(p1(X1)|r1(X1,esk13_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~p1(esk12_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14,negated_conjecture,(p1(esk12_2(esk2_0,esk4_0))|p1(esk2_0)|r1(esk2_0,esk13_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15,negated_conjecture,(p1(X1)|r1(X1,esk5_1(X1))|~r1(esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,plain,(p1(esk2_0)|r1(esk2_0,esk13_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_10]), c_0_9])])).
cnf(c_0_17,negated_conjecture,(p1(X1)|p1(esk5_1(X1))|~r1(esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18,plain,(p1(X1)|r1(X2,esk12_2(X1,X2))|p1(X4)|~epred1_1(X1)|~r1(X1,X2)|~p1(X3)|~r1(X3,X4)|~r1(esk13_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19,negated_conjecture,(p1(esk13_1(esk2_0))|p1(esk2_0)|r1(esk13_1(esk2_0),esk5_1(esk13_1(esk2_0)))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20,negated_conjecture,(p1(esk5_1(esk13_1(esk2_0)))|p1(esk13_1(esk2_0))|p1(esk2_0)), inference(spm,[status(thm)],[c_0_17, c_0_16])).
cnf(c_0_21,negated_conjecture,(p1(X1)|r1(esk5_1(X1),esk6_1(X1))|~r1(esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22,plain,(p1(esk13_1(esk2_0))|p1(esk2_0)|p1(X1)|r1(X2,esk12_2(esk2_0,X2))|~r1(esk5_1(esk13_1(esk2_0)),X1)|~r1(esk2_0,X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_10])]), c_0_20])).
cnf(c_0_23,negated_conjecture,(p1(esk13_1(esk2_0))|p1(esk2_0)|r1(esk5_1(esk13_1(esk2_0)),esk6_1(esk13_1(esk2_0)))), inference(spm,[status(thm)],[c_0_21, c_0_16])).
cnf(c_0_24,negated_conjecture,(p1(esk6_1(esk13_1(esk2_0)))|p1(esk13_1(esk2_0))|p1(esk2_0)|r1(X1,esk12_2(esk2_0,X1))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,negated_conjecture,(p1(esk6_1(esk13_1(esk2_0)))|p1(esk13_1(esk2_0))|p1(esk2_0)|r1(esk4_0,esk12_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_24, c_0_9])).
cnf(c_0_26,plain,(p1(X1)|p1(X4)|~epred1_1(X1)|~r1(X1,X2)|~p1(esk12_2(X1,X2))|~p1(X3)|~r1(X3,X4)|~r1(esk13_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_27,negated_conjecture,(p1(esk6_1(esk13_1(esk2_0)))|p1(esk12_2(esk2_0,esk4_0))|p1(esk13_1(esk2_0))|p1(esk2_0)), inference(spm,[status(thm)],[c_0_11, c_0_25])).
cnf(c_0_28,plain,(p1(esk6_1(esk13_1(esk2_0)))|p1(esk13_1(esk2_0))|p1(esk2_0)|p1(X1)|~p1(X2)|~r1(esk13_1(esk2_0),X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_10]), c_0_9])])).
cnf(c_0_29,negated_conjecture,(p1(esk6_1(esk13_1(esk2_0)))|p1(esk13_1(esk2_0))|p1(esk2_0)|p1(X1)|~r1(esk5_1(esk13_1(esk2_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_19]), c_0_20])).
cnf(c_0_30,negated_conjecture,(p1(X1)|~r1(esk2_0,X1)|~p1(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_31,negated_conjecture,(p1(esk6_1(esk13_1(esk2_0)))|p1(esk13_1(esk2_0))|p1(esk2_0)), inference(spm,[status(thm)],[c_0_29, c_0_23])).
cnf(c_0_32,plain,(p1(X1)|r1(X2,esk12_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)|~p1(esk13_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_33,negated_conjecture,(p1(esk13_1(esk2_0))|p1(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_16])).
cnf(c_0_34,plain,(p1(esk2_0)|r1(X1,esk12_2(esk2_0,X1))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_10])])).
cnf(c_0_35,negated_conjecture,(p1(esk2_0)|r1(esk4_0,esk12_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_34, c_0_9])).
cnf(c_0_36,plain,(p1(X2)|r1(X3,esk7_2(X1,X3))|r1(esk9_1(X1),esk10_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_37,plain,(p1(X1)|~epred1_1(X1)|~r1(X1,X2)|~p1(esk12_2(X1,X2))|~p1(esk13_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_38,negated_conjecture,(p1(esk12_2(esk2_0,esk4_0))|p1(esk2_0)), inference(spm,[status(thm)],[c_0_11, c_0_35])).
cnf(c_0_39,negated_conjecture,(p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))|~p1(esk2_0)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_9]), c_0_10])])).
cnf(c_0_40,plain,(p1(esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_10]), c_0_9])]), c_0_33])).
cnf(c_0_41,plain,(p1(X2)|r1(esk7_2(X1,X3),esk8_2(X1,X3))|r1(esk9_1(X1),esk10_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_42,plain,(p1(X2)|r1(X3,esk7_2(X1,X3))|r1(X1,esk9_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_43,negated_conjecture,(p1(X1)|r1(esk9_1(esk2_0),esk10_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40])])).
cnf(c_0_44,negated_conjecture,(r1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_45,negated_conjecture,(~p1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_46,negated_conjecture,(p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))|~p1(esk2_0)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_9]), c_0_10])])).
cnf(c_0_47,negated_conjecture,(p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))|~p1(esk2_0)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_9]), c_0_10])])).
cnf(c_0_48,plain,(p1(X2)|r1(esk7_2(X1,X3),esk8_2(X1,X3))|r1(X1,esk9_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_49,plain,(p1(X2)|r1(X4,esk14_2(X1,X4))|~epred1_1(X1)|~r1(X3,X2)|~r1(X4,X3)|~r1(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_50,negated_conjecture,(r1(esk4_0,esk7_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
cnf(c_0_51,negated_conjecture,(p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_40])])).
cnf(c_0_52,negated_conjecture,(p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_40])])).
cnf(c_0_53,negated_conjecture,(p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))|~p1(esk2_0)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_9]), c_0_10])])).
cnf(c_0_54,plain,(p1(X1)|r1(esk9_1(esk2_0),esk10_1(esk2_0))|r1(esk4_0,esk14_2(X2,esk4_0))|~epred1_1(X2)|~r1(esk7_2(esk2_0,esk4_0),X1)|~r1(X2,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55,negated_conjecture,(r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_44]), c_0_45])).
cnf(c_0_56,negated_conjecture,(r1(esk4_0,esk7_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_44]), c_0_45])).
cnf(c_0_57,negated_conjecture,(p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_40])])).
cnf(c_0_58,plain,(p1(esk8_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))|r1(esk4_0,esk14_2(X1,esk4_0))|~epred1_1(X1)|~r1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_59,plain,(p1(X1)|r1(esk4_0,esk14_2(X2,esk4_0))|r1(esk2_0,esk9_1(esk2_0))|~epred1_1(X2)|~r1(esk7_2(esk2_0,esk4_0),X1)|~r1(X2,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_56])).
cnf(c_0_60,negated_conjecture,(r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_44]), c_0_45])).
cnf(c_0_61,plain,(p1(X2)|r1(esk9_1(X1),esk10_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk8_2(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_62,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_9]), c_0_10])])).
cnf(c_0_63,plain,(p1(esk8_2(esk2_0,esk4_0))|r1(esk4_0,esk14_2(X1,esk4_0))|r1(esk2_0,esk9_1(esk2_0))|~epred1_1(X1)|~r1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_64,plain,(p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_65,plain,(p1(X2)|r1(X1,esk9_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk8_2(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_66,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_9]), c_0_10])])).
cnf(c_0_67,negated_conjecture,(r1(esk9_1(esk2_0),esk10_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_44]), c_0_45])).
cnf(c_0_68,plain,(p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_69,plain,(p1(X2)|~epred1_1(X1)|~r1(X3,X2)|~r1(X4,X3)|~r1(X1,X4)|~p1(esk14_2(X1,X4))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_70,negated_conjecture,(p1(esk14_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_67])).
cnf(c_0_71,negated_conjecture,(r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_44]), c_0_45])).
cnf(c_0_72,plain,(p1(X1)|r1(esk9_1(esk2_0),esk10_1(esk2_0))|~r1(esk4_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_10]), c_0_9])])).
cnf(c_0_73,negated_conjecture,(p1(esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_71])).
cnf(c_0_74,negated_conjecture,(p1(X1)|r1(esk9_1(esk2_0),esk10_1(esk2_0))|~r1(esk7_2(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_72, c_0_50])).
cnf(c_0_75,plain,(p1(X1)|r1(esk2_0,esk9_1(esk2_0))|~r1(esk4_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_73]), c_0_10]), c_0_9])])).
cnf(c_0_76,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|r1(esk9_1(esk2_0),esk10_1(esk2_0))), inference(spm,[status(thm)],[c_0_74, c_0_55])).
cnf(c_0_77,negated_conjecture,(p1(X1)|r1(esk2_0,esk9_1(esk2_0))|~r1(esk7_2(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_75, c_0_56])).
cnf(c_0_78,plain,(p1(X1)|r1(esk9_1(esk2_0),esk10_1(esk2_0))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_76]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_79,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|r1(esk2_0,esk9_1(esk2_0))), inference(spm,[status(thm)],[c_0_77, c_0_60])).
cnf(c_0_80,plain,(p1(X3)|r1(X3,esk15_3(X1,X2,X3))|r1(X1,esk17_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_81,negated_conjecture,(r1(esk9_1(esk2_0),esk10_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_44]), c_0_45])).
cnf(c_0_82,plain,(p1(X1)|r1(esk2_0,esk9_1(esk2_0))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_79]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_83,plain,(p1(X3)|p1(esk15_3(X1,X2,X3))|r1(X1,esk17_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_84,plain,(p1(X2)|r1(X3,esk7_2(X1,X3))|p1(X5)|p1(X6)|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(X4)|~r1(X4,X5)|~r1(X6,X4)|~r1(esk9_1(X1),X6)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_85,plain,(p1(esk10_1(esk2_0))|r1(esk10_1(esk2_0),esk15_3(X1,esk9_1(esk2_0),esk10_1(esk2_0)))|r1(X1,esk17_1(X1))|~epred1_1(X1)|~r1(X1,esk9_1(esk2_0))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_86,negated_conjecture,(r1(esk2_0,esk9_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_44]), c_0_45])).
cnf(c_0_87,plain,(p1(esk15_3(X1,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|r1(X1,esk17_1(X1))|~epred1_1(X1)|~r1(X1,esk9_1(esk2_0))), inference(spm,[status(thm)],[c_0_83, c_0_81])).
cnf(c_0_88,plain,(p1(X3)|r1(esk15_3(X1,X2,X3),esk16_3(X1,X2,X3))|r1(X1,esk17_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_89,plain,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(X3,esk7_2(esk2_0,X3))|~p1(X4)|~r1(esk10_1(esk2_0),X4)|~r1(esk2_0,X3)|~r1(esk2_0,X2)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_81]), c_0_10]), c_0_40])])).
cnf(c_0_90,negated_conjecture,(p1(esk10_1(esk2_0))|r1(esk10_1(esk2_0),esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|r1(esk2_0,esk17_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_10])])).
cnf(c_0_91,negated_conjecture,(p1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|r1(esk2_0,esk17_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_86]), c_0_10])])).
cnf(c_0_92,plain,(p1(esk10_1(esk2_0))|r1(esk15_3(X1,esk9_1(esk2_0),esk10_1(esk2_0)),esk16_3(X1,esk9_1(esk2_0),esk10_1(esk2_0)))|r1(X1,esk17_1(X1))|~epred1_1(X1)|~r1(X1,esk9_1(esk2_0))), inference(spm,[status(thm)],[c_0_88, c_0_81])).
cnf(c_0_93,plain,(p1(X2)|r1(esk7_2(X1,X3),esk8_2(X1,X3))|p1(X5)|p1(X6)|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(X4)|~r1(X4,X5)|~r1(X6,X4)|~r1(esk9_1(X1),X6)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_94,plain,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk2_0,esk17_1(esk2_0))|r1(X3,esk7_2(esk2_0,X3))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X2)|~r1(esk2_0,X3)|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_91])).
cnf(c_0_95,negated_conjecture,(p1(esk10_1(esk2_0))|r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|r1(esk2_0,esk17_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_86]), c_0_10])])).
cnf(c_0_96,plain,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk7_2(esk2_0,X3),esk8_2(esk2_0,X3))|~p1(X4)|~r1(esk10_1(esk2_0),X4)|~r1(esk2_0,X3)|~r1(esk2_0,X2)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_81]), c_0_10]), c_0_40])])).
cnf(c_0_97,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(X2,esk7_2(esk2_0,X2))|r1(esk2_0,esk17_1(esk2_0))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_98,negated_conjecture,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk7_2(esk2_0,X3),esk8_2(esk2_0,X3))|r1(esk2_0,esk17_1(esk2_0))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X2)|~r1(esk2_0,X3)|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_90]), c_0_91])).
cnf(c_0_99,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_97, c_0_9])).
cnf(c_0_100,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,X2),esk8_2(esk2_0,X2))|r1(esk2_0,esk17_1(esk2_0))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_98, c_0_95])).
cnf(c_0_101,plain,(p1(X3)|r1(X1,esk17_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk16_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_102,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_44]), c_0_45])).
cnf(c_0_103,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_100, c_0_9])).
cnf(c_0_104,plain,(p1(esk10_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_10]), c_0_81]), c_0_86])])).
cnf(c_0_105,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_44]), c_0_45])).
cnf(c_0_106,plain,(p1(esk10_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(X2,esk4_0))|r1(esk2_0,esk17_1(esk2_0))|~epred1_1(X2)|~r1(esk7_2(esk2_0,esk4_0),X1)|~r1(X2,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_104])).
cnf(c_0_107,plain,(p1(esk10_1(esk2_0))|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_105]), c_0_10]), c_0_81]), c_0_86])])).
cnf(c_0_108,plain,(p1(esk8_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(X1,esk4_0))|r1(esk2_0,esk17_1(esk2_0))|~epred1_1(X1)|~r1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_109,plain,(p1(X2)|p1(X5)|p1(X6)|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk8_2(X1,X3))|~p1(X4)|~r1(X4,X5)|~r1(X6,X4)|~r1(esk9_1(X1),X6)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_110,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_9]), c_0_10])])).
cnf(c_0_111,plain,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|p1(X3)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))|~p1(X4)|~r1(esk9_1(esk2_0),X1)|~r1(esk2_0,X3)|~r1(X1,X4)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_112,negated_conjecture,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))|~p1(X3)|~r1(esk10_1(esk2_0),X3)|~r1(esk2_0,X1)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_111, c_0_81])).
cnf(c_0_113,negated_conjecture,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X1)|~r1(esk2_0,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_90]), c_0_91])).
cnf(c_0_114,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_113, c_0_95])).
cnf(c_0_115,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_44]), c_0_45])).
cnf(c_0_116,plain,(p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))|r1(esk2_0,esk17_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_115]), c_0_10]), c_0_81]), c_0_86])])).
cnf(c_0_117,negated_conjecture,(p1(esk14_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))|r1(esk2_0,esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_116])).
cnf(c_0_118,plain,(p1(esk10_1(esk2_0))|p1(X1)|r1(esk2_0,esk17_1(esk2_0))|~r1(esk4_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_117]), c_0_10]), c_0_9])])).
cnf(c_0_119,plain,(p1(esk10_1(esk2_0))|p1(X1)|r1(esk2_0,esk17_1(esk2_0))|~r1(esk7_2(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_118, c_0_104])).
cnf(c_0_120,plain,(p1(esk8_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))|r1(esk2_0,esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_119, c_0_107])).
cnf(c_0_121,plain,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|p1(X3)|r1(esk2_0,esk17_1(esk2_0))|~p1(X4)|~r1(esk9_1(esk2_0),X1)|~r1(esk2_0,X3)|~r1(X1,X4)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_120]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_122,negated_conjecture,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk2_0,esk17_1(esk2_0))|~p1(X3)|~r1(esk10_1(esk2_0),X3)|~r1(esk2_0,X1)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_121, c_0_81])).
cnf(c_0_123,negated_conjecture,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk2_0,esk17_1(esk2_0))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X1)|~r1(esk2_0,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_90]), c_0_91])).
cnf(c_0_124,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk2_0,esk17_1(esk2_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_123, c_0_95])).
cnf(c_0_125,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|r1(esk2_0,esk17_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_44]), c_0_45])).
cnf(c_0_126,plain,(p1(esk10_1(esk2_0))|r1(esk2_0,esk17_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_125]), c_0_10]), c_0_81]), c_0_86])])).
cnf(c_0_127,plain,(p1(X2)|r1(X3,esk7_2(X1,X3))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk10_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_128,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_17, c_0_126])).
cnf(c_0_129,plain,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(X2,esk7_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_10]), c_0_40])])).
cnf(c_0_130,plain,(p1(X2)|r1(esk7_2(X1,X3),esk8_2(X1,X3))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk10_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_131,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_129, c_0_9])).
cnf(c_0_132,plain,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,X2),esk8_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_128]), c_0_10]), c_0_40])])).
cnf(c_0_133,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_44]), c_0_45])).
cnf(c_0_134,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_132, c_0_9])).
cnf(c_0_135,plain,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(X2,esk4_0))|~epred1_1(X2)|~r1(esk7_2(esk2_0,esk4_0),X1)|~r1(X2,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_133])).
cnf(c_0_136,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_44]), c_0_45])).
cnf(c_0_137,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(X1,esk4_0))|~epred1_1(X1)|~r1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_138,plain,(p1(X2)|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk8_2(X1,X3))|~p1(esk10_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_139,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk8_2(esk2_0,esk4_0))|p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_9]), c_0_10])])).
cnf(c_0_140,plain,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_10]), c_0_40]), c_0_9])]), c_0_128])).
cnf(c_0_141,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_44]), c_0_45])).
cnf(c_0_142,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk14_2(esk2_0,esk4_0))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_141])).
cnf(c_0_143,plain,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|~r1(esk4_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_142]), c_0_10]), c_0_9])])).
cnf(c_0_144,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|~r1(esk7_2(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_143, c_0_133])).
cnf(c_0_145,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_144, c_0_136])).
cnf(c_0_146,plain,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_145]), c_0_10]), c_0_40]), c_0_9])]), c_0_128])).
cnf(c_0_147,plain,(p1(X3)|r1(X3,esk15_3(X1,X2,X3))|p1(X5)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(X4)|~r1(X4,X5)|~r1(esk17_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_148,negated_conjecture,(p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|r1(esk17_1(esk2_0),esk5_1(esk17_1(esk2_0)))), inference(spm,[status(thm)],[c_0_15, c_0_126])).
cnf(c_0_149,negated_conjecture,(p1(esk5_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_44]), c_0_45])).
cnf(c_0_150,plain,(p1(X3)|p1(esk15_3(X1,X2,X3))|p1(X5)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(X4)|~r1(X4,X5)|~r1(esk17_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_151,plain,(p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(X2,esk15_3(esk2_0,X3,X2))|~r1(esk5_1(esk17_1(esk2_0)),X1)|~r1(esk2_0,X3)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_10])]), c_0_149])).
cnf(c_0_152,negated_conjecture,(p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|r1(esk5_1(esk17_1(esk2_0)),esk6_1(esk17_1(esk2_0)))), inference(spm,[status(thm)],[c_0_21, c_0_126])).
cnf(c_0_153,plain,(p1(esk15_3(esk2_0,X1,X2))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X3)|p1(X2)|~r1(esk5_1(esk17_1(esk2_0)),X3)|~r1(esk2_0,X1)|~r1(X1,X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_148]), c_0_10])]), c_0_149])).
cnf(c_0_154,plain,(p1(X3)|r1(esk15_3(X1,X2,X3),esk16_3(X1,X2,X3))|p1(X5)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(X4)|~r1(X4,X5)|~r1(esk17_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_155,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|r1(X1,esk15_3(esk2_0,X2,X1))|~r1(esk2_0,X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_156,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk15_3(esk2_0,X1,X2))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X2)|~r1(esk2_0,X1)|~r1(X1,X2)), inference(spm,[status(thm)],[c_0_153, c_0_152])).
cnf(c_0_157,plain,(p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk15_3(esk2_0,X3,X2),esk16_3(esk2_0,X3,X2))|~r1(esk5_1(esk17_1(esk2_0)),X1)|~r1(esk2_0,X3)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_148]), c_0_10])]), c_0_149])).
cnf(c_0_158,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|r1(X1,esk15_3(esk2_0,esk9_1(esk2_0),X1))|~r1(esk9_1(esk2_0),X1)), inference(spm,[status(thm)],[c_0_155, c_0_86])).
cnf(c_0_159,negated_conjecture,(p1(esk15_3(esk2_0,esk9_1(esk2_0),X1))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|~r1(esk9_1(esk2_0),X1)), inference(spm,[status(thm)],[c_0_156, c_0_86])).
cnf(c_0_160,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk15_3(esk2_0,X2,X1),esk16_3(esk2_0,X2,X1))|~r1(esk2_0,X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_157, c_0_152])).
cnf(c_0_161,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|r1(esk10_1(esk2_0),esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))), inference(spm,[status(thm)],[c_0_158, c_0_81])).
cnf(c_0_162,negated_conjecture,(p1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_159, c_0_81])).
cnf(c_0_163,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk15_3(esk2_0,esk9_1(esk2_0),X1),esk16_3(esk2_0,esk9_1(esk2_0),X1))|~r1(esk9_1(esk2_0),X1)), inference(spm,[status(thm)],[c_0_160, c_0_86])).
cnf(c_0_164,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(X3,esk7_2(esk2_0,X3))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X2)|~r1(esk2_0,X3)|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_161]), c_0_162])).
cnf(c_0_165,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))), inference(spm,[status(thm)],[c_0_163, c_0_81])).
cnf(c_0_166,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|r1(X2,esk7_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_167,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk7_2(esk2_0,X3),esk8_2(esk2_0,X3))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X2)|~r1(esk2_0,X3)|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_161]), c_0_162])).
cnf(c_0_168,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_166, c_0_9])).
cnf(c_0_169,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,X2),esk8_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_167, c_0_165])).
cnf(c_0_170,plain,(p1(X3)|p1(X5)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk16_3(X1,X2,X3))|~p1(X4)|~r1(X4,X5)|~r1(esk17_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_171,plain,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_168, c_0_126])).
cnf(c_0_172,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_169, c_0_9])).
cnf(c_0_173,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|~p1(X2)|~r1(esk17_1(esk2_0),X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171]), c_0_10]), c_0_81]), c_0_86])])).
cnf(c_0_174,plain,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_172, c_0_126])).
cnf(c_0_175,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|~r1(esk5_1(esk17_1(esk2_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_148]), c_0_149])).
cnf(c_0_176,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|~p1(X2)|~r1(esk17_1(esk2_0),X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_174]), c_0_10]), c_0_81]), c_0_86])])).
cnf(c_0_177,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_175, c_0_152])).
cnf(c_0_178,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|~r1(esk5_1(esk17_1(esk2_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_148]), c_0_149])).
cnf(c_0_179,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(X2,esk4_0))|~epred1_1(X2)|~r1(esk7_2(esk2_0,esk4_0),X1)|~r1(X2,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_177])).
cnf(c_0_180,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_178, c_0_152])).
cnf(c_0_181,plain,(p1(esk8_2(esk2_0,esk4_0))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(X1,esk4_0))|~epred1_1(X1)|~r1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_179, c_0_180])).
cnf(c_0_182,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk8_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_9]), c_0_10])])).
cnf(c_0_183,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|p1(X3)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~p1(X4)|~r1(esk9_1(esk2_0),X1)|~r1(esk2_0,X3)|~r1(X1,X4)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_182]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_184,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|p1(X2)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~p1(X3)|~r1(esk10_1(esk2_0),X3)|~r1(esk2_0,X1)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_183, c_0_81])).
cnf(c_0_185,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X1)|~r1(esk2_0,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_161]), c_0_162])).
cnf(c_0_186,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_185, c_0_165])).
cnf(c_0_187,plain,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_186, c_0_126])).
cnf(c_0_188,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~p1(X2)|~r1(esk17_1(esk2_0),X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_187]), c_0_10]), c_0_81]), c_0_86])])).
cnf(c_0_189,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~r1(esk5_1(esk17_1(esk2_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_148]), c_0_149])).
cnf(c_0_190,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_189, c_0_152])).
cnf(c_0_191,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk14_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_190])).
cnf(c_0_192,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|~r1(esk4_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_191]), c_0_10]), c_0_9])])).
cnf(c_0_193,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|~r1(esk7_2(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_192, c_0_177])).
cnf(c_0_194,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))), inference(spm,[status(thm)],[c_0_193, c_0_180])).
cnf(c_0_195,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|p1(X2)|p1(X3)|~p1(X4)|~r1(esk9_1(esk2_0),X1)|~r1(esk2_0,X3)|~r1(X1,X4)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_194]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_196,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|~p1(X3)|~r1(esk10_1(esk2_0),X3)|~r1(esk2_0,X1)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_195, c_0_81])).
cnf(c_0_197,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|p1(X2)|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X1)|~r1(esk2_0,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_161]), c_0_162])).
cnf(c_0_198,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_197, c_0_165])).
cnf(c_0_199,plain,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_198, c_0_126])).
cnf(c_0_200,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))|p1(X1)|~p1(X2)|~r1(esk17_1(esk2_0),X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_199]), c_0_10]), c_0_81]), c_0_86])])).
cnf(c_0_201,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))|p1(X1)|~r1(esk5_1(esk17_1(esk2_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_148]), c_0_149])).
cnf(c_0_202,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(esk10_1(esk2_0))), inference(spm,[status(thm)],[c_0_201, c_0_152])).
cnf(c_0_203,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(X2,esk7_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_202]), c_0_10]), c_0_40])])).
cnf(c_0_204,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_203, c_0_9])).
cnf(c_0_205,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,X2),esk8_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_202]), c_0_10]), c_0_40])])).
cnf(c_0_206,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_44]), c_0_45])).
cnf(c_0_207,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_205, c_0_9])).
cnf(c_0_208,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(X2,esk4_0))|~epred1_1(X2)|~r1(esk7_2(esk2_0,esk4_0),X1)|~r1(X2,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_206])).
cnf(c_0_209,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_44]), c_0_45])).
cnf(c_0_210,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(X1,esk4_0))|~epred1_1(X1)|~r1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_208, c_0_209])).
cnf(c_0_211,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk8_2(esk2_0,esk4_0))|p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_9]), c_0_10])])).
cnf(c_0_212,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_211]), c_0_10]), c_0_40]), c_0_9])]), c_0_202])).
cnf(c_0_213,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_44]), c_0_45])).
cnf(c_0_214,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk14_2(esk2_0,esk4_0))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_213])).
cnf(c_0_215,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|~r1(esk4_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_214]), c_0_10]), c_0_9])])).
cnf(c_0_216,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|~r1(esk7_2(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_215, c_0_206])).
cnf(c_0_217,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_216, c_0_209])).
cnf(c_0_218,plain,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))|p1(X1)|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_217]), c_0_10]), c_0_40]), c_0_9])]), c_0_202])).
cnf(c_0_219,negated_conjecture,(p1(esk6_1(esk17_1(esk2_0)))|p1(esk17_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_44]), c_0_45])).
cnf(c_0_220,negated_conjecture,(p1(esk17_1(esk2_0))|~r1(esk2_0,esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_30, c_0_219])).
cnf(c_0_221,plain,(p1(esk10_1(esk2_0))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_220, c_0_126])).
cnf(c_0_222,plain,(p1(esk17_1(esk2_0))|p1(X1)|r1(X2,esk7_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_221]), c_0_10]), c_0_40])])).
cnf(c_0_223,negated_conjecture,(p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_222, c_0_9])).
cnf(c_0_224,plain,(p1(esk17_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,X2),esk8_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_221]), c_0_10]), c_0_40])])).
cnf(c_0_225,negated_conjecture,(p1(esk17_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_44]), c_0_45])).
cnf(c_0_226,negated_conjecture,(p1(esk17_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_224, c_0_9])).
cnf(c_0_227,plain,(p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(X2,esk4_0))|~epred1_1(X2)|~r1(esk7_2(esk2_0,esk4_0),X1)|~r1(X2,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_225])).
cnf(c_0_228,negated_conjecture,(p1(esk17_1(esk2_0))|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_44]), c_0_45])).
cnf(c_0_229,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(X1,esk4_0))|~epred1_1(X1)|~r1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_227, c_0_228])).
cnf(c_0_230,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_9]), c_0_10])])).
cnf(c_0_231,plain,(p1(esk17_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_230]), c_0_10]), c_0_40]), c_0_9])]), c_0_221])).
cnf(c_0_232,negated_conjecture,(p1(esk17_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_44]), c_0_45])).
cnf(c_0_233,negated_conjecture,(p1(esk14_2(esk2_0,esk4_0))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_232])).
cnf(c_0_234,plain,(p1(esk17_1(esk2_0))|p1(X1)|~r1(esk4_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_233]), c_0_10]), c_0_9])])).
cnf(c_0_235,negated_conjecture,(p1(esk17_1(esk2_0))|p1(X1)|~r1(esk7_2(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_234, c_0_225])).
cnf(c_0_236,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk17_1(esk2_0))), inference(spm,[status(thm)],[c_0_235, c_0_228])).
cnf(c_0_237,plain,(p1(esk17_1(esk2_0))|p1(X1)|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_236]), c_0_10]), c_0_40]), c_0_9])]), c_0_221])).
cnf(c_0_238,plain,(p1(X3)|r1(X3,esk15_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk17_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_239,negated_conjecture,(p1(esk17_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_44]), c_0_45])).
cnf(c_0_240,plain,(p1(X3)|p1(esk15_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk17_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_241,plain,(p1(X1)|r1(X1,esk15_3(esk2_0,X2,X1))|~r1(esk2_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239]), c_0_10])])).
cnf(c_0_242,plain,(p1(esk15_3(esk2_0,X1,X2))|p1(X2)|~r1(esk2_0,X1)|~r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_239]), c_0_10])])).
cnf(c_0_243,plain,(p1(X3)|r1(esk15_3(X1,X2,X3),esk16_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk17_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_244,negated_conjecture,(p1(X1)|r1(X1,esk15_3(esk2_0,esk9_1(esk2_0),X1))|~r1(esk9_1(esk2_0),X1)), inference(spm,[status(thm)],[c_0_241, c_0_86])).
cnf(c_0_245,negated_conjecture,(p1(esk15_3(esk2_0,esk9_1(esk2_0),X1))|p1(X1)|~r1(esk9_1(esk2_0),X1)), inference(spm,[status(thm)],[c_0_242, c_0_86])).
cnf(c_0_246,plain,(p1(X1)|r1(esk15_3(esk2_0,X2,X1),esk16_3(esk2_0,X2,X1))|~r1(esk2_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_239]), c_0_10])])).
cnf(c_0_247,negated_conjecture,(p1(esk10_1(esk2_0))|r1(esk10_1(esk2_0),esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))), inference(spm,[status(thm)],[c_0_244, c_0_81])).
cnf(c_0_248,negated_conjecture,(p1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))), inference(spm,[status(thm)],[c_0_245, c_0_81])).
cnf(c_0_249,negated_conjecture,(p1(X1)|r1(esk15_3(esk2_0,esk9_1(esk2_0),X1),esk16_3(esk2_0,esk9_1(esk2_0),X1))|~r1(esk9_1(esk2_0),X1)), inference(spm,[status(thm)],[c_0_246, c_0_86])).
cnf(c_0_250,plain,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(X3,esk7_2(esk2_0,X3))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X2)|~r1(esk2_0,X3)|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_247]), c_0_248])).
cnf(c_0_251,negated_conjecture,(p1(esk10_1(esk2_0))|r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))), inference(spm,[status(thm)],[c_0_249, c_0_81])).
cnf(c_0_252,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(X2,esk7_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_250, c_0_251])).
cnf(c_0_253,plain,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk7_2(esk2_0,X3),esk8_2(esk2_0,X3))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X2)|~r1(esk2_0,X3)|~r1(esk2_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_247]), c_0_248])).
cnf(c_0_254,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_252, c_0_9])).
cnf(c_0_255,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,X2),esk8_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_253, c_0_251])).
cnf(c_0_256,plain,(p1(X3)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk16_3(X1,X2,X3))|~p1(esk17_1(X1))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_257,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_44]), c_0_45])).
cnf(c_0_258,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_255, c_0_9])).
cnf(c_0_259,plain,(p1(esk10_1(esk2_0))|r1(esk4_0,esk7_2(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_257]), c_0_10]), c_0_239]), c_0_81]), c_0_86])])).
cnf(c_0_260,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_44]), c_0_45])).
cnf(c_0_261,plain,(p1(esk10_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(X2,esk4_0))|~epred1_1(X2)|~r1(esk7_2(esk2_0,esk4_0),X1)|~r1(X2,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_259])).
cnf(c_0_262,plain,(p1(esk10_1(esk2_0))|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_260]), c_0_10]), c_0_239]), c_0_81]), c_0_86])])).
cnf(c_0_263,plain,(p1(esk8_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(X1,esk4_0))|~epred1_1(X1)|~r1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_261, c_0_262])).
cnf(c_0_264,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_9]), c_0_10])])).
cnf(c_0_265,plain,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|p1(X3)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~p1(X4)|~r1(esk9_1(esk2_0),X1)|~r1(esk2_0,X3)|~r1(X1,X4)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_264]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_266,negated_conjecture,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~p1(X3)|~r1(esk10_1(esk2_0),X3)|~r1(esk2_0,X1)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_265, c_0_81])).
cnf(c_0_267,negated_conjecture,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X1)|~r1(esk2_0,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_247]), c_0_248])).
cnf(c_0_268,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_267, c_0_251])).
cnf(c_0_269,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_44]), c_0_45])).
cnf(c_0_270,plain,(p1(esk10_1(esk2_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_269]), c_0_10]), c_0_239]), c_0_81]), c_0_86])])).
cnf(c_0_271,negated_conjecture,(p1(esk14_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_270])).
cnf(c_0_272,plain,(p1(esk10_1(esk2_0))|p1(X1)|~r1(esk4_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_271]), c_0_10]), c_0_9])])).
cnf(c_0_273,plain,(p1(esk10_1(esk2_0))|p1(X1)|~r1(esk7_2(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_272, c_0_259])).
cnf(c_0_274,plain,(p1(esk8_2(esk2_0,esk4_0))|p1(esk10_1(esk2_0))), inference(spm,[status(thm)],[c_0_273, c_0_262])).
cnf(c_0_275,plain,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|p1(X3)|~p1(X4)|~r1(esk9_1(esk2_0),X1)|~r1(esk2_0,X3)|~r1(X1,X4)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_274]), c_0_10]), c_0_40]), c_0_9])])).
cnf(c_0_276,negated_conjecture,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|~p1(X3)|~r1(esk10_1(esk2_0),X3)|~r1(esk2_0,X1)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_275, c_0_81])).
cnf(c_0_277,negated_conjecture,(p1(esk10_1(esk2_0))|p1(X1)|p1(X2)|~r1(esk15_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)),X1)|~r1(esk2_0,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_247]), c_0_248])).
cnf(c_0_278,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))|p1(X1)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_277, c_0_251])).
cnf(c_0_279,negated_conjecture,(p1(esk16_3(esk2_0,esk9_1(esk2_0),esk10_1(esk2_0)))|p1(esk10_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_44]), c_0_45])).
cnf(c_0_280,plain,(p1(esk10_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_279]), c_0_10]), c_0_239]), c_0_81]), c_0_86])])).
cnf(c_0_281,plain,(p1(X1)|r1(X2,esk7_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_280]), c_0_10]), c_0_40])])).
cnf(c_0_282,negated_conjecture,(p1(X1)|r1(esk4_0,esk7_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_281, c_0_9])).
cnf(c_0_283,plain,(p1(X1)|r1(esk7_2(esk2_0,X2),esk8_2(esk2_0,X2))|~r1(esk2_0,X2)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_280]), c_0_10]), c_0_40])])).
cnf(c_0_284,negated_conjecture,(r1(esk4_0,esk7_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_44]), c_0_45])).
cnf(c_0_285,negated_conjecture,(p1(X1)|r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_283, c_0_9])).
cnf(c_0_286,plain,(p1(X1)|r1(esk4_0,esk14_2(X2,esk4_0))|~epred1_1(X2)|~r1(esk7_2(esk2_0,esk4_0),X1)|~r1(X2,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_284])).
cnf(c_0_287,negated_conjecture,(r1(esk7_2(esk2_0,esk4_0),esk8_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_44]), c_0_45])).
cnf(c_0_288,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|r1(esk4_0,esk14_2(X1,esk4_0))|~epred1_1(X1)|~r1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_286, c_0_287])).
cnf(c_0_289,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))|r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_9]), c_0_10])])).
cnf(c_0_290,plain,(p1(X1)|r1(esk4_0,esk14_2(esk2_0,esk4_0))|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_289]), c_0_10]), c_0_280]), c_0_40]), c_0_9])])).
cnf(c_0_291,negated_conjecture,(r1(esk4_0,esk14_2(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_44]), c_0_45])).
cnf(c_0_292,negated_conjecture,(p1(esk14_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_11, c_0_291])).
cnf(c_0_293,plain,(p1(X1)|~r1(esk4_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_292]), c_0_10]), c_0_9])])).
cnf(c_0_294,negated_conjecture,(p1(X1)|~r1(esk7_2(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_293, c_0_284])).
cnf(c_0_295,negated_conjecture,(p1(esk8_2(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_294, c_0_287])).
cnf(c_0_296,plain,(p1(X1)|~r1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_295]), c_0_10]), c_0_280]), c_0_40]), c_0_9])])).
cnf(c_0_297,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_44]), c_0_45]), ['proof']).
# SZS output end CNFRefutation.
EOF
