% Problem    : SWB014+3 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n183.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:05:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:![X4]:![X5]:![X6]:![X7]:((((iext(uri_rdf_first,X4,X5)&iext(uri_rdf_rest,X4,X6))&iext(uri_rdf_first,X6,X7))&iext(uri_rdf_rest,X6,uri_rdf_nil))=>(iext(uri_owl_unionOf,X1,X4)<=>(((ic(X1)&ic(X5))&ic(X7))&![X3]:(icext(X1,X3)<=>(icext(X5,X3)|icext(X7,X3))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB002+0.ax', owl_bool_unionof_class_002)).
fof(c_0_1, axiom, (?[X22]:?[X23]:?[X24]:(((((((iext(uri_rdf_type,uri_ex_Eagle,uri_ex_Species)&iext(uri_rdf_type,uri_ex_Falcon,uri_ex_Species))&iext(uri_rdf_type,uri_ex_harry,X22))&iext(uri_owl_unionOf,X22,X23))&iext(uri_rdf_first,X23,uri_ex_Eagle))&iext(uri_rdf_rest,X23,X24))&iext(uri_rdf_first,X24,uri_ex_Falcon))&iext(uri_rdf_rest,X24,uri_rdf_nil))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_premise_fullish_014_Harry_belongs_to_some_Species)).
fof(c_0_2, axiom, (![X3]:![X2]:(iext(uri_rdf_type,X3,X2)<=>icext(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWB002+0.ax', rdfs_cext_def)).
fof(c_0_3, conjecture, (?[X21]:(iext(uri_rdf_type,uri_ex_harry,X21)&iext(uri_rdf_type,X21,uri_ex_Species))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_conclusion_fullish_014_Harry_belongs_to_some_Species)).
fof(c_0_4, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((ic(X8)|~iext(uri_owl_unionOf,X8,X9))|(((~iext(uri_rdf_first,X9,X10)|~iext(uri_rdf_rest,X9,X11))|~iext(uri_rdf_first,X11,X12))|~iext(uri_rdf_rest,X11,uri_rdf_nil)))&((ic(X10)|~iext(uri_owl_unionOf,X8,X9))|(((~iext(uri_rdf_first,X9,X10)|~iext(uri_rdf_rest,X9,X11))|~iext(uri_rdf_first,X11,X12))|~iext(uri_rdf_rest,X11,uri_rdf_nil))))&((ic(X12)|~iext(uri_owl_unionOf,X8,X9))|(((~iext(uri_rdf_first,X9,X10)|~iext(uri_rdf_rest,X9,X11))|~iext(uri_rdf_first,X11,X12))|~iext(uri_rdf_rest,X11,uri_rdf_nil))))&((((~icext(X8,X13)|(icext(X10,X13)|icext(X12,X13)))|~iext(uri_owl_unionOf,X8,X9))|(((~iext(uri_rdf_first,X9,X10)|~iext(uri_rdf_rest,X9,X11))|~iext(uri_rdf_first,X11,X12))|~iext(uri_rdf_rest,X11,uri_rdf_nil)))&((((~icext(X10,X14)|icext(X8,X14))|~iext(uri_owl_unionOf,X8,X9))|(((~iext(uri_rdf_first,X9,X10)|~iext(uri_rdf_rest,X9,X11))|~iext(uri_rdf_first,X11,X12))|~iext(uri_rdf_rest,X11,uri_rdf_nil)))&(((~icext(X12,X14)|icext(X8,X14))|~iext(uri_owl_unionOf,X8,X9))|(((~iext(uri_rdf_first,X9,X10)|~iext(uri_rdf_rest,X9,X11))|~iext(uri_rdf_first,X11,X12))|~iext(uri_rdf_rest,X11,uri_rdf_nil))))))&((((((~icext(X10,esk7_5(X8,X9,X10,X11,X12))|~icext(X8,esk7_5(X8,X9,X10,X11,X12)))|((~ic(X8)|~ic(X10))|~ic(X12)))|iext(uri_owl_unionOf,X8,X9))|(((~iext(uri_rdf_first,X9,X10)|~iext(uri_rdf_rest,X9,X11))|~iext(uri_rdf_first,X11,X12))|~iext(uri_rdf_rest,X11,uri_rdf_nil)))&((((~icext(X12,esk7_5(X8,X9,X10,X11,X12))|~icext(X8,esk7_5(X8,X9,X10,X11,X12)))|((~ic(X8)|~ic(X10))|~ic(X12)))|iext(uri_owl_unionOf,X8,X9))|(((~iext(uri_rdf_first,X9,X10)|~iext(uri_rdf_rest,X9,X11))|~iext(uri_rdf_first,X11,X12))|~iext(uri_rdf_rest,X11,uri_rdf_nil))))&((((icext(X8,esk7_5(X8,X9,X10,X11,X12))|(icext(X10,esk7_5(X8,X9,X10,X11,X12))|icext(X12,esk7_5(X8,X9,X10,X11,X12))))|((~ic(X8)|~ic(X10))|~ic(X12)))|iext(uri_owl_unionOf,X8,X9))|(((~iext(uri_rdf_first,X9,X10)|~iext(uri_rdf_rest,X9,X11))|~iext(uri_rdf_first,X11,X12))|~iext(uri_rdf_rest,X11,uri_rdf_nil)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])])).
fof(c_0_5, plain, ((((((((iext(uri_rdf_type,uri_ex_Eagle,uri_ex_Species)&iext(uri_rdf_type,uri_ex_Falcon,uri_ex_Species))&iext(uri_rdf_type,uri_ex_harry,esk19_0))&iext(uri_owl_unionOf,esk19_0,esk20_0))&iext(uri_rdf_first,esk20_0,uri_ex_Eagle))&iext(uri_rdf_rest,esk20_0,esk21_0))&iext(uri_rdf_first,esk21_0,uri_ex_Falcon))&iext(uri_rdf_rest,esk21_0,uri_rdf_nil))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_1])])).
cnf(c_0_6,plain,(icext(X2,X6)|icext(X4,X6)|~iext(uri_rdf_rest,X1,uri_rdf_nil)|~iext(uri_rdf_first,X1,X2)|~iext(uri_rdf_rest,X3,X1)|~iext(uri_rdf_first,X3,X4)|~iext(uri_owl_unionOf,X5,X3)|~icext(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(iext(uri_rdf_rest,esk21_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,plain,(icext(X1,X2)|icext(X3,X2)|~icext(X4,X2)|~iext(uri_rdf_first,esk21_0,X3)|~iext(uri_rdf_rest,X5,esk21_0)|~iext(uri_rdf_first,X5,X1)|~iext(uri_owl_unionOf,X4,X5)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_9,plain,(iext(uri_rdf_first,esk21_0,uri_ex_Falcon)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,plain,(icext(uri_ex_Falcon,X1)|icext(X2,X1)|~icext(X3,X1)|~iext(uri_rdf_rest,X4,esk21_0)|~iext(uri_rdf_first,X4,X2)|~iext(uri_owl_unionOf,X3,X4)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_11,plain,(iext(uri_rdf_rest,esk20_0,esk21_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12,plain,(icext(uri_ex_Falcon,X1)|icext(X2,X1)|~icext(X3,X1)|~iext(uri_rdf_first,esk20_0,X2)|~iext(uri_owl_unionOf,X3,esk20_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_13,plain,(iext(uri_rdf_first,esk20_0,uri_ex_Eagle)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:![X7]:((~iext(uri_rdf_type,X4,X5)|icext(X5,X4))&(~icext(X7,X6)|iext(uri_rdf_type,X6,X7)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
fof(c_0_15, negated_conjecture, (~(?[X21]:(iext(uri_rdf_type,uri_ex_harry,X21)&iext(uri_rdf_type,X21,uri_ex_Species)))), inference(assume_negation,[status(cth)],[c_0_3])).
cnf(c_0_16,plain,(icext(uri_ex_Eagle,X1)|icext(uri_ex_Falcon,X1)|~icext(X2,X1)|~iext(uri_owl_unionOf,X2,esk20_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17,plain,(iext(uri_owl_unionOf,esk19_0,esk20_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18,plain,(icext(X1,X2)|~iext(uri_rdf_type,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,plain,(iext(uri_rdf_type,uri_ex_harry,esk19_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_20, negated_conjecture, (![X22]:(~iext(uri_rdf_type,uri_ex_harry,X22)|~iext(uri_rdf_type,X22,uri_ex_Species))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
cnf(c_0_21,plain,(icext(uri_ex_Falcon,X1)|icext(uri_ex_Eagle,X1)|~icext(esk19_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22,plain,(icext(esk19_0,uri_ex_harry)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23,negated_conjecture,(~iext(uri_rdf_type,X1,uri_ex_Species)|~iext(uri_rdf_type,uri_ex_harry,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,plain,(iext(uri_rdf_type,uri_ex_Falcon,uri_ex_Species)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25,plain,(iext(uri_rdf_type,X1,X2)|~icext(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26,plain,(icext(uri_ex_Eagle,uri_ex_harry)|icext(uri_ex_Falcon,uri_ex_harry)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27,negated_conjecture,(~iext(uri_rdf_type,uri_ex_harry,uri_ex_Falcon)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28,plain,(iext(uri_rdf_type,uri_ex_Eagle,uri_ex_Species)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29,plain,(icext(uri_ex_Eagle,uri_ex_harry)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
cnf(c_0_30,negated_conjecture,(~iext(uri_rdf_type,uri_ex_harry,uri_ex_Eagle)), inference(spm,[status(thm)],[c_0_23, c_0_28])).
cnf(c_0_31,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_29]), c_0_30]), ['proof']).
# SZS output end CNFRefutation.
EOF
