% Problem    : CSR048+3 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n152.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:18:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X4]:(mtvisible(c_tptpgeo_spindlecollectormt)=>borderson(c_georegion_l4_x45_y9,X4))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', query148)).
fof(c_0_1, axiom, ((mtvisible(c_tptpgeo_member1_mt)=>borderson(c_georegion_l4_x45_y9,c_georegion_l4_x45_y10))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_648)).
fof(c_0_2, axiom, (![X21]:![X22]:((mtvisible(X21)&genlmt(X21,X22))=>mtvisible(X22))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_7997)).
fof(c_0_3, axiom, (genlmt(c_tptpgeo_spindlecollectormt,c_tptpgeo_member1_mt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_724)).
fof(c_0_4, negated_conjecture, (~(?[X4]:(mtvisible(c_tptpgeo_spindlecollectormt)=>borderson(c_georegion_l4_x45_y9,X4)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, plain, ((~mtvisible(c_tptpgeo_member1_mt)|borderson(c_georegion_l4_x45_y9,c_georegion_l4_x45_y10))), inference(fof_nnf,[status(thm)],[c_0_1])).
fof(c_0_6, negated_conjecture, (![X5]:(mtvisible(c_tptpgeo_spindlecollectormt)&~borderson(c_georegion_l4_x45_y9,X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
fof(c_0_7, plain, (![X23]:![X24]:((~mtvisible(X23)|~genlmt(X23,X24))|mtvisible(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
cnf(c_0_8,plain,(borderson(c_georegion_l4_x45_y9,c_georegion_l4_x45_y10)|~mtvisible(c_tptpgeo_member1_mt)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,negated_conjecture,(~borderson(c_georegion_l4_x45_y9,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(mtvisible(X1)|~genlmt(X2,X1)|~mtvisible(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,plain,(genlmt(c_tptpgeo_spindlecollectormt,c_tptpgeo_member1_mt)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_12,negated_conjecture,(mtvisible(c_tptpgeo_spindlecollectormt)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13,plain,(~mtvisible(c_tptpgeo_member1_mt)), inference(sr,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14,plain,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])]), c_0_13]), ['proof']).
# SZS output end CNFRefutation.
EOF
