% Problem    : CSR115+35 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n143.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:05:11 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:?[X8]:(((((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&has_card_leq(X8,int1994))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&sub(X7,jahr__1_1))&val(X2,bmw_0))&val(X3,bmw_0))&val(X7,X8))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_news_1222)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((pred(c3345,sonderschicht_1_1)&sub(c3346,sowohl_2_1))&attr(c3446,c3447))&sub(c3446,firma_1_1))&sub(c3447,name_1_1))&val(c3447,bmw_0))&attr(c3527,c3528))&sub(c3527,firma_1_1))&sub(c3528,name_1_1))&val(c3528,rover_0))&attr(c3531,c3532))&sub(c3532,jahr__1_1))&val(c3532,c3529))&prop(c3536,zwei_1_1))&subs(c3536,nachfrage__1_1))&tupl_p8(c3626,c3345,c3346,c3446,c3527,c3531,c3536,c3345))&assoc(sonderschicht_1_1,besonder_1_1))&sub(sonderschicht_1_1,schicht_1_1))&sort(c3345,d))&card(c3345,cons(x_constant,cons(int1,nil))))&etype(c3345,int1))&fact(c3345,real))&gener(c3345,gener_c))&quant(c3345,mult))&refer(c3345,indet))&varia(c3345,varia_c))&sort(sonderschicht_1_1,d))&card(sonderschicht_1_1,int1))&etype(sonderschicht_1_1,int0))&fact(sonderschicht_1_1,real))&gener(sonderschicht_1_1,ge))&quant(sonderschicht_1_1,one))&refer(sonderschicht_1_1,refer_c))&varia(sonderschicht_1_1,varia_c))&sort(c3346,o))&card(c3346,int1))&etype(c3346,int0))&fact(c3346,real))&gener(c3346,gener_c))&quant(c3346,one))&refer(c3346,refer_c))&varia(c3346,varia_c))&sort(sowohl_2_1,o))&card(sowohl_2_1,int1))&etype(sowohl_2_1,int0))&fact(sowohl_2_1,real))&gener(sowohl_2_1,ge))&quant(sowohl_2_1,one))&refer(sowohl_2_1,refer_c))&varia(sowohl_2_1,varia_c))&sort(c3446,d))&sort(c3446,io))&card(c3446,int1))&etype(c3446,int0))&fact(c3446,real))&gener(c3446,sp))&quant(c3446,one))&refer(c3446,det))&varia(c3446,con))&sort(c3447,na))&card(c3447,int1))&etype(c3447,int0))&fact(c3447,real))&gener(c3447,sp))&quant(c3447,one))&refer(c3447,indet))&varia(c3447,varia_c))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(bmw_0,fe))&sort(c3527,d))&sort(c3527,io))&card(c3527,int1))&etype(c3527,int0))&fact(c3527,real))&gener(c3527,sp))&quant(c3527,one))&refer(c3527,det))&varia(c3527,con))&sort(c3528,na))&card(c3528,int1))&etype(c3528,int0))&fact(c3528,real))&gener(c3528,sp))&quant(c3528,one))&refer(c3528,indet))&varia(c3528,varia_c))&sort(rover_0,fe))&sort(c3531,t))&card(c3531,int1))&etype(c3531,int0))&fact(c3531,real))&gener(c3531,sp))&quant(c3531,one))&refer(c3531,det))&varia(c3531,con))&sort(c3532,me))&sort(c3532,oa))&sort(c3532,ta))&card(c3532,card_c))&etype(c3532,etype_c))&fact(c3532,real))&gener(c3532,sp))&quant(c3532,quant_c))&refer(c3532,refer_c))&varia(c3532,varia_c))&sort(jahr__1_1,me))&sort(jahr__1_1,oa))&sort(jahr__1_1,ta))&card(jahr__1_1,card_c))&etype(jahr__1_1,etype_c))&fact(jahr__1_1,real))&gener(jahr__1_1,ge))&quant(jahr__1_1,quant_c))&refer(jahr__1_1,refer_c))&varia(jahr__1_1,varia_c))&sort(c3529,nu))&card(c3529,int1994))&sort(c3536,ad))&card(c3536,int1))&etype(c3536,int0))&fact(c3536,real))&gener(c3536,sp))&quant(c3536,one))&refer(c3536,det))&varia(c3536,con))&sort(zwei_1_1,nq))&sort(nachfrage__1_1,ad))&card(nachfrage__1_1,int1))&etype(nachfrage__1_1,int0))&fact(nachfrage__1_1,real))&gener(nachfrage__1_1,ge))&quant(nachfrage__1_1,one))&refer(nachfrage__1_1,refer_c))&varia(nachfrage__1_1,varia_c))&sort(c3626,ent))&card(c3626,card_c))&etype(c3626,etype_c))&fact(c3626,real))&gener(c3626,gener_c))&quant(c3626,quant_c))&refer(c3626,refer_c))&varia(c3626,varia_c))&sort(besonder_1_1,tq))&sort(schicht_1_1,d))&card(schicht_1_1,int1))&etype(schicht_1_1,int0))&fact(schicht_1_1,real))&gener(schicht_1_1,ge))&quant(schicht_1_1,one))&refer(schicht_1_1,refer_c))&varia(schicht_1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_news_1222)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_3, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_4, axiom, (![X1]:![X2]:(card(X1,X2)=>has_card_leq(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', has_card_eq)).
fof(c_0_5, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:?[X8]:(((((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&has_card_leq(X8,int1994))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&sub(X7,jahr__1_1))&val(X2,bmw_0))&val(X3,bmw_0))&val(X7,X8)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_6, negated_conjecture, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:(((((((((((~attr(X9,X10)|~attr(X12,X11))|~attr(X14,X15))|~has_card_leq(X16,int1994))|~obj(X13,X9))|~sub(X10,name_1_1))|~sub(X9,firma_1_1))|~sub(X11,name_1_1))|~sub(X15,jahr__1_1))|~val(X10,bmw_0))|~val(X11,bmw_0))|~val(X15,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
cnf(c_0_7,negated_conjecture,(~val(X1,X2)|~val(X3,bmw_0)|~val(X4,bmw_0)|~sub(X1,jahr__1_1)|~sub(X3,name_1_1)|~sub(X5,firma_1_1)|~sub(X4,name_1_1)|~obj(X6,X5)|~has_card_leq(X2,int1994)|~attr(X7,X1)|~attr(X8,X3)|~attr(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_8,hypothesis,(val(c3447,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(sub(c3447,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_10,hypothesis,(~val(X1,bmw_0)|~val(X2,X3)|~attr(X4,c3447)|~attr(X5,X1)|~attr(X6,X2)|~obj(X7,X4)|~sub(X4,firma_1_1)|~sub(X1,name_1_1)|~sub(X2,jahr__1_1)|~has_card_leq(X3,int1994)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9])])).
cnf(c_0_11,hypothesis,(~val(X1,X2)|~attr(X3,c3447)|~attr(X4,c3447)|~attr(X5,X1)|~obj(X6,X3)|~sub(X3,firma_1_1)|~sub(X1,jahr__1_1)|~has_card_leq(X2,int1994)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_8]), c_0_9])])).
cnf(c_0_12,hypothesis,(attr(c3446,c3447)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(sub(c3446,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_14,hypothesis,(~val(X1,X2)|~attr(X3,c3447)|~attr(X4,X1)|~obj(X5,c3446)|~sub(X1,jahr__1_1)|~has_card_leq(X2,int1994)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
fof(c_0_15, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk41_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk41_3(X7,X8,X9),esk42_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk40_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk40_3(X7,X8,X9),esk41_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk40_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk42_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk41_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk40_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_16,hypothesis,(~val(X1,X2)|~attr(X3,X1)|~obj(X4,c3446)|~sub(X1,jahr__1_1)|~has_card_leq(X2,int1994)), inference(spm,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_17,plain,(obj(esk40_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_18, plain, (![X4]:![X5]:(((arg1(esk43_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk43_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk43_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_19,hypothesis,(~val(X1,X2)|~subr(X3,sub_0)|~attr(X4,X1)|~arg2(X3,X5)|~arg1(X3,c3446)|~sub(X1,jahr__1_1)|~has_card_leq(X2,int1994)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,plain,(subr(esk43_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21,hypothesis,(~val(X1,X2)|~attr(X3,X1)|~arg2(esk43_2(X4,X5),X6)|~arg1(esk43_2(X4,X5),c3446)|~sub(X1,jahr__1_1)|~sub(X4,X5)|~has_card_leq(X2,int1994)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_22,plain,(arg2(esk43_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23,hypothesis,(~val(X1,X2)|~attr(X3,X1)|~arg1(esk43_2(X4,X5),c3446)|~sub(X1,jahr__1_1)|~sub(X4,X5)|~has_card_leq(X2,int1994)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_24,plain,(arg1(esk43_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25,hypothesis,(~val(X1,X2)|~attr(X3,X1)|~sub(X1,jahr__1_1)|~sub(c3446,X4)|~has_card_leq(X2,int1994)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_26, plain, (![X3]:![X4]:(~card(X3,X4)|has_card_leq(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
cnf(c_0_27,hypothesis,(~val(X1,X2)|~attr(X3,X1)|~sub(X1,jahr__1_1)|~has_card_leq(X2,int1994)), inference(spm,[status(thm)],[c_0_25, c_0_13])).
cnf(c_0_28,plain,(has_card_leq(X1,X2)|~card(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_29,hypothesis,(~val(X1,X2)|~attr(X3,X1)|~sub(X1,jahr__1_1)|~card(X2,int1994)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_30,hypothesis,(val(c3532,c3529)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_31,hypothesis,(sub(c3532,jahr__1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_32,hypothesis,(card(c3529,int1994)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_33,hypothesis,(attr(c3531,c3532)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_34,hypothesis,(~attr(X1,c3532)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_35,hypothesis,($false), inference(sr,[status(thm)],[c_0_33, c_0_34]), ['proof']).
# SZS output end CNFRefutation.
EOF
