% Problem    : CSR114+28 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n142.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:46:56 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:((((attr(X1,X2)&scar(X4,X3))&sub(X2,name_1_1))&sub(X1,stadt__1_1))&val(X2,rom_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_004_qapw_93)).
fof(c_0_1, hypothesis, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((assoc(au__337enfassade_1_1,au__337en_1_1)&sub(au__337enfassade_1_1,fassade_1_1))&sub(c573,architektur_1_1))&sub(c582,au__337enfassade_1_1))&sub(c609,sich_1_1))&sub(c614,kolosseum_1_1))&attr(c620,c621))&sub(c620,stadt__1_1))&sub(c621,name_1_1))&val(c621,rom_0))&tupl_p6(c895,c573,c582,c609,c614,c620))&sort(au__337enfassade_1_1,d))&card(au__337enfassade_1_1,int1))&etype(au__337enfassade_1_1,int0))&fact(au__337enfassade_1_1,real))&gener(au__337enfassade_1_1,ge))&quant(au__337enfassade_1_1,one))&refer(au__337enfassade_1_1,refer_c))&varia(au__337enfassade_1_1,varia_c))&sort(au__337en_1_1,l))&card(au__337en_1_1,int1))&etype(au__337en_1_1,int0))&fact(au__337en_1_1,real))&gener(au__337en_1_1,gener_c))&quant(au__337en_1_1,one))&refer(au__337en_1_1,refer_c))&varia(au__337en_1_1,varia_c))&sort(fassade_1_1,d))&card(fassade_1_1,int1))&etype(fassade_1_1,int0))&fact(fassade_1_1,real))&gener(fassade_1_1,ge))&quant(fassade_1_1,one))&refer(fassade_1_1,refer_c))&varia(fassade_1_1,varia_c))&sort(c573,io))&card(c573,int1))&etype(c573,int0))&fact(c573,real))&gener(c573,sp))&quant(c573,one))&refer(c573,det))&varia(c573,con))&sort(architektur_1_1,io))&card(architektur_1_1,int1))&etype(architektur_1_1,int0))&fact(architektur_1_1,real))&gener(architektur_1_1,ge))&quant(architektur_1_1,one))&refer(architektur_1_1,refer_c))&varia(architektur_1_1,varia_c))&sort(c582,d))&card(c582,int1))&etype(c582,int0))&fact(c582,real))&gener(c582,sp))&quant(c582,one))&refer(c582,det))&varia(c582,con))&sort(c609,o))&card(c609,int1))&etype(c609,int0))&fact(c609,real))&gener(c609,gener_c))&quant(c609,one))&refer(c609,refer_c))&varia(c609,varia_c))&sort(sich_1_1,o))&card(sich_1_1,int1))&etype(sich_1_1,int0))&fact(sich_1_1,real))&gener(sich_1_1,gener_c))&quant(sich_1_1,one))&refer(sich_1_1,refer_c))&varia(sich_1_1,varia_c))&sort(c614,d))&card(c614,int1))&etype(c614,int0))&fact(c614,real))&gener(c614,sp))&quant(c614,one))&refer(c614,det))&varia(c614,con))&sort(kolosseum_1_1,d))&card(kolosseum_1_1,int1))&etype(kolosseum_1_1,int0))&fact(kolosseum_1_1,real))&gener(kolosseum_1_1,sp))&quant(kolosseum_1_1,one))&refer(kolosseum_1_1,det))&varia(kolosseum_1_1,con))&sort(c620,d))&sort(c620,io))&card(c620,int1))&etype(c620,int0))&fact(c620,real))&gener(c620,sp))&quant(c620,one))&refer(c620,det))&varia(c620,con))&sort(c621,na))&card(c621,int1))&etype(c621,int0))&fact(c621,real))&gener(c621,sp))&quant(c621,one))&refer(c621,indet))&varia(c621,varia_c))&sort(stadt__1_1,d))&sort(stadt__1_1,io))&card(stadt__1_1,int1))&etype(stadt__1_1,int0))&fact(stadt__1_1,real))&gener(stadt__1_1,ge))&quant(stadt__1_1,one))&refer(stadt__1_1,refer_c))&varia(stadt__1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(rom_0,fe))&sort(c895,ent))&card(c895,card_c))&etype(c895,etype_c))&fact(c895,real))&gener(c895,gener_c))&quant(c895,quant_c))&refer(c895,refer_c))&varia(c895,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_004_qapw_93)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((attr(X3,X1)&member(X2,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))))&sub(X1,X2))=>?[X4]:(((mcont(X4,X3)&obj(X4,X3))&scar(X4,X3))&subs(X4,stehen_1_b)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', attr_name__abk__374rzung_stehen_1_b_f__374r)).
fof(c_0_3, axiom, (![X1]:![X2]:![X3]:(member(X1,X3)=>member(X1,cons(X2,X3)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', member_second)).
fof(c_0_4, axiom, (![X1]:![X2]:member(X1,cons(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', member_first)).
fof(c_0_5, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:((((attr(X1,X2)&scar(X4,X3))&sub(X2,name_1_1))&sub(X1,stadt__1_1))&val(X2,rom_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_6, negated_conjecture, (![X5]:![X6]:![X7]:![X8]:((((~attr(X5,X6)|~scar(X8,X7))|~sub(X6,name_1_1))|~sub(X5,stadt__1_1))|~val(X6,rom_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
cnf(c_0_7,negated_conjecture,(~val(X1,rom_0)|~sub(X2,stadt__1_1)|~sub(X1,name_1_1)|~scar(X3,X4)|~attr(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_8,hypothesis,(val(c621,rom_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(sub(c621,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:((((mcont(esk50_3(X5,X6,X7),X7)|((~attr(X7,X5)|~member(X6,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))))|~sub(X5,X6)))&(obj(esk50_3(X5,X6,X7),X7)|((~attr(X7,X5)|~member(X6,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))))|~sub(X5,X6))))&(scar(esk50_3(X5,X6,X7),X7)|((~attr(X7,X5)|~member(X6,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))))|~sub(X5,X6))))&(subs(esk50_3(X5,X6,X7),stehen_1_b)|((~attr(X7,X5)|~member(X6,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))))|~sub(X5,X6))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_11,hypothesis,(~attr(X1,c621)|~scar(X2,X3)|~sub(X1,stadt__1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9])])).
cnf(c_0_12,hypothesis,(attr(c620,c621)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(sub(c620,stadt__1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_14,plain,(scar(esk50_3(X1,X2,X3),X3)|~sub(X1,X2)|~member(X2,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil))))|~attr(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,hypothesis,(~scar(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:(~member(X4,X6)|member(X4,cons(X5,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
cnf(c_0_17,plain,(~attr(X1,X2)|~sub(X2,X3)|~member(X3,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil))))), inference(sr,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_18,plain,(member(X1,cons(X2,X3))|~member(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_19,plain,(~attr(X1,X2)|~sub(X2,X3)|~member(X3,cons(familiename_1_1,cons(name_1_1,nil)))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_20, plain, (![X3]:![X4]:member(X3,cons(X3,X4))), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_21,plain,(~attr(X1,X2)|~sub(X2,X3)|~member(X3,cons(name_1_1,nil))), inference(spm,[status(thm)],[c_0_19, c_0_18])).
cnf(c_0_22,plain,(member(X1,cons(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_23,plain,(~attr(X1,X2)|~sub(X2,name_1_1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_24,hypothesis,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_12]), c_0_9])]), ['proof']).
# SZS output end CNFRefutation.
EOF
