% Problem    : CSR115+70 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n110.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:32:21 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_wp_480)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((arg1(c11,c3)&arg2(c11,c25))&subs(c11,behandeln_1_1))&sub(c25,geschichte_1_1))&mcont(c3,c25))&attch(c40,c25))&attr(c40,c41))&sub(c40,firma_1_1))&sub(c41,name_1_1))&val(c41,bmw_0))&sort(c11,st))&fact(c11,real))&gener(c11,sp))&sort(c3,o))&card(c3,int1))&etype(c3,int0))&fact(c3,real))&gener(c3,sp))&quant(c3,one))&refer(c3,det))&varia(c3,varia_c))&sort(c25,ad))&sort(c25,d))&sort(c25,io))&card(c25,int1))&etype(c25,int0))&fact(c25,hypo))&gener(c25,sp))&quant(c25,one))&refer(c25,det))&varia(c25,con))&sort(behandeln_1_1,st))&fact(behandeln_1_1,real))&gener(behandeln_1_1,ge))&sort(geschichte_1_1,ad))&sort(geschichte_1_1,d))&sort(geschichte_1_1,io))&card(geschichte_1_1,int1))&etype(geschichte_1_1,int0))&fact(geschichte_1_1,real))&gener(geschichte_1_1,ge))&quant(geschichte_1_1,one))&refer(geschichte_1_1,refer_c))&varia(geschichte_1_1,varia_c))&sort(c40,d))&sort(c40,io))&card(c40,int1))&etype(c40,int0))&fact(c40,real))&gener(c40,sp))&quant(c40,one))&refer(c40,det))&varia(c40,con))&sort(c41,na))&card(c41,int1))&etype(c41,int0))&fact(c41,real))&gener(c41,sp))&quant(c41,one))&refer(c41,indet))&varia(c41,varia_c))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(bmw_0,fe))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_wp_480)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_3, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_4, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, negated_conjecture, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((((~attr(X8,X9)|~attr(X11,X10))|~attr(X12,X13))|~obj(X14,X8))|~sub(X9,name_1_1))|~sub(X8,firma_1_1))|~sub(X10,name_1_1))|~val(X9,bmw_0))|~val(X10,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_6,negated_conjecture,(~val(X1,bmw_0)|~val(X2,bmw_0)|~sub(X1,name_1_1)|~sub(X3,firma_1_1)|~sub(X2,name_1_1)|~obj(X4,X3)|~attr(X5,X6)|~attr(X7,X1)|~attr(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,hypothesis,(val(c41,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(sub(c41,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(~val(X1,bmw_0)|~attr(X2,c41)|~attr(X3,X1)|~attr(X4,X5)|~obj(X6,X2)|~sub(X2,firma_1_1)|~sub(X1,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_10,hypothesis,(~attr(X1,c41)|~attr(X2,c41)|~attr(X3,X4)|~obj(X5,X1)|~sub(X1,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_8])])).
cnf(c_0_11,hypothesis,(attr(c40,c41)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_12,hypothesis,(sub(c40,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(~attr(X1,c41)|~attr(X2,X3)|~obj(X4,c40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk19_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk19_3(X7,X8,X9),esk20_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk18_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk18_3(X7,X8,X9),esk19_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk18_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk20_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk19_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk18_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_15,hypothesis,(~attr(X1,X2)|~obj(X3,c40)), inference(spm,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_16,plain,(obj(esk18_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, (![X4]:![X5]:(((arg1(esk24_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk24_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk24_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_18,hypothesis,(~subr(X1,sub_0)|~attr(X2,X3)|~arg2(X1,X4)|~arg1(X1,c40)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(subr(esk24_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,hypothesis,(~attr(X1,X2)|~arg2(esk24_2(X3,X4),X5)|~arg1(esk24_2(X3,X4),c40)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(arg2(esk24_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,hypothesis,(~attr(X1,X2)|~arg1(esk24_2(X3,X4),c40)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,plain,(arg1(esk24_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,hypothesis,(~attr(X1,X2)|~sub(c40,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,hypothesis,(~attr(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_12])).
cnf(c_0_26,hypothesis,($false), inference(sr,[status(thm)],[c_0_11, c_0_25]), ['proof']).
# SZS output end CNFRefutation.
EOF
