% Problem    : CSR113+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n121.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:12:26 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:(((((((flp(X1,X3)&attr(X3,X2))&loc(X4,X1))&scar(X4,X5))&sub(X2,name_1_1))&sub(X5,freiheitsstatue_1_1))&subs(X4,stehen_1_1))&val(X2,new_york_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_003_insicht_3)).
fof(c_0_1, hypothesis, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((loc(c11,c59)&scar(c11,c18))&subs(c11,befinden_1_2))&attr(c13,c14))&sub(c13,stadt__1_1))&sub(c14,name_1_1))&val(c14,paris_0))&pred(c18,freiheitsstatue_1_1))&propr(c18,gleich_1_1))&pred(c36,gr__366__337e_1_1))&loc(c41,c57))&sub(c41,freiheitsstatue_1_1))&attr(c48,c49))&sub(c48,stadt__1_1))&sub(c49,name_1_1))&val(c49,new_york_0))&arg1(c54,c18))&arg2(c54,gleich_1_1))&assoc(c54,c36))&instr(c54,c41))&subr(c54,propr_0))&in(c57,c48))&in(c59,c13))&assoc(freiheitsstatue_1_1,freiheit_1_1))&sub(freiheitsstatue_1_1,statue_1_1))&sort(c11,st))&fact(c11,real))&gener(c11,sp))&sort(c59,l))&card(c59,int1))&etype(c59,int0))&fact(c59,real))&gener(c59,sp))&quant(c59,one))&refer(c59,det))&varia(c59,con))&sort(c18,d))&card(c18,int4))&etype(c18,int1))&fact(c18,real))&gener(c18,sp))&quant(c18,nfquant))&refer(c18,indet))&varia(c18,varia_c))&sort(befinden_1_2,st))&fact(befinden_1_2,real))&gener(befinden_1_2,ge))&sort(c13,d))&sort(c13,io))&card(c13,int1))&etype(c13,int0))&fact(c13,real))&gener(c13,sp))&quant(c13,one))&refer(c13,det))&varia(c13,con))&sort(c14,na))&card(c14,int1))&etype(c14,int0))&fact(c14,real))&gener(c14,sp))&quant(c14,one))&refer(c14,indet))&varia(c14,varia_c))&sort(stadt__1_1,d))&sort(stadt__1_1,io))&card(stadt__1_1,int1))&etype(stadt__1_1,int0))&fact(stadt__1_1,real))&gener(stadt__1_1,ge))&quant(stadt__1_1,one))&refer(stadt__1_1,refer_c))&varia(stadt__1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(paris_0,fe))&sort(freiheitsstatue_1_1,d))&card(freiheitsstatue_1_1,int1))&etype(freiheitsstatue_1_1,int0))&fact(freiheitsstatue_1_1,real))&gener(freiheitsstatue_1_1,ge))&quant(freiheitsstatue_1_1,one))&refer(freiheitsstatue_1_1,refer_c))&varia(freiheitsstatue_1_1,varia_c))&sort(gleich_1_1,rq))&sort(c36,oa))&card(c36,cons(x_constant,cons(int1,nil))))&etype(c36,int1))&fact(c36,real))&gener(c36,sp))&quant(c36,mult))&refer(c36,det))&varia(c36,con))&sort(gr__366__337e_1_1,oa))&card(gr__366__337e_1_1,int1))&etype(gr__366__337e_1_1,int0))&fact(gr__366__337e_1_1,real))&gener(gr__366__337e_1_1,ge))&quant(gr__366__337e_1_1,one))&refer(gr__366__337e_1_1,refer_c))&varia(gr__366__337e_1_1,varia_c))&sort(c41,d))&card(c41,int1))&etype(c41,int0))&fact(c41,real))&gener(c41,sp))&quant(c41,one))&refer(c41,det))&varia(c41,con))&sort(c57,l))&card(c57,int1))&etype(c57,int0))&fact(c57,real))&gener(c57,sp))&quant(c57,one))&refer(c57,det))&varia(c57,con))&sort(c48,d))&sort(c48,io))&card(c48,int1))&etype(c48,int0))&fact(c48,real))&gener(c48,sp))&quant(c48,one))&refer(c48,det))&varia(c48,con))&sort(c49,na))&card(c49,int1))&etype(c49,int0))&fact(c49,real))&gener(c49,sp))&quant(c49,one))&refer(c49,indet))&varia(c49,varia_c))&sort(new_york_0,fe))&sort(c54,st))&fact(c54,real))&gener(c54,sp))&sort(propr_0,st))&fact(propr_0,real))&gener(propr_0,gener_c))&sort(freiheit_1_1,as))&sort(freiheit_1_1,io))&card(freiheit_1_1,int1))&etype(freiheit_1_1,int0))&fact(freiheit_1_1,real))&gener(freiheit_1_1,ge))&quant(freiheit_1_1,one))&refer(freiheit_1_1,refer_c))&varia(freiheit_1_1,varia_c))&sort(statue_1_1,d))&card(statue_1_1,int1))&etype(statue_1_1,int0))&fact(statue_1_1,real))&gener(statue_1_1,ge))&quant(statue_1_1,one))&refer(statue_1_1,refer_c))&varia(statue_1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_003_insicht_3)).
fof(c_0_2, axiom, (![X1]:![X2]:(((in(X1,X2)|an(X1,X2))|bei(X1,X2))=>flp(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', local_function___flp)).
fof(c_0_3, axiom, (![X1]:![X2]:(loc(X1,X2)=>?[X3]:((loc(X3,X2)&scar(X3,X1))&subs(X3,stehen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', loc__stehen_1_1_loc)).
fof(c_0_4, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:(((((((flp(X1,X3)&attr(X3,X2))&loc(X4,X1))&scar(X4,X5))&sub(X2,name_1_1))&sub(X5,freiheitsstatue_1_1))&subs(X4,stehen_1_1))&val(X2,new_york_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, negated_conjecture, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((((~flp(X6,X8)|~attr(X8,X7))|~loc(X9,X6))|~scar(X9,X10))|~sub(X7,name_1_1))|~sub(X10,freiheitsstatue_1_1))|~subs(X9,stehen_1_1))|~val(X7,new_york_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_6,negated_conjecture,(~val(X1,new_york_0)|~subs(X2,stehen_1_1)|~sub(X3,freiheitsstatue_1_1)|~sub(X1,name_1_1)|~scar(X2,X3)|~loc(X2,X4)|~attr(X5,X1)|~flp(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,hypothesis,(val(c49,new_york_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(sub(c49,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_9, plain, (![X3]:![X4]:(((~in(X3,X4)|flp(X3,X4))&(~an(X3,X4)|flp(X3,X4)))&(~bei(X3,X4)|flp(X3,X4)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
cnf(c_0_10,hypothesis,(~attr(X1,c49)|~flp(X2,X1)|~scar(X3,X4)|~loc(X3,X2)|~subs(X3,stehen_1_1)|~sub(X4,freiheitsstatue_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_11,plain,(flp(X1,X2)|~in(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12,hypothesis,(~in(X1,X2)|~attr(X2,c49)|~scar(X3,X4)|~loc(X3,X1)|~subs(X3,stehen_1_1)|~sub(X4,freiheitsstatue_1_1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_13,hypothesis,(in(c57,c48)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_14,hypothesis,(attr(c48,c49)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_15, plain, (![X4]:![X5]:(((loc(esk12_2(X4,X5),X5)|~loc(X4,X5))&(scar(esk12_2(X4,X5),X4)|~loc(X4,X5)))&(subs(esk12_2(X4,X5),stehen_1_1)|~loc(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_16,hypothesis,(~scar(X1,X2)|~loc(X1,c57)|~subs(X1,stehen_1_1)|~sub(X2,freiheitsstatue_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
cnf(c_0_17,plain,(scar(esk12_2(X1,X2),X1)|~loc(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18,plain,(subs(esk12_2(X1,X2),stehen_1_1)|~loc(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19,hypothesis,(~loc(esk12_2(X1,X2),c57)|~loc(X1,X2)|~sub(X1,freiheitsstatue_1_1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])).
cnf(c_0_20,plain,(loc(esk12_2(X1,X2),X2)|~loc(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,hypothesis,(~loc(X1,c57)|~sub(X1,freiheitsstatue_1_1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_22,hypothesis,(loc(c41,c57)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_23,hypothesis,(sub(c41,freiheitsstatue_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_24,hypothesis,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])]), ['proof']).
# SZS output end CNFRefutation.
EOF
