% Problem    : LCL640+1.010 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n103.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:06:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____042_B03_F1_AE_Q4_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (~(?[X1]:~(((((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|(((p1(X2)|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|~(![X1]:((~(r1(X2,X1))|p1(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2)))))))&![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))))))))))))|~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|(((p1(X1)|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|~(![X2]:((~(r1(X1,X2))|p1(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1)))))))&![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))))))))))))|~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|(((p1(X2)|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|~(![X1]:((~(r1(X2,X1))|p1(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2)))))))&![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))|~(![X2]:(~(r1(X1,X2))|p1(X2))))))))))))))|~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|((((((![X2]:(~(r1(X1,X2))|p1(X2))|~(p1(X1)))|![X2]:(~(r1(X1,X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1))))))|~(![X2]:(((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2)))|~(![X1]:((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2))))|~((![X2]:(~(r1(X1,X2))|p1(X2))|~(p1(X1)))))))))&((p1(X1)|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|~(![X2]:((~(r1(X1,X2))|p1(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1))))))))&![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))&(![X2]:(~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|p1(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2))))))|~(![X2]:((~(r1(X1,X2))|p1(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1))))))))))))))))|~((((((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|~(![X2]:((~(r1(X1,X2))|p1(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1)))))))))))))&![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|(((p1(X2)|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|~(![X1]:((~(r1(X2,X1))|p1(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2)))))))&![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))))))))&![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|(((p1(X1)|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|~(![X2]:((~(r1(X1,X2))|p1(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1)))))))&![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))|~(![X1]:(~(r1(X2,X1))|p1(X1))))))))))&![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|(((p1(X2)|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|~(![X1]:((~(r1(X2,X1))|p1(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2)))))))&![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))))))&![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|(((p1(X1)|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|~(![X2]:((~(r1(X1,X2))|p1(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1)))))))&![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))|~(![X1]:(~(r1(X2,X1))|p1(X1))))))))&![X2]:(~(r1(X1,X2))|(((p1(X2)|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|~(![X1]:((~(r1(X2,X1))|p1(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1)))|~(p1(X2)))))))&![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', main)).
fof(c_0_1, plain, (![X1]:(epred1_1(X1)<=>((((((![X2]:(~r1(X1,X2)|p1(X2))|~p1(X1))|![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1)))))|~(![X2]:(((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2)))|~((![X2]:(~r1(X1,X2)|p1(X2))|~p1(X1))))))))&((p1(X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1)))))))&![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))|~(![X1]:(~r1(X2,X1)|p1(X1)))))&(![X2]:(~r1(X1,X2)|![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2)))))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))))))))), introduced(definition)).
fof(c_0_2, negated_conjecture, (~(~(?[X1]:~(((((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|(((p1(X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))))))&![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))|~(![X2]:(~r1(X1,X2)|p1(X2)))))))))))))))|~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|(((p1(X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))))))&![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))|~(![X1]:(~r1(X2,X1)|p1(X1)))))))))))))))|~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|(((p1(X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))))))&![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))|~(![X2]:(~r1(X1,X2)|p1(X2))))))))))))))|~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|epred1_1(X1)))))))))|~((((((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))))))))))))&![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|(((p1(X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))))))&![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))|~(![X2]:(~r1(X1,X2)|p1(X2)))))))))))&![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|(((p1(X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))))))&![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))|~(![X1]:(~r1(X2,X1)|p1(X1))))))))))&![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|(((p1(X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))))))&![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))|~(![X2]:(~r1(X1,X2)|p1(X2)))))))))&![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|(((p1(X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))))))&![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))|~(![X1]:(~r1(X2,X1)|p1(X1))))))))&![X2]:(~r1(X1,X2)|(((p1(X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|~(![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))))))&![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))|~(![X2]:(~r1(X1,X2)|p1(X2))))))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_1])).
fof(c_0_3, negated_conjecture, (![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X16]:![X17]:![X18]:![X19]:![X20]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X33]:![X34]:![X35]:![X36]:![X37]:![X39]:![X40]:![X41]:![X42]:![X43]:![X44]:![X45]:![X46]:![X49]:![X50]:![X51]:![X52]:![X53]:![X55]:![X56]:![X57]:![X58]:![X59]:![X60]:![X69]:![X70]:![X73]:![X74]:![X75]:![X76]:![X77]:![X78]:![X81]:![X82]:![X83]:![X84]:![X85]:![X87]:![X88]:![X89]:![X90]:![X91]:![X94]:![X95]:![X96]:![X97]:![X98]:![X100]:![X101]:![X102]:![X103]:![X106]:![X107]:![X108]:![X109]:![X110]:![X112]:![X113]:![X114]:![X117]:![X118]:![X119]:![X120]:![X121]:![X123]:![X124]:![X127]:![X128]:![X129]:![X130]:![X131]:((((((((((((((((((r1(X12,esk3_9(X4,X5,X6,X7,X8,X9,X10,X11,X12))|((r1(X13,esk2_10(X4,X5,X6,X7,X8,X9,X10,X11,X12,X13))|~r1(X12,X13))|p1(X12)))|~r1(X11,X12))|~r1(X10,X11))|~r1(X9,X10))|~r1(X8,X9))|~r1(X7,X8))|~r1(X6,X7))|~r1(X5,X6))|~r1(X4,X5))|~r1(esk1_0,X4))&((((((((((~p1(esk3_9(X4,X5,X6,X7,X8,X9,X10,X11,X12))|((r1(X13,esk2_10(X4,X5,X6,X7,X8,X9,X10,X11,X12,X13))|~r1(X12,X13))|p1(X12)))|~r1(X11,X12))|~r1(X10,X11))|~r1(X9,X10))|~r1(X8,X9))|~r1(X7,X8))|~r1(X6,X7))|~r1(X5,X6))|~r1(X4,X5))|~r1(esk1_0,X4)))&((((((((((((~r1(esk3_9(X4,X5,X6,X7,X8,X9,X10,X11,X12),X16)|(~r1(X16,X17)|p1(X17)))|~p1(X16))|((r1(X13,esk2_10(X4,X5,X6,X7,X8,X9,X10,X11,X12,X13))|~r1(X12,X13))|p1(X12)))|~r1(X11,X12))|~r1(X10,X11))|~r1(X9,X10))|~r1(X8,X9))|~r1(X7,X8))|~r1(X6,X7))|~r1(X5,X6))|~r1(X4,X5))|~r1(esk1_0,X4)))&((((((((((((r1(X12,esk3_9(X4,X5,X6,X7,X8,X9,X10,X11,X12))|((~p1(esk2_10(X4,X5,X6,X7,X8,X9,X10,X11,X12,X13))|~r1(X12,X13))|p1(X12)))|~r1(X11,X12))|~r1(X10,X11))|~r1(X9,X10))|~r1(X8,X9))|~r1(X7,X8))|~r1(X6,X7))|~r1(X5,X6))|~r1(X4,X5))|~r1(esk1_0,X4))&((((((((((~p1(esk3_9(X4,X5,X6,X7,X8,X9,X10,X11,X12))|((~p1(esk2_10(X4,X5,X6,X7,X8,X9,X10,X11,X12,X13))|~r1(X12,X13))|p1(X12)))|~r1(X11,X12))|~r1(X10,X11))|~r1(X9,X10))|~r1(X8,X9))|~r1(X7,X8))|~r1(X6,X7))|~r1(X5,X6))|~r1(X4,X5))|~r1(esk1_0,X4)))&((((((((((((~r1(esk3_9(X4,X5,X6,X7,X8,X9,X10,X11,X12),X16)|(~r1(X16,X17)|p1(X17)))|~p1(X16))|((~p1(esk2_10(X4,X5,X6,X7,X8,X9,X10,X11,X12,X13))|~r1(X12,X13))|p1(X12)))|~r1(X11,X12))|~r1(X10,X11))|~r1(X9,X10))|~r1(X8,X9))|~r1(X7,X8))|~r1(X6,X7))|~r1(X5,X6))|~r1(X4,X5))|~r1(esk1_0,X4))))&(((((((((((r1(X18,esk4_10(X4,X5,X6,X7,X8,X9,X10,X11,X12,X18))|(~r1(X12,X18)|(~r1(X18,X19)|(~r1(X19,X20)|p1(X20)))))|~r1(X11,X12))|~r1(X10,X11))|~r1(X9,X10))|~r1(X8,X9))|~r1(X7,X8))|~r1(X6,X7))|~r1(X5,X6))|~r1(X4,X5))|~r1(esk1_0,X4))&((((((((((~p1(esk4_10(X4,X5,X6,X7,X8,X9,X10,X11,X12,X18))|(~r1(X12,X18)|(~r1(X18,X19)|(~r1(X19,X20)|p1(X20)))))|~r1(X11,X12))|~r1(X10,X11))|~r1(X9,X10))|~r1(X8,X9))|~r1(X7,X8))|~r1(X6,X7))|~r1(X5,X6))|~r1(X4,X5))|~r1(esk1_0,X4))))&(((((((((((((r1(X29,esk6_8(X22,X23,X24,X25,X26,X27,X28,X29))|((r1(X30,esk5_9(X22,X23,X24,X25,X26,X27,X28,X29,X30))|~r1(X29,X30))|p1(X29)))|~r1(X28,X29))|~r1(X27,X28))|~r1(X26,X27))|~r1(X25,X26))|~r1(X24,X25))|~r1(X23,X24))|~r1(X22,X23))|~r1(esk1_0,X22))&(((((((((~p1(esk6_8(X22,X23,X24,X25,X26,X27,X28,X29))|((r1(X30,esk5_9(X22,X23,X24,X25,X26,X27,X28,X29,X30))|~r1(X29,X30))|p1(X29)))|~r1(X28,X29))|~r1(X27,X28))|~r1(X26,X27))|~r1(X25,X26))|~r1(X24,X25))|~r1(X23,X24))|~r1(X22,X23))|~r1(esk1_0,X22)))&(((((((((((~r1(esk6_8(X22,X23,X24,X25,X26,X27,X28,X29),X33)|(~r1(X33,X34)|p1(X34)))|~p1(X33))|((r1(X30,esk5_9(X22,X23,X24,X25,X26,X27,X28,X29,X30))|~r1(X29,X30))|p1(X29)))|~r1(X28,X29))|~r1(X27,X28))|~r1(X26,X27))|~r1(X25,X26))|~r1(X24,X25))|~r1(X23,X24))|~r1(X22,X23))|~r1(esk1_0,X22)))&(((((((((((r1(X29,esk6_8(X22,X23,X24,X25,X26,X27,X28,X29))|((~p1(esk5_9(X22,X23,X24,X25,X26,X27,X28,X29,X30))|~r1(X29,X30))|p1(X29)))|~r1(X28,X29))|~r1(X27,X28))|~r1(X26,X27))|~r1(X25,X26))|~r1(X24,X25))|~r1(X23,X24))|~r1(X22,X23))|~r1(esk1_0,X22))&(((((((((~p1(esk6_8(X22,X23,X24,X25,X26,X27,X28,X29))|((~p1(esk5_9(X22,X23,X24,X25,X26,X27,X28,X29,X30))|~r1(X29,X30))|p1(X29)))|~r1(X28,X29))|~r1(X27,X28))|~r1(X26,X27))|~r1(X25,X26))|~r1(X24,X25))|~r1(X23,X24))|~r1(X22,X23))|~r1(esk1_0,X22)))&(((((((((((~r1(esk6_8(X22,X23,X24,X25,X26,X27,X28,X29),X33)|(~r1(X33,X34)|p1(X34)))|~p1(X33))|((~p1(esk5_9(X22,X23,X24,X25,X26,X27,X28,X29,X30))|~r1(X29,X30))|p1(X29)))|~r1(X28,X29))|~r1(X27,X28))|~r1(X26,X27))|~r1(X25,X26))|~r1(X24,X25))|~r1(X23,X24))|~r1(X22,X23))|~r1(esk1_0,X22))))&((((((((((r1(X35,esk7_9(X22,X23,X24,X25,X26,X27,X28,X29,X35))|(~r1(X29,X35)|(~r1(X35,X36)|(~r1(X36,X37)|p1(X37)))))|~r1(X28,X29))|~r1(X27,X28))|~r1(X26,X27))|~r1(X25,X26))|~r1(X24,X25))|~r1(X23,X24))|~r1(X22,X23))|~r1(esk1_0,X22))&(((((((((~p1(esk7_9(X22,X23,X24,X25,X26,X27,X28,X29,X35))|(~r1(X29,X35)|(~r1(X35,X36)|(~r1(X36,X37)|p1(X37)))))|~r1(X28,X29))|~r1(X27,X28))|~r1(X26,X27))|~r1(X25,X26))|~r1(X24,X25))|~r1(X23,X24))|~r1(X22,X23))|~r1(esk1_0,X22)))))&((((((((((((r1(X45,esk9_7(X39,X40,X41,X42,X43,X44,X45))|((r1(X46,esk8_8(X39,X40,X41,X42,X43,X44,X45,X46))|~r1(X45,X46))|p1(X45)))|~r1(X44,X45))|~r1(X43,X44))|~r1(X42,X43))|~r1(X41,X42))|~r1(X40,X41))|~r1(X39,X40))|~r1(esk1_0,X39))&((((((((~p1(esk9_7(X39,X40,X41,X42,X43,X44,X45))|((r1(X46,esk8_8(X39,X40,X41,X42,X43,X44,X45,X46))|~r1(X45,X46))|p1(X45)))|~r1(X44,X45))|~r1(X43,X44))|~r1(X42,X43))|~r1(X41,X42))|~r1(X40,X41))|~r1(X39,X40))|~r1(esk1_0,X39)))&((((((((((~r1(esk9_7(X39,X40,X41,X42,X43,X44,X45),X49)|(~r1(X49,X50)|p1(X50)))|~p1(X49))|((r1(X46,esk8_8(X39,X40,X41,X42,X43,X44,X45,X46))|~r1(X45,X46))|p1(X45)))|~r1(X44,X45))|~r1(X43,X44))|~r1(X42,X43))|~r1(X41,X42))|~r1(X40,X41))|~r1(X39,X40))|~r1(esk1_0,X39)))&((((((((((r1(X45,esk9_7(X39,X40,X41,X42,X43,X44,X45))|((~p1(esk8_8(X39,X40,X41,X42,X43,X44,X45,X46))|~r1(X45,X46))|p1(X45)))|~r1(X44,X45))|~r1(X43,X44))|~r1(X42,X43))|~r1(X41,X42))|~r1(X40,X41))|~r1(X39,X40))|~r1(esk1_0,X39))&((((((((~p1(esk9_7(X39,X40,X41,X42,X43,X44,X45))|((~p1(esk8_8(X39,X40,X41,X42,X43,X44,X45,X46))|~r1(X45,X46))|p1(X45)))|~r1(X44,X45))|~r1(X43,X44))|~r1(X42,X43))|~r1(X41,X42))|~r1(X40,X41))|~r1(X39,X40))|~r1(esk1_0,X39)))&((((((((((~r1(esk9_7(X39,X40,X41,X42,X43,X44,X45),X49)|(~r1(X49,X50)|p1(X50)))|~p1(X49))|((~p1(esk8_8(X39,X40,X41,X42,X43,X44,X45,X46))|~r1(X45,X46))|p1(X45)))|~r1(X44,X45))|~r1(X43,X44))|~r1(X42,X43))|~r1(X41,X42))|~r1(X40,X41))|~r1(X39,X40))|~r1(esk1_0,X39))))&(((((((((r1(X51,esk10_8(X39,X40,X41,X42,X43,X44,X45,X51))|(~r1(X45,X51)|(~r1(X51,X52)|(~r1(X52,X53)|p1(X53)))))|~r1(X44,X45))|~r1(X43,X44))|~r1(X42,X43))|~r1(X41,X42))|~r1(X40,X41))|~r1(X39,X40))|~r1(esk1_0,X39))&((((((((~p1(esk10_8(X39,X40,X41,X42,X43,X44,X45,X51))|(~r1(X45,X51)|(~r1(X51,X52)|(~r1(X52,X53)|p1(X53)))))|~r1(X44,X45))|~r1(X43,X44))|~r1(X42,X43))|~r1(X41,X42))|~r1(X40,X41))|~r1(X39,X40))|~r1(esk1_0,X39)))))&(~r1(esk1_0,X55)|(~r1(X55,X56)|(~r1(X56,X57)|(~r1(X57,X58)|(~r1(X58,X59)|(~r1(X59,X60)|epred1_1(X60))))))))&((((((r1(esk1_0,esk11_0)&(r1(esk11_0,esk12_0)&(r1(esk12_0,esk13_0)&(r1(esk13_0,esk14_0)&(r1(esk14_0,esk15_0)&(((r1(esk15_0,esk16_0)&(r1(esk16_0,esk17_0)&~p1(esk17_0)))&(r1(esk16_0,esk18_0)&(~r1(esk18_0,X69)|p1(X69))))&(((r1(X70,esk19_1(X70))|(~r1(esk16_0,X70)|p1(X70)))&((r1(esk19_1(X70),esk20_1(X70))|(~r1(esk16_0,X70)|p1(X70)))&(~p1(esk20_1(X70))|(~r1(esk16_0,X70)|p1(X70)))))&(p1(esk19_1(X70))|(~r1(esk16_0,X70)|p1(X70))))))))))&((((((((((r1(X77,esk22_5(X73,X74,X75,X76,X77))|((r1(X78,esk21_6(X73,X74,X75,X76,X77,X78))|~r1(X77,X78))|p1(X77)))|~r1(X76,X77))|~r1(X75,X76))|~r1(X74,X75))|~r1(X73,X74))|~r1(esk1_0,X73))&((((((~p1(esk22_5(X73,X74,X75,X76,X77))|((r1(X78,esk21_6(X73,X74,X75,X76,X77,X78))|~r1(X77,X78))|p1(X77)))|~r1(X76,X77))|~r1(X75,X76))|~r1(X74,X75))|~r1(X73,X74))|~r1(esk1_0,X73)))&((((((((~r1(esk22_5(X73,X74,X75,X76,X77),X81)|(~r1(X81,X82)|p1(X82)))|~p1(X81))|((r1(X78,esk21_6(X73,X74,X75,X76,X77,X78))|~r1(X77,X78))|p1(X77)))|~r1(X76,X77))|~r1(X75,X76))|~r1(X74,X75))|~r1(X73,X74))|~r1(esk1_0,X73)))&((((((((r1(X77,esk22_5(X73,X74,X75,X76,X77))|((~p1(esk21_6(X73,X74,X75,X76,X77,X78))|~r1(X77,X78))|p1(X77)))|~r1(X76,X77))|~r1(X75,X76))|~r1(X74,X75))|~r1(X73,X74))|~r1(esk1_0,X73))&((((((~p1(esk22_5(X73,X74,X75,X76,X77))|((~p1(esk21_6(X73,X74,X75,X76,X77,X78))|~r1(X77,X78))|p1(X77)))|~r1(X76,X77))|~r1(X75,X76))|~r1(X74,X75))|~r1(X73,X74))|~r1(esk1_0,X73)))&((((((((~r1(esk22_5(X73,X74,X75,X76,X77),X81)|(~r1(X81,X82)|p1(X82)))|~p1(X81))|((~p1(esk21_6(X73,X74,X75,X76,X77,X78))|~r1(X77,X78))|p1(X77)))|~r1(X76,X77))|~r1(X75,X76))|~r1(X74,X75))|~r1(X73,X74))|~r1(esk1_0,X73))))&(((((((r1(X83,esk23_6(X73,X74,X75,X76,X77,X83))|(~r1(X77,X83)|(~r1(X83,X84)|(~r1(X84,X85)|p1(X85)))))|~r1(X76,X77))|~r1(X75,X76))|~r1(X74,X75))|~r1(X73,X74))|~r1(esk1_0,X73))&((((((~p1(esk23_6(X73,X74,X75,X76,X77,X83))|(~r1(X77,X83)|(~r1(X83,X84)|(~r1(X84,X85)|p1(X85)))))|~r1(X76,X77))|~r1(X75,X76))|~r1(X74,X75))|~r1(X73,X74))|~r1(esk1_0,X73)))))&(((((((((r1(X90,esk25_4(X87,X88,X89,X90))|((r1(X91,esk24_5(X87,X88,X89,X90,X91))|~r1(X90,X91))|p1(X90)))|~r1(X89,X90))|~r1(X88,X89))|~r1(X87,X88))|~r1(esk1_0,X87))&(((((~p1(esk25_4(X87,X88,X89,X90))|((r1(X91,esk24_5(X87,X88,X89,X90,X91))|~r1(X90,X91))|p1(X90)))|~r1(X89,X90))|~r1(X88,X89))|~r1(X87,X88))|~r1(esk1_0,X87)))&(((((((~r1(esk25_4(X87,X88,X89,X90),X94)|(~r1(X94,X95)|p1(X95)))|~p1(X94))|((r1(X91,esk24_5(X87,X88,X89,X90,X91))|~r1(X90,X91))|p1(X90)))|~r1(X89,X90))|~r1(X88,X89))|~r1(X87,X88))|~r1(esk1_0,X87)))&(((((((r1(X90,esk25_4(X87,X88,X89,X90))|((~p1(esk24_5(X87,X88,X89,X90,X91))|~r1(X90,X91))|p1(X90)))|~r1(X89,X90))|~r1(X88,X89))|~r1(X87,X88))|~r1(esk1_0,X87))&(((((~p1(esk25_4(X87,X88,X89,X90))|((~p1(esk24_5(X87,X88,X89,X90,X91))|~r1(X90,X91))|p1(X90)))|~r1(X89,X90))|~r1(X88,X89))|~r1(X87,X88))|~r1(esk1_0,X87)))&(((((((~r1(esk25_4(X87,X88,X89,X90),X94)|(~r1(X94,X95)|p1(X95)))|~p1(X94))|((~p1(esk24_5(X87,X88,X89,X90,X91))|~r1(X90,X91))|p1(X90)))|~r1(X89,X90))|~r1(X88,X89))|~r1(X87,X88))|~r1(esk1_0,X87))))&((((((r1(X96,esk26_5(X87,X88,X89,X90,X96))|(~r1(X90,X96)|(~r1(X96,X97)|(~r1(X97,X98)|p1(X98)))))|~r1(X89,X90))|~r1(X88,X89))|~r1(X87,X88))|~r1(esk1_0,X87))&(((((~p1(esk26_5(X87,X88,X89,X90,X96))|(~r1(X90,X96)|(~r1(X96,X97)|(~r1(X97,X98)|p1(X98)))))|~r1(X89,X90))|~r1(X88,X89))|~r1(X87,X88))|~r1(esk1_0,X87)))))&((((((((r1(X102,esk28_3(X100,X101,X102))|((r1(X103,esk27_4(X100,X101,X102,X103))|~r1(X102,X103))|p1(X102)))|~r1(X101,X102))|~r1(X100,X101))|~r1(esk1_0,X100))&((((~p1(esk28_3(X100,X101,X102))|((r1(X103,esk27_4(X100,X101,X102,X103))|~r1(X102,X103))|p1(X102)))|~r1(X101,X102))|~r1(X100,X101))|~r1(esk1_0,X100)))&((((((~r1(esk28_3(X100,X101,X102),X106)|(~r1(X106,X107)|p1(X107)))|~p1(X106))|((r1(X103,esk27_4(X100,X101,X102,X103))|~r1(X102,X103))|p1(X102)))|~r1(X101,X102))|~r1(X100,X101))|~r1(esk1_0,X100)))&((((((r1(X102,esk28_3(X100,X101,X102))|((~p1(esk27_4(X100,X101,X102,X103))|~r1(X102,X103))|p1(X102)))|~r1(X101,X102))|~r1(X100,X101))|~r1(esk1_0,X100))&((((~p1(esk28_3(X100,X101,X102))|((~p1(esk27_4(X100,X101,X102,X103))|~r1(X102,X103))|p1(X102)))|~r1(X101,X102))|~r1(X100,X101))|~r1(esk1_0,X100)))&((((((~r1(esk28_3(X100,X101,X102),X106)|(~r1(X106,X107)|p1(X107)))|~p1(X106))|((~p1(esk27_4(X100,X101,X102,X103))|~r1(X102,X103))|p1(X102)))|~r1(X101,X102))|~r1(X100,X101))|~r1(esk1_0,X100))))&(((((r1(X108,esk29_4(X100,X101,X102,X108))|(~r1(X102,X108)|(~r1(X108,X109)|(~r1(X109,X110)|p1(X110)))))|~r1(X101,X102))|~r1(X100,X101))|~r1(esk1_0,X100))&((((~p1(esk29_4(X100,X101,X102,X108))|(~r1(X102,X108)|(~r1(X108,X109)|(~r1(X109,X110)|p1(X110)))))|~r1(X101,X102))|~r1(X100,X101))|~r1(esk1_0,X100)))))&(((((((r1(X113,esk31_2(X112,X113))|((r1(X114,esk30_3(X112,X113,X114))|~r1(X113,X114))|p1(X113)))|~r1(X112,X113))|~r1(esk1_0,X112))&(((~p1(esk31_2(X112,X113))|((r1(X114,esk30_3(X112,X113,X114))|~r1(X113,X114))|p1(X113)))|~r1(X112,X113))|~r1(esk1_0,X112)))&(((((~r1(esk31_2(X112,X113),X117)|(~r1(X117,X118)|p1(X118)))|~p1(X117))|((r1(X114,esk30_3(X112,X113,X114))|~r1(X113,X114))|p1(X113)))|~r1(X112,X113))|~r1(esk1_0,X112)))&(((((r1(X113,esk31_2(X112,X113))|((~p1(esk30_3(X112,X113,X114))|~r1(X113,X114))|p1(X113)))|~r1(X112,X113))|~r1(esk1_0,X112))&(((~p1(esk31_2(X112,X113))|((~p1(esk30_3(X112,X113,X114))|~r1(X113,X114))|p1(X113)))|~r1(X112,X113))|~r1(esk1_0,X112)))&(((((~r1(esk31_2(X112,X113),X117)|(~r1(X117,X118)|p1(X118)))|~p1(X117))|((~p1(esk30_3(X112,X113,X114))|~r1(X113,X114))|p1(X113)))|~r1(X112,X113))|~r1(esk1_0,X112))))&((((r1(X119,esk32_3(X112,X113,X119))|(~r1(X113,X119)|(~r1(X119,X120)|(~r1(X120,X121)|p1(X121)))))|~r1(X112,X113))|~r1(esk1_0,X112))&(((~p1(esk32_3(X112,X113,X119))|(~r1(X113,X119)|(~r1(X119,X120)|(~r1(X120,X121)|p1(X121)))))|~r1(X112,X113))|~r1(esk1_0,X112)))))&((((((r1(X123,esk34_1(X123))|((r1(X124,esk33_2(X123,X124))|~r1(X123,X124))|p1(X123)))|~r1(esk1_0,X123))&((~p1(esk34_1(X123))|((r1(X124,esk33_2(X123,X124))|~r1(X123,X124))|p1(X123)))|~r1(esk1_0,X123)))&((((~r1(esk34_1(X123),X127)|(~r1(X127,X128)|p1(X128)))|~p1(X127))|((r1(X124,esk33_2(X123,X124))|~r1(X123,X124))|p1(X123)))|~r1(esk1_0,X123)))&((((r1(X123,esk34_1(X123))|((~p1(esk33_2(X123,X124))|~r1(X123,X124))|p1(X123)))|~r1(esk1_0,X123))&((~p1(esk34_1(X123))|((~p1(esk33_2(X123,X124))|~r1(X123,X124))|p1(X123)))|~r1(esk1_0,X123)))&((((~r1(esk34_1(X123),X127)|(~r1(X127,X128)|p1(X128)))|~p1(X127))|((~p1(esk33_2(X123,X124))|~r1(X123,X124))|p1(X123)))|~r1(esk1_0,X123))))&(((r1(X129,esk35_2(X123,X129))|(~r1(X123,X129)|(~r1(X129,X130)|(~r1(X130,X131)|p1(X131)))))|~r1(esk1_0,X123))&((~p1(esk35_2(X123,X129))|(~r1(X123,X129)|(~r1(X129,X130)|(~r1(X130,X131)|p1(X131)))))|~r1(esk1_0,X123))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_4,negated_conjecture,(epred1_1(X1)|~r1(X2,X1)|~r1(X3,X2)|~r1(X4,X3)|~r1(X5,X4)|~r1(X6,X5)|~r1(esk1_0,X6)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_5,negated_conjecture,(r1(esk1_0,esk11_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,negated_conjecture,(epred1_1(X1)|~r1(esk11_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X1)), inference(spm,[status(thm)],[c_0_4, c_0_5])).
cnf(c_0_7,negated_conjecture,(r1(esk11_0,esk12_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_8,negated_conjecture,(epred1_1(X1)|~r1(esk12_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X1)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_9,negated_conjecture,(r1(esk12_0,esk13_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_10, plain, (![X1]:(epred1_1(X1)=>((((((![X2]:(~r1(X1,X2)|p1(X2))|~p1(X1))|![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1)))))|~(![X2]:(((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2)))|~((![X2]:(~r1(X1,X2)|p1(X2))|~p1(X1))))))))&((p1(X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1)))))))&![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))|~(![X1]:(~r1(X2,X1)|p1(X1)))))&(![X2]:(~r1(X1,X2)|![X1]:((~r1(X2,X1)|p1(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1)))|~p1(X2)))))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))))))))), inference(split_equiv,[status(thm)],[c_0_1])).
cnf(c_0_11,negated_conjecture,(epred1_1(X1)|~r1(esk13_0,X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12,negated_conjecture,(r1(esk13_0,esk14_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_13, plain, (![X3]:![X4]:![X5]:![X10]:![X11]:![X12]:![X14]:![X17]:![X18]:![X19]:![X20]:![X21]:![X23]:![X24]:![X28]:![X29]:((((((((((r1(X3,esk38_1(X3))|((r1(X5,esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((r1(esk38_1(X3),esk39_1(X3))|((r1(X5,esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&((~p1(esk39_1(X3))|((r1(X5,esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((p1(esk38_1(X3))|((r1(X5,esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((((r1(X10,esk40_2(X3,X10))|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(X5,esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((~p1(esk40_2(X3,X10))|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(X5,esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((p1(X10)|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(X5,esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((((((r1(X3,esk38_1(X3))|((r1(esk36_2(X3,X5),esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((r1(esk38_1(X3),esk39_1(X3))|((r1(esk36_2(X3,X5),esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&((~p1(esk39_1(X3))|((r1(esk36_2(X3,X5),esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((p1(esk38_1(X3))|((r1(esk36_2(X3,X5),esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((((r1(X10,esk40_2(X3,X10))|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(esk36_2(X3,X5),esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((~p1(esk40_2(X3,X10))|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(esk36_2(X3,X5),esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((p1(X10)|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((r1(esk36_2(X3,X5),esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&(((((r1(X3,esk38_1(X3))|((~p1(esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((r1(esk38_1(X3),esk39_1(X3))|((~p1(esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&((~p1(esk39_1(X3))|((~p1(esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((p1(esk38_1(X3))|((~p1(esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((((r1(X10,esk40_2(X3,X10))|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((~p1(esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((~p1(esk40_2(X3,X10))|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((~p1(esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((p1(X10)|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((~p1(esk37_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))))&(((((r1(X3,esk38_1(X3))|((p1(esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((r1(esk38_1(X3),esk39_1(X3))|((p1(esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&((~p1(esk39_1(X3))|((p1(esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))))&((p1(esk38_1(X3))|((p1(esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((((r1(X10,esk40_2(X3,X10))|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((p1(esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3))&(((~p1(esk40_2(X3,X10))|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((p1(esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))&(((p1(X10)|(~r1(esk38_1(X3),X10)|((~r1(X10,X11)|(~r1(X11,X12)|p1(X12)))|~p1(X11))))|((p1(esk36_2(X3,X5))|~r1(X3,X5))|((~r1(X3,X4)|p1(X4))|~p1(X3))))|~epred1_1(X3)))))&(((((r1(X3,esk42_1(X3))|((r1(X14,esk41_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3))&((~p1(esk42_1(X3))|((r1(X14,esk41_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3)))&((((~r1(esk42_1(X3),X17)|(~r1(X17,X18)|p1(X18)))|~p1(X17))|((r1(X14,esk41_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3)))&((((r1(X3,esk42_1(X3))|((~p1(esk41_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3))&((~p1(esk42_1(X3))|((~p1(esk41_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3)))&((((~r1(esk42_1(X3),X17)|(~r1(X17,X18)|p1(X18)))|~p1(X17))|((~p1(esk41_2(X3,X14))|~r1(X3,X14))|p1(X3)))|~epred1_1(X3)))))&(((r1(X19,esk43_2(X3,X19))|(~r1(X3,X19)|(~r1(X19,X20)|(~r1(X20,X21)|p1(X21)))))|~epred1_1(X3))&((~p1(esk43_2(X3,X19))|(~r1(X3,X19)|(~r1(X19,X20)|(~r1(X20,X21)|p1(X21)))))|~epred1_1(X3))))&((((((r1(X3,esk46_1(X3))|((r1(X24,esk44_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3))&((~p1(esk46_1(X3))|((r1(X24,esk44_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((~r1(esk46_1(X3),X28)|(~r1(X28,X29)|p1(X29)))|~p1(X28))|((r1(X24,esk44_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&(((((r1(X3,esk46_1(X3))|((r1(esk44_3(X3,X23,X24),esk45_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3))&((~p1(esk46_1(X3))|((r1(esk44_3(X3,X23,X24),esk45_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((~r1(esk46_1(X3),X28)|(~r1(X28,X29)|p1(X29)))|~p1(X28))|((r1(esk44_3(X3,X23,X24),esk45_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((r1(X3,esk46_1(X3))|((~p1(esk45_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3))&((~p1(esk46_1(X3))|((~p1(esk45_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((~r1(esk46_1(X3),X28)|(~r1(X28,X29)|p1(X29)))|~p1(X28))|((~p1(esk45_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))))&((((r1(X3,esk46_1(X3))|((p1(esk44_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3))&((~p1(esk46_1(X3))|((p1(esk44_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))&((((~r1(esk46_1(X3),X28)|(~r1(X28,X29)|p1(X29)))|~p1(X28))|((p1(esk44_3(X3,X23,X24))|(~r1(X23,X24)|p1(X24)))|~r1(X3,X23)))|~epred1_1(X3)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
cnf(c_0_14,negated_conjecture,(epred1_1(X1)|~r1(esk14_0,X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15,negated_conjecture,(r1(esk14_0,esk15_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_16,plain,(p1(X1)|r1(X2,esk41_2(X1,X2))|r1(X1,esk42_1(X1))|~epred1_1(X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,negated_conjecture,(r1(esk16_0,esk18_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_18,negated_conjecture,(epred1_1(X1)|~r1(esk15_0,X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19,negated_conjecture,(r1(esk15_0,esk16_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_20,negated_conjecture,(p1(esk16_0)|r1(esk18_0,esk41_2(esk16_0,esk18_0))|r1(esk16_0,esk42_1(esk16_0))|~epred1_1(esk16_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21,negated_conjecture,(epred1_1(esk16_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22,negated_conjecture,(p1(X1)|~r1(esk18_0,X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_23,negated_conjecture,(p1(esk16_0)|r1(esk18_0,esk41_2(esk16_0,esk18_0))|r1(esk16_0,esk42_1(esk16_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21])])).
cnf(c_0_24,plain,(p1(X1)|r1(X1,esk42_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~p1(esk41_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,negated_conjecture,(p1(esk41_2(esk16_0,esk18_0))|p1(esk16_0)|r1(esk16_0,esk42_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26,negated_conjecture,(p1(X1)|r1(X1,esk19_1(X1))|~r1(esk16_0,X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_27,plain,(p1(esk16_0)|r1(esk16_0,esk42_1(esk16_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_21]), c_0_17])])).
cnf(c_0_28,plain,(p1(X1)|r1(X2,esk41_2(X1,X2))|p1(X4)|~epred1_1(X1)|~r1(X1,X2)|~p1(X3)|~r1(X3,X4)|~r1(esk42_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_29,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)|r1(esk42_1(esk16_0),esk19_1(esk42_1(esk16_0)))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_30,negated_conjecture,(p1(X1)|p1(esk19_1(X1))|~r1(esk16_0,X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_31,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)|p1(X1)|r1(X2,esk41_2(esk16_0,X2))|~p1(esk19_1(esk42_1(esk16_0)))|~r1(esk19_1(esk42_1(esk16_0)),X1)|~r1(esk16_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_21])])).
cnf(c_0_32,negated_conjecture,(p1(esk19_1(esk42_1(esk16_0)))|p1(esk42_1(esk16_0))|p1(esk16_0)), inference(spm,[status(thm)],[c_0_30, c_0_27])).
cnf(c_0_33,plain,(p1(X1)|p1(X4)|~epred1_1(X1)|~r1(X1,X2)|~p1(esk41_2(X1,X2))|~p1(X3)|~r1(X3,X4)|~r1(esk42_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_34,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)|p1(X1)|r1(X2,esk41_2(esk16_0,X2))|~r1(esk19_1(esk42_1(esk16_0)),X1)|~r1(esk16_0,X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_35,negated_conjecture,(p1(X1)|r1(esk19_1(X1),esk20_1(X1))|~r1(esk16_0,X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_36,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)|p1(X1)|~p1(esk19_1(esk42_1(esk16_0)))|~p1(esk41_2(esk16_0,X2))|~r1(esk19_1(esk42_1(esk16_0)),X1)|~r1(esk16_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_29]), c_0_21])])).
cnf(c_0_37,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)|p1(X1)|r1(esk18_0,esk41_2(esk16_0,esk18_0))|~r1(esk19_1(esk42_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_34, c_0_17])).
cnf(c_0_38,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)|r1(esk19_1(esk42_1(esk16_0)),esk20_1(esk42_1(esk16_0)))), inference(spm,[status(thm)],[c_0_35, c_0_27])).
cnf(c_0_39,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)|p1(X1)|~p1(esk41_2(esk16_0,X2))|~r1(esk19_1(esk42_1(esk16_0)),X1)|~r1(esk16_0,X2)), inference(spm,[status(thm)],[c_0_36, c_0_32])).
cnf(c_0_40,negated_conjecture,(p1(esk20_1(esk42_1(esk16_0)))|p1(esk42_1(esk16_0))|p1(esk16_0)|r1(esk18_0,esk41_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_41,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)|p1(X1)|~p1(esk41_2(esk16_0,esk18_0))|~r1(esk19_1(esk42_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_39, c_0_17])).
cnf(c_0_42,negated_conjecture,(p1(esk41_2(esk16_0,esk18_0))|p1(esk20_1(esk42_1(esk16_0)))|p1(esk42_1(esk16_0))|p1(esk16_0)), inference(spm,[status(thm)],[c_0_22, c_0_40])).
cnf(c_0_43,negated_conjecture,(p1(X1)|~r1(esk16_0,X1)|~p1(esk20_1(X1))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_44,negated_conjecture,(p1(esk20_1(esk42_1(esk16_0)))|p1(esk42_1(esk16_0))|p1(esk16_0)|p1(X1)|~r1(esk19_1(esk42_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)|~p1(esk20_1(esk42_1(esk16_0)))), inference(spm,[status(thm)],[c_0_43, c_0_27])).
cnf(c_0_46,negated_conjecture,(p1(esk20_1(esk42_1(esk16_0)))|p1(esk42_1(esk16_0))|p1(esk16_0)), inference(spm,[status(thm)],[c_0_44, c_0_38])).
cnf(c_0_47,plain,(p1(X1)|r1(X2,esk41_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)|~p1(esk42_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_48,negated_conjecture,(p1(esk42_1(esk16_0))|p1(esk16_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_49,plain,(p1(X1)|~epred1_1(X1)|~r1(X1,X2)|~p1(esk41_2(X1,X2))|~p1(esk42_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_50,plain,(p1(esk16_0)|r1(X1,esk41_2(esk16_0,X1))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_21])])).
cnf(c_0_51,plain,(p1(esk16_0)|~p1(esk41_2(esk16_0,X1))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_48]), c_0_21])])).
cnf(c_0_52,negated_conjecture,(p1(esk16_0)|r1(esk18_0,esk41_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_50, c_0_17])).
cnf(c_0_53,plain,(p1(X2)|r1(X3,esk36_2(X1,X3))|r1(esk38_1(X1),esk39_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54,negated_conjecture,(p1(esk16_0)|~p1(esk41_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_51, c_0_17])).
cnf(c_0_55,negated_conjecture,(p1(esk41_2(esk16_0,esk18_0))|p1(esk16_0)), inference(spm,[status(thm)],[c_0_22, c_0_52])).
cnf(c_0_56,negated_conjecture,(p1(X1)|r1(esk38_1(esk16_0),esk39_1(esk16_0))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~p1(esk16_0)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_17]), c_0_21])])).
cnf(c_0_57,negated_conjecture,(p1(esk16_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_58,plain,(p1(X2)|r1(X3,esk36_2(X1,X3))|r1(X1,esk38_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_59,negated_conjecture,(p1(X1)|r1(esk38_1(esk16_0),esk39_1(esk16_0))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57])])).
cnf(c_0_60,negated_conjecture,(r1(esk16_0,esk17_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_61,negated_conjecture,(~p1(esk17_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_62,negated_conjecture,(p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))|~p1(esk16_0)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_17]), c_0_21])])).
cnf(c_0_63,plain,(p1(X2)|r1(X4,esk43_2(X1,X4))|~epred1_1(X1)|~r1(X3,X2)|~r1(X4,X3)|~r1(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_64,negated_conjecture,(r1(esk38_1(esk16_0),esk39_1(esk16_0))|r1(esk18_0,esk36_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])).
cnf(c_0_65,plain,(p1(X2)|r1(esk36_2(X1,X3),esk37_2(X1,X3))|r1(esk38_1(X1),esk39_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_66,negated_conjecture,(p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_57])])).
cnf(c_0_67,plain,(p1(X1)|r1(esk38_1(esk16_0),esk39_1(esk16_0))|r1(esk18_0,esk43_2(X2,esk18_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68,negated_conjecture,(p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk38_1(esk16_0),esk39_1(esk16_0))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_17]), c_0_21]), c_0_57])])).
cnf(c_0_69,negated_conjecture,(r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_60]), c_0_61])).
cnf(c_0_70,plain,(p1(X2)|r1(esk36_2(X1,X3),esk37_2(X1,X3))|r1(X1,esk38_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_71,negated_conjecture,(p1(X1)|r1(esk38_1(esk16_0),esk39_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_17]), c_0_21])])).
cnf(c_0_72,negated_conjecture,(r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk38_1(esk16_0),esk39_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_60]), c_0_61])).
cnf(c_0_73,plain,(p1(X1)|r1(esk18_0,esk43_2(X2,esk18_0))|r1(esk16_0,esk38_1(esk16_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_69])).
cnf(c_0_74,negated_conjecture,(p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_17]), c_0_21]), c_0_57])])).
cnf(c_0_75,plain,(p1(X2)|r1(esk38_1(X1),esk39_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk37_2(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_76,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk38_1(esk16_0),esk39_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_77,negated_conjecture,(p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_17]), c_0_21])])).
cnf(c_0_78,negated_conjecture,(r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_60]), c_0_61])).
cnf(c_0_79,plain,(p1(X1)|r1(esk38_1(esk16_0),esk39_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_21]), c_0_57]), c_0_17])])).
cnf(c_0_80,plain,(p1(X2)|r1(X1,esk38_1(X1))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk37_2(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_81,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_82,negated_conjecture,(r1(esk38_1(esk16_0),esk39_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_60]), c_0_61])).
cnf(c_0_83,plain,(p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_21]), c_0_57]), c_0_17])])).
cnf(c_0_84,plain,(p1(X2)|~epred1_1(X1)|~r1(X3,X2)|~r1(X4,X3)|~r1(X1,X4)|~p1(esk43_2(X1,X4))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_85,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|r1(esk38_1(esk16_0),esk39_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_82])).
cnf(c_0_86,negated_conjecture,(r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_60]), c_0_61])).
cnf(c_0_87,plain,(p1(X1)|r1(esk38_1(esk16_0),esk39_1(esk16_0))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_21]), c_0_17])])).
cnf(c_0_88,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_86])).
cnf(c_0_89,negated_conjecture,(p1(X1)|r1(esk38_1(esk16_0),esk39_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_87, c_0_64])).
cnf(c_0_90,plain,(p1(X1)|r1(esk16_0,esk38_1(esk16_0))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_88]), c_0_21]), c_0_17])])).
cnf(c_0_91,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk38_1(esk16_0),esk39_1(esk16_0))), inference(spm,[status(thm)],[c_0_89, c_0_72])).
cnf(c_0_92,negated_conjecture,(p1(X1)|r1(esk16_0,esk38_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_90, c_0_69])).
cnf(c_0_93,plain,(p1(X1)|r1(esk38_1(esk16_0),esk39_1(esk16_0))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_91]), c_0_21]), c_0_57]), c_0_17])])).
cnf(c_0_94,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk38_1(esk16_0))), inference(spm,[status(thm)],[c_0_92, c_0_78])).
cnf(c_0_95,plain,(p1(X3)|r1(X3,esk44_3(X1,X2,X3))|r1(X1,esk46_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_96,negated_conjecture,(r1(esk38_1(esk16_0),esk39_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_60]), c_0_61])).
cnf(c_0_97,plain,(p1(X1)|r1(esk16_0,esk38_1(esk16_0))|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_94]), c_0_21]), c_0_57]), c_0_17])])).
cnf(c_0_98,plain,(p1(esk39_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(X1,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(X1,esk46_1(X1))|~epred1_1(X1)|~r1(X1,esk38_1(esk16_0))), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_99,negated_conjecture,(r1(esk16_0,esk38_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_60]), c_0_61])).
cnf(c_0_100,plain,(p1(X2)|r1(X3,esk36_2(X1,X3))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk39_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_101,negated_conjecture,(p1(esk39_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_21])])).
cnf(c_0_102,plain,(p1(X3)|r1(esk44_3(X1,X2,X3),esk45_3(X1,X2,X3))|r1(X1,esk46_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_103,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(X2,esk36_2(esk16_0,X2))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_21]), c_0_57])])).
cnf(c_0_104,negated_conjecture,(p1(esk39_1(esk16_0))|r1(esk44_3(X1,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(X1,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(X1,esk46_1(X1))|~epred1_1(X1)|~r1(X1,esk38_1(esk16_0))), inference(spm,[status(thm)],[c_0_102, c_0_96])).
cnf(c_0_105,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_103, c_0_17])).
cnf(c_0_106,plain,(p1(X2)|r1(esk36_2(X1,X3),esk37_2(X1,X3))|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk39_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_107,negated_conjecture,(p1(esk39_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_99]), c_0_21])])).
cnf(c_0_108,negated_conjecture,(r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_60]), c_0_61])).
cnf(c_0_109,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,X2),esk37_2(esk16_0,X2))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_101]), c_0_21]), c_0_57])])).
cnf(c_0_110,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(X2,esk36_2(esk16_0,X2))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_107]), c_0_21]), c_0_57])])).
cnf(c_0_111,plain,(p1(X2)|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk37_2(X1,X3))|~p1(esk39_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_112,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(X2,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_108])).
cnf(c_0_113,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_109, c_0_17])).
cnf(c_0_114,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_110, c_0_17])).
cnf(c_0_115,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~p1(esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_101]), c_0_21]), c_0_57])])).
cnf(c_0_116,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_17]), c_0_21])])).
cnf(c_0_117,negated_conjecture,(r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_60]), c_0_61])).
cnf(c_0_118,negated_conjecture,(r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_60]), c_0_61])).
cnf(c_0_119,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,X2),esk37_2(esk16_0,X2))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_21]), c_0_57])])).
cnf(c_0_120,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~p1(esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_115, c_0_17])).
cnf(c_0_121,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_122,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(X2,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_118])).
cnf(c_0_123,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_119, c_0_17])).
cnf(c_0_124,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_125,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~p1(esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_107]), c_0_21]), c_0_57])])).
cnf(c_0_126,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_17]), c_0_21])])).
cnf(c_0_127,negated_conjecture,(r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_60]), c_0_61])).
cnf(c_0_128,negated_conjecture,(r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_60]), c_0_61])).
cnf(c_0_129,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~p1(esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_125, c_0_17])).
cnf(c_0_130,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_131,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_128])).
cnf(c_0_132,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_133,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_131]), c_0_21]), c_0_17])])).
cnf(c_0_134,negated_conjecture,(r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_60]), c_0_61])).
cnf(c_0_135,plain,(p1(X2)|r1(X3,esk36_2(X1,X3))|p1(X5)|p1(X6)|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(X4)|~r1(X4,X5)|~r1(X6,X4)|~r1(esk38_1(X1),X6)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_136,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_133, c_0_108])).
cnf(c_0_137,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_134])).
cnf(c_0_138,plain,(p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|r1(X3,esk36_2(esk16_0,X3))|~p1(X4)|~r1(esk39_1(esk16_0),X4)|~r1(esk16_0,X3)|~r1(esk16_0,X2)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_96]), c_0_21]), c_0_57])])).
cnf(c_0_139,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_136, c_0_117])).
cnf(c_0_140,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_137]), c_0_21]), c_0_17])])).
cnf(c_0_141,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~p1(X3)|~r1(esk39_1(esk16_0),X3)|~r1(esk16_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_138, c_0_17])).
cnf(c_0_142,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_120, c_0_139])).
cnf(c_0_143,plain,(p1(X3)|p1(esk44_3(X1,X2,X3))|r1(X1,esk46_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_144,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_140, c_0_118])).
cnf(c_0_145,plain,(p1(X2)|r1(esk36_2(X1,X3),esk37_2(X1,X3))|p1(X5)|p1(X6)|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(X4)|~r1(X4,X5)|~r1(X6,X4)|~r1(esk38_1(X1),X6)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_146,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~p1(X2)|~r1(esk39_1(esk16_0),X2)|~r1(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_60]), c_0_61])).
cnf(c_0_147,negated_conjecture,(r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_60]), c_0_61])).
cnf(c_0_148,plain,(p1(esk44_3(X1,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|r1(X1,esk46_1(X1))|~epred1_1(X1)|~r1(X1,esk38_1(esk16_0))), inference(spm,[status(thm)],[c_0_143, c_0_96])).
cnf(c_0_149,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_144, c_0_127])).
cnf(c_0_150,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|r1(esk36_2(esk16_0,X3),esk37_2(esk16_0,X3))|~p1(X4)|~r1(esk39_1(esk16_0),X4)|~r1(esk16_0,X3)|~r1(esk16_0,X2)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_96]), c_0_21]), c_0_57])])).
cnf(c_0_151,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_152,negated_conjecture,(p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|r1(esk16_0,esk46_1(esk16_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_99]), c_0_21])])).
cnf(c_0_153,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_129, c_0_149])).
cnf(c_0_154,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~p1(X3)|~r1(esk39_1(esk16_0),X3)|~r1(esk16_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_150, c_0_17])).
cnf(c_0_155,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_156,negated_conjecture,(r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_60]), c_0_61])).
cnf(c_0_157,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~p1(X2)|~r1(esk39_1(esk16_0),X2)|~r1(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_60]), c_0_61])).
cnf(c_0_158,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_159,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_157, c_0_147])).
cnf(c_0_160,plain,(p1(X2)|p1(X5)|p1(X6)|~epred1_1(X1)|~p1(X1)|~r1(X1,X2)|~r1(X1,X3)|~p1(esk37_2(X1,X3))|~p1(X4)|~r1(X4,X5)|~r1(X6,X4)|~r1(esk38_1(X1),X6)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_161,plain,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(X2,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_158])).
cnf(c_0_162,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_159, c_0_152])).
cnf(c_0_163,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|~p1(esk37_2(esk16_0,X3))|~p1(X4)|~r1(esk39_1(esk16_0),X4)|~r1(esk16_0,X3)|~r1(esk16_0,X2)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_96]), c_0_21]), c_0_57])])).
cnf(c_0_164,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_17]), c_0_21])])).
cnf(c_0_165,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_162, c_0_156])).
cnf(c_0_166,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|~p1(esk37_2(esk16_0,esk18_0))|~p1(X3)|~r1(esk39_1(esk16_0),X3)|~r1(esk16_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_163, c_0_17])).
cnf(c_0_167,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk39_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_168,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~p1(X3)|~r1(esk39_1(esk16_0),X3)|~r1(esk16_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_166, c_0_167])).
cnf(c_0_169,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~p1(X2)|~r1(esk39_1(esk16_0),X2)|~r1(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_60]), c_0_61])).
cnf(c_0_170,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_169, c_0_147])).
cnf(c_0_171,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_170, c_0_152])).
cnf(c_0_172,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_171, c_0_156])).
cnf(c_0_173,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk43_2(esk16_0,esk18_0))|p1(esk39_1(esk16_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_172])).
cnf(c_0_174,plain,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_173]), c_0_21]), c_0_17])])).
cnf(c_0_175,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_174, c_0_158])).
cnf(c_0_176,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk39_1(esk16_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_175, c_0_165])).
cnf(c_0_177,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|r1(esk16_0,esk46_1(esk16_0))|~p1(X3)|~r1(esk39_1(esk16_0),X3)|~r1(esk16_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_166, c_0_176])).
cnf(c_0_178,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~p1(X2)|~r1(esk39_1(esk16_0),X2)|~r1(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_60]), c_0_61])).
cnf(c_0_179,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_178, c_0_147])).
cnf(c_0_180,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_179, c_0_152])).
cnf(c_0_181,plain,(p1(X3)|r1(X1,esk46_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk45_3(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_182,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_180, c_0_156])).
cnf(c_0_183,plain,(p1(esk39_1(esk16_0))|r1(esk16_0,esk46_1(esk16_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182]), c_0_21]), c_0_96]), c_0_99])])).
cnf(c_0_184,plain,(p1(X1)|r1(X2,esk36_2(esk16_0,X2))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_183]), c_0_21]), c_0_57])])).
cnf(c_0_185,negated_conjecture,(p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_184, c_0_17])).
cnf(c_0_186,negated_conjecture,(r1(esk18_0,esk36_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_60]), c_0_61])).
cnf(c_0_187,plain,(p1(X1)|r1(esk36_2(esk16_0,X2),esk37_2(esk16_0,X2))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_183]), c_0_21]), c_0_57])])).
cnf(c_0_188,plain,(p1(X1)|r1(esk18_0,esk43_2(X2,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_186])).
cnf(c_0_189,negated_conjecture,(p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_187, c_0_17])).
cnf(c_0_190,plain,(p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~p1(esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_183]), c_0_21]), c_0_57])])).
cnf(c_0_191,negated_conjecture,(p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_17]), c_0_21])])).
cnf(c_0_192,negated_conjecture,(r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_60]), c_0_61])).
cnf(c_0_193,negated_conjecture,(p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~p1(esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_190, c_0_17])).
cnf(c_0_194,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_191, c_0_192])).
cnf(c_0_195,negated_conjecture,(p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_193, c_0_194])).
cnf(c_0_196,negated_conjecture,(r1(esk18_0,esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_60]), c_0_61])).
cnf(c_0_197,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_196])).
cnf(c_0_198,plain,(p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_197]), c_0_21]), c_0_17])])).
cnf(c_0_199,negated_conjecture,(p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_198, c_0_186])).
cnf(c_0_200,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk16_0,esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_199, c_0_192])).
cnf(c_0_201,negated_conjecture,(p1(X1)|r1(esk16_0,esk46_1(esk16_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_193, c_0_200])).
cnf(c_0_202,negated_conjecture,(r1(esk16_0,esk46_1(esk16_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_60]), c_0_61])).
cnf(c_0_203,plain,(p1(X3)|r1(X3,esk44_3(X1,X2,X3))|p1(X5)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(X4)|~r1(X4,X5)|~r1(esk46_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_204,negated_conjecture,(p1(esk46_1(esk16_0))|r1(esk46_1(esk16_0),esk19_1(esk46_1(esk16_0)))), inference(spm,[status(thm)],[c_0_26, c_0_202])).
cnf(c_0_205,plain,(p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|r1(X2,esk44_3(esk16_0,X3,X2))|~p1(esk19_1(esk46_1(esk16_0)))|~r1(esk19_1(esk46_1(esk16_0)),X1)|~r1(esk16_0,X3)|~r1(X3,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_204]), c_0_21])])).
cnf(c_0_206,negated_conjecture,(p1(esk19_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_30, c_0_202])).
cnf(c_0_207,negated_conjecture,(p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|r1(X2,esk44_3(esk16_0,X3,X2))|~r1(esk19_1(esk46_1(esk16_0)),X1)|~r1(esk16_0,X3)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_205, c_0_206])).
cnf(c_0_208,plain,(p1(X3)|r1(esk44_3(X1,X2,X3),esk45_3(X1,X2,X3))|p1(X5)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(X4)|~r1(X4,X5)|~r1(esk46_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_209,negated_conjecture,(p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|r1(X2,esk44_3(esk16_0,esk38_1(esk16_0),X2))|~r1(esk19_1(esk46_1(esk16_0)),X1)|~r1(esk38_1(esk16_0),X2)), inference(spm,[status(thm)],[c_0_207, c_0_99])).
cnf(c_0_210,plain,(p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|r1(esk44_3(esk16_0,X3,X2),esk45_3(esk16_0,X3,X2))|~p1(esk19_1(esk46_1(esk16_0)))|~r1(esk19_1(esk46_1(esk16_0)),X1)|~r1(esk16_0,X3)|~r1(X3,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_204]), c_0_21])])).
cnf(c_0_211,negated_conjecture,(p1(esk46_1(esk16_0))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk19_1(esk46_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_209, c_0_96])).
cnf(c_0_212,negated_conjecture,(p1(esk46_1(esk16_0))|r1(esk19_1(esk46_1(esk16_0)),esk20_1(esk46_1(esk16_0)))), inference(spm,[status(thm)],[c_0_35, c_0_202])).
cnf(c_0_213,negated_conjecture,(p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|r1(esk44_3(esk16_0,X3,X2),esk45_3(esk16_0,X3,X2))|~r1(esk19_1(esk46_1(esk16_0)),X1)|~r1(esk16_0,X3)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_210, c_0_206])).
cnf(c_0_214,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(esk39_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_211, c_0_212])).
cnf(c_0_215,negated_conjecture,(p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|r1(esk44_3(esk16_0,esk38_1(esk16_0),X2),esk45_3(esk16_0,esk38_1(esk16_0),X2))|~r1(esk19_1(esk46_1(esk16_0)),X1)|~r1(esk38_1(esk16_0),X2)), inference(spm,[status(thm)],[c_0_213, c_0_99])).
cnf(c_0_216,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(X2,esk36_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_214]), c_0_21]), c_0_57])])).
cnf(c_0_217,negated_conjecture,(p1(esk46_1(esk16_0))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk19_1(esk46_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_215, c_0_96])).
cnf(c_0_218,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_216, c_0_17])).
cnf(c_0_219,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(esk39_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_217, c_0_212])).
cnf(c_0_220,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_218, c_0_202])).
cnf(c_0_221,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,X2),esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_214]), c_0_21]), c_0_57])])).
cnf(c_0_222,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(X2,esk36_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_219]), c_0_21]), c_0_57])])).
cnf(c_0_223,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(X2,esk18_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_220])).
cnf(c_0_224,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_221, c_0_17])).
cnf(c_0_225,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_222, c_0_17])).
cnf(c_0_226,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~p1(esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_214]), c_0_21]), c_0_57])])).
cnf(c_0_227,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_17]), c_0_21])])).
cnf(c_0_228,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_224, c_0_202])).
cnf(c_0_229,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_225, c_0_202])).
cnf(c_0_230,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,X2),esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_219]), c_0_21]), c_0_57])])).
cnf(c_0_231,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~p1(esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_226, c_0_17])).
cnf(c_0_232,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk46_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_227, c_0_228])).
cnf(c_0_233,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(X2,esk18_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_229])).
cnf(c_0_234,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_230, c_0_17])).
cnf(c_0_235,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_231, c_0_232])).
cnf(c_0_236,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~p1(esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_219]), c_0_21]), c_0_57])])).
cnf(c_0_237,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_17]), c_0_21])])).
cnf(c_0_238,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_234, c_0_202])).
cnf(c_0_239,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_235, c_0_202])).
cnf(c_0_240,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~p1(esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_236, c_0_17])).
cnf(c_0_241,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk46_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_237, c_0_238])).
cnf(c_0_242,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_22, c_0_239])).
cnf(c_0_243,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_240, c_0_241])).
cnf(c_0_244,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_242]), c_0_21]), c_0_17])])).
cnf(c_0_245,plain,(p1(X3)|p1(esk44_3(X1,X2,X3))|p1(X5)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(X4)|~r1(X4,X5)|~r1(esk46_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_246,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_243, c_0_202])).
cnf(c_0_247,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_244, c_0_220])).
cnf(c_0_248,plain,(p1(esk44_3(esk16_0,X1,X2))|p1(esk46_1(esk16_0))|p1(X3)|p1(X2)|~p1(esk19_1(esk46_1(esk16_0)))|~r1(esk19_1(esk46_1(esk16_0)),X3)|~r1(esk16_0,X1)|~r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_204]), c_0_21])])).
cnf(c_0_249,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_22, c_0_246])).
cnf(c_0_250,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk46_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_247, c_0_228])).
cnf(c_0_251,negated_conjecture,(p1(esk44_3(esk16_0,X1,X2))|p1(esk46_1(esk16_0))|p1(X3)|p1(X2)|~r1(esk19_1(esk46_1(esk16_0)),X3)|~r1(esk16_0,X1)|~r1(X1,X2)), inference(spm,[status(thm)],[c_0_248, c_0_206])).
cnf(c_0_252,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_249]), c_0_21]), c_0_17])])).
cnf(c_0_253,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_231, c_0_250])).
cnf(c_0_254,negated_conjecture,(p1(esk44_3(esk16_0,esk38_1(esk16_0),X1))|p1(esk46_1(esk16_0))|p1(X2)|p1(X1)|~r1(esk19_1(esk46_1(esk16_0)),X2)|~r1(esk38_1(esk16_0),X1)), inference(spm,[status(thm)],[c_0_251, c_0_99])).
cnf(c_0_255,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_252, c_0_229])).
cnf(c_0_256,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_253, c_0_202])).
cnf(c_0_257,negated_conjecture,(p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(esk39_1(esk16_0))|p1(X1)|~r1(esk19_1(esk46_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_254, c_0_96])).
cnf(c_0_258,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk46_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_255, c_0_238])).
cnf(c_0_259,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_146, c_0_256])).
cnf(c_0_260,negated_conjecture,(p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(esk39_1(esk16_0))), inference(spm,[status(thm)],[c_0_257, c_0_212])).
cnf(c_0_261,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_240, c_0_258])).
cnf(c_0_262,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_259, c_0_260])).
cnf(c_0_263,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_261, c_0_202])).
cnf(c_0_264,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|r1(esk18_0,esk36_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_262, c_0_263])).
cnf(c_0_265,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_157, c_0_256])).
cnf(c_0_266,plain,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(X2,esk18_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_264])).
cnf(c_0_267,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_265, c_0_260])).
cnf(c_0_268,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_17]), c_0_21])])).
cnf(c_0_269,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_267, c_0_263])).
cnf(c_0_270,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_268, c_0_269])).
cnf(c_0_271,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~p1(X3)|~r1(esk39_1(esk16_0),X3)|~r1(esk16_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_166, c_0_270])).
cnf(c_0_272,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~p1(X2)|~r1(esk39_1(esk16_0),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_271, c_0_202])).
cnf(c_0_273,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_272, c_0_256])).
cnf(c_0_274,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_273, c_0_260])).
cnf(c_0_275,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_274, c_0_263])).
cnf(c_0_276,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk43_2(esk16_0,esk18_0))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_275])).
cnf(c_0_277,plain,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_276]), c_0_21]), c_0_17])])).
cnf(c_0_278,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_277, c_0_264])).
cnf(c_0_279,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_278, c_0_269])).
cnf(c_0_280,plain,(p1(X3)|p1(X5)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk45_3(X1,X2,X3))|~p1(X4)|~r1(X4,X5)|~r1(esk46_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_281,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|~p1(X3)|~r1(esk39_1(esk16_0),X3)|~r1(esk16_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_166, c_0_279])).
cnf(c_0_282,plain,(p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|~p1(esk45_3(esk16_0,X3,X2))|~p1(esk19_1(esk46_1(esk16_0)))|~r1(esk19_1(esk46_1(esk16_0)),X1)|~r1(esk16_0,X3)|~r1(X3,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_204]), c_0_21])])).
cnf(c_0_283,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|~p1(X2)|~r1(esk39_1(esk16_0),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_281, c_0_202])).
cnf(c_0_284,negated_conjecture,(p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|~p1(esk45_3(esk16_0,X3,X2))|~r1(esk19_1(esk46_1(esk16_0)),X1)|~r1(esk16_0,X3)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_282, c_0_206])).
cnf(c_0_285,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_283, c_0_256])).
cnf(c_0_286,negated_conjecture,(p1(esk46_1(esk16_0))|p1(X1)|p1(X2)|~p1(esk45_3(esk16_0,esk38_1(esk16_0),X2))|~r1(esk19_1(esk46_1(esk16_0)),X1)|~r1(esk38_1(esk16_0),X2)), inference(spm,[status(thm)],[c_0_284, c_0_99])).
cnf(c_0_287,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))|p1(X1)|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_285, c_0_260])).
cnf(c_0_288,negated_conjecture,(p1(esk46_1(esk16_0))|p1(esk39_1(esk16_0))|p1(X1)|~p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk19_1(esk46_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_286, c_0_96])).
cnf(c_0_289,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_287, c_0_263])).
cnf(c_0_290,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(esk39_1(esk16_0))|p1(X1)|~r1(esk19_1(esk46_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_288, c_0_289])).
cnf(c_0_291,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(esk39_1(esk16_0))), inference(spm,[status(thm)],[c_0_290, c_0_212])).
cnf(c_0_292,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(X2,esk36_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_291]), c_0_21]), c_0_57])])).
cnf(c_0_293,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_292, c_0_17])).
cnf(c_0_294,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk18_0,esk36_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_293, c_0_202])).
cnf(c_0_295,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk36_2(esk16_0,X2),esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_291]), c_0_21]), c_0_57])])).
cnf(c_0_296,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(X2,esk18_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_294])).
cnf(c_0_297,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_295, c_0_17])).
cnf(c_0_298,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|~p1(esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_291]), c_0_21]), c_0_57])])).
cnf(c_0_299,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_17]), c_0_21])])).
cnf(c_0_300,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_297, c_0_202])).
cnf(c_0_301,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|~p1(esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_298, c_0_17])).
cnf(c_0_302,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk46_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_299, c_0_300])).
cnf(c_0_303,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_301, c_0_302])).
cnf(c_0_304,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_303, c_0_202])).
cnf(c_0_305,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_304])).
cnf(c_0_306,plain,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_305]), c_0_21]), c_0_17])])).
cnf(c_0_307,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_306, c_0_294])).
cnf(c_0_308,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_307, c_0_300])).
cnf(c_0_309,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))|p1(X1)|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_301, c_0_308])).
cnf(c_0_310,negated_conjecture,(p1(esk46_1(esk16_0))|~p1(esk20_1(esk46_1(esk16_0)))), inference(spm,[status(thm)],[c_0_43, c_0_202])).
cnf(c_0_311,negated_conjecture,(p1(esk20_1(esk46_1(esk16_0)))|p1(esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_309, c_0_202])).
cnf(c_0_312,plain,(p1(X3)|r1(X3,esk44_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk46_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_313,negated_conjecture,(p1(esk46_1(esk16_0))), inference(spm,[status(thm)],[c_0_310, c_0_311])).
cnf(c_0_314,plain,(p1(X1)|r1(X1,esk44_3(esk16_0,X2,X1))|~r1(esk16_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_313]), c_0_21])])).
cnf(c_0_315,negated_conjecture,(p1(X1)|r1(X1,esk44_3(esk16_0,esk38_1(esk16_0),X1))|~r1(esk38_1(esk16_0),X1)), inference(spm,[status(thm)],[c_0_314, c_0_99])).
cnf(c_0_316,plain,(p1(X3)|r1(esk44_3(X1,X2,X3),esk45_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk46_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_317,negated_conjecture,(p1(esk39_1(esk16_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_315, c_0_96])).
cnf(c_0_318,plain,(p1(X1)|r1(esk44_3(esk16_0,X2,X1),esk45_3(esk16_0,X2,X1))|~r1(esk16_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_313]), c_0_21])])).
cnf(c_0_319,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(X2,esk36_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_317]), c_0_21]), c_0_57])])).
cnf(c_0_320,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),X1),esk45_3(esk16_0,esk38_1(esk16_0),X1))|~r1(esk38_1(esk16_0),X1)), inference(spm,[status(thm)],[c_0_318, c_0_99])).
cnf(c_0_321,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_319, c_0_17])).
cnf(c_0_322,negated_conjecture,(p1(esk39_1(esk16_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_320, c_0_96])).
cnf(c_0_323,negated_conjecture,(r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_60]), c_0_61])).
cnf(c_0_324,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,X2),esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_317]), c_0_21]), c_0_57])])).
cnf(c_0_325,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(X2,esk36_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_322]), c_0_21]), c_0_57])])).
cnf(c_0_326,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(X2,esk18_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_323])).
cnf(c_0_327,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_324, c_0_17])).
cnf(c_0_328,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_325, c_0_17])).
cnf(c_0_329,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~p1(esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_317]), c_0_21]), c_0_57])])).
cnf(c_0_330,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_17]), c_0_21])])).
cnf(c_0_331,negated_conjecture,(r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_60]), c_0_61])).
cnf(c_0_332,negated_conjecture,(r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk36_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_60]), c_0_61])).
cnf(c_0_333,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,X2),esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_322]), c_0_21]), c_0_57])])).
cnf(c_0_334,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~p1(esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_329, c_0_17])).
cnf(c_0_335,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_330, c_0_331])).
cnf(c_0_336,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(X2,esk18_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_332])).
cnf(c_0_337,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_333, c_0_17])).
cnf(c_0_338,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_334, c_0_335])).
cnf(c_0_339,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~p1(esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_322]), c_0_21]), c_0_57])])).
cnf(c_0_340,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_17]), c_0_21])])).
cnf(c_0_341,negated_conjecture,(r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_60]), c_0_61])).
cnf(c_0_342,negated_conjecture,(r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_60]), c_0_61])).
cnf(c_0_343,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~p1(esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_339, c_0_17])).
cnf(c_0_344,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_340, c_0_341])).
cnf(c_0_345,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_22, c_0_342])).
cnf(c_0_346,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_343, c_0_344])).
cnf(c_0_347,plain,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_345]), c_0_21]), c_0_17])])).
cnf(c_0_348,negated_conjecture,(r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_60]), c_0_61])).
cnf(c_0_349,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_347, c_0_323])).
cnf(c_0_350,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_22, c_0_348])).
cnf(c_0_351,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_349, c_0_331])).
cnf(c_0_352,plain,(p1(X3)|p1(esk44_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk46_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_353,plain,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_350]), c_0_21]), c_0_17])])).
cnf(c_0_354,negated_conjecture,(p1(X1)|r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_334, c_0_351])).
cnf(c_0_355,plain,(p1(esk44_3(esk16_0,X1,X2))|p1(X2)|~r1(esk16_0,X1)|~r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_313]), c_0_21])])).
cnf(c_0_356,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_353, c_0_332])).
cnf(c_0_357,negated_conjecture,(r1(esk39_1(esk16_0),esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_60]), c_0_61])).
cnf(c_0_358,negated_conjecture,(p1(esk44_3(esk16_0,esk38_1(esk16_0),X1))|p1(X1)|~r1(esk38_1(esk16_0),X1)), inference(spm,[status(thm)],[c_0_355, c_0_99])).
cnf(c_0_359,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_356, c_0_341])).
cnf(c_0_360,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_146, c_0_357])).
cnf(c_0_361,negated_conjecture,(p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))), inference(spm,[status(thm)],[c_0_358, c_0_96])).
cnf(c_0_362,negated_conjecture,(p1(X1)|r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_343, c_0_359])).
cnf(c_0_363,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_360, c_0_361])).
cnf(c_0_364,negated_conjecture,(r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_362, c_0_60]), c_0_61])).
cnf(c_0_365,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|r1(esk18_0,esk36_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_363, c_0_364])).
cnf(c_0_366,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_157, c_0_357])).
cnf(c_0_367,plain,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(X2,esk18_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_365])).
cnf(c_0_368,negated_conjecture,(p1(esk39_1(esk16_0))|p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_366, c_0_361])).
cnf(c_0_369,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_17]), c_0_21])])).
cnf(c_0_370,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_368, c_0_364])).
cnf(c_0_371,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk39_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_369, c_0_370])).
cnf(c_0_372,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~p1(X3)|~r1(esk39_1(esk16_0),X3)|~r1(esk16_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_166, c_0_371])).
cnf(c_0_373,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~p1(X2)|~r1(esk39_1(esk16_0),X2)|~r1(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_60]), c_0_61])).
cnf(c_0_374,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_373, c_0_357])).
cnf(c_0_375,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_374, c_0_361])).
cnf(c_0_376,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_375, c_0_364])).
cnf(c_0_377,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk43_2(esk16_0,esk18_0))|p1(esk39_1(esk16_0))), inference(spm,[status(thm)],[c_0_22, c_0_376])).
cnf(c_0_378,plain,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_377]), c_0_21]), c_0_17])])).
cnf(c_0_379,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_378, c_0_365])).
cnf(c_0_380,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk37_2(esk16_0,esk18_0))|p1(esk39_1(esk16_0))), inference(spm,[status(thm)],[c_0_379, c_0_370])).
cnf(c_0_381,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|p1(X2)|~p1(X3)|~r1(esk39_1(esk16_0),X3)|~r1(esk16_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_166, c_0_380])).
cnf(c_0_382,plain,(p1(X3)|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p1(esk45_3(X1,X2,X3))|~p1(esk46_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_383,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|~p1(X2)|~r1(esk39_1(esk16_0),X2)|~r1(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_60]), c_0_61])).
cnf(c_0_384,plain,(p1(X1)|~p1(esk45_3(esk16_0,X2,X1))|~r1(esk16_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_313]), c_0_21])])).
cnf(c_0_385,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|~p1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_383, c_0_357])).
cnf(c_0_386,negated_conjecture,(p1(X1)|~p1(esk45_3(esk16_0,esk38_1(esk16_0),X1))|~r1(esk38_1(esk16_0),X1)), inference(spm,[status(thm)],[c_0_384, c_0_99])).
cnf(c_0_387,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))|p1(X1)|~r1(esk44_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)),X1)), inference(spm,[status(thm)],[c_0_385, c_0_361])).
cnf(c_0_388,negated_conjecture,(p1(esk39_1(esk16_0))|~p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))), inference(spm,[status(thm)],[c_0_386, c_0_96])).
cnf(c_0_389,negated_conjecture,(p1(esk45_3(esk16_0,esk38_1(esk16_0),esk39_1(esk16_0)))|p1(esk39_1(esk16_0))), inference(spm,[status(thm)],[c_0_387, c_0_364])).
cnf(c_0_390,negated_conjecture,(p1(esk39_1(esk16_0))), inference(spm,[status(thm)],[c_0_388, c_0_389])).
cnf(c_0_391,plain,(p1(X1)|r1(X2,esk36_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_390]), c_0_21]), c_0_57])])).
cnf(c_0_392,negated_conjecture,(p1(X1)|r1(esk18_0,esk36_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_391, c_0_17])).
cnf(c_0_393,negated_conjecture,(r1(esk18_0,esk36_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_60]), c_0_61])).
cnf(c_0_394,plain,(p1(X1)|r1(esk36_2(esk16_0,X2),esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_390]), c_0_21]), c_0_57])])).
cnf(c_0_395,plain,(p1(X1)|r1(esk18_0,esk43_2(X2,esk18_0))|~epred1_1(X2)|~r1(esk36_2(esk16_0,esk18_0),X1)|~r1(X2,esk18_0)), inference(spm,[status(thm)],[c_0_63, c_0_393])).
cnf(c_0_396,negated_conjecture,(p1(X1)|r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_394, c_0_17])).
cnf(c_0_397,plain,(p1(X1)|~p1(esk37_2(esk16_0,X2))|~r1(esk16_0,X2)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_390]), c_0_21]), c_0_57])])).
cnf(c_0_398,negated_conjecture,(p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_17]), c_0_21])])).
cnf(c_0_399,negated_conjecture,(r1(esk36_2(esk16_0,esk18_0),esk37_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_60]), c_0_61])).
cnf(c_0_400,negated_conjecture,(p1(X1)|~p1(esk37_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_397, c_0_17])).
cnf(c_0_401,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))|r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_398, c_0_399])).
cnf(c_0_402,negated_conjecture,(p1(X1)|r1(esk18_0,esk43_2(esk16_0,esk18_0))|~r1(esk16_0,X1)), inference(spm,[status(thm)],[c_0_400, c_0_401])).
cnf(c_0_403,negated_conjecture,(r1(esk18_0,esk43_2(esk16_0,esk18_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_60]), c_0_61])).
cnf(c_0_404,negated_conjecture,(p1(esk43_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_22, c_0_403])).
cnf(c_0_405,plain,(p1(X1)|~r1(esk18_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_404]), c_0_21]), c_0_17])])).
cnf(c_0_406,negated_conjecture,(p1(X1)|~r1(esk36_2(esk16_0,esk18_0),X1)), inference(spm,[status(thm)],[c_0_405, c_0_393])).
cnf(c_0_407,negated_conjecture,(p1(esk37_2(esk16_0,esk18_0))), inference(spm,[status(thm)],[c_0_406, c_0_399])).
cnf(c_0_408,negated_conjecture,(p1(X1)|~r1(esk16_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_400, c_0_407])])).
cnf(c_0_409,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_60]), c_0_61]), ['proof']).
# SZS output end CNFRefutation.
EOF
