% Problem    : CSR115+81 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n127.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:44:11 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:(((((((((agt(X5,X4)&attr(X1,X2))&attr(X4,X3))&attr(X6,X7))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&subs(X5,n374bernehmen_1_1))&val(X2,bmw_0))&val(X3,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_wp_491)).
fof(c_0_1, hypothesis, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((pred(c12,getriebe__1_1)&pred(c16,rotornabe_1_1))&sub(c4,konstruktion_1_1))&agt(c45,c57))&modl(c45,sollen_0))&obj(c45,c62))&subs(c45,n374bernehmen_1_1))&attr(c57,c58))&sub(c57,firma_1_1))&sub(c58,name_1_1))&val(c58,bmw_0))&itms_p4(c62,c4,c12,c16))&attch(c9,c4))&subs(c9,motoranlage_1_1))&assoc(motoranlage_1_1,motor__1_1))&subs(motoranlage_1_1,anlage_1_1))&assoc(rotornabe_1_1,rotor_1_1))&sub(rotornabe_1_1,nabe_1_1))&sort(c12,d))&card(c12,cons(x_constant,cons(int1,nil))))&etype(c12,int1))&fact(c12,real))&gener(c12,gener_c))&quant(c12,mult))&refer(c12,indet))&varia(c12,varia_c))&sort(getriebe__1_1,d))&card(getriebe__1_1,int1))&etype(getriebe__1_1,int0))&fact(getriebe__1_1,real))&gener(getriebe__1_1,ge))&quant(getriebe__1_1,one))&refer(getriebe__1_1,refer_c))&varia(getriebe__1_1,varia_c))&sort(c16,o))&card(c16,cons(x_constant,cons(int1,nil))))&etype(c16,int1))&fact(c16,real))&gener(c16,gener_c))&quant(c16,mult))&refer(c16,indet))&varia(c16,varia_c))&sort(rotornabe_1_1,o))&card(rotornabe_1_1,int1))&etype(rotornabe_1_1,int0))&fact(rotornabe_1_1,real))&gener(rotornabe_1_1,ge))&quant(rotornabe_1_1,one))&refer(rotornabe_1_1,refer_c))&varia(rotornabe_1_1,varia_c))&sort(c4,d))&card(c4,int1))&etype(c4,int0))&fact(c4,real))&gener(c4,sp))&quant(c4,one))&refer(c4,det))&varia(c4,con))&sort(konstruktion_1_1,d))&card(konstruktion_1_1,int1))&etype(konstruktion_1_1,int0))&fact(konstruktion_1_1,real))&gener(konstruktion_1_1,ge))&quant(konstruktion_1_1,one))&refer(konstruktion_1_1,refer_c))&varia(konstruktion_1_1,varia_c))&sort(c45,da))&fact(c45,real))&gener(c45,sp))&sort(c57,d))&sort(c57,io))&card(c57,int1))&etype(c57,int0))&fact(c57,real))&gener(c57,sp))&quant(c57,one))&refer(c57,det))&varia(c57,con))&sort(sollen_0,md))&fact(sollen_0,real))&gener(sollen_0,gener_c))&sort(c62,o))&card(c62,int3))&etype(c62,int1))&fact(c62,real))&gener(c62,sp))&quant(c62,nfquant))&refer(c62,det))&varia(c62,con))&sort(n374bernehmen_1_1,da))&fact(n374bernehmen_1_1,real))&gener(n374bernehmen_1_1,ge))&sort(c58,na))&card(c58,int1))&etype(c58,int0))&fact(c58,real))&gener(c58,sp))&quant(c58,one))&refer(c58,indet))&varia(c58,varia_c))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(bmw_0,fe))&sort(c9,as))&card(c9,int1))&etype(c9,int0))&fact(c9,real))&gener(c9,gener_c))&quant(c9,one))&refer(c9,refer_c))&varia(c9,varia_c))&sort(motoranlage_1_1,as))&card(motoranlage_1_1,int1))&etype(motoranlage_1_1,int0))&fact(motoranlage_1_1,real))&gener(motoranlage_1_1,ge))&quant(motoranlage_1_1,one))&refer(motoranlage_1_1,refer_c))&varia(motoranlage_1_1,varia_c))&sort(motor__1_1,d))&card(motor__1_1,int1))&etype(motor__1_1,int0))&fact(motor__1_1,real))&gener(motor__1_1,ge))&quant(motor__1_1,one))&refer(motor__1_1,refer_c))&varia(motor__1_1,varia_c))&sort(anlage_1_1,as))&card(anlage_1_1,int1))&etype(anlage_1_1,int0))&fact(anlage_1_1,real))&gener(anlage_1_1,ge))&quant(anlage_1_1,one))&refer(anlage_1_1,refer_c))&varia(anlage_1_1,varia_c))&sort(rotor_1_1,o))&card(rotor_1_1,int1))&etype(rotor_1_1,int0))&fact(rotor_1_1,real))&gener(rotor_1_1,ge))&quant(rotor_1_1,one))&refer(rotor_1_1,refer_c))&varia(rotor_1_1,varia_c))&sort(nabe_1_1,o))&card(nabe_1_1,int1))&etype(nabe_1_1,int0))&fact(nabe_1_1,real))&gener(nabe_1_1,ge))&quant(nabe_1_1,one))&refer(nabe_1_1,refer_c))&varia(nabe_1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_wp_491)).
fof(c_0_2, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:(((((((((agt(X5,X4)&attr(X1,X2))&attr(X4,X3))&attr(X6,X7))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&subs(X5,n374bernehmen_1_1))&val(X2,bmw_0))&val(X3,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_3, negated_conjecture, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:(((((((((~agt(X12,X11)|~attr(X8,X9))|~attr(X11,X10))|~attr(X13,X14))|~sub(X9,name_1_1))|~sub(X8,firma_1_1))|~sub(X10,name_1_1))|~subs(X12,n374bernehmen_1_1))|~val(X9,bmw_0))|~val(X10,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_4,negated_conjecture,(~val(X1,bmw_0)|~val(X2,bmw_0)|~subs(X3,n374bernehmen_1_1)|~sub(X1,name_1_1)|~sub(X4,firma_1_1)|~sub(X2,name_1_1)|~attr(X5,X6)|~attr(X7,X1)|~attr(X4,X2)|~agt(X3,X7)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_5,hypothesis,(val(c58,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_6,hypothesis,(sub(c58,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_7,hypothesis,(~val(X1,bmw_0)|~attr(X2,c58)|~attr(X3,X1)|~attr(X4,X5)|~agt(X6,X3)|~subs(X6,n374bernehmen_1_1)|~sub(X2,firma_1_1)|~sub(X1,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_5]), c_0_6])])).
cnf(c_0_8,hypothesis,(~attr(X1,c58)|~attr(X2,c58)|~attr(X3,X4)|~agt(X5,X2)|~subs(X5,n374bernehmen_1_1)|~sub(X1,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_5]), c_0_6])])).
cnf(c_0_9,hypothesis,(attr(c57,c58)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_10,hypothesis,(sub(c57,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_11,hypothesis,(~attr(X1,c58)|~attr(X2,X3)|~agt(X4,X1)|~subs(X4,n374bernehmen_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
cnf(c_0_12,hypothesis,(~attr(X1,X2)|~agt(X3,c57)|~subs(X3,n374bernehmen_1_1)), inference(spm,[status(thm)],[c_0_11, c_0_9])).
cnf(c_0_13,hypothesis,(agt(c45,c57)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_14,hypothesis,(subs(c45,n374bernehmen_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_15,hypothesis,(~attr(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
cnf(c_0_16,hypothesis,($false), inference(sr,[status(thm)],[c_0_9, c_0_15]), ['proof']).
# SZS output end CNFRefutation.
EOF
