% Problem    : LCL642+1.020 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n110.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:10:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic U_____043_B31_F1_PI_AE_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (~(?[X1]:~(((((((![X2]:(~(r1(X1,X2))|p3(X2))|~(![X2]:((~(r1(X1,X2))|p3(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p3(X2)))|~(p3(X1)))))))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1)))))))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(![X2]:((~(r1(X1,X2))|p1(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1)))))))|~((((((((((((((((((((((((![X2]:(~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))|~(![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1)))))))&(p2(X1)|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))|~(![X2]:((~(r1(X1,X2))|((![X1]:(~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1))))))|~(![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2)))))))&(p2(X2)|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1)))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|((![X1]:(~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1))))))|~(![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2)))))))&(p2(X2)|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1))))))))|~(((![X2]:(~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))|~(![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1)))))))&(p2(X1)|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))))))))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))))|~(![X2]:(((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))|~(![X1]:((~(r1(X2,X1))|![X2]:((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))))))&(((p2(X1)|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))))|~(![X2]:((((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))))|~(((p2(X1)|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))))))&((p1(X1)|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))))|~(![X2]:(((~(r1(X1,X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))|~(![X1]:((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))))|~((p1(X1)|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))))))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))|~(![X2]:((((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))|~(![X1]:((~(r1(X2,X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))))))))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))|~(![X2]:(((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))|~(![X1]:((~(r1(X2,X1))|![X2]:((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))))))))&(((p2(X1)|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))|~(![X2]:((((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))|~(![X1]:((~(r1(X2,X1))|![X2]:(((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))))|~(((p2(X1)|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))))))))&((p1(X1)|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))|~(![X2]:(((~(r1(X1,X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))|~(![X1]:((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))))|~((p1(X1)|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))))))))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))|~(![X2]:((((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))|~(![X1]:((~(r1(X2,X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))|~(![X2]:(((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))|~(![X1]:((~(r1(X2,X1))|![X2]:((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))))&(((p2(X1)|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))|~(![X2]:((((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))|~(![X1]:((~(r1(X2,X1))|![X2]:(((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))|~(((p2(X1)|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))))&((p1(X1)|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))|~(![X2]:(((~(r1(X1,X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))|~(![X1]:((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))))|~((p1(X1)|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~(![X2]:((((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))|~(![X1]:((~(r1(X2,X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))))))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~(![X2]:(((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))|~(![X1]:((~(r1(X2,X1))|![X2]:((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))))))&(((p2(X1)|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~(![X2]:((((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))|~(![X1]:((~(r1(X2,X1))|![X2]:(((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))|~(((p2(X1)|p1(X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))))))&((p1(X1)|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~(![X2]:(((~(r1(X1,X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))|~(![X1]:((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p1(X2))|![X1]:(((((~(r1(X2,X1))|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))))|~((p1(X1)|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))))))))))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))|~(![X2]:((((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))|~(![X1]:((~(r1(X2,X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))|~(![X2]:(((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))|~(![X1]:((~(r1(X2,X1))|![X2]:((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))&(((p2(X1)|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))|~(![X2]:((((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))|~(![X1]:((~(r1(X2,X1))|![X2]:(((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~(((p2(X1)|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))&((p1(X1)|![X2]:(~(r1(X1,X2))|~$true))|~(![X2]:(((~(r1(X1,X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))|~(![X1]:((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~((p1(X1)|![X2]:(~(r1(X1,X2))|~$true)))))))))&![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', main)).
fof(c_0_1, plain, (![X1]:(epred1_1(X1)<=>~(![X2]:((~r1(X1,X2)|((![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))&(p2(X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|((![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))&(p2(X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))))|~(((![X2]:(~r1(X1,X2)|![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))|~(![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))&(p2(X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))))))))), introduced(definition)).
fof(c_0_2, plain, (![X1]:(epred2_1(X1)<=>~(![X2]:(((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))))|~(![X1]:((~r1(X2,X1)|![X2]:((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))))))))), introduced(definition)).
fof(c_0_3, plain, (![X1]:(epred3_1(X1)<=>~(![X2]:((((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))))|~(![X1]:((~r1(X2,X1)|![X2]:(((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))))|~(((p2(X1)|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))))))))), introduced(definition)).
fof(c_0_4, plain, (![X1]:(epred4_1(X1)<=>~(![X2]:(((~r1(X1,X2)|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))))|~(![X1]:((~r1(X2,X1)|![X2]:((~r1(X1,X2)|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))))|~((p1(X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))))))))), introduced(definition)).
fof(c_0_5, plain, (![X1]:(epred5_1(X1)<=>~(![X2]:((((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))))|~(![X1]:((~r1(X2,X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))))))))), introduced(definition)).
fof(c_0_6, plain, (![X1]:(epred6_1(X1)<=>~(![X2]:(((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))))|~(![X1]:((~r1(X2,X1)|![X2]:((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))))))))), introduced(definition)).
fof(c_0_7, plain, (![X1]:(epred7_1(X1)<=>~(![X2]:((((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))))|~(![X1]:((~r1(X2,X1)|![X2]:(((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))|~(((p2(X1)|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))))))))), introduced(definition)).
fof(c_0_8, plain, (![X1]:(epred8_1(X1)<=>~(![X2]:(((~r1(X1,X2)|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))))|~(![X1]:((~r1(X2,X1)|![X2]:((~r1(X1,X2)|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))|~((p1(X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))))))))), introduced(definition)).
fof(c_0_9, plain, (![X1]:(epred9_1(X1)<=>~(![X2]:((((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))|~(![X1]:((~r1(X2,X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))))))), introduced(definition)).
fof(c_0_10, plain, (![X1]:(epred10_1(X1)<=>~(![X2]:(((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))|~(![X1]:((~r1(X2,X1)|![X2]:((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))))))), introduced(definition)).
fof(c_0_11, plain, (![X1]:(epred11_1(X1)<=>~(![X2]:((((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))|~(![X1]:((~r1(X2,X1)|![X2]:(((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))|~(((p2(X1)|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))))))), introduced(definition)).
fof(c_0_12, plain, (![X1]:(epred12_1(X1)<=>~(![X2]:(((~r1(X1,X2)|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))|~(![X1]:((~r1(X2,X1)|![X2]:((~r1(X1,X2)|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))|~((p1(X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))))))), introduced(definition)).
fof(c_0_13, negated_conjecture, (~(~(?[X1]:~(((((((![X2]:(~r1(X1,X2)|p3(X2))|~(![X2]:((~r1(X1,X2)|p3(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p3(X1))))))|![X2]:(~r1(X1,X2)|p2(X2)))|~(![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))|![X2]:(~r1(X1,X2)|p1(X2)))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))))))|~((((((((((((((((((((((((![X2]:(~r1(X1,X2)|![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))|~(![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))&(p2(X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))|epred1_1(X1))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))|epred2_1(X1)))&(((p2(X1)|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))|epred3_1(X1)))&((p1(X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))))|epred4_1(X1)))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))|epred5_1(X1)))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))|epred6_1(X1)))&(((p2(X1)|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))|epred7_1(X1)))&((p1(X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))))|epred8_1(X1)))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))|epred9_1(X1)))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))|epred10_1(X1)))&(((p2(X1)|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))|epred11_1(X1)))&((p1(X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))|epred12_1(X1)))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~(![X2]:((((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))|~(![X1]:((~r1(X2,X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))))))))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~(![X2]:(((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))|~(![X1]:((~r1(X2,X1)|![X2]:((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))))))))&(((p2(X1)|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~(![X2]:((((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))|~(![X1]:((~r1(X2,X1)|![X2]:(((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))|~(((p2(X1)|p1(X1))|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))))))))&((p1(X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~(![X2]:(((~r1(X1,X2)|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))|~(![X1]:((~r1(X2,X1)|![X2]:((~r1(X1,X2)|p1(X2))|![X1]:(((((~r1(X2,X1)|p4(X1))|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))))|~((p1(X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))))))))))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))|~(![X2]:((((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))|~(![X1]:((~r1(X2,X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))))))))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))|~(![X2]:(((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))|~(![X1]:((~r1(X2,X1)|![X2]:((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))))))))&(((p2(X1)|p1(X1))|![X2]:~r1(X1,X2))|~(![X2]:((((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))|~(![X1]:((~r1(X2,X1)|![X2]:(((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~(((p2(X1)|p1(X1))|![X2]:~r1(X1,X2)))))))))&((p1(X1)|![X2]:~r1(X1,X2))|~(![X2]:(((~r1(X1,X2)|p1(X2))|![X1]:~r1(X2,X1))|~(![X1]:((~r1(X2,X1)|![X2]:((~r1(X1,X2)|p1(X2))|![X1]:~r1(X2,X1)))|~((p1(X1)|![X2]:~r1(X1,X2)))))))))&![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))))))))), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_1]), c_0_2]), c_0_3]), c_0_4]), c_0_5]), c_0_6]), c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11]), c_0_12])).
fof(c_0_14, negated_conjecture, (![X5]:![X9]:![X13]:![X16]:![X17]:![X21]:![X22]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:![X34]:![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:![X43]:![X44]:![X45]:![X46]:![X47]:![X48]:![X49]:![X50]:![X51]:![X52]:![X53]:![X54]:![X55]:![X56]:![X57]:![X58]:![X59]:![X60]:![X61]:![X62]:![X63]:![X64]:![X65]:![X66]:![X67]:![X68]:![X69]:![X73]:![X74]:![X75]:![X76]:![X79]:![X80]:![X84]:![X85]:![X86]:![X87]:![X90]:![X91]:![X95]:![X96]:![X97]:![X98]:![X101]:![X102]:![X106]:![X107]:![X108]:![X109]:![X112]:![X115]:![X116]:![X117]:![X119]:![X122]:![X123]:![X124]:![X126]:![X129]:![X130]:![X131]:![X133]:![X136]:![X137]:![X138]:![X140]:![X142]:![X143]:(((((((r1(esk1_0,esk2_0)&~p3(esk2_0))&(((r1(X5,esk3_1(X5))|(~r1(esk1_0,X5)|p3(X5)))&((r1(esk3_1(X5),esk4_1(X5))|(~r1(esk1_0,X5)|p3(X5)))&(~p3(esk4_1(X5))|(~r1(esk1_0,X5)|p3(X5)))))&(p3(esk3_1(X5))|(~r1(esk1_0,X5)|p3(X5)))))&(r1(esk1_0,esk5_0)&~p2(esk5_0)))&(((r1(X9,esk6_1(X9))|(~r1(esk1_0,X9)|p2(X9)))&((r1(esk6_1(X9),esk7_1(X9))|(~r1(esk1_0,X9)|p2(X9)))&(~p2(esk7_1(X9))|(~r1(esk1_0,X9)|p2(X9)))))&(p2(esk6_1(X9))|(~r1(esk1_0,X9)|p2(X9)))))&(r1(esk1_0,esk8_0)&~p1(esk8_0)))&(((r1(X13,esk9_1(X13))|(~r1(esk1_0,X13)|p1(X13)))&((r1(esk9_1(X13),esk10_1(X13))|(~r1(esk1_0,X13)|p1(X13)))&(~p1(esk10_1(X13))|(~r1(esk1_0,X13)|p1(X13)))))&(p1(esk9_1(X13))|(~r1(esk1_0,X13)|p1(X13)))))&(((((((((((((((((((((((((((r1(esk1_0,esk13_0)|((r1(X17,esk11_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0))&((~p2(esk13_0)|((r1(X17,esk11_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0)))&((((~r1(esk13_0,X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk11_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0)))&(((((r1(esk1_0,esk13_0)|((r1(esk11_2(X16,X17),esk12_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0))&((~p2(esk13_0)|((r1(esk11_2(X16,X17),esk12_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0)))&((((~r1(esk13_0,X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk11_2(X16,X17),esk12_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0)))&((((r1(esk1_0,esk13_0)|((~p2(esk12_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0))&((~p2(esk13_0)|((~p2(esk12_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0)))&((((~r1(esk13_0,X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk12_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0)))))&((((r1(esk1_0,esk13_0)|((p2(esk11_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0))&((~p2(esk13_0)|((p2(esk11_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0)))&((((~r1(esk13_0,X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk11_2(X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(esk1_0,X16)))|epred1_1(esk1_0))))&((((r1(esk1_0,esk14_0)|p2(esk1_0))|epred1_1(esk1_0))&(((r1(esk14_0,esk15_0)|p2(esk1_0))|epred1_1(esk1_0))&((~p2(esk15_0)|p2(esk1_0))|epred1_1(esk1_0))))&((p2(esk14_0)|p2(esk1_0))|epred1_1(esk1_0))))&((((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X25)|p4(X25))|p3(X25))|p2(X25))|p1(X25))|(((((~r1(X25,X26)|p4(X26))|p3(X26))|p2(X26))|p1(X26))|(((((~r1(X26,X27)|p4(X27))|p3(X27))|p2(X27))|p1(X27))|(((((~r1(X27,X28)|p4(X28))|p3(X28))|p2(X28))|p1(X28))|~r1(X28,X29))))))|epred2_1(esk1_0)))&(((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X30)|p4(X30))|p3(X30))|p2(X30))|p1(X30))|(((((~r1(X30,X31)|p4(X31))|p3(X31))|p2(X31))|p1(X31))|(((((~r1(X31,X32)|p4(X32))|p3(X32))|p2(X32))|p1(X32))|(((((~r1(X32,X33)|p4(X33))|p3(X33))|p2(X33))|p1(X33))|~r1(X33,X34))))))|epred3_1(esk1_0)))&((p1(esk1_0)|(((((~r1(esk1_0,X35)|p4(X35))|p3(X35))|p2(X35))|p1(X35))|(((((~r1(X35,X36)|p4(X36))|p3(X36))|p2(X36))|p1(X36))|(((((~r1(X36,X37)|p4(X37))|p3(X37))|p2(X37))|p1(X37))|(((((~r1(X37,X38)|p4(X38))|p3(X38))|p2(X38))|p1(X38))|~r1(X38,X39))))))|epred4_1(esk1_0)))&(((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X40)|p4(X40))|p3(X40))|p2(X40))|p1(X40))|(((((~r1(X40,X41)|p4(X41))|p3(X41))|p2(X41))|p1(X41))|(((((~r1(X41,X42)|p4(X42))|p3(X42))|p2(X42))|p1(X42))|~r1(X42,X43)))))|epred5_1(esk1_0)))&((((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X44)|p4(X44))|p3(X44))|p2(X44))|p1(X44))|(((((~r1(X44,X45)|p4(X45))|p3(X45))|p2(X45))|p1(X45))|(((((~r1(X45,X46)|p4(X46))|p3(X46))|p2(X46))|p1(X46))|~r1(X46,X47)))))|epred6_1(esk1_0)))&(((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X48)|p4(X48))|p3(X48))|p2(X48))|p1(X48))|(((((~r1(X48,X49)|p4(X49))|p3(X49))|p2(X49))|p1(X49))|(((((~r1(X49,X50)|p4(X50))|p3(X50))|p2(X50))|p1(X50))|~r1(X50,X51)))))|epred7_1(esk1_0)))&((p1(esk1_0)|(((((~r1(esk1_0,X52)|p4(X52))|p3(X52))|p2(X52))|p1(X52))|(((((~r1(X52,X53)|p4(X53))|p3(X53))|p2(X53))|p1(X53))|(((((~r1(X53,X54)|p4(X54))|p3(X54))|p2(X54))|p1(X54))|~r1(X54,X55)))))|epred8_1(esk1_0)))&(((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X56)|p4(X56))|p3(X56))|p2(X56))|p1(X56))|(((((~r1(X56,X57)|p4(X57))|p3(X57))|p2(X57))|p1(X57))|~r1(X57,X58))))|epred9_1(esk1_0)))&((((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X59)|p4(X59))|p3(X59))|p2(X59))|p1(X59))|(((((~r1(X59,X60)|p4(X60))|p3(X60))|p2(X60))|p1(X60))|~r1(X60,X61))))|epred10_1(esk1_0)))&(((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X62)|p4(X62))|p3(X62))|p2(X62))|p1(X62))|(((((~r1(X62,X63)|p4(X63))|p3(X63))|p2(X63))|p1(X63))|~r1(X63,X64))))|epred11_1(esk1_0)))&((p1(esk1_0)|(((((~r1(esk1_0,X65)|p4(X65))|p3(X65))|p2(X65))|p1(X65))|(((((~r1(X65,X66)|p4(X66))|p3(X66))|p2(X66))|p1(X66))|~r1(X66,X67))))|epred12_1(esk1_0)))&(((((((r1(esk1_0,esk16_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69))))&(~p4(esk16_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&(~p3(esk16_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&(~p2(esk16_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&(~p1(esk16_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&((((((r1(esk16_0,esk17_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69))))&(~p4(esk17_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&(~p3(esk17_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&(~p2(esk17_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&(~p1(esk17_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&(r1(esk17_0,esk18_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69))))))&((((((~p4(X73)|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69))))&((~p3(X73)|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&((~p2(X73)|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&((~p1(X73)|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&(((((((r1(X73,esk19_1(X73))|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69))))&((~p4(esk19_1(X73))|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&((~p3(esk19_1(X73))|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&((~p2(esk19_1(X73))|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&((~p1(esk19_1(X73))|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69)))))&((r1(esk19_1(X73),esk20_1(X73))|(~r1(esk16_0,X73)|(((((~r1(X73,X74)|p4(X74))|p3(X74))|p2(X74))|p1(X74))|(((((~r1(X74,X75)|p4(X75))|p3(X75))|p2(X75))|p1(X75))|~r1(X75,X76)))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X68)|p4(X68))|p3(X68))|p2(X68))|p1(X68))|~r1(X68,X69))))))))&((((((r1(esk1_0,esk21_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80))))&(~p3(esk21_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&(~p2(esk21_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&(~p1(esk21_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&((((((r1(esk21_0,esk22_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80))))&(~p4(esk22_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&(~p3(esk22_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&(~p2(esk22_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&(~p1(esk22_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&(r1(esk22_0,esk23_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80))))))&(((((~p3(X84)|(~r1(esk21_0,X84)|((((~r1(X84,X85)|p3(X85))|p2(X85))|p1(X85))|(((((~r1(X85,X86)|p4(X86))|p3(X86))|p2(X86))|p1(X86))|~r1(X86,X87)))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80))))&((~p2(X84)|(~r1(esk21_0,X84)|((((~r1(X84,X85)|p3(X85))|p2(X85))|p1(X85))|(((((~r1(X85,X86)|p4(X86))|p3(X86))|p2(X86))|p1(X86))|~r1(X86,X87)))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&((~p1(X84)|(~r1(esk21_0,X84)|((((~r1(X84,X85)|p3(X85))|p2(X85))|p1(X85))|(((((~r1(X85,X86)|p4(X86))|p3(X86))|p2(X86))|p1(X86))|~r1(X86,X87)))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&(((((((r1(X84,esk24_1(X84))|(~r1(esk21_0,X84)|((((~r1(X84,X85)|p3(X85))|p2(X85))|p1(X85))|(((((~r1(X85,X86)|p4(X86))|p3(X86))|p2(X86))|p1(X86))|~r1(X86,X87)))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80))))&((~p4(esk24_1(X84))|(~r1(esk21_0,X84)|((((~r1(X84,X85)|p3(X85))|p2(X85))|p1(X85))|(((((~r1(X85,X86)|p4(X86))|p3(X86))|p2(X86))|p1(X86))|~r1(X86,X87)))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&((~p3(esk24_1(X84))|(~r1(esk21_0,X84)|((((~r1(X84,X85)|p3(X85))|p2(X85))|p1(X85))|(((((~r1(X85,X86)|p4(X86))|p3(X86))|p2(X86))|p1(X86))|~r1(X86,X87)))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&((~p2(esk24_1(X84))|(~r1(esk21_0,X84)|((((~r1(X84,X85)|p3(X85))|p2(X85))|p1(X85))|(((((~r1(X85,X86)|p4(X86))|p3(X86))|p2(X86))|p1(X86))|~r1(X86,X87)))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&((~p1(esk24_1(X84))|(~r1(esk21_0,X84)|((((~r1(X84,X85)|p3(X85))|p2(X85))|p1(X85))|(((((~r1(X85,X86)|p4(X86))|p3(X86))|p2(X86))|p1(X86))|~r1(X86,X87)))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80)))))&((r1(esk24_1(X84),esk25_1(X84))|(~r1(esk21_0,X84)|((((~r1(X84,X85)|p3(X85))|p2(X85))|p1(X85))|(((((~r1(X85,X86)|p4(X86))|p3(X86))|p2(X86))|p1(X86))|~r1(X86,X87)))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|(((((~r1(esk1_0,X79)|p4(X79))|p3(X79))|p2(X79))|p1(X79))|~r1(X79,X80))))))))&(((((r1(esk1_0,esk26_0)|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91))))&(~p2(esk26_0)|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&(~p1(esk26_0)|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&((((((r1(esk26_0,esk27_0)|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91))))&(~p4(esk27_0)|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&(~p3(esk27_0)|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&(~p2(esk27_0)|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&(~p1(esk27_0)|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&(r1(esk27_0,esk28_0)|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91))))))&((((~p2(X95)|(~r1(esk26_0,X95)|(((~r1(X95,X96)|p2(X96))|p1(X96))|(((((~r1(X96,X97)|p4(X97))|p3(X97))|p2(X97))|p1(X97))|~r1(X97,X98)))))|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91))))&((~p1(X95)|(~r1(esk26_0,X95)|(((~r1(X95,X96)|p2(X96))|p1(X96))|(((((~r1(X96,X97)|p4(X97))|p3(X97))|p2(X97))|p1(X97))|~r1(X97,X98)))))|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&(((((((r1(X95,esk29_1(X95))|(~r1(esk26_0,X95)|(((~r1(X95,X96)|p2(X96))|p1(X96))|(((((~r1(X96,X97)|p4(X97))|p3(X97))|p2(X97))|p1(X97))|~r1(X97,X98)))))|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91))))&((~p4(esk29_1(X95))|(~r1(esk26_0,X95)|(((~r1(X95,X96)|p2(X96))|p1(X96))|(((((~r1(X96,X97)|p4(X97))|p3(X97))|p2(X97))|p1(X97))|~r1(X97,X98)))))|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&((~p3(esk29_1(X95))|(~r1(esk26_0,X95)|(((~r1(X95,X96)|p2(X96))|p1(X96))|(((((~r1(X96,X97)|p4(X97))|p3(X97))|p2(X97))|p1(X97))|~r1(X97,X98)))))|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&((~p2(esk29_1(X95))|(~r1(esk26_0,X95)|(((~r1(X95,X96)|p2(X96))|p1(X96))|(((((~r1(X96,X97)|p4(X97))|p3(X97))|p2(X97))|p1(X97))|~r1(X97,X98)))))|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&((~p1(esk29_1(X95))|(~r1(esk26_0,X95)|(((~r1(X95,X96)|p2(X96))|p1(X96))|(((((~r1(X96,X97)|p4(X97))|p3(X97))|p2(X97))|p1(X97))|~r1(X97,X98)))))|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91)))))&((r1(esk29_1(X95),esk30_1(X95))|(~r1(esk26_0,X95)|(((~r1(X95,X96)|p2(X96))|p1(X96))|(((((~r1(X96,X97)|p4(X97))|p3(X97))|p2(X97))|p1(X97))|~r1(X97,X98)))))|((p2(esk1_0)|p1(esk1_0))|(((((~r1(esk1_0,X90)|p4(X90))|p3(X90))|p2(X90))|p1(X90))|~r1(X90,X91))))))))&((((r1(esk1_0,esk31_0)|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102))))&(~p1(esk31_0)|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102)))))&((((((r1(esk31_0,esk32_0)|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102))))&(~p4(esk32_0)|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102)))))&(~p3(esk32_0)|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102)))))&(~p2(esk32_0)|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102)))))&(~p1(esk32_0)|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102)))))&(r1(esk32_0,esk33_0)|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102))))))&(((~p1(X106)|(~r1(esk31_0,X106)|((~r1(X106,X107)|p1(X107))|(((((~r1(X107,X108)|p4(X108))|p3(X108))|p2(X108))|p1(X108))|~r1(X108,X109)))))|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102))))&(((((((r1(X106,esk34_1(X106))|(~r1(esk31_0,X106)|((~r1(X106,X107)|p1(X107))|(((((~r1(X107,X108)|p4(X108))|p3(X108))|p2(X108))|p1(X108))|~r1(X108,X109)))))|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102))))&((~p4(esk34_1(X106))|(~r1(esk31_0,X106)|((~r1(X106,X107)|p1(X107))|(((((~r1(X107,X108)|p4(X108))|p3(X108))|p2(X108))|p1(X108))|~r1(X108,X109)))))|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102)))))&((~p3(esk34_1(X106))|(~r1(esk31_0,X106)|((~r1(X106,X107)|p1(X107))|(((((~r1(X107,X108)|p4(X108))|p3(X108))|p2(X108))|p1(X108))|~r1(X108,X109)))))|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102)))))&((~p2(esk34_1(X106))|(~r1(esk31_0,X106)|((~r1(X106,X107)|p1(X107))|(((((~r1(X107,X108)|p4(X108))|p3(X108))|p2(X108))|p1(X108))|~r1(X108,X109)))))|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102)))))&((~p1(esk34_1(X106))|(~r1(esk31_0,X106)|((~r1(X106,X107)|p1(X107))|(((((~r1(X107,X108)|p4(X108))|p3(X108))|p2(X108))|p1(X108))|~r1(X108,X109)))))|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102)))))&((r1(esk34_1(X106),esk35_1(X106))|(~r1(esk31_0,X106)|((~r1(X106,X107)|p1(X107))|(((((~r1(X107,X108)|p4(X108))|p3(X108))|p2(X108))|p1(X108))|~r1(X108,X109)))))|(p1(esk1_0)|(((((~r1(esk1_0,X101)|p4(X101))|p3(X101))|p2(X101))|p1(X101))|~r1(X101,X102))))))))&(((((((r1(esk1_0,esk36_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112)))&(~p4(esk36_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112))))&(~p3(esk36_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112))))&(~p2(esk36_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112))))&(~p1(esk36_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112))))&(r1(esk36_0,esk37_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112))))&((((((~p4(X115)|(~r1(esk36_0,X115)|(((((~r1(X115,X116)|p4(X116))|p3(X116))|p2(X116))|p1(X116))|~r1(X116,X117))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112)))&((~p3(X115)|(~r1(esk36_0,X115)|(((((~r1(X115,X116)|p4(X116))|p3(X116))|p2(X116))|p1(X116))|~r1(X116,X117))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112))))&((~p2(X115)|(~r1(esk36_0,X115)|(((((~r1(X115,X116)|p4(X116))|p3(X116))|p2(X116))|p1(X116))|~r1(X116,X117))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112))))&((~p1(X115)|(~r1(esk36_0,X115)|(((((~r1(X115,X116)|p4(X116))|p3(X116))|p2(X116))|p1(X116))|~r1(X116,X117))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112))))&((r1(X115,esk38_1(X115))|(~r1(esk36_0,X115)|(((((~r1(X115,X116)|p4(X116))|p3(X116))|p2(X116))|p1(X116))|~r1(X116,X117))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X112))))))&((((((r1(esk1_0,esk39_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X119)))&(~p3(esk39_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X119))))&(~p2(esk39_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X119))))&(~p1(esk39_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X119))))&(r1(esk39_0,esk40_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X119))))&(((((~p3(X122)|(~r1(esk39_0,X122)|((((~r1(X122,X123)|p3(X123))|p2(X123))|p1(X123))|~r1(X123,X124))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X119)))&((~p2(X122)|(~r1(esk39_0,X122)|((((~r1(X122,X123)|p3(X123))|p2(X123))|p1(X123))|~r1(X123,X124))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X119))))&((~p1(X122)|(~r1(esk39_0,X122)|((((~r1(X122,X123)|p3(X123))|p2(X123))|p1(X123))|~r1(X123,X124))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X119))))&((r1(X122,esk41_1(X122))|(~r1(esk39_0,X122)|((((~r1(X122,X123)|p3(X123))|p2(X123))|p1(X123))|~r1(X123,X124))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X119))))))&(((((r1(esk1_0,esk42_0)|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X126)))&(~p2(esk42_0)|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X126))))&(~p1(esk42_0)|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X126))))&(r1(esk42_0,esk43_0)|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X126))))&((((~p2(X129)|(~r1(esk42_0,X129)|(((~r1(X129,X130)|p2(X130))|p1(X130))|~r1(X130,X131))))|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X126)))&((~p1(X129)|(~r1(esk42_0,X129)|(((~r1(X129,X130)|p2(X130))|p1(X130))|~r1(X130,X131))))|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X126))))&((r1(X129,esk44_1(X129))|(~r1(esk42_0,X129)|(((~r1(X129,X130)|p2(X130))|p1(X130))|~r1(X130,X131))))|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X126))))))&((((r1(esk1_0,esk45_0)|(p1(esk1_0)|~r1(esk1_0,X133)))&(~p1(esk45_0)|(p1(esk1_0)|~r1(esk1_0,X133))))&(r1(esk45_0,esk46_0)|(p1(esk1_0)|~r1(esk1_0,X133))))&(((~p1(X136)|(~r1(esk45_0,X136)|((~r1(X136,X137)|p1(X137))|~r1(X137,X138))))|(p1(esk1_0)|~r1(esk1_0,X133)))&((r1(X136,esk47_1(X136))|(~r1(esk45_0,X136)|((~r1(X136,X137)|p1(X137))|~r1(X137,X138))))|(p1(esk1_0)|~r1(esk1_0,X133))))))&(((r1(X140,esk48_1(X140))|(~r1(esk1_0,X140)|p2(X140)))&(~p2(esk48_1(X140))|(~r1(esk1_0,X140)|p2(X140))))&(((~r1(esk48_1(X140),X142)|(~r1(X142,X143)|p2(X143)))|~p2(X142))|(~r1(esk1_0,X140)|p2(X140))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])).
cnf(c_0_15,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(X2,esk11_2(X1,X2))|r1(esk1_0,esk13_0)|~r1(esk1_0,X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_16,negated_conjecture,(r1(esk1_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17,negated_conjecture,(p2(X1)|r1(X1,esk48_1(X1))|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18,negated_conjecture,(~p2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,negated_conjecture,(p2(X1)|~r1(esk1_0,X1)|~p2(esk48_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|p2(esk11_2(X1,X2))|r1(esk1_0,esk13_0)|~r1(esk1_0,X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, (![X1]:(epred1_1(X1)=>~(![X2]:((~r1(X1,X2)|((![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))&(p2(X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|((![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))&(p2(X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))))|~(((![X2]:(~r1(X1,X2)|![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))|~(![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))&(p2(X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))))))))), inference(split_equiv,[status(thm)],[c_0_1])).
cnf(c_0_22,negated_conjecture,(p2(X1)|p2(X3)|~r1(esk1_0,X1)|~p2(X2)|~r1(X2,X3)|~r1(esk48_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk11_2(esk5_0,X1))|r1(esk1_0,esk13_0)|~r1(esk5_0,X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_24,negated_conjecture,(r1(esk5_0,esk48_1(esk5_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_16]), c_0_18])).
cnf(c_0_25,negated_conjecture,(~p2(esk48_1(esk5_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_16]), c_0_18])).
cnf(c_0_26,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_2(esk5_0,X1))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk5_0,X1)), inference(spm,[status(thm)],[c_0_20, c_0_16])).
cnf(c_0_27,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(esk11_2(X1,X2),esk12_2(X1,X2))|r1(esk1_0,esk13_0)|~r1(esk1_0,X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_28,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(esk1_0,esk13_0)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk12_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_29, plain, (![X3]:![X7]:![X8]:![X9]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X21]:![X22]:![X27]:![X28]:![X29]:![X32]:![X33]:(((r1(X3,esk49_1(X3))|~epred1_1(X3))&(((((~p2(esk49_1(X3))|r1(esk49_1(X3),esk50_1(X3)))|~epred1_1(X3))&((((~r1(esk49_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|r1(esk49_1(X3),esk50_1(X3)))|~epred1_1(X3)))&(((((~p2(esk49_1(X3))|r1(esk50_1(X3),esk51_1(X3)))|~epred1_1(X3))&((((~r1(esk49_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|r1(esk50_1(X3),esk51_1(X3)))|~epred1_1(X3)))&(((~p2(esk49_1(X3))|~p2(esk51_1(X3)))|~epred1_1(X3))&((((~r1(esk49_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|~p2(esk51_1(X3)))|~epred1_1(X3))))&(((~p2(esk49_1(X3))|((~r1(esk51_1(X3),X7)|(~r1(X7,X8)|p2(X8)))|~p2(X7)))|~epred1_1(X3))&((((~r1(esk49_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|((~r1(esk51_1(X3),X7)|(~r1(X7,X8)|p2(X8)))|~p2(X7)))|~epred1_1(X3)))))&(((((~p2(esk49_1(X3))|(r1(X9,esk52_2(X3,X9))|(~r1(esk49_1(X3),X9)|p2(X9))))|~epred1_1(X3))&((((~r1(esk49_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|(r1(X9,esk52_2(X3,X9))|(~r1(esk49_1(X3),X9)|p2(X9))))|~epred1_1(X3)))&((((~p2(esk49_1(X3))|(r1(esk52_2(X3,X9),esk53_2(X3,X9))|(~r1(esk49_1(X3),X9)|p2(X9))))|~epred1_1(X3))&((((~r1(esk49_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|(r1(esk52_2(X3,X9),esk53_2(X3,X9))|(~r1(esk49_1(X3),X9)|p2(X9))))|~epred1_1(X3)))&(((~p2(esk49_1(X3))|(~p2(esk53_2(X3,X9))|(~r1(esk49_1(X3),X9)|p2(X9))))|~epred1_1(X3))&((((~r1(esk49_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|(~p2(esk53_2(X3,X9))|(~r1(esk49_1(X3),X9)|p2(X9))))|~epred1_1(X3)))))&(((~p2(esk49_1(X3))|(p2(esk52_2(X3,X9))|(~r1(esk49_1(X3),X9)|p2(X9))))|~epred1_1(X3))&((((~r1(esk49_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|(p2(esk52_2(X3,X9))|(~r1(esk49_1(X3),X9)|p2(X9))))|~epred1_1(X3))))))&(((((((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))))&(((((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk54_5(X3,X14,X15,X16,X17),esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))))&((((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk55_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))))))&((((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk56_3(X3,X14,X15))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk56_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk54_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))))&((((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&(((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk57_3(X3,X14,X15),esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk58_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))))&((((((~p2(X14)|r1(X14,esk59_2(X3,X14)))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk59_2(X3,X14)))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk59_2(X3,X14),esk60_2(X3,X14)))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk60_2(X3,X14)))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk60_2(X3,X14)))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk60_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk61_3(X3,X14,X29),esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk62_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk61_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk57_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk49_1(X3),X14)))|~epred1_1(X3))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).
cnf(c_0_30,negated_conjecture,(p2(X1)|~p2(X2)|~r1(esk48_1(esk5_0),X2)|~r1(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_16]), c_0_18])).
cnf(c_0_31,negated_conjecture,(epred1_1(esk1_0)|r1(esk48_1(esk5_0),esk11_2(esk5_0,esk48_1(esk5_0)))|r1(esk1_0,esk13_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_32,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_2(esk5_0,esk48_1(esk5_0)))|r1(esk1_0,esk13_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_24]), c_0_25])).
cnf(c_0_33,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk11_2(esk5_0,X1),esk12_2(esk5_0,X1))|r1(esk1_0,esk13_0)|~r1(esk5_0,X1)), inference(spm,[status(thm)],[c_0_27, c_0_16])).
cnf(c_0_34,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk12_2(esk5_0,X1))|~r1(esk5_0,X1)), inference(spm,[status(thm)],[c_0_28, c_0_16])).
cnf(c_0_35,negated_conjecture,(p2(X1)|r1(X1,esk6_1(X1))|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_36,plain,(r1(X1,esk49_1(X1))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk11_2(esk5_0,esk48_1(esk5_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_38,negated_conjecture,(epred1_1(esk1_0)|r1(esk11_2(esk5_0,esk48_1(esk5_0)),esk12_2(esk5_0,esk48_1(esk5_0)))|r1(esk1_0,esk13_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_24]), c_0_25])).
cnf(c_0_39,negated_conjecture,(epred1_1(esk1_0)|r1(esk1_0,esk13_0)|~p2(esk12_2(esk5_0,esk48_1(esk5_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_24]), c_0_25])).
cnf(c_0_40,plain,(p2(esk49_1(esk1_0))|r1(esk49_1(esk1_0),esk6_1(esk49_1(esk1_0)))|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41,negated_conjecture,(epred1_1(esk1_0)|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_42,negated_conjecture,(p2(X1)|p2(esk6_1(X1))|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_43,plain,(p2(X2)|r1(X2,esk52_2(X1,X2))|p2(X4)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~p2(X3)|~r1(X3,X4)|~r1(esk49_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_44,plain,(p2(esk49_1(esk1_0))|r1(esk49_1(esk1_0),esk6_1(esk49_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45,plain,(p2(esk6_1(esk49_1(esk1_0)))|p2(esk49_1(esk1_0))|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_36])).
cnf(c_0_46,plain,(p2(esk49_1(esk1_0))|r1(esk49_1(esk1_0),esk48_1(esk49_1(esk1_0)))|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_17, c_0_36])).
cnf(c_0_47,plain,(p2(esk49_1(esk1_0))|~epred1_1(esk1_0)|~p2(esk48_1(esk49_1(esk1_0)))), inference(spm,[status(thm)],[c_0_19, c_0_36])).
cnf(c_0_48,negated_conjecture,(p2(X1)|r1(esk6_1(X1),esk7_1(X1))|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_49,negated_conjecture,(p2(X1)|~r1(esk1_0,X1)|~p2(esk7_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_50,plain,(p2(X2)|p2(esk52_2(X1,X2))|p2(X4)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~p2(X3)|~r1(X3,X4)|~r1(esk49_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_51,plain,(p2(esk49_1(esk1_0))|p2(X1)|p2(X2)|r1(X2,esk52_2(esk1_0,X2))|r1(esk1_0,esk13_0)|~r1(esk6_1(esk49_1(esk1_0)),X1)|~r1(esk49_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_41])).
cnf(c_0_52,plain,(p2(esk49_1(esk1_0))|r1(esk49_1(esk1_0),esk48_1(esk49_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_46, c_0_41])).
cnf(c_0_53,plain,(p2(esk49_1(esk1_0))|r1(esk1_0,esk13_0)|~p2(esk48_1(esk49_1(esk1_0)))), inference(spm,[status(thm)],[c_0_47, c_0_41])).
cnf(c_0_54,plain,(p2(esk49_1(esk1_0))|r1(esk6_1(esk49_1(esk1_0)),esk7_1(esk49_1(esk1_0)))|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_36])).
cnf(c_0_55,plain,(p2(esk49_1(esk1_0))|~epred1_1(esk1_0)|~p2(esk7_1(esk49_1(esk1_0)))), inference(spm,[status(thm)],[c_0_49, c_0_36])).
cnf(c_0_56,plain,(p2(esk52_2(esk1_0,X1))|p2(esk49_1(esk1_0))|p2(X2)|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk6_1(esk49_1(esk1_0)),X2)|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_44]), c_0_45]), c_0_41])).
cnf(c_0_57,plain,(p2(X2)|r1(esk52_2(X1,X2),esk53_2(X1,X2))|p2(X4)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~p2(X3)|~r1(X3,X4)|~r1(esk49_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_58,plain,(p2(esk49_1(esk1_0))|p2(X1)|~epred1_1(esk1_0)|~p2(X2)|~r1(esk48_1(esk49_1(esk1_0)),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_22, c_0_36])).
cnf(c_0_59,plain,(p2(esk49_1(esk1_0))|p2(X1)|r1(esk48_1(esk49_1(esk1_0)),esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|r1(esk1_0,esk13_0)|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_60,plain,(p2(esk49_1(esk1_0))|r1(esk6_1(esk49_1(esk1_0)),esk7_1(esk49_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_54, c_0_41])).
cnf(c_0_61,plain,(p2(esk49_1(esk1_0))|r1(esk1_0,esk13_0)|~p2(esk7_1(esk49_1(esk1_0)))), inference(spm,[status(thm)],[c_0_55, c_0_41])).
cnf(c_0_62,plain,(p2(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|p2(esk49_1(esk1_0))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_52]), c_0_53])).
cnf(c_0_63,plain,(p2(esk49_1(esk1_0))|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,X2),esk53_2(esk1_0,X2))|r1(esk1_0,esk13_0)|~r1(esk6_1(esk49_1(esk1_0)),X1)|~r1(esk49_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_44]), c_0_45]), c_0_41])).
cnf(c_0_64,plain,(p2(X2)|p2(X4)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~p2(esk53_2(X1,X2))|~p2(X3)|~r1(X3,X4)|~r1(esk49_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_65,plain,(p2(esk49_1(esk1_0))|p2(X1)|r1(esk1_0,esk13_0)|~p2(X2)|~r1(esk48_1(esk49_1(esk1_0)),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_58, c_0_41])).
cnf(c_0_66,plain,(p2(esk49_1(esk1_0))|r1(esk48_1(esk49_1(esk1_0)),esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])).
cnf(c_0_67,plain,(p2(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|p2(esk49_1(esk1_0))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_60]), c_0_61])).
cnf(c_0_68,plain,(p2(esk49_1(esk1_0))|p2(X1)|r1(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))),esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|r1(esk1_0,esk13_0)|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_52]), c_0_53])).
cnf(c_0_69,plain,(p2(esk49_1(esk1_0))|p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~p2(esk53_2(esk1_0,X2))|~r1(esk6_1(esk49_1(esk1_0)),X1)|~r1(esk49_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_44]), c_0_45]), c_0_41])).
cnf(c_0_70,plain,(p2(esk49_1(esk1_0))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_71,plain,(p2(esk49_1(esk1_0))|r1(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))),esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_60]), c_0_61])).
cnf(c_0_72,plain,(p2(esk49_1(esk1_0))|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_52]), c_0_53])).
cnf(c_0_73,plain,(p2(esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|p2(esk49_1(esk1_0))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_74,plain,(p2(esk49_1(esk1_0))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_75,plain,(r1(esk49_1(X1),esk50_1(X1))|~epred1_1(X1)|~p2(esk49_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_76,plain,(p2(esk49_1(esk1_0))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_60]), c_0_61])).
cnf(c_0_77,plain,(p2(X2)|r1(X2,esk52_2(X1,X2))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~p2(esk49_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_78,plain,(p2(X3)|r1(X3,esk57_3(X1,X2,X3))|p2(X4)|r1(X4,esk61_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_79,plain,(r1(esk49_1(esk1_0),esk50_1(esk1_0))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_41])).
cnf(c_0_80,plain,(p2(X1)|r1(X1,esk52_2(esk1_0,X1))|r1(esk1_0,esk13_0)|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_76]), c_0_41])).
cnf(c_0_81,plain,(p2(X3)|r1(X3,esk57_3(X1,X2,X3))|p2(X4)|p2(esk61_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_82,plain,(p2(X1)|p2(X2)|r1(X1,esk61_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_41])).
cnf(c_0_83,plain,(p2(esk50_1(esk1_0))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_80, c_0_79])).
cnf(c_0_84,plain,(r1(esk50_1(X1),esk51_1(X1))|~epred1_1(X1)|~p2(esk49_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_85,plain,(~epred1_1(X1)|~p2(esk51_1(X1))|~p2(esk49_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_86,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_79]), c_0_41])).
cnf(c_0_87,plain,(p2(X3)|r1(X3,esk57_3(X1,X2,X3))|p2(X4)|r1(esk61_3(X1,X2,X4),esk62_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_88,plain,(p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_89,plain,(r1(esk50_1(esk1_0),esk51_1(esk1_0))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_76]), c_0_41])).
cnf(c_0_90,plain,(r1(esk1_0,esk13_0)|~p2(esk51_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_76]), c_0_41])).
cnf(c_0_91,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_86, c_0_83])).
cnf(c_0_92,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X1),esk62_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_79]), c_0_41])).
cnf(c_0_93,plain,(p2(X3)|r1(esk57_3(X1,X2,X3),esk58_3(X1,X2,X3))|p2(X4)|r1(X4,esk61_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_94,plain,(p2(X3)|r1(esk57_3(X1,X2,X3),esk58_3(X1,X2,X3))|p2(X4)|p2(esk61_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_95,plain,(p2(X3)|r1(X3,esk57_3(X1,X2,X3))|p2(X4)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk62_3(X1,X2,X4))|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_96,plain,(p2(X3)|~epred1_1(X1)|~p2(X2)|~r1(X2,X3)|~r1(esk51_1(X1),X2)|~p2(esk49_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_97,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90])).
cnf(c_0_98,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_89]), c_0_90])).
cnf(c_0_99,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_92, c_0_83])).
cnf(c_0_100,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(X1,esk61_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_79]), c_0_41])).
cnf(c_0_101,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_79]), c_0_41])).
cnf(c_0_102,plain,(p2(X3)|r1(esk57_3(X1,X2,X3),esk58_3(X1,X2,X3))|p2(X4)|r1(esk61_3(X1,X2,X4),esk62_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_103,plain,(p2(X2)|r1(esk52_2(X1,X2),esk53_2(X1,X2))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~p2(esk49_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_104,plain,(p2(X1)|p2(X2)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_79]), c_0_41])).
cnf(c_0_105,plain,(p2(X1)|r1(esk1_0,esk13_0)|~p2(X2)|~r1(esk51_1(esk1_0),X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_76]), c_0_41])).
cnf(c_0_106,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_89]), c_0_90])).
cnf(c_0_107,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_89]), c_0_90])).
cnf(c_0_108,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_89]), c_0_90])).
cnf(c_0_109,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_100, c_0_83])).
cnf(c_0_110,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_101, c_0_83])).
cnf(c_0_111,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk61_3(esk1_0,esk50_1(esk1_0),X1),esk62_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_79]), c_0_41])).
cnf(c_0_112,plain,(p2(X1)|r1(esk52_2(esk1_0,X1),esk53_2(esk1_0,X1))|r1(esk1_0,esk13_0)|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_76]), c_0_41])).
cnf(c_0_113,plain,(p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_104, c_0_83])).
cnf(c_0_114,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107])).
cnf(c_0_115,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_89]), c_0_90])).
cnf(c_0_116,plain,(p2(X3)|p2(esk57_3(X1,X2,X3))|p2(X4)|r1(X4,esk61_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_117,plain,(p2(X3)|r1(esk57_3(X1,X2,X3),esk58_3(X1,X2,X3))|p2(X4)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk62_3(X1,X2,X4))|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_118,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_89]), c_0_90])).
cnf(c_0_119,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_89]), c_0_90])).
cnf(c_0_120,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_111, c_0_83])).
cnf(c_0_121,plain,(p2(esk50_1(esk1_0))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_112, c_0_79])).
cnf(c_0_122,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_89]), c_0_90])).
cnf(c_0_123,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_124,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_79]), c_0_41])).
cnf(c_0_125,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_79]), c_0_41])).
cnf(c_0_126,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_89]), c_0_90])).
cnf(c_0_127,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_89]), c_0_90])).
cnf(c_0_128,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_89]), c_0_90])).
cnf(c_0_129,plain,(p2(X3)|p2(X4)|r1(X4,esk61_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~p2(esk58_3(X1,X2,X3))|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_130,plain,(p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_82, c_0_121])).
cnf(c_0_131,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_86, c_0_121])).
cnf(c_0_132,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_133,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_124, c_0_83])).
cnf(c_0_134,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_125, c_0_83])).
cnf(c_0_135,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_126]), c_0_127])).
cnf(c_0_136,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_89]), c_0_90])).
cnf(c_0_137,plain,(p2(X1)|p2(X2)|r1(X1,esk61_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_79]), c_0_41])).
cnf(c_0_138,plain,(p2(X3)|p2(esk57_3(X1,X2,X3))|p2(X4)|r1(esk61_3(X1,X2,X4),esk62_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_139,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_89]), c_0_90])).
cnf(c_0_140,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_89]), c_0_90])).
cnf(c_0_141,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_92, c_0_121])).
cnf(c_0_142,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_89]), c_0_90])).
cnf(c_0_143,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_89]), c_0_90])).
cnf(c_0_144,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_89]), c_0_90])).
cnf(c_0_145,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_146,plain,(p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_137, c_0_83])).
cnf(c_0_147,plain,(p2(X3)|p2(esk57_3(X1,X2,X3))|p2(X4)|p2(esk61_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_148,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_79]), c_0_41])).
cnf(c_0_149,plain,(p2(X3)|p2(X4)|r1(esk61_3(X1,X2,X4),esk62_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~p2(esk58_3(X1,X2,X3))|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_150,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_89]), c_0_90])).
cnf(c_0_151,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_89]), c_0_90])).
cnf(c_0_152,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_89]), c_0_90])).
cnf(c_0_153,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_100, c_0_121])).
cnf(c_0_154,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_101, c_0_121])).
cnf(c_0_155,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_105, c_0_142])).
cnf(c_0_156,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_89]), c_0_90])).
cnf(c_0_157,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_158,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_89]), c_0_90])).
cnf(c_0_159,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk61_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_79]), c_0_41])).
cnf(c_0_160,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_148, c_0_83])).
cnf(c_0_161,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X1),esk62_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_79]), c_0_41])).
cnf(c_0_162,plain,(p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_104, c_0_121])).
cnf(c_0_163,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_150]), c_0_151])).
cnf(c_0_164,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_89]), c_0_90])).
cnf(c_0_165,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_89]), c_0_90])).
cnf(c_0_166,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_89]), c_0_90])).
cnf(c_0_167,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_111, c_0_121])).
cnf(c_0_168,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_169,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_89]), c_0_90])).
cnf(c_0_170,plain,(r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_89]), c_0_90])).
cnf(c_0_171,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_159, c_0_83])).
cnf(c_0_172,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_89]), c_0_90])).
cnf(c_0_173,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_161, c_0_83])).
cnf(c_0_174,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_89]), c_0_90])).
cnf(c_0_175,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_163, c_0_164])).
cnf(c_0_176,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_89]), c_0_90])).
cnf(c_0_177,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_89]), c_0_90])).
cnf(c_0_178,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_89]), c_0_90])).
cnf(c_0_179,plain,(p2(X3)|p2(esk57_3(X1,X2,X3))|p2(X4)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk62_3(X1,X2,X4))|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_180,plain,(r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_170])).
cnf(c_0_181,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_89]), c_0_90])).
cnf(c_0_182,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_89]), c_0_90])).
cnf(c_0_183,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_89]), c_0_90])).
cnf(c_0_184,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_174, c_0_175])).
cnf(c_0_185,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_124, c_0_121])).
cnf(c_0_186,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_125, c_0_121])).
cnf(c_0_187,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_176]), c_0_177])).
cnf(c_0_188,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_89]), c_0_90])).
cnf(c_0_189,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_79]), c_0_41])).
cnf(c_0_190,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_105, c_0_180])).
cnf(c_0_191,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_89]), c_0_90])).
cnf(c_0_192,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_155, c_0_182])).
cnf(c_0_193,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_89]), c_0_90])).
cnf(c_0_194,plain,(r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_89]), c_0_90])).
cnf(c_0_195,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_89]), c_0_90])).
cnf(c_0_196,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_89]), c_0_90])).
cnf(c_0_197,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_187, c_0_188])).
cnf(c_0_198,plain,(p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_137, c_0_121])).
cnf(c_0_199,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_189, c_0_83])).
cnf(c_0_200,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_190, c_0_191])).
cnf(c_0_201,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_169]), c_0_193])).
cnf(c_0_202,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_105, c_0_194])).
cnf(c_0_203,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_89]), c_0_90])).
cnf(c_0_204,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_196, c_0_197])).
cnf(c_0_205,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_89]), c_0_90])).
cnf(c_0_206,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_148, c_0_121])).
cnf(c_0_207,plain,(p2(X3)|p2(X4)|p2(esk61_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~p2(esk58_3(X1,X2,X3))|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_208,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_89]), c_0_90])).
cnf(c_0_209,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_200, c_0_201])).
cnf(c_0_210,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_202, c_0_203])).
cnf(c_0_211,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_89]), c_0_90])).
cnf(c_0_212,plain,(r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_89]), c_0_90])).
cnf(c_0_213,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_159, c_0_121])).
cnf(c_0_214,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_89]), c_0_90])).
cnf(c_0_215,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_161, c_0_121])).
cnf(c_0_216,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_79]), c_0_41])).
cnf(c_0_217,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_208, c_0_209])).
cnf(c_0_218,plain,(r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_212])).
cnf(c_0_219,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_89]), c_0_90])).
cnf(c_0_220,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_89]), c_0_90])).
cnf(c_0_221,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_89]), c_0_90])).
cnf(c_0_222,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_216, c_0_83])).
cnf(c_0_223,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_89]), c_0_90])).
cnf(c_0_224,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_105, c_0_218])).
cnf(c_0_225,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_89]), c_0_90])).
cnf(c_0_226,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_202, c_0_220])).
cnf(c_0_227,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_89]), c_0_90])).
cnf(c_0_228,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_89]), c_0_90])).
cnf(c_0_229,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_155, c_0_223])).
cnf(c_0_230,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_189, c_0_121])).
cnf(c_0_231,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_224, c_0_225])).
cnf(c_0_232,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_211]), c_0_227])).
cnf(c_0_233,plain,(p2(X3)|p2(X4)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~p2(esk58_3(X1,X2,X3))|~r1(X2,X4)|~p2(esk62_3(X1,X2,X4))|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_234,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_89]), c_0_90])).
cnf(c_0_235,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_229, c_0_169])).
cnf(c_0_236,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_89]), c_0_90])).
cnf(c_0_237,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_231, c_0_232])).
cnf(c_0_238,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_79]), c_0_41])).
cnf(c_0_239,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_234, c_0_235])).
cnf(c_0_240,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_236, c_0_237])).
cnf(c_0_241,plain,(p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_238, c_0_83])).
cnf(c_0_242,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_190, c_0_239])).
cnf(c_0_243,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_216, c_0_121])).
cnf(c_0_244,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_89]), c_0_90])).
cnf(c_0_245,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_89]), c_0_90])).
cnf(c_0_246,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_242, c_0_201])).
cnf(c_0_247,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_89]), c_0_90])).
cnf(c_0_248,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_202, c_0_244])).
cnf(c_0_249,plain,(p2(X3)|r1(X3,esk57_3(X1,X2,X3))|p2(X4)|r1(X4,esk61_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_250,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_245, c_0_246])).
cnf(c_0_251,plain,(p2(X2)|p2(esk52_2(X1,X2))|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~p2(esk49_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_252,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_89]), c_0_90])).
cnf(c_0_253,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_248, c_0_211])).
cnf(c_0_254,plain,(p2(X3)|r1(X3,esk57_3(X1,X2,X3))|p2(X4)|p2(esk61_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_255,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_79]), c_0_41])).
cnf(c_0_256,plain,(r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_89]), c_0_235]), c_0_90])).
cnf(c_0_257,plain,(p2(esk52_2(esk1_0,X1))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_76]), c_0_41])).
cnf(c_0_258,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_252, c_0_253])).
cnf(c_0_259,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_79]), c_0_41])).
cnf(c_0_260,plain,(p2(X3)|r1(X3,esk57_3(X1,X2,X3))|p2(X4)|r1(esk61_3(X1,X2,X4),esk62_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_261,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_255, c_0_256])).
cnf(c_0_262,plain,(p2(esk52_2(esk1_0,esk50_1(esk1_0)))|p2(esk50_1(esk1_0))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_257, c_0_79])).
cnf(c_0_263,plain,(p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_238, c_0_121])).
cnf(c_0_264,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_224, c_0_258])).
cnf(c_0_265,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_259, c_0_256])).
cnf(c_0_266,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_79]), c_0_41])).
cnf(c_0_267,plain,(p2(X3)|r1(esk57_3(X1,X2,X3),esk58_3(X1,X2,X3))|p2(X4)|r1(X4,esk61_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_268,plain,(p2(X3)|r1(esk57_3(X1,X2,X3),esk58_3(X1,X2,X3))|p2(X4)|p2(esk61_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_269,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_262]), c_0_82])).
cnf(c_0_270,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_89]), c_0_90])).
cnf(c_0_271,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_264, c_0_232])).
cnf(c_0_272,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_262]), c_0_86])).
cnf(c_0_273,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_266, c_0_256])).
cnf(c_0_274,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_79]), c_0_41])).
cnf(c_0_275,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_79]), c_0_41])).
cnf(c_0_276,plain,(p2(X3)|r1(esk57_3(X1,X2,X3),esk58_3(X1,X2,X3))|p2(X4)|r1(esk61_3(X1,X2,X4),esk62_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_277,plain,(p2(X3)|r1(X3,esk57_3(X1,X2,X3))|p2(X4)|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk62_3(X1,X2,X4))|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_278,plain,(p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_89]), c_0_90])).
cnf(c_0_279,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_270, c_0_271])).
cnf(c_0_280,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_89]), c_0_90])).
cnf(c_0_281,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_262]), c_0_92])).
cnf(c_0_282,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_274, c_0_256])).
cnf(c_0_283,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_275, c_0_256])).
cnf(c_0_284,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_79]), c_0_41])).
cnf(c_0_285,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_79]), c_0_41])).
cnf(c_0_286,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_89]), c_0_90])).
cnf(c_0_287,plain,(r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_89]), c_0_253]), c_0_90])).
cnf(c_0_288,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_89]), c_0_90])).
cnf(c_0_289,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_89]), c_0_90])).
cnf(c_0_290,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_262]), c_0_100])).
cnf(c_0_291,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_262]), c_0_101])).
cnf(c_0_292,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_284, c_0_256])).
cnf(c_0_293,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_285, c_0_256])).
cnf(c_0_294,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_286, c_0_287])).
cnf(c_0_295,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_288, c_0_287])).
cnf(c_0_296,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_89]), c_0_90])).
cnf(c_0_297,plain,(p2(X3)|p2(esk57_3(X1,X2,X3))|p2(X4)|r1(X4,esk61_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_298,plain,(p2(X3)|r1(esk57_3(X1,X2,X3),esk58_3(X1,X2,X3))|p2(X4)|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk62_3(X1,X2,X4))|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_299,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_89]), c_0_90])).
cnf(c_0_300,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_89]), c_0_90])).
cnf(c_0_301,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_262]), c_0_111])).
cnf(c_0_302,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_262]), c_0_104])).
cnf(c_0_303,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_294]), c_0_295])).
cnf(c_0_304,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_296, c_0_287])).
cnf(c_0_305,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk61_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_79]), c_0_41])).
cnf(c_0_306,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_79]), c_0_41])).
cnf(c_0_307,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_89]), c_0_90])).
cnf(c_0_308,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_89]), c_0_90])).
cnf(c_0_309,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_89]), c_0_90])).
cnf(c_0_310,plain,(p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_89]), c_0_90])).
cnf(c_0_311,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_303, c_0_304])).
cnf(c_0_312,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_305, c_0_256])).
cnf(c_0_313,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_306, c_0_256])).
cnf(c_0_314,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_307, c_0_287])).
cnf(c_0_315,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_308, c_0_287])).
cnf(c_0_316,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_89]), c_0_90])).
cnf(c_0_317,plain,(p2(X3)|p2(esk57_3(X1,X2,X3))|p2(X4)|r1(esk61_3(X1,X2,X4),esk62_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_318,plain,(p2(X3)|p2(X4)|r1(X4,esk61_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~p2(esk58_3(X1,X2,X3))|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_319,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_310, c_0_311])).
cnf(c_0_320,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk61_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_262]), c_0_124])).
cnf(c_0_321,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_262]), c_0_125])).
cnf(c_0_322,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_314]), c_0_315])).
cnf(c_0_323,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_316, c_0_287])).
cnf(c_0_324,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X3),esk62_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_79]), c_0_41])).
cnf(c_0_325,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_79]), c_0_41])).
cnf(c_0_326,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_89]), c_0_90])).
cnf(c_0_327,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_89]), c_0_90])).
cnf(c_0_328,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_89]), c_0_90])).
cnf(c_0_329,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_322, c_0_323])).
cnf(c_0_330,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_324, c_0_256])).
cnf(c_0_331,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_325, c_0_256])).
cnf(c_0_332,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_326, c_0_287])).
cnf(c_0_333,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_89]), c_0_90])).
cnf(c_0_334,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_328, c_0_329])).
cnf(c_0_335,plain,(p2(X3)|p2(esk57_3(X1,X2,X3))|p2(X4)|p2(esk61_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_336,plain,(p2(X3)|p2(X4)|r1(esk61_3(X1,X2,X4),esk62_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~p2(esk58_3(X1,X2,X3))|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_337,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X3),esk62_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_262]), c_0_148])).
cnf(c_0_338,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_262]), c_0_137])).
cnf(c_0_339,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_105, c_0_332])).
cnf(c_0_340,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_333, c_0_287])).
cnf(c_0_341,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_89]), c_0_90])).
cnf(c_0_342,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk61_3(esk1_0,esk50_1(esk1_0),X2))|p2(X3)|p2(X2)|p2(X1)|r1(esk1_0,esk13_0)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_79]), c_0_41])).
cnf(c_0_343,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_79]), c_0_41])).
cnf(c_0_344,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_89]), c_0_90])).
cnf(c_0_345,plain,(p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_89]), c_0_90])).
cnf(c_0_346,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_339, c_0_340])).
cnf(c_0_347,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_341, c_0_287])).
cnf(c_0_348,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|p2(X3)|r1(esk1_0,esk13_0)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_342, c_0_256])).
cnf(c_0_349,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_343, c_0_256])).
cnf(c_0_350,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_89]), c_0_90])).
cnf(c_0_351,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_89]), c_0_90])).
cnf(c_0_352,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_346, c_0_347])).
cnf(c_0_353,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk61_3(esk1_0,esk50_1(esk1_0),X2))|p2(X3)|p2(X2)|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_262]), c_0_159])).
cnf(c_0_354,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_262]), c_0_161])).
cnf(c_0_355,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_350, c_0_287])).
cnf(c_0_356,plain,(p2(X3)|p2(esk57_3(X1,X2,X3))|p2(X4)|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk62_3(X1,X2,X4))|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_357,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_351, c_0_352])).
cnf(c_0_358,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_89]), c_0_90])).
cnf(c_0_359,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_89]), c_0_90])).
cnf(c_0_360,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_339, c_0_355])).
cnf(c_0_361,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_79]), c_0_41])).
cnf(c_0_362,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_357, c_0_287])).
cnf(c_0_363,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_358, c_0_89]), c_0_90])).
cnf(c_0_364,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_89]), c_0_90])).
cnf(c_0_365,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_360, c_0_347])).
cnf(c_0_366,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|p2(X3)|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_361, c_0_256])).
cnf(c_0_367,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_105, c_0_362])).
cnf(c_0_368,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_363, c_0_287])).
cnf(c_0_369,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_364, c_0_365])).
cnf(c_0_370,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_262]), c_0_189])).
cnf(c_0_371,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_367, c_0_368])).
cnf(c_0_372,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_369, c_0_287])).
cnf(c_0_373,plain,(p2(X3)|p2(X4)|p2(esk61_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~p2(esk58_3(X1,X2,X3))|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_374,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_370, c_0_89]), c_0_90])).
cnf(c_0_375,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_371, c_0_372])).
cnf(c_0_376,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_79]), c_0_41])).
cnf(c_0_377,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X2)|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_374, c_0_375])).
cnf(c_0_378,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_376, c_0_256])).
cnf(c_0_379,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_89]), c_0_90])).
cnf(c_0_380,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_262]), c_0_216])).
cnf(c_0_381,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_379, c_0_287])).
cnf(c_0_382,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_89]), c_0_90])).
cnf(c_0_383,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_339, c_0_381])).
cnf(c_0_384,plain,(p2(X3)|p2(X4)|p2(X6)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~r1(X2,X3)|~p2(esk58_3(X1,X2,X3))|~r1(X2,X4)|~p2(esk62_3(X1,X2,X4))|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_385,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_89]), c_0_90])).
cnf(c_0_386,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_383, c_0_347])).
cnf(c_0_387,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_79]), c_0_41])).
cnf(c_0_388,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_385, c_0_386])).
cnf(c_0_389,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_387, c_0_256])).
cnf(c_0_390,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_388, c_0_287])).
cnf(c_0_391,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_262]), c_0_238])).
cnf(c_0_392,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_367, c_0_390])).
cnf(c_0_393,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_89]), c_0_90])).
cnf(c_0_394,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_392, c_0_372])).
cnf(c_0_395,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_393, c_0_394])).
cnf(c_0_396,plain,(p2(X2)|~epred1_1(X1)|~r1(esk49_1(X1),X2)|~p2(esk53_2(X1,X2))|~p2(esk49_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_397,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_89]), c_0_386]), c_0_90])).
cnf(c_0_398,plain,(p2(X1)|r1(esk1_0,esk13_0)|~p2(esk53_2(esk1_0,X1))|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_76]), c_0_41])).
cnf(c_0_399,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_397, c_0_287])).
cnf(c_0_400,plain,(p2(esk50_1(esk1_0))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_399]), c_0_79])).
cnf(c_0_401,plain,(p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_82, c_0_400])).
cnf(c_0_402,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_86, c_0_400])).
cnf(c_0_403,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_89]), c_0_90])).
cnf(c_0_404,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_89]), c_0_90])).
cnf(c_0_405,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_92, c_0_400])).
cnf(c_0_406,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_403, c_0_89]), c_0_90])).
cnf(c_0_407,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_89]), c_0_90])).
cnf(c_0_408,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_405, c_0_89]), c_0_90])).
cnf(c_0_409,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_100, c_0_400])).
cnf(c_0_410,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_101, c_0_400])).
cnf(c_0_411,plain,(p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_104, c_0_400])).
cnf(c_0_412,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_406]), c_0_407])).
cnf(c_0_413,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_89]), c_0_90])).
cnf(c_0_414,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_89]), c_0_90])).
cnf(c_0_415,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_89]), c_0_90])).
cnf(c_0_416,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_111, c_0_400])).
cnf(c_0_417,plain,(p2(X1)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_411, c_0_89]), c_0_90])).
cnf(c_0_418,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_412, c_0_413])).
cnf(c_0_419,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_89]), c_0_90])).
cnf(c_0_420,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_89]), c_0_90])).
cnf(c_0_421,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_89]), c_0_90])).
cnf(c_0_422,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_417, c_0_418])).
cnf(c_0_423,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_124, c_0_400])).
cnf(c_0_424,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_125, c_0_400])).
cnf(c_0_425,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_419]), c_0_420])).
cnf(c_0_426,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_89]), c_0_90])).
cnf(c_0_427,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_89]), c_0_90])).
cnf(c_0_428,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_423, c_0_89]), c_0_90])).
cnf(c_0_429,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_424, c_0_89]), c_0_90])).
cnf(c_0_430,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_425, c_0_426])).
cnf(c_0_431,plain,(p2(X1)|p2(X2)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_137, c_0_400])).
cnf(c_0_432,plain,(p2(X1)|r1(esk1_0,esk13_0)|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_105, c_0_427])).
cnf(c_0_433,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_89]), c_0_90])).
cnf(c_0_434,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_429, c_0_430])).
cnf(c_0_435,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_89]), c_0_90])).
cnf(c_0_436,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_148, c_0_400])).
cnf(c_0_437,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_432, c_0_433])).
cnf(c_0_438,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_434, c_0_89]), c_0_90])).
cnf(c_0_439,plain,(r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_89]), c_0_90])).
cnf(c_0_440,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_159, c_0_400])).
cnf(c_0_441,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_436, c_0_89]), c_0_90])).
cnf(c_0_442,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_161, c_0_400])).
cnf(c_0_443,plain,(r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_437, c_0_438]), c_0_439])).
cnf(c_0_444,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_89]), c_0_90])).
cnf(c_0_445,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_441, c_0_89]), c_0_90])).
cnf(c_0_446,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_89]), c_0_90])).
cnf(c_0_447,plain,(p2(X1)|r1(esk1_0,esk13_0)|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_105, c_0_443])).
cnf(c_0_448,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_89]), c_0_90])).
cnf(c_0_449,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_432, c_0_445])).
cnf(c_0_450,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_446, c_0_89]), c_0_90])).
cnf(c_0_451,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_189, c_0_400])).
cnf(c_0_452,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_447, c_0_448])).
cnf(c_0_453,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_438]), c_0_450])).
cnf(c_0_454,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_89]), c_0_90])).
cnf(c_0_455,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_452, c_0_453])).
cnf(c_0_456,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk1_0,esk13_0)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_454, c_0_455])).
cnf(c_0_457,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_216, c_0_400])).
cnf(c_0_458,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_89]), c_0_90])).
cnf(c_0_459,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_457, c_0_89]), c_0_90])).
cnf(c_0_460,plain,(p2(X1)|r1(esk1_0,esk13_0)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_432, c_0_458])).
cnf(c_0_461,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_459, c_0_89]), c_0_90])).
cnf(c_0_462,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_460, c_0_438])).
cnf(c_0_463,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_461, c_0_462])).
cnf(c_0_464,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_238, c_0_400])).
cnf(c_0_465,plain,(p2(X1)|r1(esk1_0,esk13_0)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_447, c_0_463])).
cnf(c_0_466,plain,(p2(X1)|r1(esk1_0,esk13_0)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_89]), c_0_90])).
cnf(c_0_467,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk1_0,esk13_0)), inference(spm,[status(thm)],[c_0_465, c_0_453])).
cnf(c_0_468,plain,(p2(X1)|r1(esk1_0,esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_466, c_0_467])).
cnf(c_0_469,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(X2,esk11_2(X1,X2))|p2(X4)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(X3)|~r1(X3,X4)|~r1(esk13_0,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_470,plain,(r1(esk1_0,esk13_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_89]), c_0_462]), c_0_90])).
cnf(c_0_471,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|p2(esk11_2(X1,X2))|p2(X4)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(X3)|~r1(X3,X4)|~r1(esk13_0,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_472,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|p2(X2)|r1(X2,esk11_2(esk5_0,X2))|~p2(X3)|~r1(esk13_0,X3)|~r1(esk5_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_469, c_0_16])).
cnf(c_0_473,negated_conjecture,(p2(esk13_0)|r1(esk13_0,esk6_1(esk13_0))), inference(spm,[status(thm)],[c_0_35, c_0_470])).
cnf(c_0_474,negated_conjecture,(p2(esk6_1(esk13_0))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_42, c_0_470])).
cnf(c_0_475,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_2(esk5_0,X1))|p2(X2)|p2(X1)|~p2(X3)|~r1(esk13_0,X3)|~r1(esk5_0,X1)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_471, c_0_16])).
cnf(c_0_476,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(esk11_2(X1,X2),esk12_2(X1,X2))|p2(X4)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(X3)|~r1(X3,X4)|~r1(esk13_0,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_477,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|p2(X1)|p2(X2)|r1(X1,esk11_2(esk5_0,X1))|~r1(esk6_1(esk13_0),X2)|~r1(esk5_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_472, c_0_473]), c_0_474])).
cnf(c_0_478,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_2(esk5_0,X1))|p2(esk13_0)|p2(X1)|p2(X2)|~r1(esk6_1(esk13_0),X2)|~r1(esk5_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_475, c_0_473]), c_0_474])).
cnf(c_0_479,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|p2(X2)|r1(esk11_2(esk5_0,X2),esk12_2(esk5_0,X2))|~p2(X3)|~r1(esk13_0,X3)|~r1(esk5_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_476, c_0_16])).
cnf(c_0_480,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|p2(X4)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk12_2(X1,X2))|~p2(X3)|~r1(X3,X4)|~r1(esk13_0,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_481,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|p2(X1)|r1(esk48_1(esk5_0),esk11_2(esk5_0,esk48_1(esk5_0)))|~r1(esk6_1(esk13_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_477, c_0_24]), c_0_25])).
cnf(c_0_482,negated_conjecture,(p2(esk13_0)|r1(esk6_1(esk13_0),esk7_1(esk13_0))), inference(spm,[status(thm)],[c_0_48, c_0_470])).
cnf(c_0_483,negated_conjecture,(p2(esk13_0)|~p2(esk7_1(esk13_0))), inference(spm,[status(thm)],[c_0_49, c_0_470])).
cnf(c_0_484,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_2(esk5_0,esk48_1(esk5_0)))|p2(esk13_0)|p2(X1)|~r1(esk6_1(esk13_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_478, c_0_24]), c_0_25])).
cnf(c_0_485,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk11_2(esk5_0,X1),esk12_2(esk5_0,X1))|~r1(esk6_1(esk13_0),X2)|~r1(esk5_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_479, c_0_473]), c_0_474])).
cnf(c_0_486,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|p2(X2)|~p2(esk12_2(esk5_0,X2))|~p2(X3)|~r1(esk13_0,X3)|~r1(esk5_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_480, c_0_16])).
cnf(c_0_487,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|r1(esk48_1(esk5_0),esk11_2(esk5_0,esk48_1(esk5_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_482]), c_0_483])).
cnf(c_0_488,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_2(esk5_0,esk48_1(esk5_0)))|p2(esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_484, c_0_482]), c_0_483])).
cnf(c_0_489,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|p2(X1)|r1(esk11_2(esk5_0,esk48_1(esk5_0)),esk12_2(esk5_0,esk48_1(esk5_0)))|~r1(esk6_1(esk13_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_24]), c_0_25])).
cnf(c_0_490,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk12_2(esk5_0,X1))|~r1(esk6_1(esk13_0),X2)|~r1(esk5_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_486, c_0_473]), c_0_474])).
cnf(c_0_491,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|p2(X1)|~r1(esk11_2(esk5_0,esk48_1(esk5_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_487]), c_0_488])).
cnf(c_0_492,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|r1(esk11_2(esk5_0,esk48_1(esk5_0)),esk12_2(esk5_0,esk48_1(esk5_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_482]), c_0_483])).
cnf(c_0_493,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|p2(X1)|~p2(esk12_2(esk5_0,esk48_1(esk5_0)))|~r1(esk6_1(esk13_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_490, c_0_24]), c_0_25])).
cnf(c_0_494,negated_conjecture,(epred1_1(esk1_0)|p2(esk12_2(esk5_0,esk48_1(esk5_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_491, c_0_492])).
cnf(c_0_495,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)|p2(X1)|~r1(esk6_1(esk13_0),X1)), inference(spm,[status(thm)],[c_0_493, c_0_494])).
cnf(c_0_496,negated_conjecture,(epred1_1(esk1_0)|p2(esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_495, c_0_482]), c_0_483])).
cnf(c_0_497,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|r1(esk49_1(esk1_0),esk6_1(esk49_1(esk1_0)))), inference(spm,[status(thm)],[c_0_40, c_0_496])).
cnf(c_0_498,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|p2(X2)|r1(X2,esk52_2(esk1_0,X2))|~r1(esk6_1(esk49_1(esk1_0)),X1)|~r1(esk49_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_497]), c_0_45]), c_0_496])).
cnf(c_0_499,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|r1(esk49_1(esk1_0),esk48_1(esk49_1(esk1_0)))), inference(spm,[status(thm)],[c_0_46, c_0_496])).
cnf(c_0_500,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|~p2(esk48_1(esk49_1(esk1_0)))), inference(spm,[status(thm)],[c_0_47, c_0_496])).
cnf(c_0_501,plain,(p2(esk52_2(esk1_0,X1))|p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X2)|p2(X1)|~r1(esk6_1(esk49_1(esk1_0)),X2)|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_497]), c_0_45]), c_0_496])).
cnf(c_0_502,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|r1(esk48_1(esk49_1(esk1_0)),esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_498, c_0_499]), c_0_500])).
cnf(c_0_503,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|r1(esk6_1(esk49_1(esk1_0)),esk7_1(esk49_1(esk1_0)))), inference(spm,[status(thm)],[c_0_54, c_0_496])).
cnf(c_0_504,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|~p2(esk7_1(esk49_1(esk1_0)))), inference(spm,[status(thm)],[c_0_55, c_0_496])).
cnf(c_0_505,plain,(p2(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_501, c_0_499]), c_0_500])).
cnf(c_0_506,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,X2),esk53_2(esk1_0,X2))|~r1(esk6_1(esk49_1(esk1_0)),X1)|~r1(esk49_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_497]), c_0_45]), c_0_496])).
cnf(c_0_507,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|~p2(X2)|~r1(esk48_1(esk49_1(esk1_0)),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_58, c_0_496])).
cnf(c_0_508,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|r1(esk48_1(esk49_1(esk1_0)),esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_503]), c_0_504])).
cnf(c_0_509,plain,(p2(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|p2(esk49_1(esk1_0))|p2(esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_505, c_0_503]), c_0_504])).
cnf(c_0_510,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))),esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_506, c_0_499]), c_0_500])).
cnf(c_0_511,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk53_2(esk1_0,X2))|~r1(esk6_1(esk49_1(esk1_0)),X1)|~r1(esk49_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_497]), c_0_45]), c_0_496])).
cnf(c_0_512,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|~r1(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_507, c_0_508]), c_0_509])).
cnf(c_0_513,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|r1(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))),esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_510, c_0_503]), c_0_504])).
cnf(c_0_514,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|~p2(esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_511, c_0_499]), c_0_500])).
cnf(c_0_515,plain,(p2(esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|p2(esk49_1(esk1_0))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_512, c_0_513])).
cnf(c_0_516,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)|p2(X1)|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_514, c_0_515])).
cnf(c_0_517,plain,(p2(esk49_1(esk1_0))|p2(esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_516, c_0_503]), c_0_504])).
cnf(c_0_518,plain,(p2(esk13_0)|r1(esk49_1(esk1_0),esk50_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_517]), c_0_496])).
cnf(c_0_519,plain,(p2(esk13_0)|p2(X1)|r1(X1,esk52_2(esk1_0,X1))|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_517]), c_0_496])).
cnf(c_0_520,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(X1,esk61_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_518]), c_0_496])).
cnf(c_0_521,plain,(p2(esk50_1(esk1_0))|p2(esk13_0)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_519, c_0_518])).
cnf(c_0_522,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_518]), c_0_496])).
cnf(c_0_523,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_520, c_0_521])).
cnf(c_0_524,plain,(p2(esk13_0)|r1(esk50_1(esk1_0),esk51_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_517]), c_0_496])).
cnf(c_0_525,plain,(p2(esk13_0)|~p2(esk51_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_517]), c_0_496])).
cnf(c_0_526,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_522, c_0_521])).
cnf(c_0_527,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X1),esk62_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_518]), c_0_496])).
cnf(c_0_528,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_523, c_0_524]), c_0_525])).
cnf(c_0_529,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_526, c_0_524]), c_0_525])).
cnf(c_0_530,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_527, c_0_521])).
cnf(c_0_531,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(X1,esk61_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_518]), c_0_496])).
cnf(c_0_532,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_518]), c_0_496])).
cnf(c_0_533,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_518]), c_0_496])).
cnf(c_0_534,plain,(p2(esk13_0)|p2(X1)|~p2(X2)|~r1(esk51_1(esk1_0),X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_517]), c_0_496])).
cnf(c_0_535,plain,(p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_528, c_0_524]), c_0_525])).
cnf(c_0_536,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_529, c_0_524]), c_0_525])).
cnf(c_0_537,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_530, c_0_524]), c_0_525])).
cnf(c_0_538,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_531, c_0_521])).
cnf(c_0_539,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_532, c_0_521])).
cnf(c_0_540,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk61_3(esk1_0,esk50_1(esk1_0),X1),esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_518]), c_0_496])).
cnf(c_0_541,plain,(p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,X1),esk53_2(esk1_0,X1))|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_517]), c_0_496])).
cnf(c_0_542,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_533, c_0_521])).
cnf(c_0_543,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_535]), c_0_536])).
cnf(c_0_544,plain,(p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_537, c_0_524]), c_0_525])).
cnf(c_0_545,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_538, c_0_524]), c_0_525])).
cnf(c_0_546,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_539, c_0_524]), c_0_525])).
cnf(c_0_547,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_540, c_0_521])).
cnf(c_0_548,plain,(p2(esk50_1(esk1_0))|p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_541, c_0_518])).
cnf(c_0_549,plain,(p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_542, c_0_524]), c_0_525])).
cnf(c_0_550,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_543, c_0_544])).
cnf(c_0_551,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_518]), c_0_496])).
cnf(c_0_552,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_518]), c_0_496])).
cnf(c_0_553,plain,(p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_545, c_0_524]), c_0_525])).
cnf(c_0_554,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_546, c_0_524]), c_0_525])).
cnf(c_0_555,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_547, c_0_524]), c_0_525])).
cnf(c_0_556,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_520, c_0_548])).
cnf(c_0_557,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_522, c_0_548])).
cnf(c_0_558,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_549, c_0_550])).
cnf(c_0_559,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_551, c_0_521])).
cnf(c_0_560,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_552, c_0_521])).
cnf(c_0_561,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_553]), c_0_554])).
cnf(c_0_562,plain,(p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_524]), c_0_525])).
cnf(c_0_563,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(X1,esk61_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_518]), c_0_496])).
cnf(c_0_564,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_556, c_0_524]), c_0_525])).
cnf(c_0_565,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_557, c_0_524]), c_0_525])).
cnf(c_0_566,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_527, c_0_548])).
cnf(c_0_567,plain,(p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_558, c_0_524]), c_0_525])).
cnf(c_0_568,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_559, c_0_524]), c_0_525])).
cnf(c_0_569,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_560, c_0_524]), c_0_525])).
cnf(c_0_570,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_561, c_0_562])).
cnf(c_0_571,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_563, c_0_521])).
cnf(c_0_572,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_518]), c_0_496])).
cnf(c_0_573,plain,(p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_564, c_0_524]), c_0_525])).
cnf(c_0_574,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_565, c_0_524]), c_0_525])).
cnf(c_0_575,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_524]), c_0_525])).
cnf(c_0_576,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_531, c_0_548])).
cnf(c_0_577,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_532, c_0_548])).
cnf(c_0_578,plain,(p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_534, c_0_567])).
cnf(c_0_579,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_568, c_0_524]), c_0_525])).
cnf(c_0_580,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_569, c_0_570])).
cnf(c_0_581,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_571, c_0_524]), c_0_525])).
cnf(c_0_582,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk61_3(esk1_0,esk50_1(esk1_0),X2))|p2(esk13_0)|p2(X2)|p2(X1)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_518]), c_0_496])).
cnf(c_0_583,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_572, c_0_521])).
cnf(c_0_584,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X1),esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_518]), c_0_496])).
cnf(c_0_585,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_533, c_0_548])).
cnf(c_0_586,plain,(p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_573]), c_0_574])).
cnf(c_0_587,plain,(p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_575, c_0_524]), c_0_525])).
cnf(c_0_588,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_576, c_0_524]), c_0_525])).
cnf(c_0_589,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_577, c_0_524]), c_0_525])).
cnf(c_0_590,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_540, c_0_548])).
cnf(c_0_591,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_578, c_0_579])).
cnf(c_0_592,plain,(p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_580, c_0_524]), c_0_525])).
cnf(c_0_593,plain,(p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_581, c_0_524]), c_0_525])).
cnf(c_0_594,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_582, c_0_521])).
cnf(c_0_595,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_524]), c_0_525])).
cnf(c_0_596,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_584, c_0_521])).
cnf(c_0_597,plain,(p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_524]), c_0_525])).
cnf(c_0_598,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_586, c_0_587])).
cnf(c_0_599,plain,(p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_588, c_0_524]), c_0_525])).
cnf(c_0_600,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_589, c_0_524]), c_0_525])).
cnf(c_0_601,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_590, c_0_524]), c_0_525])).
cnf(c_0_602,plain,(p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_591, c_0_592]), c_0_593])).
cnf(c_0_603,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_594, c_0_524]), c_0_525])).
cnf(c_0_604,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_595, c_0_524]), c_0_525])).
cnf(c_0_605,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_596, c_0_524]), c_0_525])).
cnf(c_0_606,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_597, c_0_598])).
cnf(c_0_607,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_551, c_0_548])).
cnf(c_0_608,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_552, c_0_548])).
cnf(c_0_609,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_599]), c_0_600])).
cnf(c_0_610,plain,(p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_601, c_0_524]), c_0_525])).
cnf(c_0_611,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_518]), c_0_496])).
cnf(c_0_612,plain,(p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_534, c_0_602])).
cnf(c_0_613,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_524]), c_0_525])).
cnf(c_0_614,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_578, c_0_604])).
cnf(c_0_615,plain,(p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_605, c_0_524]), c_0_525])).
cnf(c_0_616,plain,(p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_606, c_0_524]), c_0_525])).
cnf(c_0_617,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_607, c_0_524]), c_0_525])).
cnf(c_0_618,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_608, c_0_524]), c_0_525])).
cnf(c_0_619,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_609, c_0_610])).
cnf(c_0_620,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_563, c_0_548])).
cnf(c_0_621,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_611, c_0_521])).
cnf(c_0_622,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_612, c_0_613])).
cnf(c_0_623,plain,(p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_614, c_0_592]), c_0_615])).
cnf(c_0_624,plain,(p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_534, c_0_616])).
cnf(c_0_625,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_617, c_0_524]), c_0_525])).
cnf(c_0_626,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_618, c_0_619])).
cnf(c_0_627,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_620, c_0_524]), c_0_525])).
cnf(c_0_628,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_572, c_0_548])).
cnf(c_0_629,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_621, c_0_524]), c_0_525])).
cnf(c_0_630,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_622, c_0_623])).
cnf(c_0_631,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_624, c_0_625])).
cnf(c_0_632,plain,(p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_626, c_0_524]), c_0_525])).
cnf(c_0_633,plain,(p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_524]), c_0_525])).
cnf(c_0_634,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_582, c_0_548])).
cnf(c_0_635,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_628, c_0_524]), c_0_525])).
cnf(c_0_636,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_584, c_0_548])).
cnf(c_0_637,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_518]), c_0_496])).
cnf(c_0_638,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_629, c_0_630])).
cnf(c_0_639,plain,(p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_631, c_0_632]), c_0_633])).
cnf(c_0_640,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_634, c_0_524]), c_0_525])).
cnf(c_0_641,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_635, c_0_524]), c_0_525])).
cnf(c_0_642,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_636, c_0_524]), c_0_525])).
cnf(c_0_643,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_637, c_0_521])).
cnf(c_0_644,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_638, c_0_524]), c_0_525])).
cnf(c_0_645,plain,(p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_534, c_0_639])).
cnf(c_0_646,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_640, c_0_524]), c_0_525])).
cnf(c_0_647,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_624, c_0_641])).
cnf(c_0_648,plain,(p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_642, c_0_524]), c_0_525])).
cnf(c_0_649,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_643, c_0_524]), c_0_525])).
cnf(c_0_650,plain,(p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_578, c_0_644])).
cnf(c_0_651,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_611, c_0_548])).
cnf(c_0_652,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_645, c_0_646])).
cnf(c_0_653,plain,(p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_647, c_0_632]), c_0_648])).
cnf(c_0_654,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_649, c_0_524]), c_0_525])).
cnf(c_0_655,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_650, c_0_592])).
cnf(c_0_656,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_651, c_0_524]), c_0_525])).
cnf(c_0_657,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_652, c_0_653])).
cnf(c_0_658,plain,(p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_518]), c_0_496])).
cnf(c_0_659,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_654, c_0_655])).
cnf(c_0_660,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_656, c_0_657])).
cnf(c_0_661,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_658, c_0_521])).
cnf(c_0_662,plain,(p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_612, c_0_659])).
cnf(c_0_663,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_637, c_0_548])).
cnf(c_0_664,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_660, c_0_524]), c_0_525])).
cnf(c_0_665,plain,(p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_661, c_0_524]), c_0_525])).
cnf(c_0_666,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_662, c_0_623])).
cnf(c_0_667,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_663, c_0_524]), c_0_525])).
cnf(c_0_668,plain,(p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_624, c_0_664])).
cnf(c_0_669,plain,(p2(esk13_0)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_665, c_0_666])).
cnf(c_0_670,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_667, c_0_524]), c_0_525])).
cnf(c_0_671,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_668, c_0_632])).
cnf(c_0_672,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_518]), c_0_496])).
cnf(c_0_673,plain,(p2(esk13_0)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_669, c_0_524]), c_0_655]), c_0_525])).
cnf(c_0_674,plain,(p2(esk52_2(esk1_0,X1))|p2(esk13_0)|p2(X1)|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_517]), c_0_496])).
cnf(c_0_675,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_670, c_0_671])).
cnf(c_0_676,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_518]), c_0_496])).
cnf(c_0_677,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_672, c_0_673])).
cnf(c_0_678,plain,(p2(esk52_2(esk1_0,esk50_1(esk1_0)))|p2(esk50_1(esk1_0))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_674, c_0_518])).
cnf(c_0_679,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_658, c_0_548])).
cnf(c_0_680,plain,(p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_645, c_0_675])).
cnf(c_0_681,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_676, c_0_673])).
cnf(c_0_682,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_518]), c_0_496])).
cnf(c_0_683,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_677, c_0_678]), c_0_520])).
cnf(c_0_684,plain,(p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_679, c_0_524]), c_0_525])).
cnf(c_0_685,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_680, c_0_653])).
cnf(c_0_686,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_681, c_0_678]), c_0_522])).
cnf(c_0_687,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_682, c_0_673])).
cnf(c_0_688,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_518]), c_0_496])).
cnf(c_0_689,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_518]), c_0_496])).
cnf(c_0_690,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_683, c_0_524]), c_0_525])).
cnf(c_0_691,plain,(p2(esk13_0)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_684, c_0_685])).
cnf(c_0_692,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_686, c_0_524]), c_0_525])).
cnf(c_0_693,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_687, c_0_678]), c_0_527])).
cnf(c_0_694,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_688, c_0_673])).
cnf(c_0_695,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_689, c_0_673])).
cnf(c_0_696,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_518]), c_0_496])).
cnf(c_0_697,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_518]), c_0_496])).
cnf(c_0_698,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_524]), c_0_525])).
cnf(c_0_699,plain,(p2(esk13_0)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_691, c_0_524]), c_0_671]), c_0_525])).
cnf(c_0_700,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_524]), c_0_525])).
cnf(c_0_701,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_693, c_0_524]), c_0_525])).
cnf(c_0_702,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_694, c_0_678]), c_0_531])).
cnf(c_0_703,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_678]), c_0_532])).
cnf(c_0_704,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_696, c_0_673])).
cnf(c_0_705,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_697, c_0_673])).
cnf(c_0_706,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_698, c_0_699])).
cnf(c_0_707,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_700, c_0_699])).
cnf(c_0_708,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_701, c_0_524]), c_0_525])).
cnf(c_0_709,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_524]), c_0_525])).
cnf(c_0_710,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_703, c_0_524]), c_0_525])).
cnf(c_0_711,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_704, c_0_678]), c_0_540])).
cnf(c_0_712,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_705, c_0_678]), c_0_533])).
cnf(c_0_713,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_706]), c_0_707])).
cnf(c_0_714,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_708, c_0_699])).
cnf(c_0_715,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk61_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_518]), c_0_496])).
cnf(c_0_716,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_518]), c_0_496])).
cnf(c_0_717,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_709, c_0_524]), c_0_525])).
cnf(c_0_718,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_710, c_0_524]), c_0_525])).
cnf(c_0_719,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_711, c_0_524]), c_0_525])).
cnf(c_0_720,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_712, c_0_524]), c_0_525])).
cnf(c_0_721,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_713, c_0_714])).
cnf(c_0_722,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_715, c_0_673])).
cnf(c_0_723,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_716, c_0_673])).
cnf(c_0_724,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_717, c_0_699])).
cnf(c_0_725,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_718, c_0_699])).
cnf(c_0_726,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_719, c_0_524]), c_0_525])).
cnf(c_0_727,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_720, c_0_721])).
cnf(c_0_728,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk61_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_722, c_0_678]), c_0_551])).
cnf(c_0_729,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_723, c_0_678]), c_0_552])).
cnf(c_0_730,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_724]), c_0_725])).
cnf(c_0_731,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_726, c_0_699])).
cnf(c_0_732,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X3)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X3),esk62_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_518]), c_0_496])).
cnf(c_0_733,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_518]), c_0_496])).
cnf(c_0_734,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_727, c_0_524]), c_0_525])).
cnf(c_0_735,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_728, c_0_524]), c_0_525])).
cnf(c_0_736,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_729, c_0_524]), c_0_525])).
cnf(c_0_737,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_730, c_0_731])).
cnf(c_0_738,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_732, c_0_673])).
cnf(c_0_739,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_733, c_0_673])).
cnf(c_0_740,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_734, c_0_699])).
cnf(c_0_741,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_735, c_0_524]), c_0_525])).
cnf(c_0_742,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_736, c_0_737])).
cnf(c_0_743,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X3)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X3),esk62_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_738, c_0_678]), c_0_572])).
cnf(c_0_744,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_739, c_0_678]), c_0_563])).
cnf(c_0_745,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_534, c_0_740])).
cnf(c_0_746,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_741, c_0_699])).
cnf(c_0_747,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_742, c_0_524]), c_0_525])).
cnf(c_0_748,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk61_3(esk1_0,esk50_1(esk1_0),X2))|p2(esk13_0)|p2(X3)|p2(X2)|p2(X1)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_518]), c_0_496])).
cnf(c_0_749,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_518]), c_0_496])).
cnf(c_0_750,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_743, c_0_524]), c_0_525])).
cnf(c_0_751,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_744, c_0_524]), c_0_525])).
cnf(c_0_752,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_745, c_0_746])).
cnf(c_0_753,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_747, c_0_699])).
cnf(c_0_754,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_748, c_0_673])).
cnf(c_0_755,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_749, c_0_673])).
cnf(c_0_756,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_750, c_0_524]), c_0_525])).
cnf(c_0_757,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_751, c_0_524]), c_0_525])).
cnf(c_0_758,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_752, c_0_753])).
cnf(c_0_759,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk61_3(esk1_0,esk50_1(esk1_0),X2))|p2(esk13_0)|p2(X3)|p2(X2)|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_754, c_0_678]), c_0_582])).
cnf(c_0_760,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_755, c_0_678]), c_0_584])).
cnf(c_0_761,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_756, c_0_699])).
cnf(c_0_762,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_757, c_0_758])).
cnf(c_0_763,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_759, c_0_524]), c_0_525])).
cnf(c_0_764,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_760, c_0_524]), c_0_525])).
cnf(c_0_765,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_745, c_0_761])).
cnf(c_0_766,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X3)|p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_518]), c_0_496])).
cnf(c_0_767,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_762, c_0_699])).
cnf(c_0_768,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_763, c_0_524]), c_0_525])).
cnf(c_0_769,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_524]), c_0_525])).
cnf(c_0_770,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_765, c_0_753])).
cnf(c_0_771,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_766, c_0_673])).
cnf(c_0_772,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_534, c_0_767])).
cnf(c_0_773,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_768, c_0_699])).
cnf(c_0_774,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_769, c_0_770])).
cnf(c_0_775,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X3)|p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_771, c_0_678]), c_0_611])).
cnf(c_0_776,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_772, c_0_773])).
cnf(c_0_777,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_774, c_0_699])).
cnf(c_0_778,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_775, c_0_524]), c_0_525])).
cnf(c_0_779,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_776, c_0_777])).
cnf(c_0_780,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_518]), c_0_496])).
cnf(c_0_781,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X2)|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_778, c_0_779])).
cnf(c_0_782,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_780, c_0_673])).
cnf(c_0_783,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_781, c_0_524]), c_0_525])).
cnf(c_0_784,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_782, c_0_678]), c_0_637])).
cnf(c_0_785,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_783, c_0_699])).
cnf(c_0_786,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_784, c_0_524]), c_0_525])).
cnf(c_0_787,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_745, c_0_785])).
cnf(c_0_788,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_786, c_0_524]), c_0_525])).
cnf(c_0_789,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_787, c_0_753])).
cnf(c_0_790,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_518]), c_0_496])).
cnf(c_0_791,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_788, c_0_789])).
cnf(c_0_792,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_790, c_0_673])).
cnf(c_0_793,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_791, c_0_699])).
cnf(c_0_794,plain,(p2(esk13_0)|p2(X1)|p2(X2)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_792, c_0_678]), c_0_658])).
cnf(c_0_795,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_772, c_0_793])).
cnf(c_0_796,plain,(p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_524]), c_0_525])).
cnf(c_0_797,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_795, c_0_777])).
cnf(c_0_798,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_796, c_0_797])).
cnf(c_0_799,plain,(p2(esk13_0)|p2(X1)|~p2(esk53_2(esk1_0,X1))|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_517]), c_0_496])).
cnf(c_0_800,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_798, c_0_524]), c_0_789]), c_0_525])).
cnf(c_0_801,plain,(p2(esk50_1(esk1_0))|p2(esk13_0)|~p2(esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_799, c_0_518])).
cnf(c_0_802,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_800, c_0_699])).
cnf(c_0_803,plain,(p2(esk50_1(esk1_0))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_801, c_0_802])).
cnf(c_0_804,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_520, c_0_803])).
cnf(c_0_805,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_522, c_0_803])).
cnf(c_0_806,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_804, c_0_524]), c_0_525])).
cnf(c_0_807,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_805, c_0_524]), c_0_525])).
cnf(c_0_808,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_527, c_0_803])).
cnf(c_0_809,plain,(p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_806, c_0_524]), c_0_525])).
cnf(c_0_810,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_807, c_0_524]), c_0_525])).
cnf(c_0_811,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_808, c_0_524]), c_0_525])).
cnf(c_0_812,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_531, c_0_803])).
cnf(c_0_813,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_532, c_0_803])).
cnf(c_0_814,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_533, c_0_803])).
cnf(c_0_815,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_809]), c_0_810])).
cnf(c_0_816,plain,(p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_811, c_0_524]), c_0_525])).
cnf(c_0_817,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_812, c_0_524]), c_0_525])).
cnf(c_0_818,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_813, c_0_524]), c_0_525])).
cnf(c_0_819,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_540, c_0_803])).
cnf(c_0_820,plain,(p2(esk13_0)|p2(X1)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_814, c_0_524]), c_0_525])).
cnf(c_0_821,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_815, c_0_816])).
cnf(c_0_822,plain,(p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_817, c_0_524]), c_0_525])).
cnf(c_0_823,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_818, c_0_524]), c_0_525])).
cnf(c_0_824,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_819, c_0_524]), c_0_525])).
cnf(c_0_825,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_820, c_0_821])).
cnf(c_0_826,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_551, c_0_803])).
cnf(c_0_827,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_552, c_0_803])).
cnf(c_0_828,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_822]), c_0_823])).
cnf(c_0_829,plain,(p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_824, c_0_524]), c_0_525])).
cnf(c_0_830,plain,(p2(esk13_0)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_825, c_0_524]), c_0_525])).
cnf(c_0_831,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_826, c_0_524]), c_0_525])).
cnf(c_0_832,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_827, c_0_524]), c_0_525])).
cnf(c_0_833,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_828, c_0_829])).
cnf(c_0_834,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_563, c_0_803])).
cnf(c_0_835,plain,(p2(esk13_0)|p2(X1)|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_534, c_0_830])).
cnf(c_0_836,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_831, c_0_524]), c_0_525])).
cnf(c_0_837,plain,(p2(esk13_0)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_832, c_0_833])).
cnf(c_0_838,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_834, c_0_524]), c_0_525])).
cnf(c_0_839,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_572, c_0_803])).
cnf(c_0_840,plain,(p2(esk13_0)|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_835, c_0_836])).
cnf(c_0_841,plain,(p2(esk13_0)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_837, c_0_524]), c_0_525])).
cnf(c_0_842,plain,(p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_838, c_0_524]), c_0_525])).
cnf(c_0_843,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(esk13_0)|p2(X2)|p2(X1)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_582, c_0_803])).
cnf(c_0_844,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_839, c_0_524]), c_0_525])).
cnf(c_0_845,plain,(p2(esk13_0)|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_584, c_0_803])).
cnf(c_0_846,plain,(p2(esk13_0)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_840, c_0_841]), c_0_842])).
cnf(c_0_847,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_843, c_0_524]), c_0_525])).
cnf(c_0_848,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_844, c_0_524]), c_0_525])).
cnf(c_0_849,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_845, c_0_524]), c_0_525])).
cnf(c_0_850,plain,(p2(esk13_0)|p2(X1)|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_534, c_0_846])).
cnf(c_0_851,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_847, c_0_524]), c_0_525])).
cnf(c_0_852,plain,(p2(esk13_0)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_835, c_0_848])).
cnf(c_0_853,plain,(p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_849, c_0_524]), c_0_525])).
cnf(c_0_854,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_611, c_0_803])).
cnf(c_0_855,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_850, c_0_851])).
cnf(c_0_856,plain,(p2(esk13_0)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_852, c_0_841]), c_0_853])).
cnf(c_0_857,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_854, c_0_524]), c_0_525])).
cnf(c_0_858,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_855, c_0_856])).
cnf(c_0_859,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X1)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_857, c_0_858])).
cnf(c_0_860,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk13_0)|p2(X2)|p2(X1)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_637, c_0_803])).
cnf(c_0_861,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_859, c_0_524]), c_0_525])).
cnf(c_0_862,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|p2(X1)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_860, c_0_524]), c_0_525])).
cnf(c_0_863,plain,(p2(esk13_0)|p2(X1)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_835, c_0_861])).
cnf(c_0_864,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_862, c_0_524]), c_0_525])).
cnf(c_0_865,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_863, c_0_841])).
cnf(c_0_866,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_864, c_0_865])).
cnf(c_0_867,plain,(p2(esk13_0)|p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_658, c_0_803])).
cnf(c_0_868,plain,(p2(esk13_0)|p2(X1)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_850, c_0_866])).
cnf(c_0_869,plain,(p2(esk13_0)|p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_867, c_0_524]), c_0_525])).
cnf(c_0_870,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk13_0)), inference(spm,[status(thm)],[c_0_868, c_0_856])).
cnf(c_0_871,plain,(p2(esk13_0)|p2(X1)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_869, c_0_870])).
cnf(c_0_872,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(X2,esk11_2(X1,X2))|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk13_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_873,plain,(p2(esk13_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_871, c_0_524]), c_0_865]), c_0_525])).
cnf(c_0_874,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|p2(esk11_2(X1,X2))|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk13_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_875,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk11_2(X2,X1))|~r1(esk1_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_872, c_0_873])])).
cnf(c_0_876,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_2(X1,X2))|p2(X2)|~r1(esk1_0,X1)|~r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_874, c_0_873])])).
cnf(c_0_877,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(esk11_2(X1,X2),esk12_2(X1,X2))|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk13_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_878,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk11_2(esk5_0,X1))|~r1(esk5_0,X1)), inference(spm,[status(thm)],[c_0_875, c_0_16])).
cnf(c_0_879,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_2(esk5_0,X1))|p2(X1)|~r1(esk5_0,X1)), inference(spm,[status(thm)],[c_0_876, c_0_16])).
cnf(c_0_880,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk11_2(X2,X1),esk12_2(X2,X1))|~r1(esk1_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_877, c_0_873])])).
cnf(c_0_881,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk12_2(X1,X2))|~p2(esk13_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_882,negated_conjecture,(epred1_1(esk1_0)|r1(esk48_1(esk5_0),esk11_2(esk5_0,esk48_1(esk5_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_878, c_0_24]), c_0_25])).
cnf(c_0_883,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_2(esk5_0,esk48_1(esk5_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_879, c_0_24]), c_0_25])).
cnf(c_0_884,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk11_2(esk5_0,X1),esk12_2(esk5_0,X1))|~r1(esk5_0,X1)), inference(spm,[status(thm)],[c_0_880, c_0_16])).
cnf(c_0_885,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|~p2(esk12_2(X2,X1))|~r1(esk1_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_881, c_0_873])])).
cnf(c_0_886,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|~r1(esk11_2(esk5_0,esk48_1(esk5_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_882]), c_0_883])).
cnf(c_0_887,negated_conjecture,(epred1_1(esk1_0)|r1(esk11_2(esk5_0,esk48_1(esk5_0)),esk12_2(esk5_0,esk48_1(esk5_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_884, c_0_24]), c_0_25])).
cnf(c_0_888,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|~p2(esk12_2(esk5_0,X1))|~r1(esk5_0,X1)), inference(spm,[status(thm)],[c_0_885, c_0_16])).
cnf(c_0_889,negated_conjecture,(epred1_1(esk1_0)|p2(esk12_2(esk5_0,esk48_1(esk5_0)))), inference(spm,[status(thm)],[c_0_886, c_0_887])).
cnf(c_0_890,negated_conjecture,(epred1_1(esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_888, c_0_24]), c_0_25]), c_0_889])).
cnf(c_0_891,plain,(p2(esk49_1(esk1_0))|r1(esk49_1(esk1_0),esk6_1(esk49_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_890])])).
cnf(c_0_892,plain,(p2(esk6_1(esk49_1(esk1_0)))|p2(esk49_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_890])])).
cnf(c_0_893,plain,(p2(esk49_1(esk1_0))|p2(X1)|p2(X2)|r1(X2,esk52_2(esk1_0,X2))|~r1(esk6_1(esk49_1(esk1_0)),X1)|~r1(esk49_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_891]), c_0_890])]), c_0_892])).
cnf(c_0_894,plain,(p2(esk49_1(esk1_0))|r1(esk49_1(esk1_0),esk48_1(esk49_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_890])])).
cnf(c_0_895,plain,(p2(esk49_1(esk1_0))|~p2(esk48_1(esk49_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_890])])).
cnf(c_0_896,plain,(p2(esk52_2(esk1_0,X1))|p2(esk49_1(esk1_0))|p2(X2)|p2(X1)|~r1(esk6_1(esk49_1(esk1_0)),X2)|~r1(esk49_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_891]), c_0_890])]), c_0_892])).
cnf(c_0_897,plain,(p2(esk49_1(esk1_0))|p2(X1)|r1(esk48_1(esk49_1(esk1_0)),esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_893, c_0_894]), c_0_895])).
cnf(c_0_898,plain,(p2(esk49_1(esk1_0))|r1(esk6_1(esk49_1(esk1_0)),esk7_1(esk49_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_890])])).
cnf(c_0_899,plain,(p2(esk49_1(esk1_0))|~p2(esk7_1(esk49_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_890])])).
cnf(c_0_900,plain,(p2(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|p2(esk49_1(esk1_0))|p2(X1)|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_896, c_0_894]), c_0_895])).
cnf(c_0_901,plain,(p2(esk49_1(esk1_0))|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,X2),esk53_2(esk1_0,X2))|~r1(esk6_1(esk49_1(esk1_0)),X1)|~r1(esk49_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_891]), c_0_890])]), c_0_892])).
cnf(c_0_902,plain,(p2(esk49_1(esk1_0))|p2(X1)|~p2(X2)|~r1(esk48_1(esk49_1(esk1_0)),X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_890])])).
cnf(c_0_903,plain,(p2(esk49_1(esk1_0))|r1(esk48_1(esk49_1(esk1_0)),esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_897, c_0_898]), c_0_899])).
cnf(c_0_904,plain,(p2(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))))|p2(esk49_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_900, c_0_898]), c_0_899])).
cnf(c_0_905,plain,(p2(esk49_1(esk1_0))|p2(X1)|r1(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))),esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_901, c_0_894]), c_0_895])).
cnf(c_0_906,plain,(p2(esk49_1(esk1_0))|p2(X1)|p2(X2)|~p2(esk53_2(esk1_0,X2))|~r1(esk6_1(esk49_1(esk1_0)),X1)|~r1(esk49_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_891]), c_0_890])]), c_0_892])).
cnf(c_0_907,plain,(p2(esk49_1(esk1_0))|p2(X1)|~r1(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_902, c_0_903]), c_0_904])).
cnf(c_0_908,plain,(p2(esk49_1(esk1_0))|r1(esk52_2(esk1_0,esk48_1(esk49_1(esk1_0))),esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_905, c_0_898]), c_0_899])).
cnf(c_0_909,plain,(p2(esk49_1(esk1_0))|p2(X1)|~p2(esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_906, c_0_894]), c_0_895])).
cnf(c_0_910,plain,(p2(esk53_2(esk1_0,esk48_1(esk49_1(esk1_0))))|p2(esk49_1(esk1_0))), inference(spm,[status(thm)],[c_0_907, c_0_908])).
cnf(c_0_911,plain,(p2(esk49_1(esk1_0))|p2(X1)|~r1(esk6_1(esk49_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_909, c_0_910])).
cnf(c_0_912,plain,(p2(esk49_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_911, c_0_898]), c_0_899])).
cnf(c_0_913,plain,(r1(esk49_1(esk1_0),esk50_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_912]), c_0_890])])).
cnf(c_0_914,plain,(p2(X1)|r1(X1,esk52_2(esk1_0,X1))|~r1(esk49_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_912]), c_0_890])])).
cnf(c_0_915,plain,(p2(X1)|p2(X2)|r1(X1,esk61_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_913]), c_0_890])])).
cnf(c_0_916,plain,(p2(esk50_1(esk1_0))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_914, c_0_913])).
cnf(c_0_917,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_913]), c_0_890])])).
cnf(c_0_918,plain,(p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_915, c_0_916])).
cnf(c_0_919,plain,(r1(esk50_1(esk1_0),esk51_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_912]), c_0_890])])).
cnf(c_0_920,plain,(~p2(esk51_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_912]), c_0_890])])).
cnf(c_0_921,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_917, c_0_916])).
cnf(c_0_922,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X1),esk62_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_913]), c_0_890])])).
cnf(c_0_923,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_918, c_0_919]), c_0_920])).
cnf(c_0_924,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_921, c_0_919]), c_0_920])).
cnf(c_0_925,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_922, c_0_916])).
cnf(c_0_926,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(X1,esk61_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_913]), c_0_890])])).
cnf(c_0_927,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_913]), c_0_890])])).
cnf(c_0_928,plain,(p2(X1)|p2(X2)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_913]), c_0_890])])).
cnf(c_0_929,plain,(p2(X1)|~p2(X2)|~r1(esk51_1(esk1_0),X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_912]), c_0_890])])).
cnf(c_0_930,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_923, c_0_919]), c_0_920])).
cnf(c_0_931,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_924, c_0_919]), c_0_920])).
cnf(c_0_932,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_925, c_0_919]), c_0_920])).
cnf(c_0_933,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_926, c_0_916])).
cnf(c_0_934,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_927, c_0_916])).
cnf(c_0_935,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk61_3(esk1_0,esk50_1(esk1_0),X1),esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_913]), c_0_890])])).
cnf(c_0_936,plain,(p2(X1)|r1(esk52_2(esk1_0,X1),esk53_2(esk1_0,X1))|~r1(esk49_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_912]), c_0_890])])).
cnf(c_0_937,plain,(p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_928, c_0_916])).
cnf(c_0_938,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_929, c_0_930]), c_0_931])).
cnf(c_0_939,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_932, c_0_919]), c_0_920])).
cnf(c_0_940,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_933, c_0_919]), c_0_920])).
cnf(c_0_941,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_934, c_0_919]), c_0_920])).
cnf(c_0_942,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_935, c_0_916])).
cnf(c_0_943,plain,(p2(esk50_1(esk1_0))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_936, c_0_913])).
cnf(c_0_944,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_937, c_0_919]), c_0_920])).
cnf(c_0_945,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_938, c_0_939])).
cnf(c_0_946,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_913]), c_0_890])])).
cnf(c_0_947,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_913]), c_0_890])])).
cnf(c_0_948,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_940, c_0_919]), c_0_920])).
cnf(c_0_949,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_941, c_0_919]), c_0_920])).
cnf(c_0_950,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_942, c_0_919]), c_0_920])).
cnf(c_0_951,plain,(p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_915, c_0_943])).
cnf(c_0_952,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_917, c_0_943])).
cnf(c_0_953,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_944, c_0_945])).
cnf(c_0_954,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_946, c_0_916])).
cnf(c_0_955,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_947, c_0_916])).
cnf(c_0_956,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_929, c_0_948]), c_0_949])).
cnf(c_0_957,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_950, c_0_919]), c_0_920])).
cnf(c_0_958,plain,(p2(X1)|p2(X2)|r1(X1,esk61_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_913]), c_0_890])])).
cnf(c_0_959,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_951, c_0_919]), c_0_920])).
cnf(c_0_960,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_952, c_0_919]), c_0_920])).
cnf(c_0_961,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_922, c_0_943])).
cnf(c_0_962,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_919]), c_0_920])).
cnf(c_0_963,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_954, c_0_919]), c_0_920])).
cnf(c_0_964,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_955, c_0_919]), c_0_920])).
cnf(c_0_965,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_956, c_0_957])).
cnf(c_0_966,plain,(p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_958, c_0_916])).
cnf(c_0_967,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_913]), c_0_890])])).
cnf(c_0_968,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_959, c_0_919]), c_0_920])).
cnf(c_0_969,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_960, c_0_919]), c_0_920])).
cnf(c_0_970,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_961, c_0_919]), c_0_920])).
cnf(c_0_971,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_926, c_0_943])).
cnf(c_0_972,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_927, c_0_943])).
cnf(c_0_973,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_929, c_0_962])).
cnf(c_0_974,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_963, c_0_919]), c_0_920])).
cnf(c_0_975,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_964, c_0_965])).
cnf(c_0_976,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_966, c_0_919]), c_0_920])).
cnf(c_0_977,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk61_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_913]), c_0_890])])).
cnf(c_0_978,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_967, c_0_916])).
cnf(c_0_979,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X1),esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_913]), c_0_890])])).
cnf(c_0_980,plain,(p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_928, c_0_943])).
cnf(c_0_981,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_929, c_0_968]), c_0_969])).
cnf(c_0_982,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_970, c_0_919]), c_0_920])).
cnf(c_0_983,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_971, c_0_919]), c_0_920])).
cnf(c_0_984,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_972, c_0_919]), c_0_920])).
cnf(c_0_985,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_935, c_0_943])).
cnf(c_0_986,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_973, c_0_974])).
cnf(c_0_987,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_975, c_0_919]), c_0_920])).
cnf(c_0_988,plain,(r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_976, c_0_919]), c_0_920])).
cnf(c_0_989,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_977, c_0_916])).
cnf(c_0_990,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_978, c_0_919]), c_0_920])).
cnf(c_0_991,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_979, c_0_916])).
cnf(c_0_992,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_980, c_0_919]), c_0_920])).
cnf(c_0_993,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_981, c_0_982])).
cnf(c_0_994,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_983, c_0_919]), c_0_920])).
cnf(c_0_995,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_984, c_0_919]), c_0_920])).
cnf(c_0_996,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_985, c_0_919]), c_0_920])).
cnf(c_0_997,plain,(r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_986, c_0_987]), c_0_988])).
cnf(c_0_998,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_989, c_0_919]), c_0_920])).
cnf(c_0_999,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_990, c_0_919]), c_0_920])).
cnf(c_0_1000,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_991, c_0_919]), c_0_920])).
cnf(c_0_1001,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_992, c_0_993])).
cnf(c_0_1002,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_946, c_0_943])).
cnf(c_0_1003,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_947, c_0_943])).
cnf(c_0_1004,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_929, c_0_994]), c_0_995])).
cnf(c_0_1005,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_996, c_0_919]), c_0_920])).
cnf(c_0_1006,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_913]), c_0_890])])).
cnf(c_0_1007,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_929, c_0_997])).
cnf(c_0_1008,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_998, c_0_919]), c_0_920])).
cnf(c_0_1009,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_973, c_0_999])).
cnf(c_0_1010,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1000, c_0_919]), c_0_920])).
cnf(c_0_1011,plain,(r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1001, c_0_919]), c_0_920])).
cnf(c_0_1012,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1002, c_0_919]), c_0_920])).
cnf(c_0_1013,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_919]), c_0_920])).
cnf(c_0_1014,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1004, c_0_1005])).
cnf(c_0_1015,plain,(p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_958, c_0_943])).
cnf(c_0_1016,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1006, c_0_916])).
cnf(c_0_1017,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1007, c_0_1008])).
cnf(c_0_1018,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1009, c_0_987]), c_0_1010])).
cnf(c_0_1019,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_929, c_0_1011])).
cnf(c_0_1020,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1012, c_0_919]), c_0_920])).
cnf(c_0_1021,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1013, c_0_1014])).
cnf(c_0_1022,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1015, c_0_919]), c_0_920])).
cnf(c_0_1023,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_967, c_0_943])).
cnf(c_0_1024,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1016, c_0_919]), c_0_920])).
cnf(c_0_1025,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1017, c_0_1018])).
cnf(c_0_1026,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1019, c_0_1020])).
cnf(c_0_1027,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1021, c_0_919]), c_0_920])).
cnf(c_0_1028,plain,(r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1022, c_0_919]), c_0_920])).
cnf(c_0_1029,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_977, c_0_943])).
cnf(c_0_1030,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1023, c_0_919]), c_0_920])).
cnf(c_0_1031,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_979, c_0_943])).
cnf(c_0_1032,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_913]), c_0_890])])).
cnf(c_0_1033,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1024, c_0_1025])).
cnf(c_0_1034,plain,(r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1026, c_0_1027]), c_0_1028])).
cnf(c_0_1035,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1029, c_0_919]), c_0_920])).
cnf(c_0_1036,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1030, c_0_919]), c_0_920])).
cnf(c_0_1037,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1031, c_0_919]), c_0_920])).
cnf(c_0_1038,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1032, c_0_916])).
cnf(c_0_1039,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1033, c_0_919]), c_0_920])).
cnf(c_0_1040,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_929, c_0_1034])).
cnf(c_0_1041,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1035, c_0_919]), c_0_920])).
cnf(c_0_1042,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1019, c_0_1036])).
cnf(c_0_1043,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1037, c_0_919]), c_0_920])).
cnf(c_0_1044,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1038, c_0_919]), c_0_920])).
cnf(c_0_1045,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_973, c_0_1039])).
cnf(c_0_1046,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1006, c_0_943])).
cnf(c_0_1047,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1040, c_0_1041])).
cnf(c_0_1048,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1042, c_0_1027]), c_0_1043])).
cnf(c_0_1049,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1044, c_0_919]), c_0_920])).
cnf(c_0_1050,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1045, c_0_987])).
cnf(c_0_1051,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1046, c_0_919]), c_0_920])).
cnf(c_0_1052,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1047, c_0_1048])).
cnf(c_0_1053,plain,(p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk50_1(esk1_0))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_913]), c_0_890])])).
cnf(c_0_1054,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1049, c_0_1050])).
cnf(c_0_1055,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1051, c_0_1052])).
cnf(c_0_1056,plain,(p2(X1)|p2(X2)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1053, c_0_916])).
cnf(c_0_1057,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1007, c_0_1054])).
cnf(c_0_1058,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1032, c_0_943])).
cnf(c_0_1059,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1055, c_0_919]), c_0_920])).
cnf(c_0_1060,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1056, c_0_919]), c_0_920])).
cnf(c_0_1061,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1057, c_0_1018])).
cnf(c_0_1062,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1058, c_0_919]), c_0_920])).
cnf(c_0_1063,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1019, c_0_1059])).
cnf(c_0_1064,plain,(p2(X1)|r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1060, c_0_1061])).
cnf(c_0_1065,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1062, c_0_919]), c_0_920])).
cnf(c_0_1066,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1063, c_0_1027])).
cnf(c_0_1067,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_913]), c_0_890])])).
cnf(c_0_1068,plain,(r1(esk50_1(esk1_0),esk52_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1064, c_0_919]), c_0_920]), c_0_1050])).
cnf(c_0_1069,plain,(p2(esk52_2(esk1_0,X1))|p2(X1)|~r1(esk49_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_912]), c_0_890])])).
cnf(c_0_1070,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1065, c_0_1066])).
cnf(c_0_1071,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_913]), c_0_890])])).
cnf(c_0_1072,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1067, c_0_1068])).
cnf(c_0_1073,plain,(p2(esk52_2(esk1_0,esk50_1(esk1_0)))|p2(esk50_1(esk1_0))), inference(spm,[status(thm)],[c_0_1069, c_0_913])).
cnf(c_0_1074,plain,(p2(X1)|p2(X2)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1053, c_0_943])).
cnf(c_0_1075,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1040, c_0_1070])).
cnf(c_0_1076,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1071, c_0_1068])).
cnf(c_0_1077,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_913]), c_0_890])])).
cnf(c_0_1078,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1072, c_0_1073]), c_0_915])).
cnf(c_0_1079,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1074, c_0_919]), c_0_920])).
cnf(c_0_1080,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1075, c_0_1048])).
cnf(c_0_1081,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1076, c_0_1073]), c_0_917])).
cnf(c_0_1082,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1077, c_0_1068])).
cnf(c_0_1083,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_913]), c_0_890])])).
cnf(c_0_1084,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_913]), c_0_890])])).
cnf(c_0_1085,plain,(p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1078, c_0_919]), c_0_920])).
cnf(c_0_1086,plain,(p2(X1)|r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1079, c_0_1080])).
cnf(c_0_1087,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1081, c_0_919]), c_0_920])).
cnf(c_0_1088,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1082, c_0_1073]), c_0_922])).
cnf(c_0_1089,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1083, c_0_1068])).
cnf(c_0_1090,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1084, c_0_1068])).
cnf(c_0_1091,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_913]), c_0_890])])).
cnf(c_0_1092,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_913]), c_0_890])])).
cnf(c_0_1093,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1085, c_0_919]), c_0_920])).
cnf(c_0_1094,plain,(r1(esk52_2(esk1_0,esk50_1(esk1_0)),esk53_2(esk1_0,esk50_1(esk1_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1086, c_0_919]), c_0_920]), c_0_1066])).
cnf(c_0_1095,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1087, c_0_919]), c_0_920])).
cnf(c_0_1096,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1088, c_0_919]), c_0_920])).
cnf(c_0_1097,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1089, c_0_1073]), c_0_926])).
cnf(c_0_1098,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1090, c_0_1073]), c_0_927])).
cnf(c_0_1099,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1091, c_0_1068])).
cnf(c_0_1100,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1092, c_0_1068])).
cnf(c_0_1101,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1093, c_0_1094])).
cnf(c_0_1102,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1095, c_0_1094])).
cnf(c_0_1103,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1096, c_0_919]), c_0_920])).
cnf(c_0_1104,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1097, c_0_919]), c_0_920])).
cnf(c_0_1105,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1098, c_0_919]), c_0_920])).
cnf(c_0_1106,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1099, c_0_1073]), c_0_935])).
cnf(c_0_1107,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X3,esk57_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1100, c_0_1073]), c_0_928])).
cnf(c_0_1108,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_929, c_0_1101]), c_0_1102])).
cnf(c_0_1109,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1103, c_0_1094])).
cnf(c_0_1110,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk61_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_913]), c_0_890])])).
cnf(c_0_1111,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_913]), c_0_890])])).
cnf(c_0_1112,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1104, c_0_919]), c_0_920])).
cnf(c_0_1113,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1105, c_0_919]), c_0_920])).
cnf(c_0_1114,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1106, c_0_919]), c_0_920])).
cnf(c_0_1115,plain,(p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1107, c_0_919]), c_0_920])).
cnf(c_0_1116,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1108, c_0_1109])).
cnf(c_0_1117,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1110, c_0_1068])).
cnf(c_0_1118,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1111, c_0_1068])).
cnf(c_0_1119,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1112, c_0_1094])).
cnf(c_0_1120,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1113, c_0_1094])).
cnf(c_0_1121,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1114, c_0_919]), c_0_920])).
cnf(c_0_1122,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1115, c_0_1116])).
cnf(c_0_1123,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk61_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1117, c_0_1073]), c_0_946])).
cnf(c_0_1124,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X3),esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1118, c_0_1073]), c_0_947])).
cnf(c_0_1125,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_929, c_0_1119]), c_0_1120])).
cnf(c_0_1126,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1121, c_0_1094])).
cnf(c_0_1127,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X3),esk62_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_913]), c_0_890])])).
cnf(c_0_1128,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_913]), c_0_890])])).
cnf(c_0_1129,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1122, c_0_919]), c_0_920])).
cnf(c_0_1130,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1123, c_0_919]), c_0_920])).
cnf(c_0_1131,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1124, c_0_919]), c_0_920])).
cnf(c_0_1132,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1125, c_0_1126])).
cnf(c_0_1133,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1127, c_0_1068])).
cnf(c_0_1134,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1128, c_0_1068])).
cnf(c_0_1135,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1129, c_0_1094])).
cnf(c_0_1136,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1130, c_0_919]), c_0_920])).
cnf(c_0_1137,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1131, c_0_1132])).
cnf(c_0_1138,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X3),esk62_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1133, c_0_1073]), c_0_967])).
cnf(c_0_1139,plain,(p2(X1)|p2(X2)|p2(X3)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1134, c_0_1073]), c_0_958])).
cnf(c_0_1140,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_929, c_0_1135])).
cnf(c_0_1141,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1136, c_0_1094])).
cnf(c_0_1142,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1137, c_0_919]), c_0_920])).
cnf(c_0_1143,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk61_3(esk1_0,esk50_1(esk1_0),X2))|p2(X3)|p2(X2)|p2(X1)|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_913]), c_0_890])])).
cnf(c_0_1144,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_913]), c_0_890])])).
cnf(c_0_1145,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1138, c_0_919]), c_0_920])).
cnf(c_0_1146,plain,(p2(X1)|p2(X2)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1139, c_0_919]), c_0_920])).
cnf(c_0_1147,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1140, c_0_1141])).
cnf(c_0_1148,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1142, c_0_1094])).
cnf(c_0_1149,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|p2(X3)|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1143, c_0_1068])).
cnf(c_0_1150,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1144, c_0_1068])).
cnf(c_0_1151,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1145, c_0_919]), c_0_920])).
cnf(c_0_1152,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1146, c_0_919]), c_0_920])).
cnf(c_0_1153,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1147, c_0_1148])).
cnf(c_0_1154,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk61_3(esk1_0,esk50_1(esk1_0),X2))|p2(X3)|p2(X2)|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1149, c_0_1073]), c_0_977])).
cnf(c_0_1155,plain,(p2(X1)|p2(X2)|p2(X3)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1150, c_0_1073]), c_0_979])).
cnf(c_0_1156,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1151, c_0_1094])).
cnf(c_0_1157,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1152, c_0_1153])).
cnf(c_0_1158,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1154, c_0_919]), c_0_920])).
cnf(c_0_1159,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1155, c_0_919]), c_0_920])).
cnf(c_0_1160,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1140, c_0_1156])).
cnf(c_0_1161,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_913]), c_0_890])])).
cnf(c_0_1162,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1157, c_0_1094])).
cnf(c_0_1163,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1158, c_0_919]), c_0_920])).
cnf(c_0_1164,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1159, c_0_919]), c_0_920])).
cnf(c_0_1165,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1160, c_0_1148])).
cnf(c_0_1166,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|p2(X3)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1161, c_0_1068])).
cnf(c_0_1167,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_929, c_0_1162])).
cnf(c_0_1168,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1163, c_0_1094])).
cnf(c_0_1169,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1164, c_0_1165])).
cnf(c_0_1170,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X3)|p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X3)|~r1(esk50_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1166, c_0_1073]), c_0_1006])).
cnf(c_0_1171,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1167, c_0_1168])).
cnf(c_0_1172,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1169, c_0_1094])).
cnf(c_0_1173,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1170, c_0_919]), c_0_920])).
cnf(c_0_1174,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1171, c_0_1172])).
cnf(c_0_1175,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_913]), c_0_890])])).
cnf(c_0_1176,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X2)|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1173, c_0_1174])).
cnf(c_0_1177,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1175, c_0_1068])).
cnf(c_0_1178,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1176, c_0_919]), c_0_920])).
cnf(c_0_1179,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1177, c_0_1073]), c_0_1032])).
cnf(c_0_1180,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1178, c_0_1094])).
cnf(c_0_1181,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1179, c_0_919]), c_0_920])).
cnf(c_0_1182,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1140, c_0_1180])).
cnf(c_0_1183,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1181, c_0_919]), c_0_920])).
cnf(c_0_1184,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1182, c_0_1148])).
cnf(c_0_1185,plain,(p2(X1)|p2(X2)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(X4)|~r1(esk50_1(esk1_0),X4)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)|~r1(X4,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_913]), c_0_890])])).
cnf(c_0_1186,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1183, c_0_1184])).
cnf(c_0_1187,plain,(p2(X1)|p2(X2)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk52_2(esk1_0,esk50_1(esk1_0)))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X3)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1185, c_0_1068])).
cnf(c_0_1188,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1186, c_0_1094])).
cnf(c_0_1189,plain,(p2(X1)|p2(X2)|p2(X3)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X3))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1187, c_0_1073]), c_0_1053])).
cnf(c_0_1190,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1167, c_0_1188])).
cnf(c_0_1191,plain,(p2(X1)|p2(X2)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X2)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1189, c_0_919]), c_0_920])).
cnf(c_0_1192,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1190, c_0_1172])).
cnf(c_0_1193,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)|~r1(esk50_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1191, c_0_1192])).
cnf(c_0_1194,plain,(p2(X1)|~p2(esk53_2(esk1_0,X1))|~r1(esk49_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_912]), c_0_890])])).
cnf(c_0_1195,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))|p2(X1)|~r1(esk52_2(esk1_0,esk50_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1193, c_0_919]), c_0_920]), c_0_1184])).
cnf(c_0_1196,plain,(p2(esk50_1(esk1_0))|~p2(esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1194, c_0_913])).
cnf(c_0_1197,plain,(p2(esk53_2(esk1_0,esk50_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1195, c_0_1094])).
cnf(c_0_1198,plain,(p2(esk50_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1196, c_0_1197])])).
cnf(c_0_1199,plain,(p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_915, c_0_1198])])).
cnf(c_0_1200,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(X2,esk57_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_917, c_0_1198])])).
cnf(c_0_1201,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1199, c_0_919]), c_0_920])).
cnf(c_0_1202,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1200, c_0_919]), c_0_920])).
cnf(c_0_1203,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_922, c_0_1198])])).
cnf(c_0_1204,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1201, c_0_919]), c_0_920])).
cnf(c_0_1205,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1202, c_0_919]), c_0_920])).
cnf(c_0_1206,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1203, c_0_919]), c_0_920])).
cnf(c_0_1207,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_926, c_0_1198])])).
cnf(c_0_1208,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X2),esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_927, c_0_1198])])).
cnf(c_0_1209,plain,(p2(X1)|p2(X2)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_928, c_0_1198])])).
cnf(c_0_1210,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_929, c_0_1204]), c_0_1205])).
cnf(c_0_1211,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1206, c_0_919]), c_0_920])).
cnf(c_0_1212,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1207, c_0_919]), c_0_920])).
cnf(c_0_1213,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1208, c_0_919]), c_0_920])).
cnf(c_0_1214,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_935, c_0_1198])])).
cnf(c_0_1215,plain,(p2(X1)|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1209, c_0_919]), c_0_920])).
cnf(c_0_1216,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1210, c_0_1211])).
cnf(c_0_1217,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1212, c_0_919]), c_0_920])).
cnf(c_0_1218,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1213, c_0_919]), c_0_920])).
cnf(c_0_1219,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1214, c_0_919]), c_0_920])).
cnf(c_0_1220,plain,(p2(X1)|r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(X1,esk57_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1215, c_0_1216])).
cnf(c_0_1221,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_946, c_0_1198])])).
cnf(c_0_1222,plain,(p2(X1)|p2(X2)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_947, c_0_1198])])).
cnf(c_0_1223,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_929, c_0_1217]), c_0_1218])).
cnf(c_0_1224,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1219, c_0_919]), c_0_920])).
cnf(c_0_1225,plain,(r1(esk51_1(esk1_0),esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1220, c_0_919]), c_0_920])).
cnf(c_0_1226,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1221, c_0_919]), c_0_920])).
cnf(c_0_1227,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1222, c_0_919]), c_0_920])).
cnf(c_0_1228,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1223, c_0_1224])).
cnf(c_0_1229,plain,(p2(X1)|p2(X2)|r1(X2,esk61_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_958, c_0_1198])])).
cnf(c_0_1230,plain,(p2(X1)|~p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_929, c_0_1225])).
cnf(c_0_1231,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1226, c_0_919]), c_0_920])).
cnf(c_0_1232,plain,(p2(X1)|r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk57_3(esk1_0,esk50_1(esk1_0),X1),esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1227, c_0_1228])).
cnf(c_0_1233,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1229, c_0_919]), c_0_920])).
cnf(c_0_1234,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_967, c_0_1198])])).
cnf(c_0_1235,plain,(p2(X1)|r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1230, c_0_1231])).
cnf(c_0_1236,plain,(r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1232, c_0_919]), c_0_920])).
cnf(c_0_1237,plain,(r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1233, c_0_919]), c_0_920])).
cnf(c_0_1238,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X2))|p2(X2)|p2(X1)|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_977, c_0_1198])])).
cnf(c_0_1239,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1234, c_0_919]), c_0_920])).
cnf(c_0_1240,plain,(p2(X1)|p2(X2)|r1(esk61_3(esk1_0,esk50_1(esk1_0),X2),esk62_3(esk1_0,esk50_1(esk1_0),X2))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_979, c_0_1198])])).
cnf(c_0_1241,plain,(r1(esk51_1(esk1_0),esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1235, c_0_1236]), c_0_1237])).
cnf(c_0_1242,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1238, c_0_919]), c_0_920])).
cnf(c_0_1243,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1239, c_0_919]), c_0_920])).
cnf(c_0_1244,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1240, c_0_919]), c_0_920])).
cnf(c_0_1245,plain,(p2(X1)|~p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_929, c_0_1241])).
cnf(c_0_1246,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1242, c_0_919]), c_0_920])).
cnf(c_0_1247,plain,(p2(X1)|r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1230, c_0_1243])).
cnf(c_0_1248,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1244, c_0_919]), c_0_920])).
cnf(c_0_1249,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|p2(X2)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1006, c_0_1198])])).
cnf(c_0_1250,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_1245, c_0_1246])).
cnf(c_0_1251,plain,(r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1247, c_0_1236]), c_0_1248])).
cnf(c_0_1252,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1249, c_0_919]), c_0_920])).
cnf(c_0_1253,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1250, c_0_1251])).
cnf(c_0_1254,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(esk57_3(esk1_0,esk50_1(esk1_0),X1))|p2(X1)|~r1(esk50_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1252, c_0_1253])).
cnf(c_0_1255,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),X1))|p2(X2)|p2(X1)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X1)|~r1(esk50_1(esk1_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1032, c_0_1198])])).
cnf(c_0_1256,plain,(p2(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1254, c_0_919]), c_0_920])).
cnf(c_0_1257,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|p2(X1)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1255, c_0_919]), c_0_920])).
cnf(c_0_1258,plain,(p2(X1)|~r1(esk57_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1230, c_0_1256])])).
cnf(c_0_1259,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1257, c_0_919]), c_0_920])).
cnf(c_0_1260,plain,(p2(esk58_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1258, c_0_1236])).
cnf(c_0_1261,plain,(p2(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1259, c_0_1260])])).
cnf(c_0_1262,plain,(p2(X1)|p2(X2)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~p2(esk62_3(esk1_0,esk50_1(esk1_0),X2))|~r1(esk50_1(esk1_0),X2)|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1053, c_0_1198])])).
cnf(c_0_1263,plain,(p2(X1)|~r1(esk61_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1245, c_0_1261])])).
cnf(c_0_1264,plain,(p2(X1)|~p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1262, c_0_919]), c_0_920])).
cnf(c_0_1265,plain,(p2(esk62_3(esk1_0,esk50_1(esk1_0),esk51_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1263, c_0_1251])).
cnf(c_0_1266,plain,(p2(X1)|~p2(esk58_3(esk1_0,esk50_1(esk1_0),X1))|~r1(esk50_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1264, c_0_1265])])).
cnf(c_0_1267,plain,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1266, c_0_919]), c_0_1260])]), c_0_920]), ['proof']).
# SZS output end CNFRefutation.
EOF
