% Problem    : CSR033+3 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n144.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:33:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, ((mtvisible(c_tptpgeo_spindlecollectormt)=>(inregion(c_geolocation_x76_y23,c_georegion_l1_x2_y0)&geolevel_1(c_georegion_l1_x2_y0)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', query133)).
fof(c_0_1, axiom, (![X21]:![X22]:((mtvisible(X21)&genlmt(X21,X22))=>mtvisible(X22))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_7997)).
fof(c_0_2, axiom, (genlmt(c_tptpgeo_spindlecollectormt,c_tptpgeo_member2_mt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_232)).
fof(c_0_3, axiom, (![X15]:![X16]:![X17]:((inregion(X15,X16)&inregion(X16,X17))=>inregion(X15,X17))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_7302)).
fof(c_0_4, axiom, (![X3]:![X4]:(geographicalsubregions(X3,X4)=>inregion(X4,X3))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1935)).
fof(c_0_5, axiom, ((mtvisible(c_tptpgeo_member2_mt)=>geographicalsubregions(c_georegion_l1_x2_y0,c_georegion_l2_x8_y2))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1039)).
fof(c_0_6, axiom, ((mtvisible(c_tptpgeo_member2_mt)=>geographicalsubregions(c_georegion_l2_x8_y2,c_georegion_l3_x25_y7))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_438)).
fof(c_0_7, axiom, ((mtvisible(c_tptpgeo_member2_mt)=>geographicalsubregions(c_georegion_l3_x25_y7,c_georegion_l4_x76_y23))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_2573)).
fof(c_0_8, axiom, (genlmt(c_tptpgeo_member2_mt,c_tptpgeo_spindleheadmt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1355)).
fof(c_0_9, axiom, ((mtvisible(c_tptpgeo_member2_mt)=>inregion(c_geolocation_x76_y23,c_georegion_l4_x76_y23))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_2962)).
fof(c_0_10, axiom, ((mtvisible(c_worldgeographymt)=>geolevel_1(c_georegion_l1_x2_y0))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_2024)).
fof(c_0_11, axiom, (genlmt(c_tptpgeo_spindleheadmt,c_worldgeographymt)), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+2.ax', ax2_1111)).
fof(c_0_12, negated_conjecture, (~((mtvisible(c_tptpgeo_spindlecollectormt)=>(inregion(c_geolocation_x76_y23,c_georegion_l1_x2_y0)&geolevel_1(c_georegion_l1_x2_y0))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_13, plain, (![X23]:![X24]:((~mtvisible(X23)|~genlmt(X23,X24))|mtvisible(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
fof(c_0_14, negated_conjecture, ((mtvisible(c_tptpgeo_spindlecollectormt)&(~inregion(c_geolocation_x76_y23,c_georegion_l1_x2_y0)|~geolevel_1(c_georegion_l1_x2_y0)))), inference(fof_nnf,[status(thm)],[c_0_12])).
cnf(c_0_15,plain,(mtvisible(X1)|~genlmt(X2,X1)|~mtvisible(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16,plain,(genlmt(c_tptpgeo_spindlecollectormt,c_tptpgeo_member2_mt)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_17,negated_conjecture,(mtvisible(c_tptpgeo_spindlecollectormt)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, (![X18]:![X19]:![X20]:((~inregion(X18,X19)|~inregion(X19,X20))|inregion(X18,X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
fof(c_0_19, plain, (![X5]:![X6]:(~geographicalsubregions(X5,X6)|inregion(X6,X5))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
fof(c_0_20, plain, ((~mtvisible(c_tptpgeo_member2_mt)|geographicalsubregions(c_georegion_l1_x2_y0,c_georegion_l2_x8_y2))), inference(fof_nnf,[status(thm)],[c_0_5])).
fof(c_0_21, plain, ((~mtvisible(c_tptpgeo_member2_mt)|geographicalsubregions(c_georegion_l2_x8_y2,c_georegion_l3_x25_y7))), inference(fof_nnf,[status(thm)],[c_0_6])).
fof(c_0_22, plain, ((~mtvisible(c_tptpgeo_member2_mt)|geographicalsubregions(c_georegion_l3_x25_y7,c_georegion_l4_x76_y23))), inference(fof_nnf,[status(thm)],[c_0_7])).
cnf(c_0_23,plain,(genlmt(c_tptpgeo_member2_mt,c_tptpgeo_spindleheadmt)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_24,plain,(mtvisible(c_tptpgeo_member2_mt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
cnf(c_0_25,plain,(inregion(X1,X2)|~inregion(X3,X2)|~inregion(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26,plain,(inregion(X1,X2)|~geographicalsubregions(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,plain,(geographicalsubregions(c_georegion_l1_x2_y0,c_georegion_l2_x8_y2)|~mtvisible(c_tptpgeo_member2_mt)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,plain,(geographicalsubregions(c_georegion_l2_x8_y2,c_georegion_l3_x25_y7)|~mtvisible(c_tptpgeo_member2_mt)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29,plain,(geographicalsubregions(c_georegion_l3_x25_y7,c_georegion_l4_x76_y23)|~mtvisible(c_tptpgeo_member2_mt)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_30, plain, ((~mtvisible(c_tptpgeo_member2_mt)|inregion(c_geolocation_x76_y23,c_georegion_l4_x76_y23))), inference(fof_nnf,[status(thm)],[c_0_9])).
fof(c_0_31, plain, ((~mtvisible(c_worldgeographymt)|geolevel_1(c_georegion_l1_x2_y0))), inference(fof_nnf,[status(thm)],[c_0_10])).
cnf(c_0_32,plain,(genlmt(c_tptpgeo_spindleheadmt,c_worldgeographymt)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_33,plain,(mtvisible(c_tptpgeo_spindleheadmt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_23]), c_0_24])])).
cnf(c_0_34,plain,(inregion(X1,X2)|~inregion(X1,X3)|~geographicalsubregions(X2,X3)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_35,plain,(geographicalsubregions(c_georegion_l1_x2_y0,c_georegion_l2_x8_y2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_24])])).
cnf(c_0_36,plain,(geographicalsubregions(c_georegion_l2_x8_y2,c_georegion_l3_x25_y7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_24])])).
cnf(c_0_37,plain,(geographicalsubregions(c_georegion_l3_x25_y7,c_georegion_l4_x76_y23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_24])])).
cnf(c_0_38,plain,(inregion(c_geolocation_x76_y23,c_georegion_l4_x76_y23)|~mtvisible(c_tptpgeo_member2_mt)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39,plain,(geolevel_1(c_georegion_l1_x2_y0)|~mtvisible(c_worldgeographymt)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40,plain,(mtvisible(c_worldgeographymt)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_32]), c_0_33])])).
cnf(c_0_41,plain,(inregion(X1,c_georegion_l1_x2_y0)|~inregion(X1,c_georegion_l2_x8_y2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42,plain,(inregion(X1,c_georegion_l2_x8_y2)|~inregion(X1,c_georegion_l3_x25_y7)), inference(spm,[status(thm)],[c_0_34, c_0_36])).
cnf(c_0_43,plain,(inregion(X1,c_georegion_l3_x25_y7)|~inregion(X1,c_georegion_l4_x76_y23)), inference(spm,[status(thm)],[c_0_34, c_0_37])).
cnf(c_0_44,plain,(inregion(c_geolocation_x76_y23,c_georegion_l4_x76_y23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_24])])).
cnf(c_0_45,negated_conjecture,(~geolevel_1(c_georegion_l1_x2_y0)|~inregion(c_geolocation_x76_y23,c_georegion_l1_x2_y0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_46,plain,(geolevel_1(c_georegion_l1_x2_y0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40])])).
cnf(c_0_47,plain,(inregion(X1,c_georegion_l1_x2_y0)|~inregion(X1,c_georegion_l3_x25_y7)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48,plain,(inregion(c_geolocation_x76_y23,c_georegion_l3_x25_y7)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49,negated_conjecture,(~inregion(c_geolocation_x76_y23,c_georegion_l1_x2_y0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])])).
cnf(c_0_50,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), ['proof']).
# SZS output end CNFRefutation.
EOF
