% Problem    : SWV437+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n092.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:21:56 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_B03_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X7]:![X5]:![X6]:![X8]:(system_indi_is_oca(system,X7)=>(oca_compartment_has_scg(X7,X5,X8)=>(admin_compartment_has_sso(admin,X5,X6)=>(sso_compartment_has_scg(X6,X5,X8)=>admin_compartment_has_scg(admin,X5,X8)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax7)).
fof(c_0_1, axiom, (![X5]:![X6]:(system_compartment_has_sso(system,X5,X6)=>admin_compartment_has_sso(admin,X5,X6))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax6)).
fof(c_0_2, hypothesis, (system_compartment_has_sso(system,compartmentb,sso_compartmentb)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax44)).
fof(c_0_3, axiom, (![X9]:![X2]:![X5]:![X10]:![X6]:![X8]:(admin_compartment_has_sso(admin,X5,X6)=>(admin_compartment_has_scg(admin,X5,X8)=>(sso_file_has_level(X6,X9,X2,X8)=>(admin_file_has_level_h(admin,X9,X2,X10)=>admin_file_has_level_h(admin,X9,X2,cons(X5,X10))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax13)).
fof(c_0_4, hypothesis, (sso_compartment_has_scg(sso_compartmentb,compartmentb,scg_compartmentb)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax46)).
fof(c_0_5, axiom, (![X1]:![X13]:(system_indi_has_citizenship(system,X1,X13)=>admin_indi_has_citizenship(admin,X1,X13))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax23)).
fof(c_0_6, hypothesis, (sso_file_has_level(sso_compartmentb,secretfile,secret,scg_compartmentb)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax55)).
fof(c_0_7, hypothesis, (oca_compartment_has_scg(oca,compartmentb,scg_compartmentb)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax45)).
fof(c_0_8, hypothesis, (system_indi_is_oca(system,oca)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax41)).
fof(c_0_9, axiom, (![X1]:![X2]:![X3]:![X18]:![X4]:(system_indi_needs_level(system,X1,X3)=>(admin_indi_has_citizenship(admin,X1,usa)=>(admin_indi_has_polygraph(admin,X1)=>(admin_indi_has_employment(admin,X1)=>(admin_indi_has_credit(admin,X1)=>(loca_level_below(admin,X2,X3)=>(system_indi_is_level_admin(system,X18)=>(level_admin_indi_has_level(X18,X1,X4)=>(loca_level_below(admin,X2,X4)=>(admin_indi_has_background(admin,X1,X2)=>admin_indi_has_level(admin,X1,X2)))))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax25)).
fof(c_0_10, hypothesis, (system_indi_has_citizenship(system,alice,usa)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax71)).
fof(c_0_11, axiom, (![X1]:![X17]:(system_indi_is_hr_admin(system,X17)=>(hr_admin_indi_has_employment(X17,X1)=>admin_indi_has_employment(admin,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax21)).
fof(c_0_12, axiom, (![X1]:![X5]:![X7]:![X3]:![X19]:![X21]:(system_indi_is_oca(system,X7)=>(oca_compartment_is_compartment(X7,X5,X3,X19,yes,X21)=>(admin_indi_has_credit(admin,X1)=>admin_indi_has_credit_for_compartment(admin,X1,X5))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax32)).
fof(c_0_13, axiom, (![X1]:![X15]:(system_indi_is_credit_admin(system,X15)=>(credit_admin_indi_has_credit(X15,X1)=>admin_indi_has_credit(admin,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax18)).
fof(c_0_14, axiom, (![X1]:![X5]:![X7]:![X3]:![X19]:![X20]:(system_indi_is_oca(system,X7)=>(oca_compartment_is_compartment(X7,X5,X3,X19,X20,yes)=>(admin_indi_has_polygraph(admin,X1)=>admin_indi_has_polygraph_for_compartment(admin,X1,X5))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax30)).
fof(c_0_15, axiom, (![X1]:![X14]:(system_indi_is_polygraph_admin(system,X14)=>(polygraph_admin_indi_has_polygraph(X14,X1)=>admin_indi_has_polygraph(admin,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax17)).
fof(c_0_16, hypothesis, (system_indi_is_hr_admin(system,hr_admin)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax69)).
fof(c_0_17, axiom, (![X1]:![X2]:![X3]:![X4]:(loca_level_direct_below(X1,X3,X4)=>(loca_level_below(X1,X2,X3)=>loca_level_below(X1,X2,X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax5)).
fof(c_0_18, axiom, (![X1]:loca_level_direct_below(X1,secret,topsecret)), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax3)).
fof(c_0_19, axiom, (![X9]:![X2]:![X10]:(system_file_needs_level(system,X9,X2)=>(admin_file_has_compartments(admin,X9,X10)=>(admin_file_has_level_h(admin,X9,X2,X10)=>admin_file_has_level(admin,X9,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax11)).
fof(c_0_20, axiom, (![X9]:![X10]:![X11]:![X12]:![X6]:(admin_compartment_has_sso(admin,X11,X6)=>(sso_file_has_compartments(X6,X9,X10)=>(admin_file_has_compartments_h(admin,X9,X10,X12)=>admin_file_has_compartments_h(admin,X9,X10,cons(X11,X12)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax10)).
fof(c_0_21, axiom, (![X1]:![X5]:![X10]:![X6]:(system_indi_needs_compartment(system,X1,X5)=>(admin_indi_has_employment(admin,X1)=>(admin_indi_has_citizenship(admin,X1,usa)=>(admin_indi_has_polygraph_for_compartment(admin,X1,X5)=>(admin_indi_has_credit_for_compartment(admin,X1,X5)=>(admin_compartment_has_sso(admin,X5,X6)=>(sso_indi_has_compartment(X6,X1,X5)=>(admin_indi_has_background_for_compartment(admin,X1,X5)=>(admin_indi_has_level_for_compartment(admin,X1,X5)=>(admin_indi_has_compartments(admin,X1,X10)=>admin_indi_has_compartments(admin,X1,cons(X5,X10))))))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax27)).
fof(c_0_22, hypothesis, (oca_compartment_is_compartment(oca,compartmentb,confidential,topsecret,yes,yes)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax42)).
fof(c_0_23, hypothesis, (system_indi_is_credit_admin(system,credit_admin)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax67)).
fof(c_0_24, hypothesis, (system_indi_is_polygraph_admin(system,polygraph_admin)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax66)).
fof(c_0_25, hypothesis, (hr_admin_indi_has_employment(hr_admin,alice)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax75)).
fof(c_0_26, axiom, (![X1]:![X2]:![X16]:![X3]:(system_indi_is_background_admin(system,X16)=>(background_admin_indi_has_background(X16,X1,X3)=>(loca_level_below(admin,X2,X3)=>admin_indi_has_background(admin,X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax20)).
fof(c_0_27, hypothesis, (system_file_needs_level(system,secretfile,secret)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax54)).
fof(c_0_28, axiom, (![X9]:![X10]:(system_file_needs_compartments(system,X9,X10)=>(admin_file_has_compartments_h(admin,X9,X10,X10)=>admin_file_has_compartments(admin,X9,X10)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax8)).
fof(c_0_29, hypothesis, (system_compartment_has_sso(system,compartmenta,sso_compartmenta)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax47)).
fof(c_0_30, axiom, (![X1]:![X5]:![X7]:![X3]:![X19]:![X20]:![X21]:(system_indi_is_oca(system,X7)=>(oca_compartment_is_compartment(X7,X5,X3,X19,X20,X21)=>(admin_indi_has_level(admin,X1,X3)=>admin_indi_has_level_for_compartment(admin,X1,X5))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax29)).
fof(c_0_31, hypothesis, (credit_admin_indi_has_credit(credit_admin,alice)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax73)).
fof(c_0_32, hypothesis, (sso_compartment_has_scg(sso_compartmenta,compartmenta,scg_compartmenta)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax49)).
fof(c_0_33, hypothesis, (polygraph_admin_indi_has_polygraph(polygraph_admin,alice)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax72)).
fof(c_0_34, hypothesis, (background_admin_indi_has_background(background_admin,alice,topsecret)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax74)).
fof(c_0_35, hypothesis, (system_indi_is_background_admin(system,background_admin)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax68)).
fof(c_0_36, hypothesis, (system_file_needs_compartments(system,secretfile,cons(compartmentb,cons(compartmenta,nil)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax51)).
fof(c_0_37, hypothesis, (sso_file_has_compartments(sso_compartmentb,secretfile,cons(compartmentb,cons(compartmenta,nil)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax52)).
fof(c_0_38, hypothesis, (sso_file_has_level(sso_compartmenta,secretfile,secret,scg_compartmenta)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax56)).
fof(c_0_39, hypothesis, (oca_compartment_has_scg(oca,compartmenta,scg_compartmenta)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax48)).
fof(c_0_40, hypothesis, (system_indi_needs_compartment(system,alice,compartmentb)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax78)).
fof(c_0_41, hypothesis, (system_indi_needs_level(system,alice,secret)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax76)).
fof(c_0_42, hypothesis, (sso_file_has_compartments(sso_compartmenta,secretfile,cons(compartmentb,cons(compartmenta,nil)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax53)).
fof(c_0_43, axiom, (![X9]:![X10]:admin_file_has_compartments_h(admin,X9,X10,nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax9)).
fof(c_0_44, axiom, (![X1]:![X9]:![X2]:(admin_file_has_level(admin,X9,X2)=>(admin_indi_has_level(admin,X1,X2)=>admin_indi_has_level_for_file(admin,X1,X9)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax35)).
fof(c_0_45, hypothesis, (level_admin_indi_has_level(level_admin,alice,topsecret)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax77)).
fof(c_0_46, hypothesis, (system_indi_is_level_admin(system,level_admin)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax70)).
fof(c_0_47, axiom, (![X9]:![X2]:admin_file_has_level_h(admin,X9,X2,nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax12)).
fof(c_0_48, axiom, (![X1]:![X9]:![X10]:(admin_file_has_compartments(admin,X9,X10)=>(admin_indi_has_compartments(admin,X1,X10)=>admin_indi_has_compartments_for_file(admin,X1,X9)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax34)).
fof(c_0_49, hypothesis, (sso_indi_has_compartment(sso_compartmentb,alice,compartmentb)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax80)).
fof(c_0_50, axiom, (![X1]:![X2]:loca_level_below(X1,X2,X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax4)).
fof(c_0_51, axiom, (![X1]:![X9]:(admin_indi_has_citizenship(admin,X1,usa)=>admin_indi_has_citizenship_for_file(admin,X1,X9))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax38)).
fof(c_0_52, axiom, (![X1]:![X9]:![X22]:(state_file_has_owner(X9,X22)=>(owner_indi_has_need_to_know(X22,X1,X9)=>admin_indi_has_need_to_know_for_file(admin,X1,X9)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax36)).
fof(c_0_53, axiom, (![X1]:![X9]:(state_file_is_not_working_paper(X9)=>(admin_indi_has_citizenship_for_file(admin,X1,X9)=>(admin_indi_has_need_to_know_for_file(admin,X1,X9)=>(admin_indi_has_level_for_file(admin,X1,X9)=>(admin_indi_has_compartments_for_file(admin,X1,X9)=>admin_indi_may_file(admin,X1,X9,read))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax39)).
fof(c_0_54, hypothesis, (owner_indi_has_need_to_know(owner_secretfile,alice,secretfile)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax82)).
fof(c_0_55, hypothesis, (state_file_has_owner(secretfile,owner_secretfile)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax60)).
fof(c_0_56, conjecture, (admin_indi_may_file(admin,alice,secretfile,read)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', alicereadsecret)).
fof(c_0_57, axiom, (![X1]:![X5]:![X7]:![X3]:![X19]:![X20]:![X21]:(system_indi_is_oca(system,X7)=>(oca_compartment_is_compartment(X7,X5,X3,X19,X20,X21)=>(admin_indi_has_background(admin,X1,X19)=>admin_indi_has_background_for_compartment(admin,X1,X5))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax28)).
fof(c_0_58, axiom, (![X1]:admin_indi_has_background(admin,X1,unclassified)), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax19)).
fof(c_0_59, axiom, (![X1]:![X5]:![X7]:![X3]:![X19]:![X21]:(system_indi_is_oca(system,X7)=>(oca_compartment_is_compartment(X7,X5,X3,X19,no,X21)=>admin_indi_has_credit_for_compartment(admin,X1,X5)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax33)).
fof(c_0_60, axiom, (![X1]:![X5]:![X7]:![X3]:![X19]:![X20]:(system_indi_is_oca(system,X7)=>(oca_compartment_is_compartment(X7,X5,X3,X19,X20,no)=>admin_indi_has_polygraph_for_compartment(admin,X1,X5)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax31)).
fof(c_0_61, hypothesis, (state_file_is_not_working_paper(secretfile)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax50)).
fof(c_0_62, hypothesis, (oca_compartment_is_compartment(oca,compartmenta,sbu,unclassified,no,no)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax43)).
fof(c_0_63, hypothesis, (system_indi_needs_compartment(system,alice,compartmenta)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax79)).
fof(c_0_64, hypothesis, (sso_indi_has_compartment(sso_compartmenta,alice,compartmenta)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax81)).
fof(c_0_65, axiom, (![X1]:admin_indi_has_compartments(admin,X1,nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax26)).
fof(c_0_66, axiom, (![X1]:loca_level_direct_below(X1,confidential,secret)), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax2)).
fof(c_0_67, axiom, (![X1]:loca_level_direct_below(X1,sbu,confidential)), file('/export/starexec/sandbox/benchmark/Axioms/SWV009+0.ax', ax1)).
fof(c_0_68, plain, (![X9]:![X10]:![X11]:![X12]:(~system_indi_is_oca(system,X9)|(~oca_compartment_has_scg(X9,X10,X12)|(~admin_compartment_has_sso(admin,X10,X11)|(~sso_compartment_has_scg(X11,X10,X12)|admin_compartment_has_scg(admin,X10,X12)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])).
fof(c_0_69, plain, (![X7]:![X8]:(~system_compartment_has_sso(system,X7,X8)|admin_compartment_has_sso(admin,X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
cnf(c_0_70,plain,(admin_compartment_has_scg(admin,X1,X2)|~sso_compartment_has_scg(X3,X1,X2)|~admin_compartment_has_sso(admin,X1,X3)|~oca_compartment_has_scg(X4,X1,X2)|~system_indi_is_oca(system,X4)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_71,plain,(admin_compartment_has_sso(admin,X1,X2)|~system_compartment_has_sso(system,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_72,plain,(admin_compartment_has_scg(admin,X1,X2)|~sso_compartment_has_scg(X3,X1,X2)|~oca_compartment_has_scg(X4,X1,X2)|~system_indi_is_oca(system,X4)|~system_compartment_has_sso(system,X1,X3)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_73,hypothesis,(system_compartment_has_sso(system,compartmentb,sso_compartmentb)), inference(split_conjunct,[status(thm)],[c_0_2])).
fof(c_0_74, plain, (![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:(~admin_compartment_has_sso(admin,X13,X15)|(~admin_compartment_has_scg(admin,X13,X16)|(~sso_file_has_level(X15,X11,X12,X16)|(~admin_file_has_level_h(admin,X11,X12,X14)|admin_file_has_level_h(admin,X11,X12,cons(X13,X14))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_75,hypothesis,(admin_compartment_has_scg(admin,compartmentb,X1)|~sso_compartment_has_scg(sso_compartmentb,compartmentb,X1)|~oca_compartment_has_scg(X2,compartmentb,X1)|~system_indi_is_oca(system,X2)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_76,hypothesis,(sso_compartment_has_scg(sso_compartmentb,compartmentb,scg_compartmentb)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_77, plain, (![X14]:![X15]:(~system_indi_has_citizenship(system,X14,X15)|admin_indi_has_citizenship(admin,X14,X15))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
cnf(c_0_78,plain,(admin_file_has_level_h(admin,X1,X2,cons(X3,X4))|~admin_file_has_level_h(admin,X1,X2,X4)|~sso_file_has_level(X5,X1,X2,X6)|~admin_compartment_has_scg(admin,X3,X6)|~admin_compartment_has_sso(admin,X3,X5)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_79,hypothesis,(sso_file_has_level(sso_compartmentb,secretfile,secret,scg_compartmentb)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_80,hypothesis,(admin_compartment_has_scg(admin,compartmentb,scg_compartmentb)|~oca_compartment_has_scg(X1,compartmentb,scg_compartmentb)|~system_indi_is_oca(system,X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_81,hypothesis,(oca_compartment_has_scg(oca,compartmentb,scg_compartmentb)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_82,hypothesis,(system_indi_is_oca(system,oca)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_83, plain, (![X19]:![X20]:![X21]:![X22]:![X23]:(~system_indi_needs_level(system,X19,X21)|(~admin_indi_has_citizenship(admin,X19,usa)|(~admin_indi_has_polygraph(admin,X19)|(~admin_indi_has_employment(admin,X19)|(~admin_indi_has_credit(admin,X19)|(~loca_level_below(admin,X20,X21)|(~system_indi_is_level_admin(system,X22)|(~level_admin_indi_has_level(X22,X19,X23)|(~loca_level_below(admin,X20,X23)|(~admin_indi_has_background(admin,X19,X20)|admin_indi_has_level(admin,X19,X20)))))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
cnf(c_0_84,plain,(admin_indi_has_citizenship(admin,X1,X2)|~system_indi_has_citizenship(system,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_85,hypothesis,(system_indi_has_citizenship(system,alice,usa)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_86, plain, (![X18]:![X19]:(~system_indi_is_hr_admin(system,X19)|(~hr_admin_indi_has_employment(X19,X18)|admin_indi_has_employment(admin,X18)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).
cnf(c_0_87,hypothesis,(admin_file_has_level_h(admin,secretfile,secret,cons(X1,X2))|~admin_file_has_level_h(admin,secretfile,secret,X2)|~admin_compartment_has_scg(admin,X1,scg_compartmentb)|~admin_compartment_has_sso(admin,X1,sso_compartmentb)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_88,hypothesis,(admin_compartment_has_scg(admin,compartmentb,scg_compartmentb)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82])])).
fof(c_0_89, plain, (![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:(~system_indi_is_oca(system,X24)|(~oca_compartment_is_compartment(X24,X23,X25,X26,yes,X27)|(~admin_indi_has_credit(admin,X22)|admin_indi_has_credit_for_compartment(admin,X22,X23))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
fof(c_0_90, plain, (![X16]:![X17]:(~system_indi_is_credit_admin(system,X17)|(~credit_admin_indi_has_credit(X17,X16)|admin_indi_has_credit(admin,X16)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).
fof(c_0_91, plain, (![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:(~system_indi_is_oca(system,X23)|(~oca_compartment_is_compartment(X23,X22,X24,X25,X26,yes)|(~admin_indi_has_polygraph(admin,X21)|admin_indi_has_polygraph_for_compartment(admin,X21,X22))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).
fof(c_0_92, plain, (![X15]:![X16]:(~system_indi_is_polygraph_admin(system,X16)|(~polygraph_admin_indi_has_polygraph(X16,X15)|admin_indi_has_polygraph(admin,X15)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
cnf(c_0_93,plain,(admin_indi_has_level(admin,X1,X2)|~admin_indi_has_background(admin,X1,X2)|~loca_level_below(admin,X2,X3)|~level_admin_indi_has_level(X4,X1,X3)|~system_indi_is_level_admin(system,X4)|~loca_level_below(admin,X2,X5)|~admin_indi_has_credit(admin,X1)|~admin_indi_has_employment(admin,X1)|~admin_indi_has_polygraph(admin,X1)|~admin_indi_has_citizenship(admin,X1,usa)|~system_indi_needs_level(system,X1,X5)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_94,hypothesis,(admin_indi_has_citizenship(admin,alice,usa)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_95,plain,(admin_indi_has_employment(admin,X1)|~hr_admin_indi_has_employment(X2,X1)|~system_indi_is_hr_admin(system,X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_96,hypothesis,(system_indi_is_hr_admin(system,hr_admin)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_97, plain, (![X5]:![X6]:![X7]:![X8]:(~loca_level_direct_below(X5,X7,X8)|(~loca_level_below(X5,X6,X7)|loca_level_below(X5,X6,X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
fof(c_0_98, plain, (![X2]:loca_level_direct_below(X2,secret,topsecret)), inference(variable_rename,[status(thm)],[c_0_18])).
fof(c_0_99, plain, (![X11]:![X12]:![X13]:(~system_file_needs_level(system,X11,X12)|(~admin_file_has_compartments(admin,X11,X13)|(~admin_file_has_level_h(admin,X11,X12,X13)|admin_file_has_level(admin,X11,X12))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
cnf(c_0_100,hypothesis,(admin_file_has_level_h(admin,secretfile,secret,cons(compartmentb,X1))|~admin_file_has_level_h(admin,secretfile,secret,X1)|~admin_compartment_has_sso(admin,compartmentb,sso_compartmentb)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
fof(c_0_101, plain, (![X13]:![X14]:![X15]:![X16]:![X17]:(~admin_compartment_has_sso(admin,X15,X17)|(~sso_file_has_compartments(X17,X13,X14)|(~admin_file_has_compartments_h(admin,X13,X14,X16)|admin_file_has_compartments_h(admin,X13,X14,cons(X15,X16)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).
fof(c_0_102, plain, (![X11]:![X12]:![X13]:![X14]:(~system_indi_needs_compartment(system,X11,X12)|(~admin_indi_has_employment(admin,X11)|(~admin_indi_has_citizenship(admin,X11,usa)|(~admin_indi_has_polygraph_for_compartment(admin,X11,X12)|(~admin_indi_has_credit_for_compartment(admin,X11,X12)|(~admin_compartment_has_sso(admin,X12,X14)|(~sso_indi_has_compartment(X14,X11,X12)|(~admin_indi_has_background_for_compartment(admin,X11,X12)|(~admin_indi_has_level_for_compartment(admin,X11,X12)|(~admin_indi_has_compartments(admin,X11,X13)|admin_indi_has_compartments(admin,X11,cons(X12,X13))))))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])).
cnf(c_0_103,plain,(admin_indi_has_credit_for_compartment(admin,X1,X2)|~admin_indi_has_credit(admin,X1)|~oca_compartment_is_compartment(X3,X2,X4,X5,yes,X6)|~system_indi_is_oca(system,X3)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_104,hypothesis,(oca_compartment_is_compartment(oca,compartmentb,confidential,topsecret,yes,yes)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_105,plain,(admin_indi_has_credit(admin,X1)|~credit_admin_indi_has_credit(X2,X1)|~system_indi_is_credit_admin(system,X2)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_106,hypothesis,(system_indi_is_credit_admin(system,credit_admin)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_107,plain,(admin_indi_has_polygraph_for_compartment(admin,X1,X2)|~admin_indi_has_polygraph(admin,X1)|~oca_compartment_is_compartment(X3,X2,X4,X5,X6,yes)|~system_indi_is_oca(system,X3)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_108,plain,(admin_indi_has_polygraph(admin,X1)|~polygraph_admin_indi_has_polygraph(X2,X1)|~system_indi_is_polygraph_admin(system,X2)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_109,hypothesis,(system_indi_is_polygraph_admin(system,polygraph_admin)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_110,hypothesis,(admin_indi_has_level(admin,alice,X1)|~level_admin_indi_has_level(X2,alice,X3)|~system_indi_is_level_admin(system,X2)|~system_indi_needs_level(system,alice,X4)|~admin_indi_has_employment(admin,alice)|~admin_indi_has_background(admin,alice,X1)|~admin_indi_has_credit(admin,alice)|~admin_indi_has_polygraph(admin,alice)|~loca_level_below(admin,X1,X4)|~loca_level_below(admin,X1,X3)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_111,hypothesis,(admin_indi_has_employment(admin,X1)|~hr_admin_indi_has_employment(hr_admin,X1)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_112,hypothesis,(hr_admin_indi_has_employment(hr_admin,alice)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_113, plain, (![X17]:![X18]:![X19]:![X20]:(~system_indi_is_background_admin(system,X19)|(~background_admin_indi_has_background(X19,X17,X20)|(~loca_level_below(admin,X18,X20)|admin_indi_has_background(admin,X17,X18))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])])).
cnf(c_0_114,plain,(loca_level_below(X1,X2,X3)|~loca_level_below(X1,X2,X4)|~loca_level_direct_below(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_115,plain,(loca_level_direct_below(X1,secret,topsecret)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_116,plain,(admin_file_has_level(admin,X1,X2)|~admin_file_has_level_h(admin,X1,X2,X3)|~admin_file_has_compartments(admin,X1,X3)|~system_file_needs_level(system,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_117,hypothesis,(admin_file_has_level_h(admin,secretfile,secret,cons(compartmentb,X1))|~admin_file_has_level_h(admin,secretfile,secret,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_71]), c_0_73])])).
cnf(c_0_118,hypothesis,(system_file_needs_level(system,secretfile,secret)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_119, plain, (![X11]:![X12]:(~system_file_needs_compartments(system,X11,X12)|(~admin_file_has_compartments_h(admin,X11,X12,X12)|admin_file_has_compartments(admin,X11,X12)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])).
cnf(c_0_120,plain,(admin_file_has_compartments_h(admin,X1,X2,cons(X3,X4))|~admin_file_has_compartments_h(admin,X1,X2,X4)|~sso_file_has_compartments(X5,X1,X2)|~admin_compartment_has_sso(admin,X3,X5)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_121,hypothesis,(system_compartment_has_sso(system,compartmenta,sso_compartmenta)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_122,plain,(admin_indi_has_compartments(admin,X1,cons(X2,X3))|~admin_indi_has_compartments(admin,X1,X3)|~admin_indi_has_level_for_compartment(admin,X1,X2)|~admin_indi_has_background_for_compartment(admin,X1,X2)|~sso_indi_has_compartment(X4,X1,X2)|~admin_compartment_has_sso(admin,X2,X4)|~admin_indi_has_credit_for_compartment(admin,X1,X2)|~admin_indi_has_polygraph_for_compartment(admin,X1,X2)|~admin_indi_has_citizenship(admin,X1,usa)|~admin_indi_has_employment(admin,X1)|~system_indi_needs_compartment(system,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
fof(c_0_123, plain, (![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:(~system_indi_is_oca(system,X24)|(~oca_compartment_is_compartment(X24,X23,X25,X26,X27,X28)|(~admin_indi_has_level(admin,X22,X25)|admin_indi_has_level_for_compartment(admin,X22,X23))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])).
cnf(c_0_124,hypothesis,(admin_indi_has_credit_for_compartment(admin,X1,compartmentb)|~admin_indi_has_credit(admin,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_82])])).
cnf(c_0_125,hypothesis,(admin_indi_has_credit(admin,X1)|~credit_admin_indi_has_credit(credit_admin,X1)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_126,hypothesis,(admin_indi_has_polygraph_for_compartment(admin,X1,compartmentb)|~admin_indi_has_polygraph(admin,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_104]), c_0_82])])).
cnf(c_0_127,hypothesis,(admin_indi_has_polygraph(admin,X1)|~polygraph_admin_indi_has_polygraph(polygraph_admin,X1)), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_128,hypothesis,(admin_indi_has_level(admin,alice,X1)|~level_admin_indi_has_level(X2,alice,X3)|~system_indi_is_level_admin(system,X2)|~system_indi_needs_level(system,alice,X4)|~admin_indi_has_background(admin,alice,X1)|~admin_indi_has_credit(admin,alice)|~admin_indi_has_polygraph(admin,alice)|~loca_level_below(admin,X1,X4)|~loca_level_below(admin,X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112])])).
cnf(c_0_129,hypothesis,(credit_admin_indi_has_credit(credit_admin,alice)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_130,plain,(admin_indi_has_background(admin,X1,X2)|~loca_level_below(admin,X2,X3)|~background_admin_indi_has_background(X4,X1,X3)|~system_indi_is_background_admin(system,X4)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_131,plain,(loca_level_below(X1,X2,topsecret)|~loca_level_below(X1,X2,secret)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_132,hypothesis,(admin_file_has_level(admin,secretfile,secret)|~admin_file_has_level_h(admin,secretfile,secret,X1)|~admin_file_has_compartments(admin,secretfile,cons(compartmentb,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_118])])).
cnf(c_0_133,plain,(admin_file_has_compartments(admin,X1,X2)|~admin_file_has_compartments_h(admin,X1,X2,X2)|~system_file_needs_compartments(system,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_134,plain,(admin_file_has_compartments_h(admin,X1,X2,cons(X3,X4))|~sso_file_has_compartments(X5,X1,X2)|~admin_file_has_compartments_h(admin,X1,X2,X4)|~system_compartment_has_sso(system,X3,X5)), inference(spm,[status(thm)],[c_0_120, c_0_71])).
cnf(c_0_135,hypothesis,(admin_compartment_has_scg(admin,compartmenta,X1)|~sso_compartment_has_scg(sso_compartmenta,compartmenta,X1)|~oca_compartment_has_scg(X2,compartmenta,X1)|~system_indi_is_oca(system,X2)), inference(spm,[status(thm)],[c_0_72, c_0_121])).
cnf(c_0_136,hypothesis,(sso_compartment_has_scg(sso_compartmenta,compartmenta,scg_compartmenta)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_137,hypothesis,(admin_indi_has_compartments(admin,alice,cons(X1,X2))|~admin_indi_has_level_for_compartment(admin,alice,X1)|~admin_indi_has_background_for_compartment(admin,alice,X1)|~sso_indi_has_compartment(X3,alice,X1)|~admin_indi_has_credit_for_compartment(admin,alice,X1)|~admin_indi_has_polygraph_for_compartment(admin,alice,X1)|~system_indi_needs_compartment(system,alice,X1)|~admin_indi_has_compartments(admin,alice,X2)|~admin_indi_has_employment(admin,alice)|~admin_compartment_has_sso(admin,X1,X3)), inference(spm,[status(thm)],[c_0_122, c_0_94])).
cnf(c_0_138,plain,(admin_indi_has_level_for_compartment(admin,X1,X2)|~admin_indi_has_level(admin,X1,X3)|~oca_compartment_is_compartment(X4,X2,X3,X5,X6,X7)|~system_indi_is_oca(system,X4)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_139,hypothesis,(admin_indi_has_credit_for_compartment(admin,X1,compartmentb)|~credit_admin_indi_has_credit(credit_admin,X1)), inference(spm,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_140,hypothesis,(admin_indi_has_polygraph_for_compartment(admin,X1,compartmentb)|~polygraph_admin_indi_has_polygraph(polygraph_admin,X1)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_141,hypothesis,(polygraph_admin_indi_has_polygraph(polygraph_admin,alice)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_142,hypothesis,(admin_indi_has_level(admin,alice,X1)|~level_admin_indi_has_level(X2,alice,X3)|~system_indi_is_level_admin(system,X2)|~system_indi_needs_level(system,alice,X4)|~admin_indi_has_background(admin,alice,X1)|~admin_indi_has_polygraph(admin,alice)|~loca_level_below(admin,X1,X4)|~loca_level_below(admin,X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_125]), c_0_129])])).
cnf(c_0_143,plain,(admin_indi_has_background(admin,X1,X2)|~background_admin_indi_has_background(X3,X1,topsecret)|~system_indi_is_background_admin(system,X3)|~loca_level_below(admin,X2,secret)), inference(spm,[status(thm)],[c_0_130, c_0_131])).
cnf(c_0_144,hypothesis,(background_admin_indi_has_background(background_admin,alice,topsecret)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_145,hypothesis,(system_indi_is_background_admin(system,background_admin)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_146,hypothesis,(admin_file_has_level(admin,secretfile,secret)|~admin_file_has_level_h(admin,secretfile,secret,X1)|~admin_file_has_compartments_h(admin,secretfile,cons(compartmentb,X1),cons(compartmentb,X1))|~system_file_needs_compartments(system,secretfile,cons(compartmentb,X1))), inference(spm,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_147,hypothesis,(system_file_needs_compartments(system,secretfile,cons(compartmentb,cons(compartmenta,nil)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_148,hypothesis,(admin_file_has_compartments_h(admin,X1,X2,cons(compartmentb,X3))|~sso_file_has_compartments(sso_compartmentb,X1,X2)|~admin_file_has_compartments_h(admin,X1,X2,X3)), inference(spm,[status(thm)],[c_0_134, c_0_73])).
cnf(c_0_149,hypothesis,(sso_file_has_compartments(sso_compartmentb,secretfile,cons(compartmentb,cons(compartmenta,nil)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_150,hypothesis,(sso_file_has_level(sso_compartmenta,secretfile,secret,scg_compartmenta)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_151,hypothesis,(admin_compartment_has_scg(admin,compartmenta,scg_compartmenta)|~oca_compartment_has_scg(X1,compartmenta,scg_compartmenta)|~system_indi_is_oca(system,X1)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_152,hypothesis,(oca_compartment_has_scg(oca,compartmenta,scg_compartmenta)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_153,hypothesis,(admin_indi_has_compartments(admin,alice,cons(X1,X2))|~admin_indi_has_level_for_compartment(admin,alice,X1)|~admin_indi_has_background_for_compartment(admin,alice,X1)|~sso_indi_has_compartment(X3,alice,X1)|~admin_indi_has_credit_for_compartment(admin,alice,X1)|~admin_indi_has_polygraph_for_compartment(admin,alice,X1)|~system_indi_needs_compartment(system,alice,X1)|~admin_indi_has_compartments(admin,alice,X2)|~admin_compartment_has_sso(admin,X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_111]), c_0_112])])).
cnf(c_0_154,hypothesis,(admin_indi_has_level_for_compartment(admin,X1,compartmentb)|~admin_indi_has_level(admin,X1,confidential)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_104]), c_0_82])])).
cnf(c_0_155,hypothesis,(admin_indi_has_credit_for_compartment(admin,alice,compartmentb)), inference(spm,[status(thm)],[c_0_139, c_0_129])).
cnf(c_0_156,hypothesis,(admin_indi_has_polygraph_for_compartment(admin,alice,compartmentb)), inference(spm,[status(thm)],[c_0_140, c_0_141])).
cnf(c_0_157,hypothesis,(system_indi_needs_compartment(system,alice,compartmentb)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_158,hypothesis,(admin_indi_has_level(admin,alice,X1)|~level_admin_indi_has_level(X2,alice,X3)|~system_indi_is_level_admin(system,X2)|~system_indi_needs_level(system,alice,X4)|~admin_indi_has_background(admin,alice,X1)|~loca_level_below(admin,X1,X4)|~loca_level_below(admin,X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_127]), c_0_141])])).
cnf(c_0_159,hypothesis,(system_indi_needs_level(system,alice,secret)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_160,hypothesis,(admin_indi_has_background(admin,alice,X1)|~loca_level_below(admin,X1,secret)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_145])])).
cnf(c_0_161,hypothesis,(admin_file_has_level(admin,secretfile,secret)|~admin_file_has_level_h(admin,secretfile,secret,cons(compartmenta,nil))|~admin_file_has_compartments_h(admin,secretfile,cons(compartmentb,cons(compartmenta,nil)),cons(compartmentb,cons(compartmenta,nil)))), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_162,hypothesis,(admin_file_has_compartments_h(admin,secretfile,cons(compartmentb,cons(compartmenta,nil)),cons(compartmentb,X1))|~admin_file_has_compartments_h(admin,secretfile,cons(compartmentb,cons(compartmenta,nil)),X1)), inference(spm,[status(thm)],[c_0_148, c_0_149])).
cnf(c_0_163,hypothesis,(admin_file_has_compartments_h(admin,X1,X2,cons(compartmenta,X3))|~sso_file_has_compartments(sso_compartmenta,X1,X2)|~admin_file_has_compartments_h(admin,X1,X2,X3)), inference(spm,[status(thm)],[c_0_134, c_0_121])).
cnf(c_0_164,hypothesis,(sso_file_has_compartments(sso_compartmenta,secretfile,cons(compartmentb,cons(compartmenta,nil)))), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_165, plain, (![X11]:![X12]:admin_file_has_compartments_h(admin,X11,X12,nil)), inference(variable_rename,[status(thm)],[c_0_43])).
cnf(c_0_166,hypothesis,(admin_file_has_level_h(admin,secretfile,secret,cons(X1,X2))|~admin_file_has_level_h(admin,secretfile,secret,X2)|~admin_compartment_has_scg(admin,X1,scg_compartmenta)|~admin_compartment_has_sso(admin,X1,sso_compartmenta)), inference(spm,[status(thm)],[c_0_78, c_0_150])).
cnf(c_0_167,hypothesis,(admin_compartment_has_scg(admin,compartmenta,scg_compartmenta)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_82])])).
cnf(c_0_168,hypothesis,(admin_indi_has_compartments(admin,alice,cons(compartmentb,X1))|~admin_indi_has_background_for_compartment(admin,alice,compartmentb)|~sso_indi_has_compartment(X2,alice,compartmentb)|~admin_indi_has_compartments(admin,alice,X1)|~admin_indi_has_level(admin,alice,confidential)|~admin_compartment_has_sso(admin,compartmentb,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_155]), c_0_156]), c_0_157])])).
fof(c_0_169, plain, (![X10]:![X11]:![X12]:(~admin_file_has_level(admin,X11,X12)|(~admin_indi_has_level(admin,X10,X12)|admin_indi_has_level_for_file(admin,X10,X11)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_44])])).
cnf(c_0_170,hypothesis,(admin_indi_has_level(admin,alice,X1)|~level_admin_indi_has_level(X2,alice,X3)|~system_indi_is_level_admin(system,X2)|~loca_level_below(admin,X1,secret)|~loca_level_below(admin,X1,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_160])).
cnf(c_0_171,hypothesis,(level_admin_indi_has_level(level_admin,alice,topsecret)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_172,hypothesis,(system_indi_is_level_admin(system,level_admin)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_173,hypothesis,(admin_file_has_level(admin,secretfile,secret)|~admin_file_has_level_h(admin,secretfile,secret,cons(compartmenta,nil))|~admin_file_has_compartments_h(admin,secretfile,cons(compartmentb,cons(compartmenta,nil)),cons(compartmenta,nil))), inference(spm,[status(thm)],[c_0_161, c_0_162])).
cnf(c_0_174,hypothesis,(admin_file_has_compartments_h(admin,secretfile,cons(compartmentb,cons(compartmenta,nil)),cons(compartmenta,X1))|~admin_file_has_compartments_h(admin,secretfile,cons(compartmentb,cons(compartmenta,nil)),X1)), inference(spm,[status(thm)],[c_0_163, c_0_164])).
cnf(c_0_175,plain,(admin_file_has_compartments_h(admin,X1,X2,nil)), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_176,hypothesis,(admin_file_has_level_h(admin,secretfile,secret,cons(compartmenta,X1))|~admin_file_has_level_h(admin,secretfile,secret,X1)|~admin_compartment_has_sso(admin,compartmenta,sso_compartmenta)), inference(spm,[status(thm)],[c_0_166, c_0_167])).
fof(c_0_177, plain, (![X10]:![X11]:admin_file_has_level_h(admin,X10,X11,nil)), inference(variable_rename,[status(thm)],[c_0_47])).
fof(c_0_178, plain, (![X11]:![X12]:![X13]:(~admin_file_has_compartments(admin,X12,X13)|(~admin_indi_has_compartments(admin,X11,X13)|admin_indi_has_compartments_for_file(admin,X11,X12)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_48])])).
cnf(c_0_179,hypothesis,(admin_indi_has_compartments(admin,alice,cons(compartmentb,X1))|~admin_indi_has_background_for_compartment(admin,alice,compartmentb)|~sso_indi_has_compartment(X2,alice,compartmentb)|~admin_indi_has_compartments(admin,alice,X1)|~admin_indi_has_level(admin,alice,confidential)|~system_compartment_has_sso(system,compartmentb,X2)), inference(spm,[status(thm)],[c_0_168, c_0_71])).
cnf(c_0_180,hypothesis,(sso_indi_has_compartment(sso_compartmentb,alice,compartmentb)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_181,plain,(admin_indi_has_level_for_file(admin,X1,X2)|~admin_indi_has_level(admin,X1,X3)|~admin_file_has_level(admin,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_182,hypothesis,(admin_indi_has_level(admin,alice,X1)|~loca_level_below(admin,X1,secret)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171]), c_0_172])]), c_0_131])).
cnf(c_0_183,hypothesis,(admin_file_has_level(admin,secretfile,secret)|~admin_file_has_level_h(admin,secretfile,secret,cons(compartmenta,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_175])])).
cnf(c_0_184,hypothesis,(admin_file_has_level_h(admin,secretfile,secret,cons(compartmenta,X1))|~admin_file_has_level_h(admin,secretfile,secret,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_71]), c_0_121])])).
cnf(c_0_185,plain,(admin_file_has_level_h(admin,X1,X2,nil)), inference(split_conjunct,[status(thm)],[c_0_177])).
fof(c_0_186, plain, (![X3]:![X4]:loca_level_below(X3,X4,X4)), inference(variable_rename,[status(thm)],[c_0_50])).
fof(c_0_187, plain, (![X10]:![X11]:(~admin_indi_has_citizenship(admin,X10,usa)|admin_indi_has_citizenship_for_file(admin,X10,X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_51])])])])).
fof(c_0_188, plain, (![X23]:![X24]:![X25]:(~state_file_has_owner(X24,X25)|(~owner_indi_has_need_to_know(X25,X23,X24)|admin_indi_has_need_to_know_for_file(admin,X23,X24)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_52])])).
cnf(c_0_189,plain,(admin_indi_has_compartments_for_file(admin,X1,X2)|~admin_indi_has_compartments(admin,X1,X3)|~admin_file_has_compartments(admin,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_178])).
cnf(c_0_190,hypothesis,(admin_indi_has_compartments(admin,alice,cons(compartmentb,X1))|~admin_indi_has_background_for_compartment(admin,alice,compartmentb)|~admin_indi_has_compartments(admin,alice,X1)|~admin_indi_has_level(admin,alice,confidential)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_73])])).
fof(c_0_191, plain, (![X10]:![X11]:(~state_file_is_not_working_paper(X11)|(~admin_indi_has_citizenship_for_file(admin,X10,X11)|(~admin_indi_has_need_to_know_for_file(admin,X10,X11)|(~admin_indi_has_level_for_file(admin,X10,X11)|(~admin_indi_has_compartments_for_file(admin,X10,X11)|admin_indi_may_file(admin,X10,X11,read))))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_53])])).
cnf(c_0_192,hypothesis,(admin_indi_has_level_for_file(admin,alice,X1)|~admin_file_has_level(admin,X1,X2)|~loca_level_below(admin,X2,secret)), inference(spm,[status(thm)],[c_0_181, c_0_182])).
cnf(c_0_193,hypothesis,(admin_file_has_level(admin,secretfile,secret)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184]), c_0_185])])).
cnf(c_0_194,plain,(loca_level_below(X1,X2,X2)), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_195,plain,(admin_indi_has_citizenship_for_file(admin,X1,X2)|~admin_indi_has_citizenship(admin,X1,usa)), inference(split_conjunct,[status(thm)],[c_0_187])).
cnf(c_0_196,plain,(admin_indi_has_need_to_know_for_file(admin,X1,X2)|~owner_indi_has_need_to_know(X3,X1,X2)|~state_file_has_owner(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_197,hypothesis,(owner_indi_has_need_to_know(owner_secretfile,alice,secretfile)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_198,hypothesis,(state_file_has_owner(secretfile,owner_secretfile)), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_199, negated_conjecture, (~admin_indi_may_file(admin,alice,secretfile,read)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_56])])).
fof(c_0_200, plain, (![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:(~system_indi_is_oca(system,X24)|(~oca_compartment_is_compartment(X24,X23,X25,X26,X27,X28)|(~admin_indi_has_background(admin,X22,X26)|admin_indi_has_background_for_compartment(admin,X22,X23))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_57])])])])).
fof(c_0_201, plain, (![X2]:admin_indi_has_background(admin,X2,unclassified)), inference(variable_rename,[status(thm)],[c_0_58])).
fof(c_0_202, plain, (![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:(~system_indi_is_oca(system,X24)|(~oca_compartment_is_compartment(X24,X23,X25,X26,no,X27)|admin_indi_has_credit_for_compartment(admin,X22,X23)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_59])])])])).
fof(c_0_203, plain, (![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:(~system_indi_is_oca(system,X23)|(~oca_compartment_is_compartment(X23,X22,X24,X25,X26,no)|admin_indi_has_polygraph_for_compartment(admin,X21,X22)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_60])])])])).
cnf(c_0_204,hypothesis,(admin_indi_has_compartments_for_file(admin,alice,X1)|~admin_indi_has_background_for_compartment(admin,alice,compartmentb)|~admin_indi_has_compartments(admin,alice,X2)|~admin_indi_has_level(admin,alice,confidential)|~admin_file_has_compartments(admin,X1,cons(compartmentb,X2))), inference(spm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_205,plain,(admin_indi_may_file(admin,X1,X2,read)|~admin_indi_has_compartments_for_file(admin,X1,X2)|~admin_indi_has_level_for_file(admin,X1,X2)|~admin_indi_has_need_to_know_for_file(admin,X1,X2)|~admin_indi_has_citizenship_for_file(admin,X1,X2)|~state_file_is_not_working_paper(X2)), inference(split_conjunct,[status(thm)],[c_0_191])).
cnf(c_0_206,hypothesis,(admin_indi_has_level_for_file(admin,alice,secretfile)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193]), c_0_194])])).
cnf(c_0_207,hypothesis,(state_file_is_not_working_paper(secretfile)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_208,hypothesis,(admin_indi_has_citizenship_for_file(admin,alice,X1)), inference(spm,[status(thm)],[c_0_195, c_0_94])).
cnf(c_0_209,hypothesis,(admin_indi_has_need_to_know_for_file(admin,alice,secretfile)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_197]), c_0_198])])).
cnf(c_0_210,negated_conjecture,(~admin_indi_may_file(admin,alice,secretfile,read)), inference(split_conjunct,[status(thm)],[c_0_199])).
cnf(c_0_211,hypothesis,(oca_compartment_is_compartment(oca,compartmenta,sbu,unclassified,no,no)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_212,plain,(admin_indi_has_background_for_compartment(admin,X1,X2)|~admin_indi_has_background(admin,X1,X3)|~oca_compartment_is_compartment(X4,X2,X5,X3,X6,X7)|~system_indi_is_oca(system,X4)), inference(split_conjunct,[status(thm)],[c_0_200])).
cnf(c_0_213,plain,(admin_indi_has_background(admin,X1,unclassified)), inference(split_conjunct,[status(thm)],[c_0_201])).
cnf(c_0_214,plain,(admin_indi_has_credit_for_compartment(admin,X1,X2)|~oca_compartment_is_compartment(X3,X2,X4,X5,no,X6)|~system_indi_is_oca(system,X3)), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_215,plain,(admin_indi_has_polygraph_for_compartment(admin,X1,X2)|~oca_compartment_is_compartment(X3,X2,X4,X5,X6,no)|~system_indi_is_oca(system,X3)), inference(split_conjunct,[status(thm)],[c_0_203])).
cnf(c_0_216,hypothesis,(admin_indi_has_compartments_for_file(admin,alice,X1)|~admin_indi_has_background_for_compartment(admin,alice,compartmentb)|~admin_indi_has_compartments(admin,alice,X2)|~admin_indi_has_level(admin,alice,confidential)|~admin_file_has_compartments_h(admin,X1,cons(compartmentb,X2),cons(compartmentb,X2))|~system_file_needs_compartments(system,X1,cons(compartmentb,X2))), inference(spm,[status(thm)],[c_0_204, c_0_133])).
cnf(c_0_217,hypothesis,(~admin_indi_has_compartments_for_file(admin,alice,secretfile)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206]), c_0_207]), c_0_208]), c_0_209])]), c_0_210])).
cnf(c_0_218,hypothesis,(admin_indi_has_level_for_compartment(admin,X1,compartmenta)|~admin_indi_has_level(admin,X1,sbu)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_211]), c_0_82])])).
cnf(c_0_219,hypothesis,(admin_indi_has_background_for_compartment(admin,X1,compartmenta)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_211]), c_0_213]), c_0_82])])).
cnf(c_0_220,hypothesis,(admin_indi_has_credit_for_compartment(admin,X1,compartmenta)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_211]), c_0_82])])).
cnf(c_0_221,hypothesis,(admin_indi_has_polygraph_for_compartment(admin,X1,compartmenta)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_211]), c_0_82])])).
cnf(c_0_222,hypothesis,(system_indi_needs_compartment(system,alice,compartmenta)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_223,hypothesis,(~admin_indi_has_background_for_compartment(admin,alice,compartmentb)|~admin_indi_has_compartments(admin,alice,cons(compartmenta,nil))|~admin_indi_has_level(admin,alice,confidential)|~admin_file_has_compartments_h(admin,secretfile,cons(compartmentb,cons(compartmenta,nil)),cons(compartmentb,cons(compartmenta,nil)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_147]), c_0_217])).
cnf(c_0_224,hypothesis,(admin_indi_has_compartments(admin,alice,cons(compartmenta,X1))|~sso_indi_has_compartment(X2,alice,compartmenta)|~admin_indi_has_compartments(admin,alice,X1)|~admin_indi_has_level(admin,alice,sbu)|~admin_compartment_has_sso(admin,compartmenta,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_218]), c_0_219]), c_0_220]), c_0_221]), c_0_222])])).
cnf(c_0_225,hypothesis,(~admin_indi_has_background_for_compartment(admin,alice,compartmentb)|~admin_indi_has_compartments(admin,alice,cons(compartmenta,nil))|~admin_indi_has_level(admin,alice,confidential)|~admin_file_has_compartments_h(admin,secretfile,cons(compartmentb,cons(compartmenta,nil)),cons(compartmenta,nil))), inference(spm,[status(thm)],[c_0_223, c_0_162])).
cnf(c_0_226,hypothesis,(admin_indi_has_compartments(admin,alice,cons(compartmenta,X1))|~sso_indi_has_compartment(X2,alice,compartmenta)|~admin_indi_has_compartments(admin,alice,X1)|~admin_indi_has_level(admin,alice,sbu)|~system_compartment_has_sso(system,compartmenta,X2)), inference(spm,[status(thm)],[c_0_224, c_0_71])).
cnf(c_0_227,hypothesis,(sso_indi_has_compartment(sso_compartmenta,alice,compartmenta)), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_228, plain, (![X2]:admin_indi_has_compartments(admin,X2,nil)), inference(variable_rename,[status(thm)],[c_0_65])).
cnf(c_0_229,hypothesis,(~admin_indi_has_background_for_compartment(admin,alice,compartmentb)|~admin_indi_has_compartments(admin,alice,cons(compartmenta,nil))|~admin_indi_has_level(admin,alice,confidential)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_174]), c_0_175])])).
cnf(c_0_230,hypothesis,(admin_indi_has_compartments(admin,alice,cons(compartmenta,X1))|~admin_indi_has_compartments(admin,alice,X1)|~admin_indi_has_level(admin,alice,sbu)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227]), c_0_121])])).
cnf(c_0_231,plain,(admin_indi_has_compartments(admin,X1,nil)), inference(split_conjunct,[status(thm)],[c_0_228])).
cnf(c_0_232,plain,(admin_indi_has_background(admin,X1,X2)|~background_admin_indi_has_background(X3,X1,X2)|~system_indi_is_background_admin(system,X3)), inference(spm,[status(thm)],[c_0_130, c_0_194])).
cnf(c_0_233,hypothesis,(~admin_indi_has_background_for_compartment(admin,alice,compartmentb)|~admin_indi_has_level(admin,alice,confidential)|~admin_indi_has_level(admin,alice,sbu)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230]), c_0_231])])).
cnf(c_0_234,hypothesis,(admin_indi_has_background_for_compartment(admin,X1,compartmentb)|~admin_indi_has_background(admin,X1,topsecret)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_104]), c_0_82])])).
cnf(c_0_235,hypothesis,(admin_indi_has_background(admin,alice,topsecret)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_144]), c_0_145])])).
cnf(c_0_236,hypothesis,(~admin_indi_has_level(admin,alice,confidential)|~admin_indi_has_level(admin,alice,sbu)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_234]), c_0_235])])).
fof(c_0_237, plain, (![X2]:loca_level_direct_below(X2,confidential,secret)), inference(variable_rename,[status(thm)],[c_0_66])).
cnf(c_0_238,hypothesis,(~admin_indi_has_level(admin,alice,sbu)|~loca_level_below(admin,confidential,secret)), inference(spm,[status(thm)],[c_0_236, c_0_182])).
cnf(c_0_239,plain,(loca_level_direct_below(X1,confidential,secret)), inference(split_conjunct,[status(thm)],[c_0_237])).
cnf(c_0_240,hypothesis,(~loca_level_below(admin,confidential,secret)|~loca_level_below(admin,sbu,secret)), inference(spm,[status(thm)],[c_0_238, c_0_182])).
cnf(c_0_241,plain,(loca_level_below(X1,X2,secret)|~loca_level_below(X1,X2,confidential)), inference(spm,[status(thm)],[c_0_114, c_0_239])).
fof(c_0_242, plain, (![X2]:loca_level_direct_below(X2,sbu,confidential)), inference(variable_rename,[status(thm)],[c_0_67])).
cnf(c_0_243,hypothesis,(~loca_level_below(admin,sbu,secret)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241]), c_0_194])])).
cnf(c_0_244,plain,(loca_level_direct_below(X1,sbu,confidential)), inference(split_conjunct,[status(thm)],[c_0_242])).
cnf(c_0_245,hypothesis,(~loca_level_below(admin,sbu,confidential)), inference(spm,[status(thm)],[c_0_243, c_0_241])).
cnf(c_0_246,plain,(loca_level_below(X1,X2,confidential)|~loca_level_below(X1,X2,sbu)), inference(spm,[status(thm)],[c_0_114, c_0_244])).
cnf(c_0_247,hypothesis,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246]), c_0_194])]), ['proof']).
# SZS output end CNFRefutation.
EOF
