% Problem    : KRS234+1 : TPTP v6.1.0. Bugfixed v5.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n100.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:19:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C18_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X7]:![X8]:(![X1]:![X2]:(status(X1,X2,X7)=>~(status(X1,X2,X8)))<=>nevera(X7,X8))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', nevera)).
fof(c_0_1, axiom, (![X9]:![X10]:(model(X9,X10)<~>model(X9,not(X10)))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', completeness)).
fof(c_0_2, axiom, (![X1]:![X2]:((![X3]:model(X3,X1)&![X4]:model(X4,not(X2)))<=>status(X1,X2,uns))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', uns)).
fof(c_0_3, axiom, (![X1]:![X2]:(![X3]:(model(X3,X1)=>model(X3,X2))<=>status(X1,X2,thm))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', thm)).
fof(c_0_4, conjecture, (nevera(uns,thm)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', nevera_uns_thm)).
fof(c_0_5, plain, (![X7]:![X8]:(![X1]:![X2]:(status(X1,X2,X7)=>~status(X1,X2,X8))<=>nevera(X7,X8))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_6, plain, (![X9]:![X10]:~((model(X9,X10)<=>model(X9,not(X10))))), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_7, plain, (![X5]:![X6]:![X9]:![X10]:![X11]:![X12]:(((~model(esk39_2(X5,X6),X5)|~model(esk40_2(X5,X6),not(X6)))|status(X5,X6,uns))&((model(X11,X9)|~status(X9,X10,uns))&(model(X12,not(X10))|~status(X9,X10,uns))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
fof(c_0_8, plain, (![X9]:![X10]:![X13]:![X14]:![X15]:![X16]:(((status(esk49_2(X9,X10),esk50_2(X9,X10),X9)|nevera(X9,X10))&(status(esk49_2(X9,X10),esk50_2(X9,X10),X10)|nevera(X9,X10)))&(~nevera(X13,X14)|(~status(X15,X16,X13)|~status(X15,X16,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).
fof(c_0_9, plain, (![X4]:![X5]:![X7]:![X8]:![X9]:(((model(esk10_2(X4,X5),X4)|status(X4,X5,thm))&(~model(esk10_2(X4,X5),X5)|status(X4,X5,thm)))&(~status(X7,X8,thm)|(~model(X9,X7)|model(X9,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])])).
fof(c_0_10, plain, (![X11]:![X12]:![X13]:![X14]:((~model(X11,X12)|~model(X11,not(X12)))&(model(X13,X14)|model(X13,not(X14))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_11,plain,(model(X3,not(X2))|~status(X1,X2,uns)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(nevera(X1,X2)|status(esk49_2(X1,X2),esk50_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(model(X1,X2)|~model(X1,X3)|~status(X3,X2,thm)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(nevera(X1,X2)|status(esk49_2(X1,X2),esk50_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15,plain,(model(X3,X1)|~status(X1,X2,uns)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_16, negated_conjecture, (~nevera(uns,thm)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_4])])).
cnf(c_0_17,plain,(~model(X1,not(X2))|~model(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,plain,(nevera(uns,X1)|model(X2,not(esk50_2(uns,X1)))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_19,plain,(nevera(X1,thm)|model(X2,esk50_2(X1,thm))|~model(X2,esk49_2(X1,thm))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20,plain,(nevera(uns,X1)|model(X2,esk49_2(uns,X1))), inference(spm,[status(thm)],[c_0_15, c_0_12])).
cnf(c_0_21,negated_conjecture,(~nevera(uns,thm)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(nevera(uns,X1)|~model(X2,esk50_2(uns,X1))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23,plain,(model(X1,esk50_2(uns,thm))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_24,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_21]), ['proof']).
# SZS output end CNFRefutation.
EOF
