% Problem    : CSR115+27 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n192.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:58:21 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:?[X8]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&has_card_leq(X8,int1994))&obj(X5,X1))&sub(X2,name_1_1))&sub(X3,name_1_1))&sub(X7,jahr__1_1))&val(X7,X8))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_news_1195)).
fof(c_0_1, axiom, (![X1]:![X2]:(card(X1,X2)=>has_card_leq(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', has_card_eq)).
fof(c_0_2, hypothesis, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((attr(c11,c12)&sub(c11,stadt__1_1))&sub(c12,name_1_1))&val(c12,m__374nchen_0))&attr(c18,c19))&attr(c18,c20))&sub(c19,tag_1_1))&val(c19,c16))&sub(c20,monat_1_1))&val(c20,c17))&tupl_p7(c319,c68,c72,c79,c84,c91,c95))&sub(c324,mark_1_1))&tupl(c47,c11,c18))&sub(c68,bmw_1_1))&sub(c72,firmengruppe_1_1))&sub(c79,konzerngewinn_1_1))&attr(c84,c85))&sub(c85,jahr__1_1))&val(c85,c80))&quant_p3(c91,c86,hundertstel__1_1))&quant_p3(c95,c92,million_1_1))&assoc(konzerngewinn_1_1,firmengruppe_1_1))&sub(konzerngewinn_1_1,gewinn_1_1))&sort(c11,d))&sort(c11,io))&card(c11,int1))&etype(c11,int0))&fact(c11,real))&gener(c11,sp))&quant(c11,one))&refer(c11,det))&varia(c11,con))&sort(c12,na))&card(c12,int1))&etype(c12,int0))&fact(c12,real))&gener(c12,sp))&quant(c12,one))&refer(c12,indet))&varia(c12,varia_c))&sort(stadt__1_1,d))&sort(stadt__1_1,io))&card(stadt__1_1,int1))&etype(stadt__1_1,int0))&fact(stadt__1_1,real))&gener(stadt__1_1,ge))&quant(stadt__1_1,one))&refer(stadt__1_1,refer_c))&varia(stadt__1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(m__374nchen_0,fe))&sort(c18,t))&card(c18,int1))&etype(c18,int0))&fact(c18,real))&gener(c18,sp))&quant(c18,one))&refer(c18,det))&varia(c18,con))&sort(c19,me))&sort(c19,oa))&sort(c19,ta))&card(c19,card_c))&etype(c19,etype_c))&fact(c19,real))&gener(c19,sp))&quant(c19,quant_c))&refer(c19,refer_c))&varia(c19,varia_c))&sort(c20,me))&sort(c20,oa))&sort(c20,ta))&card(c20,card_c))&etype(c20,etype_c))&fact(c20,real))&gener(c20,sp))&quant(c20,quant_c))&refer(c20,refer_c))&varia(c20,varia_c))&sort(tag_1_1,me))&sort(tag_1_1,oa))&sort(tag_1_1,ta))&card(tag_1_1,card_c))&etype(tag_1_1,etype_c))&fact(tag_1_1,real))&gener(tag_1_1,ge))&quant(tag_1_1,quant_c))&refer(tag_1_1,refer_c))&varia(tag_1_1,varia_c))&sort(c16,nu))&card(c16,int16))&sort(monat_1_1,me))&sort(monat_1_1,oa))&sort(monat_1_1,ta))&card(monat_1_1,card_c))&etype(monat_1_1,etype_c))&fact(monat_1_1,real))&gener(monat_1_1,ge))&quant(monat_1_1,quant_c))&refer(monat_1_1,refer_c))&varia(monat_1_1,varia_c))&sort(c17,nu))&card(c17,int3))&sort(c319,ent))&card(c319,card_c))&etype(c319,etype_c))&fact(c319,real))&gener(c319,gener_c))&quant(c319,quant_c))&refer(c319,refer_c))&varia(c319,varia_c))&sort(c68,d))&card(c68,int1))&etype(c68,int0))&fact(c68,real))&gener(c68,sp))&quant(c68,one))&refer(c68,det))&varia(c68,con))&sort(c72,d))&sort(c72,io))&card(c72,int1))&etype(c72,int0))&fact(c72,real))&gener(c72,gener_c))&quant(c72,one))&refer(c72,refer_c))&varia(c72,varia_c))&sort(c79,d))&card(c79,int1))&etype(c79,int0))&fact(c79,real))&gener(c79,sp))&quant(c79,one))&refer(c79,det))&varia(c79,con))&sort(c84,t))&card(c84,int1))&etype(c84,int0))&fact(c84,real))&gener(c84,sp))&quant(c84,one))&refer(c84,det))&varia(c84,con))&sort(c91,co))&sort(c91,m))&card(c91,card_c))&etype(c91,etype_c))&fact(c91,real))&gener(c91,gener_c))&quant(c91,quant_c))&refer(c91,refer_c))&varia(c91,con))&sort(c95,co))&sort(c95,m))&card(c95,card_c))&etype(c95,etype_c))&fact(c95,real))&gener(c95,gener_c))&quant(c95,quant_c))&refer(c95,refer_c))&varia(c95,con))&sort(c324,me))&fact(c324,real))&sort(mark_1_1,me))&gener(mark_1_1,ge))&sort(c47,ent))&card(c47,card_c))&etype(c47,etype_c))&fact(c47,real))&gener(c47,gener_c))&quant(c47,quant_c))&refer(c47,refer_c))&varia(c47,varia_c))&sort(bmw_1_1,d))&card(bmw_1_1,int1))&etype(bmw_1_1,int0))&fact(bmw_1_1,real))&gener(bmw_1_1,ge))&quant(bmw_1_1,one))&refer(bmw_1_1,refer_c))&varia(bmw_1_1,varia_c))&sort(firmengruppe_1_1,d))&sort(firmengruppe_1_1,io))&card(firmengruppe_1_1,int1))&etype(firmengruppe_1_1,int0))&fact(firmengruppe_1_1,real))&gener(firmengruppe_1_1,ge))&quant(firmengruppe_1_1,one))&refer(firmengruppe_1_1,refer_c))&varia(firmengruppe_1_1,varia_c))&sort(konzerngewinn_1_1,d))&card(konzerngewinn_1_1,int1))&etype(konzerngewinn_1_1,int0))&fact(konzerngewinn_1_1,real))&gener(konzerngewinn_1_1,ge))&quant(konzerngewinn_1_1,one))&refer(konzerngewinn_1_1,refer_c))&varia(konzerngewinn_1_1,varia_c))&sort(c85,me))&sort(c85,oa))&sort(c85,ta))&card(c85,card_c))&etype(c85,etype_c))&fact(c85,real))&gener(c85,sp))&quant(c85,quant_c))&refer(c85,refer_c))&varia(c85,varia_c))&sort(jahr__1_1,me))&sort(jahr__1_1,oa))&sort(jahr__1_1,ta))&card(jahr__1_1,card_c))&etype(jahr__1_1,etype_c))&fact(jahr__1_1,real))&gener(jahr__1_1,ge))&quant(jahr__1_1,quant_c))&refer(jahr__1_1,refer_c))&varia(jahr__1_1,varia_c))&sort(c80,nu))&card(c80,int1994))&sort(c86,nu))&card(c86,int35))&sort(hundertstel__1_1,me))&gener(hundertstel__1_1,ge))&sort(c92,nu))&card(c92,int697))&sort(million_1_1,me))&gener(million_1_1,ge))&sort(gewinn_1_1,d))&card(gewinn_1_1,int1))&etype(gewinn_1_1,int0))&fact(gewinn_1_1,real))&gener(gewinn_1_1,ge))&quant(gewinn_1_1,one))&refer(gewinn_1_1,refer_c))&varia(gewinn_1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_news_1195)).
fof(c_0_3, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_4, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_5, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:?[X8]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&has_card_leq(X8,int1994))&obj(X5,X1))&sub(X2,name_1_1))&sub(X3,name_1_1))&sub(X7,jahr__1_1))&val(X7,X8)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_6, negated_conjecture, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:((((((((~attr(X9,X10)|~attr(X12,X11))|~attr(X14,X15))|~has_card_leq(X16,int1994))|~obj(X13,X9))|~sub(X10,name_1_1))|~sub(X11,name_1_1))|~sub(X15,jahr__1_1))|~val(X15,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
fof(c_0_7, plain, (![X3]:![X4]:(~card(X3,X4)|has_card_leq(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
cnf(c_0_8,negated_conjecture,(~val(X1,X2)|~sub(X1,jahr__1_1)|~sub(X3,name_1_1)|~sub(X4,name_1_1)|~obj(X5,X6)|~has_card_leq(X2,int1994)|~attr(X7,X1)|~attr(X8,X3)|~attr(X6,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,plain,(has_card_leq(X1,X2)|~card(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,negated_conjecture,(~val(X1,X2)|~attr(X3,X4)|~attr(X5,X1)|~attr(X6,X7)|~obj(X8,X6)|~sub(X7,name_1_1)|~sub(X4,name_1_1)|~sub(X1,jahr__1_1)|~card(X2,int1994)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_11,hypothesis,(val(c85,c80)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_12,hypothesis,(sub(c85,jahr__1_1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_13,hypothesis,(card(c80,int1994)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_14,hypothesis,(~attr(X1,c85)|~attr(X2,X3)|~attr(X4,X5)|~obj(X6,X4)|~sub(X5,name_1_1)|~sub(X3,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_15,hypothesis,(attr(c84,c85)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_16,hypothesis,(~attr(X1,X2)|~attr(X3,X4)|~obj(X5,X3)|~sub(X4,name_1_1)|~sub(X2,name_1_1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_17,hypothesis,(attr(c11,c12)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_18,hypothesis,(sub(c12,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_19,hypothesis,(~attr(X1,X2)|~obj(X3,X1)|~sub(X2,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])])).
fof(c_0_20, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk45_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk45_3(X7,X8,X9),esk46_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk44_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk44_3(X7,X8,X9),esk45_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk44_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk46_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk45_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk44_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_21,hypothesis,(~obj(X1,c11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_17]), c_0_18])])).
cnf(c_0_22,plain,(obj(esk44_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_23, plain, (![X4]:![X5]:(((arg1(esk47_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk47_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk47_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_24,hypothesis,(~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,c11)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25,plain,(subr(esk47_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_26,hypothesis,(~arg2(esk47_2(X1,X2),X3)|~arg1(esk47_2(X1,X2),c11)|~sub(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_27,plain,(arg2(esk47_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,hypothesis,(~arg1(esk47_2(X1,X2),c11)|~sub(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_29,plain,(arg1(esk47_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30,hypothesis,(sub(c11,stadt__1_1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_31,hypothesis,(~sub(c11,X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_32,hypothesis,($false), inference(sr,[status(thm)],[c_0_30, c_0_31]), ['proof']).
# SZS output end CNFRefutation.
EOF
