% Problem    : CSR115+42 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n191.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:37:16 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_news_1246)).
fof(c_0_1, hypothesis, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((attr(c2253,c2254)&sub(c2253,firma_1_1))&sub(c2254,name_1_1))&val(c2254,bmw_0))&prop(c2257,britisch__1_1))&sub(c2257,c2259))&pred(c2258,es_2_1))&pmod(c2259,letzt_1_1,massenhersteller_1_1))&sub(c2279,seite_1_1))&attch(c2283,c2279))&sub(c2283,medaille_1_1))&tupl_p5(c2419,c2253,c2257,c2258,c2279))&assoc(massenhersteller_1_1,masse_1_1))&sub(massenhersteller_1_1,fabrikant_1_1))&sort(c2253,d))&sort(c2253,io))&card(c2253,int1))&etype(c2253,int0))&fact(c2253,real))&gener(c2253,sp))&quant(c2253,one))&refer(c2253,det))&varia(c2253,con))&sort(c2254,na))&card(c2254,int1))&etype(c2254,int0))&fact(c2254,real))&gener(c2254,sp))&quant(c2254,one))&refer(c2254,indet))&varia(c2254,varia_c))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(bmw_0,fe))&sort(c2257,d))&sort(c2257,io))&card(c2257,int1))&etype(c2257,int0))&fact(c2257,real))&gener(c2257,gener_c))&quant(c2257,one))&refer(c2257,refer_c))&varia(c2257,varia_c))&sort(britisch__1_1,nq))&sort(c2259,d))&sort(c2259,io))&card(c2259,int1))&etype(c2259,int0))&fact(c2259,real))&gener(c2259,ge))&quant(c2259,one))&refer(c2259,refer_c))&varia(c2259,varia_c))&sort(c2258,o))&card(c2258,cons(x_constant,cons(int1,nil))))&etype(c2258,int1))&fact(c2258,real))&gener(c2258,gener_c))&quant(c2258,mult))&refer(c2258,indet))&varia(c2258,varia_c))&sort(es_2_1,o))&card(es_2_1,int1))&etype(es_2_1,int0))&fact(es_2_1,real))&gener(es_2_1,ge))&quant(es_2_1,one))&refer(es_2_1,refer_c))&varia(es_2_1,varia_c))&sort(letzt_1_1,oq))&card(letzt_1_1,card_c))&sort(massenhersteller_1_1,d))&sort(massenhersteller_1_1,io))&card(massenhersteller_1_1,int1))&etype(massenhersteller_1_1,int0))&fact(massenhersteller_1_1,real))&gener(massenhersteller_1_1,ge))&quant(massenhersteller_1_1,one))&refer(massenhersteller_1_1,refer_c))&varia(massenhersteller_1_1,varia_c))&sort(c2279,d))&sort(c2279,io))&card(c2279,int1))&etype(c2279,int0))&fact(c2279,real))&gener(c2279,sp))&quant(c2279,one))&refer(c2279,indet))&varia(c2279,varia_c))&sort(seite_1_1,d))&sort(seite_1_1,io))&card(seite_1_1,int1))&etype(seite_1_1,int0))&fact(seite_1_1,real))&gener(seite_1_1,ge))&quant(seite_1_1,one))&refer(seite_1_1,refer_c))&varia(seite_1_1,varia_c))&sort(c2283,d))&card(c2283,int1))&etype(c2283,int0))&fact(c2283,real))&gener(c2283,sp))&quant(c2283,one))&refer(c2283,det))&varia(c2283,con))&sort(medaille_1_1,d))&card(medaille_1_1,int1))&etype(medaille_1_1,int0))&fact(medaille_1_1,real))&gener(medaille_1_1,ge))&quant(medaille_1_1,one))&refer(medaille_1_1,refer_c))&varia(medaille_1_1,varia_c))&sort(c2419,ent))&card(c2419,card_c))&etype(c2419,etype_c))&fact(c2419,real))&gener(c2419,gener_c))&quant(c2419,quant_c))&refer(c2419,refer_c))&varia(c2419,varia_c))&sort(masse_1_1,io))&card(masse_1_1,card_c))&etype(masse_1_1,int1))&fact(masse_1_1,real))&gener(masse_1_1,ge))&quant(masse_1_1,quant_c))&refer(masse_1_1,refer_c))&varia(masse_1_1,varia_c))&sort(fabrikant_1_1,d))&sort(fabrikant_1_1,io))&card(fabrikant_1_1,int1))&etype(fabrikant_1_1,int0))&fact(fabrikant_1_1,real))&gener(fabrikant_1_1,ge))&quant(fabrikant_1_1,one))&refer(fabrikant_1_1,refer_c))&varia(fabrikant_1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_news_1246)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_3, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_4, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, negated_conjecture, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((((~attr(X8,X9)|~attr(X11,X10))|~attr(X12,X13))|~obj(X14,X8))|~sub(X9,name_1_1))|~sub(X8,firma_1_1))|~sub(X10,name_1_1))|~val(X9,bmw_0))|~val(X10,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_6,negated_conjecture,(~val(X1,bmw_0)|~val(X2,bmw_0)|~sub(X1,name_1_1)|~sub(X3,firma_1_1)|~sub(X2,name_1_1)|~obj(X4,X3)|~attr(X5,X6)|~attr(X7,X1)|~attr(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,hypothesis,(val(c2254,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(sub(c2254,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(~val(X1,bmw_0)|~attr(X2,c2254)|~attr(X3,X1)|~attr(X4,X5)|~obj(X6,X2)|~sub(X2,firma_1_1)|~sub(X1,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_10,hypothesis,(~attr(X1,c2254)|~attr(X2,c2254)|~attr(X3,X4)|~obj(X5,X1)|~sub(X1,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_8])])).
cnf(c_0_11,hypothesis,(attr(c2253,c2254)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_12,hypothesis,(sub(c2253,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(~attr(X1,c2254)|~attr(X2,X3)|~obj(X4,c2253)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk41_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk41_3(X7,X8,X9),esk42_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk40_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk40_3(X7,X8,X9),esk41_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk40_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk42_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk41_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk40_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_15,hypothesis,(~attr(X1,X2)|~obj(X3,c2253)), inference(spm,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_16,plain,(obj(esk40_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, (![X4]:![X5]:(((arg1(esk43_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk43_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk43_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_18,hypothesis,(~subr(X1,sub_0)|~attr(X2,X3)|~arg2(X1,X4)|~arg1(X1,c2253)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(subr(esk43_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,hypothesis,(~attr(X1,X2)|~arg2(esk43_2(X3,X4),X5)|~arg1(esk43_2(X3,X4),c2253)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(arg2(esk43_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,hypothesis,(~attr(X1,X2)|~arg1(esk43_2(X3,X4),c2253)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,plain,(arg1(esk43_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,hypothesis,(~attr(X1,X2)|~sub(c2253,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,hypothesis,(~attr(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_12])).
cnf(c_0_26,hypothesis,($false), inference(sr,[status(thm)],[c_0_11, c_0_25]), ['proof']).
# SZS output end CNFRefutation.
EOF
