% Problem    : KRS217+1 : TPTP v6.1.0. Bugfixed v5.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n109.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:18:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C18_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X7]:![X8]:(![X1]:![X2]:(status(X1,X2,X7)=>~(status(X1,X2,X8)))<=>nevera(X7,X8))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', nevera)).
fof(c_0_1, axiom, (![X7]:![X8]:(?[X1]:?[X2]:(status(X1,X2,X7)&~(status(X1,X2,X8)))<=>nota(X7,X8))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', nota)).
fof(c_0_2, axiom, (?[X1]:?[X2]:(((?[X3]:model(X3,X1)&![X4]:(model(X4,X1)=>model(X4,X2)))&?[X5]:(~(model(X5,X1))&model(X5,X2)))&?[X6]:~(model(X6,X2)))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', mixed_pair)).
fof(c_0_3, axiom, (![X1]:![X2]:((~(?[X3]:model(X3,X1))=>~(?[X4]:model(X4,X2)))<=>status(X1,X2,unp))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', unp)).
fof(c_0_4, axiom, (![X1]:![X2]:((~(?[X3]:model(X3,X1))&?[X4]:model(X4,X2))<=>status(X1,X2,sca))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', sca)).
fof(c_0_5, axiom, (![X9]:![X10]:(model(X9,X10)<~>model(X9,not(X10)))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', completeness)).
fof(c_0_6, axiom, (![X1]:![X2]:((?[X3]:(model(X3,X1)&model(X3,X2))&?[X4]:(model(X4,X1)&model(X4,not(X2))))<=>status(X1,X2,noc))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', noc)).
fof(c_0_7, axiom, (![X1]:![X2]:(~(?[X3]:model(X3,X1))<=>status(X1,X2,cax))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', cax)).
fof(c_0_8, axiom, (?[X3]:?[X1]:?[X2]:((model(X3,X1)&~(model(X3,X2)))&?[X4]:model(X4,X2))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', non_thm_spt)).
fof(c_0_9, axiom, (![X1]:![X2]:(![X3]:(model(X3,X1)=>model(X3,X2))<=>status(X1,X2,thm))), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+0.ax', thm)).
fof(c_0_10, axiom, (?[X10]:![X9]:model(X9,X10)), file('/export/starexec/sandbox/benchmark/Axioms/KRS001+1.ax', tautology)).
fof(c_0_11, conjecture, (nota(noc,thm)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', nota_noc_thm)).
fof(c_0_12, plain, (![X7]:![X8]:(![X1]:![X2]:(status(X1,X2,X7)=>~status(X1,X2,X8))<=>nevera(X7,X8))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_13, plain, (![X7]:![X8]:(?[X1]:?[X2]:(status(X1,X2,X7)&~status(X1,X2,X8))<=>nota(X7,X8))), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_14, plain, (?[X1]:?[X2]:(((?[X3]:model(X3,X1)&![X4]:(model(X4,X1)=>model(X4,X2)))&?[X5]:(~model(X5,X1)&model(X5,X2)))&?[X6]:~model(X6,X2))), inference(fof_simplification,[status(thm)],[c_0_2])).
fof(c_0_15, plain, (![X9]:![X10]:![X13]:![X14]:![X15]:![X16]:(((status(esk49_2(X9,X10),esk50_2(X9,X10),X9)|nevera(X9,X10))&(status(esk49_2(X9,X10),esk50_2(X9,X10),X10)|nevera(X9,X10)))&(~nevera(X13,X14)|(~status(X15,X16,X13)|~status(X15,X16,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])])).
fof(c_0_16, plain, (![X5]:![X6]:![X7]:![X9]:![X10]:![X12]:(((~model(X7,X5)|status(X5,X6,unp))&(model(esk1_2(X5,X6),X6)|status(X5,X6,unp)))&(~status(X9,X10,unp)|(model(esk2_2(X9,X10),X9)|~model(X12,X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])])).
fof(c_0_17, plain, (![X5]:![X6]:![X8]:![X9]:![X10]:![X11]:(((model(esk31_2(X5,X6),X5)|~model(X8,X6))|status(X5,X6,sca))&((~model(X11,X9)|~status(X9,X10,sca))&(model(esk32_2(X9,X10),X10)|~status(X9,X10,sca))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
fof(c_0_18, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:(((~status(X11,X12,X9)|status(X11,X12,X10))|nota(X9,X10))&((status(esk47_2(X13,X14),esk48_2(X13,X14),X13)|~nota(X13,X14))&(~status(esk47_2(X13,X14),esk48_2(X13,X14),X14)|~nota(X13,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])])).
fof(c_0_19, plain, (![X9]:![X10]:~((model(X9,X10)<=>model(X9,not(X10))))), inference(fof_simplification,[status(thm)],[c_0_5])).
fof(c_0_20, plain, (![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:((((~model(X7,X5)|~model(X7,X6))|(~model(X8,X5)|~model(X8,not(X6))))|status(X5,X6,noc))&(((model(esk41_2(X9,X10),X9)|~status(X9,X10,noc))&(model(esk41_2(X9,X10),X10)|~status(X9,X10,noc)))&((model(esk42_2(X9,X10),X9)|~status(X9,X10,noc))&(model(esk42_2(X9,X10),not(X10))|~status(X9,X10,noc)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])])).
fof(c_0_21, plain, (![X4]:![X6]:![X7]:![X8]:![X9]:((model(esk30_1(X4),X4)|status(X4,X6,cax))&(~status(X7,X8,cax)|~model(X9,X7)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
fof(c_0_22, plain, (?[X3]:?[X1]:?[X2]:((model(X3,X1)&~model(X3,X2))&?[X4]:model(X4,X2))), inference(fof_simplification,[status(thm)],[c_0_8])).
fof(c_0_23, plain, (![X10]:(((model(esk65_0,esk63_0)&(~model(X10,esk63_0)|model(X10,esk64_0)))&(~model(esk66_0,esk63_0)&model(esk66_0,esk64_0)))&~model(esk67_0,esk64_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).
cnf(c_0_24,plain,(~status(X1,X2,X3)|~status(X1,X2,X4)|~nevera(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(nevera(X1,X2)|status(esk49_2(X1,X2),esk50_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26,plain,(model(esk2_2(X3,X2),X3)|~model(X1,X2)|~status(X3,X2,unp)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27,plain,(model(esk32_2(X1,X2),X2)|~status(X1,X2,sca)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28,plain,(nevera(X1,X2)|status(esk49_2(X1,X2),esk50_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29,plain,(~status(X1,X2,sca)|~model(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30,plain,(nota(X1,X2)|status(X3,X4,X2)|~status(X3,X4,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_31, plain, (![X4]:![X5]:![X7]:![X8]:![X9]:(((model(esk10_2(X4,X5),X4)|status(X4,X5,thm))&(~model(esk10_2(X4,X5),X5)|status(X4,X5,thm)))&(~status(X7,X8,thm)|(~model(X9,X7)|model(X9,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])])).
fof(c_0_32, plain, (![X11]:![X12]:![X13]:![X14]:((~model(X11,X12)|~model(X11,not(X12)))&(model(X13,X14)|model(X13,not(X14))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
cnf(c_0_33,plain,(model(esk42_2(X1,X2),not(X2))|~status(X1,X2,noc)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_34,plain,(~model(X1,X2)|~status(X2,X3,cax)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_35,plain,(status(X1,X2,cax)|model(esk30_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_36, plain, (((model(esk68_0,esk69_0)&~model(esk68_0,esk70_0))&model(esk71_0,esk70_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_22])])).
cnf(c_0_37,plain,(status(X1,X2,unp)|~model(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_38,plain,(model(esk65_0,esk63_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_39,plain,(nevera(X1,X2)|~nevera(X2,X3)|~status(esk49_2(X1,X2),esk50_2(X1,X2),X3)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_40,plain,(nevera(X1,unp)|model(esk2_2(esk49_2(X1,unp),esk50_2(X1,unp)),esk49_2(X1,unp))|~model(X2,esk50_2(X1,unp))), inference(spm,[status(thm)],[c_0_26, c_0_25])).
cnf(c_0_41,plain,(nevera(sca,X1)|model(esk32_2(esk49_2(sca,X1),esk50_2(sca,X1)),esk50_2(sca,X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_42,plain,(nevera(sca,X1)|~model(X2,esk49_2(sca,X1))), inference(spm,[status(thm)],[c_0_29, c_0_28])).
fof(c_0_43, plain, (![X12]:model(X12,esk53_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_10])])).
cnf(c_0_44,plain,(nevera(X1,X2)|nota(X2,X3)|status(esk49_2(X1,X2),esk50_2(X1,X2),X3)), inference(spm,[status(thm)],[c_0_30, c_0_25])).
cnf(c_0_45,plain,(model(X1,X2)|~model(X1,X3)|~status(X3,X2,thm)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_46,plain,(model(esk42_2(X1,X2),X1)|~status(X1,X2,noc)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_47,plain,(~model(X1,not(X2))|~model(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_48,plain,(nevera(X1,noc)|model(esk42_2(esk49_2(X1,noc),esk50_2(X1,noc)),not(esk50_2(X1,noc)))), inference(spm,[status(thm)],[c_0_33, c_0_25])).
cnf(c_0_49,plain,(model(esk30_1(X1),X1)|~model(X2,X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_50,plain,(status(X1,X2,thm)|model(esk10_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_51,plain,(status(X1,X2,sca)|model(esk31_2(X1,X2),X1)|~model(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_52,plain,(model(esk71_0,esk70_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_53,plain,(status(esk63_0,X1,unp)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_54,plain,(nevera(X1,X2)|~nevera(X2,X1)), inference(spm,[status(thm)],[c_0_39, c_0_28])).
cnf(c_0_55,plain,(nevera(sca,unp)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_56,plain,(model(X1,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_57,plain,(nevera(X1,X2)|nota(X2,X3)|~nevera(X3,X4)|~status(esk49_2(X1,X2),esk50_2(X1,X2),X4)), inference(spm,[status(thm)],[c_0_24, c_0_44])).
cnf(c_0_58,plain,(nevera(thm,X1)|model(X2,esk50_2(thm,X1))|~model(X2,esk49_2(thm,X1))), inference(spm,[status(thm)],[c_0_45, c_0_28])).
cnf(c_0_59,plain,(nevera(X1,noc)|model(esk42_2(esk49_2(X1,noc),esk50_2(X1,noc)),esk49_2(X1,noc))), inference(spm,[status(thm)],[c_0_46, c_0_25])).
cnf(c_0_60,plain,(nevera(X1,noc)|~model(esk42_2(esk49_2(X1,noc),esk50_2(X1,noc)),esk50_2(X1,noc))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
fof(c_0_61, negated_conjecture, (~nota(noc,thm)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_11])])).
cnf(c_0_62,plain,(status(X1,X2,noc)|~model(X3,not(X2))|~model(X3,X1)|~model(X4,X2)|~model(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_63,plain,(status(X1,X2,thm)|model(esk30_1(X1),X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_64,plain,(model(X1,esk64_0)|~model(X1,esk63_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_65,plain,(status(X1,esk70_0,sca)|model(esk31_2(X1,esk70_0),X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_66,plain,(~nevera(unp,X1)|~status(esk63_0,X2,X1)), inference(spm,[status(thm)],[c_0_24, c_0_53])).
cnf(c_0_67,plain,(nevera(unp,sca)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_68,plain,(status(esk53_0,X1,unp)), inference(spm,[status(thm)],[c_0_37, c_0_56])).
cnf(c_0_69,plain,(nevera(X1,X2)|nota(X2,X3)|~nevera(X3,X2)), inference(spm,[status(thm)],[c_0_57, c_0_25])).
cnf(c_0_70,plain,(nevera(thm,noc)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60])).
cnf(c_0_71,negated_conjecture,(~nota(noc,thm)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_72,plain,(status(not(X1),X2,thm)|status(X3,X1,noc)|~model(esk30_1(not(X1)),X3)|~model(X4,X1)|~model(X4,X3)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_73,plain,(status(esk63_0,esk70_0,sca)|model(esk31_2(esk63_0,esk70_0),esk64_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74,plain,(~status(esk63_0,X1,sca)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_75,plain,(~nevera(unp,X1)|~status(esk53_0,X2,X1)), inference(spm,[status(thm)],[c_0_24, c_0_68])).
cnf(c_0_76,plain,(nevera(X1,noc)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])).
cnf(c_0_77,plain,(status(not(X1),X2,thm)|status(esk53_0,X1,noc)|~model(X3,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_56]), c_0_56])])).
cnf(c_0_78,plain,(model(esk31_2(esk63_0,esk70_0),esk64_0)), inference(sr,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_79,plain,(~status(esk53_0,X1,noc)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_80,plain,(status(not(esk64_0),X1,thm)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])).
cnf(c_0_81,plain,(model(X1,X2)|~model(X1,not(esk64_0))), inference(spm,[status(thm)],[c_0_45, c_0_80])).
cnf(c_0_82,plain,(model(X1,not(X2))|model(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_83,plain,(model(X1,esk64_0)|model(X1,X2)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_84,plain,(~model(esk67_0,esk64_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_85,plain,(model(X1,esk64_0)), inference(ef,[status(thm)],[c_0_83])).
cnf(c_0_86,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_85])]), ['proof']).
# SZS output end CNFRefutation.
EOF
