% Problem    : CSR115+100 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n100.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:44:46 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mw3_233_a462)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((sub(c36629,firma_1_1)&sub(c36633,linde_1_1))&pred(c36646,gas__1_1))&prop(c36646,technisch__1_1))&sub(c36652,teil_1_1))&sub(c36657,bmw_1_1))&pred(c36709,kabelbaum_1_1))&attr(c36711,c36712))&sub(c36711,firma_1_1))&sub(c36712,name_1_1))&val(c36712,audi_0))&sub(c36719,timi_1_1))&tupl_p9(c36870,c36629,c36633,c36646,c36652,c36657,c36711,c36709,c36719))&pred(c36914,dr__344xelmaier_1_1))&tupl(c36937,c36711,c36914))&assoc(kabelbaum_1_1,kabeln_1_1))&sub(kabelbaum_1_1,baum_1_1))&sort(c36629,d))&sort(c36629,io))&card(c36629,int1))&etype(c36629,int0))&fact(c36629,real))&gener(c36629,sp))&quant(c36629,one))&refer(c36629,det))&varia(c36629,con))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(c36633,d))&card(c36633,int1))&etype(c36633,int0))&fact(c36633,real))&gener(c36633,gener_c))&quant(c36633,one))&refer(c36633,refer_c))&varia(c36633,varia_c))&sort(linde_1_1,d))&card(linde_1_1,int1))&etype(linde_1_1,int0))&fact(linde_1_1,real))&gener(linde_1_1,ge))&quant(linde_1_1,one))&refer(linde_1_1,refer_c))&varia(linde_1_1,varia_c))&sort(c36646,s))&card(c36646,cons(x_constant,cons(int1,nil))))&etype(c36646,int1))&fact(c36646,real))&gener(c36646,gener_c))&quant(c36646,mult))&refer(c36646,refer_c))&varia(c36646,varia_c))&sort(gas__1_1,s))&card(gas__1_1,int1))&etype(gas__1_1,int0))&fact(gas__1_1,real))&gener(gas__1_1,ge))&quant(gas__1_1,one))&refer(gas__1_1,refer_c))&varia(gas__1_1,varia_c))&sort(technisch__1_1,nq))&sort(c36652,co))&card(c36652,card_c))&etype(c36652,etype_c))&fact(c36652,real))&gener(c36652,sp))&quant(c36652,quant_c))&refer(c36652,indet))&varia(c36652,varia_c))&sort(teil_1_1,co))&card(teil_1_1,card_c))&etype(teil_1_1,etype_c))&fact(teil_1_1,real))&gener(teil_1_1,ge))&quant(teil_1_1,quant_c))&refer(teil_1_1,refer_c))&varia(teil_1_1,varia_c))&sort(c36657,d))&card(c36657,int1))&etype(c36657,int0))&fact(c36657,real))&gener(c36657,sp))&quant(c36657,one))&refer(c36657,det))&varia(c36657,con))&sort(bmw_1_1,d))&card(bmw_1_1,int1))&etype(bmw_1_1,int0))&fact(bmw_1_1,real))&gener(bmw_1_1,ge))&quant(bmw_1_1,one))&refer(bmw_1_1,refer_c))&varia(bmw_1_1,varia_c))&sort(c36709,d))&card(c36709,cons(x_constant,cons(int1,nil))))&etype(c36709,int1))&fact(c36709,real))&gener(c36709,gener_c))&quant(c36709,mult))&refer(c36709,indet))&varia(c36709,varia_c))&sort(kabelbaum_1_1,d))&card(kabelbaum_1_1,int1))&etype(kabelbaum_1_1,int0))&fact(kabelbaum_1_1,real))&gener(kabelbaum_1_1,ge))&quant(kabelbaum_1_1,one))&refer(kabelbaum_1_1,refer_c))&varia(kabelbaum_1_1,varia_c))&sort(c36711,d))&sort(c36711,io))&card(c36711,int1))&etype(c36711,int0))&fact(c36711,real))&gener(c36711,sp))&quant(c36711,one))&refer(c36711,det))&varia(c36711,con))&sort(c36712,na))&card(c36712,int1))&etype(c36712,int0))&fact(c36712,real))&gener(c36712,sp))&quant(c36712,one))&refer(c36712,indet))&varia(c36712,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(audi_0,fe))&sort(c36719,o))&card(c36719,int1))&etype(c36719,int0))&fact(c36719,real))&gener(c36719,gener_c))&quant(c36719,one))&refer(c36719,refer_c))&varia(c36719,varia_c))&sort(timi_1_1,o))&card(timi_1_1,int1))&etype(timi_1_1,int0))&fact(timi_1_1,real))&gener(timi_1_1,ge))&quant(timi_1_1,one))&refer(timi_1_1,refer_c))&varia(timi_1_1,varia_c))&sort(c36870,ent))&card(c36870,card_c))&etype(c36870,etype_c))&fact(c36870,real))&gener(c36870,gener_c))&quant(c36870,quant_c))&refer(c36870,refer_c))&varia(c36870,varia_c))&sort(c36914,o))&card(c36914,cons(x_constant,cons(int1,nil))))&etype(c36914,int1))&fact(c36914,real))&gener(c36914,gener_c))&quant(c36914,mult))&refer(c36914,indet))&varia(c36914,varia_c))&sort(dr__344xelmaier_1_1,o))&card(dr__344xelmaier_1_1,int1))&etype(dr__344xelmaier_1_1,int0))&fact(dr__344xelmaier_1_1,real))&gener(dr__344xelmaier_1_1,ge))&quant(dr__344xelmaier_1_1,one))&refer(dr__344xelmaier_1_1,refer_c))&varia(dr__344xelmaier_1_1,varia_c))&sort(c36937,ent))&card(c36937,card_c))&etype(c36937,etype_c))&fact(c36937,real))&gener(c36937,gener_c))&quant(c36937,quant_c))&refer(c36937,refer_c))&varia(c36937,varia_c))&sort(kabeln_1_1,dn))&fact(kabeln_1_1,real))&gener(kabeln_1_1,ge))&sort(baum_1_1,d))&card(baum_1_1,int1))&etype(baum_1_1,int0))&fact(baum_1_1,real))&gener(baum_1_1,ge))&quant(baum_1_1,one))&refer(baum_1_1,refer_c))&varia(baum_1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mw3_233_a462)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_3, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_4, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, negated_conjecture, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((~attr(X8,X9)|~attr(X11,X10))|~attr(X12,X13))|~obj(X14,X8))|~sub(X9,name_1_1))|~sub(X8,firma_1_1))|~sub(X10,name_1_1))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_6,negated_conjecture,(~sub(X1,name_1_1)|~sub(X2,firma_1_1)|~sub(X3,name_1_1)|~obj(X4,X2)|~attr(X5,X6)|~attr(X7,X1)|~attr(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,hypothesis,(attr(c36711,c36712)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(sub(c36712,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(~attr(X1,X2)|~attr(X3,X4)|~obj(X5,X3)|~sub(X4,name_1_1)|~sub(X3,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_10,hypothesis,(~attr(X1,X2)|~obj(X3,X1)|~sub(X2,name_1_1)|~sub(X1,firma_1_1)), inference(spm,[status(thm)],[c_0_9, c_0_7])).
cnf(c_0_11,hypothesis,(sub(c36711,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_12, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk41_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk41_3(X7,X8,X9),esk42_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk40_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk40_3(X7,X8,X9),esk41_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk40_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk42_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk41_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk40_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_13,hypothesis,(~obj(X1,c36711)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_7]), c_0_8]), c_0_11])])).
cnf(c_0_14,plain,(obj(esk40_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_15, plain, (![X4]:![X5]:(((arg1(esk43_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk43_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk43_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_16,hypothesis,(~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,c36711)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17,plain,(subr(esk43_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18,hypothesis,(~arg2(esk43_2(X1,X2),X3)|~arg1(esk43_2(X1,X2),c36711)|~sub(X1,X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_19,plain,(arg2(esk43_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20,hypothesis,(~arg1(esk43_2(X1,X2),c36711)|~sub(X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(arg1(esk43_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,hypothesis,(~sub(c36711,X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,hypothesis,($false), inference(sr,[status(thm)],[c_0_11, c_0_22]), ['proof']).
# SZS output end CNFRefutation.
EOF
