% Problem    : CSR114+4 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n107.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:01:31 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:((((((((in(X3,X1)&attr(X1,X2))&loc(X5,X3))&scar(X5,X4))&sub(X2,name_1_1))&sub(X1,stadt__1_1))&sub(X4,kolosseum_1_1))&subs(X5,stehen_1_1))&val(X2,rom_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_004_insicht_3)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((chsp2(bekommen_1_1,c0)&attch(c12058,c7))&attr(c12058,c12059))&sub(c12058,stadt__1_1))&sub(c12059,name_1_1))&val(c12059,verona_0))&modp(c12087,zwei_1_1,c0))&arg1(c12093,c7))&arg2(c12093,c12602))&assoc(c12093,c12583))&semrel(c12093,c16))&subr(c12093,equ_0))&loc(c12583,c12608))&sub(c12583,kolosseum_1_1))&attr(c12594,c12595))&sub(c12594,stadt__1_1))&sub(c12595,name_1_1))&val(c12595,rom_0))&prop(c12602,c0))&prop(c12602,c12605))&sub(c12602,amphi_theater_1_1))&supl(c12605,zweitgro__337_1_1,c12606))&in(c12608,c12594))&arg1(c16,c7))&arg2(c16,c7))&subr(c16,sub_0))&equ(c7,c12602))&prop(c7,c12087))&prop(c7,r__366misch__1_1))&sub(c7,amphi_theater_1_1))&sub(c7,arena__1_1))&sub(c7,c7))&sort(bekommen_1_1,dn))&fact(bekommen_1_1,real))&gener(bekommen_1_1,ge))&sort(c0,tq))&sort(c12058,d))&sort(c12058,io))&card(c12058,int1))&etype(c12058,int0))&fact(c12058,real))&gener(c12058,sp))&quant(c12058,one))&refer(c12058,det))&varia(c12058,con))&sort(c7,d))&card(c7,int1))&etype(c7,int0))&fact(c7,real))&gener(c7,sp))&quant(c7,one))&refer(c7,det))&varia(c7,con))&sort(c12059,na))&card(c12059,int1))&etype(c12059,int0))&fact(c12059,real))&gener(c12059,sp))&quant(c12059,one))&refer(c12059,indet))&varia(c12059,varia_c))&sort(stadt__1_1,d))&sort(stadt__1_1,io))&card(stadt__1_1,int1))&etype(stadt__1_1,int0))&fact(stadt__1_1,real))&gener(stadt__1_1,ge))&quant(stadt__1_1,one))&refer(stadt__1_1,refer_c))&varia(stadt__1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(verona_0,fe))&sort(c12087,tq))&sort(zwei_1_1,nq))&sort(c12093,st))&fact(c12093,real))&gener(c12093,sp))&sort(c12602,co))&card(c12602,card_c))&etype(c12602,etype_c))&fact(c12602,real))&gener(c12602,sp))&quant(c12602,quant_c))&refer(c12602,det))&varia(c12602,con))&sort(c12583,d))&card(c12583,int1))&etype(c12583,int0))&fact(c12583,real))&gener(c12583,sp))&quant(c12583,one))&refer(c12583,det))&varia(c12583,con))&sort(c16,st))&fact(c16,real))&gener(c16,sp))&sort(equ_0,st))&fact(equ_0,real))&gener(equ_0,gener_c))&sort(c12608,l))&card(c12608,int1))&etype(c12608,int0))&fact(c12608,real))&gener(c12608,sp))&quant(c12608,one))&refer(c12608,det))&varia(c12608,con))&sort(kolosseum_1_1,d))&card(kolosseum_1_1,int1))&etype(kolosseum_1_1,int0))&fact(kolosseum_1_1,real))&gener(kolosseum_1_1,sp))&quant(kolosseum_1_1,one))&refer(kolosseum_1_1,det))&varia(kolosseum_1_1,con))&sort(c12594,d))&sort(c12594,io))&card(c12594,int1))&etype(c12594,int0))&fact(c12594,real))&gener(c12594,sp))&quant(c12594,one))&refer(c12594,det))&varia(c12594,con))&sort(c12595,na))&card(c12595,int1))&etype(c12595,int0))&fact(c12595,real))&gener(c12595,sp))&quant(c12595,one))&refer(c12595,indet))&varia(c12595,varia_c))&sort(rom_0,fe))&sort(c12605,tq))&sort(amphi_theater_1_1,o))&card(amphi_theater_1_1,int1))&etype(amphi_theater_1_1,int0))&fact(amphi_theater_1_1,real))&gener(amphi_theater_1_1,ge))&quant(amphi_theater_1_1,one))&refer(amphi_theater_1_1,refer_c))&varia(amphi_theater_1_1,varia_c))&sort(zweitgro__337_1_1,mq))&sort(c12606,o))&card(c12606,card_c))&etype(c12606,int1))&etype(c12606,int2))&fact(c12606,real))&gener(c12606,gener_c))&quant(c12606,quant_c))&refer(c12606,refer_c))&varia(c12606,varia_c))&sort(sub_0,st))&fact(sub_0,real))&gener(sub_0,gener_c))&sort(r__366misch__1_1,tq))&sort(arena__1_1,d))&card(arena__1_1,int1))&etype(arena__1_1,int0))&fact(arena__1_1,real))&gener(arena__1_1,ge))&quant(arena__1_1,one))&refer(arena__1_1,refer_c))&varia(arena__1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_004_insicht_3)).
fof(c_0_2, axiom, (![X1]:![X2]:(loc(X1,X2)=>?[X3]:((loc(X3,X2)&scar(X3,X1))&subs(X3,stehen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', loc__stehen_1_1_loc)).
fof(c_0_3, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:((((((((in(X3,X1)&attr(X1,X2))&loc(X5,X3))&scar(X5,X4))&sub(X2,name_1_1))&sub(X1,stadt__1_1))&sub(X4,kolosseum_1_1))&subs(X5,stehen_1_1))&val(X2,rom_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_4, negated_conjecture, (![X6]:![X7]:![X8]:![X9]:![X10]:((((((((~in(X8,X6)|~attr(X6,X7))|~loc(X10,X8))|~scar(X10,X9))|~sub(X7,name_1_1))|~sub(X6,stadt__1_1))|~sub(X9,kolosseum_1_1))|~subs(X10,stehen_1_1))|~val(X7,rom_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_5,negated_conjecture,(~val(X1,rom_0)|~subs(X2,stehen_1_1)|~sub(X3,kolosseum_1_1)|~sub(X4,stadt__1_1)|~sub(X1,name_1_1)|~scar(X2,X3)|~loc(X2,X5)|~attr(X4,X1)|~in(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_6,hypothesis,(val(c12595,rom_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_7,hypothesis,(sub(c12595,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(~in(X1,X2)|~attr(X2,c12595)|~scar(X3,X4)|~loc(X3,X1)|~subs(X3,stehen_1_1)|~sub(X2,stadt__1_1)|~sub(X4,kolosseum_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7])])).
cnf(c_0_9,hypothesis,(in(c12608,c12594)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_10,hypothesis,(attr(c12594,c12595)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_11,hypothesis,(sub(c12594,stadt__1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_12, plain, (![X4]:![X5]:(((loc(esk12_2(X4,X5),X5)|~loc(X4,X5))&(scar(esk12_2(X4,X5),X4)|~loc(X4,X5)))&(subs(esk12_2(X4,X5),stehen_1_1)|~loc(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_13,hypothesis,(~scar(X1,X2)|~loc(X1,c12608)|~subs(X1,stehen_1_1)|~sub(X2,kolosseum_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11])])).
cnf(c_0_14,plain,(scar(esk12_2(X1,X2),X1)|~loc(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15,plain,(subs(esk12_2(X1,X2),stehen_1_1)|~loc(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,hypothesis,(~loc(esk12_2(X1,X2),c12608)|~loc(X1,X2)|~sub(X1,kolosseum_1_1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])).
cnf(c_0_17,plain,(loc(esk12_2(X1,X2),X2)|~loc(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,hypothesis,(~loc(X1,c12608)|~sub(X1,kolosseum_1_1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_19,hypothesis,(loc(c12583,c12608)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_20,hypothesis,(sub(c12583,kolosseum_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_21,hypothesis,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])]), ['proof']).
# SZS output end CNFRefutation.
EOF
