% Problem    : CSR115+87 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n094.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:54:41 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_wp_508)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((attr(c109,c110)&sub(c109,firma_1_1))&sub(c110,name_1_1))&val(c110,porsche_0))&attr(c125,c126))&sub(c125,firma_1_1))&sub(c126,name_1_1))&val(c126,bmw_0))&tupl_p8(c222,c57,c58,c68,c73,c68,c109,c125))&pred(c57,zasada_1_1))&pred(c58,fabrikat_1_1))&attch(c63,c58))&sub(c63,firma_1_1))&attr(c68,c69))&sub(c68,stadt__1_1))&sub(c69,name_1_1))&val(c69,steyr_0))&attr(c73,c74))&sub(c73,mensch_1_1))&sub(c74,familiename_1_1))&val(c74,puch_0))&sort(c109,d))&sort(c109,io))&card(c109,int1))&etype(c109,int0))&fact(c109,real))&gener(c109,sp))&quant(c109,one))&refer(c109,det))&varia(c109,con))&sort(c110,na))&card(c110,int1))&etype(c110,int0))&fact(c110,real))&gener(c110,sp))&quant(c110,one))&refer(c110,indet))&varia(c110,varia_c))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(porsche_0,fe))&sort(c125,d))&sort(c125,io))&card(c125,int1))&etype(c125,int0))&fact(c125,real))&gener(c125,sp))&quant(c125,one))&refer(c125,det))&varia(c125,con))&sort(c126,na))&card(c126,int1))&etype(c126,int0))&fact(c126,real))&gener(c126,sp))&quant(c126,one))&refer(c126,indet))&varia(c126,varia_c))&sort(bmw_0,fe))&sort(c222,ent))&card(c222,card_c))&etype(c222,etype_c))&fact(c222,real))&gener(c222,gener_c))&quant(c222,quant_c))&refer(c222,refer_c))&varia(c222,varia_c))&sort(c57,o))&card(c57,cons(x_constant,cons(int1,nil))))&etype(c57,int1))&fact(c57,real))&gener(c57,gener_c))&quant(c57,mult))&refer(c57,indet))&varia(c57,varia_c))&sort(c58,o))&card(c58,cons(x_constant,cons(int1,nil))))&etype(c58,int1))&fact(c58,real))&gener(c58,sp))&quant(c58,mult))&refer(c58,indet))&varia(c58,varia_c))&sort(c68,d))&sort(c68,io))&card(c68,int1))&etype(c68,int0))&fact(c68,real))&gener(c68,sp))&quant(c68,one))&refer(c68,det))&varia(c68,con))&sort(c73,d))&card(c73,int1))&etype(c73,int0))&fact(c73,real))&gener(c73,sp))&quant(c73,one))&refer(c73,det))&varia(c73,con))&sort(zasada_1_1,o))&card(zasada_1_1,int1))&etype(zasada_1_1,int0))&fact(zasada_1_1,real))&gener(zasada_1_1,ge))&quant(zasada_1_1,one))&refer(zasada_1_1,refer_c))&varia(zasada_1_1,varia_c))&sort(fabrikat_1_1,o))&card(fabrikat_1_1,int1))&etype(fabrikat_1_1,int0))&fact(fabrikat_1_1,real))&gener(fabrikat_1_1,ge))&quant(fabrikat_1_1,one))&refer(fabrikat_1_1,refer_c))&varia(fabrikat_1_1,varia_c))&sort(c63,d))&sort(c63,io))&card(c63,int1))&etype(c63,int0))&fact(c63,real))&gener(c63,sp))&quant(c63,one))&refer(c63,det))&varia(c63,con))&sort(c69,na))&card(c69,int1))&etype(c69,int0))&fact(c69,real))&gener(c69,sp))&quant(c69,one))&refer(c69,indet))&varia(c69,varia_c))&sort(stadt__1_1,d))&sort(stadt__1_1,io))&card(stadt__1_1,int1))&etype(stadt__1_1,int0))&fact(stadt__1_1,real))&gener(stadt__1_1,ge))&quant(stadt__1_1,one))&refer(stadt__1_1,refer_c))&varia(stadt__1_1,varia_c))&sort(steyr_0,fe))&sort(c74,na))&card(c74,int1))&etype(c74,int0))&fact(c74,real))&gener(c74,sp))&quant(c74,one))&refer(c74,indet))&varia(c74,varia_c))&sort(mensch_1_1,d))&card(mensch_1_1,int1))&etype(mensch_1_1,int0))&fact(mensch_1_1,real))&gener(mensch_1_1,ge))&quant(mensch_1_1,one))&refer(mensch_1_1,refer_c))&varia(mensch_1_1,varia_c))&sort(familiename_1_1,na))&card(familiename_1_1,int1))&etype(familiename_1_1,int0))&fact(familiename_1_1,real))&gener(familiename_1_1,ge))&quant(familiename_1_1,one))&refer(familiename_1_1,refer_c))&varia(familiename_1_1,varia_c))&sort(puch_0,fe))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_wp_508)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_3, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_4, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, negated_conjecture, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((((~attr(X8,X9)|~attr(X11,X10))|~attr(X12,X13))|~obj(X14,X8))|~sub(X9,name_1_1))|~sub(X8,firma_1_1))|~sub(X10,name_1_1))|~val(X9,bmw_0))|~val(X10,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_6,negated_conjecture,(~val(X1,bmw_0)|~val(X2,bmw_0)|~sub(X1,name_1_1)|~sub(X3,firma_1_1)|~sub(X2,name_1_1)|~obj(X4,X3)|~attr(X5,X6)|~attr(X7,X1)|~attr(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,hypothesis,(val(c126,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(sub(c126,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(~val(X1,bmw_0)|~attr(X2,c126)|~attr(X3,X1)|~attr(X4,X5)|~obj(X6,X2)|~sub(X2,firma_1_1)|~sub(X1,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_10,hypothesis,(~attr(X1,c126)|~attr(X2,c126)|~attr(X3,X4)|~obj(X5,X1)|~sub(X1,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_8])])).
cnf(c_0_11,hypothesis,(attr(c125,c126)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_12,hypothesis,(sub(c125,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(~attr(X1,c126)|~attr(X2,X3)|~obj(X4,c125)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk35_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk35_3(X7,X8,X9),esk36_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk34_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk34_3(X7,X8,X9),esk35_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk34_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk36_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk35_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk34_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_15,hypothesis,(~attr(X1,X2)|~obj(X3,c125)), inference(spm,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_16,plain,(obj(esk34_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, (![X4]:![X5]:(((arg1(esk37_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk37_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk37_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_18,hypothesis,(~subr(X1,sub_0)|~attr(X2,X3)|~arg2(X1,X4)|~arg1(X1,c125)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(subr(esk37_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,hypothesis,(~attr(X1,X2)|~arg2(esk37_2(X3,X4),X5)|~arg1(esk37_2(X3,X4),c125)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(arg2(esk37_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,hypothesis,(~attr(X1,X2)|~arg1(esk37_2(X3,X4),c125)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,plain,(arg1(esk37_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,hypothesis,(~attr(X1,X2)|~sub(c125,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,hypothesis,(attr(c109,c110)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_26,hypothesis,(~attr(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_12])).
cnf(c_0_27,hypothesis,($false), inference(sr,[status(thm)],[c_0_25, c_0_26]), ['proof']).
# SZS output end CNFRefutation.
EOF
