% Problem    : CSR115+5 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n141.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:10:41 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:((((((attr(X3,X2)&attr(X5,X6))&obj(X4,X1))&prop(X1,britisch__1_1))&sub(X1,firma_1_1))&sub(X2,name_1_1))&val(X2,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_news_1091_a19984)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((assoc(autokonzern_1_1,auto__1_1)&sub(autokonzern_1_1,firmengruppe_1_1))&attr(c18,c19))&sub(c18,stadt__1_1))&sub(c19,name_1_1))&val(c19,london_0))&sub(c21,feb_1_1))&tupl(c34,c18,c21))&attr(c55667,c55668))&prop(c55667,bundesdeutsch_1_1))&sub(c55667,autokonzern_1_1))&sub(c55668,name_1_1))&val(c55668,bmw_0))&prop(c55687,pass__351_1_1))&sub(c55687,woche_1_1))&subs(c55693,ankauf__1_1))&attch(c55986,c55693))&attr(c55986,c55987))&prop(c55986,britisch__1_1))&sub(c55986,firma_1_1))&sub(c55987,name_1_1))&val(c55987,rover_0))&subs(c55995,interesse_1_1))&subs(c55999,ankauf__1_1))&attch(c56008,c55999))&prop(c56008,britisch__1_1))&sub(c56008,luxusmarke_1_1))&attr(c56019,c56008))&attr(c56019,c56020))&sub(c56019,mensch_1_1))&sub(c56020,familiename_1_1))&val(c56020,roll_0))&attr(c56024,c56025))&sub(c56024,mensch_1_1))&sub(c56025,eigenname_1_1))&val(c56025,royce_0))&tupl_p8(c56046,c55667,c55669,c55687,c55693,c55995,c55999,c56024))&assoc(luxusmarke_1_1,luxus__1_1))&sub(luxusmarke_1_1,marke_1_1))&sort(autokonzern_1_1,d))&sort(autokonzern_1_1,io))&card(autokonzern_1_1,int1))&etype(autokonzern_1_1,int0))&fact(autokonzern_1_1,real))&gener(autokonzern_1_1,ge))&quant(autokonzern_1_1,one))&refer(autokonzern_1_1,refer_c))&varia(autokonzern_1_1,varia_c))&sort(auto__1_1,d))&card(auto__1_1,int1))&etype(auto__1_1,int0))&fact(auto__1_1,real))&gener(auto__1_1,ge))&quant(auto__1_1,one))&refer(auto__1_1,refer_c))&varia(auto__1_1,varia_c))&sort(firmengruppe_1_1,d))&sort(firmengruppe_1_1,io))&card(firmengruppe_1_1,int1))&etype(firmengruppe_1_1,int0))&fact(firmengruppe_1_1,real))&gener(firmengruppe_1_1,ge))&quant(firmengruppe_1_1,one))&refer(firmengruppe_1_1,refer_c))&varia(firmengruppe_1_1,varia_c))&sort(c18,d))&sort(c18,io))&card(c18,int1))&etype(c18,int0))&fact(c18,real))&gener(c18,sp))&quant(c18,one))&refer(c18,det))&varia(c18,con))&sort(c19,na))&card(c19,int1))&etype(c19,int0))&fact(c19,real))&gener(c19,sp))&quant(c19,one))&refer(c19,indet))&varia(c19,varia_c))&sort(stadt__1_1,d))&sort(stadt__1_1,io))&card(stadt__1_1,int1))&etype(stadt__1_1,int0))&fact(stadt__1_1,real))&gener(stadt__1_1,ge))&quant(stadt__1_1,one))&refer(stadt__1_1,refer_c))&varia(stadt__1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(london_0,fe))&sort(c21,o))&card(c21,int1))&etype(c21,int0))&fact(c21,real))&gener(c21,gener_c))&quant(c21,one))&refer(c21,refer_c))&varia(c21,varia_c))&sort(feb_1_1,o))&card(feb_1_1,int1))&etype(feb_1_1,int0))&fact(feb_1_1,real))&gener(feb_1_1,ge))&quant(feb_1_1,one))&refer(feb_1_1,refer_c))&varia(feb_1_1,varia_c))&sort(c34,ent))&card(c34,card_c))&etype(c34,etype_c))&fact(c34,real))&gener(c34,gener_c))&quant(c34,quant_c))&refer(c34,refer_c))&varia(c34,varia_c))&sort(c55667,d))&sort(c55667,io))&card(c55667,int1))&etype(c55667,int0))&fact(c55667,real))&gener(c55667,sp))&quant(c55667,one))&refer(c55667,det))&varia(c55667,con))&sort(c55668,na))&card(c55668,int1))&etype(c55668,int0))&fact(c55668,real))&gener(c55668,sp))&quant(c55668,one))&refer(c55668,indet))&varia(c55668,varia_c))&sort(bundesdeutsch_1_1,tq))&sort(bmw_0,fe))&sort(c55687,me))&sort(c55687,oa))&sort(c55687,ta))&card(c55687,card_c))&etype(c55687,etype_c))&fact(c55687,real))&gener(c55687,sp))&quant(c55687,quant_c))&refer(c55687,det))&varia(c55687,con))&sort(pass__351_1_1,tq))&sort(woche_1_1,me))&sort(woche_1_1,oa))&sort(woche_1_1,ta))&card(woche_1_1,card_c))&etype(woche_1_1,etype_c))&fact(woche_1_1,real))&gener(woche_1_1,ge))&quant(woche_1_1,quant_c))&refer(woche_1_1,refer_c))&varia(woche_1_1,varia_c))&sort(c55693,ad))&card(c55693,int1))&etype(c55693,int0))&fact(c55693,real))&gener(c55693,sp))&quant(c55693,one))&refer(c55693,det))&varia(c55693,con))&sort(ankauf__1_1,ad))&card(ankauf__1_1,int1))&etype(ankauf__1_1,int0))&fact(ankauf__1_1,real))&gener(ankauf__1_1,ge))&quant(ankauf__1_1,one))&refer(ankauf__1_1,refer_c))&varia(ankauf__1_1,varia_c))&sort(c55986,d))&sort(c55986,io))&card(c55986,int1))&etype(c55986,int0))&fact(c55986,real))&gener(c55986,sp))&quant(c55986,one))&refer(c55986,det))&varia(c55986,con))&sort(c55987,na))&card(c55987,int1))&etype(c55987,int0))&fact(c55987,real))&gener(c55987,sp))&quant(c55987,one))&refer(c55987,indet))&varia(c55987,varia_c))&sort(britisch__1_1,nq))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(rover_0,fe))&sort(c55995,as))&card(c55995,int1))&etype(c55995,int0))&fact(c55995,real))&gener(c55995,gener_c))&quant(c55995,one))&refer(c55995,refer_c))&varia(c55995,varia_c))&sort(interesse_1_1,as))&card(interesse_1_1,int1))&etype(interesse_1_1,int0))&fact(interesse_1_1,real))&gener(interesse_1_1,ge))&quant(interesse_1_1,one))&refer(interesse_1_1,refer_c))&varia(interesse_1_1,varia_c))&sort(c55999,ad))&card(c55999,int1))&etype(c55999,int0))&fact(c55999,real))&gener(c55999,sp))&quant(c55999,one))&refer(c55999,det))&varia(c55999,con))&sort(c56008,io))&sort(c56008,oa))&card(c56008,int1))&etype(c56008,int0))&fact(c56008,real))&gener(c56008,sp))&quant(c56008,one))&refer(c56008,det))&varia(c56008,con))&sort(luxusmarke_1_1,io))&sort(luxusmarke_1_1,oa))&card(luxusmarke_1_1,int1))&etype(luxusmarke_1_1,int0))&fact(luxusmarke_1_1,real))&gener(luxusmarke_1_1,ge))&quant(luxusmarke_1_1,one))&refer(luxusmarke_1_1,refer_c))&varia(luxusmarke_1_1,varia_c))&sort(c56019,d))&card(c56019,int1))&etype(c56019,int0))&fact(c56019,real))&gener(c56019,sp))&quant(c56019,one))&refer(c56019,det))&varia(c56019,con))&sort(c56020,na))&card(c56020,int1))&etype(c56020,int0))&fact(c56020,real))&gener(c56020,sp))&quant(c56020,one))&refer(c56020,indet))&varia(c56020,varia_c))&sort(mensch_1_1,d))&card(mensch_1_1,int1))&etype(mensch_1_1,int0))&fact(mensch_1_1,real))&gener(mensch_1_1,ge))&quant(mensch_1_1,one))&refer(mensch_1_1,refer_c))&varia(mensch_1_1,varia_c))&sort(familiename_1_1,na))&card(familiename_1_1,int1))&etype(familiename_1_1,int0))&fact(familiename_1_1,real))&gener(familiename_1_1,ge))&quant(familiename_1_1,one))&refer(familiename_1_1,refer_c))&varia(familiename_1_1,varia_c))&sort(roll_0,fe))&sort(c56024,d))&card(c56024,int1))&etype(c56024,int0))&fact(c56024,real))&gener(c56024,sp))&quant(c56024,one))&refer(c56024,det))&varia(c56024,con))&sort(c56025,na))&card(c56025,int1))&etype(c56025,int0))&fact(c56025,real))&gener(c56025,sp))&quant(c56025,one))&refer(c56025,indet))&varia(c56025,varia_c))&sort(eigenname_1_1,na))&card(eigenname_1_1,int1))&etype(eigenname_1_1,int0))&fact(eigenname_1_1,real))&gener(eigenname_1_1,ge))&quant(eigenname_1_1,one))&refer(eigenname_1_1,refer_c))&varia(eigenname_1_1,varia_c))&sort(royce_0,fe))&sort(c56046,ent))&card(c56046,card_c))&etype(c56046,etype_c))&fact(c56046,real))&gener(c56046,gener_c))&quant(c56046,quant_c))&refer(c56046,refer_c))&varia(c56046,varia_c))&sort(c55669,o))&card(c55669,int1))&etype(c55669,int0))&fact(c55669,real))&gener(c55669,sp))&quant(c55669,one))&refer(c55669,det))&varia(c55669,varia_c))&sort(luxus__1_1,io))&card(luxus__1_1,int1))&etype(luxus__1_1,int0))&fact(luxus__1_1,real))&gener(luxus__1_1,ge))&quant(luxus__1_1,one))&refer(luxus__1_1,refer_c))&varia(luxus__1_1,varia_c))&sort(marke_1_1,io))&sort(marke_1_1,oa))&card(marke_1_1,int1))&etype(marke_1_1,int0))&fact(marke_1_1,real))&gener(marke_1_1,ge))&quant(marke_1_1,one))&refer(marke_1_1,refer_c))&varia(marke_1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_news_1091_a19984)).
fof(c_0_2, axiom, (![X1]:![X2]:((has_fact_leq(X2,real)&loc(X2,X1))=>?[X3]:((loc(X3,X1)&obj(X3,X2))&subs(X3,geben_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', loc__geben_1_1_loc)).
fof(c_0_3, axiom, (![X1]:![X2]:(fact(X1,X2)=>has_fact_leq(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', has_fact_eq)).
fof(c_0_4, axiom, (![X1]:![X2]:![X3]:((prop(X1,X2)&state_adjective_state_binding(X2,X3))=>?[X4]:?[X5]:?[X6]:(((((in(X6,X4)&attr(X4,X5))&loc(X1,X6))&sub(X4,land_1_1))&sub(X5,name_1_1))&val(X5,X3)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', state_adjective__in_state)).
fof(c_0_5, axiom, (state_adjective_state_binding(britisch__1_1,grossbritannien_0)), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', fact_8846)).
fof(c_0_6, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:((((((attr(X3,X2)&attr(X5,X6))&obj(X4,X1))&prop(X1,britisch__1_1))&sub(X1,firma_1_1))&sub(X2,name_1_1))&val(X2,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_7, negated_conjecture, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((((((~attr(X9,X8)|~attr(X10,X11))|~obj(X12,X7))|~prop(X7,britisch__1_1))|~sub(X7,firma_1_1))|~sub(X8,name_1_1))|~val(X8,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_8,negated_conjecture,(~val(X1,bmw_0)|~sub(X1,name_1_1)|~sub(X2,firma_1_1)|~prop(X2,britisch__1_1)|~obj(X3,X2)|~attr(X4,X5)|~attr(X6,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_9,hypothesis,(val(c55668,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_10,hypothesis,(sub(c55668,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_11,hypothesis,(~prop(X1,britisch__1_1)|~attr(X2,c55668)|~attr(X3,X4)|~obj(X5,X1)|~sub(X1,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
cnf(c_0_12,hypothesis,(prop(c55986,britisch__1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(sub(c55986,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_14,hypothesis,(~attr(X1,c55668)|~attr(X2,X3)|~obj(X4,c55986)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_15,hypothesis,(attr(c55667,c55668)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_16, plain, (![X4]:![X5]:(((loc(esk12_2(X4,X5),X4)|(~has_fact_leq(X5,real)|~loc(X5,X4)))&(obj(esk12_2(X4,X5),X5)|(~has_fact_leq(X5,real)|~loc(X5,X4))))&(subs(esk12_2(X4,X5),geben_1_1)|(~has_fact_leq(X5,real)|~loc(X5,X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_17,hypothesis,(~attr(X1,X2)|~obj(X3,c55986)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_18,plain,(obj(esk12_2(X2,X1),X1)|~loc(X1,X2)|~has_fact_leq(X1,real)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, plain, (![X3]:![X4]:(~fact(X3,X4)|has_fact_leq(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
cnf(c_0_20,hypothesis,(~attr(X1,X2)|~loc(c55986,X3)|~has_fact_leq(c55986,real)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21,plain,(has_fact_leq(X1,X2)|~fact(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_22,hypothesis,(fact(c55986,real)), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_23, plain, (![X7]:![X8]:![X9]:((((((in(esk9_3(X7,X8,X9),esk7_3(X7,X8,X9))|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9)))&(attr(esk7_3(X7,X8,X9),esk8_3(X7,X8,X9))|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(loc(X7,esk9_3(X7,X8,X9))|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(sub(esk7_3(X7,X8,X9),land_1_1)|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(sub(esk8_3(X7,X8,X9),name_1_1)|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))&(val(esk8_3(X7,X8,X9),X9)|(~prop(X7,X8)|~state_adjective_state_binding(X8,X9))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
cnf(c_0_24,hypothesis,(~attr(X1,X2)|~loc(c55986,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
cnf(c_0_25,plain,(loc(X3,esk9_3(X3,X1,X2))|~state_adjective_state_binding(X1,X2)|~prop(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_26,hypothesis,(~state_adjective_state_binding(X1,X2)|~prop(c55986,X1)|~attr(X3,X4)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_27,plain,(state_adjective_state_binding(britisch__1_1,grossbritannien_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_28,hypothesis,(attr(c18,c19)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_29,hypothesis,(~attr(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_12])])).
cnf(c_0_30,hypothesis,($false), inference(sr,[status(thm)],[c_0_28, c_0_29]), ['proof']).
# SZS output end CNFRefutation.
EOF
