% Problem    : SWB012+3 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n146.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:07:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:((((((iext(uri_rdf_first,X4,X5)&iext(uri_rdf_rest,X4,X6))&iext(uri_rdf_first,X6,X7))&iext(uri_rdf_rest,X6,X8))&iext(uri_rdf_first,X8,X9))&iext(uri_rdf_rest,X8,uri_rdf_nil))=>(iext(uri_owl_intersectionOf,X1,X4)<=>((((ic(X1)&ic(X5))&ic(X7))&ic(X9))&![X3]:(icext(X1,X3)<=>((icext(X5,X3)&icext(X7,X3))&icext(X9,X3))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB002+0.ax', owl_bool_intersectionof_class_003)).
fof(c_0_1, axiom, (?[X21]:?[X22]:?[X23]:?[X24]:((((((((((((iext(uri_rdf_type,uri_foaf_Person,uri_owl_Class)&iext(uri_owl_intersectionOf,uri_ex_PersonAttribute,X21))&iext(uri_rdf_first,X21,uri_owl_DatatypeProperty))&iext(uri_rdf_rest,X21,X22))&iext(uri_rdf_first,X22,uri_owl_FunctionalProperty))&iext(uri_rdf_rest,X22,X23))&iext(uri_rdf_first,X23,X24))&iext(uri_rdf_rest,X23,uri_rdf_nil))&iext(uri_rdf_type,X24,uri_owl_Restriction))&iext(uri_owl_onProperty,X24,uri_rdfs_domain))&iext(uri_owl_hasValue,X24,uri_foaf_Person))&iext(uri_rdf_type,uri_ex_name,uri_ex_PersonAttribute))&iext(uri_ex_name,uri_ex_alice,literal_plain(dat_str_alice)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_premise_fullish_012_Template_Class)).
fof(c_0_2, axiom, (![X3]:![X2]:(iext(uri_rdf_type,X3,X2)<=>icext(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWB002+0.ax', rdfs_cext_def)).
fof(c_0_3, axiom, (![X1]:![X11]:![X14]:((iext(uri_owl_hasValue,X1,X14)&iext(uri_owl_onProperty,X1,X11))=>![X3]:(icext(X1,X3)<=>iext(X11,X3,X14)))), file('/export/starexec/sandbox/benchmark/Axioms/SWB002+0.ax', owl_restrict_hasvalue)).
fof(c_0_4, axiom, (![X11]:![X2]:(iext(uri_rdfs_domain,X11,X2)<=>((ip(X11)&ic(X2))&![X3]:![X10]:(iext(X11,X3,X10)=>icext(X2,X3))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB002+0.ax', owl_rdfsext_domain)).
fof(c_0_5, conjecture, ((iext(uri_rdf_type,uri_ex_name,uri_owl_FunctionalProperty)&iext(uri_rdf_type,uri_ex_alice,uri_foaf_Person))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_conclusion_fullish_012_Template_Class)).
fof(c_0_6, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:(((((((ic(X10)|~iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil)))&((ic(X12)|~iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil))))&((ic(X14)|~iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil))))&((ic(X16)|~iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil))))&((((((icext(X12,X17)|~icext(X10,X17))|~iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil)))&(((icext(X14,X17)|~icext(X10,X17))|~iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil))))&(((icext(X16,X17)|~icext(X10,X17))|~iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil))))&(((((~icext(X12,X18)|~icext(X14,X18))|~icext(X16,X18))|icext(X10,X18))|~iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil)))))&(((((~icext(X10,esk4_7(X10,X11,X12,X13,X14,X15,X16))|((~icext(X12,esk4_7(X10,X11,X12,X13,X14,X15,X16))|~icext(X14,esk4_7(X10,X11,X12,X13,X14,X15,X16)))|~icext(X16,esk4_7(X10,X11,X12,X13,X14,X15,X16))))|(((~ic(X10)|~ic(X12))|~ic(X14))|~ic(X16)))|iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil)))&((((((icext(X12,esk4_7(X10,X11,X12,X13,X14,X15,X16))|icext(X10,esk4_7(X10,X11,X12,X13,X14,X15,X16)))|(((~ic(X10)|~ic(X12))|~ic(X14))|~ic(X16)))|iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil)))&((((icext(X14,esk4_7(X10,X11,X12,X13,X14,X15,X16))|icext(X10,esk4_7(X10,X11,X12,X13,X14,X15,X16)))|(((~ic(X10)|~ic(X12))|~ic(X14))|~ic(X16)))|iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil))))&((((icext(X16,esk4_7(X10,X11,X12,X13,X14,X15,X16))|icext(X10,esk4_7(X10,X11,X12,X13,X14,X15,X16)))|(((~ic(X10)|~ic(X12))|~ic(X14))|~ic(X16)))|iext(uri_owl_intersectionOf,X10,X11))|(((((~iext(uri_rdf_first,X11,X12)|~iext(uri_rdf_rest,X11,X13))|~iext(uri_rdf_first,X13,X14))|~iext(uri_rdf_rest,X13,X15))|~iext(uri_rdf_first,X15,X16))|~iext(uri_rdf_rest,X15,uri_rdf_nil))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])])).
fof(c_0_7, plain, (((((((((((((iext(uri_rdf_type,uri_foaf_Person,uri_owl_Class)&iext(uri_owl_intersectionOf,uri_ex_PersonAttribute,esk19_0))&iext(uri_rdf_first,esk19_0,uri_owl_DatatypeProperty))&iext(uri_rdf_rest,esk19_0,esk20_0))&iext(uri_rdf_first,esk20_0,uri_owl_FunctionalProperty))&iext(uri_rdf_rest,esk20_0,esk21_0))&iext(uri_rdf_first,esk21_0,esk22_0))&iext(uri_rdf_rest,esk21_0,uri_rdf_nil))&iext(uri_rdf_type,esk22_0,uri_owl_Restriction))&iext(uri_owl_onProperty,esk22_0,uri_rdfs_domain))&iext(uri_owl_hasValue,esk22_0,uri_foaf_Person))&iext(uri_rdf_type,uri_ex_name,uri_ex_PersonAttribute))&iext(uri_ex_name,uri_ex_alice,literal_plain(dat_str_alice)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[c_0_1])])])).
cnf(c_0_8,plain,(icext(X2,X8)|~iext(uri_rdf_rest,X1,uri_rdf_nil)|~iext(uri_rdf_first,X1,X2)|~iext(uri_rdf_rest,X3,X1)|~iext(uri_rdf_first,X3,X4)|~iext(uri_rdf_rest,X5,X3)|~iext(uri_rdf_first,X5,X6)|~iext(uri_owl_intersectionOf,X7,X5)|~icext(X7,X8)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,plain,(iext(uri_rdf_rest,esk21_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,plain,(icext(X1,X2)|~icext(X3,X2)|~iext(uri_rdf_first,esk21_0,X1)|~iext(uri_rdf_rest,X4,esk21_0)|~iext(uri_owl_intersectionOf,X3,X5)|~iext(uri_rdf_first,X5,X6)|~iext(uri_rdf_first,X4,X7)|~iext(uri_rdf_rest,X5,X4)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_11,plain,(iext(uri_rdf_first,esk21_0,esk22_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(icext(esk22_0,X1)|~icext(X2,X1)|~iext(uri_rdf_rest,X3,esk21_0)|~iext(uri_owl_intersectionOf,X2,X4)|~iext(uri_rdf_first,X4,X5)|~iext(uri_rdf_first,X3,X6)|~iext(uri_rdf_rest,X4,X3)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_13,plain,(iext(uri_rdf_rest,esk20_0,esk21_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14,plain,(icext(X4,X8)|~iext(uri_rdf_rest,X1,uri_rdf_nil)|~iext(uri_rdf_first,X1,X2)|~iext(uri_rdf_rest,X3,X1)|~iext(uri_rdf_first,X3,X4)|~iext(uri_rdf_rest,X5,X3)|~iext(uri_rdf_first,X5,X6)|~iext(uri_owl_intersectionOf,X7,X5)|~icext(X7,X8)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,plain,(icext(esk22_0,X1)|~icext(X2,X1)|~iext(uri_rdf_first,esk20_0,X3)|~iext(uri_rdf_rest,X4,esk20_0)|~iext(uri_owl_intersectionOf,X2,X4)|~iext(uri_rdf_first,X4,X5)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16,plain,(iext(uri_rdf_first,esk20_0,uri_owl_FunctionalProperty)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,plain,(icext(X1,X2)|~icext(X3,X2)|~iext(uri_rdf_first,esk21_0,X4)|~iext(uri_rdf_rest,X5,esk21_0)|~iext(uri_owl_intersectionOf,X3,X6)|~iext(uri_rdf_first,X6,X7)|~iext(uri_rdf_first,X5,X1)|~iext(uri_rdf_rest,X6,X5)), inference(spm,[status(thm)],[c_0_14, c_0_9])).
cnf(c_0_18,plain,(icext(esk22_0,X1)|~icext(X2,X1)|~iext(uri_rdf_rest,X3,esk20_0)|~iext(uri_owl_intersectionOf,X2,X3)|~iext(uri_rdf_first,X3,X4)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(iext(uri_rdf_rest,esk19_0,esk20_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_20, plain, (![X4]:![X5]:![X6]:![X7]:((~iext(uri_rdf_type,X4,X5)|icext(X5,X4))&(~icext(X7,X6)|iext(uri_rdf_type,X6,X7)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_21,plain,(icext(X1,X2)|~icext(X3,X2)|~iext(uri_rdf_rest,X4,esk21_0)|~iext(uri_owl_intersectionOf,X3,X5)|~iext(uri_rdf_first,X5,X6)|~iext(uri_rdf_first,X4,X1)|~iext(uri_rdf_rest,X5,X4)), inference(spm,[status(thm)],[c_0_17, c_0_11])).
cnf(c_0_22,plain,(icext(esk22_0,X1)|~icext(X2,X1)|~iext(uri_owl_intersectionOf,X2,esk19_0)|~iext(uri_rdf_first,esk19_0,X3)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23,plain,(iext(uri_owl_intersectionOf,uri_ex_PersonAttribute,esk19_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24,plain,(icext(X1,X2)|~iext(uri_rdf_type,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25,plain,(iext(uri_rdf_type,uri_ex_name,uri_ex_PersonAttribute)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_26,plain,(icext(X1,X2)|~icext(X3,X2)|~iext(uri_rdf_first,esk20_0,X1)|~iext(uri_rdf_rest,X4,esk20_0)|~iext(uri_owl_intersectionOf,X3,X4)|~iext(uri_rdf_first,X4,X5)), inference(spm,[status(thm)],[c_0_21, c_0_13])).
fof(c_0_27, plain, (![X15]:![X16]:![X17]:![X18]:![X19]:(((~icext(X15,X18)|iext(X16,X18,X17))|(~iext(uri_owl_hasValue,X15,X17)|~iext(uri_owl_onProperty,X15,X16)))&((~iext(X16,X19,X17)|icext(X15,X19))|(~iext(uri_owl_hasValue,X15,X17)|~iext(uri_owl_onProperty,X15,X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_28,plain,(icext(esk22_0,X1)|~icext(uri_ex_PersonAttribute,X1)|~iext(uri_rdf_first,esk19_0,X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29,plain,(icext(uri_ex_PersonAttribute,uri_ex_name)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30,plain,(icext(uri_owl_FunctionalProperty,X1)|~icext(X2,X1)|~iext(uri_rdf_rest,X3,esk20_0)|~iext(uri_owl_intersectionOf,X2,X3)|~iext(uri_rdf_first,X3,X4)), inference(spm,[status(thm)],[c_0_26, c_0_16])).
cnf(c_0_31,plain,(iext(X2,X4,X3)|~iext(uri_owl_onProperty,X1,X2)|~iext(uri_owl_hasValue,X1,X3)|~icext(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32,plain,(iext(uri_owl_hasValue,esk22_0,uri_foaf_Person)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_33,plain,(icext(esk22_0,uri_ex_name)|~iext(uri_rdf_first,esk19_0,X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34,plain,(iext(uri_rdf_first,esk19_0,uri_owl_DatatypeProperty)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_35,plain,(icext(uri_owl_FunctionalProperty,X1)|~icext(X2,X1)|~iext(uri_owl_intersectionOf,X2,esk19_0)|~iext(uri_rdf_first,esk19_0,X3)), inference(spm,[status(thm)],[c_0_30, c_0_19])).
cnf(c_0_36,plain,(iext(X1,X2,uri_foaf_Person)|~icext(esk22_0,X2)|~iext(uri_owl_onProperty,esk22_0,X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37,plain,(icext(esk22_0,uri_ex_name)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38,plain,(icext(uri_owl_FunctionalProperty,X1)|~icext(uri_ex_PersonAttribute,X1)|~iext(uri_rdf_first,esk19_0,X2)), inference(spm,[status(thm)],[c_0_35, c_0_23])).
fof(c_0_39, plain, (![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:((((ip(X12)|~iext(uri_rdfs_domain,X12,X13))&(ic(X13)|~iext(uri_rdfs_domain,X12,X13)))&((~iext(X12,X14,X15)|icext(X13,X14))|~iext(uri_rdfs_domain,X12,X13)))&(((iext(X16,esk10_2(X16,X17),esk11_2(X16,X17))|(~ip(X16)|~ic(X17)))|iext(uri_rdfs_domain,X16,X17))&((~icext(X17,esk10_2(X16,X17))|(~ip(X16)|~ic(X17)))|iext(uri_rdfs_domain,X16,X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
cnf(c_0_40,plain,(iext(X1,uri_ex_name,uri_foaf_Person)|~iext(uri_owl_onProperty,esk22_0,X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41,plain,(iext(uri_owl_onProperty,esk22_0,uri_rdfs_domain)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_42, negated_conjecture, (~((iext(uri_rdf_type,uri_ex_name,uri_owl_FunctionalProperty)&iext(uri_rdf_type,uri_ex_alice,uri_foaf_Person)))), inference(assume_negation,[status(cth)],[c_0_5])).
cnf(c_0_43,plain,(icext(uri_owl_FunctionalProperty,uri_ex_name)|~iext(uri_rdf_first,esk19_0,X1)), inference(spm,[status(thm)],[c_0_38, c_0_29])).
cnf(c_0_44,plain,(icext(X2,X3)|~iext(uri_rdfs_domain,X1,X2)|~iext(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45,plain,(iext(uri_rdfs_domain,uri_ex_name,uri_foaf_Person)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_46, negated_conjecture, ((~iext(uri_rdf_type,uri_ex_name,uri_owl_FunctionalProperty)|~iext(uri_rdf_type,uri_ex_alice,uri_foaf_Person))), inference(fof_nnf,[status(thm)],[c_0_42])).
cnf(c_0_47,plain,(iext(uri_rdf_type,X1,X2)|~icext(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_48,plain,(icext(uri_owl_FunctionalProperty,uri_ex_name)), inference(spm,[status(thm)],[c_0_43, c_0_34])).
cnf(c_0_49,plain,(icext(uri_foaf_Person,X1)|~iext(uri_ex_name,X1,X2)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50,plain,(iext(uri_ex_name,uri_ex_alice,literal_plain(dat_str_alice))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_51,negated_conjecture,(~iext(uri_rdf_type,uri_ex_alice,uri_foaf_Person)|~iext(uri_rdf_type,uri_ex_name,uri_owl_FunctionalProperty)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52,plain,(iext(uri_rdf_type,uri_ex_name,uri_owl_FunctionalProperty)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53,plain,(icext(uri_foaf_Person,uri_ex_alice)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54,negated_conjecture,(~iext(uri_rdf_type,uri_ex_alice,uri_foaf_Person)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52])])).
cnf(c_0_55,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_53]), c_0_54]), ['proof']).
# SZS output end CNFRefutation.
EOF
