% Problem    : CSR115+3 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n140.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:40:01 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', synth_qa07_007_mira_news_1086)).
fof(c_0_1, hypothesis, (((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((attr(c3505,c3506)&sub(c3505,mensch_1_1))&sub(c3506,familiename_1_1))&val(c3506,roll_0))&pred(c3508,techniker__1_1))&attr(c3510,c3511))&sub(c3510,mensch_1_1))&sub(c3511,familiename_1_1))&val(c3511,royce_0))&attr(c3590,c3591))&sub(c3590,firma_1_1))&sub(c3591,name_1_1))&val(c3591,bmw_0))&pred(c3592,experte_1_1))&sub(c3595,fahrzeugkonzept_1_1))&prop(c3609,britisch__1_1))&sub(c3609,firma_1_1))&subs(c3613,schub__1_1))&tupl_p10(c3898,c3505,c3510,c3508,c3590,c3592,c3595,c3598,c3609,c3613))&assoc(fahrzeugkonzept_1_1,fahrzeug__1_1))&sub(fahrzeugkonzept_1_1,konzept__1_1))&sort(c3505,d))&card(c3505,int1))&etype(c3505,int0))&fact(c3505,real))&gener(c3505,sp))&quant(c3505,one))&refer(c3505,det))&varia(c3505,con))&sort(c3506,na))&card(c3506,int1))&etype(c3506,int0))&fact(c3506,real))&gener(c3506,sp))&quant(c3506,one))&refer(c3506,indet))&varia(c3506,varia_c))&sort(mensch_1_1,d))&card(mensch_1_1,int1))&etype(mensch_1_1,int0))&fact(mensch_1_1,real))&gener(mensch_1_1,ge))&quant(mensch_1_1,one))&refer(mensch_1_1,refer_c))&varia(mensch_1_1,varia_c))&sort(familiename_1_1,na))&card(familiename_1_1,int1))&etype(familiename_1_1,int0))&fact(familiename_1_1,real))&gener(familiename_1_1,ge))&quant(familiename_1_1,one))&refer(familiename_1_1,refer_c))&varia(familiename_1_1,varia_c))&sort(roll_0,fe))&sort(c3508,d))&card(c3508,cons(x_constant,cons(int1,nil))))&etype(c3508,int1))&fact(c3508,real))&gener(c3508,gener_c))&quant(c3508,mult))&refer(c3508,indet))&varia(c3508,varia_c))&sort(techniker__1_1,d))&card(techniker__1_1,int1))&etype(techniker__1_1,int0))&fact(techniker__1_1,real))&gener(techniker__1_1,ge))&quant(techniker__1_1,one))&refer(techniker__1_1,refer_c))&varia(techniker__1_1,varia_c))&sort(c3510,d))&card(c3510,int1))&etype(c3510,int0))&fact(c3510,real))&gener(c3510,sp))&quant(c3510,one))&refer(c3510,det))&varia(c3510,con))&sort(c3511,na))&card(c3511,int1))&etype(c3511,int0))&fact(c3511,real))&gener(c3511,sp))&quant(c3511,one))&refer(c3511,indet))&varia(c3511,varia_c))&sort(royce_0,fe))&sort(c3590,d))&sort(c3590,io))&card(c3590,int1))&etype(c3590,int0))&fact(c3590,real))&gener(c3590,sp))&quant(c3590,one))&refer(c3590,det))&varia(c3590,con))&sort(c3591,na))&card(c3591,int1))&etype(c3591,int0))&fact(c3591,real))&gener(c3591,sp))&quant(c3591,one))&refer(c3591,indet))&varia(c3591,varia_c))&sort(firma_1_1,d))&sort(firma_1_1,io))&card(firma_1_1,int1))&etype(firma_1_1,int0))&fact(firma_1_1,real))&gener(firma_1_1,ge))&quant(firma_1_1,one))&refer(firma_1_1,refer_c))&varia(firma_1_1,varia_c))&sort(name_1_1,na))&card(name_1_1,int1))&etype(name_1_1,int0))&fact(name_1_1,real))&gener(name_1_1,ge))&quant(name_1_1,one))&refer(name_1_1,refer_c))&varia(name_1_1,varia_c))&sort(bmw_0,fe))&sort(c3592,d))&card(c3592,cons(x_constant,cons(int1,nil))))&etype(c3592,int1))&fact(c3592,real))&gener(c3592,gener_c))&quant(c3592,mult))&refer(c3592,indet))&varia(c3592,varia_c))&sort(experte_1_1,d))&card(experte_1_1,int1))&etype(experte_1_1,int0))&fact(experte_1_1,real))&gener(experte_1_1,ge))&quant(experte_1_1,one))&refer(experte_1_1,refer_c))&varia(experte_1_1,varia_c))&sort(c3595,d))&sort(c3595,io))&card(c3595,int1))&etype(c3595,int0))&fact(c3595,real))&gener(c3595,sp))&quant(c3595,one))&refer(c3595,indet))&varia(c3595,varia_c))&sort(fahrzeugkonzept_1_1,d))&sort(fahrzeugkonzept_1_1,io))&card(fahrzeugkonzept_1_1,int1))&etype(fahrzeugkonzept_1_1,int0))&fact(fahrzeugkonzept_1_1,real))&gener(fahrzeugkonzept_1_1,ge))&quant(fahrzeugkonzept_1_1,one))&refer(fahrzeugkonzept_1_1,refer_c))&varia(fahrzeugkonzept_1_1,varia_c))&sort(c3609,d))&sort(c3609,io))&card(c3609,int1))&etype(c3609,int0))&fact(c3609,real))&gener(c3609,sp))&quant(c3609,one))&refer(c3609,det))&varia(c3609,con))&sort(britisch__1_1,nq))&sort(c3613,ad))&card(c3613,int1))&etype(c3613,int0))&fact(c3613,real))&gener(c3613,gener_c))&quant(c3613,one))&refer(c3613,refer_c))&varia(c3613,varia_c))&sort(schub__1_1,ad))&card(schub__1_1,int1))&etype(schub__1_1,int0))&fact(schub__1_1,real))&gener(schub__1_1,ge))&quant(schub__1_1,one))&refer(schub__1_1,refer_c))&varia(schub__1_1,varia_c))&sort(c3898,ent))&card(c3898,card_c))&etype(c3898,etype_c))&fact(c3898,real))&gener(c3898,gener_c))&quant(c3898,quant_c))&refer(c3898,refer_c))&varia(c3898,varia_c))&sort(c3598,o))&card(c3598,int1))&etype(c3598,int0))&fact(c3598,real))&gener(c3598,sp))&quant(c3598,one))&refer(c3598,det))&varia(c3598,varia_c))&sort(fahrzeug__1_1,d))&card(fahrzeug__1_1,int1))&etype(fahrzeug__1_1,int0))&fact(fahrzeug__1_1,real))&gener(fahrzeug__1_1,ge))&quant(fahrzeug__1_1,one))&refer(fahrzeug__1_1,refer_c))&varia(fahrzeug__1_1,varia_c))&sort(konzept__1_1,d))&sort(konzept__1_1,io))&card(konzept__1_1,int1))&etype(konzept__1_1,int0))&fact(konzept__1_1,real))&gener(konzept__1_1,ge))&quant(konzept__1_1,one))&refer(konzept__1_1,refer_c))&varia(konzept__1_1,varia_c))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ave07_era5_synth_qa07_007_mira_news_1086)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(((arg1(X1,X2)&arg2(X1,X3))&subr(X1,sub_0))=>?[X4]:?[X5]:?[X6]:(((((((arg1(X5,X2)&arg2(X5,X6))&hsit(X1,X4))&mcont(X4,X5))&obj(X4,X2))&sub(X6,X3))&subr(X5,rprs_0))&subs(X4,bezeichnen_1_1)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__bezeichnen_1_1_als)).
fof(c_0_3, axiom, (![X1]:![X2]:(sub(X1,X2)=>?[X3]:((arg1(X3,X1)&arg2(X3,X2))&subr(X3,sub_0)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR004+0.ax', sub__sub_0_expansion)).
fof(c_0_4, negated_conjecture, (~(?[X1]:?[X2]:?[X3]:?[X4]:?[X5]:?[X6]:?[X7]:((((((((attr(X1,X2)&attr(X4,X3))&attr(X6,X7))&obj(X5,X1))&sub(X2,name_1_1))&sub(X1,firma_1_1))&sub(X3,name_1_1))&val(X2,bmw_0))&val(X3,bmw_0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, negated_conjecture, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((((~attr(X8,X9)|~attr(X11,X10))|~attr(X12,X13))|~obj(X14,X8))|~sub(X9,name_1_1))|~sub(X8,firma_1_1))|~sub(X10,name_1_1))|~val(X9,bmw_0))|~val(X10,bmw_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_6,negated_conjecture,(~val(X1,bmw_0)|~val(X2,bmw_0)|~sub(X1,name_1_1)|~sub(X3,firma_1_1)|~sub(X2,name_1_1)|~obj(X4,X3)|~attr(X5,X6)|~attr(X7,X1)|~attr(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7,hypothesis,(val(c3591,bmw_0)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_8,hypothesis,(sub(c3591,name_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_9,hypothesis,(~val(X1,bmw_0)|~attr(X2,c3591)|~attr(X3,X1)|~attr(X4,X5)|~obj(X6,X2)|~sub(X2,firma_1_1)|~sub(X1,name_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_10,hypothesis,(~attr(X1,c3591)|~attr(X2,c3591)|~attr(X3,X4)|~obj(X5,X1)|~sub(X1,firma_1_1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_8])])).
cnf(c_0_11,hypothesis,(attr(c3590,c3591)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_12,hypothesis,(sub(c3590,firma_1_1)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_13,hypothesis,(~attr(X1,c3591)|~attr(X2,X3)|~obj(X4,c3590)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
fof(c_0_14, plain, (![X7]:![X8]:![X9]:((((((((arg1(esk43_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0)))&(arg2(esk43_3(X7,X8,X9),esk44_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(hsit(X7,esk42_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(mcont(esk42_3(X7,X8,X9),esk43_3(X7,X8,X9))|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(obj(esk42_3(X7,X8,X9),X8)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(sub(esk44_3(X7,X8,X9),X9)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subr(esk43_3(X7,X8,X9),rprs_0)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))&(subs(esk42_3(X7,X8,X9),bezeichnen_1_1)|((~arg1(X7,X8)|~arg2(X7,X9))|~subr(X7,sub_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_15,hypothesis,(~attr(X1,X2)|~obj(X3,c3590)), inference(spm,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_16,plain,(obj(esk42_3(X1,X3,X2),X3)|~subr(X1,sub_0)|~arg2(X1,X2)|~arg1(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, (![X4]:![X5]:(((arg1(esk45_2(X4,X5),X4)|~sub(X4,X5))&(arg2(esk45_2(X4,X5),X5)|~sub(X4,X5)))&(subr(esk45_2(X4,X5),sub_0)|~sub(X4,X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_18,hypothesis,(~subr(X1,sub_0)|~attr(X2,X3)|~arg2(X1,X4)|~arg1(X1,c3590)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(subr(esk45_2(X1,X2),sub_0)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,hypothesis,(~attr(X1,X2)|~arg2(esk45_2(X3,X4),X5)|~arg1(esk45_2(X3,X4),c3590)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21,plain,(arg2(esk45_2(X1,X2),X2)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22,hypothesis,(~attr(X1,X2)|~arg1(esk45_2(X3,X4),c3590)|~sub(X3,X4)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23,plain,(arg1(esk45_2(X1,X2),X1)|~sub(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,hypothesis,(~attr(X1,X2)|~sub(c3590,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25,hypothesis,(attr(c3505,c3506)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_26,hypothesis,(~attr(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_12])).
cnf(c_0_27,hypothesis,($false), inference(sr,[status(thm)],[c_0_25, c_0_26]), ['proof']).
# SZS output end CNFRefutation.
EOF
