% Problem    : LCL656+1.020 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n116.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:52:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_B02_F1_AE_CS_SP_PI_S0S
# and selection function SelectComplexG.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (~(?[X1]:~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p8(X1))))))))))))))))))))))|~(((![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((~(![X2]:(~(r1(X1,X2))|~(((~(p21(X2))&~(p121(X2)))&p120(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p21(X2)&~(p121(X2)))&p120(X2))))))|~((~(p120(X1))&p119(X1))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p20(X2))&~(p120(X2)))&p119(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p20(X2)&~(p120(X2)))&p119(X2))))))|~((~(p119(X1))&p118(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p19(X2))&~(p119(X2)))&p118(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p19(X2)&~(p119(X2)))&p118(X2))))))|~((~(p118(X1))&p117(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p18(X2))&~(p118(X2)))&p117(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p18(X2)&~(p118(X2)))&p117(X2))))))|~((~(p117(X1))&p116(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p17(X2))&~(p117(X2)))&p116(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p17(X2)&~(p117(X2)))&p116(X2))))))|~((~(p116(X1))&p115(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p16(X2))&~(p116(X2)))&p115(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p16(X2)&~(p116(X2)))&p115(X2))))))|~((~(p115(X1))&p114(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p15(X2))&~(p115(X2)))&p114(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p15(X2)&~(p115(X2)))&p114(X2))))))|~((~(p114(X1))&p113(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p14(X2))&~(p114(X2)))&p113(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p14(X2)&~(p114(X2)))&p113(X2))))))|~((~(p113(X1))&p112(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p13(X2))&~(p113(X2)))&p112(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p13(X2)&~(p113(X2)))&p112(X2))))))|~((~(p112(X1))&p111(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p12(X2))&~(p112(X2)))&p111(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p12(X2)&~(p112(X2)))&p111(X2))))))|~((~(p111(X1))&p110(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p11(X2))&~(p111(X2)))&p110(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p11(X2)&~(p111(X2)))&p110(X2))))))|~((~(p110(X1))&p109(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p10(X2))&~(p110(X2)))&p109(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p10(X2)&~(p110(X2)))&p109(X2))))))|~((~(p109(X1))&p108(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p9(X2))&~(p109(X2)))&p108(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p9(X2)&~(p109(X2)))&p108(X2))))))|~((~(p108(X1))&p107(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p8(X2))&~(p108(X2)))&p107(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p8(X2)&~(p108(X2)))&p107(X2))))))|~((~(p107(X1))&p106(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p7(X2))&~(p107(X2)))&p106(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p7(X2)&~(p107(X2)))&p106(X2))))))|~((~(p106(X1))&p105(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p6(X2))&~(p106(X2)))&p105(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p6(X2)&~(p106(X2)))&p105(X2))))))|~((~(p105(X1))&p104(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p5(X2))&~(p105(X2)))&p104(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p5(X2)&~(p105(X2)))&p104(X2))))))|~((~(p104(X1))&p103(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p4(X2))&~(p104(X2)))&p103(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p4(X2)&~(p104(X2)))&p103(X2))))))|~((~(p103(X1))&p102(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p3(X2))&~(p103(X2)))&p102(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p3(X2)&~(p103(X2)))&p102(X2))))))|~((~(p102(X1))&p101(X1)))))&((~(![X2]:(~(r1(X1,X2))|~(((~(p2(X2))&~(p102(X2)))&p101(X2)))))&~(![X2]:(~(r1(X1,X2))|~(((p2(X2)&~(p102(X2)))&p101(X2))))))|~((~(p101(X1))&p100(X1)))))&(((![X2]:((~(r1(X1,X2))|~(p21(X2)))|~(p120(X2)))|p21(X1))&(![X2]:((~(r1(X1,X2))|p21(X2))|~(p120(X2)))|~(p21(X1))))|~(p120(X1))))&(((![X2]:((~(r1(X1,X2))|~(p20(X2)))|~(p119(X2)))|p20(X1))&(![X2]:((~(r1(X1,X2))|p20(X2))|~(p119(X2)))|~(p20(X1))))|~(p119(X1))))&(((![X2]:((~(r1(X1,X2))|~(p19(X2)))|~(p118(X2)))|p19(X1))&(![X2]:((~(r1(X1,X2))|p19(X2))|~(p118(X2)))|~(p19(X1))))|~(p118(X1))))&(((![X2]:((~(r1(X1,X2))|~(p18(X2)))|~(p117(X2)))|p18(X1))&(![X2]:((~(r1(X1,X2))|p18(X2))|~(p117(X2)))|~(p18(X1))))|~(p117(X1))))&(((![X2]:((~(r1(X1,X2))|~(p17(X2)))|~(p116(X2)))|p17(X1))&(![X2]:((~(r1(X1,X2))|p17(X2))|~(p116(X2)))|~(p17(X1))))|~(p116(X1))))&(((![X2]:((~(r1(X1,X2))|~(p16(X2)))|~(p115(X2)))|p16(X1))&(![X2]:((~(r1(X1,X2))|p16(X2))|~(p115(X2)))|~(p16(X1))))|~(p115(X1))))&(((![X2]:((~(r1(X1,X2))|~(p15(X2)))|~(p114(X2)))|p15(X1))&(![X2]:((~(r1(X1,X2))|p15(X2))|~(p114(X2)))|~(p15(X1))))|~(p114(X1))))&(((![X2]:((~(r1(X1,X2))|~(p14(X2)))|~(p113(X2)))|p14(X1))&(![X2]:((~(r1(X1,X2))|p14(X2))|~(p113(X2)))|~(p14(X1))))|~(p113(X1))))&(((![X2]:((~(r1(X1,X2))|~(p13(X2)))|~(p112(X2)))|p13(X1))&(![X2]:((~(r1(X1,X2))|p13(X2))|~(p112(X2)))|~(p13(X1))))|~(p112(X1))))&(((![X2]:((~(r1(X1,X2))|~(p12(X2)))|~(p111(X2)))|p12(X1))&(![X2]:((~(r1(X1,X2))|p12(X2))|~(p111(X2)))|~(p12(X1))))|~(p111(X1))))&(((![X2]:((~(r1(X1,X2))|~(p11(X2)))|~(p110(X2)))|p11(X1))&(![X2]:((~(r1(X1,X2))|p11(X2))|~(p110(X2)))|~(p11(X1))))|~(p110(X1))))&(((![X2]:((~(r1(X1,X2))|~(p10(X2)))|~(p109(X2)))|p10(X1))&(![X2]:((~(r1(X1,X2))|p10(X2))|~(p109(X2)))|~(p10(X1))))|~(p109(X1))))&(((![X2]:((~(r1(X1,X2))|~(p9(X2)))|~(p108(X2)))|p9(X1))&(![X2]:((~(r1(X1,X2))|p9(X2))|~(p108(X2)))|~(p9(X1))))|~(p108(X1))))&(((![X2]:((~(r1(X1,X2))|~(p8(X2)))|~(p107(X2)))|p8(X1))&(![X2]:((~(r1(X1,X2))|p8(X2))|~(p107(X2)))|~(p8(X1))))|~(p107(X1))))&(((![X2]:((~(r1(X1,X2))|~(p7(X2)))|~(p106(X2)))|p7(X1))&(![X2]:((~(r1(X1,X2))|p7(X2))|~(p106(X2)))|~(p7(X1))))|~(p106(X1))))&(((![X2]:((~(r1(X1,X2))|~(p6(X2)))|~(p105(X2)))|p6(X1))&(![X2]:((~(r1(X1,X2))|p6(X2))|~(p105(X2)))|~(p6(X1))))|~(p105(X1))))&(((![X2]:((~(r1(X1,X2))|~(p5(X2)))|~(p104(X2)))|p5(X1))&(![X2]:((~(r1(X1,X2))|p5(X2))|~(p104(X2)))|~(p5(X1))))|~(p104(X1))))&(((![X2]:((~(r1(X1,X2))|~(p4(X2)))|~(p103(X2)))|p4(X1))&(![X2]:((~(r1(X1,X2))|p4(X2))|~(p103(X2)))|~(p4(X1))))|~(p103(X1))))&(((![X2]:((~(r1(X1,X2))|~(p3(X2)))|~(p102(X2)))|p3(X1))&(![X2]:((~(r1(X1,X2))|p3(X2))|~(p102(X2)))|~(p3(X1))))|~(p102(X1))))&(((![X2]:((~(r1(X1,X2))|~(p2(X2)))|~(p101(X2)))|p2(X1))&(![X2]:((~(r1(X1,X2))|p2(X2))|~(p101(X2)))|~(p2(X1))))|~(p101(X1))))&(((![X2]:((~(r1(X1,X2))|~(p1(X2)))|~(p100(X2)))|p1(X1))&(![X2]:((~(r1(X1,X2))|p1(X2))|~(p100(X2)))|~(p1(X1))))|~(p100(X1))))&(p120(X1)|~(p121(X1))))&(p119(X1)|~(p120(X1))))&(p118(X1)|~(p119(X1))))&(p117(X1)|~(p118(X1))))&(p116(X1)|~(p117(X1))))&(p115(X1)|~(p116(X1))))&(p114(X1)|~(p115(X1))))&(p113(X1)|~(p114(X1))))&(p112(X1)|~(p113(X1))))&(p111(X1)|~(p112(X1))))&(p110(X1)|~(p111(X1))))&(p109(X1)|~(p110(X1))))&(p108(X1)|~(p109(X1))))&(p107(X1)|~(p108(X1))))&(p106(X1)|~(p107(X1))))&(p105(X1)|~(p106(X1))))&(p104(X1)|~(p105(X1))))&(p103(X1)|~(p104(X1))))&(p102(X1)|~(p103(X1))))&(p101(X1)|~(p102(X1))))&(p100(X1)|~(p101(X1))))))))))))))))))))))))&~(p101(X1)))&p100(X1))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', main)).
fof(c_0_1, axiom, (![X1]:r1(X1,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', reflexivity)).
fof(c_0_2, plain, (![X1]:(epred1_1(X1)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((~(![X2]:(~r1(X1,X2)|~(((~p21(X2)&~p121(X2))&p120(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p21(X2)&~p121(X2))&p120(X2))))))|~((~p120(X1)&p119(X1))))&((~(![X2]:(~r1(X1,X2)|~(((~p20(X2)&~p120(X2))&p119(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p20(X2)&~p120(X2))&p119(X2))))))|~((~p119(X1)&p118(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p19(X2)&~p119(X2))&p118(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p19(X2)&~p119(X2))&p118(X2))))))|~((~p118(X1)&p117(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p18(X2)&~p118(X2))&p117(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p18(X2)&~p118(X2))&p117(X2))))))|~((~p117(X1)&p116(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p17(X2)&~p117(X2))&p116(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p17(X2)&~p117(X2))&p116(X2))))))|~((~p116(X1)&p115(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p16(X2)&~p116(X2))&p115(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p16(X2)&~p116(X2))&p115(X2))))))|~((~p115(X1)&p114(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p15(X2)&~p115(X2))&p114(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p15(X2)&~p115(X2))&p114(X2))))))|~((~p114(X1)&p113(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p14(X2)&~p114(X2))&p113(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p14(X2)&~p114(X2))&p113(X2))))))|~((~p113(X1)&p112(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p13(X2)&~p113(X2))&p112(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p13(X2)&~p113(X2))&p112(X2))))))|~((~p112(X1)&p111(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p12(X2)&~p112(X2))&p111(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p12(X2)&~p112(X2))&p111(X2))))))|~((~p111(X1)&p110(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p11(X2)&~p111(X2))&p110(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p11(X2)&~p111(X2))&p110(X2))))))|~((~p110(X1)&p109(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p10(X2)&~p110(X2))&p109(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p10(X2)&~p110(X2))&p109(X2))))))|~((~p109(X1)&p108(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p9(X2)&~p109(X2))&p108(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p9(X2)&~p109(X2))&p108(X2))))))|~((~p108(X1)&p107(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p8(X2)&~p108(X2))&p107(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p8(X2)&~p108(X2))&p107(X2))))))|~((~p107(X1)&p106(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p7(X2)&~p107(X2))&p106(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p7(X2)&~p107(X2))&p106(X2))))))|~((~p106(X1)&p105(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p6(X2)&~p106(X2))&p105(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p6(X2)&~p106(X2))&p105(X2))))))|~((~p105(X1)&p104(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p5(X2)&~p105(X2))&p104(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p5(X2)&~p105(X2))&p104(X2))))))|~((~p104(X1)&p103(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p4(X2)&~p104(X2))&p103(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p4(X2)&~p104(X2))&p103(X2))))))|~((~p103(X1)&p102(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p3(X2)&~p103(X2))&p102(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p3(X2)&~p103(X2))&p102(X2))))))|~((~p102(X1)&p101(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p2(X2)&~p102(X2))&p101(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p2(X2)&~p102(X2))&p101(X2))))))|~((~p101(X1)&p100(X1)))))&(((![X2]:((~r1(X1,X2)|~p21(X2))|~p120(X2))|p21(X1))&(![X2]:((~r1(X1,X2)|p21(X2))|~p120(X2))|~p21(X1)))|~p120(X1)))&(((![X2]:((~r1(X1,X2)|~p20(X2))|~p119(X2))|p20(X1))&(![X2]:((~r1(X1,X2)|p20(X2))|~p119(X2))|~p20(X1)))|~p119(X1)))&(((![X2]:((~r1(X1,X2)|~p19(X2))|~p118(X2))|p19(X1))&(![X2]:((~r1(X1,X2)|p19(X2))|~p118(X2))|~p19(X1)))|~p118(X1)))&(((![X2]:((~r1(X1,X2)|~p18(X2))|~p117(X2))|p18(X1))&(![X2]:((~r1(X1,X2)|p18(X2))|~p117(X2))|~p18(X1)))|~p117(X1)))&(((![X2]:((~r1(X1,X2)|~p17(X2))|~p116(X2))|p17(X1))&(![X2]:((~r1(X1,X2)|p17(X2))|~p116(X2))|~p17(X1)))|~p116(X1)))&(((![X2]:((~r1(X1,X2)|~p16(X2))|~p115(X2))|p16(X1))&(![X2]:((~r1(X1,X2)|p16(X2))|~p115(X2))|~p16(X1)))|~p115(X1)))&(((![X2]:((~r1(X1,X2)|~p15(X2))|~p114(X2))|p15(X1))&(![X2]:((~r1(X1,X2)|p15(X2))|~p114(X2))|~p15(X1)))|~p114(X1)))&(((![X2]:((~r1(X1,X2)|~p14(X2))|~p113(X2))|p14(X1))&(![X2]:((~r1(X1,X2)|p14(X2))|~p113(X2))|~p14(X1)))|~p113(X1)))&(((![X2]:((~r1(X1,X2)|~p13(X2))|~p112(X2))|p13(X1))&(![X2]:((~r1(X1,X2)|p13(X2))|~p112(X2))|~p13(X1)))|~p112(X1)))&(((![X2]:((~r1(X1,X2)|~p12(X2))|~p111(X2))|p12(X1))&(![X2]:((~r1(X1,X2)|p12(X2))|~p111(X2))|~p12(X1)))|~p111(X1)))&(((![X2]:((~r1(X1,X2)|~p11(X2))|~p110(X2))|p11(X1))&(![X2]:((~r1(X1,X2)|p11(X2))|~p110(X2))|~p11(X1)))|~p110(X1)))&(((![X2]:((~r1(X1,X2)|~p10(X2))|~p109(X2))|p10(X1))&(![X2]:((~r1(X1,X2)|p10(X2))|~p109(X2))|~p10(X1)))|~p109(X1)))&(((![X2]:((~r1(X1,X2)|~p9(X2))|~p108(X2))|p9(X1))&(![X2]:((~r1(X1,X2)|p9(X2))|~p108(X2))|~p9(X1)))|~p108(X1)))&(((![X2]:((~r1(X1,X2)|~p8(X2))|~p107(X2))|p8(X1))&(![X2]:((~r1(X1,X2)|p8(X2))|~p107(X2))|~p8(X1)))|~p107(X1)))&(((![X2]:((~r1(X1,X2)|~p7(X2))|~p106(X2))|p7(X1))&(![X2]:((~r1(X1,X2)|p7(X2))|~p106(X2))|~p7(X1)))|~p106(X1)))&(((![X2]:((~r1(X1,X2)|~p6(X2))|~p105(X2))|p6(X1))&(![X2]:((~r1(X1,X2)|p6(X2))|~p105(X2))|~p6(X1)))|~p105(X1)))&(((![X2]:((~r1(X1,X2)|~p5(X2))|~p104(X2))|p5(X1))&(![X2]:((~r1(X1,X2)|p5(X2))|~p104(X2))|~p5(X1)))|~p104(X1)))&(((![X2]:((~r1(X1,X2)|~p4(X2))|~p103(X2))|p4(X1))&(![X2]:((~r1(X1,X2)|p4(X2))|~p103(X2))|~p4(X1)))|~p103(X1)))&(((![X2]:((~r1(X1,X2)|~p3(X2))|~p102(X2))|p3(X1))&(![X2]:((~r1(X1,X2)|p3(X2))|~p102(X2))|~p3(X1)))|~p102(X1)))&(((![X2]:((~r1(X1,X2)|~p2(X2))|~p101(X2))|p2(X1))&(![X2]:((~r1(X1,X2)|p2(X2))|~p101(X2))|~p2(X1)))|~p101(X1)))&(((![X2]:((~r1(X1,X2)|~p1(X2))|~p100(X2))|p1(X1))&(![X2]:((~r1(X1,X2)|p1(X2))|~p100(X2))|~p1(X1)))|~p100(X1)))&(p120(X1)|~p121(X1)))&(p119(X1)|~p120(X1)))&(p118(X1)|~p119(X1)))&(p117(X1)|~p118(X1)))&(p116(X1)|~p117(X1)))&(p115(X1)|~p116(X1)))&(p114(X1)|~p115(X1)))&(p113(X1)|~p114(X1)))&(p112(X1)|~p113(X1)))&(p111(X1)|~p112(X1)))&(p110(X1)|~p111(X1)))&(p109(X1)|~p110(X1)))&(p108(X1)|~p109(X1)))&(p107(X1)|~p108(X1)))&(p106(X1)|~p107(X1)))&(p105(X1)|~p106(X1)))&(p104(X1)|~p105(X1)))&(p103(X1)|~p104(X1)))&(p102(X1)|~p103(X1)))&(p101(X1)|~p102(X1)))&(p100(X1)|~p101(X1))))), introduced(definition)).
fof(c_0_3, negated_conjecture, (~(~(?[X1]:~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p8(X1))))))))))))))))))))))|~(((![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|epred1_1(X1)))))))))))))))))))))&~p101(X1))&p100(X1)))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_2])).
fof(c_0_4, negated_conjecture, (![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:![X34]:![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:![X43]:((~r1(esk1_0,X4)|(~r1(X4,X5)|(~r1(X5,X6)|(~r1(X6,X7)|(~r1(X7,X8)|(~r1(X8,X9)|(~r1(X9,X10)|(~r1(X10,X11)|(~r1(X11,X12)|(~r1(X12,X13)|(~r1(X13,X14)|(~r1(X14,X15)|(~r1(X15,X16)|(~r1(X16,X17)|(~r1(X17,X18)|(~r1(X18,X19)|(~r1(X19,X20)|(~r1(X20,X21)|(~r1(X21,X22)|(~r1(X22,X23)|p8(X23)))))))))))))))))))))&(((~r1(esk1_0,X24)|(~r1(X24,X25)|(~r1(X25,X26)|(~r1(X26,X27)|(~r1(X27,X28)|(~r1(X28,X29)|(~r1(X29,X30)|(~r1(X30,X31)|(~r1(X31,X32)|(~r1(X32,X33)|(~r1(X33,X34)|(~r1(X34,X35)|(~r1(X35,X36)|(~r1(X36,X37)|(~r1(X37,X38)|(~r1(X38,X39)|(~r1(X39,X40)|(~r1(X40,X41)|(~r1(X41,X42)|(~r1(X42,X43)|epred1_1(X43)))))))))))))))))))))&~p101(esk1_0))&p100(esk1_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
fof(c_0_5, plain, (![X2]:r1(X2,X2)), inference(variable_rename,[status(thm)],[c_0_1])).
cnf(c_0_6,negated_conjecture,(epred1_1(X1)|~r1(X2,X1)|~r1(X3,X2)|~r1(X4,X3)|~r1(X5,X4)|~r1(X6,X5)|~r1(X7,X6)|~r1(X8,X7)|~r1(X9,X8)|~r1(X10,X9)|~r1(X11,X10)|~r1(X12,X11)|~r1(X13,X12)|~r1(X14,X13)|~r1(X15,X14)|~r1(X16,X15)|~r1(X17,X16)|~r1(X18,X17)|~r1(X19,X18)|~r1(X20,X19)|~r1(esk1_0,X20)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(r1(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X16)|~r1(X16,X17)|~r1(X17,X18)|~r1(X18,X19)|~r1(X19,X1)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_9,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X16)|~r1(X16,X17)|~r1(X17,X18)|~r1(X18,X1)), inference(spm,[status(thm)],[c_0_8, c_0_7])).
cnf(c_0_10,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X16)|~r1(X16,X17)|~r1(X17,X1)), inference(spm,[status(thm)],[c_0_9, c_0_7])).
cnf(c_0_11,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X16)|~r1(X16,X1)), inference(spm,[status(thm)],[c_0_10, c_0_7])).
cnf(c_0_12,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X1)), inference(spm,[status(thm)],[c_0_11, c_0_7])).
cnf(c_0_13,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X1)), inference(spm,[status(thm)],[c_0_12, c_0_7])).
cnf(c_0_14,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X1)), inference(spm,[status(thm)],[c_0_13, c_0_7])).
cnf(c_0_15,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X1)), inference(spm,[status(thm)],[c_0_14, c_0_7])).
cnf(c_0_16,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X1)), inference(spm,[status(thm)],[c_0_15, c_0_7])).
cnf(c_0_17,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X1)), inference(spm,[status(thm)],[c_0_16, c_0_7])).
cnf(c_0_18,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X1)), inference(spm,[status(thm)],[c_0_17, c_0_7])).
cnf(c_0_19,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X1)), inference(spm,[status(thm)],[c_0_18, c_0_7])).
cnf(c_0_20,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X1)), inference(spm,[status(thm)],[c_0_19, c_0_7])).
cnf(c_0_21,negated_conjecture,(p8(X1)|~r1(X2,X1)|~r1(X3,X2)|~r1(X4,X3)|~r1(X5,X4)|~r1(X6,X5)|~r1(X7,X6)|~r1(X8,X7)|~r1(X9,X8)|~r1(X10,X9)|~r1(X11,X10)|~r1(X12,X11)|~r1(X13,X12)|~r1(X14,X13)|~r1(X15,X14)|~r1(X16,X15)|~r1(X17,X16)|~r1(X18,X17)|~r1(X19,X18)|~r1(X20,X19)|~r1(esk1_0,X20)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X1)), inference(spm,[status(thm)],[c_0_20, c_0_7])).
cnf(c_0_23,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X16)|~r1(X16,X17)|~r1(X17,X18)|~r1(X18,X19)|~r1(X19,X1)), inference(spm,[status(thm)],[c_0_21, c_0_7])).
cnf(c_0_24,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X1)), inference(spm,[status(thm)],[c_0_22, c_0_7])).
cnf(c_0_25,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X16)|~r1(X16,X17)|~r1(X17,X18)|~r1(X18,X1)), inference(spm,[status(thm)],[c_0_23, c_0_7])).
cnf(c_0_26,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X1)), inference(spm,[status(thm)],[c_0_24, c_0_7])).
cnf(c_0_27,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X16)|~r1(X16,X17)|~r1(X17,X1)), inference(spm,[status(thm)],[c_0_25, c_0_7])).
cnf(c_0_28,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_26, c_0_7])).
cnf(c_0_29,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X16)|~r1(X16,X1)), inference(spm,[status(thm)],[c_0_27, c_0_7])).
fof(c_0_30, plain, (![X1]:(epred1_1(X1)=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((~(![X2]:(~r1(X1,X2)|~(((~p21(X2)&~p121(X2))&p120(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p21(X2)&~p121(X2))&p120(X2))))))|~((~p120(X1)&p119(X1))))&((~(![X2]:(~r1(X1,X2)|~(((~p20(X2)&~p120(X2))&p119(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p20(X2)&~p120(X2))&p119(X2))))))|~((~p119(X1)&p118(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p19(X2)&~p119(X2))&p118(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p19(X2)&~p119(X2))&p118(X2))))))|~((~p118(X1)&p117(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p18(X2)&~p118(X2))&p117(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p18(X2)&~p118(X2))&p117(X2))))))|~((~p117(X1)&p116(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p17(X2)&~p117(X2))&p116(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p17(X2)&~p117(X2))&p116(X2))))))|~((~p116(X1)&p115(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p16(X2)&~p116(X2))&p115(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p16(X2)&~p116(X2))&p115(X2))))))|~((~p115(X1)&p114(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p15(X2)&~p115(X2))&p114(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p15(X2)&~p115(X2))&p114(X2))))))|~((~p114(X1)&p113(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p14(X2)&~p114(X2))&p113(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p14(X2)&~p114(X2))&p113(X2))))))|~((~p113(X1)&p112(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p13(X2)&~p113(X2))&p112(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p13(X2)&~p113(X2))&p112(X2))))))|~((~p112(X1)&p111(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p12(X2)&~p112(X2))&p111(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p12(X2)&~p112(X2))&p111(X2))))))|~((~p111(X1)&p110(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p11(X2)&~p111(X2))&p110(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p11(X2)&~p111(X2))&p110(X2))))))|~((~p110(X1)&p109(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p10(X2)&~p110(X2))&p109(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p10(X2)&~p110(X2))&p109(X2))))))|~((~p109(X1)&p108(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p9(X2)&~p109(X2))&p108(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p9(X2)&~p109(X2))&p108(X2))))))|~((~p108(X1)&p107(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p8(X2)&~p108(X2))&p107(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p8(X2)&~p108(X2))&p107(X2))))))|~((~p107(X1)&p106(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p7(X2)&~p107(X2))&p106(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p7(X2)&~p107(X2))&p106(X2))))))|~((~p106(X1)&p105(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p6(X2)&~p106(X2))&p105(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p6(X2)&~p106(X2))&p105(X2))))))|~((~p105(X1)&p104(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p5(X2)&~p105(X2))&p104(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p5(X2)&~p105(X2))&p104(X2))))))|~((~p104(X1)&p103(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p4(X2)&~p104(X2))&p103(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p4(X2)&~p104(X2))&p103(X2))))))|~((~p103(X1)&p102(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p3(X2)&~p103(X2))&p102(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p3(X2)&~p103(X2))&p102(X2))))))|~((~p102(X1)&p101(X1)))))&((~(![X2]:(~r1(X1,X2)|~(((~p2(X2)&~p102(X2))&p101(X2)))))&~(![X2]:(~r1(X1,X2)|~(((p2(X2)&~p102(X2))&p101(X2))))))|~((~p101(X1)&p100(X1)))))&(((![X2]:((~r1(X1,X2)|~p21(X2))|~p120(X2))|p21(X1))&(![X2]:((~r1(X1,X2)|p21(X2))|~p120(X2))|~p21(X1)))|~p120(X1)))&(((![X2]:((~r1(X1,X2)|~p20(X2))|~p119(X2))|p20(X1))&(![X2]:((~r1(X1,X2)|p20(X2))|~p119(X2))|~p20(X1)))|~p119(X1)))&(((![X2]:((~r1(X1,X2)|~p19(X2))|~p118(X2))|p19(X1))&(![X2]:((~r1(X1,X2)|p19(X2))|~p118(X2))|~p19(X1)))|~p118(X1)))&(((![X2]:((~r1(X1,X2)|~p18(X2))|~p117(X2))|p18(X1))&(![X2]:((~r1(X1,X2)|p18(X2))|~p117(X2))|~p18(X1)))|~p117(X1)))&(((![X2]:((~r1(X1,X2)|~p17(X2))|~p116(X2))|p17(X1))&(![X2]:((~r1(X1,X2)|p17(X2))|~p116(X2))|~p17(X1)))|~p116(X1)))&(((![X2]:((~r1(X1,X2)|~p16(X2))|~p115(X2))|p16(X1))&(![X2]:((~r1(X1,X2)|p16(X2))|~p115(X2))|~p16(X1)))|~p115(X1)))&(((![X2]:((~r1(X1,X2)|~p15(X2))|~p114(X2))|p15(X1))&(![X2]:((~r1(X1,X2)|p15(X2))|~p114(X2))|~p15(X1)))|~p114(X1)))&(((![X2]:((~r1(X1,X2)|~p14(X2))|~p113(X2))|p14(X1))&(![X2]:((~r1(X1,X2)|p14(X2))|~p113(X2))|~p14(X1)))|~p113(X1)))&(((![X2]:((~r1(X1,X2)|~p13(X2))|~p112(X2))|p13(X1))&(![X2]:((~r1(X1,X2)|p13(X2))|~p112(X2))|~p13(X1)))|~p112(X1)))&(((![X2]:((~r1(X1,X2)|~p12(X2))|~p111(X2))|p12(X1))&(![X2]:((~r1(X1,X2)|p12(X2))|~p111(X2))|~p12(X1)))|~p111(X1)))&(((![X2]:((~r1(X1,X2)|~p11(X2))|~p110(X2))|p11(X1))&(![X2]:((~r1(X1,X2)|p11(X2))|~p110(X2))|~p11(X1)))|~p110(X1)))&(((![X2]:((~r1(X1,X2)|~p10(X2))|~p109(X2))|p10(X1))&(![X2]:((~r1(X1,X2)|p10(X2))|~p109(X2))|~p10(X1)))|~p109(X1)))&(((![X2]:((~r1(X1,X2)|~p9(X2))|~p108(X2))|p9(X1))&(![X2]:((~r1(X1,X2)|p9(X2))|~p108(X2))|~p9(X1)))|~p108(X1)))&(((![X2]:((~r1(X1,X2)|~p8(X2))|~p107(X2))|p8(X1))&(![X2]:((~r1(X1,X2)|p8(X2))|~p107(X2))|~p8(X1)))|~p107(X1)))&(((![X2]:((~r1(X1,X2)|~p7(X2))|~p106(X2))|p7(X1))&(![X2]:((~r1(X1,X2)|p7(X2))|~p106(X2))|~p7(X1)))|~p106(X1)))&(((![X2]:((~r1(X1,X2)|~p6(X2))|~p105(X2))|p6(X1))&(![X2]:((~r1(X1,X2)|p6(X2))|~p105(X2))|~p6(X1)))|~p105(X1)))&(((![X2]:((~r1(X1,X2)|~p5(X2))|~p104(X2))|p5(X1))&(![X2]:((~r1(X1,X2)|p5(X2))|~p104(X2))|~p5(X1)))|~p104(X1)))&(((![X2]:((~r1(X1,X2)|~p4(X2))|~p103(X2))|p4(X1))&(![X2]:((~r1(X1,X2)|p4(X2))|~p103(X2))|~p4(X1)))|~p103(X1)))&(((![X2]:((~r1(X1,X2)|~p3(X2))|~p102(X2))|p3(X1))&(![X2]:((~r1(X1,X2)|p3(X2))|~p102(X2))|~p3(X1)))|~p102(X1)))&(((![X2]:((~r1(X1,X2)|~p2(X2))|~p101(X2))|p2(X1))&(![X2]:((~r1(X1,X2)|p2(X2))|~p101(X2))|~p2(X1)))|~p101(X1)))&(((![X2]:((~r1(X1,X2)|~p1(X2))|~p100(X2))|p1(X1))&(![X2]:((~r1(X1,X2)|p1(X2))|~p100(X2))|~p1(X1)))|~p100(X1)))&(p120(X1)|~p121(X1)))&(p119(X1)|~p120(X1)))&(p118(X1)|~p119(X1)))&(p117(X1)|~p118(X1)))&(p116(X1)|~p117(X1)))&(p115(X1)|~p116(X1)))&(p114(X1)|~p115(X1)))&(p113(X1)|~p114(X1)))&(p112(X1)|~p113(X1)))&(p111(X1)|~p112(X1)))&(p110(X1)|~p111(X1)))&(p109(X1)|~p110(X1)))&(p108(X1)|~p109(X1)))&(p107(X1)|~p108(X1)))&(p106(X1)|~p107(X1)))&(p105(X1)|~p106(X1)))&(p104(X1)|~p105(X1)))&(p103(X1)|~p104(X1)))&(p102(X1)|~p103(X1)))&(p101(X1)|~p102(X1)))&(p100(X1)|~p101(X1))))), inference(split_equiv,[status(thm)],[c_0_2])).
cnf(c_0_31,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_28, c_0_7])).
cnf(c_0_32,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X15)|~r1(X15,X1)), inference(spm,[status(thm)],[c_0_29, c_0_7])).
fof(c_0_33, plain, (![X3]:![X44]:![X45]:![X46]:![X47]:![X48]:![X49]:![X50]:![X51]:![X52]:![X53]:![X54]:![X55]:![X56]:![X57]:![X58]:![X59]:![X60]:![X61]:![X62]:![X63]:![X64]:![X65]:![X66]:![X67]:![X68]:![X69]:![X70]:![X71]:![X72]:![X73]:![X74]:![X75]:![X76]:![X77]:![X78]:![X79]:![X80]:![X81]:![X82]:![X83]:![X84]:![X85]:(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((r1(X3,esk2_1(X3))|(p120(X3)|~p119(X3)))|~epred1_1(X3))&((((~p21(esk2_1(X3))|(p120(X3)|~p119(X3)))|~epred1_1(X3))&((~p121(esk2_1(X3))|(p120(X3)|~p119(X3)))|~epred1_1(X3)))&((p120(esk2_1(X3))|(p120(X3)|~p119(X3)))|~epred1_1(X3))))&(((r1(X3,esk3_1(X3))|(p120(X3)|~p119(X3)))|~epred1_1(X3))&((((p21(esk3_1(X3))|(p120(X3)|~p119(X3)))|~epred1_1(X3))&((~p121(esk3_1(X3))|(p120(X3)|~p119(X3)))|~epred1_1(X3)))&((p120(esk3_1(X3))|(p120(X3)|~p119(X3)))|~epred1_1(X3)))))&((((r1(X3,esk4_1(X3))|(p119(X3)|~p118(X3)))|~epred1_1(X3))&((((~p20(esk4_1(X3))|(p119(X3)|~p118(X3)))|~epred1_1(X3))&((~p120(esk4_1(X3))|(p119(X3)|~p118(X3)))|~epred1_1(X3)))&((p119(esk4_1(X3))|(p119(X3)|~p118(X3)))|~epred1_1(X3))))&(((r1(X3,esk5_1(X3))|(p119(X3)|~p118(X3)))|~epred1_1(X3))&((((p20(esk5_1(X3))|(p119(X3)|~p118(X3)))|~epred1_1(X3))&((~p120(esk5_1(X3))|(p119(X3)|~p118(X3)))|~epred1_1(X3)))&((p119(esk5_1(X3))|(p119(X3)|~p118(X3)))|~epred1_1(X3))))))&((((r1(X3,esk6_1(X3))|(p118(X3)|~p117(X3)))|~epred1_1(X3))&((((~p19(esk6_1(X3))|(p118(X3)|~p117(X3)))|~epred1_1(X3))&((~p119(esk6_1(X3))|(p118(X3)|~p117(X3)))|~epred1_1(X3)))&((p118(esk6_1(X3))|(p118(X3)|~p117(X3)))|~epred1_1(X3))))&(((r1(X3,esk7_1(X3))|(p118(X3)|~p117(X3)))|~epred1_1(X3))&((((p19(esk7_1(X3))|(p118(X3)|~p117(X3)))|~epred1_1(X3))&((~p119(esk7_1(X3))|(p118(X3)|~p117(X3)))|~epred1_1(X3)))&((p118(esk7_1(X3))|(p118(X3)|~p117(X3)))|~epred1_1(X3))))))&((((r1(X3,esk8_1(X3))|(p117(X3)|~p116(X3)))|~epred1_1(X3))&((((~p18(esk8_1(X3))|(p117(X3)|~p116(X3)))|~epred1_1(X3))&((~p118(esk8_1(X3))|(p117(X3)|~p116(X3)))|~epred1_1(X3)))&((p117(esk8_1(X3))|(p117(X3)|~p116(X3)))|~epred1_1(X3))))&(((r1(X3,esk9_1(X3))|(p117(X3)|~p116(X3)))|~epred1_1(X3))&((((p18(esk9_1(X3))|(p117(X3)|~p116(X3)))|~epred1_1(X3))&((~p118(esk9_1(X3))|(p117(X3)|~p116(X3)))|~epred1_1(X3)))&((p117(esk9_1(X3))|(p117(X3)|~p116(X3)))|~epred1_1(X3))))))&((((r1(X3,esk10_1(X3))|(p116(X3)|~p115(X3)))|~epred1_1(X3))&((((~p17(esk10_1(X3))|(p116(X3)|~p115(X3)))|~epred1_1(X3))&((~p117(esk10_1(X3))|(p116(X3)|~p115(X3)))|~epred1_1(X3)))&((p116(esk10_1(X3))|(p116(X3)|~p115(X3)))|~epred1_1(X3))))&(((r1(X3,esk11_1(X3))|(p116(X3)|~p115(X3)))|~epred1_1(X3))&((((p17(esk11_1(X3))|(p116(X3)|~p115(X3)))|~epred1_1(X3))&((~p117(esk11_1(X3))|(p116(X3)|~p115(X3)))|~epred1_1(X3)))&((p116(esk11_1(X3))|(p116(X3)|~p115(X3)))|~epred1_1(X3))))))&((((r1(X3,esk12_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))&((((~p16(esk12_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))&((~p116(esk12_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3)))&((p115(esk12_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))))&(((r1(X3,esk13_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))&((((p16(esk13_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))&((~p116(esk13_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3)))&((p115(esk13_1(X3))|(p115(X3)|~p114(X3)))|~epred1_1(X3))))))&((((r1(X3,esk14_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))&((((~p15(esk14_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))&((~p115(esk14_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3)))&((p114(esk14_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))))&(((r1(X3,esk15_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))&((((p15(esk15_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))&((~p115(esk15_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3)))&((p114(esk15_1(X3))|(p114(X3)|~p113(X3)))|~epred1_1(X3))))))&((((r1(X3,esk16_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))&((((~p14(esk16_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))&((~p114(esk16_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3)))&((p113(esk16_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))))&(((r1(X3,esk17_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))&((((p14(esk17_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))&((~p114(esk17_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3)))&((p113(esk17_1(X3))|(p113(X3)|~p112(X3)))|~epred1_1(X3))))))&((((r1(X3,esk18_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))&((((~p13(esk18_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))&((~p113(esk18_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3)))&((p112(esk18_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))))&(((r1(X3,esk19_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))&((((p13(esk19_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))&((~p113(esk19_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3)))&((p112(esk19_1(X3))|(p112(X3)|~p111(X3)))|~epred1_1(X3))))))&((((r1(X3,esk20_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))&((((~p12(esk20_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))&((~p112(esk20_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3)))&((p111(esk20_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))))&(((r1(X3,esk21_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))&((((p12(esk21_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))&((~p112(esk21_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3)))&((p111(esk21_1(X3))|(p111(X3)|~p110(X3)))|~epred1_1(X3))))))&((((r1(X3,esk22_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))&((((~p11(esk22_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))&((~p111(esk22_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3)))&((p110(esk22_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))))&(((r1(X3,esk23_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))&((((p11(esk23_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))&((~p111(esk23_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3)))&((p110(esk23_1(X3))|(p110(X3)|~p109(X3)))|~epred1_1(X3))))))&((((r1(X3,esk24_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))&((((~p10(esk24_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))&((~p110(esk24_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3)))&((p109(esk24_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))))&(((r1(X3,esk25_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))&((((p10(esk25_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))&((~p110(esk25_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3)))&((p109(esk25_1(X3))|(p109(X3)|~p108(X3)))|~epred1_1(X3))))))&((((r1(X3,esk26_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))&((((~p9(esk26_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))&((~p109(esk26_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3)))&((p108(esk26_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))))&(((r1(X3,esk27_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))&((((p9(esk27_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))&((~p109(esk27_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3)))&((p108(esk27_1(X3))|(p108(X3)|~p107(X3)))|~epred1_1(X3))))))&((((r1(X3,esk28_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))&((((~p8(esk28_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))&((~p108(esk28_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3)))&((p107(esk28_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))))&(((r1(X3,esk29_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))&((((p8(esk29_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))&((~p108(esk29_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3)))&((p107(esk29_1(X3))|(p107(X3)|~p106(X3)))|~epred1_1(X3))))))&((((r1(X3,esk30_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))&((((~p7(esk30_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))&((~p107(esk30_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3)))&((p106(esk30_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))))&(((r1(X3,esk31_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))&((((p7(esk31_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))&((~p107(esk31_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3)))&((p106(esk31_1(X3))|(p106(X3)|~p105(X3)))|~epred1_1(X3))))))&((((r1(X3,esk32_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))&((((~p6(esk32_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))&((~p106(esk32_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3)))&((p105(esk32_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))))&(((r1(X3,esk33_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))&((((p6(esk33_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))&((~p106(esk33_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3)))&((p105(esk33_1(X3))|(p105(X3)|~p104(X3)))|~epred1_1(X3))))))&((((r1(X3,esk34_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))&((((~p5(esk34_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))&((~p105(esk34_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3)))&((p104(esk34_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))))&(((r1(X3,esk35_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))&((((p5(esk35_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))&((~p105(esk35_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3)))&((p104(esk35_1(X3))|(p104(X3)|~p103(X3)))|~epred1_1(X3))))))&((((r1(X3,esk36_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))&((((~p4(esk36_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))&((~p104(esk36_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3)))&((p103(esk36_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))))&(((r1(X3,esk37_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))&((((p4(esk37_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))&((~p104(esk37_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3)))&((p103(esk37_1(X3))|(p103(X3)|~p102(X3)))|~epred1_1(X3))))))&((((r1(X3,esk38_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))&((((~p3(esk38_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))&((~p103(esk38_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3)))&((p102(esk38_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))))&(((r1(X3,esk39_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))&((((p3(esk39_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))&((~p103(esk39_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3)))&((p102(esk39_1(X3))|(p102(X3)|~p101(X3)))|~epred1_1(X3))))))&((((r1(X3,esk40_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))&((((~p2(esk40_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))&((~p102(esk40_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3)))&((p101(esk40_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))))&(((r1(X3,esk41_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))&((((p2(esk41_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))&((~p102(esk41_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3)))&((p101(esk41_1(X3))|(p101(X3)|~p100(X3)))|~epred1_1(X3))))))&((((((~r1(X3,X44)|~p21(X44))|~p120(X44))|p21(X3))|~p120(X3))|~epred1_1(X3))&(((((~r1(X3,X45)|p21(X45))|~p120(X45))|~p21(X3))|~p120(X3))|~epred1_1(X3))))&((((((~r1(X3,X46)|~p20(X46))|~p119(X46))|p20(X3))|~p119(X3))|~epred1_1(X3))&(((((~r1(X3,X47)|p20(X47))|~p119(X47))|~p20(X3))|~p119(X3))|~epred1_1(X3))))&((((((~r1(X3,X48)|~p19(X48))|~p118(X48))|p19(X3))|~p118(X3))|~epred1_1(X3))&(((((~r1(X3,X49)|p19(X49))|~p118(X49))|~p19(X3))|~p118(X3))|~epred1_1(X3))))&((((((~r1(X3,X50)|~p18(X50))|~p117(X50))|p18(X3))|~p117(X3))|~epred1_1(X3))&(((((~r1(X3,X51)|p18(X51))|~p117(X51))|~p18(X3))|~p117(X3))|~epred1_1(X3))))&((((((~r1(X3,X52)|~p17(X52))|~p116(X52))|p17(X3))|~p116(X3))|~epred1_1(X3))&(((((~r1(X3,X53)|p17(X53))|~p116(X53))|~p17(X3))|~p116(X3))|~epred1_1(X3))))&((((((~r1(X3,X54)|~p16(X54))|~p115(X54))|p16(X3))|~p115(X3))|~epred1_1(X3))&(((((~r1(X3,X55)|p16(X55))|~p115(X55))|~p16(X3))|~p115(X3))|~epred1_1(X3))))&((((((~r1(X3,X56)|~p15(X56))|~p114(X56))|p15(X3))|~p114(X3))|~epred1_1(X3))&(((((~r1(X3,X57)|p15(X57))|~p114(X57))|~p15(X3))|~p114(X3))|~epred1_1(X3))))&((((((~r1(X3,X58)|~p14(X58))|~p113(X58))|p14(X3))|~p113(X3))|~epred1_1(X3))&(((((~r1(X3,X59)|p14(X59))|~p113(X59))|~p14(X3))|~p113(X3))|~epred1_1(X3))))&((((((~r1(X3,X60)|~p13(X60))|~p112(X60))|p13(X3))|~p112(X3))|~epred1_1(X3))&(((((~r1(X3,X61)|p13(X61))|~p112(X61))|~p13(X3))|~p112(X3))|~epred1_1(X3))))&((((((~r1(X3,X62)|~p12(X62))|~p111(X62))|p12(X3))|~p111(X3))|~epred1_1(X3))&(((((~r1(X3,X63)|p12(X63))|~p111(X63))|~p12(X3))|~p111(X3))|~epred1_1(X3))))&((((((~r1(X3,X64)|~p11(X64))|~p110(X64))|p11(X3))|~p110(X3))|~epred1_1(X3))&(((((~r1(X3,X65)|p11(X65))|~p110(X65))|~p11(X3))|~p110(X3))|~epred1_1(X3))))&((((((~r1(X3,X66)|~p10(X66))|~p109(X66))|p10(X3))|~p109(X3))|~epred1_1(X3))&(((((~r1(X3,X67)|p10(X67))|~p109(X67))|~p10(X3))|~p109(X3))|~epred1_1(X3))))&((((((~r1(X3,X68)|~p9(X68))|~p108(X68))|p9(X3))|~p108(X3))|~epred1_1(X3))&(((((~r1(X3,X69)|p9(X69))|~p108(X69))|~p9(X3))|~p108(X3))|~epred1_1(X3))))&((((((~r1(X3,X70)|~p8(X70))|~p107(X70))|p8(X3))|~p107(X3))|~epred1_1(X3))&(((((~r1(X3,X71)|p8(X71))|~p107(X71))|~p8(X3))|~p107(X3))|~epred1_1(X3))))&((((((~r1(X3,X72)|~p7(X72))|~p106(X72))|p7(X3))|~p106(X3))|~epred1_1(X3))&(((((~r1(X3,X73)|p7(X73))|~p106(X73))|~p7(X3))|~p106(X3))|~epred1_1(X3))))&((((((~r1(X3,X74)|~p6(X74))|~p105(X74))|p6(X3))|~p105(X3))|~epred1_1(X3))&(((((~r1(X3,X75)|p6(X75))|~p105(X75))|~p6(X3))|~p105(X3))|~epred1_1(X3))))&((((((~r1(X3,X76)|~p5(X76))|~p104(X76))|p5(X3))|~p104(X3))|~epred1_1(X3))&(((((~r1(X3,X77)|p5(X77))|~p104(X77))|~p5(X3))|~p104(X3))|~epred1_1(X3))))&((((((~r1(X3,X78)|~p4(X78))|~p103(X78))|p4(X3))|~p103(X3))|~epred1_1(X3))&(((((~r1(X3,X79)|p4(X79))|~p103(X79))|~p4(X3))|~p103(X3))|~epred1_1(X3))))&((((((~r1(X3,X80)|~p3(X80))|~p102(X80))|p3(X3))|~p102(X3))|~epred1_1(X3))&(((((~r1(X3,X81)|p3(X81))|~p102(X81))|~p3(X3))|~p102(X3))|~epred1_1(X3))))&((((((~r1(X3,X82)|~p2(X82))|~p101(X82))|p2(X3))|~p101(X3))|~epred1_1(X3))&(((((~r1(X3,X83)|p2(X83))|~p101(X83))|~p2(X3))|~p101(X3))|~epred1_1(X3))))&((((((~r1(X3,X84)|~p1(X84))|~p100(X84))|p1(X3))|~p100(X3))|~epred1_1(X3))&(((((~r1(X3,X85)|p1(X85))|~p100(X85))|~p1(X3))|~p100(X3))|~epred1_1(X3))))&((p120(X3)|~p121(X3))|~epred1_1(X3)))&((p119(X3)|~p120(X3))|~epred1_1(X3)))&((p118(X3)|~p119(X3))|~epred1_1(X3)))&((p117(X3)|~p118(X3))|~epred1_1(X3)))&((p116(X3)|~p117(X3))|~epred1_1(X3)))&((p115(X3)|~p116(X3))|~epred1_1(X3)))&((p114(X3)|~p115(X3))|~epred1_1(X3)))&((p113(X3)|~p114(X3))|~epred1_1(X3)))&((p112(X3)|~p113(X3))|~epred1_1(X3)))&((p111(X3)|~p112(X3))|~epred1_1(X3)))&((p110(X3)|~p111(X3))|~epred1_1(X3)))&((p109(X3)|~p110(X3))|~epred1_1(X3)))&((p108(X3)|~p109(X3))|~epred1_1(X3)))&((p107(X3)|~p108(X3))|~epred1_1(X3)))&((p106(X3)|~p107(X3))|~epred1_1(X3)))&((p105(X3)|~p106(X3))|~epred1_1(X3)))&((p104(X3)|~p105(X3))|~epred1_1(X3)))&((p103(X3)|~p104(X3))|~epred1_1(X3)))&((p102(X3)|~p103(X3))|~epred1_1(X3)))&((p101(X3)|~p102(X3))|~epred1_1(X3)))&((p100(X3)|~p101(X3))|~epred1_1(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])])).
cnf(c_0_34,negated_conjecture,(epred1_1(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_31, c_0_7])).
cnf(c_0_35,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X14)|~r1(X14,X1)), inference(spm,[status(thm)],[c_0_32, c_0_7])).
cnf(c_0_36,plain,(p101(X1)|r1(X1,esk40_1(X1))|~epred1_1(X1)|~p100(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_37,negated_conjecture,(epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_34, c_0_7])).
cnf(c_0_38,negated_conjecture,(p100(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_39,negated_conjecture,(~p101(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_40,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X13)|~r1(X13,X1)), inference(spm,[status(thm)],[c_0_35, c_0_7])).
cnf(c_0_41,plain,(r1(esk1_0,esk40_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_42,plain,(p101(X1)|p101(esk40_1(X1))|~epred1_1(X1)|~p100(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X12)|~r1(X12,X1)), inference(spm,[status(thm)],[c_0_40, c_0_7])).
cnf(c_0_44,plain,(p102(X1)|r1(X1,esk38_1(X1))|~epred1_1(X1)|~p101(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_45,negated_conjecture,(epred1_1(esk40_1(esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_41])).
cnf(c_0_46,plain,(p101(esk40_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_47,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X11)|~r1(X11,X1)), inference(spm,[status(thm)],[c_0_43, c_0_7])).
cnf(c_0_48,negated_conjecture,(epred1_1(X1)|~r1(esk40_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_31, c_0_41])).
cnf(c_0_49,plain,(p102(esk40_1(esk1_0))|r1(esk40_1(esk1_0),esk38_1(esk40_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])])).
cnf(c_0_50,plain,(p102(X1)|p102(esk38_1(X1))|~epred1_1(X1)|~p101(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_51,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X10)|~r1(X10,X1)), inference(spm,[status(thm)],[c_0_47, c_0_7])).
cnf(c_0_52,negated_conjecture,(epred1_1(X1)|~r1(esk40_1(esk1_0),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_28, c_0_41])).
cnf(c_0_53,plain,(p103(X1)|r1(X1,esk36_1(X1))|~epred1_1(X1)|~p102(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_54,negated_conjecture,(epred1_1(esk38_1(esk40_1(esk1_0)))|p102(esk40_1(esk1_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_55,plain,(p102(esk38_1(esk40_1(esk1_0)))|p102(esk40_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_45]), c_0_46])])).
cnf(c_0_56,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X9)|~r1(X9,X1)), inference(spm,[status(thm)],[c_0_51, c_0_7])).
cnf(c_0_57,negated_conjecture,(epred1_1(X1)|~r1(esk40_1(esk1_0),X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_26, c_0_41])).
cnf(c_0_58,negated_conjecture,(epred1_1(X1)|p102(esk40_1(esk1_0))|~r1(esk38_1(esk40_1(esk1_0)),X1)), inference(spm,[status(thm)],[c_0_52, c_0_49])).
cnf(c_0_59,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|r1(esk38_1(esk40_1(esk1_0)),esk36_1(esk38_1(esk40_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).
cnf(c_0_60,plain,(p103(X1)|p103(esk36_1(X1))|~epred1_1(X1)|~p102(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_61,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X8)|~r1(X8,X1)), inference(spm,[status(thm)],[c_0_56, c_0_7])).
cnf(c_0_62,negated_conjecture,(epred1_1(X1)|~r1(esk40_1(esk1_0),X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X1)), inference(spm,[status(thm)],[c_0_24, c_0_41])).
cnf(c_0_63,plain,(epred1_1(X1)|p102(esk40_1(esk1_0))|~r1(esk38_1(esk40_1(esk1_0)),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_57, c_0_49])).
cnf(c_0_64,plain,(p104(X1)|r1(X1,esk34_1(X1))|~epred1_1(X1)|~p103(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_65,negated_conjecture,(epred1_1(esk36_1(esk38_1(esk40_1(esk1_0))))|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_66,plain,(p102(esk40_1(esk1_0))|p103(esk36_1(esk38_1(esk40_1(esk1_0))))|p103(esk38_1(esk40_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_54]), c_0_55])).
cnf(c_0_67,negated_conjecture,(p8(X1)|~r1(esk1_0,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X7)|~r1(X7,X1)), inference(spm,[status(thm)],[c_0_61, c_0_7])).
cnf(c_0_68,negated_conjecture,(epred1_1(X1)|~r1(esk40_1(esk1_0),X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X1)), inference(spm,[status(thm)],[c_0_22, c_0_41])).
cnf(c_0_69,plain,(epred1_1(X1)|p102(esk40_1(esk1_0))|~r1(esk38_1(esk40_1(esk1_0)),X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_62, c_0_49])).
cnf(c_0_70,plain,(epred1_1(X1)|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|~r1(esk36_1(esk38_1(esk40_1(esk1_0))),X1)), inference(spm,[status(thm)],[c_0_63, c_0_59])).
cnf(c_0_71,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|r1(esk36_1(esk38_1(esk40_1(esk1_0))),esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])).
cnf(c_0_72,plain,(p104(X1)|p104(esk34_1(X1))|~epred1_1(X1)|~p103(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_73,negated_conjecture,(p8(X1)|~r1(esk40_1(esk1_0),X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X6)|~r1(X6,X1)), inference(spm,[status(thm)],[c_0_67, c_0_41])).
cnf(c_0_74,plain,(epred1_1(X1)|p102(esk40_1(esk1_0))|~r1(esk38_1(esk40_1(esk1_0)),X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X1)), inference(spm,[status(thm)],[c_0_68, c_0_49])).
cnf(c_0_75,plain,(epred1_1(X1)|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|~r1(esk36_1(esk38_1(esk40_1(esk1_0))),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_69, c_0_59])).
cnf(c_0_76,plain,(p105(X1)|r1(X1,esk32_1(X1))|~epred1_1(X1)|~p104(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_77,plain,(epred1_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_78,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_65]), c_0_66])).
cnf(c_0_79,plain,(p102(esk40_1(esk1_0))|p8(X1)|~r1(esk38_1(esk40_1(esk1_0)),X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)|~r1(X5,X1)), inference(spm,[status(thm)],[c_0_73, c_0_49])).
cnf(c_0_80,plain,(epred1_1(X1)|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|~r1(esk36_1(esk38_1(esk40_1(esk1_0))),X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_74, c_0_59])).
cnf(c_0_81,plain,(epred1_1(X1)|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|~r1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))),X1)), inference(spm,[status(thm)],[c_0_75, c_0_71])).
cnf(c_0_82,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|r1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))),esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])).
cnf(c_0_83,plain,(p105(X1)|p105(esk32_1(X1))|~epred1_1(X1)|~p104(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_84,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p8(X1)|~r1(esk36_1(esk38_1(esk40_1(esk1_0))),X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X1)), inference(spm,[status(thm)],[c_0_79, c_0_59])).
cnf(c_0_85,plain,(epred1_1(X1)|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|~r1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_80, c_0_71])).
cnf(c_0_86,plain,(p106(X1)|r1(X1,esk30_1(X1))|~epred1_1(X1)|~p105(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_87,plain,(epred1_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_88,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_77]), c_0_78])).
cnf(c_0_89,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p8(X1)|~r1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))),X2)|~r1(X2,X3)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_84, c_0_71])).
cnf(c_0_90,plain,(epred1_1(X1)|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|~r1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))),X1)), inference(spm,[status(thm)],[c_0_85, c_0_82])).
cnf(c_0_91,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p106(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))|r1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))),esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88])).
cnf(c_0_92,plain,(p106(X1)|p106(esk30_1(X1))|~epred1_1(X1)|~p105(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_93,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p8(X1)|~r1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))),X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_89, c_0_82])).
cnf(c_0_94,plain,(p107(X1)|r1(X1,esk28_1(X1))|~epred1_1(X1)|~p106(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_95,plain,(epred1_1(esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))))|p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p106(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_96,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p106(esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))))|p106(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_87]), c_0_88])).
cnf(c_0_97,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p106(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))|p8(X1)|~r1(esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))),X1)), inference(spm,[status(thm)],[c_0_93, c_0_91])).
cnf(c_0_98,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p106(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))|p107(esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))))|r1(esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))),esk28_1(esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])).
cnf(c_0_99,plain,(p107(X1)|~epred1_1(X1)|~p106(X1)|~p8(esk28_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_100,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p106(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))|p107(esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))))|p8(esk28_1(esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_101,plain,(p106(X1)|~epred1_1(X1)|~p105(X1)|~p107(esk30_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_102,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p106(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))|p107(esk30_1(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_96]), c_0_95])).
cnf(c_0_103,plain,(p105(X1)|~epred1_1(X1)|~p104(X1)|~p106(esk32_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_104,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))|p106(esk32_1(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0))))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_88]), c_0_87])).
cnf(c_0_105,plain,(p104(X1)|~epred1_1(X1)|~p103(X1)|~p105(esk34_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_106,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))|p105(esk34_1(esk36_1(esk38_1(esk40_1(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_78]), c_0_77])).
cnf(c_0_107,plain,(p103(X1)|~epred1_1(X1)|~p102(X1)|~p104(esk36_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_108,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))|p104(esk36_1(esk38_1(esk40_1(esk1_0))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_66]), c_0_65])).
cnf(c_0_109,plain,(p102(X1)|~epred1_1(X1)|~p101(X1)|~p103(esk38_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_110,plain,(p102(esk40_1(esk1_0))|p103(esk38_1(esk40_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_55]), c_0_54])).
cnf(c_0_111,plain,(p101(X1)|~epred1_1(X1)|~p100(X1)|~p102(esk40_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_112,plain,(p102(esk40_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_45]), c_0_46])])).
cnf(c_0_113,plain,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_37]), c_0_38])]), c_0_39]), ['proof']).
# SZS output end CNFRefutation.
EOF
