% Problem    : LCL660+1.005 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n098.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:56:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____102_C18_F1_PI_AE_Q4_CS_SP_S1S
# and selection function SelectComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (~(?[X1]:~((((((((~(![X2]:(~(r1(X1,X2))|~(p5(X2))))&(![X2]:(~(r1(X1,X2))|p2(X2))|~(![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1))))))))|![X2]:(~(r1(X1,X2))|p3(X2)))|~(![X2]:((~(r1(X1,X2))|p3(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p3(X2)))|~(p3(X1)))))))|(~(![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p5(X1))))))&(![X2]:(~(r1(X1,X2))|p2(X2))|~(![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1)))))))))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(![X2]:((~(r1(X1,X2))|p1(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2)))|~(p1(X1)))))))|~(((((((((![X2]:(~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))|~(![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1)))))))&(p2(X1)|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))|~(![X2]:((~(r1(X1,X2))|((![X1]:(~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1))))))|~(![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2)))))))&(p2(X2)|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1)))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|((![X1]:(~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1))))))|~(![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2)))))))&(p2(X2)|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1))))))))|~(((![X2]:(~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))|~(![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p2(X2)))|~(p2(X1)))))))&(p2(X1)|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))))))))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))|~(![X2]:((((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))|~(![X1]:((~(r1(X2,X1))|![X2]:(((((~(r1(X1,X2))|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))|~(![X2]:(((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))|~(![X1]:((~(r1(X2,X1))|![X2]:((((~(r1(X1,X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))&(((p2(X1)|p1(X1))|![X2]:(~(r1(X1,X2))|~$true))|~(![X2]:((((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))|~(![X1]:((~(r1(X2,X1))|![X2]:(((~(r1(X1,X2))|p2(X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~(((p2(X1)|p1(X1))|![X2]:(~(r1(X1,X2))|~$true)))))))))&((p1(X1)|![X2]:(~(r1(X1,X2))|~$true))|~(![X2]:(((~(r1(X1,X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true))|~(![X1]:((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|p1(X2))|![X1]:(~(r1(X2,X1))|~$true)))|~((p1(X1)|![X2]:(~(r1(X1,X2))|~$true)))))))))&![X2]:((~(r1(X1,X2))|p2(X2))|~(![X1]:((~(r1(X2,X1))|p2(X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p2(X1)))|~(p2(X2))))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', main)).
fof(c_0_1, axiom, (![X1]:r1(X1,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', reflexivity)).
fof(c_0_2, plain, (![X1]:(epred1_1(X1)<=>~(![X2]:((~r1(X1,X2)|((![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))&(p2(X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|((![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))&(p2(X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))))|~(((![X2]:(~r1(X1,X2)|![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))|~(![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))&(p2(X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))))))))), introduced(definition)).
fof(c_0_3, negated_conjecture, (~(~(?[X1]:~((((((((~(![X2]:(~r1(X1,X2)|~p5(X2)))&(![X2]:(~r1(X1,X2)|p2(X2))|~(![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))))|![X2]:(~r1(X1,X2)|p3(X2)))|~(![X2]:((~r1(X1,X2)|p3(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p3(X1))))))|(~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p5(X1)))))&(![X2]:(~r1(X1,X2)|p2(X2))|~(![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))))|![X2]:(~r1(X1,X2)|p1(X2)))|~(![X2]:((~r1(X1,X2)|p1(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2)))|~p1(X1))))))|~(((((((((![X2]:(~r1(X1,X2)|![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))|~(![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))&(p2(X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))|epred1_1(X1))&(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))|~(![X2]:((((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))|~(![X1]:((~r1(X2,X1)|![X2]:(((((~r1(X1,X2)|p4(X2))|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~(((((p4(X1)|p3(X1))|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))))))))&((((p3(X1)|p2(X1))|p1(X1))|![X2]:~r1(X1,X2))|~(![X2]:(((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))|~(![X1]:((~r1(X2,X1)|![X2]:((((~r1(X1,X2)|p3(X2))|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~((((p3(X1)|p2(X1))|p1(X1))|![X2]:~r1(X1,X2)))))))))&(((p2(X1)|p1(X1))|![X2]:~r1(X1,X2))|~(![X2]:((((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:~r1(X2,X1))|~(![X1]:((~r1(X2,X1)|![X2]:(((~r1(X1,X2)|p2(X2))|p1(X2))|![X1]:~r1(X2,X1)))|~(((p2(X1)|p1(X1))|![X2]:~r1(X1,X2)))))))))&((p1(X1)|![X2]:~r1(X1,X2))|~(![X2]:(((~r1(X1,X2)|p1(X2))|![X1]:~r1(X2,X1))|~(![X1]:((~r1(X2,X1)|![X2]:((~r1(X1,X2)|p1(X2))|![X1]:~r1(X2,X1)))|~((p1(X1)|![X2]:~r1(X1,X2)))))))))&![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_2])).
fof(c_0_4, negated_conjecture, (![X4]:![X6]:![X10]:![X13]:![X16]:![X20]:![X23]:![X24]:![X28]:![X29]:![X32]:![X35]:![X36]:![X37]:![X39]:![X42]:![X43]:![X44]:![X46]:![X49]:![X50]:![X51]:![X53]:![X56]:![X57]:![X58]:![X60]:![X62]:![X63]:(((((((((r1(esk1_0,esk2_0)|(~r1(esk1_0,X4)|~p5(X4)))&(~p2(esk2_0)|(~r1(esk1_0,X4)|~p5(X4))))&((((r1(X6,esk3_1(X6))|(~r1(esk1_0,X6)|p2(X6)))|(~r1(esk1_0,X4)|~p5(X4)))&(((r1(esk3_1(X6),esk4_1(X6))|(~r1(esk1_0,X6)|p2(X6)))|(~r1(esk1_0,X4)|~p5(X4)))&((~p2(esk4_1(X6))|(~r1(esk1_0,X6)|p2(X6)))|(~r1(esk1_0,X4)|~p5(X4)))))&((p2(esk3_1(X6))|(~r1(esk1_0,X6)|p2(X6)))|(~r1(esk1_0,X4)|~p5(X4)))))&(r1(esk1_0,esk5_0)&~p3(esk5_0)))&(((r1(X10,esk6_1(X10))|(~r1(esk1_0,X10)|p3(X10)))&((r1(esk6_1(X10),esk7_1(X10))|(~r1(esk1_0,X10)|p3(X10)))&(~p3(esk7_1(X10))|(~r1(esk1_0,X10)|p3(X10)))))&(p3(esk6_1(X10))|(~r1(esk1_0,X10)|p3(X10)))))&((((r1(esk1_0,esk9_0)|(r1(X13,esk8_1(X13))|~r1(esk1_0,X13)))&(~p2(esk9_0)|(r1(X13,esk8_1(X13))|~r1(esk1_0,X13))))&((((r1(X16,esk10_1(X16))|(~r1(esk1_0,X16)|p2(X16)))|(r1(X13,esk8_1(X13))|~r1(esk1_0,X13)))&(((r1(esk10_1(X16),esk11_1(X16))|(~r1(esk1_0,X16)|p2(X16)))|(r1(X13,esk8_1(X13))|~r1(esk1_0,X13)))&((~p2(esk11_1(X16))|(~r1(esk1_0,X16)|p2(X16)))|(r1(X13,esk8_1(X13))|~r1(esk1_0,X13)))))&((p2(esk10_1(X16))|(~r1(esk1_0,X16)|p2(X16)))|(r1(X13,esk8_1(X13))|~r1(esk1_0,X13)))))&(((r1(esk1_0,esk9_0)|(p5(esk8_1(X13))|~r1(esk1_0,X13)))&(~p2(esk9_0)|(p5(esk8_1(X13))|~r1(esk1_0,X13))))&((((r1(X16,esk10_1(X16))|(~r1(esk1_0,X16)|p2(X16)))|(p5(esk8_1(X13))|~r1(esk1_0,X13)))&(((r1(esk10_1(X16),esk11_1(X16))|(~r1(esk1_0,X16)|p2(X16)))|(p5(esk8_1(X13))|~r1(esk1_0,X13)))&((~p2(esk11_1(X16))|(~r1(esk1_0,X16)|p2(X16)))|(p5(esk8_1(X13))|~r1(esk1_0,X13)))))&((p2(esk10_1(X16))|(~r1(esk1_0,X16)|p2(X16)))|(p5(esk8_1(X13))|~r1(esk1_0,X13)))))))&(r1(esk1_0,esk12_0)&~p1(esk12_0)))&(((r1(X20,esk13_1(X20))|(~r1(esk1_0,X20)|p1(X20)))&((r1(esk13_1(X20),esk14_1(X20))|(~r1(esk1_0,X20)|p1(X20)))&(~p1(esk14_1(X20))|(~r1(esk1_0,X20)|p1(X20)))))&(p1(esk13_1(X20))|(~r1(esk1_0,X20)|p1(X20)))))&((((((((((((r1(esk1_0,esk17_0)|((r1(X24,esk15_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0))&((~p2(esk17_0)|((r1(X24,esk15_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0)))&((((~r1(esk17_0,X28)|(~r1(X28,X29)|p2(X29)))|~p2(X28))|((r1(X24,esk15_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0)))&(((((r1(esk1_0,esk17_0)|((r1(esk15_2(X23,X24),esk16_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0))&((~p2(esk17_0)|((r1(esk15_2(X23,X24),esk16_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0)))&((((~r1(esk17_0,X28)|(~r1(X28,X29)|p2(X29)))|~p2(X28))|((r1(esk15_2(X23,X24),esk16_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0)))&((((r1(esk1_0,esk17_0)|((~p2(esk16_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0))&((~p2(esk17_0)|((~p2(esk16_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0)))&((((~r1(esk17_0,X28)|(~r1(X28,X29)|p2(X29)))|~p2(X28))|((~p2(esk16_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0)))))&((((r1(esk1_0,esk17_0)|((p2(esk15_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0))&((~p2(esk17_0)|((p2(esk15_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0)))&((((~r1(esk17_0,X28)|(~r1(X28,X29)|p2(X29)))|~p2(X28))|((p2(esk15_2(X23,X24))|(~r1(X23,X24)|p2(X24)))|~r1(esk1_0,X23)))|epred1_1(esk1_0))))&((((r1(esk1_0,esk18_0)|p2(esk1_0))|epred1_1(esk1_0))&(((r1(esk18_0,esk19_0)|p2(esk1_0))|epred1_1(esk1_0))&((~p2(esk19_0)|p2(esk1_0))|epred1_1(esk1_0))))&((p2(esk18_0)|p2(esk1_0))|epred1_1(esk1_0))))&(((((((r1(esk1_0,esk20_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32)))&(~p4(esk20_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32))))&(~p3(esk20_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32))))&(~p2(esk20_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32))))&(~p1(esk20_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32))))&(r1(esk20_0,esk21_0)|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32))))&((((((~p4(X35)|(~r1(esk20_0,X35)|(((((~r1(X35,X36)|p4(X36))|p3(X36))|p2(X36))|p1(X36))|~r1(X36,X37))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32)))&((~p3(X35)|(~r1(esk20_0,X35)|(((((~r1(X35,X36)|p4(X36))|p3(X36))|p2(X36))|p1(X36))|~r1(X36,X37))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32))))&((~p2(X35)|(~r1(esk20_0,X35)|(((((~r1(X35,X36)|p4(X36))|p3(X36))|p2(X36))|p1(X36))|~r1(X36,X37))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32))))&((~p1(X35)|(~r1(esk20_0,X35)|(((((~r1(X35,X36)|p4(X36))|p3(X36))|p2(X36))|p1(X36))|~r1(X36,X37))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32))))&((r1(X35,esk22_1(X35))|(~r1(esk20_0,X35)|(((((~r1(X35,X36)|p4(X36))|p3(X36))|p2(X36))|p1(X36))|~r1(X36,X37))))|((((p4(esk1_0)|p3(esk1_0))|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X32))))))&((((((r1(esk1_0,esk23_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X39)))&(~p3(esk23_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X39))))&(~p2(esk23_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X39))))&(~p1(esk23_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X39))))&(r1(esk23_0,esk24_0)|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X39))))&(((((~p3(X42)|(~r1(esk23_0,X42)|((((~r1(X42,X43)|p3(X43))|p2(X43))|p1(X43))|~r1(X43,X44))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X39)))&((~p2(X42)|(~r1(esk23_0,X42)|((((~r1(X42,X43)|p3(X43))|p2(X43))|p1(X43))|~r1(X43,X44))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X39))))&((~p1(X42)|(~r1(esk23_0,X42)|((((~r1(X42,X43)|p3(X43))|p2(X43))|p1(X43))|~r1(X43,X44))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X39))))&((r1(X42,esk25_1(X42))|(~r1(esk23_0,X42)|((((~r1(X42,X43)|p3(X43))|p2(X43))|p1(X43))|~r1(X43,X44))))|(((p3(esk1_0)|p2(esk1_0))|p1(esk1_0))|~r1(esk1_0,X39))))))&(((((r1(esk1_0,esk26_0)|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X46)))&(~p2(esk26_0)|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X46))))&(~p1(esk26_0)|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X46))))&(r1(esk26_0,esk27_0)|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X46))))&((((~p2(X49)|(~r1(esk26_0,X49)|(((~r1(X49,X50)|p2(X50))|p1(X50))|~r1(X50,X51))))|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X46)))&((~p1(X49)|(~r1(esk26_0,X49)|(((~r1(X49,X50)|p2(X50))|p1(X50))|~r1(X50,X51))))|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X46))))&((r1(X49,esk28_1(X49))|(~r1(esk26_0,X49)|(((~r1(X49,X50)|p2(X50))|p1(X50))|~r1(X50,X51))))|((p2(esk1_0)|p1(esk1_0))|~r1(esk1_0,X46))))))&((((r1(esk1_0,esk29_0)|(p1(esk1_0)|~r1(esk1_0,X53)))&(~p1(esk29_0)|(p1(esk1_0)|~r1(esk1_0,X53))))&(r1(esk29_0,esk30_0)|(p1(esk1_0)|~r1(esk1_0,X53))))&(((~p1(X56)|(~r1(esk29_0,X56)|((~r1(X56,X57)|p1(X57))|~r1(X57,X58))))|(p1(esk1_0)|~r1(esk1_0,X53)))&((r1(X56,esk31_1(X56))|(~r1(esk29_0,X56)|((~r1(X56,X57)|p1(X57))|~r1(X57,X58))))|(p1(esk1_0)|~r1(esk1_0,X53))))))&(((r1(X60,esk32_1(X60))|(~r1(esk1_0,X60)|p2(X60)))&(~p2(esk32_1(X60))|(~r1(esk1_0,X60)|p2(X60))))&(((~r1(esk32_1(X60),X62)|(~r1(X62,X63)|p2(X63)))|~p2(X62))|(~r1(esk1_0,X60)|p2(X60))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
fof(c_0_5, plain, (![X2]:r1(X2,X2)), inference(variable_rename,[status(thm)],[c_0_1])).
cnf(c_0_6,negated_conjecture,(r1(X1,esk8_1(X1))|r1(esk1_0,esk9_0)|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(r1(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(p5(esk8_1(X1))|r1(esk1_0,esk9_0)|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9,negated_conjecture,(r1(esk1_0,esk2_0)|~p5(X1)|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_11,negated_conjecture,(p5(esk8_1(esk1_0))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_8, c_0_7])).
cnf(c_0_12,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(X2,esk15_2(X1,X2))|r1(esk1_0,esk17_0)|~r1(esk1_0,X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13,negated_conjecture,(r1(esk1_0,esk9_0)|r1(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])).
cnf(c_0_14,negated_conjecture,(p2(X1)|r1(X1,esk32_1(X1))|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15,negated_conjecture,(p2(X1)|~r1(esk1_0,X1)|~p2(esk32_1(X1))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|p2(esk15_2(X1,X2))|r1(esk1_0,esk17_0)|~r1(esk1_0,X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk15_2(esk2_0,X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18,negated_conjecture,(p2(esk2_0)|r1(esk2_0,esk32_1(esk2_0))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_14, c_0_13])).
cnf(c_0_19,negated_conjecture,(p2(esk2_0)|r1(esk1_0,esk9_0)|~p2(esk32_1(esk2_0))), inference(spm,[status(thm)],[c_0_15, c_0_13])).
cnf(c_0_20,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk2_0,X1))|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_13])).
cnf(c_0_21,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(esk15_2(X1,X2),esk16_2(X1,X2))|r1(esk1_0,esk17_0)|~r1(esk1_0,X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22,negated_conjecture,(p2(X1)|p2(X3)|~r1(esk1_0,X1)|~p2(X2)|~r1(X2,X3)|~r1(esk32_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_23,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk32_1(esk2_0),esk15_2(esk2_0,esk32_1(esk2_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_24,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk2_0,esk32_1(esk2_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_18]), c_0_19])).
cnf(c_0_25,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk15_2(esk2_0,X1),esk16_2(esk2_0,X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_21, c_0_13])).
fof(c_0_26, plain, (![X1]:(epred1_1(X1)=>~(![X2]:((~r1(X1,X2)|((![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))&(p2(X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|((![X1]:(~r1(X2,X1)|![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))|~(![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2))))))&(p2(X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1)))))))|~(((![X2]:(~r1(X1,X2)|![X1]:((~r1(X2,X1)|p2(X1))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))|~(![X2]:((~r1(X1,X2)|p2(X2))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p2(X2)))|~p2(X1))))))&(p2(X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p2(X1)))|~p2(X2)))))))))))), inference(split_equiv,[status(thm)],[c_0_2])).
cnf(c_0_27,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk15_2(esk2_0,esk32_1(esk2_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_13]), c_0_24])).
cnf(c_0_28,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk15_2(esk2_0,esk32_1(esk2_0)),esk16_2(esk2_0,esk32_1(esk2_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_18]), c_0_19])).
fof(c_0_29, plain, (![X3]:![X7]:![X8]:![X9]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:![X21]:![X22]:![X27]:![X28]:![X29]:![X32]:![X33]:(((r1(X3,esk33_1(X3))|~epred1_1(X3))&(((((~p2(esk33_1(X3))|r1(esk33_1(X3),esk34_1(X3)))|~epred1_1(X3))&((((~r1(esk33_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|r1(esk33_1(X3),esk34_1(X3)))|~epred1_1(X3)))&(((((~p2(esk33_1(X3))|r1(esk34_1(X3),esk35_1(X3)))|~epred1_1(X3))&((((~r1(esk33_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|r1(esk34_1(X3),esk35_1(X3)))|~epred1_1(X3)))&(((~p2(esk33_1(X3))|~p2(esk35_1(X3)))|~epred1_1(X3))&((((~r1(esk33_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|~p2(esk35_1(X3)))|~epred1_1(X3))))&(((~p2(esk33_1(X3))|((~r1(esk35_1(X3),X7)|(~r1(X7,X8)|p2(X8)))|~p2(X7)))|~epred1_1(X3))&((((~r1(esk33_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|((~r1(esk35_1(X3),X7)|(~r1(X7,X8)|p2(X8)))|~p2(X7)))|~epred1_1(X3)))))&(((((~p2(esk33_1(X3))|(r1(X9,esk36_2(X3,X9))|(~r1(esk33_1(X3),X9)|p2(X9))))|~epred1_1(X3))&((((~r1(esk33_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|(r1(X9,esk36_2(X3,X9))|(~r1(esk33_1(X3),X9)|p2(X9))))|~epred1_1(X3)))&((((~p2(esk33_1(X3))|(r1(esk36_2(X3,X9),esk37_2(X3,X9))|(~r1(esk33_1(X3),X9)|p2(X9))))|~epred1_1(X3))&((((~r1(esk33_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|(r1(esk36_2(X3,X9),esk37_2(X3,X9))|(~r1(esk33_1(X3),X9)|p2(X9))))|~epred1_1(X3)))&(((~p2(esk33_1(X3))|(~p2(esk37_2(X3,X9))|(~r1(esk33_1(X3),X9)|p2(X9))))|~epred1_1(X3))&((((~r1(esk33_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|(~p2(esk37_2(X3,X9))|(~r1(esk33_1(X3),X9)|p2(X9))))|~epred1_1(X3)))))&(((~p2(esk33_1(X3))|(p2(esk36_2(X3,X9))|(~r1(esk33_1(X3),X9)|p2(X9))))|~epred1_1(X3))&((((~r1(esk33_1(X3),X12)|(~r1(X12,X13)|p2(X13)))|~p2(X12))|(p2(esk36_2(X3,X9))|(~r1(esk33_1(X3),X9)|p2(X9))))|~epred1_1(X3))))))&(((((((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(X17,esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))))&(((((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((r1(esk38_5(X3,X14,X15,X16,X17),esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))))&((((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((~p2(esk39_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))))))&((((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk40_3(X3,X14,X15))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((((~r1(esk40_3(X3,X14,X15),X21)|(~r1(X21,X22)|p2(X22)))|~p2(X21))|((p2(esk38_5(X3,X14,X15,X16,X17))|(~r1(X16,X17)|p2(X17)))|~r1(X15,X16)))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))))&((((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(X15,esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&(((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((r1(esk41_3(X3,X14,X15),esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((~p2(esk42_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))))&((((((~p2(X14)|r1(X14,esk43_2(X3,X14)))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(X14,esk43_2(X3,X14)))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((((~p2(X14)|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|r1(esk43_2(X3,X14),esk44_2(X3,X14)))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|~p2(esk44_2(X3,X14)))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|~p2(esk44_2(X3,X14)))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))&((((~p2(X14)|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|((~r1(esk44_2(X3,X14),X27)|(~r1(X27,X28)|p2(X28)))|~p2(X27)))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((((~p2(X14)|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(X29,esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&(((((~p2(X14)|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(r1(esk45_3(X3,X14,X29),esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))&((((~p2(X14)|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(~p2(esk46_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3)))))&((((~p2(X14)|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))&(((((~r1(X14,X32)|(~r1(X32,X33)|p2(X33)))|~p2(X32))|(p2(esk45_3(X3,X14,X29))|(~r1(X14,X29)|p2(X29))))|(((p2(esk41_3(X3,X14,X15))|p2(X15))|~r1(X14,X15))|~r1(esk33_1(X3),X14)))|~epred1_1(X3))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])])])).
cnf(c_0_30,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(esk1_0,esk17_0)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk16_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_31,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk2_0,esk32_1(esk2_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32,plain,(r1(X1,esk33_1(X1))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_18]), c_0_13]), c_0_19])).
cnf(c_0_34,negated_conjecture,(p2(X2)|r1(X2,esk3_1(X2))|~p5(X1)|~r1(esk1_0,X1)|~r1(esk1_0,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_35,plain,(p2(esk2_0)|r1(esk1_0,esk33_1(esk1_0))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36,negated_conjecture,(~p5(X1)|~r1(esk1_0,X1)|~p2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_37,negated_conjecture,(p2(X2)|p2(esk3_1(X2))|~p5(X1)|~r1(esk1_0,X1)|~r1(esk1_0,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_38,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_39,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_35]), c_0_36])).
cnf(c_0_40,plain,(p2(X2)|r1(X2,esk36_2(X1,X2))|p2(X4)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~p2(X3)|~r1(X3,X4)|~r1(esk33_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_10]), c_0_11])).
cnf(c_0_42,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_10]), c_0_11])).
cnf(c_0_43,plain,(p2(X2)|r1(X2,esk36_2(X1,X2))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~p2(esk33_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_44,negated_conjecture,(p2(X2)|r1(esk3_1(X2),esk4_1(X2))|~p5(X1)|~r1(esk1_0,X1)|~r1(esk1_0,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_45,plain,(p2(X2)|p2(esk36_2(X1,X2))|p2(X4)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~p2(X3)|~r1(X3,X4)|~r1(esk33_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_46,plain,(p2(X2)|p2(esk36_2(X1,X2))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~p2(esk33_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_47,plain,(p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_48,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_35]), c_0_36])).
cnf(c_0_49,plain,(p2(esk36_2(esk1_0,X1))|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_41]), c_0_42]), c_0_46])).
cnf(c_0_50,plain,(p2(X2)|r1(esk36_2(X1,X2),esk37_2(X1,X2))|p2(X4)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~p2(X3)|~r1(X3,X4)|~r1(esk33_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_51,plain,(p2(X2)|r1(esk36_2(X1,X2),esk37_2(X1,X2))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~p2(esk33_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_52,plain,(p2(X2)|p2(X4)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~p2(esk37_2(X1,X2))|~p2(X3)|~r1(X3,X4)|~r1(esk33_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_53,plain,(p2(X2)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~p2(esk37_2(X1,X2))|~p2(esk33_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_54,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_47, c_0_33])).
cnf(c_0_55,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_10]), c_0_11])).
cnf(c_0_56,negated_conjecture,(p2(esk36_2(esk1_0,X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_49, c_0_33])).
cnf(c_0_57,plain,(p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_41]), c_0_42]), c_0_51])).
cnf(c_0_58,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_41]), c_0_42]), c_0_53])).
cnf(c_0_59,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_60,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk33_1(esk1_0),esk32_1(esk33_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_14, c_0_35])).
cnf(c_0_61,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_15, c_0_35])).
cnf(c_0_62,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_56, c_0_55])).
cnf(c_0_63,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_57, c_0_33])).
cnf(c_0_64,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_58, c_0_33])).
cnf(c_0_65,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])).
cnf(c_0_66,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_60]), c_0_61])).
cnf(c_0_67,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_63, c_0_55])).
cnf(c_0_68,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_64, c_0_55])).
cnf(c_0_69,negated_conjecture,(p2(X2)|~p5(X1)|~r1(esk1_0,X1)|~r1(esk1_0,X2)|~p2(esk4_1(X2))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_70,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_65]), c_0_35]), c_0_66])).
cnf(c_0_71,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_60]), c_0_61])).
cnf(c_0_72,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_60]), c_0_61])).
cnf(c_0_73,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_35]), c_0_36])).
cnf(c_0_74,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])).
cnf(c_0_75,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_36])).
cnf(c_0_76,plain,(r1(esk33_1(X1),esk34_1(X1))|~epred1_1(X1)|~p2(esk33_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_77,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_10]), c_0_11])).
cnf(c_0_78,plain,(r1(esk33_1(esk1_0),esk34_1(esk1_0))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_79,plain,(p2(X3)|r1(X3,esk41_3(X1,X2,X3))|p2(X4)|p2(esk45_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_80,negated_conjecture,(p2(esk2_0)|r1(esk33_1(esk1_0),esk34_1(esk1_0))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_78, c_0_33])).
cnf(c_0_81,plain,(p2(X3)|r1(X3,esk41_3(X1,X2,X3))|p2(X4)|p2(esk45_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_82,plain,(p2(X3)|p2(esk41_3(X1,X2,X3))|p2(X4)|p2(esk45_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_83,plain,(p2(X3)|p2(esk41_3(X1,X2,X3))|p2(X4)|p2(esk45_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_84,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_33])).
cnf(c_0_85,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk34_1(esk1_0),esk36_2(esk1_0,esk34_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_80]), c_0_77]), c_0_33])).
cnf(c_0_86,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_80]), c_0_77]), c_0_33])).
cnf(c_0_87,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_80]), c_0_33])).
cnf(c_0_88,plain,(r1(esk34_1(X1),esk35_1(X1))|~epred1_1(X1)|~p2(esk33_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_89,plain,(~epred1_1(X1)|~p2(esk35_1(X1))|~p2(esk33_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_90,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X3)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_80]), c_0_33])).
cnf(c_0_91,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_80]), c_0_33])).
cnf(c_0_92,plain,(p2(X3)|r1(esk41_3(X1,X2,X3),esk42_3(X1,X2,X3))|p2(X4)|p2(esk45_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_93,plain,(p2(X3)|r1(esk41_3(X1,X2,X3),esk42_3(X1,X2,X3))|p2(X4)|p2(esk45_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_94,plain,(p2(X3)|r1(X3,esk41_3(X1,X2,X3))|p2(X4)|r1(X4,esk45_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_95,plain,(p2(X3)|r1(X3,esk41_3(X1,X2,X3))|p2(X4)|r1(X4,esk45_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_96,plain,(p2(X3)|p2(esk41_3(X1,X2,X3))|p2(X4)|r1(X4,esk45_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_97,plain,(p2(X3)|p2(esk41_3(X1,X2,X3))|p2(X4)|r1(X4,esk45_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_98,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86]), c_0_87])).
cnf(c_0_99,plain,(r1(esk34_1(esk1_0),esk35_1(esk1_0))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_88, c_0_77])).
cnf(c_0_100,plain,(r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~epred1_1(esk1_0)|~p2(esk35_1(esk1_0))), inference(spm,[status(thm)],[c_0_89, c_0_77])).
cnf(c_0_101,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_85]), c_0_86]), c_0_91])).
cnf(c_0_102,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_80]), c_0_33])).
cnf(c_0_103,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_80]), c_0_33])).
cnf(c_0_104,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_80]), c_0_33])).
cnf(c_0_105,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_80]), c_0_33])).
cnf(c_0_106,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk45_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_80]), c_0_33])).
cnf(c_0_107,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_80]), c_0_33])).
cnf(c_0_108,plain,(p2(X3)|r1(esk41_3(X1,X2,X3),esk42_3(X1,X2,X3))|p2(X4)|r1(X4,esk45_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_109,plain,(p2(X3)|r1(esk41_3(X1,X2,X3),esk42_3(X1,X2,X3))|p2(X4)|r1(X4,esk45_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_110,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_98, c_0_7])).
cnf(c_0_111,negated_conjecture,(p2(esk2_0)|r1(esk34_1(esk1_0),esk35_1(esk1_0))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_99, c_0_33])).
cnf(c_0_112,negated_conjecture,(p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk35_1(esk1_0))), inference(spm,[status(thm)],[c_0_100, c_0_33])).
cnf(c_0_113,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_101, c_0_7])).
cnf(c_0_114,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_85]), c_0_86]), c_0_103])).
cnf(c_0_115,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_85]), c_0_86]), c_0_105])).
cnf(c_0_116,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_85]), c_0_86]), c_0_107])).
cnf(c_0_117,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_80]), c_0_33])).
cnf(c_0_118,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_80]), c_0_33])).
cnf(c_0_119,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112])).
cnf(c_0_120,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_111]), c_0_112])).
cnf(c_0_121,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_114, c_0_7])).
cnf(c_0_122,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_115, c_0_7])).
cnf(c_0_123,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_116, c_0_7])).
cnf(c_0_124,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_85]), c_0_86]), c_0_118])).
cnf(c_0_125,plain,(p2(X3)|~epred1_1(X1)|~p2(X2)|~r1(X2,X3)|~r1(esk35_1(X1),X2)|~p2(esk33_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_126,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_111]), c_0_112])).
cnf(c_0_127,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_111]), c_0_112])).
cnf(c_0_128,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_111]), c_0_112])).
cnf(c_0_129,plain,(p2(X3)|r1(X3,esk41_3(X1,X2,X3))|p2(X4)|r1(esk45_3(X1,X2,X4),esk46_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_130,plain,(p2(X3)|r1(X3,esk41_3(X1,X2,X3))|p2(X4)|r1(esk45_3(X1,X2,X4),esk46_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_131,plain,(p2(X3)|p2(esk41_3(X1,X2,X3))|p2(X4)|r1(esk45_3(X1,X2,X4),esk46_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_132,plain,(p2(X3)|p2(esk41_3(X1,X2,X3))|p2(X4)|r1(esk45_3(X1,X2,X4),esk46_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_133,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_111]), c_0_112])).
cnf(c_0_134,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_111]), c_0_112])).
cnf(c_0_135,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_124, c_0_7])).
cnf(c_0_136,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_127]), c_0_77]), c_0_33])).
cnf(c_0_137,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_111]), c_0_112])).
cnf(c_0_138,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_80]), c_0_33])).
cnf(c_0_139,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_80]), c_0_33])).
cnf(c_0_140,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X3)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X3),esk46_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_80]), c_0_33])).
cnf(c_0_141,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_80]), c_0_33])).
cnf(c_0_142,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_111]), c_0_112])).
cnf(c_0_143,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_111]), c_0_112])).
cnf(c_0_144,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_111]), c_0_112])).
cnf(c_0_145,plain,(p2(X3)|p2(X4)|p2(esk45_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~p2(esk42_3(X1,X2,X3))|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_146,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_136, c_0_137])).
cnf(c_0_147,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_85]), c_0_86]), c_0_139])).
cnf(c_0_148,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_85]), c_0_86]), c_0_141])).
cnf(c_0_149,plain,(p2(X3)|r1(esk41_3(X1,X2,X3),esk42_3(X1,X2,X3))|p2(X4)|r1(esk45_3(X1,X2,X4),esk46_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_150,plain,(p2(X3)|r1(esk41_3(X1,X2,X3),esk42_3(X1,X2,X3))|p2(X4)|r1(esk45_3(X1,X2,X4),esk46_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_151,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_142]), c_0_143]), c_0_77]), c_0_33])).
cnf(c_0_152,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_111]), c_0_112])).
cnf(c_0_153,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_111]), c_0_80]), c_0_86]), c_0_33]), c_0_112])).
cnf(c_0_154,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_147, c_0_7])).
cnf(c_0_155,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_148, c_0_7])).
cnf(c_0_156,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_80]), c_0_33])).
cnf(c_0_157,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_80]), c_0_33])).
cnf(c_0_158,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_159,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_111]), c_0_112])).
cnf(c_0_160,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_111]), c_0_112])).
cnf(c_0_161,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_111]), c_0_112])).
cnf(c_0_162,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_85]), c_0_86]), c_0_157])).
cnf(c_0_163,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_158]), c_0_159]), c_0_77]), c_0_33])).
cnf(c_0_164,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_111]), c_0_112])).
cnf(c_0_165,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_143]), c_0_127]), c_0_77]), c_0_33])).
cnf(c_0_166,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_111]), c_0_112])).
cnf(c_0_167,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_162, c_0_7])).
cnf(c_0_168,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_163, c_0_164])).
cnf(c_0_169,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_170,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_111]), c_0_112])).
cnf(c_0_171,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk36_2(esk1_0,esk34_1(esk1_0)),esk37_2(esk1_0,esk34_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_80]), c_0_77]), c_0_33])).
cnf(c_0_172,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_168]), c_0_169]), c_0_77]), c_0_33])).
cnf(c_0_173,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_111]), c_0_112])).
cnf(c_0_174,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_171]), c_0_105])).
cnf(c_0_175,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_171]), c_0_107])).
cnf(c_0_176,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_171]), c_0_87])).
cnf(c_0_177,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_171]), c_0_91])).
cnf(c_0_178,plain,(p2(X3)|p2(X4)|r1(X4,esk45_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~p2(esk42_3(X1,X2,X3))|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_179,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173]), c_0_163])).
cnf(c_0_180,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_111]), c_0_112])).
cnf(c_0_181,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_111]), c_0_112])).
cnf(c_0_182,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_171]), c_0_118])).
cnf(c_0_183,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_111]), c_0_112])).
cnf(c_0_184,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_111]), c_0_112])).
cnf(c_0_185,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_171]), c_0_103])).
cnf(c_0_186,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_179]), c_0_111]), c_0_80]), c_0_86]), c_0_33]), c_0_112])).
cnf(c_0_187,plain,(p2(X3)|p2(X4)|r1(esk45_3(X1,X2,X4),esk46_3(X1,X2,X4))|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~p2(esk42_3(X1,X2,X3))|~r1(X2,X4)|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_188,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_111]), c_0_112])).
cnf(c_0_189,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_111]), c_0_112])).
cnf(c_0_190,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_111]), c_0_112])).
cnf(c_0_191,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_111]), c_0_112])).
cnf(c_0_192,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_111]), c_0_112])).
cnf(c_0_193,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_111]), c_0_112])).
cnf(c_0_194,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_111]), c_0_112])).
cnf(c_0_195,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_179]), c_0_111]), c_0_80]), c_0_86]), c_0_33]), c_0_112])).
cnf(c_0_196,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_188]), c_0_189]), c_0_77]), c_0_33])).
cnf(c_0_197,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_111]), c_0_112])).
cnf(c_0_198,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_191]), c_0_192]), c_0_77]), c_0_33])).
cnf(c_0_199,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_111]), c_0_112])).
cnf(c_0_200,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_171]), c_0_139])).
cnf(c_0_201,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_171]), c_0_141])).
cnf(c_0_202,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_194]), c_0_159]), c_0_77]), c_0_33])).
cnf(c_0_203,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_111]), c_0_112])).
cnf(c_0_204,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_196, c_0_197])).
cnf(c_0_205,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_198, c_0_199])).
cnf(c_0_206,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_111]), c_0_112])).
cnf(c_0_207,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_111]), c_0_112])).
cnf(c_0_208,plain,(p2(X3)|r1(X3,esk41_3(X1,X2,X3))|p2(X4)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk46_3(X1,X2,X4))|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_209,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_202, c_0_203])).
cnf(c_0_210,plain,(p2(X3)|p2(esk41_3(X1,X2,X3))|p2(X4)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk46_3(X1,X2,X4))|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_211,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_204]), c_0_205]), c_0_77]), c_0_33])).
cnf(c_0_212,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_111]), c_0_112])).
cnf(c_0_213,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_189]), c_0_192]), c_0_77]), c_0_33])).
cnf(c_0_214,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_111]), c_0_112])).
cnf(c_0_215,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_171]), c_0_157])).
cnf(c_0_216,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209]), c_0_111]), c_0_80]), c_0_86]), c_0_33]), c_0_112])).
cnf(c_0_217,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_209]), c_0_111]), c_0_80]), c_0_86]), c_0_33]), c_0_112])).
cnf(c_0_218,plain,(p2(X3)|r1(esk41_3(X1,X2,X3),esk42_3(X1,X2,X3))|p2(X4)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk46_3(X1,X2,X4))|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_219,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_211, c_0_212])).
cnf(c_0_220,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_213, c_0_214])).
cnf(c_0_221,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_111]), c_0_112])).
cnf(c_0_222,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_111]), c_0_112])).
cnf(c_0_223,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_111]), c_0_112])).
cnf(c_0_224,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_209]), c_0_111]), c_0_80]), c_0_86]), c_0_33]), c_0_112])).
cnf(c_0_225,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_219]), c_0_220]), c_0_77]), c_0_33])).
cnf(c_0_226,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_111]), c_0_112])).
cnf(c_0_227,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_222]), c_0_223]), c_0_77]), c_0_33])).
cnf(c_0_228,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_111]), c_0_112])).
cnf(c_0_229,plain,(p2(X3)|p2(X4)|r1(X4,esk45_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~p2(esk42_3(X1,X2,X3))|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_230,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_226]), c_0_211])).
cnf(c_0_231,plain,(p2(X3)|p2(X4)|p2(esk45_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~p2(esk42_3(X1,X2,X3))|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_232,plain,(p2(X3)|p2(X4)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~p2(esk42_3(X1,X2,X3))|~r1(X2,X4)|~p2(esk46_3(X1,X2,X4))|~p2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_233,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_227, c_0_228])).
cnf(c_0_234,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_235,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_230]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_236,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_205]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_237,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_205]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_238,plain,(p2(X3)|p2(X4)|r1(esk45_3(X1,X2,X4),esk46_3(X1,X2,X4))|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~p2(esk42_3(X1,X2,X3))|~r1(X2,X4)|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_239,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233]), c_0_111]), c_0_80]), c_0_86]), c_0_33]), c_0_112])).
cnf(c_0_240,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_85]), c_0_86]), c_0_235])).
cnf(c_0_241,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_85]), c_0_86]), c_0_237])).
cnf(c_0_242,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_230]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_243,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_209]), c_0_111]), c_0_112])).
cnf(c_0_244,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_230]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_245,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_171]), c_0_235])).
cnf(c_0_246,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_171]), c_0_237])).
cnf(c_0_247,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_85]), c_0_243]), c_0_244])).
cnf(c_0_248,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_111]), c_0_112])).
cnf(c_0_249,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_111]), c_0_112])).
cnf(c_0_250,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_171]), c_0_244])).
cnf(c_0_251,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_248]), c_0_249]), c_0_77]), c_0_33])).
cnf(c_0_252,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_111]), c_0_112])).
cnf(c_0_253,plain,(p2(X3)|r1(X3,esk41_3(X1,X2,X3))|p2(X4)|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk46_3(X1,X2,X4))|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_254,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_251, c_0_252])).
cnf(c_0_255,plain,(p2(X3)|p2(esk41_3(X1,X2,X3))|p2(X4)|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk46_3(X1,X2,X4))|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_256,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_257,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_254]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_258,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_254]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_259,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_254]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_260,plain,(p2(X3)|r1(esk41_3(X1,X2,X3),esk42_3(X1,X2,X3))|p2(X4)|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(esk46_3(X1,X2,X4))|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_261,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_85]), c_0_243]), c_0_257])).
cnf(c_0_262,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_85]), c_0_243]), c_0_259])).
cnf(c_0_263,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_254]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_264,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_254]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_265,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_171]), c_0_257])).
cnf(c_0_266,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_171]), c_0_259])).
cnf(c_0_267,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_85]), c_0_243]), c_0_264])).
cnf(c_0_268,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_111]), c_0_112])).
cnf(c_0_269,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_111]), c_0_112])).
cnf(c_0_270,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_171]), c_0_264])).
cnf(c_0_271,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_268]), c_0_269]), c_0_77]), c_0_33])).
cnf(c_0_272,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_111]), c_0_112])).
cnf(c_0_273,plain,(p2(X3)|p2(X4)|p2(X6)|~epred1_1(X1)|~r1(esk33_1(X1),X2)|~r1(X2,X3)|~p2(esk42_3(X1,X2,X3))|~r1(X2,X4)|~p2(esk46_3(X1,X2,X4))|~p2(X5)|~r1(X5,X6)|~r1(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_274,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_271, c_0_272])).
cnf(c_0_275,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_274]), c_0_111]), c_0_80]), c_0_33]), c_0_112])).
cnf(c_0_276,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(X2)|~r1(esk34_1(esk1_0),X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_254]), c_0_111]), c_0_112])).
cnf(c_0_277,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_80]), c_0_77]), c_0_33])).
cnf(c_0_278,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_85]), c_0_243]), c_0_277])).
cnf(c_0_279,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_171]), c_0_277])).
cnf(c_0_280,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_105, c_0_279])).
cnf(c_0_281,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_107, c_0_279])).
cnf(c_0_282,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_87, c_0_279])).
cnf(c_0_283,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_91, c_0_279])).
cnf(c_0_284,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_111]), c_0_112])).
cnf(c_0_285,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_111]), c_0_112])).
cnf(c_0_286,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_118, c_0_279])).
cnf(c_0_287,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_111]), c_0_112])).
cnf(c_0_288,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_111]), c_0_112])).
cnf(c_0_289,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_103, c_0_279])).
cnf(c_0_290,negated_conjecture,(p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_111]), c_0_112])).
cnf(c_0_291,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_111]), c_0_112])).
cnf(c_0_292,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_111]), c_0_112])).
cnf(c_0_293,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_111]), c_0_112])).
cnf(c_0_294,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_111]), c_0_112])).
cnf(c_0_295,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_111]), c_0_112])).
cnf(c_0_296,plain,(p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_290]), c_0_291]), c_0_77]), c_0_33])).
cnf(c_0_297,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_111]), c_0_112])).
cnf(c_0_298,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_293]), c_0_294]), c_0_77]), c_0_33])).
cnf(c_0_299,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_111]), c_0_112])).
cnf(c_0_300,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_139, c_0_279])).
cnf(c_0_301,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_141, c_0_279])).
cnf(c_0_302,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_296, c_0_297])).
cnf(c_0_303,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_298, c_0_299])).
cnf(c_0_304,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_111]), c_0_112])).
cnf(c_0_305,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_111]), c_0_112])).
cnf(c_0_306,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_302]), c_0_303]), c_0_77]), c_0_33])).
cnf(c_0_307,negated_conjecture,(p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_111]), c_0_112])).
cnf(c_0_308,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_291]), c_0_294]), c_0_77]), c_0_33])).
cnf(c_0_309,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_111]), c_0_112])).
cnf(c_0_310,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_157, c_0_279])).
cnf(c_0_311,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_306, c_0_307])).
cnf(c_0_312,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_308, c_0_309])).
cnf(c_0_313,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_111]), c_0_112])).
cnf(c_0_314,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_311]), c_0_312]), c_0_77]), c_0_33])).
cnf(c_0_315,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_111]), c_0_112])).
cnf(c_0_316,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_315]), c_0_306])).
cnf(c_0_317,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_316]), c_0_111]), c_0_80]), c_0_279]), c_0_33]), c_0_112])).
cnf(c_0_318,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_303]), c_0_111]), c_0_80]), c_0_279]), c_0_33]), c_0_112])).
cnf(c_0_319,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_111]), c_0_112])).
cnf(c_0_320,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_111]), c_0_112])).
cnf(c_0_321,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_316]), c_0_111]), c_0_80]), c_0_279]), c_0_33]), c_0_112])).
cnf(c_0_322,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_319]), c_0_320]), c_0_77]), c_0_33])).
cnf(c_0_323,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_111]), c_0_112])).
cnf(c_0_324,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_322, c_0_323])).
cnf(c_0_325,plain,(p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_324]), c_0_111]), c_0_80]), c_0_279]), c_0_33]), c_0_112])).
cnf(c_0_326,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_324]), c_0_111]), c_0_80]), c_0_279]), c_0_33]), c_0_112])).
cnf(c_0_327,negated_conjecture,(p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_111]), c_0_112])).
cnf(c_0_328,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_111]), c_0_112])).
cnf(c_0_329,plain,(p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_324]), c_0_111]), c_0_80]), c_0_279]), c_0_33]), c_0_112])).
cnf(c_0_330,plain,(p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_327]), c_0_328]), c_0_77]), c_0_33])).
cnf(c_0_331,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_111]), c_0_112])).
cnf(c_0_332,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_330, c_0_331])).
cnf(c_0_333,plain,(p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_332]), c_0_111]), c_0_80]), c_0_279]), c_0_33]), c_0_112])).
cnf(c_0_334,plain,(p2(esk2_0)|r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_324]), c_0_111]), c_0_112])).
cnf(c_0_335,negated_conjecture,(r1(esk1_0,esk17_0)|r1(esk1_0,esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_36, c_0_334])).
cnf(c_0_336,negated_conjecture,(r1(esk1_0,esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_10]), c_0_11])).
cnf(c_0_337,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk15_2(esk9_0,X1))|r1(esk1_0,esk17_0)|~r1(esk9_0,X1)), inference(spm,[status(thm)],[c_0_12, c_0_336])).
cnf(c_0_338,negated_conjecture,(p2(esk9_0)|r1(esk9_0,esk32_1(esk9_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_14, c_0_336])).
cnf(c_0_339,negated_conjecture,(p2(esk9_0)|r1(esk1_0,esk17_0)|~p2(esk32_1(esk9_0))), inference(spm,[status(thm)],[c_0_15, c_0_336])).
cnf(c_0_340,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk9_0,X1))|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk9_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_336])).
cnf(c_0_341,negated_conjecture,(epred1_1(esk1_0)|p2(esk9_0)|r1(esk32_1(esk9_0),esk15_2(esk9_0,esk32_1(esk9_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_338]), c_0_339])).
cnf(c_0_342,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk9_0,esk32_1(esk9_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_338]), c_0_339])).
cnf(c_0_343,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk15_2(esk9_0,X1),esk16_2(esk9_0,X1))|r1(esk1_0,esk17_0)|~r1(esk9_0,X1)), inference(spm,[status(thm)],[c_0_21, c_0_336])).
cnf(c_0_344,negated_conjecture,(epred1_1(esk1_0)|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk15_2(esk9_0,esk32_1(esk9_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_341]), c_0_336]), c_0_342])).
cnf(c_0_345,negated_conjecture,(epred1_1(esk1_0)|p2(esk9_0)|r1(esk15_2(esk9_0,esk32_1(esk9_0)),esk16_2(esk9_0,esk32_1(esk9_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_338]), c_0_339])).
cnf(c_0_346,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk9_0,esk32_1(esk9_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_344, c_0_345])).
cnf(c_0_347,negated_conjecture,(epred1_1(esk1_0)|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_346]), c_0_14]), c_0_336]), c_0_339])).
cnf(c_0_348,negated_conjecture,(r1(X1,esk8_1(X1))|p2(X2)|r1(X2,esk10_1(X2))|~r1(esk1_0,X1)|~r1(esk1_0,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_349,plain,(p2(esk9_0)|r1(esk1_0,esk33_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_32, c_0_347])).
cnf(c_0_350,negated_conjecture,(r1(X1,esk8_1(X1))|~r1(esk1_0,X1)|~p2(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_351,negated_conjecture,(r1(X1,esk8_1(X1))|p2(X2)|p2(esk10_1(X2))|~r1(esk1_0,X1)|~r1(esk1_0,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_352,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_349]), c_0_350])).
cnf(c_0_353,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(esk1_0,esk17_0)|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_349]), c_0_350])).
cnf(c_0_354,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_352, c_0_7])).
cnf(c_0_355,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_353, c_0_7])).
cnf(c_0_356,negated_conjecture,(r1(X1,esk8_1(X1))|p2(X2)|r1(esk10_1(X2),esk11_1(X2))|~r1(esk1_0,X1)|~r1(esk1_0,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_357,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_354]), c_0_355]), c_0_43])).
cnf(c_0_358,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_349]), c_0_350])).
cnf(c_0_359,plain,(p2(esk36_2(esk1_0,X1))|p2(X2)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_354]), c_0_355]), c_0_46])).
cnf(c_0_360,negated_conjecture,(p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_357, c_0_347])).
cnf(c_0_361,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_358, c_0_7])).
cnf(c_0_362,negated_conjecture,(p2(esk36_2(esk1_0,X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_359, c_0_347])).
cnf(c_0_363,plain,(p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_354]), c_0_355]), c_0_51])).
cnf(c_0_364,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_354]), c_0_355]), c_0_53])).
cnf(c_0_365,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_360, c_0_361])).
cnf(c_0_366,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk32_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_14, c_0_349])).
cnf(c_0_367,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p2(esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_15, c_0_349])).
cnf(c_0_368,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_362, c_0_361])).
cnf(c_0_369,negated_conjecture,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_363, c_0_347])).
cnf(c_0_370,negated_conjecture,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_364, c_0_347])).
cnf(c_0_371,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_366]), c_0_367])).
cnf(c_0_372,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_366]), c_0_367])).
cnf(c_0_373,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_369, c_0_361])).
cnf(c_0_374,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_370, c_0_361])).
cnf(c_0_375,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_371]), c_0_349]), c_0_372])).
cnf(c_0_376,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_366]), c_0_367])).
cnf(c_0_377,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_366]), c_0_367])).
cnf(c_0_378,negated_conjecture,(r1(X1,esk8_1(X1))|p2(X2)|~r1(esk1_0,X1)|~r1(esk1_0,X2)|~p2(esk11_1(X2))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_379,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_375, c_0_376]), c_0_377])).
cnf(c_0_380,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_379]), c_0_349]), c_0_350])).
cnf(c_0_381,negated_conjecture,(p5(esk8_1(X1))|p2(X2)|r1(X2,esk10_1(X2))|~r1(esk1_0,X1)|~r1(esk1_0,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_382,negated_conjecture,(p5(esk8_1(X1))|~r1(esk1_0,X1)|~p2(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_383,negated_conjecture,(p5(esk8_1(X1))|p2(X2)|p2(esk10_1(X2))|~r1(esk1_0,X1)|~r1(esk1_0,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_384,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_34, c_0_349])).
cnf(c_0_385,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_380, c_0_7])).
cnf(c_0_386,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(X1))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_349]), c_0_382])).
cnf(c_0_387,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_37, c_0_349])).
cnf(c_0_388,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p5(esk8_1(X1))|r1(esk1_0,esk17_0)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_349]), c_0_382])).
cnf(c_0_389,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_384, c_0_385])).
cnf(c_0_390,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(esk1_0))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_386, c_0_7])).
cnf(c_0_391,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_387, c_0_385])).
cnf(c_0_392,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_349])).
cnf(c_0_393,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p5(esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_388, c_0_7])).
cnf(c_0_394,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_389, c_0_390])).
cnf(c_0_395,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_391, c_0_390])).
cnf(c_0_396,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_392, c_0_385])).
cnf(c_0_397,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_389, c_0_393])).
cnf(c_0_398,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_391, c_0_393])).
cnf(c_0_399,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_394]), c_0_395]), c_0_347])).
cnf(c_0_400,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_396, c_0_390])).
cnf(c_0_401,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_394]), c_0_395]), c_0_347])).
cnf(c_0_402,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_397]), c_0_398]), c_0_347])).
cnf(c_0_403,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_396, c_0_393])).
cnf(c_0_404,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_397]), c_0_398]), c_0_347])).
cnf(c_0_405,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_399, c_0_400])).
cnf(c_0_406,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_401, c_0_400])).
cnf(c_0_407,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_394]), c_0_395]), c_0_347])).
cnf(c_0_408,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_394]), c_0_395]), c_0_347])).
cnf(c_0_409,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_402, c_0_403])).
cnf(c_0_410,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_404, c_0_403])).
cnf(c_0_411,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_397]), c_0_398]), c_0_347])).
cnf(c_0_412,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_397]), c_0_398]), c_0_347])).
cnf(c_0_413,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_405, c_0_366]), c_0_367])).
cnf(c_0_414,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_366]), c_0_367])).
cnf(c_0_415,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_407, c_0_400])).
cnf(c_0_416,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_408, c_0_400])).
cnf(c_0_417,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_366]), c_0_367])).
cnf(c_0_418,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_366]), c_0_367])).
cnf(c_0_419,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_411, c_0_403])).
cnf(c_0_420,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_412, c_0_403])).
cnf(c_0_421,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_413]), c_0_349]), c_0_414])).
cnf(c_0_422,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_366]), c_0_367])).
cnf(c_0_423,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_366]), c_0_367])).
cnf(c_0_424,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_417]), c_0_349]), c_0_418])).
cnf(c_0_425,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_366]), c_0_367])).
cnf(c_0_426,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_366]), c_0_367])).
cnf(c_0_427,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p2(esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_69, c_0_349])).
cnf(c_0_428,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_422]), c_0_423])).
cnf(c_0_429,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_424, c_0_425]), c_0_426])).
cnf(c_0_430,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_427, c_0_428])).
cnf(c_0_431,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_427, c_0_429])).
cnf(c_0_432,negated_conjecture,(p5(esk8_1(X1))|p2(X2)|r1(esk10_1(X2),esk11_1(X2))|~r1(esk1_0,X1)|~r1(esk1_0,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_433,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_385]), c_0_390])).
cnf(c_0_434,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_385]), c_0_393])).
cnf(c_0_435,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(X1))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_432, c_0_349]), c_0_382])).
cnf(c_0_436,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_433]), c_0_434]), c_0_347])).
cnf(c_0_437,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(esk1_0))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_435, c_0_7])).
cnf(c_0_438,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_433]), c_0_434]), c_0_347])).
cnf(c_0_439,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_436, c_0_437])).
cnf(c_0_440,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_438, c_0_437])).
cnf(c_0_441,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_433]), c_0_434]), c_0_347])).
cnf(c_0_442,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_433]), c_0_434]), c_0_347])).
cnf(c_0_443,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_395]), c_0_398]), c_0_347])).
cnf(c_0_444,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_395]), c_0_398]), c_0_347])).
cnf(c_0_445,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_439, c_0_366]), c_0_367])).
cnf(c_0_446,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_366]), c_0_367])).
cnf(c_0_447,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_441, c_0_437])).
cnf(c_0_448,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_442, c_0_437])).
cnf(c_0_449,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_443, c_0_437]), c_0_391])).
cnf(c_0_450,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_437]), c_0_391])).
cnf(c_0_451,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_395]), c_0_398]), c_0_347])).
cnf(c_0_452,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_395]), c_0_398]), c_0_347])).
cnf(c_0_453,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_445]), c_0_349]), c_0_446])).
cnf(c_0_454,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_447, c_0_366]), c_0_367])).
cnf(c_0_455,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_448, c_0_366]), c_0_367])).
cnf(c_0_456,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_366]), c_0_367])).
cnf(c_0_457,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_366]), c_0_367])).
cnf(c_0_458,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_437]), c_0_391])).
cnf(c_0_459,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_452, c_0_437]), c_0_391])).
cnf(c_0_460,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_453, c_0_454]), c_0_455])).
cnf(c_0_461,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_456]), c_0_349]), c_0_457])).
cnf(c_0_462,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_458, c_0_366]), c_0_367])).
cnf(c_0_463,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_459, c_0_366]), c_0_367])).
cnf(c_0_464,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_389, c_0_460])).
cnf(c_0_465,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_462]), c_0_463])).
cnf(c_0_466,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_464]), c_0_465]), c_0_347])).
cnf(c_0_467,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_396, c_0_460])).
cnf(c_0_468,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_464]), c_0_465]), c_0_347])).
cnf(c_0_469,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_466, c_0_467])).
cnf(c_0_470,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_468, c_0_467])).
cnf(c_0_471,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_464]), c_0_465]), c_0_347])).
cnf(c_0_472,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_464]), c_0_465]), c_0_347])).
cnf(c_0_473,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_469, c_0_366]), c_0_367])).
cnf(c_0_474,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_470, c_0_366]), c_0_367])).
cnf(c_0_475,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_471, c_0_467])).
cnf(c_0_476,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_472, c_0_467])).
cnf(c_0_477,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_473]), c_0_349]), c_0_474])).
cnf(c_0_478,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_475, c_0_366]), c_0_367])).
cnf(c_0_479,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_366]), c_0_367])).
cnf(c_0_480,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_477, c_0_478]), c_0_479])).
cnf(c_0_481,negated_conjecture,(p5(esk8_1(X1))|p2(X2)|~r1(esk1_0,X1)|~r1(esk1_0,X2)|~p2(esk11_1(X2))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_482,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_427, c_0_480])).
cnf(c_0_483,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(X1))|r1(esk1_0,esk17_0)|~p2(esk11_1(esk33_1(esk1_0)))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_349]), c_0_382])).
cnf(c_0_484,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_385]), c_0_460])).
cnf(c_0_485,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(X1))|r1(esk1_0,esk17_0)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_483, c_0_484]), c_0_382])).
cnf(c_0_486,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_485, c_0_7])).
cnf(c_0_487,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_389, c_0_486])).
cnf(c_0_488,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_391, c_0_486])).
cnf(c_0_489,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_487]), c_0_488]), c_0_347])).
cnf(c_0_490,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_396, c_0_486])).
cnf(c_0_491,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_487]), c_0_488]), c_0_347])).
cnf(c_0_492,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_489, c_0_490])).
cnf(c_0_493,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_491, c_0_490])).
cnf(c_0_494,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_487]), c_0_488]), c_0_347])).
cnf(c_0_495,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_487]), c_0_488]), c_0_347])).
cnf(c_0_496,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_366]), c_0_367])).
cnf(c_0_497,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_366]), c_0_367])).
cnf(c_0_498,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_494, c_0_490])).
cnf(c_0_499,plain,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_495, c_0_490])).
cnf(c_0_500,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_496]), c_0_349]), c_0_497])).
cnf(c_0_501,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_498, c_0_366]), c_0_367])).
cnf(c_0_502,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_366]), c_0_367])).
cnf(c_0_503,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_501]), c_0_502])).
cnf(c_0_504,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_427, c_0_503])).
cnf(c_0_505,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_504, c_0_385]), c_0_486])).
cnf(c_0_506,plain,(p2(esk9_0)|r1(esk33_1(esk1_0),esk34_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_505]), c_0_347])).
cnf(c_0_507,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_506]), c_0_347])).
cnf(c_0_508,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)|r1(esk34_1(esk1_0),esk36_2(esk1_0,esk34_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_506]), c_0_505]), c_0_347])).
cnf(c_0_509,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk34_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_506]), c_0_505]), c_0_347])).
cnf(c_0_510,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_506]), c_0_347])).
cnf(c_0_511,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk9_0)|p2(X3)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_506]), c_0_347])).
cnf(c_0_512,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_506]), c_0_347])).
cnf(c_0_513,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_507, c_0_508]), c_0_509]), c_0_510])).
cnf(c_0_514,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_511, c_0_508]), c_0_509]), c_0_512])).
cnf(c_0_515,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_506]), c_0_347])).
cnf(c_0_516,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_506]), c_0_347])).
cnf(c_0_517,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_506]), c_0_347])).
cnf(c_0_518,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_506]), c_0_347])).
cnf(c_0_519,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk45_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_506]), c_0_347])).
cnf(c_0_520,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_506]), c_0_347])).
cnf(c_0_521,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_513, c_0_7])).
cnf(c_0_522,plain,(p2(esk9_0)|r1(esk34_1(esk1_0),esk35_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_505]), c_0_347])).
cnf(c_0_523,plain,(p2(esk9_0)|r1(esk1_0,esk17_0)|~p2(esk35_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_505]), c_0_347])).
cnf(c_0_524,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_514, c_0_7])).
cnf(c_0_525,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_515, c_0_508]), c_0_509]), c_0_516])).
cnf(c_0_526,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_508]), c_0_509]), c_0_518])).
cnf(c_0_527,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_519, c_0_508]), c_0_509]), c_0_520])).
cnf(c_0_528,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_506]), c_0_347])).
cnf(c_0_529,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_506]), c_0_347])).
cnf(c_0_530,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_521, c_0_522]), c_0_523])).
cnf(c_0_531,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_524, c_0_522]), c_0_523])).
cnf(c_0_532,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_525, c_0_7])).
cnf(c_0_533,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_526, c_0_7])).
cnf(c_0_534,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_527, c_0_7])).
cnf(c_0_535,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_528, c_0_508]), c_0_509]), c_0_529])).
cnf(c_0_536,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_530, c_0_522]), c_0_523])).
cnf(c_0_537,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_531, c_0_522]), c_0_523])).
cnf(c_0_538,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_522]), c_0_523])).
cnf(c_0_539,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_533, c_0_522]), c_0_523])).
cnf(c_0_540,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_522]), c_0_523])).
cnf(c_0_541,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_535, c_0_7])).
cnf(c_0_542,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_536]), c_0_537]), c_0_505]), c_0_347])).
cnf(c_0_543,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_538, c_0_522]), c_0_523])).
cnf(c_0_544,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_506]), c_0_347])).
cnf(c_0_545,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_506]), c_0_347])).
cnf(c_0_546,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X3)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X3),esk46_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_506]), c_0_347])).
cnf(c_0_547,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_506]), c_0_347])).
cnf(c_0_548,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_539, c_0_522]), c_0_523])).
cnf(c_0_549,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_540, c_0_522]), c_0_523])).
cnf(c_0_550,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_541, c_0_522]), c_0_523])).
cnf(c_0_551,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_542, c_0_543])).
cnf(c_0_552,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_544, c_0_508]), c_0_509]), c_0_545])).
cnf(c_0_553,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_546, c_0_508]), c_0_509]), c_0_547])).
cnf(c_0_554,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_548]), c_0_549]), c_0_505]), c_0_347])).
cnf(c_0_555,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_522]), c_0_523])).
cnf(c_0_556,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_551]), c_0_522]), c_0_506]), c_0_509]), c_0_347]), c_0_523])).
cnf(c_0_557,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_552, c_0_7])).
cnf(c_0_558,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_553, c_0_7])).
cnf(c_0_559,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_506]), c_0_347])).
cnf(c_0_560,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_506]), c_0_347])).
cnf(c_0_561,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_554, c_0_555])).
cnf(c_0_562,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_556, c_0_522]), c_0_523])).
cnf(c_0_563,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_557, c_0_522]), c_0_523])).
cnf(c_0_564,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_558, c_0_522]), c_0_523])).
cnf(c_0_565,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_559, c_0_508]), c_0_509]), c_0_560])).
cnf(c_0_566,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_561]), c_0_562]), c_0_505]), c_0_347])).
cnf(c_0_567,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_563, c_0_522]), c_0_523])).
cnf(c_0_568,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_549]), c_0_537]), c_0_505]), c_0_347])).
cnf(c_0_569,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_564, c_0_522]), c_0_523])).
cnf(c_0_570,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_565, c_0_7])).
cnf(c_0_571,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_566, c_0_567])).
cnf(c_0_572,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_568, c_0_569])).
cnf(c_0_573,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_570, c_0_522]), c_0_523])).
cnf(c_0_574,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_571]), c_0_572]), c_0_505]), c_0_347])).
cnf(c_0_575,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_573, c_0_522]), c_0_523])).
cnf(c_0_576,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk34_1(esk1_0)),esk37_2(esk1_0,esk34_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_506]), c_0_505]), c_0_347])).
cnf(c_0_577,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_575]), c_0_566])).
cnf(c_0_578,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_526, c_0_576]), c_0_518])).
cnf(c_0_579,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_527, c_0_576]), c_0_520])).
cnf(c_0_580,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_576]), c_0_510])).
cnf(c_0_581,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_514, c_0_576]), c_0_512])).
cnf(c_0_582,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_577]), c_0_522]), c_0_506]), c_0_509]), c_0_347]), c_0_523])).
cnf(c_0_583,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_578, c_0_522]), c_0_523])).
cnf(c_0_584,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_579, c_0_522]), c_0_523])).
cnf(c_0_585,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_535, c_0_576]), c_0_529])).
cnf(c_0_586,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_580, c_0_522]), c_0_523])).
cnf(c_0_587,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_581, c_0_522]), c_0_523])).
cnf(c_0_588,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_525, c_0_576]), c_0_516])).
cnf(c_0_589,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_582, c_0_522]), c_0_523])).
cnf(c_0_590,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_577]), c_0_522]), c_0_506]), c_0_509]), c_0_347]), c_0_523])).
cnf(c_0_591,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_522]), c_0_523])).
cnf(c_0_592,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_584, c_0_522]), c_0_523])).
cnf(c_0_593,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_522]), c_0_523])).
cnf(c_0_594,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_522]), c_0_523])).
cnf(c_0_595,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_587, c_0_522]), c_0_523])).
cnf(c_0_596,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_588, c_0_522]), c_0_523])).
cnf(c_0_597,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_589]), c_0_562]), c_0_505]), c_0_347])).
cnf(c_0_598,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_590, c_0_522]), c_0_523])).
cnf(c_0_599,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_591]), c_0_592]), c_0_505]), c_0_347])).
cnf(c_0_600,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_593, c_0_522]), c_0_523])).
cnf(c_0_601,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_594]), c_0_595]), c_0_505]), c_0_347])).
cnf(c_0_602,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_596, c_0_522]), c_0_523])).
cnf(c_0_603,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_552, c_0_576]), c_0_545])).
cnf(c_0_604,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_553, c_0_576]), c_0_547])).
cnf(c_0_605,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_597, c_0_598])).
cnf(c_0_606,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_599, c_0_600])).
cnf(c_0_607,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_601, c_0_602])).
cnf(c_0_608,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_603, c_0_522]), c_0_523])).
cnf(c_0_609,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_604, c_0_522]), c_0_523])).
cnf(c_0_610,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_605]), c_0_522]), c_0_506]), c_0_509]), c_0_347]), c_0_523])).
cnf(c_0_611,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_605]), c_0_522]), c_0_506]), c_0_509]), c_0_347]), c_0_523])).
cnf(c_0_612,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_606]), c_0_607]), c_0_505]), c_0_347])).
cnf(c_0_613,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_608, c_0_522]), c_0_523])).
cnf(c_0_614,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_592]), c_0_595]), c_0_505]), c_0_347])).
cnf(c_0_615,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_609, c_0_522]), c_0_523])).
cnf(c_0_616,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_565, c_0_576]), c_0_560])).
cnf(c_0_617,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_610, c_0_522]), c_0_523])).
cnf(c_0_618,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_611, c_0_522]), c_0_523])).
cnf(c_0_619,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_605]), c_0_522]), c_0_506]), c_0_509]), c_0_347]), c_0_523])).
cnf(c_0_620,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_612, c_0_613])).
cnf(c_0_621,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_614, c_0_615])).
cnf(c_0_622,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_616, c_0_522]), c_0_523])).
cnf(c_0_623,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_617]), c_0_618]), c_0_505]), c_0_347])).
cnf(c_0_624,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_619, c_0_522]), c_0_523])).
cnf(c_0_625,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_620]), c_0_621]), c_0_505]), c_0_347])).
cnf(c_0_626,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_622, c_0_522]), c_0_523])).
cnf(c_0_627,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_623, c_0_624])).
cnf(c_0_628,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_625, c_0_626]), c_0_612])).
cnf(c_0_629,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_627]), c_0_522]), c_0_506]), c_0_509]), c_0_347]), c_0_523])).
cnf(c_0_630,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_628]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_631,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_605]), c_0_522]), c_0_523])).
cnf(c_0_632,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_628]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_633,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_607]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_634,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_607]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_635,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_630, c_0_508]), c_0_631]), c_0_632])).
cnf(c_0_636,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_633, c_0_508]), c_0_509]), c_0_634])).
cnf(c_0_637,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_628]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_638,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_628]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_639,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_635, c_0_576]), c_0_632])).
cnf(c_0_640,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_636, c_0_576]), c_0_634])).
cnf(c_0_641,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_637, c_0_508]), c_0_631]), c_0_638])).
cnf(c_0_642,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_639, c_0_522]), c_0_523])).
cnf(c_0_643,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_640, c_0_522]), c_0_523])).
cnf(c_0_644,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_641, c_0_576]), c_0_638])).
cnf(c_0_645,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_642]), c_0_643]), c_0_505]), c_0_347])).
cnf(c_0_646,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_644, c_0_522]), c_0_523])).
cnf(c_0_647,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_645, c_0_646])).
cnf(c_0_648,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_647]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_649,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_647]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_650,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_647]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_651,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_647]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_652,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_648, c_0_508]), c_0_631]), c_0_649])).
cnf(c_0_653,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_650, c_0_508]), c_0_631]), c_0_651])).
cnf(c_0_654,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_647]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_655,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_647]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_656,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_576]), c_0_649])).
cnf(c_0_657,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_653, c_0_576]), c_0_651])).
cnf(c_0_658,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_654, c_0_508]), c_0_631]), c_0_655])).
cnf(c_0_659,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_656, c_0_522]), c_0_523])).
cnf(c_0_660,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_522]), c_0_523])).
cnf(c_0_661,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_658, c_0_576]), c_0_655])).
cnf(c_0_662,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_659]), c_0_660]), c_0_505]), c_0_347])).
cnf(c_0_663,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_661, c_0_522]), c_0_523])).
cnf(c_0_664,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_662, c_0_663])).
cnf(c_0_665,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_664]), c_0_522]), c_0_506]), c_0_347]), c_0_523])).
cnf(c_0_666,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(X2)|~r1(esk34_1(esk1_0),X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_665, c_0_647]), c_0_522]), c_0_523])).
cnf(c_0_667,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_506]), c_0_505]), c_0_347])).
cnf(c_0_668,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_666, c_0_508]), c_0_631]), c_0_667])).
cnf(c_0_669,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_576]), c_0_667])).
cnf(c_0_670,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_518, c_0_669])).
cnf(c_0_671,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_520, c_0_669])).
cnf(c_0_672,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_510, c_0_669])).
cnf(c_0_673,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_512, c_0_669])).
cnf(c_0_674,plain,(p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_670, c_0_522]), c_0_523])).
cnf(c_0_675,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_671, c_0_522]), c_0_523])).
cnf(c_0_676,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_529, c_0_669])).
cnf(c_0_677,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_672, c_0_522]), c_0_523])).
cnf(c_0_678,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_673, c_0_522]), c_0_523])).
cnf(c_0_679,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_516, c_0_669])).
cnf(c_0_680,plain,(p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_674, c_0_522]), c_0_523])).
cnf(c_0_681,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_675, c_0_522]), c_0_523])).
cnf(c_0_682,plain,(p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_676, c_0_522]), c_0_523])).
cnf(c_0_683,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_677, c_0_522]), c_0_523])).
cnf(c_0_684,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_522]), c_0_523])).
cnf(c_0_685,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_679, c_0_522]), c_0_523])).
cnf(c_0_686,plain,(p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_680]), c_0_681]), c_0_505]), c_0_347])).
cnf(c_0_687,plain,(p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_682, c_0_522]), c_0_523])).
cnf(c_0_688,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_683]), c_0_684]), c_0_505]), c_0_347])).
cnf(c_0_689,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_685, c_0_522]), c_0_523])).
cnf(c_0_690,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_545, c_0_669])).
cnf(c_0_691,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_547, c_0_669])).
cnf(c_0_692,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_686, c_0_687])).
cnf(c_0_693,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_688, c_0_689])).
cnf(c_0_694,plain,(p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_522]), c_0_523])).
cnf(c_0_695,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_691, c_0_522]), c_0_523])).
cnf(c_0_696,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_692]), c_0_693]), c_0_505]), c_0_347])).
cnf(c_0_697,plain,(p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_694, c_0_522]), c_0_523])).
cnf(c_0_698,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_681]), c_0_684]), c_0_505]), c_0_347])).
cnf(c_0_699,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_522]), c_0_523])).
cnf(c_0_700,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_560, c_0_669])).
cnf(c_0_701,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_696, c_0_697])).
cnf(c_0_702,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_698, c_0_699])).
cnf(c_0_703,plain,(p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_700, c_0_522]), c_0_523])).
cnf(c_0_704,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_701]), c_0_702]), c_0_505]), c_0_347])).
cnf(c_0_705,plain,(p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_703, c_0_522]), c_0_523])).
cnf(c_0_706,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_704, c_0_705]), c_0_696])).
cnf(c_0_707,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_706]), c_0_522]), c_0_506]), c_0_669]), c_0_347]), c_0_523])).
cnf(c_0_708,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_693]), c_0_522]), c_0_506]), c_0_669]), c_0_347]), c_0_523])).
cnf(c_0_709,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_522]), c_0_523])).
cnf(c_0_710,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_708, c_0_522]), c_0_523])).
cnf(c_0_711,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_706]), c_0_522]), c_0_506]), c_0_669]), c_0_347]), c_0_523])).
cnf(c_0_712,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_709]), c_0_710]), c_0_505]), c_0_347])).
cnf(c_0_713,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_711, c_0_522]), c_0_523])).
cnf(c_0_714,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_712, c_0_713])).
cnf(c_0_715,plain,(p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_714]), c_0_522]), c_0_506]), c_0_669]), c_0_347]), c_0_523])).
cnf(c_0_716,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_714]), c_0_522]), c_0_506]), c_0_669]), c_0_347]), c_0_523])).
cnf(c_0_717,plain,(p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_715, c_0_522]), c_0_523])).
cnf(c_0_718,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_716, c_0_522]), c_0_523])).
cnf(c_0_719,plain,(p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_714]), c_0_522]), c_0_506]), c_0_669]), c_0_347]), c_0_523])).
cnf(c_0_720,plain,(p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_717]), c_0_718]), c_0_505]), c_0_347])).
cnf(c_0_721,plain,(p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_719, c_0_522]), c_0_523])).
cnf(c_0_722,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_720, c_0_721])).
cnf(c_0_723,plain,(p2(esk9_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_722]), c_0_522]), c_0_506]), c_0_669]), c_0_347]), c_0_523])).
cnf(c_0_724,plain,(p2(esk9_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_723, c_0_714]), c_0_522]), c_0_523])).
cnf(c_0_725,negated_conjecture,(r1(esk1_0,esk17_0)|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_350, c_0_724])).
cnf(c_0_726,negated_conjecture,(p5(esk8_1(X1))|r1(esk1_0,esk17_0)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_382, c_0_724])).
cnf(c_0_727,negated_conjecture,(r1(esk1_0,esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_725, c_0_7])).
cnf(c_0_728,negated_conjecture,(p5(esk8_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_726, c_0_7])).
cnf(c_0_729,negated_conjecture,(r1(esk1_0,esk17_0)|r1(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_727]), c_0_728])).
cnf(c_0_730,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk15_2(esk2_0,X1))|r1(esk1_0,esk17_0)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_12, c_0_729])).
cnf(c_0_731,negated_conjecture,(p2(esk2_0)|r1(esk2_0,esk32_1(esk2_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_14, c_0_729])).
cnf(c_0_732,negated_conjecture,(p2(esk2_0)|r1(esk1_0,esk17_0)|~p2(esk32_1(esk2_0))), inference(spm,[status(thm)],[c_0_15, c_0_729])).
cnf(c_0_733,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk2_0,X1))|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_729])).
cnf(c_0_734,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk32_1(esk2_0),esk15_2(esk2_0,esk32_1(esk2_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_731]), c_0_732])).
cnf(c_0_735,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk2_0,esk32_1(esk2_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_733, c_0_731]), c_0_732])).
cnf(c_0_736,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk15_2(esk2_0,X1),esk16_2(esk2_0,X1))|r1(esk1_0,esk17_0)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_21, c_0_729])).
cnf(c_0_737,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk15_2(esk2_0,esk32_1(esk2_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_734]), c_0_729]), c_0_735])).
cnf(c_0_738,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk15_2(esk2_0,esk32_1(esk2_0)),esk16_2(esk2_0,esk32_1(esk2_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_736, c_0_731]), c_0_732])).
cnf(c_0_739,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk2_0,esk32_1(esk2_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_737, c_0_738])).
cnf(c_0_740,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_739]), c_0_14]), c_0_729]), c_0_732])).
cnf(c_0_741,plain,(p2(esk2_0)|r1(esk1_0,esk33_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_32, c_0_740])).
cnf(c_0_742,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_741]), c_0_36])).
cnf(c_0_743,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_741]), c_0_36])).
cnf(c_0_744,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_742, c_0_727]), c_0_728])).
cnf(c_0_745,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_743, c_0_727]), c_0_728])).
cnf(c_0_746,plain,(p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_744]), c_0_745]), c_0_43])).
cnf(c_0_747,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_741]), c_0_36])).
cnf(c_0_748,plain,(p2(esk36_2(esk1_0,X1))|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_744]), c_0_745]), c_0_46])).
cnf(c_0_749,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_746, c_0_740])).
cnf(c_0_750,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_747, c_0_727]), c_0_728])).
cnf(c_0_751,negated_conjecture,(p2(esk36_2(esk1_0,X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_748, c_0_740])).
cnf(c_0_752,plain,(p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_744]), c_0_745]), c_0_51])).
cnf(c_0_753,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_744]), c_0_745]), c_0_53])).
cnf(c_0_754,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_749, c_0_750])).
cnf(c_0_755,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk33_1(esk1_0),esk32_1(esk33_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_14, c_0_741])).
cnf(c_0_756,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)|~p2(esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_15, c_0_741])).
cnf(c_0_757,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_751, c_0_750])).
cnf(c_0_758,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_752, c_0_740])).
cnf(c_0_759,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_753, c_0_740])).
cnf(c_0_760,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_754, c_0_755]), c_0_756])).
cnf(c_0_761,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_757, c_0_755]), c_0_756])).
cnf(c_0_762,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk17_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_758, c_0_750])).
cnf(c_0_763,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_759, c_0_750])).
cnf(c_0_764,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_760]), c_0_741]), c_0_761])).
cnf(c_0_765,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_762, c_0_755]), c_0_756])).
cnf(c_0_766,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_763, c_0_755]), c_0_756])).
cnf(c_0_767,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk17_0)|~p2(esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_741]), c_0_36])).
cnf(c_0_768,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_765]), c_0_766])).
cnf(c_0_769,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_767, c_0_768]), c_0_36])).
cnf(c_0_770,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_769, c_0_727]), c_0_728])).
cnf(c_0_771,plain,(r1(esk33_1(esk1_0),esk34_1(esk1_0))|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_76, c_0_770])).
cnf(c_0_772,negated_conjecture,(p2(esk2_0)|r1(esk33_1(esk1_0),esk34_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_771, c_0_740])).
cnf(c_0_773,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_772]), c_0_740])).
cnf(c_0_774,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk34_1(esk1_0),esk36_2(esk1_0,esk34_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_772]), c_0_770]), c_0_740])).
cnf(c_0_775,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_772]), c_0_770]), c_0_740])).
cnf(c_0_776,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_772]), c_0_740])).
cnf(c_0_777,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X3)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_772]), c_0_740])).
cnf(c_0_778,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_772]), c_0_740])).
cnf(c_0_779,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_773, c_0_774]), c_0_775]), c_0_776])).
cnf(c_0_780,plain,(r1(esk34_1(esk1_0),esk35_1(esk1_0))|r1(esk1_0,esk17_0)|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_88, c_0_770])).
cnf(c_0_781,plain,(r1(esk1_0,esk17_0)|~epred1_1(esk1_0)|~p2(esk35_1(esk1_0))), inference(spm,[status(thm)],[c_0_89, c_0_770])).
cnf(c_0_782,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_777, c_0_774]), c_0_775]), c_0_778])).
cnf(c_0_783,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_772]), c_0_740])).
cnf(c_0_784,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_772]), c_0_740])).
cnf(c_0_785,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_772]), c_0_740])).
cnf(c_0_786,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_772]), c_0_740])).
cnf(c_0_787,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk45_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_772]), c_0_740])).
cnf(c_0_788,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_772]), c_0_740])).
cnf(c_0_789,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_779, c_0_7])).
cnf(c_0_790,negated_conjecture,(p2(esk2_0)|r1(esk34_1(esk1_0),esk35_1(esk1_0))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_780, c_0_740])).
cnf(c_0_791,negated_conjecture,(p2(esk2_0)|r1(esk1_0,esk17_0)|~p2(esk35_1(esk1_0))), inference(spm,[status(thm)],[c_0_781, c_0_740])).
cnf(c_0_792,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_782, c_0_7])).
cnf(c_0_793,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_783, c_0_774]), c_0_775]), c_0_784])).
cnf(c_0_794,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_785, c_0_774]), c_0_775]), c_0_786])).
cnf(c_0_795,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_787, c_0_774]), c_0_775]), c_0_788])).
cnf(c_0_796,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_772]), c_0_740])).
cnf(c_0_797,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_772]), c_0_740])).
cnf(c_0_798,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_789, c_0_790]), c_0_791])).
cnf(c_0_799,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_792, c_0_790]), c_0_791])).
cnf(c_0_800,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_793, c_0_7])).
cnf(c_0_801,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_794, c_0_7])).
cnf(c_0_802,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_795, c_0_7])).
cnf(c_0_803,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_796, c_0_774]), c_0_775]), c_0_797])).
cnf(c_0_804,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_798, c_0_790]), c_0_791])).
cnf(c_0_805,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_799, c_0_790]), c_0_791])).
cnf(c_0_806,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_800, c_0_790]), c_0_791])).
cnf(c_0_807,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_801, c_0_790]), c_0_791])).
cnf(c_0_808,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_802, c_0_790]), c_0_791])).
cnf(c_0_809,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_803, c_0_7])).
cnf(c_0_810,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_804]), c_0_805]), c_0_770]), c_0_740])).
cnf(c_0_811,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_806, c_0_790]), c_0_791])).
cnf(c_0_812,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_772]), c_0_740])).
cnf(c_0_813,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_772]), c_0_740])).
cnf(c_0_814,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X3)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X3),esk46_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_772]), c_0_740])).
cnf(c_0_815,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_772]), c_0_740])).
cnf(c_0_816,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_807, c_0_790]), c_0_791])).
cnf(c_0_817,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_808, c_0_790]), c_0_791])).
cnf(c_0_818,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_809, c_0_790]), c_0_791])).
cnf(c_0_819,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_810, c_0_811])).
cnf(c_0_820,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_812, c_0_774]), c_0_775]), c_0_813])).
cnf(c_0_821,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_814, c_0_774]), c_0_775]), c_0_815])).
cnf(c_0_822,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_816]), c_0_817]), c_0_770]), c_0_740])).
cnf(c_0_823,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_818, c_0_790]), c_0_791])).
cnf(c_0_824,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_819]), c_0_790]), c_0_772]), c_0_775]), c_0_740]), c_0_791])).
cnf(c_0_825,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_820, c_0_7])).
cnf(c_0_826,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_821, c_0_7])).
cnf(c_0_827,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_772]), c_0_740])).
cnf(c_0_828,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_772]), c_0_740])).
cnf(c_0_829,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_822, c_0_823])).
cnf(c_0_830,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_824, c_0_790]), c_0_791])).
cnf(c_0_831,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_825, c_0_790]), c_0_791])).
cnf(c_0_832,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_826, c_0_790]), c_0_791])).
cnf(c_0_833,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_827, c_0_774]), c_0_775]), c_0_828])).
cnf(c_0_834,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_829]), c_0_830]), c_0_770]), c_0_740])).
cnf(c_0_835,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_831, c_0_790]), c_0_791])).
cnf(c_0_836,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_817]), c_0_805]), c_0_770]), c_0_740])).
cnf(c_0_837,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_832, c_0_790]), c_0_791])).
cnf(c_0_838,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_833, c_0_7])).
cnf(c_0_839,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_834, c_0_835])).
cnf(c_0_840,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_836, c_0_837])).
cnf(c_0_841,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_838, c_0_790]), c_0_791])).
cnf(c_0_842,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk36_2(esk1_0,esk34_1(esk1_0)),esk37_2(esk1_0,esk34_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_772]), c_0_770]), c_0_740])).
cnf(c_0_843,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_839]), c_0_840]), c_0_770]), c_0_740])).
cnf(c_0_844,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_841, c_0_790]), c_0_791])).
cnf(c_0_845,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_842]), c_0_786])).
cnf(c_0_846,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_795, c_0_842]), c_0_788])).
cnf(c_0_847,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_779, c_0_842]), c_0_776])).
cnf(c_0_848,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_782, c_0_842]), c_0_778])).
cnf(c_0_849,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_843, c_0_844]), c_0_834])).
cnf(c_0_850,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_845, c_0_790]), c_0_791])).
cnf(c_0_851,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_846, c_0_790]), c_0_791])).
cnf(c_0_852,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_803, c_0_842]), c_0_797])).
cnf(c_0_853,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_847, c_0_790]), c_0_791])).
cnf(c_0_854,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_848, c_0_790]), c_0_791])).
cnf(c_0_855,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_793, c_0_842]), c_0_784])).
cnf(c_0_856,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_849]), c_0_790]), c_0_772]), c_0_775]), c_0_740]), c_0_791])).
cnf(c_0_857,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_850, c_0_790]), c_0_791])).
cnf(c_0_858,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_851, c_0_790]), c_0_791])).
cnf(c_0_859,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_852, c_0_790]), c_0_791])).
cnf(c_0_860,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_853, c_0_790]), c_0_791])).
cnf(c_0_861,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_854, c_0_790]), c_0_791])).
cnf(c_0_862,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_855, c_0_790]), c_0_791])).
cnf(c_0_863,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_856, c_0_790]), c_0_791])).
cnf(c_0_864,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_849]), c_0_790]), c_0_772]), c_0_775]), c_0_740]), c_0_791])).
cnf(c_0_865,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_857]), c_0_858]), c_0_770]), c_0_740])).
cnf(c_0_866,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_859, c_0_790]), c_0_791])).
cnf(c_0_867,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_860]), c_0_861]), c_0_770]), c_0_740])).
cnf(c_0_868,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_862, c_0_790]), c_0_791])).
cnf(c_0_869,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_820, c_0_842]), c_0_813])).
cnf(c_0_870,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_821, c_0_842]), c_0_815])).
cnf(c_0_871,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_863]), c_0_830]), c_0_770]), c_0_740])).
cnf(c_0_872,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_864, c_0_790]), c_0_791])).
cnf(c_0_873,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_865, c_0_866])).
cnf(c_0_874,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_867, c_0_868])).
cnf(c_0_875,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_869, c_0_790]), c_0_791])).
cnf(c_0_876,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_870, c_0_790]), c_0_791])).
cnf(c_0_877,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_871, c_0_872])).
cnf(c_0_878,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_873]), c_0_874]), c_0_770]), c_0_740])).
cnf(c_0_879,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_875, c_0_790]), c_0_791])).
cnf(c_0_880,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_858]), c_0_861]), c_0_770]), c_0_740])).
cnf(c_0_881,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_876, c_0_790]), c_0_791])).
cnf(c_0_882,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_833, c_0_842]), c_0_828])).
cnf(c_0_883,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_877]), c_0_790]), c_0_772]), c_0_775]), c_0_740]), c_0_791])).
cnf(c_0_884,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_877]), c_0_790]), c_0_772]), c_0_775]), c_0_740]), c_0_791])).
cnf(c_0_885,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_878, c_0_879])).
cnf(c_0_886,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_880, c_0_881])).
cnf(c_0_887,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_882, c_0_790]), c_0_791])).
cnf(c_0_888,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_883, c_0_790]), c_0_791])).
cnf(c_0_889,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_884, c_0_790]), c_0_791])).
cnf(c_0_890,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_877]), c_0_790]), c_0_772]), c_0_775]), c_0_740]), c_0_791])).
cnf(c_0_891,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_885]), c_0_886]), c_0_770]), c_0_740])).
cnf(c_0_892,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_887, c_0_790]), c_0_791])).
cnf(c_0_893,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_888]), c_0_889]), c_0_770]), c_0_740])).
cnf(c_0_894,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_890, c_0_790]), c_0_791])).
cnf(c_0_895,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_891, c_0_892]), c_0_878])).
cnf(c_0_896,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_893, c_0_894])).
cnf(c_0_897,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_895]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_898,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_895]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_899,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_874]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_900,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_874]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_901,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_896]), c_0_790]), c_0_772]), c_0_775]), c_0_740]), c_0_791])).
cnf(c_0_902,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_897, c_0_774]), c_0_775]), c_0_898])).
cnf(c_0_903,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_899, c_0_774]), c_0_775]), c_0_900])).
cnf(c_0_904,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_895]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_905,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_901, c_0_877]), c_0_790]), c_0_791])).
cnf(c_0_906,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_895]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_907,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_902, c_0_842]), c_0_898])).
cnf(c_0_908,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_903, c_0_842]), c_0_900])).
cnf(c_0_909,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_904, c_0_774]), c_0_905]), c_0_906])).
cnf(c_0_910,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_907, c_0_790]), c_0_791])).
cnf(c_0_911,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_908, c_0_790]), c_0_791])).
cnf(c_0_912,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_909, c_0_842]), c_0_906])).
cnf(c_0_913,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_910]), c_0_911]), c_0_770]), c_0_740])).
cnf(c_0_914,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_912, c_0_790]), c_0_791])).
cnf(c_0_915,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_913, c_0_914])).
cnf(c_0_916,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_915]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_917,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_915]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_918,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_915]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_919,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_915]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_920,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_916, c_0_774]), c_0_905]), c_0_917])).
cnf(c_0_921,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_918, c_0_774]), c_0_905]), c_0_919])).
cnf(c_0_922,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_915]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_923,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_915]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_924,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_920, c_0_842]), c_0_917])).
cnf(c_0_925,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_921, c_0_842]), c_0_919])).
cnf(c_0_926,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_922, c_0_774]), c_0_905]), c_0_923])).
cnf(c_0_927,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_924, c_0_790]), c_0_791])).
cnf(c_0_928,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_925, c_0_790]), c_0_791])).
cnf(c_0_929,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_926, c_0_842]), c_0_923])).
cnf(c_0_930,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_927]), c_0_928]), c_0_770]), c_0_740])).
cnf(c_0_931,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_929, c_0_790]), c_0_791])).
cnf(c_0_932,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_930, c_0_931])).
cnf(c_0_933,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk17_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_932]), c_0_790]), c_0_772]), c_0_740]), c_0_791])).
cnf(c_0_934,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(X2)|~r1(esk34_1(esk1_0),X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_933, c_0_915]), c_0_790]), c_0_791])).
cnf(c_0_935,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)|~p2(esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_772]), c_0_770]), c_0_740])).
cnf(c_0_936,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_934, c_0_774]), c_0_905]), c_0_935])).
cnf(c_0_937,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_842]), c_0_935])).
cnf(c_0_938,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_786, c_0_937])).
cnf(c_0_939,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_788, c_0_937])).
cnf(c_0_940,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_776, c_0_937])).
cnf(c_0_941,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_778, c_0_937])).
cnf(c_0_942,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_938, c_0_790]), c_0_791])).
cnf(c_0_943,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_939, c_0_790]), c_0_791])).
cnf(c_0_944,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_797, c_0_937])).
cnf(c_0_945,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_940, c_0_790]), c_0_791])).
cnf(c_0_946,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_941, c_0_790]), c_0_791])).
cnf(c_0_947,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_784, c_0_937])).
cnf(c_0_948,negated_conjecture,(p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_942, c_0_790]), c_0_791])).
cnf(c_0_949,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_943, c_0_790]), c_0_791])).
cnf(c_0_950,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_944, c_0_790]), c_0_791])).
cnf(c_0_951,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_945, c_0_790]), c_0_791])).
cnf(c_0_952,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_946, c_0_790]), c_0_791])).
cnf(c_0_953,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_947, c_0_790]), c_0_791])).
cnf(c_0_954,plain,(p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_948]), c_0_949]), c_0_770]), c_0_740])).
cnf(c_0_955,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_950, c_0_790]), c_0_791])).
cnf(c_0_956,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_951]), c_0_952]), c_0_770]), c_0_740])).
cnf(c_0_957,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_953, c_0_790]), c_0_791])).
cnf(c_0_958,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_813, c_0_937])).
cnf(c_0_959,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_815, c_0_937])).
cnf(c_0_960,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_954, c_0_955])).
cnf(c_0_961,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_956, c_0_957])).
cnf(c_0_962,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_958, c_0_790]), c_0_791])).
cnf(c_0_963,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_959, c_0_790]), c_0_791])).
cnf(c_0_964,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_960]), c_0_961]), c_0_770]), c_0_740])).
cnf(c_0_965,negated_conjecture,(p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_962, c_0_790]), c_0_791])).
cnf(c_0_966,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_949]), c_0_952]), c_0_770]), c_0_740])).
cnf(c_0_967,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_963, c_0_790]), c_0_791])).
cnf(c_0_968,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_828, c_0_937])).
cnf(c_0_969,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_964, c_0_965])).
cnf(c_0_970,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_966, c_0_967])).
cnf(c_0_971,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_968, c_0_790]), c_0_791])).
cnf(c_0_972,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_969]), c_0_970]), c_0_770]), c_0_740])).
cnf(c_0_973,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_971, c_0_790]), c_0_791])).
cnf(c_0_974,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_972, c_0_973]), c_0_964])).
cnf(c_0_975,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_974]), c_0_790]), c_0_772]), c_0_937]), c_0_740]), c_0_791])).
cnf(c_0_976,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_961]), c_0_790]), c_0_772]), c_0_937]), c_0_740]), c_0_791])).
cnf(c_0_977,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_975, c_0_790]), c_0_791])).
cnf(c_0_978,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_976, c_0_790]), c_0_791])).
cnf(c_0_979,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_974]), c_0_790]), c_0_772]), c_0_937]), c_0_740]), c_0_791])).
cnf(c_0_980,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_977]), c_0_978]), c_0_770]), c_0_740])).
cnf(c_0_981,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_979, c_0_790]), c_0_791])).
cnf(c_0_982,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_980, c_0_981])).
cnf(c_0_983,plain,(p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_982]), c_0_790]), c_0_772]), c_0_937]), c_0_740]), c_0_791])).
cnf(c_0_984,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_982]), c_0_790]), c_0_772]), c_0_937]), c_0_740]), c_0_791])).
cnf(c_0_985,negated_conjecture,(p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_983, c_0_790]), c_0_791])).
cnf(c_0_986,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_984, c_0_790]), c_0_791])).
cnf(c_0_987,plain,(p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk17_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_982]), c_0_790]), c_0_772]), c_0_937]), c_0_740]), c_0_791])).
cnf(c_0_988,plain,(p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_985]), c_0_986]), c_0_770]), c_0_740])).
cnf(c_0_989,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_987, c_0_790]), c_0_791])).
cnf(c_0_990,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk17_0)), inference(spm,[status(thm)],[c_0_988, c_0_989])).
cnf(c_0_991,plain,(p2(esk2_0)|p2(X1)|r1(esk1_0,esk17_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_990]), c_0_790]), c_0_772]), c_0_937]), c_0_740]), c_0_791])).
cnf(c_0_992,plain,(p2(esk2_0)|r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_991, c_0_982]), c_0_790]), c_0_791])).
cnf(c_0_993,negated_conjecture,(r1(esk1_0,esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_36, c_0_992])).
cnf(c_0_994,negated_conjecture,(r1(esk1_0,esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_993, c_0_727]), c_0_728])).
cnf(c_0_995,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(X2,esk15_2(X1,X2))|p2(X4)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(X3)|~r1(X3,X4)|~r1(esk17_0,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_996,negated_conjecture,(p2(esk17_0)|r1(esk17_0,esk10_1(esk17_0))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_348, c_0_994])).
cnf(c_0_997,negated_conjecture,(p2(esk10_1(esk17_0))|p2(esk17_0)|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_351, c_0_994])).
cnf(c_0_998,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|p2(esk15_2(X1,X2))|p2(X4)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(X3)|~r1(X3,X4)|~r1(esk17_0,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_999,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|p2(X2)|r1(X2,esk15_2(esk17_0,X2))|~p2(X3)|~r1(esk17_0,X3)|~r1(esk17_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_995, c_0_994])).
cnf(c_0_1000,negated_conjecture,(p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))|r1(esk17_0,esk10_1(esk17_0))), inference(spm,[status(thm)],[c_0_996, c_0_7])).
cnf(c_0_1001,negated_conjecture,(p2(esk10_1(esk17_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_997, c_0_7])).
cnf(c_0_1002,negated_conjecture,(p2(esk17_0)|r1(esk10_1(esk17_0),esk11_1(esk17_0))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_356, c_0_994])).
cnf(c_0_1003,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(X2)|p2(X1)|~p2(X3)|~r1(esk17_0,X3)|~r1(esk17_0,X1)|~r1(X3,X2)), inference(spm,[status(thm)],[c_0_998, c_0_994])).
cnf(c_0_1004,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(esk15_2(X1,X2),esk16_2(X1,X2))|p2(X4)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(X3)|~r1(X3,X4)|~r1(esk17_0,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_1005,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|r1(X1,esk15_2(esk17_0,X1))|~r1(esk10_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1000]), c_0_1001])).
cnf(c_0_1006,negated_conjecture,(p2(esk17_0)|r1(esk10_1(esk17_0),esk11_1(esk17_0))|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1002, c_0_7])).
cnf(c_0_1007,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk10_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1000]), c_0_1001])).
cnf(c_0_1008,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|p2(X2)|r1(esk15_2(esk17_0,X2),esk16_2(esk17_0,X2))|~p2(X3)|~r1(esk17_0,X3)|~r1(esk17_0,X2)|~r1(X3,X1)), inference(spm,[status(thm)],[c_0_1004, c_0_994])).
cnf(c_0_1009,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(X1,esk15_2(esk17_0,X1))|r1(esk1_0,esk8_1(esk1_0))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1005, c_0_1006])).
cnf(c_0_1010,negated_conjecture,(p2(esk17_0)|r1(esk17_0,esk32_1(esk17_0))), inference(spm,[status(thm)],[c_0_14, c_0_994])).
cnf(c_0_1011,negated_conjecture,(p2(esk17_0)|~p2(esk32_1(esk17_0))), inference(spm,[status(thm)],[c_0_15, c_0_994])).
cnf(c_0_1012,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1007, c_0_1006])).
cnf(c_0_1013,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|r1(esk1_0,esk8_1(esk1_0))|~r1(esk10_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1008, c_0_1000]), c_0_1001])).
cnf(c_0_1014,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|r1(esk32_1(esk17_0),esk15_2(esk17_0,esk32_1(esk17_0)))|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1009, c_0_1010]), c_0_1011])).
cnf(c_0_1015,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,esk32_1(esk17_0)))|p2(esk11_1(esk17_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1012, c_0_1010]), c_0_1011])).
cnf(c_0_1016,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|r1(esk1_0,esk8_1(esk1_0))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1013, c_0_1006])).
cnf(c_0_1017,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk15_2(esk17_0,esk32_1(esk17_0)),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1014]), c_0_994])]), c_0_1015])).
cnf(c_0_1018,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|r1(esk15_2(esk17_0,esk32_1(esk17_0)),esk16_2(esk17_0,esk32_1(esk17_0)))|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1016, c_0_1010]), c_0_1011])).
cnf(c_0_1019,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|p2(X4)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk16_2(X1,X2))|~p2(X3)|~r1(X3,X4)|~r1(esk17_0,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_1020,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk17_0,esk32_1(esk17_0)))|p2(esk11_1(esk17_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1017, c_0_1018])).
cnf(c_0_1021,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~p2(X2)|~r1(esk17_0,X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1019, c_0_1020]), c_0_994])]), c_0_1010]), c_0_1011])).
cnf(c_0_1022,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk10_1(esk17_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1021, c_0_1000]), c_0_1001])).
cnf(c_0_1023,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1022, c_0_1006])).
cnf(c_0_1024,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_1023]), c_0_994])])).
cnf(c_0_1025,negated_conjecture,(p2(esk17_0)|r1(esk17_0,esk3_1(esk17_0))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_34, c_0_994])).
cnf(c_0_1026,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1024, c_0_7])).
cnf(c_0_1027,negated_conjecture,(p2(esk17_0)|p5(esk8_1(X1))|r1(esk17_0,esk10_1(esk17_0))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_381, c_0_994])).
cnf(c_0_1028,negated_conjecture,(p2(esk3_1(esk17_0))|p2(esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_37, c_0_994])).
cnf(c_0_1029,negated_conjecture,(p2(esk10_1(esk17_0))|p2(esk17_0)|p5(esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_383, c_0_994])).
cnf(c_0_1030,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk17_0,esk3_1(esk17_0))|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1025, c_0_1026])).
cnf(c_0_1031,negated_conjecture,(p2(esk17_0)|p5(esk8_1(esk1_0))|r1(esk17_0,esk10_1(esk17_0))), inference(spm,[status(thm)],[c_0_1027, c_0_7])).
cnf(c_0_1032,negated_conjecture,(epred1_1(esk1_0)|p2(esk3_1(esk17_0))|p2(esk17_0)|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1028, c_0_1026])).
cnf(c_0_1033,negated_conjecture,(p2(esk17_0)|r1(esk3_1(esk17_0),esk4_1(esk17_0))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_994])).
cnf(c_0_1034,negated_conjecture,(p2(esk10_1(esk17_0))|p2(esk17_0)|p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1029, c_0_7])).
cnf(c_0_1035,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk17_0,esk10_1(esk17_0))|r1(esk17_0,esk3_1(esk17_0))), inference(spm,[status(thm)],[c_0_1030, c_0_1031])).
cnf(c_0_1036,negated_conjecture,(epred1_1(esk1_0)|p2(esk3_1(esk17_0))|p2(esk17_0)|r1(esk17_0,esk10_1(esk17_0))), inference(spm,[status(thm)],[c_0_1032, c_0_1031])).
cnf(c_0_1037,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk3_1(esk17_0),esk4_1(esk17_0))|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1033, c_0_1026])).
cnf(c_0_1038,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk17_0)|r1(esk17_0,esk3_1(esk17_0))), inference(spm,[status(thm)],[c_0_1030, c_0_1034])).
cnf(c_0_1039,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk3_1(esk17_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1032, c_0_1034])).
cnf(c_0_1040,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk17_0,esk10_1(esk17_0))|r1(X1,esk15_2(esk17_0,X1))|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1035]), c_0_1036])).
cnf(c_0_1041,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk3_1(esk17_0),esk4_1(esk17_0))|r1(esk17_0,esk10_1(esk17_0))), inference(spm,[status(thm)],[c_0_1037, c_0_1031])).
cnf(c_0_1042,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk17_0,esk10_1(esk17_0))|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1035]), c_0_1036])).
cnf(c_0_1043,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk15_2(esk17_0,X1))|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1038]), c_0_1039])).
cnf(c_0_1044,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk17_0)|r1(esk3_1(esk17_0),esk4_1(esk17_0))), inference(spm,[status(thm)],[c_0_1037, c_0_1034])).
cnf(c_0_1045,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk10_1(esk17_0))|p2(esk17_0)|p2(X1)|p2(X2)|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1038]), c_0_1039])).
cnf(c_0_1046,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(X1,esk15_2(esk17_0,X1))|r1(esk17_0,esk10_1(esk17_0))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1040, c_0_1041])).
cnf(c_0_1047,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk17_0,esk10_1(esk17_0))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1042, c_0_1041])).
cnf(c_0_1048,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|r1(esk17_0,esk10_1(esk17_0))|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1008, c_0_1035]), c_0_1036])).
cnf(c_0_1049,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk10_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(X1,esk15_2(esk17_0,X1))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1043, c_0_1044])).
cnf(c_0_1050,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk4_1(esk17_0))|p2(esk10_1(esk17_0))|p2(esk17_0)|p2(X1)|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1045, c_0_1044])).
cnf(c_0_1051,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1008, c_0_1038]), c_0_1039])).
cnf(c_0_1052,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk32_1(esk17_0),esk15_2(esk17_0,esk32_1(esk17_0)))|r1(esk17_0,esk10_1(esk17_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1046, c_0_1010]), c_0_1011])).
cnf(c_0_1053,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,esk32_1(esk17_0)))|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk17_0,esk10_1(esk17_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1047, c_0_1010]), c_0_1011])).
cnf(c_0_1054,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|r1(esk17_0,esk10_1(esk17_0))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1048, c_0_1041])).
cnf(c_0_1055,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk32_1(esk17_0),esk15_2(esk17_0,esk32_1(esk17_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1049, c_0_1010]), c_0_1011])).
cnf(c_0_1056,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,esk32_1(esk17_0)))|p2(esk10_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1050, c_0_1010]), c_0_1011])).
cnf(c_0_1057,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk10_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1051, c_0_1044])).
cnf(c_0_1058,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk17_0,esk10_1(esk17_0))|~r1(esk15_2(esk17_0,esk32_1(esk17_0)),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1052]), c_0_994])]), c_0_1053])).
cnf(c_0_1059,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk15_2(esk17_0,esk32_1(esk17_0)),esk16_2(esk17_0,esk32_1(esk17_0)))|r1(esk17_0,esk10_1(esk17_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1054, c_0_1010]), c_0_1011])).
cnf(c_0_1060,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk10_1(esk17_0))|p2(esk17_0)|p2(X1)|~r1(esk15_2(esk17_0,esk32_1(esk17_0)),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1055]), c_0_994])]), c_0_1056])).
cnf(c_0_1061,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk15_2(esk17_0,esk32_1(esk17_0)),esk16_2(esk17_0,esk32_1(esk17_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1057, c_0_1010]), c_0_1011])).
cnf(c_0_1062,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk17_0,esk32_1(esk17_0)))|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk17_0,esk10_1(esk17_0))), inference(spm,[status(thm)],[c_0_1058, c_0_1059])).
cnf(c_0_1063,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk17_0,esk32_1(esk17_0)))|p2(esk10_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1060, c_0_1061])).
cnf(c_0_1064,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk17_0,esk10_1(esk17_0))|~p2(X2)|~r1(esk17_0,X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1019, c_0_1062]), c_0_994])]), c_0_1010]), c_0_1011])).
cnf(c_0_1065,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk10_1(esk17_0))|p2(esk17_0)|p2(X1)|~p2(X2)|~r1(esk17_0,X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1019, c_0_1063]), c_0_994])]), c_0_1010]), c_0_1011])).
cnf(c_0_1066,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk17_0,esk10_1(esk17_0))|~r1(esk3_1(esk17_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1064, c_0_1035]), c_0_1036])).
cnf(c_0_1067,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|~r1(esk3_1(esk17_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1065, c_0_1038]), c_0_1039])).
cnf(c_0_1068,negated_conjecture,(p2(esk17_0)|~p2(esk4_1(esk17_0))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_69, c_0_994])).
cnf(c_0_1069,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk17_0,esk10_1(esk17_0))), inference(spm,[status(thm)],[c_0_1066, c_0_1041])).
cnf(c_0_1070,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk10_1(esk17_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1067, c_0_1044])).
cnf(c_0_1071,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk17_0,esk10_1(esk17_0))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1068, c_0_1069])).
cnf(c_0_1072,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1068, c_0_1070])).
cnf(c_0_1073,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk17_0,esk10_1(esk17_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1071, c_0_1026]), c_0_1031])).
cnf(c_0_1074,negated_conjecture,(epred1_1(esk1_0)|p2(esk10_1(esk17_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1072, c_0_1026]), c_0_1034])).
cnf(c_0_1075,negated_conjecture,(p2(esk17_0)|p5(esk8_1(X1))|r1(esk10_1(esk17_0),esk11_1(esk17_0))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_432, c_0_994])).
cnf(c_0_1076,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk15_2(esk17_0,X1))|~r1(esk10_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1073]), c_0_1074])).
cnf(c_0_1077,negated_conjecture,(p2(esk17_0)|p5(esk8_1(esk1_0))|r1(esk10_1(esk17_0),esk11_1(esk17_0))), inference(spm,[status(thm)],[c_0_1075, c_0_7])).
cnf(c_0_1078,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk17_0)|p2(X1)|p2(X2)|~r1(esk10_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1073]), c_0_1074])).
cnf(c_0_1079,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(X1,esk15_2(esk17_0,X1))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1076, c_0_1077])).
cnf(c_0_1080,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1078, c_0_1077])).
cnf(c_0_1081,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|~r1(esk10_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1008, c_0_1073]), c_0_1074])).
cnf(c_0_1082,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p5(esk8_1(esk1_0))|r1(esk32_1(esk17_0),esk15_2(esk17_0,esk32_1(esk17_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1079, c_0_1010]), c_0_1011])).
cnf(c_0_1083,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,esk32_1(esk17_0)))|p2(esk11_1(esk17_0))|p2(esk17_0)|p5(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1080, c_0_1010]), c_0_1011])).
cnf(c_0_1084,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1081, c_0_1077])).
cnf(c_0_1085,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|~r1(esk15_2(esk17_0,esk32_1(esk17_0)),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1082]), c_0_994])]), c_0_1083])).
cnf(c_0_1086,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p5(esk8_1(esk1_0))|r1(esk15_2(esk17_0,esk32_1(esk17_0)),esk16_2(esk17_0,esk32_1(esk17_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1084, c_0_1010]), c_0_1011])).
cnf(c_0_1087,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk17_0,esk32_1(esk17_0)))|p2(esk11_1(esk17_0))|p2(esk17_0)|p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1085, c_0_1086])).
cnf(c_0_1088,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|~p2(X2)|~r1(esk17_0,X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1019, c_0_1087]), c_0_994])]), c_0_1010]), c_0_1011])).
cnf(c_0_1089,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|~r1(esk10_1(esk17_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1088, c_0_1073]), c_0_1034])).
cnf(c_0_1090,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1089, c_0_1077])).
cnf(c_0_1091,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|r1(esk17_0,esk3_1(esk17_0))), inference(spm,[status(thm)],[c_0_1030, c_0_1090])).
cnf(c_0_1092,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk3_1(esk17_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1032, c_0_1090])).
cnf(c_0_1093,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk15_2(esk17_0,X1))|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1091]), c_0_1092])).
cnf(c_0_1094,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|r1(esk3_1(esk17_0),esk4_1(esk17_0))), inference(spm,[status(thm)],[c_0_1037, c_0_1090])).
cnf(c_0_1095,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|p2(X2)|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1091]), c_0_1092])).
cnf(c_0_1096,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(X1,esk15_2(esk17_0,X1))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1093, c_0_1094])).
cnf(c_0_1097,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk4_1(esk17_0))|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1095, c_0_1094])).
cnf(c_0_1098,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1008, c_0_1091]), c_0_1092])).
cnf(c_0_1099,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk32_1(esk17_0),esk15_2(esk17_0,esk32_1(esk17_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1096, c_0_1010]), c_0_1011])).
cnf(c_0_1100,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,esk32_1(esk17_0)))|p2(esk11_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1097, c_0_1010]), c_0_1011])).
cnf(c_0_1101,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1098, c_0_1094])).
cnf(c_0_1102,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|~r1(esk15_2(esk17_0,esk32_1(esk17_0)),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1099]), c_0_994])]), c_0_1100])).
cnf(c_0_1103,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk15_2(esk17_0,esk32_1(esk17_0)),esk16_2(esk17_0,esk32_1(esk17_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1101, c_0_1010]), c_0_1011])).
cnf(c_0_1104,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk17_0,esk32_1(esk17_0)))|p2(esk11_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1102, c_0_1103])).
cnf(c_0_1105,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk11_1(esk17_0))|p2(esk17_0)|p2(X1)|~p2(X2)|~r1(esk17_0,X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1019, c_0_1104]), c_0_994])]), c_0_1010]), c_0_1011])).
cnf(c_0_1106,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|~r1(esk3_1(esk17_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1105, c_0_1091]), c_0_1092])).
cnf(c_0_1107,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk11_1(esk17_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1106, c_0_1094])).
cnf(c_0_1108,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1068, c_0_1107])).
cnf(c_0_1109,negated_conjecture,(p2(esk17_0)|p5(esk8_1(X1))|~p2(esk11_1(esk17_0))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_481, c_0_994])).
cnf(c_0_1110,negated_conjecture,(epred1_1(esk1_0)|p2(esk11_1(esk17_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1108, c_0_1026]), c_0_1090])).
cnf(c_0_1111,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p5(esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1109, c_0_1110])).
cnf(c_0_1112,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1111, c_0_7])).
cnf(c_0_1113,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk17_0,esk3_1(esk17_0))), inference(spm,[status(thm)],[c_0_1030, c_0_1112])).
cnf(c_0_1114,negated_conjecture,(epred1_1(esk1_0)|p2(esk3_1(esk17_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1032, c_0_1112])).
cnf(c_0_1115,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk15_2(esk17_0,X1))|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_999, c_0_1113]), c_0_1114])).
cnf(c_0_1116,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|r1(esk3_1(esk17_0),esk4_1(esk17_0))), inference(spm,[status(thm)],[c_0_1037, c_0_1112])).
cnf(c_0_1117,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk17_0)|p2(X1)|p2(X2)|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1113]), c_0_1114])).
cnf(c_0_1118,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(X1,esk15_2(esk17_0,X1))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1115, c_0_1116])).
cnf(c_0_1119,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,X1))|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1117, c_0_1116])).
cnf(c_0_1120,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|~r1(esk3_1(esk17_0),X2)|~r1(esk17_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1008, c_0_1113]), c_0_1114])).
cnf(c_0_1121,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk32_1(esk17_0),esk15_2(esk17_0,esk32_1(esk17_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1118, c_0_1010]), c_0_1011])).
cnf(c_0_1122,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk17_0,esk32_1(esk17_0)))|p2(esk4_1(esk17_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1119, c_0_1010]), c_0_1011])).
cnf(c_0_1123,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|r1(esk15_2(esk17_0,X1),esk16_2(esk17_0,X1))|~r1(esk17_0,X1)), inference(spm,[status(thm)],[c_0_1120, c_0_1116])).
cnf(c_0_1124,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|~r1(esk15_2(esk17_0,esk32_1(esk17_0)),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1121]), c_0_994])]), c_0_1122])).
cnf(c_0_1125,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|r1(esk15_2(esk17_0,esk32_1(esk17_0)),esk16_2(esk17_0,esk32_1(esk17_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1123, c_0_1010]), c_0_1011])).
cnf(c_0_1126,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk17_0,esk32_1(esk17_0)))|p2(esk4_1(esk17_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1124, c_0_1125])).
cnf(c_0_1127,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|~p2(X2)|~r1(esk17_0,X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1019, c_0_1126]), c_0_994])]), c_0_1010]), c_0_1011])).
cnf(c_0_1128,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)|p2(X1)|~r1(esk3_1(esk17_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1127, c_0_1113]), c_0_1114])).
cnf(c_0_1129,negated_conjecture,(epred1_1(esk1_0)|p2(esk4_1(esk17_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1128, c_0_1116])).
cnf(c_0_1130,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1068, c_0_1129])).
cnf(c_0_1131,negated_conjecture,(epred1_1(esk1_0)|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1130, c_0_1026]), c_0_1112])).
cnf(c_0_1132,plain,(p2(esk17_0)|r1(esk1_0,esk33_1(esk1_0))), inference(spm,[status(thm)],[c_0_32, c_0_1131])).
cnf(c_0_1133,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_348, c_0_1132])).
cnf(c_0_1134,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_351, c_0_1132])).
cnf(c_0_1135,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1133, c_0_7])).
cnf(c_0_1136,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1134, c_0_7])).
cnf(c_0_1137,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_356, c_0_1132])).
cnf(c_0_1138,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|r1(X2,esk36_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1135]), c_0_1136]), c_0_1131])).
cnf(c_0_1139,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1137, c_0_7])).
cnf(c_0_1140,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X2)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1135]), c_0_1136]), c_0_1131])).
cnf(c_0_1141,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk8_1(esk1_0))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1138, c_0_1139])).
cnf(c_0_1142,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_14, c_0_1132])).
cnf(c_0_1143,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|~p2(esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_15, c_0_1132])).
cnf(c_0_1144,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1140, c_0_1139])).
cnf(c_0_1145,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk8_1(esk1_0))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1135]), c_0_1136]), c_0_1131])).
cnf(c_0_1146,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|~p2(esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1135]), c_0_1136]), c_0_1131])).
cnf(c_0_1147,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1141, c_0_1142]), c_0_1143])).
cnf(c_0_1148,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1144, c_0_1142]), c_0_1143])).
cnf(c_0_1149,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk8_1(esk1_0))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1145, c_0_1139])).
cnf(c_0_1150,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1146, c_0_1139])).
cnf(c_0_1151,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1147]), c_0_1132]), c_0_1148])).
cnf(c_0_1152,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1149, c_0_1142]), c_0_1143])).
cnf(c_0_1153,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1150, c_0_1142]), c_0_1143])).
cnf(c_0_1154,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1151, c_0_1152]), c_0_1153])).
cnf(c_0_1155,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_1154]), c_0_1132])).
cnf(c_0_1156,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_34, c_0_1132])).
cnf(c_0_1157,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1155, c_0_7])).
cnf(c_0_1158,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(X1))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_381, c_0_1132])).
cnf(c_0_1159,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_37, c_0_1132])).
cnf(c_0_1160,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_383, c_0_1132])).
cnf(c_0_1161,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1156, c_0_1157])).
cnf(c_0_1162,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(esk1_0))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1158, c_0_7])).
cnf(c_0_1163,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1159, c_0_1157])).
cnf(c_0_1164,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_1132])).
cnf(c_0_1165,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1160, c_0_7])).
cnf(c_0_1166,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1161, c_0_1162])).
cnf(c_0_1167,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1163, c_0_1162])).
cnf(c_0_1168,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1164, c_0_1157])).
cnf(c_0_1169,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1161, c_0_1165])).
cnf(c_0_1170,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1163, c_0_1165])).
cnf(c_0_1171,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(X2,esk36_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1166]), c_0_1167]), c_0_1131])).
cnf(c_0_1172,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1168, c_0_1162])).
cnf(c_0_1173,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X2)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1166]), c_0_1167]), c_0_1131])).
cnf(c_0_1174,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1169]), c_0_1170]), c_0_1131])).
cnf(c_0_1175,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1168, c_0_1165])).
cnf(c_0_1176,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X2)|p2(X1)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1169]), c_0_1170]), c_0_1131])).
cnf(c_0_1177,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1171, c_0_1172])).
cnf(c_0_1178,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1173, c_0_1172])).
cnf(c_0_1179,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1166]), c_0_1167]), c_0_1131])).
cnf(c_0_1180,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1166]), c_0_1167]), c_0_1131])).
cnf(c_0_1181,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1174, c_0_1175])).
cnf(c_0_1182,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1176, c_0_1175])).
cnf(c_0_1183,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1169]), c_0_1170]), c_0_1131])).
cnf(c_0_1184,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1169]), c_0_1170]), c_0_1131])).
cnf(c_0_1185,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1177, c_0_1142]), c_0_1143])).
cnf(c_0_1186,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1178, c_0_1142]), c_0_1143])).
cnf(c_0_1187,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1179, c_0_1172])).
cnf(c_0_1188,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1180, c_0_1172])).
cnf(c_0_1189,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1181, c_0_1142]), c_0_1143])).
cnf(c_0_1190,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1182, c_0_1142]), c_0_1143])).
cnf(c_0_1191,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1183, c_0_1175])).
cnf(c_0_1192,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1184, c_0_1175])).
cnf(c_0_1193,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1185]), c_0_1132]), c_0_1186])).
cnf(c_0_1194,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1187, c_0_1142]), c_0_1143])).
cnf(c_0_1195,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1188, c_0_1142]), c_0_1143])).
cnf(c_0_1196,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1189]), c_0_1132]), c_0_1190])).
cnf(c_0_1197,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1191, c_0_1142]), c_0_1143])).
cnf(c_0_1198,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1192, c_0_1142]), c_0_1143])).
cnf(c_0_1199,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|~p2(esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_69, c_0_1132])).
cnf(c_0_1200,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1193, c_0_1194]), c_0_1195])).
cnf(c_0_1201,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1196, c_0_1197]), c_0_1198])).
cnf(c_0_1202,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1199, c_0_1200])).
cnf(c_0_1203,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1199, c_0_1201])).
cnf(c_0_1204,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1202, c_0_1157]), c_0_1162])).
cnf(c_0_1205,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1203, c_0_1157]), c_0_1165])).
cnf(c_0_1206,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(X1))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_432, c_0_1132])).
cnf(c_0_1207,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1204]), c_0_1205]), c_0_1131])).
cnf(c_0_1208,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(esk1_0))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1206, c_0_7])).
cnf(c_0_1209,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X2)|p2(X1)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1204]), c_0_1205]), c_0_1131])).
cnf(c_0_1210,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1207, c_0_1208])).
cnf(c_0_1211,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1209, c_0_1208])).
cnf(c_0_1212,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1204]), c_0_1205]), c_0_1131])).
cnf(c_0_1213,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1204]), c_0_1205]), c_0_1131])).
cnf(c_0_1214,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1167]), c_0_1170]), c_0_1131])).
cnf(c_0_1215,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X2)|p2(X1)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1167]), c_0_1170]), c_0_1131])).
cnf(c_0_1216,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(esk1_0))|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1210, c_0_1142]), c_0_1143])).
cnf(c_0_1217,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1211, c_0_1142]), c_0_1143])).
cnf(c_0_1218,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1212, c_0_1208])).
cnf(c_0_1219,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1213, c_0_1208])).
cnf(c_0_1220,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1214, c_0_1208]), c_0_1163])).
cnf(c_0_1221,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1215, c_0_1208]), c_0_1163])).
cnf(c_0_1222,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1167]), c_0_1170]), c_0_1131])).
cnf(c_0_1223,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1167]), c_0_1170]), c_0_1131])).
cnf(c_0_1224,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p5(esk8_1(esk1_0))|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1216]), c_0_1132]), c_0_1217])).
cnf(c_0_1225,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(esk1_0))|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1218, c_0_1142]), c_0_1143])).
cnf(c_0_1226,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(esk1_0))|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1219, c_0_1142]), c_0_1143])).
cnf(c_0_1227,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1220, c_0_1142]), c_0_1143])).
cnf(c_0_1228,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1221, c_0_1142]), c_0_1143])).
cnf(c_0_1229,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1222, c_0_1208]), c_0_1163])).
cnf(c_0_1230,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1223, c_0_1208]), c_0_1163])).
cnf(c_0_1231,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1224, c_0_1225]), c_0_1226])).
cnf(c_0_1232,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1227]), c_0_1132]), c_0_1228])).
cnf(c_0_1233,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1229, c_0_1142]), c_0_1143])).
cnf(c_0_1234,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1230, c_0_1142]), c_0_1143])).
cnf(c_0_1235,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1161, c_0_1231])).
cnf(c_0_1236,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1232, c_0_1233]), c_0_1234])).
cnf(c_0_1237,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1235]), c_0_1236]), c_0_1131])).
cnf(c_0_1238,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1168, c_0_1231])).
cnf(c_0_1239,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X2)|p2(X1)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1235]), c_0_1236]), c_0_1131])).
cnf(c_0_1240,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1237, c_0_1238])).
cnf(c_0_1241,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1239, c_0_1238])).
cnf(c_0_1242,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1235]), c_0_1236]), c_0_1131])).
cnf(c_0_1243,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1235]), c_0_1236]), c_0_1131])).
cnf(c_0_1244,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1240, c_0_1142]), c_0_1143])).
cnf(c_0_1245,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1241, c_0_1142]), c_0_1143])).
cnf(c_0_1246,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1242, c_0_1238])).
cnf(c_0_1247,plain,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1243, c_0_1238])).
cnf(c_0_1248,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1244]), c_0_1132]), c_0_1245])).
cnf(c_0_1249,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1246, c_0_1142]), c_0_1143])).
cnf(c_0_1250,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1247, c_0_1142]), c_0_1143])).
cnf(c_0_1251,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1248, c_0_1249]), c_0_1250])).
cnf(c_0_1252,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1199, c_0_1251])).
cnf(c_0_1253,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(X1))|~p2(esk11_1(esk33_1(esk1_0)))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_481, c_0_1132])).
cnf(c_0_1254,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1252, c_0_1157]), c_0_1231])).
cnf(c_0_1255,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(X1))|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1253, c_0_1254])).
cnf(c_0_1256,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1255, c_0_7])).
cnf(c_0_1257,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1161, c_0_1256])).
cnf(c_0_1258,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1163, c_0_1256])).
cnf(c_0_1259,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1257]), c_0_1258]), c_0_1131])).
cnf(c_0_1260,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1168, c_0_1256])).
cnf(c_0_1261,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X2)|p2(X1)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1257]), c_0_1258]), c_0_1131])).
cnf(c_0_1262,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1259, c_0_1260])).
cnf(c_0_1263,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1261, c_0_1260])).
cnf(c_0_1264,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1257]), c_0_1258]), c_0_1131])).
cnf(c_0_1265,plain,(p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1257]), c_0_1258]), c_0_1131])).
cnf(c_0_1266,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1262, c_0_1142]), c_0_1143])).
cnf(c_0_1267,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1263, c_0_1142]), c_0_1143])).
cnf(c_0_1268,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1264, c_0_1260])).
cnf(c_0_1269,plain,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1265, c_0_1260])).
cnf(c_0_1270,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|p2(X1)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1266]), c_0_1132]), c_0_1267])).
cnf(c_0_1271,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1268, c_0_1142]), c_0_1143])).
cnf(c_0_1272,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1269, c_0_1142]), c_0_1143])).
cnf(c_0_1273,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1270, c_0_1271]), c_0_1272])).
cnf(c_0_1274,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1199, c_0_1273])).
cnf(c_0_1275,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1274, c_0_1157]), c_0_1256])).
cnf(c_0_1276,plain,(p2(esk17_0)|r1(esk33_1(esk1_0),esk34_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_1275]), c_0_1131])).
cnf(c_0_1277,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_1276]), c_0_1131])).
cnf(c_0_1278,plain,(p2(esk34_1(esk1_0))|p2(esk17_0)|r1(esk34_1(esk1_0),esk36_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_1276]), c_0_1275]), c_0_1131])).
cnf(c_0_1279,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk34_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_1276]), c_0_1275]), c_0_1131])).
cnf(c_0_1280,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_1276]), c_0_1131])).
cnf(c_0_1281,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk17_0)|p2(X3)|p2(X2)|p2(X1)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_1276]), c_0_1131])).
cnf(c_0_1282,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk17_0)|p2(X2)|p2(X1)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_1276]), c_0_1131])).
cnf(c_0_1283,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1277, c_0_1278]), c_0_1279]), c_0_1280])).
cnf(c_0_1284,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk17_0)|p2(X2)|p2(X1)|p2(X3)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1281, c_0_1278]), c_0_1279]), c_0_1282])).
cnf(c_0_1285,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_1276]), c_0_1131])).
cnf(c_0_1286,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_1276]), c_0_1131])).
cnf(c_0_1287,plain,(p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_1276]), c_0_1131])).
cnf(c_0_1288,plain,(p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_1276]), c_0_1131])).
cnf(c_0_1289,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk45_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_1276]), c_0_1131])).
cnf(c_0_1290,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_1276]), c_0_1131])).
cnf(c_0_1291,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1283, c_0_7])).
cnf(c_0_1292,plain,(p2(esk17_0)|r1(esk34_1(esk1_0),esk35_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_1275]), c_0_1131])).
cnf(c_0_1293,plain,(p2(esk17_0)|~p2(esk35_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_1275]), c_0_1131])).
cnf(c_0_1294,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X2)|p2(X1)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1284, c_0_7])).
cnf(c_0_1295,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1285, c_0_1278]), c_0_1279]), c_0_1286])).
cnf(c_0_1296,plain,(p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1287, c_0_1278]), c_0_1279]), c_0_1288])).
cnf(c_0_1297,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1289, c_0_1278]), c_0_1279]), c_0_1290])).
cnf(c_0_1298,plain,(p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_1276]), c_0_1131])).
cnf(c_0_1299,plain,(p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_1276]), c_0_1131])).
cnf(c_0_1300,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1291, c_0_1292]), c_0_1293])).
cnf(c_0_1301,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1294, c_0_1292]), c_0_1293])).
cnf(c_0_1302,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1295, c_0_7])).
cnf(c_0_1303,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1296, c_0_7])).
cnf(c_0_1304,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1297, c_0_7])).
cnf(c_0_1305,plain,(p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1298, c_0_1278]), c_0_1279]), c_0_1299])).
cnf(c_0_1306,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1300, c_0_1292]), c_0_1293])).
cnf(c_0_1307,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1301, c_0_1292]), c_0_1293])).
cnf(c_0_1308,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1302, c_0_1292]), c_0_1293])).
cnf(c_0_1309,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1303, c_0_1292]), c_0_1293])).
cnf(c_0_1310,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1304, c_0_1292]), c_0_1293])).
cnf(c_0_1311,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1305, c_0_7])).
cnf(c_0_1312,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1306]), c_0_1307]), c_0_1275]), c_0_1131])).
cnf(c_0_1313,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1308, c_0_1292]), c_0_1293])).
cnf(c_0_1314,plain,(p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_1276]), c_0_1131])).
cnf(c_0_1315,plain,(p2(esk17_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_1276]), c_0_1131])).
cnf(c_0_1316,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X3)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X3),esk46_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_1276]), c_0_1131])).
cnf(c_0_1317,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_1276]), c_0_1131])).
cnf(c_0_1318,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1309, c_0_1292]), c_0_1293])).
cnf(c_0_1319,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1310, c_0_1292]), c_0_1293])).
cnf(c_0_1320,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1311, c_0_1292]), c_0_1293])).
cnf(c_0_1321,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1312, c_0_1313])).
cnf(c_0_1322,plain,(p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1314, c_0_1278]), c_0_1279]), c_0_1315])).
cnf(c_0_1323,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1316, c_0_1278]), c_0_1279]), c_0_1317])).
cnf(c_0_1324,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1318]), c_0_1319]), c_0_1275]), c_0_1131])).
cnf(c_0_1325,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1320, c_0_1292]), c_0_1293])).
cnf(c_0_1326,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_1321]), c_0_1292]), c_0_1276]), c_0_1279]), c_0_1131]), c_0_1293])).
cnf(c_0_1327,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1322, c_0_7])).
cnf(c_0_1328,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1323, c_0_7])).
cnf(c_0_1329,plain,(p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_1276]), c_0_1131])).
cnf(c_0_1330,plain,(p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_1276]), c_0_1131])).
cnf(c_0_1331,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1324, c_0_1325])).
cnf(c_0_1332,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1326, c_0_1292]), c_0_1293])).
cnf(c_0_1333,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1327, c_0_1292]), c_0_1293])).
cnf(c_0_1334,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1328, c_0_1292]), c_0_1293])).
cnf(c_0_1335,plain,(p2(esk17_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1329, c_0_1278]), c_0_1279]), c_0_1330])).
cnf(c_0_1336,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1331]), c_0_1332]), c_0_1275]), c_0_1131])).
cnf(c_0_1337,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1333, c_0_1292]), c_0_1293])).
cnf(c_0_1338,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1319]), c_0_1307]), c_0_1275]), c_0_1131])).
cnf(c_0_1339,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1334, c_0_1292]), c_0_1293])).
cnf(c_0_1340,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1335, c_0_7])).
cnf(c_0_1341,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1336, c_0_1337])).
cnf(c_0_1342,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1338, c_0_1339])).
cnf(c_0_1343,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1340, c_0_1292]), c_0_1293])).
cnf(c_0_1344,plain,(p2(esk34_1(esk1_0))|p2(esk17_0)|r1(esk36_2(esk1_0,esk34_1(esk1_0)),esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_1276]), c_0_1275]), c_0_1131])).
cnf(c_0_1345,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1341]), c_0_1342]), c_0_1275]), c_0_1131])).
cnf(c_0_1346,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1343, c_0_1292]), c_0_1293])).
cnf(c_0_1347,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1296, c_0_1344]), c_0_1288])).
cnf(c_0_1348,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1297, c_0_1344]), c_0_1290])).
cnf(c_0_1349,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1283, c_0_1344]), c_0_1280])).
cnf(c_0_1350,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X2)|p2(X1)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1284, c_0_1344]), c_0_1282])).
cnf(c_0_1351,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1345, c_0_1346]), c_0_1336])).
cnf(c_0_1352,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1347, c_0_1292]), c_0_1293])).
cnf(c_0_1353,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1348, c_0_1292]), c_0_1293])).
cnf(c_0_1354,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1305, c_0_1344]), c_0_1299])).
cnf(c_0_1355,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1349, c_0_1292]), c_0_1293])).
cnf(c_0_1356,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1350, c_0_1292]), c_0_1293])).
cnf(c_0_1357,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1295, c_0_1344]), c_0_1286])).
cnf(c_0_1358,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_1351]), c_0_1292]), c_0_1276]), c_0_1279]), c_0_1131]), c_0_1293])).
cnf(c_0_1359,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1352, c_0_1292]), c_0_1293])).
cnf(c_0_1360,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1353, c_0_1292]), c_0_1293])).
cnf(c_0_1361,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1354, c_0_1292]), c_0_1293])).
cnf(c_0_1362,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1355, c_0_1292]), c_0_1293])).
cnf(c_0_1363,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1356, c_0_1292]), c_0_1293])).
cnf(c_0_1364,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1357, c_0_1292]), c_0_1293])).
cnf(c_0_1365,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1358, c_0_1292]), c_0_1293])).
cnf(c_0_1366,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_1351]), c_0_1292]), c_0_1276]), c_0_1279]), c_0_1131]), c_0_1293])).
cnf(c_0_1367,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1359]), c_0_1360]), c_0_1275]), c_0_1131])).
cnf(c_0_1368,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1361, c_0_1292]), c_0_1293])).
cnf(c_0_1369,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1362]), c_0_1363]), c_0_1275]), c_0_1131])).
cnf(c_0_1370,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1364, c_0_1292]), c_0_1293])).
cnf(c_0_1371,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1322, c_0_1344]), c_0_1315])).
cnf(c_0_1372,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1323, c_0_1344]), c_0_1317])).
cnf(c_0_1373,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1365]), c_0_1332]), c_0_1275]), c_0_1131])).
cnf(c_0_1374,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1366, c_0_1292]), c_0_1293])).
cnf(c_0_1375,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1367, c_0_1368])).
cnf(c_0_1376,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1369, c_0_1370])).
cnf(c_0_1377,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1371, c_0_1292]), c_0_1293])).
cnf(c_0_1378,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1372, c_0_1292]), c_0_1293])).
cnf(c_0_1379,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1373, c_0_1374])).
cnf(c_0_1380,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1375]), c_0_1376]), c_0_1275]), c_0_1131])).
cnf(c_0_1381,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1377, c_0_1292]), c_0_1293])).
cnf(c_0_1382,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1360]), c_0_1363]), c_0_1275]), c_0_1131])).
cnf(c_0_1383,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1378, c_0_1292]), c_0_1293])).
cnf(c_0_1384,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1335, c_0_1344]), c_0_1330])).
cnf(c_0_1385,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_1379]), c_0_1292]), c_0_1276]), c_0_1279]), c_0_1131]), c_0_1293])).
cnf(c_0_1386,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_1379]), c_0_1292]), c_0_1276]), c_0_1279]), c_0_1131]), c_0_1293])).
cnf(c_0_1387,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1380, c_0_1381])).
cnf(c_0_1388,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1382, c_0_1383])).
cnf(c_0_1389,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1384, c_0_1292]), c_0_1293])).
cnf(c_0_1390,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1385, c_0_1292]), c_0_1293])).
cnf(c_0_1391,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1386, c_0_1292]), c_0_1293])).
cnf(c_0_1392,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_1379]), c_0_1292]), c_0_1276]), c_0_1279]), c_0_1131]), c_0_1293])).
cnf(c_0_1393,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1387]), c_0_1388]), c_0_1275]), c_0_1131])).
cnf(c_0_1394,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1389, c_0_1292]), c_0_1293])).
cnf(c_0_1395,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1390]), c_0_1391]), c_0_1275]), c_0_1131])).
cnf(c_0_1396,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1392, c_0_1292]), c_0_1293])).
cnf(c_0_1397,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1393, c_0_1394]), c_0_1380])).
cnf(c_0_1398,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1395, c_0_1396])).
cnf(c_0_1399,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_1397]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1400,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_1397]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1401,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X2)|p2(X1)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_1376]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1402,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_1376]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1403,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_1398]), c_0_1292]), c_0_1276]), c_0_1279]), c_0_1131]), c_0_1293])).
cnf(c_0_1404,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1399, c_0_1278]), c_0_1279]), c_0_1400])).
cnf(c_0_1405,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1401, c_0_1278]), c_0_1279]), c_0_1402])).
cnf(c_0_1406,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_1397]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1407,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1403, c_0_1379]), c_0_1292]), c_0_1293])).
cnf(c_0_1408,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_1397]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1409,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1404, c_0_1344]), c_0_1400])).
cnf(c_0_1410,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1405, c_0_1344]), c_0_1402])).
cnf(c_0_1411,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1406, c_0_1278]), c_0_1407]), c_0_1408])).
cnf(c_0_1412,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1409, c_0_1292]), c_0_1293])).
cnf(c_0_1413,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1410, c_0_1292]), c_0_1293])).
cnf(c_0_1414,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1411, c_0_1344]), c_0_1408])).
cnf(c_0_1415,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1412]), c_0_1413]), c_0_1275]), c_0_1131])).
cnf(c_0_1416,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1414, c_0_1292]), c_0_1293])).
cnf(c_0_1417,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1415, c_0_1416])).
cnf(c_0_1418,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_1417]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1419,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_1417]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1420,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X2)|p2(X1)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_1417]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1421,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_1417]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1422,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1418, c_0_1278]), c_0_1407]), c_0_1419])).
cnf(c_0_1423,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1420, c_0_1278]), c_0_1407]), c_0_1421])).
cnf(c_0_1424,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_1417]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1425,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_1417]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1426,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1422, c_0_1344]), c_0_1419])).
cnf(c_0_1427,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1423, c_0_1344]), c_0_1421])).
cnf(c_0_1428,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1424, c_0_1278]), c_0_1407]), c_0_1425])).
cnf(c_0_1429,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1426, c_0_1292]), c_0_1293])).
cnf(c_0_1430,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1427, c_0_1292]), c_0_1293])).
cnf(c_0_1431,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1428, c_0_1344]), c_0_1425])).
cnf(c_0_1432,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1429]), c_0_1430]), c_0_1275]), c_0_1131])).
cnf(c_0_1433,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1431, c_0_1292]), c_0_1293])).
cnf(c_0_1434,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1432, c_0_1433])).
cnf(c_0_1435,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|p2(X2)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_1434]), c_0_1292]), c_0_1276]), c_0_1131]), c_0_1293])).
cnf(c_0_1436,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk17_0)|p2(X1)|~p2(X2)|~r1(esk34_1(esk1_0),X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1435, c_0_1417]), c_0_1292]), c_0_1293])).
cnf(c_0_1437,plain,(p2(esk34_1(esk1_0))|p2(esk17_0)|~p2(esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_1276]), c_0_1275]), c_0_1131])).
cnf(c_0_1438,plain,(p2(esk34_1(esk1_0))|p2(esk17_0)|p2(X1)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1436, c_0_1278]), c_0_1407]), c_0_1437])).
cnf(c_0_1439,plain,(p2(esk34_1(esk1_0))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1438, c_0_1344]), c_0_1437])).
cnf(c_0_1440,plain,(p2(esk17_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1288, c_0_1439])).
cnf(c_0_1441,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1290, c_0_1439])).
cnf(c_0_1442,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X1)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1280, c_0_1439])).
cnf(c_0_1443,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk17_0)|p2(X2)|p2(X1)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1282, c_0_1439])).
cnf(c_0_1444,plain,(p2(esk17_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1440, c_0_1292]), c_0_1293])).
cnf(c_0_1445,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1441, c_0_1292]), c_0_1293])).
cnf(c_0_1446,plain,(p2(esk17_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1299, c_0_1439])).
cnf(c_0_1447,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1442, c_0_1292]), c_0_1293])).
cnf(c_0_1448,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1443, c_0_1292]), c_0_1293])).
cnf(c_0_1449,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X2)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1286, c_0_1439])).
cnf(c_0_1450,plain,(p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1444, c_0_1292]), c_0_1293])).
cnf(c_0_1451,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1445, c_0_1292]), c_0_1293])).
cnf(c_0_1452,plain,(p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1446, c_0_1292]), c_0_1293])).
cnf(c_0_1453,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1447, c_0_1292]), c_0_1293])).
cnf(c_0_1454,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1448, c_0_1292]), c_0_1293])).
cnf(c_0_1455,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1449, c_0_1292]), c_0_1293])).
cnf(c_0_1456,plain,(p2(esk17_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1450]), c_0_1451]), c_0_1275]), c_0_1131])).
cnf(c_0_1457,plain,(p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1452, c_0_1292]), c_0_1293])).
cnf(c_0_1458,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1453]), c_0_1454]), c_0_1275]), c_0_1131])).
cnf(c_0_1459,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1455, c_0_1292]), c_0_1293])).
cnf(c_0_1460,plain,(p2(esk17_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1315, c_0_1439])).
cnf(c_0_1461,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1317, c_0_1439])).
cnf(c_0_1462,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1456, c_0_1457])).
cnf(c_0_1463,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1458, c_0_1459])).
cnf(c_0_1464,plain,(p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1460, c_0_1292]), c_0_1293])).
cnf(c_0_1465,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1461, c_0_1292]), c_0_1293])).
cnf(c_0_1466,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1462]), c_0_1463]), c_0_1275]), c_0_1131])).
cnf(c_0_1467,plain,(p2(esk17_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1464, c_0_1292]), c_0_1293])).
cnf(c_0_1468,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1451]), c_0_1454]), c_0_1275]), c_0_1131])).
cnf(c_0_1469,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1465, c_0_1292]), c_0_1293])).
cnf(c_0_1470,plain,(p2(esk17_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1330, c_0_1439])).
cnf(c_0_1471,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1466, c_0_1467])).
cnf(c_0_1472,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1468, c_0_1469])).
cnf(c_0_1473,plain,(p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1470, c_0_1292]), c_0_1293])).
cnf(c_0_1474,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1471]), c_0_1472]), c_0_1275]), c_0_1131])).
cnf(c_0_1475,plain,(p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1473, c_0_1292]), c_0_1293])).
cnf(c_0_1476,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1474, c_0_1475]), c_0_1466])).
cnf(c_0_1477,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_1476]), c_0_1292]), c_0_1276]), c_0_1439]), c_0_1131]), c_0_1293])).
cnf(c_0_1478,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_1463]), c_0_1292]), c_0_1276]), c_0_1439]), c_0_1131]), c_0_1293])).
cnf(c_0_1479,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1477, c_0_1292]), c_0_1293])).
cnf(c_0_1480,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1478, c_0_1292]), c_0_1293])).
cnf(c_0_1481,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_1476]), c_0_1292]), c_0_1276]), c_0_1439]), c_0_1131]), c_0_1293])).
cnf(c_0_1482,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1479]), c_0_1480]), c_0_1275]), c_0_1131])).
cnf(c_0_1483,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1481, c_0_1292]), c_0_1293])).
cnf(c_0_1484,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1482, c_0_1483])).
cnf(c_0_1485,plain,(p2(esk17_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_1484]), c_0_1292]), c_0_1276]), c_0_1439]), c_0_1131]), c_0_1293])).
cnf(c_0_1486,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk17_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_1484]), c_0_1292]), c_0_1276]), c_0_1439]), c_0_1131]), c_0_1293])).
cnf(c_0_1487,plain,(p2(esk17_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1485, c_0_1292]), c_0_1293])).
cnf(c_0_1488,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1486, c_0_1292]), c_0_1293])).
cnf(c_0_1489,plain,(p2(esk17_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_1484]), c_0_1292]), c_0_1276]), c_0_1439]), c_0_1131]), c_0_1293])).
cnf(c_0_1490,plain,(p2(esk17_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1487]), c_0_1488]), c_0_1275]), c_0_1131])).
cnf(c_0_1491,plain,(p2(esk17_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1489, c_0_1292]), c_0_1293])).
cnf(c_0_1492,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk17_0)), inference(spm,[status(thm)],[c_0_1490, c_0_1491])).
cnf(c_0_1493,plain,(p2(esk17_0)|p2(X1)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_1492]), c_0_1292]), c_0_1276]), c_0_1439]), c_0_1131]), c_0_1293])).
cnf(c_0_1494,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(X2,esk15_2(X1,X2))|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk17_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_1495,plain,(p2(esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1493, c_0_1484]), c_0_1292]), c_0_1293])).
cnf(c_0_1496,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|p2(esk15_2(X1,X2))|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk17_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_1497,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk15_2(X2,X1))|~r1(esk1_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1494, c_0_1495])])).
cnf(c_0_1498,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(X1,X2))|p2(X2)|~r1(esk1_0,X1)|~r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1496, c_0_1495])])).
cnf(c_0_1499,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|r1(esk15_2(X1,X2),esk16_2(X1,X2))|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk17_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_1500,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk15_2(esk2_0,X1))|r1(esk1_0,esk9_0)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_1497, c_0_13])).
cnf(c_0_1501,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk2_0,X1))|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_1498, c_0_13])).
cnf(c_0_1502,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk15_2(X2,X1),esk16_2(X2,X1))|~r1(esk1_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1499, c_0_1495])])).
cnf(c_0_1503,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk32_1(esk2_0),esk15_2(esk2_0,esk32_1(esk2_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1500, c_0_18]), c_0_19])).
cnf(c_0_1504,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk2_0,esk32_1(esk2_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1501, c_0_18]), c_0_19])).
cnf(c_0_1505,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk15_2(esk2_0,X1),esk16_2(esk2_0,X1))|r1(esk1_0,esk9_0)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_1502, c_0_13])).
cnf(c_0_1506,negated_conjecture,(epred1_1(esk1_0)|p2(X2)|~r1(esk1_0,X1)|~r1(X1,X2)|~p2(esk16_2(X1,X2))|~p2(esk17_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_1507,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk15_2(esk2_0,esk32_1(esk2_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1503]), c_0_13]), c_0_1504])).
cnf(c_0_1508,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk15_2(esk2_0,esk32_1(esk2_0)),esk16_2(esk2_0,esk32_1(esk2_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1505, c_0_18]), c_0_19])).
cnf(c_0_1509,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|~p2(esk16_2(X2,X1))|~r1(esk1_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1506, c_0_1495])])).
cnf(c_0_1510,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk2_0,esk32_1(esk2_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1507, c_0_1508])).
cnf(c_0_1511,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1509, c_0_1510]), c_0_18]), c_0_13]), c_0_19])).
cnf(c_0_1512,plain,(p2(esk2_0)|r1(esk1_0,esk33_1(esk1_0))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_32, c_0_1511])).
cnf(c_0_1513,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_1512]), c_0_36])).
cnf(c_0_1514,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(esk1_0,esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1512]), c_0_36])).
cnf(c_0_1515,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1513, c_0_10]), c_0_11])).
cnf(c_0_1516,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1514, c_0_10]), c_0_11])).
cnf(c_0_1517,plain,(p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|r1(esk1_0,esk9_0)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1515]), c_0_1516]), c_0_43])).
cnf(c_0_1518,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_1512]), c_0_36])).
cnf(c_0_1519,plain,(p2(esk36_2(esk1_0,X1))|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1515]), c_0_1516]), c_0_46])).
cnf(c_0_1520,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk9_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1517, c_0_1511])).
cnf(c_0_1521,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1518, c_0_10]), c_0_11])).
cnf(c_0_1522,negated_conjecture,(p2(esk36_2(esk1_0,X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk9_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1519, c_0_1511])).
cnf(c_0_1523,plain,(p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk9_0)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1515]), c_0_1516]), c_0_51])).
cnf(c_0_1524,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk9_0)|~epred1_1(esk1_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1515]), c_0_1516]), c_0_53])).
cnf(c_0_1525,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk9_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1520, c_0_1521])).
cnf(c_0_1526,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk33_1(esk1_0),esk32_1(esk33_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_14, c_0_1512])).
cnf(c_0_1527,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk9_0)|~p2(esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_15, c_0_1512])).
cnf(c_0_1528,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1522, c_0_1521])).
cnf(c_0_1529,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk9_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1523, c_0_1511])).
cnf(c_0_1530,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk9_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1524, c_0_1511])).
cnf(c_0_1531,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1525, c_0_1526]), c_0_1527])).
cnf(c_0_1532,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1528, c_0_1526]), c_0_1527])).
cnf(c_0_1533,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk9_0)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1529, c_0_1521])).
cnf(c_0_1534,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1530, c_0_1521])).
cnf(c_0_1535,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1531]), c_0_1512]), c_0_1532])).
cnf(c_0_1536,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1533, c_0_1526]), c_0_1527])).
cnf(c_0_1537,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk9_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1534, c_0_1526]), c_0_1527])).
cnf(c_0_1538,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk9_0)|~p2(esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_1512]), c_0_36])).
cnf(c_0_1539,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1535, c_0_1536]), c_0_1537])).
cnf(c_0_1540,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1538, c_0_1539]), c_0_36])).
cnf(c_0_1541,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1540, c_0_10]), c_0_11])).
cnf(c_0_1542,plain,(r1(esk33_1(esk1_0),esk34_1(esk1_0))|r1(esk1_0,esk9_0)|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_76, c_0_1541])).
cnf(c_0_1543,negated_conjecture,(p2(esk2_0)|r1(esk33_1(esk1_0),esk34_1(esk1_0))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1542, c_0_1511])).
cnf(c_0_1544,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_1543]), c_0_1511])).
cnf(c_0_1545,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk34_1(esk1_0),esk36_2(esk1_0,esk34_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_1543]), c_0_1541]), c_0_1511])).
cnf(c_0_1546,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_1543]), c_0_1541]), c_0_1511])).
cnf(c_0_1547,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_1543]), c_0_1511])).
cnf(c_0_1548,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X3)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_1543]), c_0_1511])).
cnf(c_0_1549,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_1543]), c_0_1511])).
cnf(c_0_1550,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1544, c_0_1545]), c_0_1546]), c_0_1547])).
cnf(c_0_1551,plain,(r1(esk34_1(esk1_0),esk35_1(esk1_0))|r1(esk1_0,esk9_0)|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_88, c_0_1541])).
cnf(c_0_1552,plain,(r1(esk1_0,esk9_0)|~epred1_1(esk1_0)|~p2(esk35_1(esk1_0))), inference(spm,[status(thm)],[c_0_89, c_0_1541])).
cnf(c_0_1553,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1548, c_0_1545]), c_0_1546]), c_0_1549])).
cnf(c_0_1554,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_1543]), c_0_1511])).
cnf(c_0_1555,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_1543]), c_0_1511])).
cnf(c_0_1556,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_1543]), c_0_1511])).
cnf(c_0_1557,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_1543]), c_0_1511])).
cnf(c_0_1558,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk45_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_1543]), c_0_1511])).
cnf(c_0_1559,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_1543]), c_0_1511])).
cnf(c_0_1560,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1550, c_0_7])).
cnf(c_0_1561,negated_conjecture,(p2(esk2_0)|r1(esk34_1(esk1_0),esk35_1(esk1_0))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1551, c_0_1511])).
cnf(c_0_1562,negated_conjecture,(p2(esk2_0)|r1(esk1_0,esk9_0)|~p2(esk35_1(esk1_0))), inference(spm,[status(thm)],[c_0_1552, c_0_1511])).
cnf(c_0_1563,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1553, c_0_7])).
cnf(c_0_1564,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1554, c_0_1545]), c_0_1546]), c_0_1555])).
cnf(c_0_1565,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1556, c_0_1545]), c_0_1546]), c_0_1557])).
cnf(c_0_1566,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1558, c_0_1545]), c_0_1546]), c_0_1559])).
cnf(c_0_1567,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_1543]), c_0_1511])).
cnf(c_0_1568,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_1543]), c_0_1511])).
cnf(c_0_1569,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1560, c_0_1561]), c_0_1562])).
cnf(c_0_1570,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1563, c_0_1561]), c_0_1562])).
cnf(c_0_1571,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1564, c_0_7])).
cnf(c_0_1572,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1565, c_0_7])).
cnf(c_0_1573,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1566, c_0_7])).
cnf(c_0_1574,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1567, c_0_1545]), c_0_1546]), c_0_1568])).
cnf(c_0_1575,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1569, c_0_1561]), c_0_1562])).
cnf(c_0_1576,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1570, c_0_1561]), c_0_1562])).
cnf(c_0_1577,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1571, c_0_1561]), c_0_1562])).
cnf(c_0_1578,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1572, c_0_1561]), c_0_1562])).
cnf(c_0_1579,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1573, c_0_1561]), c_0_1562])).
cnf(c_0_1580,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1574, c_0_7])).
cnf(c_0_1581,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1575]), c_0_1576]), c_0_1541]), c_0_1511])).
cnf(c_0_1582,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1577, c_0_1561]), c_0_1562])).
cnf(c_0_1583,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_1543]), c_0_1511])).
cnf(c_0_1584,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_1543]), c_0_1511])).
cnf(c_0_1585,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X3)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X3),esk46_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_1543]), c_0_1511])).
cnf(c_0_1586,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_1543]), c_0_1511])).
cnf(c_0_1587,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1578, c_0_1561]), c_0_1562])).
cnf(c_0_1588,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1579, c_0_1561]), c_0_1562])).
cnf(c_0_1589,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1580, c_0_1561]), c_0_1562])).
cnf(c_0_1590,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1581, c_0_1582])).
cnf(c_0_1591,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1583, c_0_1545]), c_0_1546]), c_0_1584])).
cnf(c_0_1592,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1585, c_0_1545]), c_0_1546]), c_0_1586])).
cnf(c_0_1593,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1587]), c_0_1588]), c_0_1541]), c_0_1511])).
cnf(c_0_1594,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1589, c_0_1561]), c_0_1562])).
cnf(c_0_1595,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_1590]), c_0_1561]), c_0_1543]), c_0_1546]), c_0_1511]), c_0_1562])).
cnf(c_0_1596,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1591, c_0_7])).
cnf(c_0_1597,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1592, c_0_7])).
cnf(c_0_1598,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_1543]), c_0_1511])).
cnf(c_0_1599,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_1543]), c_0_1511])).
cnf(c_0_1600,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1593, c_0_1594])).
cnf(c_0_1601,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1595, c_0_1561]), c_0_1562])).
cnf(c_0_1602,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1596, c_0_1561]), c_0_1562])).
cnf(c_0_1603,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1597, c_0_1561]), c_0_1562])).
cnf(c_0_1604,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1598, c_0_1545]), c_0_1546]), c_0_1599])).
cnf(c_0_1605,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1600]), c_0_1601]), c_0_1541]), c_0_1511])).
cnf(c_0_1606,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1602, c_0_1561]), c_0_1562])).
cnf(c_0_1607,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1588]), c_0_1576]), c_0_1541]), c_0_1511])).
cnf(c_0_1608,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1603, c_0_1561]), c_0_1562])).
cnf(c_0_1609,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1604, c_0_7])).
cnf(c_0_1610,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1605, c_0_1606])).
cnf(c_0_1611,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1607, c_0_1608])).
cnf(c_0_1612,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1609, c_0_1561]), c_0_1562])).
cnf(c_0_1613,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1610]), c_0_1611]), c_0_1541]), c_0_1511])).
cnf(c_0_1614,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1612, c_0_1561]), c_0_1562])).
cnf(c_0_1615,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk36_2(esk1_0,esk34_1(esk1_0)),esk37_2(esk1_0,esk34_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_1543]), c_0_1541]), c_0_1511])).
cnf(c_0_1616,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1613, c_0_1614]), c_0_1605])).
cnf(c_0_1617,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1565, c_0_1615]), c_0_1557])).
cnf(c_0_1618,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1566, c_0_1615]), c_0_1559])).
cnf(c_0_1619,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1550, c_0_1615]), c_0_1547])).
cnf(c_0_1620,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1553, c_0_1615]), c_0_1549])).
cnf(c_0_1621,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_1616]), c_0_1561]), c_0_1543]), c_0_1546]), c_0_1511]), c_0_1562])).
cnf(c_0_1622,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1617, c_0_1561]), c_0_1562])).
cnf(c_0_1623,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1618, c_0_1561]), c_0_1562])).
cnf(c_0_1624,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1574, c_0_1615]), c_0_1568])).
cnf(c_0_1625,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1619, c_0_1561]), c_0_1562])).
cnf(c_0_1626,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1620, c_0_1561]), c_0_1562])).
cnf(c_0_1627,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1564, c_0_1615]), c_0_1555])).
cnf(c_0_1628,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1621, c_0_1561]), c_0_1562])).
cnf(c_0_1629,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_1616]), c_0_1561]), c_0_1543]), c_0_1546]), c_0_1511]), c_0_1562])).
cnf(c_0_1630,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1622, c_0_1561]), c_0_1562])).
cnf(c_0_1631,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1623, c_0_1561]), c_0_1562])).
cnf(c_0_1632,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1624, c_0_1561]), c_0_1562])).
cnf(c_0_1633,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1625, c_0_1561]), c_0_1562])).
cnf(c_0_1634,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1626, c_0_1561]), c_0_1562])).
cnf(c_0_1635,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1627, c_0_1561]), c_0_1562])).
cnf(c_0_1636,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1628]), c_0_1601]), c_0_1541]), c_0_1511])).
cnf(c_0_1637,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1629, c_0_1561]), c_0_1562])).
cnf(c_0_1638,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1630]), c_0_1631]), c_0_1541]), c_0_1511])).
cnf(c_0_1639,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1632, c_0_1561]), c_0_1562])).
cnf(c_0_1640,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1633]), c_0_1634]), c_0_1541]), c_0_1511])).
cnf(c_0_1641,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1635, c_0_1561]), c_0_1562])).
cnf(c_0_1642,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1591, c_0_1615]), c_0_1584])).
cnf(c_0_1643,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1592, c_0_1615]), c_0_1586])).
cnf(c_0_1644,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1636, c_0_1637])).
cnf(c_0_1645,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1638, c_0_1639])).
cnf(c_0_1646,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1640, c_0_1641])).
cnf(c_0_1647,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1642, c_0_1561]), c_0_1562])).
cnf(c_0_1648,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1643, c_0_1561]), c_0_1562])).
cnf(c_0_1649,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_1644]), c_0_1561]), c_0_1543]), c_0_1546]), c_0_1511]), c_0_1562])).
cnf(c_0_1650,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_1644]), c_0_1561]), c_0_1543]), c_0_1546]), c_0_1511]), c_0_1562])).
cnf(c_0_1651,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1645]), c_0_1646]), c_0_1541]), c_0_1511])).
cnf(c_0_1652,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1647, c_0_1561]), c_0_1562])).
cnf(c_0_1653,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1631]), c_0_1634]), c_0_1541]), c_0_1511])).
cnf(c_0_1654,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1648, c_0_1561]), c_0_1562])).
cnf(c_0_1655,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1604, c_0_1615]), c_0_1599])).
cnf(c_0_1656,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1649, c_0_1561]), c_0_1562])).
cnf(c_0_1657,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1650, c_0_1561]), c_0_1562])).
cnf(c_0_1658,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_1644]), c_0_1561]), c_0_1543]), c_0_1546]), c_0_1511]), c_0_1562])).
cnf(c_0_1659,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1651, c_0_1652])).
cnf(c_0_1660,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1653, c_0_1654])).
cnf(c_0_1661,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1655, c_0_1561]), c_0_1562])).
cnf(c_0_1662,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1656]), c_0_1657]), c_0_1541]), c_0_1511])).
cnf(c_0_1663,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1658, c_0_1561]), c_0_1562])).
cnf(c_0_1664,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1659]), c_0_1660]), c_0_1541]), c_0_1511])).
cnf(c_0_1665,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1661, c_0_1561]), c_0_1562])).
cnf(c_0_1666,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1662, c_0_1663])).
cnf(c_0_1667,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1664, c_0_1665]), c_0_1651])).
cnf(c_0_1668,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_1666]), c_0_1561]), c_0_1543]), c_0_1546]), c_0_1511]), c_0_1562])).
cnf(c_0_1669,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_1667]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1670,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1668, c_0_1644]), c_0_1561]), c_0_1562])).
cnf(c_0_1671,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_1667]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1672,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_1646]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1673,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_1646]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1674,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1669, c_0_1545]), c_0_1670]), c_0_1671])).
cnf(c_0_1675,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1672, c_0_1545]), c_0_1546]), c_0_1673])).
cnf(c_0_1676,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_1667]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1677,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_1667]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1678,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1674, c_0_1615]), c_0_1671])).
cnf(c_0_1679,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1675, c_0_1615]), c_0_1673])).
cnf(c_0_1680,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1676, c_0_1545]), c_0_1670]), c_0_1677])).
cnf(c_0_1681,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1678, c_0_1561]), c_0_1562])).
cnf(c_0_1682,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1679, c_0_1561]), c_0_1562])).
cnf(c_0_1683,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1680, c_0_1615]), c_0_1677])).
cnf(c_0_1684,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1681]), c_0_1682]), c_0_1541]), c_0_1511])).
cnf(c_0_1685,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1683, c_0_1561]), c_0_1562])).
cnf(c_0_1686,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1684, c_0_1685])).
cnf(c_0_1687,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_1686]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1688,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_1686]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1689,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_1686]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1690,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_1686]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1691,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1687, c_0_1545]), c_0_1670]), c_0_1688])).
cnf(c_0_1692,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1689, c_0_1545]), c_0_1670]), c_0_1690])).
cnf(c_0_1693,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_1686]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1694,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_1686]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1695,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1691, c_0_1615]), c_0_1688])).
cnf(c_0_1696,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1692, c_0_1615]), c_0_1690])).
cnf(c_0_1697,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1693, c_0_1545]), c_0_1670]), c_0_1694])).
cnf(c_0_1698,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1695, c_0_1561]), c_0_1562])).
cnf(c_0_1699,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1696, c_0_1561]), c_0_1562])).
cnf(c_0_1700,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1697, c_0_1615]), c_0_1694])).
cnf(c_0_1701,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1698]), c_0_1699]), c_0_1541]), c_0_1511])).
cnf(c_0_1702,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1700, c_0_1561]), c_0_1562])).
cnf(c_0_1703,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1701, c_0_1702])).
cnf(c_0_1704,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk1_0,esk9_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_1703]), c_0_1561]), c_0_1543]), c_0_1511]), c_0_1562])).
cnf(c_0_1705,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~p2(X2)|~r1(esk34_1(esk1_0),X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1704, c_0_1686]), c_0_1561]), c_0_1562])).
cnf(c_0_1706,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk9_0)|~p2(esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_1543]), c_0_1541]), c_0_1511])).
cnf(c_0_1707,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1705, c_0_1545]), c_0_1670]), c_0_1706])).
cnf(c_0_1708,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1707, c_0_1615]), c_0_1706])).
cnf(c_0_1709,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1557, c_0_1708])).
cnf(c_0_1710,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1559, c_0_1708])).
cnf(c_0_1711,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1547, c_0_1708])).
cnf(c_0_1712,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1549, c_0_1708])).
cnf(c_0_1713,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1709, c_0_1561]), c_0_1562])).
cnf(c_0_1714,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1710, c_0_1561]), c_0_1562])).
cnf(c_0_1715,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1568, c_0_1708])).
cnf(c_0_1716,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1711, c_0_1561]), c_0_1562])).
cnf(c_0_1717,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1712, c_0_1561]), c_0_1562])).
cnf(c_0_1718,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1555, c_0_1708])).
cnf(c_0_1719,negated_conjecture,(p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1713, c_0_1561]), c_0_1562])).
cnf(c_0_1720,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1714, c_0_1561]), c_0_1562])).
cnf(c_0_1721,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1715, c_0_1561]), c_0_1562])).
cnf(c_0_1722,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1716, c_0_1561]), c_0_1562])).
cnf(c_0_1723,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1717, c_0_1561]), c_0_1562])).
cnf(c_0_1724,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1718, c_0_1561]), c_0_1562])).
cnf(c_0_1725,plain,(p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1719]), c_0_1720]), c_0_1541]), c_0_1511])).
cnf(c_0_1726,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1721, c_0_1561]), c_0_1562])).
cnf(c_0_1727,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1722]), c_0_1723]), c_0_1541]), c_0_1511])).
cnf(c_0_1728,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1724, c_0_1561]), c_0_1562])).
cnf(c_0_1729,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1584, c_0_1708])).
cnf(c_0_1730,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1586, c_0_1708])).
cnf(c_0_1731,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1725, c_0_1726])).
cnf(c_0_1732,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1727, c_0_1728])).
cnf(c_0_1733,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1729, c_0_1561]), c_0_1562])).
cnf(c_0_1734,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1730, c_0_1561]), c_0_1562])).
cnf(c_0_1735,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1731]), c_0_1732]), c_0_1541]), c_0_1511])).
cnf(c_0_1736,negated_conjecture,(p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1733, c_0_1561]), c_0_1562])).
cnf(c_0_1737,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1720]), c_0_1723]), c_0_1541]), c_0_1511])).
cnf(c_0_1738,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1734, c_0_1561]), c_0_1562])).
cnf(c_0_1739,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1599, c_0_1708])).
cnf(c_0_1740,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1735, c_0_1736])).
cnf(c_0_1741,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1737, c_0_1738])).
cnf(c_0_1742,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1739, c_0_1561]), c_0_1562])).
cnf(c_0_1743,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1740]), c_0_1741]), c_0_1541]), c_0_1511])).
cnf(c_0_1744,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1742, c_0_1561]), c_0_1562])).
cnf(c_0_1745,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1743, c_0_1744]), c_0_1735])).
cnf(c_0_1746,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_1745]), c_0_1561]), c_0_1543]), c_0_1708]), c_0_1511]), c_0_1562])).
cnf(c_0_1747,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_1732]), c_0_1561]), c_0_1543]), c_0_1708]), c_0_1511]), c_0_1562])).
cnf(c_0_1748,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1746, c_0_1561]), c_0_1562])).
cnf(c_0_1749,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1747, c_0_1561]), c_0_1562])).
cnf(c_0_1750,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_1745]), c_0_1561]), c_0_1543]), c_0_1708]), c_0_1511]), c_0_1562])).
cnf(c_0_1751,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1748]), c_0_1749]), c_0_1541]), c_0_1511])).
cnf(c_0_1752,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1750, c_0_1561]), c_0_1562])).
cnf(c_0_1753,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1751, c_0_1752])).
cnf(c_0_1754,plain,(p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_1753]), c_0_1561]), c_0_1543]), c_0_1708]), c_0_1511]), c_0_1562])).
cnf(c_0_1755,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_1753]), c_0_1561]), c_0_1543]), c_0_1708]), c_0_1511]), c_0_1562])).
cnf(c_0_1756,negated_conjecture,(p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1754, c_0_1561]), c_0_1562])).
cnf(c_0_1757,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1755, c_0_1561]), c_0_1562])).
cnf(c_0_1758,plain,(p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk1_0,esk9_0)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_1753]), c_0_1561]), c_0_1543]), c_0_1708]), c_0_1511]), c_0_1562])).
cnf(c_0_1759,plain,(p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1756]), c_0_1757]), c_0_1541]), c_0_1511])).
cnf(c_0_1760,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1758, c_0_1561]), c_0_1562])).
cnf(c_0_1761,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk1_0,esk9_0)), inference(spm,[status(thm)],[c_0_1759, c_0_1760])).
cnf(c_0_1762,plain,(p2(esk2_0)|p2(X1)|r1(esk1_0,esk9_0)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_1761]), c_0_1561]), c_0_1543]), c_0_1708]), c_0_1511]), c_0_1562])).
cnf(c_0_1763,plain,(p2(esk2_0)|r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1762, c_0_1753]), c_0_1561]), c_0_1562])).
cnf(c_0_1764,negated_conjecture,(r1(esk1_0,esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_36, c_0_1763])).
cnf(c_0_1765,negated_conjecture,(r1(esk1_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1764, c_0_10]), c_0_11])).
cnf(c_0_1766,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk15_2(esk9_0,X1))|~r1(esk9_0,X1)), inference(spm,[status(thm)],[c_0_1497, c_0_1765])).
cnf(c_0_1767,negated_conjecture,(p2(esk9_0)|r1(esk9_0,esk32_1(esk9_0))), inference(spm,[status(thm)],[c_0_14, c_0_1765])).
cnf(c_0_1768,negated_conjecture,(p2(esk9_0)|~p2(esk32_1(esk9_0))), inference(spm,[status(thm)],[c_0_15, c_0_1765])).
cnf(c_0_1769,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk9_0,X1))|p2(X1)|~r1(esk9_0,X1)), inference(spm,[status(thm)],[c_0_1498, c_0_1765])).
cnf(c_0_1770,negated_conjecture,(epred1_1(esk1_0)|p2(esk9_0)|r1(esk32_1(esk9_0),esk15_2(esk9_0,esk32_1(esk9_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1766, c_0_1767]), c_0_1768])).
cnf(c_0_1771,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk9_0,esk32_1(esk9_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1769, c_0_1767]), c_0_1768])).
cnf(c_0_1772,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk15_2(esk9_0,X1),esk16_2(esk9_0,X1))|~r1(esk9_0,X1)), inference(spm,[status(thm)],[c_0_1502, c_0_1765])).
cnf(c_0_1773,negated_conjecture,(epred1_1(esk1_0)|p2(esk9_0)|p2(X1)|~r1(esk15_2(esk9_0,esk32_1(esk9_0)),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1770]), c_0_1765])]), c_0_1771])).
cnf(c_0_1774,negated_conjecture,(epred1_1(esk1_0)|p2(esk9_0)|r1(esk15_2(esk9_0,esk32_1(esk9_0)),esk16_2(esk9_0,esk32_1(esk9_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1772, c_0_1767]), c_0_1768])).
cnf(c_0_1775,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk9_0,esk32_1(esk9_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_1773, c_0_1774])).
cnf(c_0_1776,negated_conjecture,(epred1_1(esk1_0)|p2(esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1509, c_0_1775]), c_0_1765])]), c_0_1767]), c_0_1768])).
cnf(c_0_1777,plain,(p2(esk9_0)|r1(esk1_0,esk33_1(esk1_0))), inference(spm,[status(thm)],[c_0_32, c_0_1776])).
cnf(c_0_1778,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_1777]), c_0_350])).
cnf(c_0_1779,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_1777]), c_0_350])).
cnf(c_0_1780,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1778, c_0_7])).
cnf(c_0_1781,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1779, c_0_7])).
cnf(c_0_1782,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|r1(X2,esk36_2(esk1_0,X2))|~epred1_1(esk1_0)|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1780]), c_0_1781]), c_0_43])).
cnf(c_0_1783,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_1777]), c_0_350])).
cnf(c_0_1784,plain,(p2(esk36_2(esk1_0,X1))|p2(X2)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~epred1_1(esk1_0)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1780]), c_0_1781]), c_0_46])).
cnf(c_0_1785,negated_conjecture,(p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk36_2(esk1_0,X1))|r1(esk1_0,esk8_1(esk1_0))|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1782, c_0_1776])).
cnf(c_0_1786,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1783, c_0_7])).
cnf(c_0_1787,negated_conjecture,(p2(esk36_2(esk1_0,X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1784, c_0_1776])).
cnf(c_0_1788,plain,(p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|r1(esk1_0,esk8_1(esk1_0))|~epred1_1(esk1_0)|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1780]), c_0_1781]), c_0_51])).
cnf(c_0_1789,plain,(p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|~epred1_1(esk1_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1780]), c_0_1781]), c_0_53])).
cnf(c_0_1790,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1785, c_0_1786])).
cnf(c_0_1791,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_14, c_0_1777])).
cnf(c_0_1792,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|~p2(esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_15, c_0_1777])).
cnf(c_0_1793,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1787, c_0_1786])).
cnf(c_0_1794,negated_conjecture,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk8_1(esk1_0))|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1788, c_0_1776])).
cnf(c_0_1795,negated_conjecture,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk1_0,esk8_1(esk1_0))|~p2(esk37_2(esk1_0,X1))|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1789, c_0_1776])).
cnf(c_0_1796,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1790, c_0_1791]), c_0_1792])).
cnf(c_0_1797,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1793, c_0_1791]), c_0_1792])).
cnf(c_0_1798,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk1_0,esk8_1(esk1_0))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1794, c_0_1786])).
cnf(c_0_1799,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1795, c_0_1786])).
cnf(c_0_1800,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk1_0,esk8_1(esk1_0))|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1796]), c_0_1777]), c_0_1797])).
cnf(c_0_1801,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1798, c_0_1791]), c_0_1792])).
cnf(c_0_1802,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk8_1(esk1_0))|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1799, c_0_1791]), c_0_1792])).
cnf(c_0_1803,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk1_0,esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1800, c_0_1801]), c_0_1802])).
cnf(c_0_1804,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk8_1(esk1_0))|r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_1803]), c_0_1777]), c_0_350])).
cnf(c_0_1805,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_34, c_0_1777])).
cnf(c_0_1806,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1804, c_0_7])).
cnf(c_0_1807,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(X1))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_1777]), c_0_382])).
cnf(c_0_1808,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_37, c_0_1777])).
cnf(c_0_1809,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p5(esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_1777]), c_0_382])).
cnf(c_0_1810,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1805, c_0_1806])).
cnf(c_0_1811,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(esk1_0))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1807, c_0_7])).
cnf(c_0_1812,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1808, c_0_1806])).
cnf(c_0_1813,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_1777])).
cnf(c_0_1814,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1809, c_0_7])).
cnf(c_0_1815,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1810, c_0_1811])).
cnf(c_0_1816,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1812, c_0_1811])).
cnf(c_0_1817,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|~p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1813, c_0_1806])).
cnf(c_0_1818,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1810, c_0_1814])).
cnf(c_0_1819,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_1812, c_0_1814])).
cnf(c_0_1820,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(X2,esk36_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1815]), c_0_1816]), c_0_1776])).
cnf(c_0_1821,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1817, c_0_1811])).
cnf(c_0_1822,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1815]), c_0_1816]), c_0_1776])).
cnf(c_0_1823,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1818]), c_0_1819]), c_0_1776])).
cnf(c_0_1824,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1817, c_0_1814])).
cnf(c_0_1825,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1818]), c_0_1819]), c_0_1776])).
cnf(c_0_1826,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1820, c_0_1821])).
cnf(c_0_1827,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1822, c_0_1821])).
cnf(c_0_1828,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1815]), c_0_1816]), c_0_1776])).
cnf(c_0_1829,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1815]), c_0_1816]), c_0_1776])).
cnf(c_0_1830,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1823, c_0_1824])).
cnf(c_0_1831,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1825, c_0_1824])).
cnf(c_0_1832,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1818]), c_0_1819]), c_0_1776])).
cnf(c_0_1833,plain,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1818]), c_0_1819]), c_0_1776])).
cnf(c_0_1834,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1826, c_0_1791]), c_0_1792])).
cnf(c_0_1835,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1827, c_0_1791]), c_0_1792])).
cnf(c_0_1836,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1828, c_0_1821])).
cnf(c_0_1837,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1829, c_0_1821])).
cnf(c_0_1838,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1830, c_0_1791]), c_0_1792])).
cnf(c_0_1839,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1831, c_0_1791]), c_0_1792])).
cnf(c_0_1840,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1832, c_0_1824])).
cnf(c_0_1841,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1833, c_0_1824])).
cnf(c_0_1842,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1834]), c_0_1777]), c_0_1835])).
cnf(c_0_1843,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1836, c_0_1791]), c_0_1792])).
cnf(c_0_1844,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1837, c_0_1791]), c_0_1792])).
cnf(c_0_1845,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1838]), c_0_1777]), c_0_1839])).
cnf(c_0_1846,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1840, c_0_1791]), c_0_1792])).
cnf(c_0_1847,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1841, c_0_1791]), c_0_1792])).
cnf(c_0_1848,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|~p2(esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_69, c_0_1777])).
cnf(c_0_1849,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1842, c_0_1843]), c_0_1844])).
cnf(c_0_1850,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1845, c_0_1846]), c_0_1847])).
cnf(c_0_1851,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1848, c_0_1849])).
cnf(c_0_1852,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1848, c_0_1850])).
cnf(c_0_1853,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk10_1(esk33_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1851, c_0_1806]), c_0_1811])).
cnf(c_0_1854,negated_conjecture,(p2(esk10_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1852, c_0_1806]), c_0_1814])).
cnf(c_0_1855,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(X1))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_432, c_0_1777]), c_0_382])).
cnf(c_0_1856,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1853]), c_0_1854]), c_0_1776])).
cnf(c_0_1857,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(esk1_0))|r1(esk10_1(esk33_1(esk1_0)),esk11_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1855, c_0_7])).
cnf(c_0_1858,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1853]), c_0_1854]), c_0_1776])).
cnf(c_0_1859,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1856, c_0_1857])).
cnf(c_0_1860,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1858, c_0_1857])).
cnf(c_0_1861,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1853]), c_0_1854]), c_0_1776])).
cnf(c_0_1862,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1853]), c_0_1854]), c_0_1776])).
cnf(c_0_1863,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1816]), c_0_1819]), c_0_1776])).
cnf(c_0_1864,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|~r1(esk10_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1816]), c_0_1819]), c_0_1776])).
cnf(c_0_1865,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1859, c_0_1791]), c_0_1792])).
cnf(c_0_1866,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1860, c_0_1791]), c_0_1792])).
cnf(c_0_1867,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1861, c_0_1857])).
cnf(c_0_1868,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1862, c_0_1857])).
cnf(c_0_1869,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1863, c_0_1857]), c_0_1812])).
cnf(c_0_1870,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1864, c_0_1857]), c_0_1812])).
cnf(c_0_1871,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1816]), c_0_1819]), c_0_1776])).
cnf(c_0_1872,plain,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk10_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1816]), c_0_1819]), c_0_1776])).
cnf(c_0_1873,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p5(esk8_1(esk1_0))|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1865]), c_0_1777]), c_0_1866])).
cnf(c_0_1874,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1867, c_0_1791]), c_0_1792])).
cnf(c_0_1875,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1868, c_0_1791]), c_0_1792])).
cnf(c_0_1876,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1869, c_0_1791]), c_0_1792])).
cnf(c_0_1877,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1870, c_0_1791]), c_0_1792])).
cnf(c_0_1878,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1871, c_0_1857]), c_0_1812])).
cnf(c_0_1879,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1872, c_0_1857]), c_0_1812])).
cnf(c_0_1880,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p5(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1873, c_0_1874]), c_0_1875])).
cnf(c_0_1881,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1876]), c_0_1777]), c_0_1877])).
cnf(c_0_1882,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1878, c_0_1791]), c_0_1792])).
cnf(c_0_1883,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1879, c_0_1791]), c_0_1792])).
cnf(c_0_1884,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1810, c_0_1880])).
cnf(c_0_1885,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1881, c_0_1882]), c_0_1883])).
cnf(c_0_1886,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1884]), c_0_1885]), c_0_1776])).
cnf(c_0_1887,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1817, c_0_1880])).
cnf(c_0_1888,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1884]), c_0_1885]), c_0_1776])).
cnf(c_0_1889,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1886, c_0_1887])).
cnf(c_0_1890,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1888, c_0_1887])).
cnf(c_0_1891,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1884]), c_0_1885]), c_0_1776])).
cnf(c_0_1892,plain,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1884]), c_0_1885]), c_0_1776])).
cnf(c_0_1893,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1889, c_0_1791]), c_0_1792])).
cnf(c_0_1894,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1890, c_0_1791]), c_0_1792])).
cnf(c_0_1895,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1891, c_0_1887])).
cnf(c_0_1896,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1892, c_0_1887])).
cnf(c_0_1897,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1893]), c_0_1777]), c_0_1894])).
cnf(c_0_1898,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1895, c_0_1791]), c_0_1792])).
cnf(c_0_1899,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1896, c_0_1791]), c_0_1792])).
cnf(c_0_1900,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1897, c_0_1898]), c_0_1899])).
cnf(c_0_1901,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1848, c_0_1900])).
cnf(c_0_1902,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(X1))|~p2(esk11_1(esk33_1(esk1_0)))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_1777]), c_0_382])).
cnf(c_0_1903,negated_conjecture,(p2(esk11_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1901, c_0_1806]), c_0_1880])).
cnf(c_0_1904,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(X1))|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1902, c_0_1903]), c_0_382])).
cnf(c_0_1905,negated_conjecture,(p2(esk33_1(esk1_0))|p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_1904, c_0_7])).
cnf(c_0_1906,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1810, c_0_1905])).
cnf(c_0_1907,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_1812, c_0_1905])).
cnf(c_0_1908,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1906]), c_0_1907]), c_0_1776])).
cnf(c_0_1909,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1817, c_0_1905])).
cnf(c_0_1910,plain,(p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X2)|p2(X1)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1906]), c_0_1907]), c_0_1776])).
cnf(c_0_1911,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1908, c_0_1909])).
cnf(c_0_1912,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1910, c_0_1909])).
cnf(c_0_1913,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1906]), c_0_1907]), c_0_1776])).
cnf(c_0_1914,plain,(p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1906]), c_0_1907]), c_0_1776])).
cnf(c_0_1915,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1911, c_0_1791]), c_0_1792])).
cnf(c_0_1916,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1912, c_0_1791]), c_0_1792])).
cnf(c_0_1917,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1913, c_0_1909])).
cnf(c_0_1918,plain,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1914, c_0_1909])).
cnf(c_0_1919,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|p2(X1)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_1915]), c_0_1777]), c_0_1916])).
cnf(c_0_1920,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1917, c_0_1791]), c_0_1792])).
cnf(c_0_1921,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1918, c_0_1791]), c_0_1792])).
cnf(c_0_1922,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1919, c_0_1920]), c_0_1921])).
cnf(c_0_1923,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)|~p5(X1)|~r1(esk1_0,X1)), inference(spm,[status(thm)],[c_0_1848, c_0_1922])).
cnf(c_0_1924,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1923, c_0_1806]), c_0_1905])).
cnf(c_0_1925,plain,(p2(esk9_0)|r1(esk33_1(esk1_0),esk34_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_1924]), c_0_1776])).
cnf(c_0_1926,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_1925]), c_0_1776])).
cnf(c_0_1927,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)|r1(esk34_1(esk1_0),esk36_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_1925]), c_0_1924]), c_0_1776])).
cnf(c_0_1928,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk34_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_1925]), c_0_1924]), c_0_1776])).
cnf(c_0_1929,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_1925]), c_0_1776])).
cnf(c_0_1930,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk9_0)|p2(X3)|p2(X2)|p2(X1)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_1925]), c_0_1776])).
cnf(c_0_1931,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk9_0)|p2(X2)|p2(X1)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_1925]), c_0_1776])).
cnf(c_0_1932,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1926, c_0_1927]), c_0_1928]), c_0_1929])).
cnf(c_0_1933,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1930, c_0_1927]), c_0_1928]), c_0_1931])).
cnf(c_0_1934,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_1925]), c_0_1776])).
cnf(c_0_1935,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_1925]), c_0_1776])).
cnf(c_0_1936,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_1925]), c_0_1776])).
cnf(c_0_1937,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_1925]), c_0_1776])).
cnf(c_0_1938,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk45_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_1925]), c_0_1776])).
cnf(c_0_1939,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_1925]), c_0_1776])).
cnf(c_0_1940,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1932, c_0_7])).
cnf(c_0_1941,plain,(p2(esk9_0)|r1(esk34_1(esk1_0),esk35_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_1924]), c_0_1776])).
cnf(c_0_1942,plain,(p2(esk9_0)|~p2(esk35_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_1924]), c_0_1776])).
cnf(c_0_1943,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1933, c_0_7])).
cnf(c_0_1944,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1934, c_0_1927]), c_0_1928]), c_0_1935])).
cnf(c_0_1945,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1936, c_0_1927]), c_0_1928]), c_0_1937])).
cnf(c_0_1946,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1938, c_0_1927]), c_0_1928]), c_0_1939])).
cnf(c_0_1947,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_1925]), c_0_1776])).
cnf(c_0_1948,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_1925]), c_0_1776])).
cnf(c_0_1949,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1940, c_0_1941]), c_0_1942])).
cnf(c_0_1950,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1943, c_0_1941]), c_0_1942])).
cnf(c_0_1951,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1944, c_0_7])).
cnf(c_0_1952,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1945, c_0_7])).
cnf(c_0_1953,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1946, c_0_7])).
cnf(c_0_1954,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1947, c_0_1927]), c_0_1928]), c_0_1948])).
cnf(c_0_1955,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1949, c_0_1941]), c_0_1942])).
cnf(c_0_1956,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1950, c_0_1941]), c_0_1942])).
cnf(c_0_1957,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1951, c_0_1941]), c_0_1942])).
cnf(c_0_1958,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1952, c_0_1941]), c_0_1942])).
cnf(c_0_1959,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1953, c_0_1941]), c_0_1942])).
cnf(c_0_1960,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1954, c_0_7])).
cnf(c_0_1961,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1955]), c_0_1956]), c_0_1924]), c_0_1776])).
cnf(c_0_1962,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1957, c_0_1941]), c_0_1942])).
cnf(c_0_1963,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_1925]), c_0_1776])).
cnf(c_0_1964,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_1925]), c_0_1776])).
cnf(c_0_1965,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X3)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X3),esk46_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_1925]), c_0_1776])).
cnf(c_0_1966,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_1925]), c_0_1776])).
cnf(c_0_1967,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1958, c_0_1941]), c_0_1942])).
cnf(c_0_1968,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1959, c_0_1941]), c_0_1942])).
cnf(c_0_1969,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1960, c_0_1941]), c_0_1942])).
cnf(c_0_1970,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_1961, c_0_1962])).
cnf(c_0_1971,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1963, c_0_1927]), c_0_1928]), c_0_1964])).
cnf(c_0_1972,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1965, c_0_1927]), c_0_1928]), c_0_1966])).
cnf(c_0_1973,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1967]), c_0_1968]), c_0_1924]), c_0_1776])).
cnf(c_0_1974,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1969, c_0_1941]), c_0_1942])).
cnf(c_0_1975,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_1970]), c_0_1941]), c_0_1925]), c_0_1928]), c_0_1776]), c_0_1942])).
cnf(c_0_1976,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1971, c_0_7])).
cnf(c_0_1977,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1972, c_0_7])).
cnf(c_0_1978,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_1925]), c_0_1776])).
cnf(c_0_1979,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_1925]), c_0_1776])).
cnf(c_0_1980,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1973, c_0_1974])).
cnf(c_0_1981,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1975, c_0_1941]), c_0_1942])).
cnf(c_0_1982,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1976, c_0_1941]), c_0_1942])).
cnf(c_0_1983,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1977, c_0_1941]), c_0_1942])).
cnf(c_0_1984,plain,(p2(esk9_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1978, c_0_1927]), c_0_1928]), c_0_1979])).
cnf(c_0_1985,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1980]), c_0_1981]), c_0_1924]), c_0_1776])).
cnf(c_0_1986,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1982, c_0_1941]), c_0_1942])).
cnf(c_0_1987,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1968]), c_0_1956]), c_0_1924]), c_0_1776])).
cnf(c_0_1988,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1983, c_0_1941]), c_0_1942])).
cnf(c_0_1989,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1984, c_0_7])).
cnf(c_0_1990,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_1985, c_0_1986])).
cnf(c_0_1991,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_1987, c_0_1988])).
cnf(c_0_1992,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1989, c_0_1941]), c_0_1942])).
cnf(c_0_1993,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)|r1(esk36_2(esk1_0,esk34_1(esk1_0)),esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_1925]), c_0_1924]), c_0_1776])).
cnf(c_0_1994,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_1990]), c_0_1991]), c_0_1924]), c_0_1776])).
cnf(c_0_1995,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1992, c_0_1941]), c_0_1942])).
cnf(c_0_1996,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1945, c_0_1993]), c_0_1937])).
cnf(c_0_1997,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1946, c_0_1993]), c_0_1939])).
cnf(c_0_1998,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1932, c_0_1993]), c_0_1929])).
cnf(c_0_1999,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1933, c_0_1993]), c_0_1931])).
cnf(c_0_2000,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1994, c_0_1995]), c_0_1985])).
cnf(c_0_2001,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1996, c_0_1941]), c_0_1942])).
cnf(c_0_2002,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1997, c_0_1941]), c_0_1942])).
cnf(c_0_2003,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1954, c_0_1993]), c_0_1948])).
cnf(c_0_2004,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1998, c_0_1941]), c_0_1942])).
cnf(c_0_2005,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1999, c_0_1941]), c_0_1942])).
cnf(c_0_2006,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1944, c_0_1993]), c_0_1935])).
cnf(c_0_2007,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_2000]), c_0_1941]), c_0_1925]), c_0_1928]), c_0_1776]), c_0_1942])).
cnf(c_0_2008,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2001, c_0_1941]), c_0_1942])).
cnf(c_0_2009,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2002, c_0_1941]), c_0_1942])).
cnf(c_0_2010,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2003, c_0_1941]), c_0_1942])).
cnf(c_0_2011,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2004, c_0_1941]), c_0_1942])).
cnf(c_0_2012,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2005, c_0_1941]), c_0_1942])).
cnf(c_0_2013,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2006, c_0_1941]), c_0_1942])).
cnf(c_0_2014,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2007, c_0_1941]), c_0_1942])).
cnf(c_0_2015,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_2000]), c_0_1941]), c_0_1925]), c_0_1928]), c_0_1776]), c_0_1942])).
cnf(c_0_2016,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2008]), c_0_2009]), c_0_1924]), c_0_1776])).
cnf(c_0_2017,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2010, c_0_1941]), c_0_1942])).
cnf(c_0_2018,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2011]), c_0_2012]), c_0_1924]), c_0_1776])).
cnf(c_0_2019,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2013, c_0_1941]), c_0_1942])).
cnf(c_0_2020,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1971, c_0_1993]), c_0_1964])).
cnf(c_0_2021,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1972, c_0_1993]), c_0_1966])).
cnf(c_0_2022,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2014]), c_0_1981]), c_0_1924]), c_0_1776])).
cnf(c_0_2023,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2015, c_0_1941]), c_0_1942])).
cnf(c_0_2024,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2016, c_0_2017])).
cnf(c_0_2025,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2018, c_0_2019])).
cnf(c_0_2026,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2020, c_0_1941]), c_0_1942])).
cnf(c_0_2027,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2021, c_0_1941]), c_0_1942])).
cnf(c_0_2028,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2022, c_0_2023])).
cnf(c_0_2029,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2024]), c_0_2025]), c_0_1924]), c_0_1776])).
cnf(c_0_2030,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2026, c_0_1941]), c_0_1942])).
cnf(c_0_2031,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2009]), c_0_2012]), c_0_1924]), c_0_1776])).
cnf(c_0_2032,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2027, c_0_1941]), c_0_1942])).
cnf(c_0_2033,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1984, c_0_1993]), c_0_1979])).
cnf(c_0_2034,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_2028]), c_0_1941]), c_0_1925]), c_0_1928]), c_0_1776]), c_0_1942])).
cnf(c_0_2035,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_2028]), c_0_1941]), c_0_1925]), c_0_1928]), c_0_1776]), c_0_1942])).
cnf(c_0_2036,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2029, c_0_2030])).
cnf(c_0_2037,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2031, c_0_2032])).
cnf(c_0_2038,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2033, c_0_1941]), c_0_1942])).
cnf(c_0_2039,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2034, c_0_1941]), c_0_1942])).
cnf(c_0_2040,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2035, c_0_1941]), c_0_1942])).
cnf(c_0_2041,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_2028]), c_0_1941]), c_0_1925]), c_0_1928]), c_0_1776]), c_0_1942])).
cnf(c_0_2042,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2036]), c_0_2037]), c_0_1924]), c_0_1776])).
cnf(c_0_2043,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2038, c_0_1941]), c_0_1942])).
cnf(c_0_2044,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2039]), c_0_2040]), c_0_1924]), c_0_1776])).
cnf(c_0_2045,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2041, c_0_1941]), c_0_1942])).
cnf(c_0_2046,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2042, c_0_2043]), c_0_2029])).
cnf(c_0_2047,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2044, c_0_2045])).
cnf(c_0_2048,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2046]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2049,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_2046]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2050,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_2025]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2051,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_2025]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2052,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_2047]), c_0_1941]), c_0_1925]), c_0_1928]), c_0_1776]), c_0_1942])).
cnf(c_0_2053,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2048, c_0_1927]), c_0_1928]), c_0_2049])).
cnf(c_0_2054,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2050, c_0_1927]), c_0_1928]), c_0_2051])).
cnf(c_0_2055,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_2046]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2056,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2052, c_0_2028]), c_0_1941]), c_0_1942])).
cnf(c_0_2057,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_2046]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2058,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2053, c_0_1993]), c_0_2049])).
cnf(c_0_2059,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2054, c_0_1993]), c_0_2051])).
cnf(c_0_2060,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2055, c_0_1927]), c_0_2056]), c_0_2057])).
cnf(c_0_2061,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2058, c_0_1941]), c_0_1942])).
cnf(c_0_2062,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2059, c_0_1941]), c_0_1942])).
cnf(c_0_2063,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2060, c_0_1993]), c_0_2057])).
cnf(c_0_2064,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2061]), c_0_2062]), c_0_1924]), c_0_1776])).
cnf(c_0_2065,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2063, c_0_1941]), c_0_1942])).
cnf(c_0_2066,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2064, c_0_2065])).
cnf(c_0_2067,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_2066]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2068,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_2066]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2069,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X2)|p2(X1)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_2066]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2070,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_2066]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2071,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2067, c_0_1927]), c_0_2056]), c_0_2068])).
cnf(c_0_2072,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2069, c_0_1927]), c_0_2056]), c_0_2070])).
cnf(c_0_2073,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_2066]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2074,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_2066]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2075,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2071, c_0_1993]), c_0_2068])).
cnf(c_0_2076,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2072, c_0_1993]), c_0_2070])).
cnf(c_0_2077,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2073, c_0_1927]), c_0_2056]), c_0_2074])).
cnf(c_0_2078,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2075, c_0_1941]), c_0_1942])).
cnf(c_0_2079,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2076, c_0_1941]), c_0_1942])).
cnf(c_0_2080,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2077, c_0_1993]), c_0_2074])).
cnf(c_0_2081,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2078]), c_0_2079]), c_0_1924]), c_0_1776])).
cnf(c_0_2082,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2080, c_0_1941]), c_0_1942])).
cnf(c_0_2083,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2081, c_0_2082])).
cnf(c_0_2084,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|p2(X2)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_2083]), c_0_1941]), c_0_1925]), c_0_1776]), c_0_1942])).
cnf(c_0_2085,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk9_0)|p2(X1)|~p2(X2)|~r1(esk34_1(esk1_0),X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2084, c_0_2066]), c_0_1941]), c_0_1942])).
cnf(c_0_2086,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)|~p2(esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_1925]), c_0_1924]), c_0_1776])).
cnf(c_0_2087,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)|p2(X1)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2085, c_0_1927]), c_0_2056]), c_0_2086])).
cnf(c_0_2088,plain,(p2(esk34_1(esk1_0))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2087, c_0_1993]), c_0_2086])).
cnf(c_0_2089,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1937, c_0_2088])).
cnf(c_0_2090,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1939, c_0_2088])).
cnf(c_0_2091,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1929, c_0_2088])).
cnf(c_0_2092,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk9_0)|p2(X2)|p2(X1)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1931, c_0_2088])).
cnf(c_0_2093,plain,(p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2089, c_0_1941]), c_0_1942])).
cnf(c_0_2094,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2090, c_0_1941]), c_0_1942])).
cnf(c_0_2095,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1948, c_0_2088])).
cnf(c_0_2096,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2091, c_0_1941]), c_0_1942])).
cnf(c_0_2097,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2092, c_0_1941]), c_0_1942])).
cnf(c_0_2098,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X2)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_1935, c_0_2088])).
cnf(c_0_2099,plain,(p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2093, c_0_1941]), c_0_1942])).
cnf(c_0_2100,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2094, c_0_1941]), c_0_1942])).
cnf(c_0_2101,plain,(p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2095, c_0_1941]), c_0_1942])).
cnf(c_0_2102,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2096, c_0_1941]), c_0_1942])).
cnf(c_0_2103,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2097, c_0_1941]), c_0_1942])).
cnf(c_0_2104,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2098, c_0_1941]), c_0_1942])).
cnf(c_0_2105,plain,(p2(esk9_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2099]), c_0_2100]), c_0_1924]), c_0_1776])).
cnf(c_0_2106,plain,(p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2101, c_0_1941]), c_0_1942])).
cnf(c_0_2107,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2102]), c_0_2103]), c_0_1924]), c_0_1776])).
cnf(c_0_2108,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2104, c_0_1941]), c_0_1942])).
cnf(c_0_2109,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1964, c_0_2088])).
cnf(c_0_2110,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1966, c_0_2088])).
cnf(c_0_2111,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2105, c_0_2106])).
cnf(c_0_2112,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2107, c_0_2108])).
cnf(c_0_2113,plain,(p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2109, c_0_1941]), c_0_1942])).
cnf(c_0_2114,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2110, c_0_1941]), c_0_1942])).
cnf(c_0_2115,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2111]), c_0_2112]), c_0_1924]), c_0_1776])).
cnf(c_0_2116,plain,(p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2113, c_0_1941]), c_0_1942])).
cnf(c_0_2117,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2100]), c_0_2103]), c_0_1924]), c_0_1776])).
cnf(c_0_2118,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2114, c_0_1941]), c_0_1942])).
cnf(c_0_2119,plain,(p2(esk9_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_1979, c_0_2088])).
cnf(c_0_2120,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2115, c_0_2116])).
cnf(c_0_2121,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2117, c_0_2118])).
cnf(c_0_2122,plain,(p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2119, c_0_1941]), c_0_1942])).
cnf(c_0_2123,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2120]), c_0_2121]), c_0_1924]), c_0_1776])).
cnf(c_0_2124,plain,(p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2122, c_0_1941]), c_0_1942])).
cnf(c_0_2125,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2123, c_0_2124]), c_0_2115])).
cnf(c_0_2126,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_2125]), c_0_1941]), c_0_1925]), c_0_2088]), c_0_1776]), c_0_1942])).
cnf(c_0_2127,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_2112]), c_0_1941]), c_0_1925]), c_0_2088]), c_0_1776]), c_0_1942])).
cnf(c_0_2128,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2126, c_0_1941]), c_0_1942])).
cnf(c_0_2129,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2127, c_0_1941]), c_0_1942])).
cnf(c_0_2130,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_2125]), c_0_1941]), c_0_1925]), c_0_2088]), c_0_1776]), c_0_1942])).
cnf(c_0_2131,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2128]), c_0_2129]), c_0_1924]), c_0_1776])).
cnf(c_0_2132,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2130, c_0_1941]), c_0_1942])).
cnf(c_0_2133,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2131, c_0_2132])).
cnf(c_0_2134,plain,(p2(esk9_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_2133]), c_0_1941]), c_0_1925]), c_0_2088]), c_0_1776]), c_0_1942])).
cnf(c_0_2135,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk9_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_2133]), c_0_1941]), c_0_1925]), c_0_2088]), c_0_1776]), c_0_1942])).
cnf(c_0_2136,plain,(p2(esk9_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2134, c_0_1941]), c_0_1942])).
cnf(c_0_2137,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2135, c_0_1941]), c_0_1942])).
cnf(c_0_2138,plain,(p2(esk9_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_2133]), c_0_1941]), c_0_1925]), c_0_2088]), c_0_1776]), c_0_1942])).
cnf(c_0_2139,plain,(p2(esk9_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2136]), c_0_2137]), c_0_1924]), c_0_1776])).
cnf(c_0_2140,plain,(p2(esk9_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2138, c_0_1941]), c_0_1942])).
cnf(c_0_2141,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk9_0)), inference(spm,[status(thm)],[c_0_2139, c_0_2140])).
cnf(c_0_2142,plain,(p2(esk9_0)|p2(X1)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_2141]), c_0_1941]), c_0_1925]), c_0_2088]), c_0_1776]), c_0_1942])).
cnf(c_0_2143,plain,(p2(esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2142, c_0_2133]), c_0_1941]), c_0_1942])).
cnf(c_0_2144,negated_conjecture,(r1(X1,esk8_1(X1))|~r1(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_350, c_0_2143])])).
cnf(c_0_2145,negated_conjecture,(p5(esk8_1(X1))|~r1(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_382, c_0_2143])])).
cnf(c_0_2146,negated_conjecture,(r1(esk1_0,esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_2144, c_0_7])).
cnf(c_0_2147,negated_conjecture,(p5(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_2145, c_0_7])).
cnf(c_0_2148,negated_conjecture,(r1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_2146]), c_0_2147])])).
cnf(c_0_2149,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(X1,esk15_2(esk2_0,X1))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_1497, c_0_2148])).
cnf(c_0_2150,negated_conjecture,(p2(esk2_0)|r1(esk2_0,esk32_1(esk2_0))), inference(spm,[status(thm)],[c_0_14, c_0_2148])).
cnf(c_0_2151,negated_conjecture,(p2(esk2_0)|~p2(esk32_1(esk2_0))), inference(spm,[status(thm)],[c_0_15, c_0_2148])).
cnf(c_0_2152,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk2_0,X1))|p2(X1)|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_1498, c_0_2148])).
cnf(c_0_2153,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk32_1(esk2_0),esk15_2(esk2_0,esk32_1(esk2_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2149, c_0_2150]), c_0_2151])).
cnf(c_0_2154,negated_conjecture,(epred1_1(esk1_0)|p2(esk15_2(esk2_0,esk32_1(esk2_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2152, c_0_2150]), c_0_2151])).
cnf(c_0_2155,negated_conjecture,(epred1_1(esk1_0)|p2(X1)|r1(esk15_2(esk2_0,X1),esk16_2(esk2_0,X1))|~r1(esk2_0,X1)), inference(spm,[status(thm)],[c_0_1502, c_0_2148])).
cnf(c_0_2156,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|p2(X1)|~r1(esk15_2(esk2_0,esk32_1(esk2_0)),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_2153]), c_0_2148])]), c_0_2154])).
cnf(c_0_2157,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)|r1(esk15_2(esk2_0,esk32_1(esk2_0)),esk16_2(esk2_0,esk32_1(esk2_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2155, c_0_2150]), c_0_2151])).
cnf(c_0_2158,negated_conjecture,(epred1_1(esk1_0)|p2(esk16_2(esk2_0,esk32_1(esk2_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2156, c_0_2157])).
cnf(c_0_2159,negated_conjecture,(epred1_1(esk1_0)|p2(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_1509, c_0_2158]), c_0_2148])]), c_0_2150]), c_0_2151])).
cnf(c_0_2160,plain,(p2(esk2_0)|r1(esk1_0,esk33_1(esk1_0))), inference(spm,[status(thm)],[c_0_32, c_0_2159])).
cnf(c_0_2161,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_2160]), c_0_36])).
cnf(c_0_2162,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_2160]), c_0_36])).
cnf(c_0_2163,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk33_1(esk1_0),esk3_1(esk33_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2161, c_0_2146]), c_0_2147])])).
cnf(c_0_2164,negated_conjecture,(p2(esk3_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2162, c_0_2146]), c_0_2147])])).
cnf(c_0_2165,plain,(p2(X1)|p2(X2)|r1(X2,esk36_2(esk1_0,X2))|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_2163]), c_0_2164]), c_0_43])).
cnf(c_0_2166,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_2160]), c_0_36])).
cnf(c_0_2167,plain,(p2(esk36_2(esk1_0,X1))|p2(X2)|p2(X1)|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_2163]), c_0_2164]), c_0_46])).
cnf(c_0_2168,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2165, c_0_2159])).
cnf(c_0_2169,negated_conjecture,(p2(esk33_1(esk1_0))|r1(esk3_1(esk33_1(esk1_0)),esk4_1(esk33_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2166, c_0_2146]), c_0_2147])])).
cnf(c_0_2170,negated_conjecture,(p2(esk36_2(esk1_0,X1))|p2(esk2_0)|p2(X1)|p2(X2)|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2167, c_0_2159])).
cnf(c_0_2171,plain,(p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X2),esk37_2(esk1_0,X2))|~epred1_1(esk1_0)|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_2163]), c_0_2164]), c_0_51])).
cnf(c_0_2172,plain,(p2(X1)|p2(X2)|~epred1_1(esk1_0)|~p2(esk37_2(esk1_0,X2))|~r1(esk3_1(esk33_1(esk1_0)),X1)|~r1(esk33_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_2163]), c_0_2164]), c_0_53])).
cnf(c_0_2173,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(X1,esk36_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2168, c_0_2169])).
cnf(c_0_2174,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk33_1(esk1_0),esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_14, c_0_2160])).
cnf(c_0_2175,negated_conjecture,(p2(esk33_1(esk1_0))|p2(esk2_0)|~p2(esk32_1(esk33_1(esk1_0)))), inference(spm,[status(thm)],[c_0_15, c_0_2160])).
cnf(c_0_2176,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk36_2(esk1_0,X1))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2170, c_0_2169])).
cnf(c_0_2177,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2171, c_0_2159])).
cnf(c_0_2178,negated_conjecture,(p2(esk2_0)|p2(X1)|p2(X2)|~p2(esk37_2(esk1_0,X1))|~r1(esk3_1(esk33_1(esk1_0)),X2)|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2172, c_0_2159])).
cnf(c_0_2179,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk32_1(esk33_1(esk1_0)),esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2173, c_0_2174]), c_0_2175])).
cnf(c_0_2180,negated_conjecture,(p2(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))))|p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2176, c_0_2174]), c_0_2175])).
cnf(c_0_2181,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|r1(esk36_2(esk1_0,X1),esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2177, c_0_2169])).
cnf(c_0_2182,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|~p2(esk37_2(esk1_0,X1))|~r1(esk33_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2178, c_0_2169])).
cnf(c_0_2183,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|p2(X1)|~r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_2179]), c_0_2160]), c_0_2180])).
cnf(c_0_2184,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|r1(esk36_2(esk1_0,esk32_1(esk33_1(esk1_0))),esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2181, c_0_2174]), c_0_2175])).
cnf(c_0_2185,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)|~p2(esk37_2(esk1_0,esk32_1(esk33_1(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2182, c_0_2174]), c_0_2175])).
cnf(c_0_2186,negated_conjecture,(p2(esk33_1(esk1_0))|~p2(esk4_1(esk33_1(esk1_0)))|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_2160]), c_0_36])).
cnf(c_0_2187,negated_conjecture,(p2(esk4_1(esk33_1(esk1_0)))|p2(esk33_1(esk1_0))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2183, c_0_2184]), c_0_2185])).
cnf(c_0_2188,negated_conjecture,(p2(esk33_1(esk1_0))|~p5(X1)|~r1(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2186, c_0_2187]), c_0_36])).
cnf(c_0_2189,negated_conjecture,(p2(esk33_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2188, c_0_2146]), c_0_2147])])).
cnf(c_0_2190,plain,(r1(esk33_1(esk1_0),esk34_1(esk1_0))|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_76, c_0_2189])).
cnf(c_0_2191,negated_conjecture,(p2(esk2_0)|r1(esk33_1(esk1_0),esk34_1(esk1_0))), inference(spm,[status(thm)],[c_0_2190, c_0_2159])).
cnf(c_0_2192,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_2191]), c_0_2159])).
cnf(c_0_2193,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk34_1(esk1_0),esk36_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_2191]), c_0_2189])]), c_0_2159])).
cnf(c_0_2194,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk34_1(esk1_0))|p2(esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_2191]), c_0_2189])]), c_0_2159])).
cnf(c_0_2195,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_2191]), c_0_2159])).
cnf(c_0_2196,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X3)|p2(X2)|p2(X1)|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_2191]), c_0_2159])).
cnf(c_0_2197,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_2191]), c_0_2159])).
cnf(c_0_2198,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2192, c_0_2193]), c_0_2194]), c_0_2195])).
cnf(c_0_2199,plain,(r1(esk34_1(esk1_0),esk35_1(esk1_0))|~epred1_1(esk1_0)), inference(spm,[status(thm)],[c_0_88, c_0_2189])).
cnf(c_0_2200,plain,(~epred1_1(esk1_0)|~p2(esk35_1(esk1_0))), inference(spm,[status(thm)],[c_0_89, c_0_2189])).
cnf(c_0_2201,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2196, c_0_2193]), c_0_2194]), c_0_2197])).
cnf(c_0_2202,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_2191]), c_0_2159])).
cnf(c_0_2203,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_2191]), c_0_2159])).
cnf(c_0_2204,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_2191]), c_0_2159])).
cnf(c_0_2205,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_2191]), c_0_2159])).
cnf(c_0_2206,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X3)|p2(X1)|r1(X3,esk45_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_2191]), c_0_2159])).
cnf(c_0_2207,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_2191]), c_0_2159])).
cnf(c_0_2208,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_2198, c_0_7])).
cnf(c_0_2209,negated_conjecture,(p2(esk2_0)|r1(esk34_1(esk1_0),esk35_1(esk1_0))), inference(spm,[status(thm)],[c_0_2199, c_0_2159])).
cnf(c_0_2210,negated_conjecture,(p2(esk2_0)|~p2(esk35_1(esk1_0))), inference(spm,[status(thm)],[c_0_2200, c_0_2159])).
cnf(c_0_2211,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2201, c_0_7])).
cnf(c_0_2212,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2202, c_0_2193]), c_0_2194]), c_0_2203])).
cnf(c_0_2213,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2204, c_0_2193]), c_0_2194]), c_0_2205])).
cnf(c_0_2214,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2206, c_0_2193]), c_0_2194]), c_0_2207])).
cnf(c_0_2215,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_2191]), c_0_2159])).
cnf(c_0_2216,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_2191]), c_0_2159])).
cnf(c_0_2217,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2208, c_0_2209]), c_0_2210])).
cnf(c_0_2218,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2211, c_0_2209]), c_0_2210])).
cnf(c_0_2219,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_2212, c_0_7])).
cnf(c_0_2220,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_2213, c_0_7])).
cnf(c_0_2221,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2214, c_0_7])).
cnf(c_0_2222,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2215, c_0_2193]), c_0_2194]), c_0_2216])).
cnf(c_0_2223,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2217, c_0_2209]), c_0_2210])).
cnf(c_0_2224,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2218, c_0_2209]), c_0_2210])).
cnf(c_0_2225,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2219, c_0_2209]), c_0_2210])).
cnf(c_0_2226,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2220, c_0_2209]), c_0_2210])).
cnf(c_0_2227,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2221, c_0_2209]), c_0_2210])).
cnf(c_0_2228,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_2222, c_0_7])).
cnf(c_0_2229,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2223]), c_0_2189])]), c_0_2224]), c_0_2159])).
cnf(c_0_2230,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2225, c_0_2209]), c_0_2210])).
cnf(c_0_2231,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X3,esk41_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_2191]), c_0_2159])).
cnf(c_0_2232,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_2191]), c_0_2159])).
cnf(c_0_2233,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X3)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X3),esk46_3(esk1_0,esk34_1(esk1_0),X3))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X4,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_2191]), c_0_2159])).
cnf(c_0_2234,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_2191]), c_0_2159])).
cnf(c_0_2235,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2226, c_0_2209]), c_0_2210])).
cnf(c_0_2236,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2227, c_0_2209]), c_0_2210])).
cnf(c_0_2237,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2228, c_0_2209]), c_0_2210])).
cnf(c_0_2238,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2229, c_0_2230])).
cnf(c_0_2239,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2231, c_0_2193]), c_0_2194]), c_0_2232])).
cnf(c_0_2240,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2233, c_0_2193]), c_0_2194]), c_0_2234])).
cnf(c_0_2241,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2235]), c_0_2189])]), c_0_2236]), c_0_2159])).
cnf(c_0_2242,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2237, c_0_2209]), c_0_2210])).
cnf(c_0_2243,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_2238]), c_0_2209]), c_0_2191]), c_0_2194]), c_0_2159]), c_0_2210])).
cnf(c_0_2244,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_2239, c_0_7])).
cnf(c_0_2245,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2240, c_0_7])).
cnf(c_0_2246,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X3),esk42_3(esk1_0,esk34_1(esk1_0),X3))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X4)|~r1(esk34_1(esk1_0),X4)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X3)|~r1(X4,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_2191]), c_0_2159])).
cnf(c_0_2247,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_2191]), c_0_2159])).
cnf(c_0_2248,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2241, c_0_2242])).
cnf(c_0_2249,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2243, c_0_2209]), c_0_2210])).
cnf(c_0_2250,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2244, c_0_2209]), c_0_2210])).
cnf(c_0_2251,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2245, c_0_2209]), c_0_2210])).
cnf(c_0_2252,plain,(p2(esk2_0)|p2(X1)|p2(X2)|p2(X3)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X3)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2246, c_0_2193]), c_0_2194]), c_0_2247])).
cnf(c_0_2253,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2248]), c_0_2189])]), c_0_2249]), c_0_2159])).
cnf(c_0_2254,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2250, c_0_2209]), c_0_2210])).
cnf(c_0_2255,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2236]), c_0_2189])]), c_0_2224]), c_0_2159])).
cnf(c_0_2256,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2251, c_0_2209]), c_0_2210])).
cnf(c_0_2257,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_2252, c_0_7])).
cnf(c_0_2258,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2253, c_0_2254])).
cnf(c_0_2259,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2255, c_0_2256])).
cnf(c_0_2260,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2257, c_0_2209]), c_0_2210])).
cnf(c_0_2261,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|r1(esk36_2(esk1_0,esk34_1(esk1_0)),esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_2191]), c_0_2189])]), c_0_2159])).
cnf(c_0_2262,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2258]), c_0_2189])]), c_0_2259]), c_0_2159])).
cnf(c_0_2263,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2260, c_0_2209]), c_0_2210])).
cnf(c_0_2264,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2213, c_0_2261]), c_0_2205])).
cnf(c_0_2265,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2214, c_0_2261]), c_0_2207])).
cnf(c_0_2266,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2198, c_0_2261]), c_0_2195])).
cnf(c_0_2267,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2201, c_0_2261]), c_0_2197])).
cnf(c_0_2268,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2262, c_0_2263]), c_0_2253])).
cnf(c_0_2269,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2264, c_0_2209]), c_0_2210])).
cnf(c_0_2270,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2265, c_0_2209]), c_0_2210])).
cnf(c_0_2271,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2222, c_0_2261]), c_0_2216])).
cnf(c_0_2272,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2266, c_0_2209]), c_0_2210])).
cnf(c_0_2273,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2267, c_0_2209]), c_0_2210])).
cnf(c_0_2274,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2212, c_0_2261]), c_0_2203])).
cnf(c_0_2275,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_2268]), c_0_2209]), c_0_2191]), c_0_2194]), c_0_2159]), c_0_2210])).
cnf(c_0_2276,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2269, c_0_2209]), c_0_2210])).
cnf(c_0_2277,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2270, c_0_2209]), c_0_2210])).
cnf(c_0_2278,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2271, c_0_2209]), c_0_2210])).
cnf(c_0_2279,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2272, c_0_2209]), c_0_2210])).
cnf(c_0_2280,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2273, c_0_2209]), c_0_2210])).
cnf(c_0_2281,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2274, c_0_2209]), c_0_2210])).
cnf(c_0_2282,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2275, c_0_2209]), c_0_2210])).
cnf(c_0_2283,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_2268]), c_0_2209]), c_0_2191]), c_0_2194]), c_0_2159]), c_0_2210])).
cnf(c_0_2284,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2276]), c_0_2189])]), c_0_2277]), c_0_2159])).
cnf(c_0_2285,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2278, c_0_2209]), c_0_2210])).
cnf(c_0_2286,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2279]), c_0_2189])]), c_0_2280]), c_0_2159])).
cnf(c_0_2287,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2281, c_0_2209]), c_0_2210])).
cnf(c_0_2288,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2239, c_0_2261]), c_0_2232])).
cnf(c_0_2289,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2240, c_0_2261]), c_0_2234])).
cnf(c_0_2290,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2282]), c_0_2189])]), c_0_2249]), c_0_2159])).
cnf(c_0_2291,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2283, c_0_2209]), c_0_2210])).
cnf(c_0_2292,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2284, c_0_2285])).
cnf(c_0_2293,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2286, c_0_2287])).
cnf(c_0_2294,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2288, c_0_2209]), c_0_2210])).
cnf(c_0_2295,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2289, c_0_2209]), c_0_2210])).
cnf(c_0_2296,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2290, c_0_2291])).
cnf(c_0_2297,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2292]), c_0_2189])]), c_0_2293]), c_0_2159])).
cnf(c_0_2298,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2294, c_0_2209]), c_0_2210])).
cnf(c_0_2299,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2277]), c_0_2189])]), c_0_2280]), c_0_2159])).
cnf(c_0_2300,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2295, c_0_2209]), c_0_2210])).
cnf(c_0_2301,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2252, c_0_2261]), c_0_2247])).
cnf(c_0_2302,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_2296]), c_0_2209]), c_0_2191]), c_0_2194]), c_0_2159]), c_0_2210])).
cnf(c_0_2303,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_2296]), c_0_2209]), c_0_2191]), c_0_2194]), c_0_2159]), c_0_2210])).
cnf(c_0_2304,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2297, c_0_2298])).
cnf(c_0_2305,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2299, c_0_2300])).
cnf(c_0_2306,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2301, c_0_2209]), c_0_2210])).
cnf(c_0_2307,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2302, c_0_2209]), c_0_2210])).
cnf(c_0_2308,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2303, c_0_2209]), c_0_2210])).
cnf(c_0_2309,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_2296]), c_0_2209]), c_0_2191]), c_0_2194]), c_0_2159]), c_0_2210])).
cnf(c_0_2310,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2304]), c_0_2189])]), c_0_2305]), c_0_2159])).
cnf(c_0_2311,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2306, c_0_2209]), c_0_2210])).
cnf(c_0_2312,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2307]), c_0_2189])]), c_0_2308]), c_0_2159])).
cnf(c_0_2313,negated_conjecture,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2309, c_0_2209]), c_0_2210])).
cnf(c_0_2314,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2310, c_0_2311]), c_0_2297])).
cnf(c_0_2315,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2312, c_0_2313])).
cnf(c_0_2316,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_2314]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2317,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_2314]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2318,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_2293]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2319,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_2293]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2320,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_2315]), c_0_2209]), c_0_2191]), c_0_2194]), c_0_2159]), c_0_2210])).
cnf(c_0_2321,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2316, c_0_2193]), c_0_2194]), c_0_2317])).
cnf(c_0_2322,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2318, c_0_2193]), c_0_2194]), c_0_2319])).
cnf(c_0_2323,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_2314]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2324,plain,(p2(esk36_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2320, c_0_2296]), c_0_2209]), c_0_2210])).
cnf(c_0_2325,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_2314]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2326,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2321, c_0_2261]), c_0_2317])).
cnf(c_0_2327,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2322, c_0_2261]), c_0_2319])).
cnf(c_0_2328,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2323, c_0_2193]), c_0_2324]), c_0_2325])).
cnf(c_0_2329,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2326, c_0_2209]), c_0_2210])).
cnf(c_0_2330,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2327, c_0_2209]), c_0_2210])).
cnf(c_0_2331,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2328, c_0_2261]), c_0_2325])).
cnf(c_0_2332,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2329]), c_0_2189])]), c_0_2330]), c_0_2159])).
cnf(c_0_2333,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2331, c_0_2209]), c_0_2210])).
cnf(c_0_2334,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2332, c_0_2333])).
cnf(c_0_2335,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_2334]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2336,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_2334]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2337,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X2)|p2(X1)|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X1)|~r1(X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_2334]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2338,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_2334]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2339,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2335, c_0_2193]), c_0_2324]), c_0_2336])).
cnf(c_0_2340,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2337, c_0_2193]), c_0_2324]), c_0_2338])).
cnf(c_0_2341,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_2334]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2342,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~p2(esk34_1(esk1_0))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_2334]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2343,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2339, c_0_2261]), c_0_2336])).
cnf(c_0_2344,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2340, c_0_2261]), c_0_2338])).
cnf(c_0_2345,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X2)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2341, c_0_2193]), c_0_2324]), c_0_2342])).
cnf(c_0_2346,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2343, c_0_2209]), c_0_2210])).
cnf(c_0_2347,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2344, c_0_2209]), c_0_2210])).
cnf(c_0_2348,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2345, c_0_2261]), c_0_2342])).
cnf(c_0_2349,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2346]), c_0_2189])]), c_0_2347]), c_0_2159])).
cnf(c_0_2350,negated_conjecture,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2348, c_0_2209]), c_0_2210])).
cnf(c_0_2351,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2349, c_0_2350])).
cnf(c_0_2352,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|p2(X2)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X2))|~p2(X3)|~r1(esk34_1(esk1_0),X3)|~r1(esk34_1(esk1_0),X2)|~r1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_2351]), c_0_2209]), c_0_2191]), c_0_2159]), c_0_2210])).
cnf(c_0_2353,plain,(p2(esk37_2(esk1_0,esk34_1(esk1_0)))|p2(esk2_0)|p2(X1)|~p2(X2)|~r1(esk34_1(esk1_0),X2)|~r1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2352, c_0_2334]), c_0_2209]), c_0_2210])).
cnf(c_0_2354,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|~p2(esk37_2(esk1_0,esk34_1(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_2191]), c_0_2189])]), c_0_2159])).
cnf(c_0_2355,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)|p2(X1)|~r1(esk36_2(esk1_0,esk34_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2353, c_0_2193]), c_0_2324]), c_0_2354])).
cnf(c_0_2356,plain,(p2(esk34_1(esk1_0))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2355, c_0_2261]), c_0_2354])).
cnf(c_0_2357,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(X2,esk41_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_2195, c_0_2356])).
cnf(c_0_2358,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X2))|p2(esk2_0)|p2(X2)|p2(X1)|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_2197, c_0_2356])).
cnf(c_0_2359,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2357, c_0_2209]), c_0_2210])).
cnf(c_0_2360,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2358, c_0_2209]), c_0_2210])).
cnf(c_0_2361,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X2)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X2),esk42_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X1)|~r1(esk34_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_2203, c_0_2356])).
cnf(c_0_2362,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2205, c_0_2356])).
cnf(c_0_2363,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2207, c_0_2356])).
cnf(c_0_2364,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2359, c_0_2209]), c_0_2210])).
cnf(c_0_2365,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2360, c_0_2209]), c_0_2210])).
cnf(c_0_2366,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2361, c_0_2209]), c_0_2210])).
cnf(c_0_2367,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2362, c_0_2209]), c_0_2210])).
cnf(c_0_2368,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2363, c_0_2209]), c_0_2210])).
cnf(c_0_2369,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(X2,esk45_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2216, c_0_2356])).
cnf(c_0_2370,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2364]), c_0_2189])]), c_0_2365]), c_0_2159])).
cnf(c_0_2371,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2366, c_0_2209]), c_0_2210])).
cnf(c_0_2372,negated_conjecture,(p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2367, c_0_2209]), c_0_2210])).
cnf(c_0_2373,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2368, c_0_2209]), c_0_2210])).
cnf(c_0_2374,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2369, c_0_2209]), c_0_2210])).
cnf(c_0_2375,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2370, c_0_2371])).
cnf(c_0_2376,plain,(p2(esk2_0)|p2(X1)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2372]), c_0_2189])]), c_0_2373]), c_0_2159])).
cnf(c_0_2377,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2374, c_0_2209]), c_0_2210])).
cnf(c_0_2378,plain,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk45_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_2375]), c_0_2209]), c_0_2191]), c_0_2356]), c_0_2159]), c_0_2210])).
cnf(c_0_2379,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2232, c_0_2356])).
cnf(c_0_2380,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2234, c_0_2356])).
cnf(c_0_2381,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2376, c_0_2377])).
cnf(c_0_2382,negated_conjecture,(p2(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2378, c_0_2209]), c_0_2210])).
cnf(c_0_2383,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2379, c_0_2209]), c_0_2210])).
cnf(c_0_2384,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2380, c_0_2209]), c_0_2210])).
cnf(c_0_2385,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2381]), c_0_2189])]), c_0_2382]), c_0_2159])).
cnf(c_0_2386,negated_conjecture,(p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2383, c_0_2209]), c_0_2210])).
cnf(c_0_2387,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2373]), c_0_2189])]), c_0_2365]), c_0_2159])).
cnf(c_0_2388,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2384, c_0_2209]), c_0_2210])).
cnf(c_0_2389,plain,(p2(esk2_0)|p2(X1)|p2(X2)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X2),esk46_3(esk1_0,esk34_1(esk1_0),X2))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X2)|~r1(esk34_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_2247, c_0_2356])).
cnf(c_0_2390,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(spm,[status(thm)],[c_0_2385, c_0_2386])).
cnf(c_0_2391,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2387, c_0_2388])).
cnf(c_0_2392,negated_conjecture,(p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2389, c_0_2209]), c_0_2210])).
cnf(c_0_2393,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2390]), c_0_2189])]), c_0_2391]), c_0_2159])).
cnf(c_0_2394,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2392, c_0_2209]), c_0_2210])).
cnf(c_0_2395,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2393, c_0_2394]), c_0_2385])).
cnf(c_0_2396,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(X1,esk45_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_2395]), c_0_2209]), c_0_2191]), c_0_2356]), c_0_2159]), c_0_2210])).
cnf(c_0_2397,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk35_1(esk1_0),esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2396, c_0_2209]), c_0_2210])).
cnf(c_0_2398,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|r1(esk45_3(esk1_0,esk34_1(esk1_0),X1),esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_2395]), c_0_2209]), c_0_2191]), c_0_2356]), c_0_2159]), c_0_2210])).
cnf(c_0_2399,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|p2(X1)|~r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2397]), c_0_2189])]), c_0_2382]), c_0_2159])).
cnf(c_0_2400,negated_conjecture,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)|r1(esk45_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2398, c_0_2209]), c_0_2210])).
cnf(c_0_2401,plain,(p2(esk46_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2399, c_0_2400])).
cnf(c_0_2402,plain,(p2(esk2_0)|p2(X1)|r1(X1,esk41_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_2401]), c_0_2209]), c_0_2191]), c_0_2356]), c_0_2159]), c_0_2210])).
cnf(c_0_2403,plain,(p2(esk41_3(esk1_0,esk34_1(esk1_0),X1))|p2(esk2_0)|p2(X1)|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_2401]), c_0_2209]), c_0_2191]), c_0_2356]), c_0_2159]), c_0_2210])).
cnf(c_0_2404,negated_conjecture,(p2(esk2_0)|r1(esk35_1(esk1_0),esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2402, c_0_2209]), c_0_2210])).
cnf(c_0_2405,negated_conjecture,(p2(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2403, c_0_2209]), c_0_2210])).
cnf(c_0_2406,plain,(p2(esk2_0)|p2(X1)|r1(esk41_3(esk1_0,esk34_1(esk1_0),X1),esk42_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_2401]), c_0_2209]), c_0_2191]), c_0_2356]), c_0_2159]), c_0_2210])).
cnf(c_0_2407,plain,(p2(esk2_0)|p2(X1)|~r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_2404]), c_0_2189])]), c_0_2405]), c_0_2159])).
cnf(c_0_2408,negated_conjecture,(p2(esk2_0)|r1(esk41_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)),esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2406, c_0_2209]), c_0_2210])).
cnf(c_0_2409,plain,(p2(esk42_3(esk1_0,esk34_1(esk1_0),esk35_1(esk1_0)))|p2(esk2_0)), inference(spm,[status(thm)],[c_0_2407, c_0_2408])).
cnf(c_0_2410,plain,(p2(esk2_0)|p2(X1)|~p2(esk46_3(esk1_0,esk34_1(esk1_0),X1))|~r1(esk34_1(esk1_0),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_2409]), c_0_2209]), c_0_2191]), c_0_2356]), c_0_2159]), c_0_2210])).
cnf(c_0_2411,plain,(p2(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_2410, c_0_2401]), c_0_2209]), c_0_2210])).
cnf(c_0_2412,negated_conjecture,(~p5(X1)|~r1(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_2411])])).
cnf(c_0_2413,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2412, c_0_2146]), c_0_2147])]), ['proof']).
# SZS output end CNFRefutation.
EOF
