% Problem    : LCL652+1.005 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n127.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:19:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____042_B03_F1_AE_Q4_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (~(?[X1]:~(((((((~(![X2]:(~(r1(X1,X2))|~(p4(X2))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((((~(![X1]:(~(r1(X2,X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p3(X1)))|~(p2(X2))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((![X1]:(~(r1(X2,X1))|p3(X1))|~(p2(X2))))))&p2(X1))&~(![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p2(X2)))))))))))))|(~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p3(X2)))|~(p2(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p2(X2)))))))|~(![X1]:(~(r1(X2,X1))|~((![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p2(X1)))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p3(X1)))|~(p2(X2))))))))))))|(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p3(X2)))|~(p2(X1)))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p2(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p2(X2)))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', main)).
fof(c_0_1, plain, (![X2]:(epred2_1(X2)<=>~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((((~(![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2)))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((![X1]:(~r1(X2,X1)|p3(X1))|~p2(X2)))))&p2(X1))&~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))))))))))|(~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))))|~(![X1]:(~r1(X2,X1)|~((![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2)))))))))))|(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))), introduced(definition)).
fof(c_0_2, plain, (![X2]:(epred1_1(X2)<=>~(((((~(![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2)))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((![X1]:(~r1(X2,X1)|p3(X1))|~p2(X2)))))&p2(X1))&~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))))))))))|(~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))))|~(![X1]:(~r1(X2,X1)|~((![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2)))))))))))|(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))))))), introduced(definition)).
fof(c_0_3, plain, (![X2]:(epred2_1(X2)=>~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|epred1_1(X2))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_1]), c_0_2])).
fof(c_0_4, negated_conjecture, (~(~(?[X1]:~(((((((~(![X2]:(~r1(X1,X2)|~p4(X2)))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))))|~(![X2]:(~r1(X1,X2)|epred2_1(X2))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p2(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_1])).
fof(c_0_5, plain, (![X3]:![X4]:![X5]:![X6]:![X9]:![X10]:![X11]:![X12]:![X13]:![X16]:![X17]:![X18]:![X19]:![X20]:![X23]:![X24]:![X25]:![X26]:![X27]:![X30]:![X31]:![X32]:![X33]:![X34]:![X36]:![X37]:![X38]:![X40]:![X41]:![X42]:![X44]:![X45]:![X46]:(((((((r1(X4,esk46_2(X3,X4))|((r1(X6,esk45_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred2_1(X3))&((((~r1(esk46_2(X3,X4),X9)|~p1(X9))|((r1(X6,esk45_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred2_1(X3)))&((((r1(X4,esk46_2(X3,X4))|((p1(esk45_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred2_1(X3))&((((~r1(esk46_2(X3,X4),X9)|~p1(X9))|((p1(esk45_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred2_1(X3))))&((((((((r1(X11,esk48_3(X3,X10,X11))|((r1(X13,esk47_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred2_1(X3))&(((((~r1(esk48_3(X3,X10,X11),X16)|~p1(X16))|((r1(X13,esk47_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred2_1(X3)))&(((((r1(X11,esk48_3(X3,X10,X11))|((p1(esk47_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred2_1(X3))&(((((~r1(esk48_3(X3,X10,X11),X16)|~p1(X16))|((p1(esk47_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred2_1(X3))))&(((((((((r1(X18,esk50_4(X3,X10,X17,X18))|((r1(X20,esk49_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&((((((~r1(esk50_4(X3,X10,X17,X18),X23)|~p1(X23))|((r1(X20,esk49_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))&((((((r1(X18,esk50_4(X3,X10,X17,X18))|((p1(esk49_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&((((((~r1(esk50_4(X3,X10,X17,X18),X23)|~p1(X23))|((p1(esk49_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))))&((((((((((r1(X25,esk52_5(X3,X10,X17,X24,X25))|((r1(X27,esk51_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&(((((((~r1(esk52_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((r1(X27,esk51_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))&(((((((r1(X25,esk52_5(X3,X10,X17,X24,X25))|((p1(esk51_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&(((((((~r1(esk52_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((p1(esk51_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))))&(((((~r1(X24,X31)|epred1_1(X31))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))&(((((((r1(X32,esk53_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&((((((~p1(esk53_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))))&((((((r1(X36,esk54_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&(((((~p1(esk54_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))))&(((((r1(X40,esk55_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred2_1(X3))&((((~p1(esk55_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred2_1(X3)))))&((((r1(X44,esk56_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred2_1(X3))&(((~p1(esk56_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred2_1(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
fof(c_0_6, negated_conjecture, (![X4]:![X5]:![X6]:![X7]:![X10]:![X11]:![X12]:![X13]:![X14]:![X18]:![X23]:![X28]:![X33]:![X38]:![X41]:![X43]:![X44]:![X45]:![X47]:![X48]:![X49]:![X51]:![X52]:![X53]:![X55]:![X56]:![X57]:![X59]:![X60]:![X61]:(((((((~r1(esk1_0,X4)|~p4(X4))&((((r1(X5,esk3_1(X5))|((r1(X7,esk2_3(X5,X6,X7))|~r1(X6,X7))|~r1(X5,X6)))|~r1(esk1_0,X5))&(((~r1(esk3_1(X5),X10)|~p1(X10))|((r1(X7,esk2_3(X5,X6,X7))|~r1(X6,X7))|~r1(X5,X6)))|~r1(esk1_0,X5)))&(((r1(X5,esk3_1(X5))|((p1(esk2_3(X5,X6,X7))|~r1(X6,X7))|~r1(X5,X6)))|~r1(esk1_0,X5))&(((~r1(esk3_1(X5),X10)|~p1(X10))|((p1(esk2_3(X5,X6,X7))|~r1(X6,X7))|~r1(X5,X6)))|~r1(esk1_0,X5)))))&(~r1(esk1_0,X11)|epred2_1(X11)))&(((r1(X12,esk4_1(X12))|(~r1(X12,X13)|(~r1(X13,X14)|p1(X14))))|~r1(esk1_0,X12))&((~p1(esk4_1(X12))|(~r1(X12,X13)|(~r1(X13,X14)|p1(X14))))|~r1(esk1_0,X12))))&((r1(esk1_0,esk5_0)&(r1(esk5_0,esk6_0)&(~r1(esk6_0,X18)|p1(X18))))&(r1(esk5_0,esk7_0)&~p1(esk7_0))))&(((r1(esk1_0,esk8_0)&((r1(esk8_0,esk9_0)&(r1(esk9_0,esk10_0)&(~r1(esk10_0,X23)|p1(X23))))&(r1(esk9_0,esk11_0)&~p1(esk11_0))))&(((r1(esk8_0,esk12_0)&((r1(esk12_0,esk13_0)&(r1(esk13_0,esk14_0)&(~r1(esk14_0,X28)|p1(X28))))&(r1(esk13_0,esk15_0)&~p1(esk15_0))))&(((r1(esk12_0,esk16_0)&((r1(esk16_0,esk17_0)&(r1(esk17_0,esk18_0)&(~r1(esk18_0,X33)|p1(X33))))&(r1(esk17_0,esk19_0)&~p1(esk19_0))))&(((r1(esk16_0,esk20_0)&((r1(esk20_0,esk21_0)&(r1(esk21_0,esk22_0)&(~r1(esk22_0,X38)|p1(X38))))&(r1(esk21_0,esk23_0)&~p1(esk23_0))))&(r1(esk20_0,esk24_0)&((r1(X41,esk25_1(X41))|~r1(esk24_0,X41))&(~p2(esk25_1(X41))|~r1(esk24_0,X41)))))&((r1(X43,esk26_1(X43))|(~r1(esk20_0,X43)|(~r1(X43,X44)|(~r1(X44,X45)|~p1(X45)))))&(p1(esk26_1(X43))|(~r1(esk20_0,X43)|(~r1(X43,X44)|(~r1(X44,X45)|~p1(X45))))))))&((r1(X47,esk27_1(X47))|(~r1(esk16_0,X47)|(~r1(X47,X48)|(~r1(X48,X49)|~p1(X49)))))&(p1(esk27_1(X47))|(~r1(esk16_0,X47)|(~r1(X47,X48)|(~r1(X48,X49)|~p1(X49))))))))&((r1(X51,esk28_1(X51))|(~r1(esk12_0,X51)|(~r1(X51,X52)|(~r1(X52,X53)|~p1(X53)))))&(p1(esk28_1(X51))|(~r1(esk12_0,X51)|(~r1(X51,X52)|(~r1(X52,X53)|~p1(X53))))))))&((r1(X55,esk29_1(X55))|(~r1(esk8_0,X55)|(~r1(X55,X56)|(~r1(X56,X57)|~p1(X57)))))&(p1(esk29_1(X55))|(~r1(esk8_0,X55)|(~r1(X55,X56)|(~r1(X56,X57)|~p1(X57))))))))&((r1(X59,esk30_1(X59))|(~r1(esk1_0,X59)|(~r1(X59,X60)|(~r1(X60,X61)|~p1(X61)))))&(p1(esk30_1(X59))|(~r1(esk1_0,X59)|(~r1(X59,X60)|(~r1(X60,X61)|~p1(X61)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
cnf(c_0_7,plain,(epred1_1(X5)|~epred2_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(r1(esk20_0,esk24_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,negated_conjecture,(epred1_1(esk24_0)|~epred2_1(X1)|~r1(X2,esk20_0)|~r1(X3,X2)|~r1(X1,X3)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_10,negated_conjecture,(r1(esk16_0,esk20_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_11, plain, (![X2]:(epred1_1(X2)=>~(((((~(![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2)))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((![X1]:(~r1(X2,X1)|p3(X1))|~p2(X2)))))&p2(X1))&~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))))))))))|(~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))))|~(![X1]:(~r1(X2,X1)|~((![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2)))))))))))|(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))))))), inference(split_equiv,[status(thm)],[c_0_2])).
cnf(c_0_12,negated_conjecture,(epred1_1(esk24_0)|~epred2_1(X1)|~r1(X2,esk16_0)|~r1(X1,X2)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13,negated_conjecture,(r1(esk12_0,esk16_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14,negated_conjecture,(epred2_1(X1)|~r1(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,negated_conjecture,(r1(esk1_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_16, plain, (![X3]:![X4]:![X7]:![X8]:![X10]:![X12]:![X13]:![X14]:![X18]:![X19]:![X21]:![X23]:![X29]:((((((((r1(X4,esk31_2(X3,X4))|~r1(X3,X4))|~epred1_1(X3))&(((r1(esk31_2(X3,X4),esk32_2(X3,X4))|~r1(X3,X4))|~epred1_1(X3))&((~p3(esk32_2(X3,X4))|~r1(X3,X4))|~epred1_1(X3))))&((p2(esk31_2(X3,X4))|~r1(X3,X4))|~epred1_1(X3)))&(((((((r1(X10,esk34_3(X3,X7,X10))|~r1(X7,X10))|((r1(X8,esk33_3(X3,X7,X8))|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3))&(((((~r1(esk34_3(X3,X7,X10),X12)|~p2(X12))|~r1(X7,X10))|((r1(X8,esk33_3(X3,X7,X8))|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3)))&(((((r1(X10,esk34_3(X3,X7,X10))|~r1(X7,X10))|((~p3(esk33_3(X3,X7,X8))|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3))&(((((~r1(esk34_3(X3,X7,X10),X12)|~p2(X12))|~r1(X7,X10))|((~p3(esk33_3(X3,X7,X8))|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3))))&(((((r1(X10,esk34_3(X3,X7,X10))|~r1(X7,X10))|((p2(X8)|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3))&(((((~r1(esk34_3(X3,X7,X10),X12)|~p2(X12))|~r1(X7,X10))|((p2(X8)|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3)))))&(((((r1(X3,esk37_1(X3))|((r1(X14,esk35_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3))&(((~r1(esk37_1(X3),X18)|~p2(X18))|((r1(X14,esk35_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3)))&((((r1(X3,esk37_1(X3))|((r1(esk35_3(X3,X13,X14),esk36_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3))&(((~r1(esk37_1(X3),X18)|~p2(X18))|((r1(esk35_3(X3,X13,X14),esk36_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3)))&(((r1(X3,esk37_1(X3))|((~p3(esk36_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3))&(((~r1(esk37_1(X3),X18)|~p2(X18))|((~p3(esk36_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3)))))&(((r1(X3,esk37_1(X3))|((p2(esk35_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3))&(((~r1(esk37_1(X3),X18)|~p2(X18))|((p2(esk35_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3)))))&(((((r1(X19,esk39_2(X3,X19))|r1(X19,esk38_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))&((((((r1(X23,esk40_3(X3,X19,X23))|~r1(esk39_2(X3,X19),X23))|r1(X19,esk38_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))&(((((r1(esk40_3(X3,X19,X23),esk41_3(X3,X19,X23))|~r1(esk39_2(X3,X19),X23))|r1(X19,esk38_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))&((((~p3(esk41_3(X3,X19,X23))|~r1(esk39_2(X3,X19),X23))|r1(X19,esk38_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))))&((((p2(esk40_3(X3,X19,X23))|~r1(esk39_2(X3,X19),X23))|r1(X19,esk38_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))))&((((r1(X19,esk39_2(X3,X19))|(~r1(esk38_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))&((((((r1(X23,esk40_3(X3,X19,X23))|~r1(esk39_2(X3,X19),X23))|(~r1(esk38_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))&(((((r1(esk40_3(X3,X19,X23),esk41_3(X3,X19,X23))|~r1(esk39_2(X3,X19),X23))|(~r1(esk38_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))&((((~p3(esk41_3(X3,X19,X23))|~r1(esk39_2(X3,X19),X23))|(~r1(esk38_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))))&((((p2(esk40_3(X3,X19,X23))|~r1(esk39_2(X3,X19),X23))|(~r1(esk38_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))))))&(((((r1(X3,esk44_1(X3))|r1(X3,esk42_1(X3)))|~epred1_1(X3))&(((~r1(esk44_1(X3),X29)|~p2(X29))|r1(X3,esk42_1(X3)))|~epred1_1(X3)))&((((r1(X3,esk44_1(X3))|r1(esk42_1(X3),esk43_1(X3)))|~epred1_1(X3))&(((~r1(esk44_1(X3),X29)|~p2(X29))|r1(esk42_1(X3),esk43_1(X3)))|~epred1_1(X3)))&(((r1(X3,esk44_1(X3))|~p3(esk43_1(X3)))|~epred1_1(X3))&(((~r1(esk44_1(X3),X29)|~p2(X29))|~p3(esk43_1(X3)))|~epred1_1(X3)))))&(((r1(X3,esk44_1(X3))|p2(esk42_1(X3)))|~epred1_1(X3))&(((~r1(esk44_1(X3),X29)|~p2(X29))|p2(esk42_1(X3)))|~epred1_1(X3)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
cnf(c_0_17,negated_conjecture,(epred1_1(esk24_0)|~epred2_1(X1)|~r1(X1,esk12_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18,negated_conjecture,(r1(esk8_0,esk12_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_19,negated_conjecture,(epred2_1(esk8_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20,plain,(r1(X1,esk42_1(X1))|r1(X1,esk44_1(X1))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21,negated_conjecture,(epred1_1(esk24_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
cnf(c_0_22,plain,(r1(X2,esk31_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(r1(esk24_0,esk42_1(esk24_0))|r1(esk24_0,esk44_1(esk24_0))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24,plain,(p2(esk42_1(X1))|r1(X1,esk44_1(X1))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,plain,(r1(X1,esk42_1(X1))|~epred1_1(X1)|~p2(X2)|~r1(esk44_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26,plain,(r1(esk44_1(esk24_0),esk31_2(esk24_0,esk44_1(esk24_0)))|r1(esk24_0,esk42_1(esk24_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_21])])).
cnf(c_0_27,plain,(p2(esk31_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28,plain,(p2(esk42_1(esk24_0))|r1(esk24_0,esk44_1(esk24_0))), inference(spm,[status(thm)],[c_0_24, c_0_21])).
cnf(c_0_29,plain,(r1(esk24_0,esk42_1(esk24_0))|~p2(esk31_2(esk24_0,esk44_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_21])])).
cnf(c_0_30,plain,(p2(esk31_2(esk24_0,esk44_1(esk24_0)))|r1(esk24_0,esk42_1(esk24_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_23]), c_0_21])])).
cnf(c_0_31,plain,(p2(esk42_1(X1))|~epred1_1(X1)|~p2(X2)|~r1(esk44_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_32,plain,(p2(esk42_1(esk24_0))|r1(esk44_1(esk24_0),esk31_2(esk24_0,esk44_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_28]), c_0_21])])).
cnf(c_0_33,plain,(r1(X2,esk38_2(X1,X2))|r1(X2,esk39_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_34,plain,(r1(esk24_0,esk42_1(esk24_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35,plain,(p2(esk42_1(esk24_0))|~p2(esk31_2(esk24_0,esk44_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_21])])).
cnf(c_0_36,plain,(p2(esk31_2(esk24_0,esk44_1(esk24_0)))|p2(esk42_1(esk24_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_21])])).
cnf(c_0_37,plain,(p2(X3)|r1(X4,esk34_3(X1,X2,X4))|~epred1_1(X1)|~r1(X1,X2)|~p2(X2)|~r1(X2,X3)|~r1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_38,plain,(r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))|r1(esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_21])])).
cnf(c_0_39,plain,(p2(esk42_1(esk24_0))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40,negated_conjecture,(r1(X1,esk25_1(X1))|~r1(esk24_0,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_41,plain,(p2(X1)|r1(esk39_2(esk24_0,esk42_1(esk24_0)),esk34_3(X2,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))))|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))|~epred1_1(X2)|~r1(esk42_1(esk24_0),X1)|~r1(X2,esk42_1(esk24_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
cnf(c_0_42,negated_conjecture,(r1(esk42_1(esk24_0),esk25_1(esk42_1(esk24_0)))), inference(spm,[status(thm)],[c_0_40, c_0_34])).
cnf(c_0_43,negated_conjecture,(p2(esk25_1(esk42_1(esk24_0)))|r1(esk39_2(esk24_0,esk42_1(esk24_0)),esk34_3(X1,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))))|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))|~epred1_1(X1)|~r1(X1,esk42_1(esk24_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_44,negated_conjecture,(~r1(esk24_0,X1)|~p2(esk25_1(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_45,plain,(p2(esk25_1(esk42_1(esk24_0)))|r1(esk39_2(esk24_0,esk42_1(esk24_0)),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))))|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_34]), c_0_21])])).
cnf(c_0_46,plain,(r1(X2,esk38_2(X1,X2))|r1(X3,esk40_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(esk39_2(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_47,negated_conjecture,(r1(esk39_2(esk24_0,esk42_1(esk24_0)),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))))|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_34])])).
cnf(c_0_48,plain,(p2(X3)|~epred1_1(X1)|~r1(X1,X2)|~p2(X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(esk34_3(X1,X2,X4),X5)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_49,plain,(r1(esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))),esk40_3(esk24_0,esk42_1(esk24_0),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0)))))|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_21]), c_0_34])])).
cnf(c_0_50,plain,(p2(X1)|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))|~p2(esk40_3(esk24_0,esk42_1(esk24_0),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0)))))|~r1(esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0)))|~r1(esk42_1(esk24_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_21]), c_0_39]), c_0_34])])).
cnf(c_0_51,plain,(r1(X2,esk38_2(X1,X2))|p2(esk40_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(esk39_2(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_52,plain,(p2(X1)|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))|~p2(esk40_3(esk24_0,esk42_1(esk24_0),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0)))))|~r1(esk42_1(esk24_0),X1)), inference(spm,[status(thm)],[c_0_50, c_0_38])).
cnf(c_0_53,plain,(p2(esk40_3(esk24_0,esk42_1(esk24_0),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0)))))|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_47]), c_0_21]), c_0_34])])).
cnf(c_0_54,plain,(p2(X1)|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))|~r1(esk42_1(esk24_0),X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_55,negated_conjecture,(p2(esk25_1(esk42_1(esk24_0)))|r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))), inference(spm,[status(thm)],[c_0_54, c_0_42])).
cnf(c_0_56,plain,(r1(X3,esk35_3(X1,X2,X3))|r1(X1,esk37_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_57,negated_conjecture,(r1(esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_55]), c_0_34])])).
cnf(c_0_58,plain,(p2(esk35_3(X1,X2,X3))|r1(X1,esk37_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_59,plain,(r1(esk38_2(esk24_0,esk42_1(esk24_0)),esk35_3(X1,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|r1(X1,esk37_1(X1))|~epred1_1(X1)|~r1(X1,esk42_1(esk24_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_60,plain,(p2(esk35_3(X1,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|r1(X1,esk37_1(X1))|~epred1_1(X1)|~r1(X1,esk42_1(esk24_0))), inference(spm,[status(thm)],[c_0_58, c_0_57])).
cnf(c_0_61,plain,(r1(esk38_2(esk24_0,esk42_1(esk24_0)),esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|r1(esk24_0,esk37_1(esk24_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_34]), c_0_21])])).
cnf(c_0_62,plain,(p2(esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|r1(esk24_0,esk37_1(esk24_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_34]), c_0_21])])).
cnf(c_0_63,plain,(r1(X3,esk35_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p2(X4)|~r1(esk37_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_64,plain,(r1(esk38_2(esk24_0,esk42_1(esk24_0)),esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|r1(esk37_1(esk24_0),esk31_2(esk24_0,esk37_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_61]), c_0_21])])).
cnf(c_0_65,plain,(p2(esk35_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p2(X4)|~r1(esk37_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_66,plain,(p2(esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|r1(esk37_1(esk24_0),esk31_2(esk24_0,esk37_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_62]), c_0_21])])).
cnf(c_0_67,plain,(r1(esk38_2(esk24_0,esk42_1(esk24_0)),esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|r1(X1,esk35_3(esk24_0,X2,X1))|~p2(esk31_2(esk24_0,esk37_1(esk24_0)))|~r1(esk24_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_21])])).
cnf(c_0_68,plain,(p2(esk31_2(esk24_0,esk37_1(esk24_0)))|r1(esk38_2(esk24_0,esk42_1(esk24_0)),esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_61]), c_0_21])])).
cnf(c_0_69,plain,(p2(esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|p2(esk35_3(esk24_0,X1,X2))|~p2(esk31_2(esk24_0,esk37_1(esk24_0)))|~r1(esk24_0,X1)|~r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_21])])).
cnf(c_0_70,plain,(p2(esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|p2(esk31_2(esk24_0,esk37_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_62]), c_0_21])])).
cnf(c_0_71,plain,(r1(esk38_2(esk24_0,esk42_1(esk24_0)),esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|r1(X1,esk35_3(esk24_0,X2,X1))|~r1(esk24_0,X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_72,plain,(p2(esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|p2(esk35_3(esk24_0,X1,X2))|~r1(esk24_0,X1)|~r1(X1,X2)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_73,plain,(r1(esk38_2(esk24_0,esk42_1(esk24_0)),esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|r1(X1,esk35_3(esk24_0,esk42_1(esk24_0),X1))|~r1(esk42_1(esk24_0),X1)), inference(spm,[status(thm)],[c_0_71, c_0_34])).
cnf(c_0_74,plain,(p2(esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))|p2(esk35_3(esk24_0,esk42_1(esk24_0),X1))|~r1(esk42_1(esk24_0),X1)), inference(spm,[status(thm)],[c_0_72, c_0_34])).
cnf(c_0_75,plain,(r1(X2,esk39_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)|~p2(X3)|~r1(esk38_2(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_76,negated_conjecture,(r1(esk38_2(esk24_0,esk42_1(esk24_0)),esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))), inference(spm,[status(thm)],[c_0_73, c_0_57])).
cnf(c_0_77,negated_conjecture,(p2(esk35_3(esk24_0,esk42_1(esk24_0),esk38_2(esk24_0,esk42_1(esk24_0))))), inference(spm,[status(thm)],[c_0_74, c_0_57])).
cnf(c_0_78,plain,(r1(esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_21]), c_0_77]), c_0_34])])).
cnf(c_0_79,plain,(p2(X1)|r1(esk39_2(esk24_0,esk42_1(esk24_0)),esk34_3(X2,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))))|~epred1_1(X2)|~r1(esk42_1(esk24_0),X1)|~r1(X2,esk42_1(esk24_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_78]), c_0_39])])).
cnf(c_0_80,negated_conjecture,(p2(esk25_1(esk42_1(esk24_0)))|r1(esk39_2(esk24_0,esk42_1(esk24_0)),esk34_3(X1,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))))|~epred1_1(X1)|~r1(X1,esk42_1(esk24_0))), inference(spm,[status(thm)],[c_0_79, c_0_42])).
cnf(c_0_81,plain,(r1(X4,esk40_3(X1,X2,X4))|~epred1_1(X1)|~r1(X1,X2)|~p2(X3)|~r1(esk38_2(X1,X2),X3)|~r1(esk39_2(X1,X2),X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_82,plain,(p2(esk25_1(esk42_1(esk24_0)))|r1(esk39_2(esk24_0,esk42_1(esk24_0)),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_34]), c_0_21])])).
cnf(c_0_83,plain,(p2(esk40_3(X1,X2,X4))|~epred1_1(X1)|~r1(X1,X2)|~p2(X3)|~r1(esk38_2(X1,X2),X3)|~r1(esk39_2(X1,X2),X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_84,plain,(r1(X1,esk40_3(esk24_0,esk42_1(esk24_0),X1))|~r1(esk39_2(esk24_0,esk42_1(esk24_0)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_76]), c_0_21]), c_0_77]), c_0_34])])).
cnf(c_0_85,negated_conjecture,(r1(esk39_2(esk24_0,esk42_1(esk24_0)),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_82]), c_0_34])])).
cnf(c_0_86,plain,(p2(esk40_3(esk24_0,esk42_1(esk24_0),X1))|~r1(esk39_2(esk24_0,esk42_1(esk24_0)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_76]), c_0_21]), c_0_77]), c_0_34])])).
cnf(c_0_87,plain,(r1(esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0))),esk40_3(esk24_0,esk42_1(esk24_0),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0)))))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_88,plain,(p2(esk40_3(esk24_0,esk42_1(esk24_0),esk34_3(esk24_0,esk42_1(esk24_0),esk39_2(esk24_0,esk42_1(esk24_0)))))), inference(spm,[status(thm)],[c_0_86, c_0_85])).
cnf(c_0_89,plain,(p2(X1)|~r1(esk42_1(esk24_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_87]), c_0_21]), c_0_88]), c_0_39]), c_0_78]), c_0_34])])).
cnf(c_0_90,negated_conjecture,(p2(esk25_1(esk42_1(esk24_0)))), inference(spm,[status(thm)],[c_0_89, c_0_42])).
cnf(c_0_91,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_90]), c_0_34])]), ['proof']).
# SZS output end CNFRefutation.
EOF
