% Problem    : LCL652+1.020 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n138.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:45:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____042_B03_F1_AE_Q4_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (~(?[X1]:~(((((((~(![X2]:(~(r1(X1,X2))|~(p4(X2))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p1(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))|~(![X1]:(~(r1(X2,X1))|~(((((~(![X2]:(~(r1(X1,X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p3(X2)))|~(p2(X1))))))|~(![X2]:(~(r1(X1,X2))|~(((~(![X1]:(~(r1(X2,X1))|~((![X2]:(~(r1(X1,X2))|p3(X2))|~(p2(X1))))))&p2(X2))&~(![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p2(X1)))))))))))))|(~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p3(X1)))|~(p2(X2)))))))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p2(X1)))))))|~(![X2]:(~(r1(X1,X2))|~((![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(p2(X2)))))&![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p3(X2)))|~(p2(X1))))))))))))|(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p3(X1)))|~(p2(X2)))&![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|~(p2(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))))))|~(![X1]:(~(r1(X2,X1))|~((~(![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|p1(X1))))&![X2]:(~(r1(X1,X2))|p1(X2)))))))))))|~(![X2]:(~(r1(X1,X2))|~((~(![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|p1(X2))))&![X1]:(~(r1(X2,X1))|p1(X1)))))))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(((~(r1(X2,X1))|![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|p1(X2)))))|![X1]:(~(r1(X2,X1))|p1(X1))))|![X2]:(((~(r1(X1,X2))|![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p1(X1)))))|![X2]:(~(r1(X1,X2))|p1(X2))))|![X1]:(~(r1(X2,X1))|~(![X2]:(~(r1(X1,X2))|~(![X1]:(~(r1(X2,X1))|p2(X1)))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1))))))))|~(![X1]:((~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|~(p1(X1)))))|~(![X2]:(~(r1(X1,X2))|~(p1(X2))))))))|~(![X2]:((~(r1(X1,X2))|![X1]:(~(r1(X2,X1))|![X2]:(~(r1(X1,X2))|~(p1(X2)))))|~(![X1]:(~(r1(X2,X1))|~(p1(X1)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', main)).
fof(c_0_1, plain, (![X1]:(epred5_1(X1)<=>~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((((~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((![X2]:(~r1(X1,X2)|p3(X2))|~p2(X1)))))&p2(X2))&~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))))|(~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))|~(![X2]:(~r1(X1,X2)|~((![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))))))))|(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))), introduced(definition)).
fof(c_0_2, plain, (![X1]:(epred4_1(X1)<=>~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((((~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((![X2]:(~r1(X1,X2)|p3(X2))|~p2(X1)))))&p2(X2))&~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))))|(~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))|~(![X2]:(~r1(X1,X2)|~((![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))))))))|(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))), introduced(definition)).
fof(c_0_3, plain, (![X1]:(epred3_1(X1)<=>~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((((~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((![X2]:(~r1(X1,X2)|p3(X2))|~p2(X1)))))&p2(X2))&~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))))|(~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))|~(![X2]:(~r1(X1,X2)|~((![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))))))))|(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))), introduced(definition)).
fof(c_0_4, negated_conjecture, (~(~(?[X1]:~(((((((~(![X2]:(~r1(X1,X2)|~p4(X2)))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|epred5_1(X1))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(((~r1(X2,X1)|![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|p1(X2)))))|![X1]:(~r1(X2,X1)|p1(X1))))|![X2]:(((~r1(X1,X2)|![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p1(X1)))))|![X2]:(~r1(X1,X2)|p1(X2))))|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|p2(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~p1(X1))))|~(![X2]:(~r1(X1,X2)|~p1(X2)))))))|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~p1(X2))))|~(![X1]:(~r1(X2,X1)|~p1(X1)))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_1])).
fof(c_0_5, plain, (![X1]:(epred2_1(X1)<=>~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((((~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((![X2]:(~r1(X1,X2)|p3(X2))|~p2(X1)))))&p2(X2))&~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))))|(~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))|~(![X2]:(~r1(X1,X2)|~((![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))))))))|(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))), introduced(definition)).
fof(c_0_6, plain, (![X1]:(epred4_1(X1)=>~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|epred3_1(X1))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_2]), c_0_3])).
fof(c_0_7, plain, (![X1]:(epred5_1(X1)=>~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|epred4_1(X1))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_1]), c_0_2])).
fof(c_0_8, negated_conjecture, (![X4]:![X5]:![X6]:![X7]:![X10]:![X11]:![X12]:![X13]:![X14]:![X17]:![X18]:![X19]:![X20]:![X21]:![X24]:![X25]:![X26]:![X27]:![X28]:![X31]:![X32]:![X33]:![X34]:![X35]:![X37]:![X38]:![X39]:![X41]:![X42]:![X43]:![X45]:![X46]:![X47]:![X51]:![X56]:![X61]:![X66]:![X71]:![X76]:![X81]:![X86]:![X91]:![X96]:![X101]:![X106]:![X111]:![X116]:![X121]:![X126]:![X131]:![X136]:![X141]:![X146]:![X149]:![X151]:![X152]:![X153]:![X155]:![X156]:![X157]:![X159]:![X160]:![X161]:![X163]:![X164]:![X165]:![X167]:![X168]:![X169]:![X171]:![X172]:![X173]:![X175]:![X176]:![X177]:![X179]:![X180]:![X181]:![X183]:![X184]:![X185]:![X187]:![X188]:![X189]:![X191]:![X192]:![X193]:![X195]:![X196]:![X197]:![X199]:![X200]:![X201]:![X203]:![X204]:![X205]:![X207]:![X208]:![X209]:![X211]:![X212]:![X213]:![X215]:![X216]:![X217]:![X219]:![X220]:![X221]:![X223]:![X224]:![X225]:![X227]:![X228]:![X229]:(((((((~r1(esk1_0,X4)|~p4(X4))&((((r1(X5,esk3_1(X5))|((r1(X7,esk2_3(X5,X6,X7))|~r1(X6,X7))|~r1(X5,X6)))|~r1(esk1_0,X5))&(((~r1(esk3_1(X5),X10)|~p1(X10))|((r1(X7,esk2_3(X5,X6,X7))|~r1(X6,X7))|~r1(X5,X6)))|~r1(esk1_0,X5)))&(((r1(X5,esk3_1(X5))|((p1(esk2_3(X5,X6,X7))|~r1(X6,X7))|~r1(X5,X6)))|~r1(esk1_0,X5))&(((~r1(esk3_1(X5),X10)|~p1(X10))|((p1(esk2_3(X5,X6,X7))|~r1(X6,X7))|~r1(X5,X6)))|~r1(esk1_0,X5)))))&(((((((r1(X12,esk5_2(X11,X12))|((r1(X14,esk4_4(X11,X12,X13,X14))|~r1(X13,X14))|~r1(X12,X13)))|~r1(X11,X12))|~r1(esk1_0,X11))&((((~r1(esk5_2(X11,X12),X17)|~p1(X17))|((r1(X14,esk4_4(X11,X12,X13,X14))|~r1(X13,X14))|~r1(X12,X13)))|~r1(X11,X12))|~r1(esk1_0,X11)))&((((r1(X12,esk5_2(X11,X12))|((p1(esk4_4(X11,X12,X13,X14))|~r1(X13,X14))|~r1(X12,X13)))|~r1(X11,X12))|~r1(esk1_0,X11))&((((~r1(esk5_2(X11,X12),X17)|~p1(X17))|((p1(esk4_4(X11,X12,X13,X14))|~r1(X13,X14))|~r1(X12,X13)))|~r1(X11,X12))|~r1(esk1_0,X11))))&((((((((r1(X19,esk7_3(X11,X18,X19))|((r1(X21,esk6_5(X11,X18,X19,X20,X21))|~r1(X20,X21))|~r1(X19,X20)))|~r1(X18,X19))|~r1(X11,X18))|~r1(esk1_0,X11))&(((((~r1(esk7_3(X11,X18,X19),X24)|~p1(X24))|((r1(X21,esk6_5(X11,X18,X19,X20,X21))|~r1(X20,X21))|~r1(X19,X20)))|~r1(X18,X19))|~r1(X11,X18))|~r1(esk1_0,X11)))&(((((r1(X19,esk7_3(X11,X18,X19))|((p1(esk6_5(X11,X18,X19,X20,X21))|~r1(X20,X21))|~r1(X19,X20)))|~r1(X18,X19))|~r1(X11,X18))|~r1(esk1_0,X11))&(((((~r1(esk7_3(X11,X18,X19),X24)|~p1(X24))|((p1(esk6_5(X11,X18,X19,X20,X21))|~r1(X20,X21))|~r1(X19,X20)))|~r1(X18,X19))|~r1(X11,X18))|~r1(esk1_0,X11))))&(((((((((r1(X26,esk9_4(X11,X18,X25,X26))|((r1(X28,esk8_6(X11,X18,X25,X26,X27,X28))|~r1(X27,X28))|~r1(X26,X27)))|~r1(X25,X26))|~r1(X18,X25))|~r1(X11,X18))|~r1(esk1_0,X11))&((((((~r1(esk9_4(X11,X18,X25,X26),X31)|~p1(X31))|((r1(X28,esk8_6(X11,X18,X25,X26,X27,X28))|~r1(X27,X28))|~r1(X26,X27)))|~r1(X25,X26))|~r1(X18,X25))|~r1(X11,X18))|~r1(esk1_0,X11)))&((((((r1(X26,esk9_4(X11,X18,X25,X26))|((p1(esk8_6(X11,X18,X25,X26,X27,X28))|~r1(X27,X28))|~r1(X26,X27)))|~r1(X25,X26))|~r1(X18,X25))|~r1(X11,X18))|~r1(esk1_0,X11))&((((((~r1(esk9_4(X11,X18,X25,X26),X31)|~p1(X31))|((p1(esk8_6(X11,X18,X25,X26,X27,X28))|~r1(X27,X28))|~r1(X26,X27)))|~r1(X25,X26))|~r1(X18,X25))|~r1(X11,X18))|~r1(esk1_0,X11))))&((((~r1(X25,X32)|epred5_1(X32))|~r1(X18,X25))|~r1(X11,X18))|~r1(esk1_0,X11)))&((((((r1(X33,esk10_4(X11,X18,X25,X33))|(~r1(X33,X34)|(~r1(X34,X35)|p1(X35))))|~r1(X25,X33))|~r1(X18,X25))|~r1(X11,X18))|~r1(esk1_0,X11))&(((((~p1(esk10_4(X11,X18,X25,X33))|(~r1(X33,X34)|(~r1(X34,X35)|p1(X35))))|~r1(X25,X33))|~r1(X18,X25))|~r1(X11,X18))|~r1(esk1_0,X11)))))&(((((r1(X37,esk11_3(X11,X18,X37))|(~r1(X37,X38)|(~r1(X38,X39)|p1(X39))))|~r1(X18,X37))|~r1(X11,X18))|~r1(esk1_0,X11))&((((~p1(esk11_3(X11,X18,X37))|(~r1(X37,X38)|(~r1(X38,X39)|p1(X39))))|~r1(X18,X37))|~r1(X11,X18))|~r1(esk1_0,X11)))))&((((r1(X41,esk12_2(X11,X41))|(~r1(X41,X42)|(~r1(X42,X43)|p1(X43))))|~r1(X11,X41))|~r1(esk1_0,X11))&(((~p1(esk12_2(X11,X41))|(~r1(X41,X42)|(~r1(X42,X43)|p1(X43))))|~r1(X11,X41))|~r1(esk1_0,X11)))))&(((r1(X45,esk13_1(X45))|(~r1(X45,X46)|(~r1(X46,X47)|p1(X47))))|~r1(esk1_0,X45))&((~p1(esk13_1(X45))|(~r1(X45,X46)|(~r1(X46,X47)|p1(X47))))|~r1(esk1_0,X45))))&((r1(esk1_0,esk14_0)&(r1(esk14_0,esk15_0)&(~r1(esk15_0,X51)|p1(X51))))&(r1(esk14_0,esk16_0)&~p1(esk16_0))))&(((r1(esk1_0,esk17_0)&((r1(esk17_0,esk18_0)&(r1(esk18_0,esk19_0)&(~r1(esk19_0,X56)|p1(X56))))&(r1(esk18_0,esk20_0)&~p1(esk20_0))))&(((r1(esk17_0,esk21_0)&((r1(esk21_0,esk22_0)&(r1(esk22_0,esk23_0)&(~r1(esk23_0,X61)|p1(X61))))&(r1(esk22_0,esk24_0)&~p1(esk24_0))))&(((r1(esk21_0,esk25_0)&((r1(esk25_0,esk26_0)&(r1(esk26_0,esk27_0)&(~r1(esk27_0,X66)|p1(X66))))&(r1(esk26_0,esk28_0)&~p1(esk28_0))))&(((r1(esk25_0,esk29_0)&((r1(esk29_0,esk30_0)&(r1(esk30_0,esk31_0)&(~r1(esk31_0,X71)|p1(X71))))&(r1(esk30_0,esk32_0)&~p1(esk32_0))))&(((r1(esk29_0,esk33_0)&((r1(esk33_0,esk34_0)&(r1(esk34_0,esk35_0)&(~r1(esk35_0,X76)|p1(X76))))&(r1(esk34_0,esk36_0)&~p1(esk36_0))))&(((r1(esk33_0,esk37_0)&((r1(esk37_0,esk38_0)&(r1(esk38_0,esk39_0)&(~r1(esk39_0,X81)|p1(X81))))&(r1(esk38_0,esk40_0)&~p1(esk40_0))))&(((r1(esk37_0,esk41_0)&((r1(esk41_0,esk42_0)&(r1(esk42_0,esk43_0)&(~r1(esk43_0,X86)|p1(X86))))&(r1(esk42_0,esk44_0)&~p1(esk44_0))))&(((r1(esk41_0,esk45_0)&((r1(esk45_0,esk46_0)&(r1(esk46_0,esk47_0)&(~r1(esk47_0,X91)|p1(X91))))&(r1(esk46_0,esk48_0)&~p1(esk48_0))))&(((r1(esk45_0,esk49_0)&((r1(esk49_0,esk50_0)&(r1(esk50_0,esk51_0)&(~r1(esk51_0,X96)|p1(X96))))&(r1(esk50_0,esk52_0)&~p1(esk52_0))))&(((r1(esk49_0,esk53_0)&((r1(esk53_0,esk54_0)&(r1(esk54_0,esk55_0)&(~r1(esk55_0,X101)|p1(X101))))&(r1(esk54_0,esk56_0)&~p1(esk56_0))))&(((r1(esk53_0,esk57_0)&((r1(esk57_0,esk58_0)&(r1(esk58_0,esk59_0)&(~r1(esk59_0,X106)|p1(X106))))&(r1(esk58_0,esk60_0)&~p1(esk60_0))))&(((r1(esk57_0,esk61_0)&((r1(esk61_0,esk62_0)&(r1(esk62_0,esk63_0)&(~r1(esk63_0,X111)|p1(X111))))&(r1(esk62_0,esk64_0)&~p1(esk64_0))))&(((r1(esk61_0,esk65_0)&((r1(esk65_0,esk66_0)&(r1(esk66_0,esk67_0)&(~r1(esk67_0,X116)|p1(X116))))&(r1(esk66_0,esk68_0)&~p1(esk68_0))))&(((r1(esk65_0,esk69_0)&((r1(esk69_0,esk70_0)&(r1(esk70_0,esk71_0)&(~r1(esk71_0,X121)|p1(X121))))&(r1(esk70_0,esk72_0)&~p1(esk72_0))))&(((r1(esk69_0,esk73_0)&((r1(esk73_0,esk74_0)&(r1(esk74_0,esk75_0)&(~r1(esk75_0,X126)|p1(X126))))&(r1(esk74_0,esk76_0)&~p1(esk76_0))))&(((r1(esk73_0,esk77_0)&((r1(esk77_0,esk78_0)&(r1(esk78_0,esk79_0)&(~r1(esk79_0,X131)|p1(X131))))&(r1(esk78_0,esk80_0)&~p1(esk80_0))))&(((r1(esk77_0,esk81_0)&((r1(esk81_0,esk82_0)&(r1(esk82_0,esk83_0)&(~r1(esk83_0,X136)|p1(X136))))&(r1(esk82_0,esk84_0)&~p1(esk84_0))))&(((r1(esk81_0,esk85_0)&((r1(esk85_0,esk86_0)&(r1(esk86_0,esk87_0)&(~r1(esk87_0,X141)|p1(X141))))&(r1(esk86_0,esk88_0)&~p1(esk88_0))))&(((r1(esk85_0,esk89_0)&((r1(esk89_0,esk90_0)&(r1(esk90_0,esk91_0)&(~r1(esk91_0,X146)|p1(X146))))&(r1(esk90_0,esk92_0)&~p1(esk92_0))))&(r1(esk89_0,esk93_0)&((r1(X149,esk94_1(X149))|~r1(esk93_0,X149))&(~p2(esk94_1(X149))|~r1(esk93_0,X149)))))&((r1(X151,esk95_1(X151))|(~r1(esk89_0,X151)|(~r1(X151,X152)|(~r1(X152,X153)|~p1(X153)))))&(p1(esk95_1(X151))|(~r1(esk89_0,X151)|(~r1(X151,X152)|(~r1(X152,X153)|~p1(X153))))))))&((r1(X155,esk96_1(X155))|(~r1(esk85_0,X155)|(~r1(X155,X156)|(~r1(X156,X157)|~p1(X157)))))&(p1(esk96_1(X155))|(~r1(esk85_0,X155)|(~r1(X155,X156)|(~r1(X156,X157)|~p1(X157))))))))&((r1(X159,esk97_1(X159))|(~r1(esk81_0,X159)|(~r1(X159,X160)|(~r1(X160,X161)|~p1(X161)))))&(p1(esk97_1(X159))|(~r1(esk81_0,X159)|(~r1(X159,X160)|(~r1(X160,X161)|~p1(X161))))))))&((r1(X163,esk98_1(X163))|(~r1(esk77_0,X163)|(~r1(X163,X164)|(~r1(X164,X165)|~p1(X165)))))&(p1(esk98_1(X163))|(~r1(esk77_0,X163)|(~r1(X163,X164)|(~r1(X164,X165)|~p1(X165))))))))&((r1(X167,esk99_1(X167))|(~r1(esk73_0,X167)|(~r1(X167,X168)|(~r1(X168,X169)|~p1(X169)))))&(p1(esk99_1(X167))|(~r1(esk73_0,X167)|(~r1(X167,X168)|(~r1(X168,X169)|~p1(X169))))))))&((r1(X171,esk100_1(X171))|(~r1(esk69_0,X171)|(~r1(X171,X172)|(~r1(X172,X173)|~p1(X173)))))&(p1(esk100_1(X171))|(~r1(esk69_0,X171)|(~r1(X171,X172)|(~r1(X172,X173)|~p1(X173))))))))&((r1(X175,esk101_1(X175))|(~r1(esk65_0,X175)|(~r1(X175,X176)|(~r1(X176,X177)|~p1(X177)))))&(p1(esk101_1(X175))|(~r1(esk65_0,X175)|(~r1(X175,X176)|(~r1(X176,X177)|~p1(X177))))))))&((r1(X179,esk102_1(X179))|(~r1(esk61_0,X179)|(~r1(X179,X180)|(~r1(X180,X181)|~p1(X181)))))&(p1(esk102_1(X179))|(~r1(esk61_0,X179)|(~r1(X179,X180)|(~r1(X180,X181)|~p1(X181))))))))&((r1(X183,esk103_1(X183))|(~r1(esk57_0,X183)|(~r1(X183,X184)|(~r1(X184,X185)|~p1(X185)))))&(p1(esk103_1(X183))|(~r1(esk57_0,X183)|(~r1(X183,X184)|(~r1(X184,X185)|~p1(X185))))))))&((r1(X187,esk104_1(X187))|(~r1(esk53_0,X187)|(~r1(X187,X188)|(~r1(X188,X189)|~p1(X189)))))&(p1(esk104_1(X187))|(~r1(esk53_0,X187)|(~r1(X187,X188)|(~r1(X188,X189)|~p1(X189))))))))&((r1(X191,esk105_1(X191))|(~r1(esk49_0,X191)|(~r1(X191,X192)|(~r1(X192,X193)|~p1(X193)))))&(p1(esk105_1(X191))|(~r1(esk49_0,X191)|(~r1(X191,X192)|(~r1(X192,X193)|~p1(X193))))))))&((r1(X195,esk106_1(X195))|(~r1(esk45_0,X195)|(~r1(X195,X196)|(~r1(X196,X197)|~p1(X197)))))&(p1(esk106_1(X195))|(~r1(esk45_0,X195)|(~r1(X195,X196)|(~r1(X196,X197)|~p1(X197))))))))&((r1(X199,esk107_1(X199))|(~r1(esk41_0,X199)|(~r1(X199,X200)|(~r1(X200,X201)|~p1(X201)))))&(p1(esk107_1(X199))|(~r1(esk41_0,X199)|(~r1(X199,X200)|(~r1(X200,X201)|~p1(X201))))))))&((r1(X203,esk108_1(X203))|(~r1(esk37_0,X203)|(~r1(X203,X204)|(~r1(X204,X205)|~p1(X205)))))&(p1(esk108_1(X203))|(~r1(esk37_0,X203)|(~r1(X203,X204)|(~r1(X204,X205)|~p1(X205))))))))&((r1(X207,esk109_1(X207))|(~r1(esk33_0,X207)|(~r1(X207,X208)|(~r1(X208,X209)|~p1(X209)))))&(p1(esk109_1(X207))|(~r1(esk33_0,X207)|(~r1(X207,X208)|(~r1(X208,X209)|~p1(X209))))))))&((r1(X211,esk110_1(X211))|(~r1(esk29_0,X211)|(~r1(X211,X212)|(~r1(X212,X213)|~p1(X213)))))&(p1(esk110_1(X211))|(~r1(esk29_0,X211)|(~r1(X211,X212)|(~r1(X212,X213)|~p1(X213))))))))&((r1(X215,esk111_1(X215))|(~r1(esk25_0,X215)|(~r1(X215,X216)|(~r1(X216,X217)|~p1(X217)))))&(p1(esk111_1(X215))|(~r1(esk25_0,X215)|(~r1(X215,X216)|(~r1(X216,X217)|~p1(X217))))))))&((r1(X219,esk112_1(X219))|(~r1(esk21_0,X219)|(~r1(X219,X220)|(~r1(X220,X221)|~p1(X221)))))&(p1(esk112_1(X219))|(~r1(esk21_0,X219)|(~r1(X219,X220)|(~r1(X220,X221)|~p1(X221))))))))&((r1(X223,esk113_1(X223))|(~r1(esk17_0,X223)|(~r1(X223,X224)|(~r1(X224,X225)|~p1(X225)))))&(p1(esk113_1(X223))|(~r1(esk17_0,X223)|(~r1(X223,X224)|(~r1(X224,X225)|~p1(X225))))))))&((r1(X227,esk114_1(X227))|(~r1(esk1_0,X227)|(~r1(X227,X228)|(~r1(X228,X229)|~p1(X229)))))&(p1(esk114_1(X227))|(~r1(esk1_0,X227)|(~r1(X227,X228)|(~r1(X228,X229)|~p1(X229)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
fof(c_0_9, plain, (![X1]:(epred1_1(X1)<=>~(((((~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((![X2]:(~r1(X1,X2)|p3(X2))|~p2(X1)))))&p2(X2))&~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))))|(~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))|~(![X2]:(~r1(X1,X2)|~((![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))))))))|(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))), introduced(definition)).
fof(c_0_10, plain, (![X1]:(epred3_1(X1)=>~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|epred2_1(X1))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_3]), c_0_5])).
fof(c_0_11, plain, (![X3]:![X4]:![X5]:![X6]:![X9]:![X10]:![X11]:![X12]:![X13]:![X16]:![X17]:![X18]:![X19]:![X20]:![X23]:![X24]:![X25]:![X26]:![X27]:![X30]:![X31]:![X32]:![X33]:![X34]:![X36]:![X37]:![X38]:![X40]:![X41]:![X42]:![X44]:![X45]:![X46]:(((((((r1(X4,esk154_2(X3,X4))|((r1(X6,esk153_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred4_1(X3))&((((~r1(esk154_2(X3,X4),X9)|~p1(X9))|((r1(X6,esk153_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred4_1(X3)))&((((r1(X4,esk154_2(X3,X4))|((p1(esk153_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred4_1(X3))&((((~r1(esk154_2(X3,X4),X9)|~p1(X9))|((p1(esk153_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred4_1(X3))))&((((((((r1(X11,esk156_3(X3,X10,X11))|((r1(X13,esk155_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred4_1(X3))&(((((~r1(esk156_3(X3,X10,X11),X16)|~p1(X16))|((r1(X13,esk155_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred4_1(X3)))&(((((r1(X11,esk156_3(X3,X10,X11))|((p1(esk155_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred4_1(X3))&(((((~r1(esk156_3(X3,X10,X11),X16)|~p1(X16))|((p1(esk155_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred4_1(X3))))&(((((((((r1(X18,esk158_4(X3,X10,X17,X18))|((r1(X20,esk157_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3))&((((((~r1(esk158_4(X3,X10,X17,X18),X23)|~p1(X23))|((r1(X20,esk157_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3)))&((((((r1(X18,esk158_4(X3,X10,X17,X18))|((p1(esk157_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3))&((((((~r1(esk158_4(X3,X10,X17,X18),X23)|~p1(X23))|((p1(esk157_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3))))&((((((((((r1(X25,esk160_5(X3,X10,X17,X24,X25))|((r1(X27,esk159_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3))&(((((((~r1(esk160_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((r1(X27,esk159_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3)))&(((((((r1(X25,esk160_5(X3,X10,X17,X24,X25))|((p1(esk159_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3))&(((((((~r1(esk160_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((p1(esk159_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3))))&(((((~r1(X24,X31)|epred3_1(X31))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3)))&(((((((r1(X32,esk161_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3))&((((((~p1(esk161_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3)))))&((((((r1(X36,esk162_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3))&(((((~p1(esk162_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred4_1(X3)))))&(((((r1(X40,esk163_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred4_1(X3))&((((~p1(esk163_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred4_1(X3)))))&((((r1(X44,esk164_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred4_1(X3))&(((~p1(esk164_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred4_1(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
fof(c_0_12, plain, (![X3]:![X4]:![X5]:![X6]:![X9]:![X10]:![X11]:![X12]:![X13]:![X16]:![X17]:![X18]:![X19]:![X20]:![X23]:![X24]:![X25]:![X26]:![X27]:![X30]:![X31]:![X32]:![X33]:![X34]:![X36]:![X37]:![X38]:![X40]:![X41]:![X42]:![X44]:![X45]:![X46]:(((((((r1(X4,esk166_2(X3,X4))|((r1(X6,esk165_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred5_1(X3))&((((~r1(esk166_2(X3,X4),X9)|~p1(X9))|((r1(X6,esk165_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred5_1(X3)))&((((r1(X4,esk166_2(X3,X4))|((p1(esk165_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred5_1(X3))&((((~r1(esk166_2(X3,X4),X9)|~p1(X9))|((p1(esk165_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred5_1(X3))))&((((((((r1(X11,esk168_3(X3,X10,X11))|((r1(X13,esk167_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred5_1(X3))&(((((~r1(esk168_3(X3,X10,X11),X16)|~p1(X16))|((r1(X13,esk167_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred5_1(X3)))&(((((r1(X11,esk168_3(X3,X10,X11))|((p1(esk167_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred5_1(X3))&(((((~r1(esk168_3(X3,X10,X11),X16)|~p1(X16))|((p1(esk167_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred5_1(X3))))&(((((((((r1(X18,esk170_4(X3,X10,X17,X18))|((r1(X20,esk169_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3))&((((((~r1(esk170_4(X3,X10,X17,X18),X23)|~p1(X23))|((r1(X20,esk169_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3)))&((((((r1(X18,esk170_4(X3,X10,X17,X18))|((p1(esk169_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3))&((((((~r1(esk170_4(X3,X10,X17,X18),X23)|~p1(X23))|((p1(esk169_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3))))&((((((((((r1(X25,esk172_5(X3,X10,X17,X24,X25))|((r1(X27,esk171_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3))&(((((((~r1(esk172_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((r1(X27,esk171_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3)))&(((((((r1(X25,esk172_5(X3,X10,X17,X24,X25))|((p1(esk171_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3))&(((((((~r1(esk172_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((p1(esk171_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3))))&(((((~r1(X24,X31)|epred4_1(X31))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3)))&(((((((r1(X32,esk173_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3))&((((((~p1(esk173_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3)))))&((((((r1(X36,esk174_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3))&(((((~p1(esk174_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred5_1(X3)))))&(((((r1(X40,esk175_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred5_1(X3))&((((~p1(esk175_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred5_1(X3)))))&((((r1(X44,esk176_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred5_1(X3))&(((~p1(esk176_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred5_1(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
cnf(c_0_13,negated_conjecture,(epred5_1(X4)|~r1(esk1_0,X1)|~r1(X1,X2)|~r1(X2,X3)|~r1(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,negated_conjecture,(r1(esk1_0,esk17_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_15, plain, (![X1]:(epred2_1(X1)=>~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|~(((~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p1(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p1(X1))))))))|~(![X1]:(~r1(X2,X1)|epred1_1(X1))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))))|~(![X1]:(~r1(X2,X1)|~((~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|p1(X1))))&![X2]:(~r1(X1,X2)|p1(X2)))))))))))|~(![X2]:(~r1(X1,X2)|~((~(![X1]:(~r1(X2,X1)|![X2]:(~r1(X1,X2)|p1(X2))))&![X1]:(~r1(X2,X1)|p1(X1)))))))))), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_5]), c_0_9])).
fof(c_0_16, plain, (![X3]:![X4]:![X5]:![X6]:![X9]:![X10]:![X11]:![X12]:![X13]:![X16]:![X17]:![X18]:![X19]:![X20]:![X23]:![X24]:![X25]:![X26]:![X27]:![X30]:![X31]:![X32]:![X33]:![X34]:![X36]:![X37]:![X38]:![X40]:![X41]:![X42]:![X44]:![X45]:![X46]:(((((((r1(X4,esk142_2(X3,X4))|((r1(X6,esk141_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred3_1(X3))&((((~r1(esk142_2(X3,X4),X9)|~p1(X9))|((r1(X6,esk141_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred3_1(X3)))&((((r1(X4,esk142_2(X3,X4))|((p1(esk141_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred3_1(X3))&((((~r1(esk142_2(X3,X4),X9)|~p1(X9))|((p1(esk141_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred3_1(X3))))&((((((((r1(X11,esk144_3(X3,X10,X11))|((r1(X13,esk143_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred3_1(X3))&(((((~r1(esk144_3(X3,X10,X11),X16)|~p1(X16))|((r1(X13,esk143_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred3_1(X3)))&(((((r1(X11,esk144_3(X3,X10,X11))|((p1(esk143_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred3_1(X3))&(((((~r1(esk144_3(X3,X10,X11),X16)|~p1(X16))|((p1(esk143_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred3_1(X3))))&(((((((((r1(X18,esk146_4(X3,X10,X17,X18))|((r1(X20,esk145_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3))&((((((~r1(esk146_4(X3,X10,X17,X18),X23)|~p1(X23))|((r1(X20,esk145_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3)))&((((((r1(X18,esk146_4(X3,X10,X17,X18))|((p1(esk145_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3))&((((((~r1(esk146_4(X3,X10,X17,X18),X23)|~p1(X23))|((p1(esk145_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3))))&((((((((((r1(X25,esk148_5(X3,X10,X17,X24,X25))|((r1(X27,esk147_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3))&(((((((~r1(esk148_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((r1(X27,esk147_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3)))&(((((((r1(X25,esk148_5(X3,X10,X17,X24,X25))|((p1(esk147_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3))&(((((((~r1(esk148_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((p1(esk147_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3))))&(((((~r1(X24,X31)|epred2_1(X31))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3)))&(((((((r1(X32,esk149_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3))&((((((~p1(esk149_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3)))))&((((((r1(X36,esk150_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3))&(((((~p1(esk150_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred3_1(X3)))))&(((((r1(X40,esk151_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred3_1(X3))&((((~p1(esk151_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred3_1(X3)))))&((((r1(X44,esk152_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred3_1(X3))&(((~p1(esk152_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred3_1(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
cnf(c_0_17,plain,(epred3_1(X5)|~epred4_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,negated_conjecture,(r1(esk57_0,esk61_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19,plain,(epred4_1(X5)|~epred5_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,negated_conjecture,(r1(esk41_0,esk45_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21,negated_conjecture,(epred5_1(X1)|~r1(esk17_0,X2)|~r1(X3,X1)|~r1(X2,X3)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_22,negated_conjecture,(r1(esk17_0,esk21_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_23, plain, (![X3]:![X4]:![X5]:![X6]:![X9]:![X10]:![X11]:![X12]:![X13]:![X16]:![X17]:![X18]:![X19]:![X20]:![X23]:![X24]:![X25]:![X26]:![X27]:![X30]:![X31]:![X32]:![X33]:![X34]:![X36]:![X37]:![X38]:![X40]:![X41]:![X42]:![X44]:![X45]:![X46]:(((((((r1(X4,esk130_2(X3,X4))|((r1(X6,esk129_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred2_1(X3))&((((~r1(esk130_2(X3,X4),X9)|~p1(X9))|((r1(X6,esk129_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred2_1(X3)))&((((r1(X4,esk130_2(X3,X4))|((p1(esk129_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred2_1(X3))&((((~r1(esk130_2(X3,X4),X9)|~p1(X9))|((p1(esk129_4(X3,X4,X5,X6))|~r1(X5,X6))|~r1(X4,X5)))|~r1(X3,X4))|~epred2_1(X3))))&((((((((r1(X11,esk132_3(X3,X10,X11))|((r1(X13,esk131_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred2_1(X3))&(((((~r1(esk132_3(X3,X10,X11),X16)|~p1(X16))|((r1(X13,esk131_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred2_1(X3)))&(((((r1(X11,esk132_3(X3,X10,X11))|((p1(esk131_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred2_1(X3))&(((((~r1(esk132_3(X3,X10,X11),X16)|~p1(X16))|((p1(esk131_5(X3,X10,X11,X12,X13))|~r1(X12,X13))|~r1(X11,X12)))|~r1(X10,X11))|~r1(X3,X10))|~epred2_1(X3))))&(((((((((r1(X18,esk134_4(X3,X10,X17,X18))|((r1(X20,esk133_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&((((((~r1(esk134_4(X3,X10,X17,X18),X23)|~p1(X23))|((r1(X20,esk133_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))&((((((r1(X18,esk134_4(X3,X10,X17,X18))|((p1(esk133_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&((((((~r1(esk134_4(X3,X10,X17,X18),X23)|~p1(X23))|((p1(esk133_6(X3,X10,X17,X18,X19,X20))|~r1(X19,X20))|~r1(X18,X19)))|~r1(X17,X18))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))))&((((((((((r1(X25,esk136_5(X3,X10,X17,X24,X25))|((r1(X27,esk135_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&(((((((~r1(esk136_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((r1(X27,esk135_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))&(((((((r1(X25,esk136_5(X3,X10,X17,X24,X25))|((p1(esk135_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&(((((((~r1(esk136_5(X3,X10,X17,X24,X25),X30)|~p1(X30))|((p1(esk135_7(X3,X10,X17,X24,X25,X26,X27))|~r1(X26,X27))|~r1(X25,X26)))|~r1(X24,X25))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))))&(((((~r1(X24,X31)|epred1_1(X31))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))&(((((((r1(X32,esk137_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&((((((~p1(esk137_5(X3,X10,X17,X24,X32))|(~r1(X32,X33)|(~r1(X33,X34)|p1(X34))))|~r1(X24,X32))|~r1(X17,X24))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))))&((((((r1(X36,esk138_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3))&(((((~p1(esk138_4(X3,X10,X17,X36))|(~r1(X36,X37)|(~r1(X37,X38)|p1(X38))))|~r1(X17,X36))|~r1(X10,X17))|~r1(X3,X10))|~epred2_1(X3)))))&(((((r1(X40,esk139_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred2_1(X3))&((((~p1(esk139_3(X3,X10,X40))|(~r1(X40,X41)|(~r1(X41,X42)|p1(X42))))|~r1(X10,X40))|~r1(X3,X10))|~epred2_1(X3)))))&((((r1(X44,esk140_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred2_1(X3))&(((~p1(esk140_2(X3,X44))|(~r1(X44,X45)|(~r1(X45,X46)|p1(X46))))|~r1(X3,X44))|~epred2_1(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).
cnf(c_0_24,plain,(epred2_1(X5)|~epred3_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,negated_conjecture,(r1(esk73_0,esk77_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_26,negated_conjecture,(epred3_1(esk61_0)|~epred4_1(X1)|~r1(X2,esk57_0)|~r1(X3,X2)|~r1(X1,X3)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_27,negated_conjecture,(r1(esk53_0,esk57_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_28,negated_conjecture,(epred4_1(esk45_0)|~epred5_1(X1)|~r1(X2,esk41_0)|~r1(X3,X2)|~r1(X1,X3)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_29,negated_conjecture,(r1(esk37_0,esk41_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_30,negated_conjecture,(epred5_1(X1)|~r1(esk21_0,X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_31,negated_conjecture,(r1(esk21_0,esk25_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_32,plain,(epred1_1(X5)|~epred2_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~r1(X3,X4)|~r1(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33,negated_conjecture,(r1(esk89_0,esk93_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_34,negated_conjecture,(epred2_1(esk77_0)|~epred3_1(X1)|~r1(X2,esk73_0)|~r1(X3,X2)|~r1(X1,X3)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_35,negated_conjecture,(r1(esk69_0,esk73_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_36,negated_conjecture,(epred3_1(esk61_0)|~epred4_1(X1)|~r1(X2,esk53_0)|~r1(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_37,negated_conjecture,(r1(esk49_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_38,negated_conjecture,(epred4_1(esk45_0)|~epred5_1(X1)|~r1(X2,esk37_0)|~r1(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_39,negated_conjecture,(r1(esk33_0,esk37_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_40,negated_conjecture,(epred5_1(X1)|~r1(esk25_0,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_41,negated_conjecture,(r1(esk25_0,esk29_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_42,negated_conjecture,(epred1_1(esk93_0)|~epred2_1(X1)|~r1(X2,esk89_0)|~r1(X3,X2)|~r1(X1,X3)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_43,negated_conjecture,(r1(esk85_0,esk89_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_44,negated_conjecture,(epred2_1(esk77_0)|~epred3_1(X1)|~r1(X2,esk69_0)|~r1(X1,X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_45,negated_conjecture,(r1(esk65_0,esk69_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_46,negated_conjecture,(epred3_1(esk61_0)|~epred4_1(X1)|~r1(X1,esk49_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_47,negated_conjecture,(r1(esk45_0,esk49_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_48,negated_conjecture,(epred4_1(esk45_0)|~epred5_1(X1)|~r1(X1,esk33_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_49,negated_conjecture,(r1(esk29_0,esk33_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_50,negated_conjecture,(epred5_1(esk29_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51,negated_conjecture,(epred1_1(esk93_0)|~epred2_1(X1)|~r1(X2,esk85_0)|~r1(X1,X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52,negated_conjecture,(r1(esk81_0,esk85_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_53,negated_conjecture,(epred2_1(esk77_0)|~epred3_1(X1)|~r1(X1,esk65_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54,negated_conjecture,(r1(esk61_0,esk65_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_55,negated_conjecture,(epred3_1(esk61_0)|~epred4_1(esk45_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_56,negated_conjecture,(epred4_1(esk45_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
fof(c_0_57, plain, (![X1]:(epred1_1(X1)=>~(((((~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))|~(![X2]:(~r1(X1,X2)|~(((~(![X1]:(~r1(X2,X1)|~((![X2]:(~r1(X1,X2)|p3(X2))|~p2(X1)))))&p2(X2))&~(![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))))))|(~(![X2]:(~r1(X1,X2)|![X1]:(~r1(X2,X1)|~(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))))))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))|~(![X2]:(~r1(X1,X2)|~((![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~p2(X2))))&![X1]:(~r1(X2,X1)|~(![X2]:(~r1(X1,X2)|~(![X1]:((~r1(X2,X1)|![X2]:(~r1(X1,X2)|p3(X2)))|~p2(X1)))))))))))|(![X2]:((~r1(X1,X2)|![X1]:(~r1(X2,X1)|p3(X1)))|~p2(X2))&![X2]:(~r1(X1,X2)|~(![X1]:(~r1(X2,X1)|~p2(X1))))))))), inference(split_equiv,[status(thm)],[c_0_9])).
cnf(c_0_58,negated_conjecture,(epred1_1(esk93_0)|~epred2_1(X1)|~r1(X1,esk81_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59,negated_conjecture,(r1(esk77_0,esk81_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_60,negated_conjecture,(epred2_1(esk77_0)|~epred3_1(esk61_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61,negated_conjecture,(epred3_1(esk61_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56])])).
fof(c_0_62, plain, (![X3]:![X4]:![X7]:![X8]:![X10]:![X12]:![X13]:![X14]:![X18]:![X19]:![X21]:![X23]:![X29]:((((((((r1(X4,esk115_2(X3,X4))|~r1(X3,X4))|~epred1_1(X3))&(((r1(esk115_2(X3,X4),esk116_2(X3,X4))|~r1(X3,X4))|~epred1_1(X3))&((~p3(esk116_2(X3,X4))|~r1(X3,X4))|~epred1_1(X3))))&((p2(esk115_2(X3,X4))|~r1(X3,X4))|~epred1_1(X3)))&(((((((r1(X10,esk118_3(X3,X7,X10))|~r1(X7,X10))|((r1(X8,esk117_3(X3,X7,X8))|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3))&(((((~r1(esk118_3(X3,X7,X10),X12)|~p2(X12))|~r1(X7,X10))|((r1(X8,esk117_3(X3,X7,X8))|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3)))&(((((r1(X10,esk118_3(X3,X7,X10))|~r1(X7,X10))|((~p3(esk117_3(X3,X7,X8))|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3))&(((((~r1(esk118_3(X3,X7,X10),X12)|~p2(X12))|~r1(X7,X10))|((~p3(esk117_3(X3,X7,X8))|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3))))&(((((r1(X10,esk118_3(X3,X7,X10))|~r1(X7,X10))|((p2(X8)|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3))&(((((~r1(esk118_3(X3,X7,X10),X12)|~p2(X12))|~r1(X7,X10))|((p2(X8)|~r1(X7,X8))|~p2(X7)))|~r1(X3,X7))|~epred1_1(X3)))))&(((((r1(X3,esk121_1(X3))|((r1(X14,esk119_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3))&(((~r1(esk121_1(X3),X18)|~p2(X18))|((r1(X14,esk119_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3)))&((((r1(X3,esk121_1(X3))|((r1(esk119_3(X3,X13,X14),esk120_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3))&(((~r1(esk121_1(X3),X18)|~p2(X18))|((r1(esk119_3(X3,X13,X14),esk120_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3)))&(((r1(X3,esk121_1(X3))|((~p3(esk120_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3))&(((~r1(esk121_1(X3),X18)|~p2(X18))|((~p3(esk120_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3)))))&(((r1(X3,esk121_1(X3))|((p2(esk119_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3))&(((~r1(esk121_1(X3),X18)|~p2(X18))|((p2(esk119_3(X3,X13,X14))|~r1(X13,X14))|~r1(X3,X13)))|~epred1_1(X3)))))&(((((r1(X19,esk123_2(X3,X19))|r1(X19,esk122_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))&((((((r1(X23,esk124_3(X3,X19,X23))|~r1(esk123_2(X3,X19),X23))|r1(X19,esk122_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))&(((((r1(esk124_3(X3,X19,X23),esk125_3(X3,X19,X23))|~r1(esk123_2(X3,X19),X23))|r1(X19,esk122_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))&((((~p3(esk125_3(X3,X19,X23))|~r1(esk123_2(X3,X19),X23))|r1(X19,esk122_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))))&((((p2(esk124_3(X3,X19,X23))|~r1(esk123_2(X3,X19),X23))|r1(X19,esk122_2(X3,X19)))|~r1(X3,X19))|~epred1_1(X3))))&((((r1(X19,esk123_2(X3,X19))|(~r1(esk122_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))&((((((r1(X23,esk124_3(X3,X19,X23))|~r1(esk123_2(X3,X19),X23))|(~r1(esk122_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))&(((((r1(esk124_3(X3,X19,X23),esk125_3(X3,X19,X23))|~r1(esk123_2(X3,X19),X23))|(~r1(esk122_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))&((((~p3(esk125_3(X3,X19,X23))|~r1(esk123_2(X3,X19),X23))|(~r1(esk122_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))))&((((p2(esk124_3(X3,X19,X23))|~r1(esk123_2(X3,X19),X23))|(~r1(esk122_2(X3,X19),X21)|~p2(X21)))|~r1(X3,X19))|~epred1_1(X3))))))&(((((r1(X3,esk128_1(X3))|r1(X3,esk126_1(X3)))|~epred1_1(X3))&(((~r1(esk128_1(X3),X29)|~p2(X29))|r1(X3,esk126_1(X3)))|~epred1_1(X3)))&((((r1(X3,esk128_1(X3))|r1(esk126_1(X3),esk127_1(X3)))|~epred1_1(X3))&(((~r1(esk128_1(X3),X29)|~p2(X29))|r1(esk126_1(X3),esk127_1(X3)))|~epred1_1(X3)))&(((r1(X3,esk128_1(X3))|~p3(esk127_1(X3)))|~epred1_1(X3))&(((~r1(esk128_1(X3),X29)|~p2(X29))|~p3(esk127_1(X3)))|~epred1_1(X3)))))&(((r1(X3,esk128_1(X3))|p2(esk126_1(X3)))|~epred1_1(X3))&(((~r1(esk128_1(X3),X29)|~p2(X29))|p2(esk126_1(X3)))|~epred1_1(X3)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_57])])])])])).
cnf(c_0_63,negated_conjecture,(epred1_1(esk93_0)|~epred2_1(esk77_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_64,negated_conjecture,(epred2_1(esk77_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61])])).
cnf(c_0_65,plain,(r1(X1,esk126_1(X1))|r1(X1,esk128_1(X1))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_66,negated_conjecture,(epred1_1(esk93_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_67,plain,(r1(X2,esk115_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_68,plain,(r1(esk93_0,esk126_1(esk93_0))|r1(esk93_0,esk128_1(esk93_0))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_69,plain,(p2(esk126_1(X1))|r1(X1,esk128_1(X1))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_70,plain,(r1(X1,esk126_1(X1))|~epred1_1(X1)|~p2(X2)|~r1(esk128_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_71,plain,(r1(esk128_1(esk93_0),esk115_2(esk93_0,esk128_1(esk93_0)))|r1(esk93_0,esk126_1(esk93_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_66])])).
cnf(c_0_72,plain,(p2(esk115_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_73,plain,(p2(esk126_1(esk93_0))|r1(esk93_0,esk128_1(esk93_0))), inference(spm,[status(thm)],[c_0_69, c_0_66])).
cnf(c_0_74,plain,(r1(esk93_0,esk126_1(esk93_0))|~p2(esk115_2(esk93_0,esk128_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_66])])).
cnf(c_0_75,plain,(p2(esk115_2(esk93_0,esk128_1(esk93_0)))|r1(esk93_0,esk126_1(esk93_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_68]), c_0_66])])).
cnf(c_0_76,plain,(p2(esk126_1(X1))|~epred1_1(X1)|~p2(X2)|~r1(esk128_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_77,plain,(p2(esk126_1(esk93_0))|r1(esk128_1(esk93_0),esk115_2(esk93_0,esk128_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_73]), c_0_66])])).
cnf(c_0_78,plain,(r1(X2,esk122_2(X1,X2))|r1(X2,esk123_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_79,plain,(r1(esk93_0,esk126_1(esk93_0))), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_80,plain,(p2(esk126_1(esk93_0))|~p2(esk115_2(esk93_0,esk128_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_66])])).
cnf(c_0_81,plain,(p2(esk115_2(esk93_0,esk128_1(esk93_0)))|p2(esk126_1(esk93_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_66])])).
cnf(c_0_82,plain,(p2(X3)|r1(X4,esk118_3(X1,X2,X4))|~epred1_1(X1)|~r1(X1,X2)|~p2(X2)|~r1(X2,X3)|~r1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_83,plain,(r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))|r1(esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_66])])).
cnf(c_0_84,plain,(p2(esk126_1(esk93_0))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85,negated_conjecture,(r1(X1,esk94_1(X1))|~r1(esk93_0,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_86,plain,(p2(X1)|r1(esk123_2(esk93_0,esk126_1(esk93_0)),esk118_3(X2,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))))|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))|~epred1_1(X2)|~r1(esk126_1(esk93_0),X1)|~r1(X2,esk126_1(esk93_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84])])).
cnf(c_0_87,negated_conjecture,(r1(esk126_1(esk93_0),esk94_1(esk126_1(esk93_0)))), inference(spm,[status(thm)],[c_0_85, c_0_79])).
cnf(c_0_88,negated_conjecture,(p2(esk94_1(esk126_1(esk93_0)))|r1(esk123_2(esk93_0,esk126_1(esk93_0)),esk118_3(X1,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))))|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))|~epred1_1(X1)|~r1(X1,esk126_1(esk93_0))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_89,negated_conjecture,(~r1(esk93_0,X1)|~p2(esk94_1(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_90,plain,(p2(esk94_1(esk126_1(esk93_0)))|r1(esk123_2(esk93_0,esk126_1(esk93_0)),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))))|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_79]), c_0_66])])).
cnf(c_0_91,plain,(r1(X2,esk122_2(X1,X2))|r1(X3,esk124_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(esk123_2(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_92,negated_conjecture,(r1(esk123_2(esk93_0,esk126_1(esk93_0)),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))))|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_79])])).
cnf(c_0_93,plain,(p2(X3)|~epred1_1(X1)|~r1(X1,X2)|~p2(X2)|~r1(X2,X3)|~r1(X2,X4)|~p2(X5)|~r1(esk118_3(X1,X2,X4),X5)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_94,plain,(r1(esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))),esk124_3(esk93_0,esk126_1(esk93_0),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0)))))|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_66]), c_0_79])])).
cnf(c_0_95,plain,(p2(X1)|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))|~p2(esk124_3(esk93_0,esk126_1(esk93_0),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0)))))|~r1(esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0)))|~r1(esk126_1(esk93_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_66]), c_0_84]), c_0_79])])).
cnf(c_0_96,plain,(r1(X2,esk122_2(X1,X2))|p2(esk124_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(esk123_2(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_97,plain,(p2(X1)|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))|~p2(esk124_3(esk93_0,esk126_1(esk93_0),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0)))))|~r1(esk126_1(esk93_0),X1)), inference(spm,[status(thm)],[c_0_95, c_0_83])).
cnf(c_0_98,plain,(p2(esk124_3(esk93_0,esk126_1(esk93_0),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0)))))|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_92]), c_0_66]), c_0_79])])).
cnf(c_0_99,plain,(p2(X1)|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))|~r1(esk126_1(esk93_0),X1)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_100,negated_conjecture,(p2(esk94_1(esk126_1(esk93_0)))|r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))), inference(spm,[status(thm)],[c_0_99, c_0_87])).
cnf(c_0_101,plain,(r1(X3,esk119_3(X1,X2,X3))|r1(X1,esk121_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_102,negated_conjecture,(r1(esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_100]), c_0_79])])).
cnf(c_0_103,plain,(p2(esk119_3(X1,X2,X3))|r1(X1,esk121_1(X1))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_104,plain,(r1(esk122_2(esk93_0,esk126_1(esk93_0)),esk119_3(X1,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|r1(X1,esk121_1(X1))|~epred1_1(X1)|~r1(X1,esk126_1(esk93_0))), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_105,plain,(p2(esk119_3(X1,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|r1(X1,esk121_1(X1))|~epred1_1(X1)|~r1(X1,esk126_1(esk93_0))), inference(spm,[status(thm)],[c_0_103, c_0_102])).
cnf(c_0_106,plain,(r1(esk122_2(esk93_0,esk126_1(esk93_0)),esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|r1(esk93_0,esk121_1(esk93_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_79]), c_0_66])])).
cnf(c_0_107,plain,(p2(esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|r1(esk93_0,esk121_1(esk93_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_79]), c_0_66])])).
cnf(c_0_108,plain,(r1(X3,esk119_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p2(X4)|~r1(esk121_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_109,plain,(r1(esk122_2(esk93_0,esk126_1(esk93_0)),esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|r1(esk121_1(esk93_0),esk115_2(esk93_0,esk121_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_106]), c_0_66])])).
cnf(c_0_110,plain,(p2(esk119_3(X1,X2,X3))|~epred1_1(X1)|~r1(X1,X2)|~r1(X2,X3)|~p2(X4)|~r1(esk121_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_111,plain,(p2(esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|r1(esk121_1(esk93_0),esk115_2(esk93_0,esk121_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_107]), c_0_66])])).
cnf(c_0_112,plain,(r1(esk122_2(esk93_0,esk126_1(esk93_0)),esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|r1(X1,esk119_3(esk93_0,X2,X1))|~p2(esk115_2(esk93_0,esk121_1(esk93_0)))|~r1(esk93_0,X2)|~r1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_66])])).
cnf(c_0_113,plain,(p2(esk115_2(esk93_0,esk121_1(esk93_0)))|r1(esk122_2(esk93_0,esk126_1(esk93_0)),esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_106]), c_0_66])])).
cnf(c_0_114,plain,(p2(esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|p2(esk119_3(esk93_0,X1,X2))|~p2(esk115_2(esk93_0,esk121_1(esk93_0)))|~r1(esk93_0,X1)|~r1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_66])])).
cnf(c_0_115,plain,(p2(esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|p2(esk115_2(esk93_0,esk121_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_107]), c_0_66])])).
cnf(c_0_116,plain,(r1(esk122_2(esk93_0,esk126_1(esk93_0)),esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|r1(X1,esk119_3(esk93_0,X2,X1))|~r1(esk93_0,X2)|~r1(X2,X1)), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_117,plain,(p2(esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|p2(esk119_3(esk93_0,X1,X2))|~r1(esk93_0,X1)|~r1(X1,X2)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_118,plain,(r1(esk122_2(esk93_0,esk126_1(esk93_0)),esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|r1(X1,esk119_3(esk93_0,esk126_1(esk93_0),X1))|~r1(esk126_1(esk93_0),X1)), inference(spm,[status(thm)],[c_0_116, c_0_79])).
cnf(c_0_119,plain,(p2(esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))|p2(esk119_3(esk93_0,esk126_1(esk93_0),X1))|~r1(esk126_1(esk93_0),X1)), inference(spm,[status(thm)],[c_0_117, c_0_79])).
cnf(c_0_120,plain,(r1(X2,esk123_2(X1,X2))|~epred1_1(X1)|~r1(X1,X2)|~p2(X3)|~r1(esk122_2(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_121,negated_conjecture,(r1(esk122_2(esk93_0,esk126_1(esk93_0)),esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))), inference(spm,[status(thm)],[c_0_118, c_0_102])).
cnf(c_0_122,negated_conjecture,(p2(esk119_3(esk93_0,esk126_1(esk93_0),esk122_2(esk93_0,esk126_1(esk93_0))))), inference(spm,[status(thm)],[c_0_119, c_0_102])).
cnf(c_0_123,plain,(r1(esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_66]), c_0_122]), c_0_79])])).
cnf(c_0_124,plain,(p2(X1)|r1(esk123_2(esk93_0,esk126_1(esk93_0)),esk118_3(X2,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))))|~epred1_1(X2)|~r1(esk126_1(esk93_0),X1)|~r1(X2,esk126_1(esk93_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_123]), c_0_84])])).
cnf(c_0_125,negated_conjecture,(p2(esk94_1(esk126_1(esk93_0)))|r1(esk123_2(esk93_0,esk126_1(esk93_0)),esk118_3(X1,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))))|~epred1_1(X1)|~r1(X1,esk126_1(esk93_0))), inference(spm,[status(thm)],[c_0_124, c_0_87])).
cnf(c_0_126,plain,(r1(X4,esk124_3(X1,X2,X4))|~epred1_1(X1)|~r1(X1,X2)|~p2(X3)|~r1(esk122_2(X1,X2),X3)|~r1(esk123_2(X1,X2),X4)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_127,plain,(p2(esk94_1(esk126_1(esk93_0)))|r1(esk123_2(esk93_0,esk126_1(esk93_0)),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_79]), c_0_66])])).
cnf(c_0_128,plain,(p2(esk124_3(X1,X2,X4))|~epred1_1(X1)|~r1(X1,X2)|~p2(X3)|~r1(esk122_2(X1,X2),X3)|~r1(esk123_2(X1,X2),X4)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_129,plain,(r1(X1,esk124_3(esk93_0,esk126_1(esk93_0),X1))|~r1(esk123_2(esk93_0,esk126_1(esk93_0)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_121]), c_0_66]), c_0_122]), c_0_79])])).
cnf(c_0_130,negated_conjecture,(r1(esk123_2(esk93_0,esk126_1(esk93_0)),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_127]), c_0_79])])).
cnf(c_0_131,plain,(p2(esk124_3(esk93_0,esk126_1(esk93_0),X1))|~r1(esk123_2(esk93_0,esk126_1(esk93_0)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_121]), c_0_66]), c_0_122]), c_0_79])])).
cnf(c_0_132,plain,(r1(esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0))),esk124_3(esk93_0,esk126_1(esk93_0),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0)))))), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_133,plain,(p2(esk124_3(esk93_0,esk126_1(esk93_0),esk118_3(esk93_0,esk126_1(esk93_0),esk123_2(esk93_0,esk126_1(esk93_0)))))), inference(spm,[status(thm)],[c_0_131, c_0_130])).
cnf(c_0_134,plain,(p2(X1)|~r1(esk126_1(esk93_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_132]), c_0_66]), c_0_133]), c_0_84]), c_0_123]), c_0_79])])).
cnf(c_0_135,negated_conjecture,(p2(esk94_1(esk126_1(esk93_0)))), inference(spm,[status(thm)],[c_0_134, c_0_87])).
cnf(c_0_136,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_135]), c_0_79])]), ['proof']).
# SZS output end CNFRefutation.
EOF
