% Problem    : GEO324+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n128.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:41:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X77]:![X78]:![X79]:![X80]:![X81]:![X82]:![X83]:![X84]:((((((((((((ron(X79,X83)&ron(X77,X83))&X77!=X81)&X77!=X79)&X83=X84)&rline(X84))&X81=X82)&rpoint(X82))&X79=X80)&rpoint(X80))&X77=X78)&rpoint(X78))=>((~(rR(X77,X79,X81))&ron(X81,X83))<=>vangle(X79,X77,X81)=v0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(188), 0), imp(cond(axiom(188), 0)))')).
fof(c_0_1, axiom, (![X219]:![X220]:((X219=X220&rreal(X220))=>vplus(X219,v0)=X219)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(141), 0), imp(cond(axiom(141), 0)))')).
fof(c_0_2, axiom, (![X215]:![X216]:![X217]:![X218]:((((X217=X218&rreal(X218))&X215=X216)&rreal(X216))=>vplus(X215,X217)=vplus(X217,X215))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(143), 0), imp(cond(axiom(143), 0)))')).
fof(c_0_3, axiom, (((((((vf(vd1287,vd1288)=vf(vd1287,vd1289)&vd1288!=vd1289)&vd1287!=vd1289)&vd1287!=vd1288)&?[X1]:(vd1289=X1&rpoint(X1)))&?[X2]:(vd1288=X2&rpoint(X2)))&?[X3]:(vd1287=X3&rpoint(X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(holds(conjunct2(345), 1294, 0), and(pred(conjunct1(345), 9), and(pred(conjunct1(345), 8), and(pred(conjunct1(345), 7), and(qe(s3(plural(345))), and(qe(s2(plural(345))), qe(s1(plural(345)))))))))')).
fof(c_0_4, axiom, (![X205]:![X206]:![X207]:![X208]:![X209]:![X210]:(((((((vplus(X205,X209)=vplus(X207,X209)&X209=X210)&rreal(X210))&X207=X208)&rreal(X208))&X205=X206)&rreal(X206))=>X205=X207)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(156, 0), 0)')).
fof(c_0_5, axiom, (rreal(vd684)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(axiom(137), 1)')).
fof(c_0_6, axiom, (v0=vd684), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(axiom(137), 2)')).
fof(c_0_7, axiom, (![X221]:![X222]:![X223]:![X224]:![X225]:(((((((((((X224=X225&ron(X223,X224))&ron(X222,X224))&ron(X221,X224))&rline(X225))&X222!=X223)&X221!=X223)&X221!=X222)&?[X226]:(X223=X226&rpoint(X226)))&?[X227]:(X222=X227&rpoint(X227)))&?[X228]:(X221=X228&rpoint(X228)))=>((rR(X222,X221,X223)|rR(X221,X222,X223))|rR(X223,X221,X222)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(91), 0), imp(cond(axiom(91), 0)))')).
fof(c_0_8, axiom, (~(rR(vd1287,vd1288,vd1329))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'neg(neg(conjunct2(372)))')).
fof(c_0_9, axiom, (ron(vd1329,vd1295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(conjunct1(368), 0)')).
fof(c_0_10, axiom, (rline(vd1295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(346, 0)')).
fof(c_0_11, axiom, (rpoint(vd1329)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(367, 0)')).
fof(c_0_12, axiom, (![X63]:![X64]:![X65]:![X66]:![X67]:((((((((X66=X67&ron(X65,X66))&ron(X63,X66))&rR(X64,X63,X65))&rline(X67))&?[X68]:(X65=X68&rpoint(X68)))&?[X69]:(X64=X69&rpoint(X69)))&?[X70]:(X63=X70&rpoint(X70)))=>ron(X64,X66))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(85), 0), imp(cond(axiom(85), 0)))')).
fof(c_0_13, axiom, (![X89]:![X90]:((?[X91]:(X90=X91&rpoint(X91))&?[X92]:(X89=X92&rpoint(X92)))=>?[X93]:(vf(X89,X90)=X93&rreal(X93)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(160), 0), imp(cond(axiom(160), 0)))')).
fof(c_0_14, axiom, ((((vd1302=vd1303&rR(vd1289,vd1287,vd1302))&rpoint(vd1303))&rR(vd1288,vd1287,vd1299))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(pred(conjunct2(348), 4), and(holds(conjunct2(348), 1304, 0), and(pred(conjunct2(348), 1), holds(conjunct1(348), 1301, 0))))')).
fof(c_0_15, axiom, (ron(vd1288,vd1295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(s3(plural(370)), 0)')).
fof(c_0_16, axiom, (ron(vd1287,vd1295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(s1(plural(346)), 0)')).
fof(c_0_17, axiom, (vd1329!=vd1287), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct2(conjunct2(conjunct2(371))), 1332, 0)')).
fof(c_0_18, axiom, (![X266]:![X267]:![X268]:((((rR(X267,X266,X268)&?[X269]:(X268=X269&rpoint(X269)))&?[X270]:(X267=X270&rpoint(X270)))&?[X271]:(X266=X271&rpoint(X271)))=>vplus(vf(X266,X267),vf(X267,X268))=vf(X266,X268))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(180), 0), imp(cond(axiom(180), 0)))')).
fof(c_0_19, axiom, (![X85]:![X86]:((?[X87]:(X86=X87&rpoint(X87))&?[X88]:(X85=X88&rpoint(X88)))=>vf(X85,X86)=vf(X86,X85))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(166), 0), imp(cond(axiom(166), 0)))')).
fof(c_0_20, axiom, (rpoint(vd1300)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(348, 1)')).
fof(c_0_21, axiom, (vd1299=vd1300), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(348, 4)')).
fof(c_0_22, axiom, (ron(vd1302,vd1297)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(conjunct2(349), 0)')).
fof(c_0_23, axiom, (rline(vd1297)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(347, 0)')).
fof(c_0_24, axiom, (rR(vd1329,vd1287,vd1299)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct1(367), 1330, 0)')).
fof(c_0_25, axiom, (ron(vd1287,vd1297)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(s1(plural(347)), 0)')).
fof(c_0_26, axiom, (vd1302!=vd1287), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct1(conjunct2(conjunct2(371))), 1334, 0)')).
fof(c_0_27, axiom, (![X26]:![X27]:![X28]:((((rR(X27,X26,X28)&?[X29]:(X28=X29&rpoint(X29)))&?[X30]:(X27=X30&rpoint(X30)))&?[X31]:(X26=X31&rpoint(X31)))=>(((~(rR(X26,X27,X28))&X26!=X28)&X26!=X27)&rR(X27,X28,X26)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(81), 0), imp(cond(axiom(81), 0)))')).
fof(c_0_28, axiom, (vd1329!=vd1302), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct1(conjunct2(conjunct2(368))), 1333, 0)')).
fof(c_0_29, axiom, (![X102]:![X103]:![X104]:![X105]:![X106]:((((((((X105=X106&ron(X103,X105))&ron(X102,X105))&rR(X103,X102,X104))&rline(X106))&?[X107]:(X104=X107&rpoint(X107)))&?[X108]:(X103=X108&rpoint(X108)))&?[X109]:(X102=X109&rpoint(X109)))=>ron(X104,X105))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(83), 0), imp(cond(axiom(83), 0)))')).
fof(c_0_30, axiom, (![X149]:![X150]:((?[X151]:(X150=X151&rpoint(X151))&?[X152]:(X149=X152&rpoint(X152)))=>(vf(X149,X150)=v0<=>X149=X150))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(162), 0), imp(cond(axiom(162), 0)))')).
fof(c_0_31, axiom, (vf(vd1287,vd1302)=vf(vd1287,vd1329)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct2(367), 1331, 0)')).
fof(c_0_32, axiom, (![X347]:![X348]:![X349]:![X350]:![X351]:((((((((rR(X349,X347,X348)&~(rinside(X349,X350)))&(ron(X347,X350)|rinside(X347,X350)))&X350=X351)&rcircle(X351))&?[X352]:(X349=X352&rpoint(X352)))&?[X353]:(X348=X353&rpoint(X353)))&?[X354]:(X347=X354&rpoint(X354)))=>(~(ron(X348,X350))&~(ron(X348,X350))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(123), 0), imp(cond(axiom(123), 0)))')).
fof(c_0_33, axiom, (![X181]:![X182]:![X183]:![X184]:(((((rinside(X181,X183)&X183=X184)&rcircle(X184))&X181=X182)&rpoint(X182))=>~(ron(X181,X183)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(79), 0), imp(cond(axiom(79), 0)))')).
fof(c_0_34, axiom, (![X200]:![X201]:(((X200!=X201&?[X202]:(X201=X202&rpoint(X202)))&?[X203]:(X200=X203&rpoint(X203)))=>?[X204]:((ron(X201,X204)&rcenter(X200,X204))&rcircle(X204)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(53), 0), imp(cond(axiom(53), 0)))')).
fof(c_0_35, axiom, (![X94]:![X95]:((X94=X95&rreal(X95))=>~(rless(X94,X94)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(149), 0), imp(cond(axiom(149), 0)))')).
fof(c_0_36, axiom, (![X290]:![X291]:![X292]:![X293]:![X294]:![X295]:(((((((rless(X290,X292)&X294=X295)&rreal(X295))&X292=X293)&rreal(X293))&X290=X291)&rreal(X291))=>rless(vplus(X290,X294),vplus(X292,X294)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(153), 0), imp(cond(axiom(153), 0)))')).
fof(c_0_37, axiom, (![X280]:![X281]:(rcenter(X280,X281)=>rinside(X280,X281))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(77), 0), imp(cond(axiom(77), 0)))')).
fof(c_0_38, axiom, (![X464]:![X465]:![X466]:![X467]:((((((rR(X465,X464,X467)&rR(X465,X464,X466))&?[X468]:(X467=X468&rpoint(X468)))&?[X469]:(X466=X469&rpoint(X469)))&?[X470]:(X465=X470&rpoint(X470)))&?[X471]:(X464=X471&rpoint(X471)))=>~(rR(X465,X466,X467)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(93), 0), imp(cond(axiom(93), 0)))')).
fof(c_0_39, axiom, (![X18]:![X19]:![X20]:![X21]:((((((rR(X20,X19,X21)&rR(X19,X18,X20))&?[X22]:(X21=X22&rpoint(X22)))&?[X23]:(X20=X23&rpoint(X23)))&?[X24]:(X19=X24&rpoint(X24)))&?[X25]:(X18=X25&rpoint(X25)))=>rR(X19,X18,X21))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(89), 0), imp(cond(axiom(89), 0)))')).
fof(c_0_40, axiom, (![X110]:![X111]:(((X110!=X111&?[X112]:(X111=X112&rpoint(X112)))&?[X113]:(X110=X113&rpoint(X113)))=>?[X114]:((ron(X111,X114)&ron(X110,X114))&rline(X114)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(51), 0), imp(cond(axiom(51), 0)))')).
fof(c_0_41, axiom, (![X137]:![X138]:![X139]:![X140]:![X141]:![X142]:((((((ron(X141,X140)&X141=X142)&rpoint(X142))&rcenter(X139,X140))&X137=X138)&rpoint(X138))=>(vf(X139,X137)=vf(X139,X141)<=>ron(X137,X140)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(184), 0), imp(cond(axiom(184), 0)))')).
fof(c_0_42, axiom, (![X696]:![X697]:![X698]:![X699]:![X700]:((((((((ron(X697,X699)&ron(X696,X699))&X696!=X697)&X699=X700)&rline(X700))&?[X701]:(X698=X701&rpoint(X701)))&?[X702]:(X697=X702&rpoint(X702)))&?[X703]:(X696=X703&rpoint(X703)))=>(vg(X696,X697,X698)=v0<=>ron(X698,X699)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(200), 0), imp(cond(axiom(200), 0)))')).
fof(c_0_43, axiom, (![X302]:![X303]:![X304]:![X305]:((((X304=X305&rreal(X305))&X302=X303)&rreal(X303))=>((rless(X302,X304)|X302=X304)|rless(X304,X302)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(145), 0), imp(cond(axiom(145), 0)))')).
fof(c_0_44, axiom, (![X173]:![X174]:![X175]:![X176]:((((((rR(X176,X173,X174)&rR(X174,X173,X175))&?[X177]:(X176=X177&rpoint(X177)))&?[X178]:(X175=X178&rpoint(X178)))&?[X179]:(X174=X179&rpoint(X179)))&?[X180]:(X173=X180&rpoint(X180)))=>rR(X176,X173,X175))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(87), 0), imp(cond(axiom(87), 0)))')).
fof(c_0_45, axiom, (![X306]:![X307]:![X308]:![X309]:((((X308=X309&rreal(X309))&X306=X307)&rreal(X307))=>?[X310]:(vplus(X306,X308)=X310&rreal(X310)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(139), 0), imp(cond(axiom(139), 0)))')).
fof(c_0_46, axiom, (![X675]:![X676]:![X677]:![X678]:![X679]:![X680]:(((((((((((((~(ron(X678,X679))&X679=X680)&rline(X680))&X677!=X678)&X676!=X678)&X676!=X677)&X675!=X678)&X675!=X677)&X675!=X676)&?[X681]:(X678=X681&rpoint(X681)))&?[X682]:(X677=X682&rpoint(X682)))&?[X683]:(X676=X683&rpoint(X683)))&?[X684]:(X675=X684&rpoint(X684)))=>(rR(X677,X675,X676)<=>vplus(vg(X675,X677,X678),vg(X678,X677,X676))=vg(X675,X678,X676)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(202), 0), imp(cond(axiom(202), 0)))')).
fof(c_0_47, axiom, (~(rless(vf(vd1287,vd1299),vf(vd1287,vd1302)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'neg(neg(366))')).
fof(c_0_48, axiom, (![X739]:![X740]:![X741]:(((?[X742]:(X741=X742&rpoint(X742))&?[X743]:(X740=X743&rpoint(X743)))&?[X744]:(X739=X744&rpoint(X744)))=>(vg(X739,X740,X741)=vg(X739,X741,X740)&vg(X739,X740,X741)=vg(X741,X739,X740)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(176), 0), imp(cond(axiom(176), 0)))')).
fof(c_0_49, axiom, (![X143]:![X144]:![X145]:(((((X143!=X145&X143!=X144)&?[X146]:(X145=X146&rpoint(X146)))&?[X147]:(X144=X147&rpoint(X147)))&?[X148]:(X143=X148&rpoint(X148)))=>vangle(X143,X144,X145)=vangle(X145,X144,X143))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(168), 0), imp(cond(axiom(168), 0)))')).
fof(c_0_50, conjecture, (vangle(vd1287,vd1289,vd1329)=vangle(vd1287,vd1288,vd1302)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct2(conjunct2(375)), 1345, 0)')).
fof(c_0_51, axiom, (ron(vd1289,vd1297)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(s2(plural(347)), 0)')).
fof(c_0_52, axiom, (ron(vd1299,vd1295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(conjunct1(349), 0)')).
fof(c_0_53, plain, (![X77]:![X78]:![X79]:![X80]:![X81]:![X82]:![X83]:![X84]:((((((((((((ron(X79,X83)&ron(X77,X83))&X77!=X81)&X77!=X79)&X83=X84)&rline(X84))&X81=X82)&rpoint(X82))&X79=X80)&rpoint(X80))&X77=X78)&rpoint(X78))=>((~rR(X77,X79,X81)&ron(X81,X83))<=>vangle(X79,X77,X81)=v0))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_54, plain, (![X85]:![X86]:![X87]:![X88]:![X89]:![X90]:![X91]:![X92]:((((rR(X85,X87,X89)|~ron(X89,X91))|vangle(X87,X85,X89)=v0)|(((((((((((~ron(X87,X91)|~ron(X85,X91))|X85=X89)|X85=X87)|X91!=X92)|~rline(X92))|X89!=X90)|~rpoint(X90))|X87!=X88)|~rpoint(X88))|X85!=X86)|~rpoint(X86)))&(((~rR(X85,X87,X89)|vangle(X87,X85,X89)!=v0)|(((((((((((~ron(X87,X91)|~ron(X85,X91))|X85=X89)|X85=X87)|X91!=X92)|~rline(X92))|X89!=X90)|~rpoint(X90))|X87!=X88)|~rpoint(X88))|X85!=X86)|~rpoint(X86)))&((ron(X89,X91)|vangle(X87,X85,X89)!=v0)|(((((((((((~ron(X87,X91)|~ron(X85,X91))|X85=X89)|X85=X87)|X91!=X92)|~rline(X92))|X89!=X90)|~rpoint(X90))|X87!=X88)|~rpoint(X88))|X85!=X86)|~rpoint(X86)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_53])])])])])).
fof(c_0_55, plain, (![X221]:![X222]:((X221!=X222|~rreal(X222))|vplus(X221,v0)=X221)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_56, plain, (![X219]:![X220]:![X221]:![X222]:((((X221!=X222|~rreal(X222))|X219!=X220)|~rreal(X220))|vplus(X219,X221)=vplus(X221,X219))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_57,plain,(X2=X4|X2=X6|vangle(X4,X2,X6)=v0|rR(X2,X4,X6)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rline(X7)|X8!=X7|~ron(X2,X8)|~ron(X4,X8)|~ron(X6,X8)), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_58, plain, (((((((vf(vd1287,vd1288)=vf(vd1287,vd1289)&vd1288!=vd1289)&vd1287!=vd1289)&vd1287!=vd1288)&(vd1289=esk1_0&rpoint(esk1_0)))&(vd1288=esk2_0&rpoint(esk2_0)))&(vd1287=esk3_0&rpoint(esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_3])])).
fof(c_0_59, plain, (![X211]:![X212]:![X213]:![X214]:![X215]:![X216]:(((((((vplus(X211,X215)!=vplus(X213,X215)|X215!=X216)|~rreal(X216))|X213!=X214)|~rreal(X214))|X211!=X212)|~rreal(X212))|X211=X213)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_60,plain,(vplus(X1,v0)=X1|~rreal(X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_61,plain,(vplus(X1,X2)=vplus(X2,X1)|~rreal(X3)|X1!=X3|~rreal(X4)|X2!=X4), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_62,plain,(rreal(vd684)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_63,plain,(v0=vd684), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_64, plain, (![X229]:![X230]:![X231]:![X232]:![X233]:![X234]:![X235]:![X236]:(((((((((((X232!=X233|~ron(X231,X232))|~ron(X230,X232))|~ron(X229,X232))|~rline(X233))|X230=X231)|X229=X231)|X229=X230)|(X231!=X234|~rpoint(X234)))|(X230!=X235|~rpoint(X235)))|(X229!=X236|~rpoint(X236)))|((rR(X230,X229,X231)|rR(X229,X230,X231))|rR(X231,X229,X230)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
fof(c_0_65, plain, (~rR(vd1287,vd1288,vd1329)), inference(fof_simplification,[status(thm)],[c_0_8])).
cnf(c_0_66,plain,(vangle(X1,X2,X3)=v0|X2=X3|X2=X1|rR(X2,X1,X3)|~rline(X4)|~ron(X3,X4)|~ron(X1,X4)|~ron(X2,X4)|~rpoint(X3)|~rpoint(X1)|~rpoint(X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_57])])])])).
cnf(c_0_67,plain,(ron(vd1329,vd1295)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_68,plain,(rline(vd1295)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_69,plain,(rpoint(vd1329)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_70,plain,(rpoint(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_71,plain,(vd1288=esk2_0), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_72,plain,(rpoint(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_73,plain,(vd1287=esk3_0), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_74, plain, (![X71]:![X72]:![X73]:![X74]:![X75]:![X76]:![X77]:![X78]:((((((((X74!=X75|~ron(X73,X74))|~ron(X71,X74))|~rR(X72,X71,X73))|~rline(X75))|(X73!=X76|~rpoint(X76)))|(X72!=X77|~rpoint(X77)))|(X71!=X78|~rpoint(X78)))|ron(X72,X74))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
cnf(c_0_75,plain,(X1=X2|~rreal(X3)|X1!=X3|~rreal(X4)|X2!=X4|~rreal(X5)|X6!=X5|vplus(X1,X6)!=vplus(X2,X6)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_76,plain,(vplus(X1,v0)=X1|~rreal(X1)), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_77,plain,(vplus(X1,X2)=vplus(X2,X1)|~rreal(X2)|~rreal(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_61])])).
cnf(c_0_78,plain,(rreal(v0)), inference(rw,[status(thm)],[c_0_62, c_0_63])).
fof(c_0_79, plain, (![X94]:![X95]:![X96]:![X97]:((vf(X94,X95)=esk7_2(X94,X95)|((X95!=X96|~rpoint(X96))|(X94!=X97|~rpoint(X97))))&(rreal(esk7_2(X94,X95))|((X95!=X96|~rpoint(X96))|(X94!=X97|~rpoint(X97)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])).
cnf(c_0_80,plain,(rR(X1,X2,X3)|rR(X2,X3,X1)|rR(X3,X2,X1)|X2=X3|X2=X1|X3=X1|~rpoint(X4)|X2!=X4|~rpoint(X5)|X3!=X5|~rpoint(X6)|X1!=X6|~rline(X7)|~ron(X2,X8)|~ron(X3,X8)|~ron(X1,X8)|X8!=X7), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_81,plain,(rpoint(vd1303)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_82,plain,(vd1302=vd1303), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_83,plain,(X2=X4|X2=X6|ron(X6,X8)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rline(X7)|X8!=X7|~ron(X2,X8)|~ron(X4,X8)|vangle(X4,X2,X6)!=v0), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_84,plain,(~rR(vd1287,vd1288,vd1329)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_85,plain,(vangle(X1,X2,vd1329)=v0|X2=vd1329|X2=X1|rR(X2,X1,vd1329)|~ron(X1,vd1295)|~ron(X2,vd1295)|~rpoint(X1)|~rpoint(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), c_0_69])])).
cnf(c_0_86,plain,(ron(vd1288,vd1295)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_87,plain,(ron(vd1287,vd1295)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_88,plain,(rpoint(vd1288)), inference(rw,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_89,plain,(rpoint(vd1287)), inference(rw,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_90,plain,(vd1329!=vd1287), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_91,plain,(vd1287!=vd1288), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_92, plain, (![X272]:![X273]:![X274]:![X275]:![X276]:![X277]:((((~rR(X273,X272,X274)|(X274!=X275|~rpoint(X275)))|(X273!=X276|~rpoint(X276)))|(X272!=X277|~rpoint(X277)))|vplus(vf(X272,X273),vf(X273,X274))=vf(X272,X274))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
fof(c_0_93, plain, (![X89]:![X90]:![X91]:![X92]:(((X90!=X91|~rpoint(X91))|(X89!=X92|~rpoint(X92)))|vf(X89,X90)=vf(X90,X89))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
cnf(c_0_94,plain,(ron(X1,X2)|~rpoint(X3)|X4!=X3|~rpoint(X5)|X1!=X5|~rpoint(X6)|X7!=X6|~rline(X8)|~rR(X1,X4,X7)|~ron(X4,X2)|~ron(X7,X2)|X2!=X8), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_95,plain,(rpoint(vd1300)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_96,plain,(vd1299=vd1300), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_97,plain,(X1=X2|vplus(X1,X3)!=vplus(X2,X3)|~rreal(X3)|~rreal(X2)|~rreal(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_75])])])).
cnf(c_0_98,plain,(vplus(v0,X1)=X1|~rreal(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])])).
cnf(c_0_99,plain,(rreal(esk7_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_100,plain,(vf(X2,X4)=esk7_2(X2,X4)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_101,plain,(X1=X2|X3=X2|X3=X1|rR(X2,X1,X3)|rR(X1,X2,X3)|rR(X3,X1,X2)|~rline(X4)|~ron(X2,X4)|~ron(X1,X4)|~ron(X3,X4)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_80])])])])).
cnf(c_0_102,plain,(ron(vd1302,vd1297)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_103,plain,(rline(vd1297)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_104,plain,(rpoint(vd1302)), inference(rw,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_105,plain,(X1=X2|X1=X3|ron(X2,X4)|vangle(X3,X1,X2)!=v0|~rline(X4)|~ron(X3,X4)|~ron(X1,X4)|~rpoint(X2)|~rpoint(X3)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_83])])])])).
cnf(c_0_106,plain,(vangle(vd1288,vd1287,vd1329)=v0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89])]), c_0_90]), c_0_91])).
cnf(c_0_107,plain,(vplus(vf(X1,X2),vf(X2,X3))=vf(X1,X3)|~rpoint(X4)|X1!=X4|~rpoint(X5)|X2!=X5|~rpoint(X6)|X3!=X6|~rR(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_108,plain,(vf(X1,X2)=vf(X2,X1)|~rpoint(X3)|X1!=X3|~rpoint(X4)|X2!=X4), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_109,plain,(ron(X1,X2)|~rline(X2)|~ron(X3,X2)|~ron(X4,X2)|~rpoint(X3)|~rpoint(X1)|~rpoint(X4)|~rR(X1,X4,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_94])])])])).
cnf(c_0_110,plain,(rR(vd1329,vd1287,vd1299)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_111,plain,(rpoint(vd1299)), inference(rw,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_112,plain,(v0=X1|vplus(X1,X2)!=X2|~rreal(X2)|~rreal(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_78])])).
cnf(c_0_113,plain,(rreal(esk7_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_99])])).
cnf(c_0_114,plain,(esk7_2(X1,X2)=vf(X1,X2)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_100])])).
cnf(c_0_115,plain,(X1=vd1302|X2=vd1302|X2=X1|rR(vd1302,X1,X2)|rR(X1,vd1302,X2)|rR(X2,X1,vd1302)|~ron(X1,vd1297)|~ron(X2,vd1297)|~rpoint(X2)|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_116,plain,(ron(vd1287,vd1297)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_117,plain,(vd1302!=vd1287), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_118,plain,(ron(vd1329,X1)|~rline(X1)|~ron(vd1288,X1)|~ron(vd1287,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_69]), c_0_88]), c_0_89])]), c_0_91]), c_0_90])).
fof(c_0_119, plain, (![X26]:![X27]:![X28]:((((rR(X27,X26,X28)&?[X29]:(X28=X29&rpoint(X29)))&?[X30]:(X27=X30&rpoint(X30)))&?[X31]:(X26=X31&rpoint(X31)))=>(((~rR(X26,X27,X28)&X26!=X28)&X26!=X27)&rR(X27,X28,X26)))), inference(fof_simplification,[status(thm)],[c_0_27])).
cnf(c_0_120,plain,(vplus(vf(X1,X2),vf(X2,X3))=vf(X1,X3)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~rR(X2,X1,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_107])])])).
cnf(c_0_121,plain,(vf(X1,X2)=vf(X2,X1)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_108])])).
cnf(c_0_122,plain,(ron(vd1329,X1)|~rline(X1)|~ron(vd1299,X1)|~ron(vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_111]), c_0_69]), c_0_89])])).
cnf(c_0_123,plain,(v0=X1|vplus(X2,X1)!=X2|~rreal(X2)|~rreal(X1)), inference(spm,[status(thm)],[c_0_112, c_0_77])).
cnf(c_0_124,plain,(rreal(vf(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_125,plain,(X1=vd1287|X1=vd1302|rR(X1,vd1287,vd1302)|rR(vd1287,vd1302,X1)|rR(vd1302,vd1287,X1)|~ron(X1,vd1297)|~rpoint(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_89])]), c_0_117])).
cnf(c_0_126,plain,(ron(vd1329,vd1297)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_103]), c_0_116])])).
cnf(c_0_127,plain,(vd1329!=vd1302), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_128, plain, (![X32]:![X33]:![X34]:![X35]:![X36]:![X37]:((((~rR(X32,X33,X34)|(((~rR(X33,X32,X34)|(X34!=X35|~rpoint(X35)))|(X33!=X36|~rpoint(X36)))|(X32!=X37|~rpoint(X37))))&(X32!=X34|(((~rR(X33,X32,X34)|(X34!=X35|~rpoint(X35)))|(X33!=X36|~rpoint(X36)))|(X32!=X37|~rpoint(X37)))))&(X32!=X33|(((~rR(X33,X32,X34)|(X34!=X35|~rpoint(X35)))|(X33!=X36|~rpoint(X36)))|(X32!=X37|~rpoint(X37)))))&(rR(X33,X34,X32)|(((~rR(X33,X32,X34)|(X34!=X35|~rpoint(X35)))|(X33!=X36|~rpoint(X36)))|(X32!=X37|~rpoint(X37)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_119])])])])).
cnf(c_0_129,plain,(vplus(vf(X1,X2),vf(X3,X2))=vf(X1,X3)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~rR(X2,X1,X3)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_130,plain,(ron(vd1329,vd1297)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_103]), c_0_116])])).
fof(c_0_131, plain, (![X110]:![X111]:![X112]:![X113]:![X114]:![X115]:![X116]:![X117]:((((((((X113!=X114|~ron(X111,X113))|~ron(X110,X113))|~rR(X111,X110,X112))|~rline(X114))|(X112!=X115|~rpoint(X115)))|(X111!=X116|~rpoint(X116)))|(X110!=X117|~rpoint(X117)))|ron(X112,X113))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])).
fof(c_0_132, plain, (![X153]:![X154]:![X155]:![X156]:(((vf(X153,X154)!=v0|X153=X154)|((X154!=X155|~rpoint(X155))|(X153!=X156|~rpoint(X156))))&((X153!=X154|vf(X153,X154)=v0)|((X154!=X155|~rpoint(X155))|(X153!=X156|~rpoint(X156)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])).
cnf(c_0_133,plain,(vf(vd1287,vd1302)=vf(vd1287,vd1329)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_134,plain,(vf(X1,X2)=v0|vf(X3,X2)!=vf(X3,X1)|~rpoint(X2)|~rpoint(X1)|~rpoint(X3)|~rR(X1,X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_120]), c_0_124]), c_0_124])).
cnf(c_0_135,plain,(rR(vd1302,vd1287,vd1329)|rR(vd1287,vd1302,vd1329)|rR(vd1329,vd1287,vd1302)|~ron(vd1288,vd1297)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_69])]), c_0_90]), c_0_127])).
fof(c_0_136, plain, (![X347]:![X348]:![X349]:![X350]:![X351]:((((((((rR(X349,X347,X348)&~rinside(X349,X350))&(ron(X347,X350)|rinside(X347,X350)))&X350=X351)&rcircle(X351))&?[X352]:(X349=X352&rpoint(X352)))&?[X353]:(X348=X353&rpoint(X353)))&?[X354]:(X347=X354&rpoint(X354)))=>~ron(X348,X350))), inference(fof_simplification,[status(thm)],[c_0_32])).
cnf(c_0_137,plain,(~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rR(X4,X2,X6)|~rR(X2,X4,X6)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_138,plain,(vf(X1,X2)=v0|vf(X3,X1)!=vf(X3,X2)|~rpoint(X1)|~rpoint(X2)|~rpoint(X3)|~rR(X2,X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_129]), c_0_124]), c_0_124])).
cnf(c_0_139,plain,(rR(vd1302,vd1287,vd1329)|rR(vd1287,vd1302,vd1329)|rR(vd1329,vd1287,vd1302)|~ron(vd1299,vd1297)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_130]), c_0_69])]), c_0_90]), c_0_127])).
cnf(c_0_140,plain,(ron(X1,X2)|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rpoint(X7)|X1!=X7|~rline(X8)|~rR(X6,X4,X1)|~ron(X4,X2)|~ron(X6,X2)|X2!=X8), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_141,plain,(X2=X4|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|vf(X2,X4)!=v0), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_142,plain,(vplus(vf(X1,vd1287),vf(vd1287,vd1302))=vf(X1,vd1329)|~rpoint(X1)|~rR(vd1287,X1,vd1329)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_133]), c_0_69]), c_0_89])])).
cnf(c_0_143,plain,(vplus(vf(X1,X2),vf(X1,X3))=vf(X2,X3)|~rpoint(X3)|~rpoint(X1)|~rpoint(X2)|~rR(X1,X2,X3)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_144,plain,(rR(X4,X6,X2)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rR(X4,X2,X6)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_145,plain,(vf(vd1329,vd1302)=v0|rR(vd1287,vd1302,vd1329)|rR(vd1302,vd1287,vd1329)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_133]), c_0_104]), c_0_69]), c_0_89])])).
fof(c_0_146, plain, (![X181]:![X182]:![X183]:![X184]:(((((rinside(X181,X183)&X183=X184)&rcircle(X184))&X181=X182)&rpoint(X182))=>~ron(X181,X183))), inference(fof_simplification,[status(thm)],[c_0_33])).
fof(c_0_147, plain, (![X355]:![X356]:![X357]:![X358]:![X359]:![X360]:![X361]:![X362]:((((((((~ron(X355,X358)|(~rR(X357,X355,X356)|rinside(X357,X358)))|X358!=X359)|~rcircle(X359))|(X357!=X360|~rpoint(X360)))|(X356!=X361|~rpoint(X361)))|(X355!=X362|~rpoint(X362)))|~ron(X356,X358))&(((((((~rinside(X355,X358)|(~rR(X357,X355,X356)|rinside(X357,X358)))|X358!=X359)|~rcircle(X359))|(X357!=X360|~rpoint(X360)))|(X356!=X361|~rpoint(X361)))|(X355!=X362|~rpoint(X362)))|~ron(X356,X358)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_136])])])])])).
cnf(c_0_148,plain,(~rpoint(X1)|~rpoint(X2)|~rpoint(X3)|~rR(X2,X3,X1)|~rR(X3,X2,X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_137])])])).
cnf(c_0_149,plain,(vf(vd1302,vd1329)=v0|rR(vd1287,vd1302,vd1329)|rR(vd1302,vd1287,vd1329)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_133]), c_0_104]), c_0_69]), c_0_89])])).
cnf(c_0_150,plain,(ron(X1,X2)|~rline(X2)|~ron(X3,X2)|~ron(X4,X2)|~rpoint(X1)|~rpoint(X3)|~rpoint(X4)|~rR(X3,X4,X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_140])])])])).
cnf(c_0_151,plain,(rR(vd1288,vd1287,vd1299)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_152,plain,(vplus(vf(X1,vd1329),vf(vd1287,vd1302))=vf(X1,vd1287)|~rpoint(X1)|~rR(vd1329,X1,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_133]), c_0_89]), c_0_69])])).
cnf(c_0_153,plain,(X1=X2|vf(X1,X2)!=v0|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_141])])).
cnf(c_0_154,plain,(vplus(vf(vd1287,X1),vf(vd1287,vd1302))=vf(X1,vd1329)|~rpoint(X1)|~rR(vd1287,X1,vd1329)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_121]), c_0_89])])).
cnf(c_0_155,plain,(vplus(vf(vd1287,vd1302),vf(vd1287,X1))=vf(vd1329,X1)|~rpoint(X1)|~rR(vd1287,vd1329,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_133]), c_0_89]), c_0_69])])).
cnf(c_0_156,plain,(rR(X1,X2,X3)|~rpoint(X2)|~rpoint(X1)|~rpoint(X3)|~rR(X1,X3,X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_144])])])).
cnf(c_0_157,plain,(vf(vd1329,vd1302)=v0|rR(vd1287,vd1302,vd1329)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_145]), c_0_133]), c_0_69]), c_0_104]), c_0_89])])).
cnf(c_0_158,plain,(vplus(vf(X1,X2),vf(X3,X1))=vf(X3,X2)|~rpoint(X2)|~rpoint(X1)|~rpoint(X3)|~rR(X1,X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_120]), c_0_124]), c_0_124])).
cnf(c_0_159,plain,(vf(vd1287,vd1288)=vf(vd1287,vd1289)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_160, plain, (![X185]:![X186]:![X187]:![X188]:(((((~rinside(X185,X187)|X187!=X188)|~rcircle(X188))|X185!=X186)|~rpoint(X186))|~ron(X185,X187))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_146])])])])).
cnf(c_0_161,plain,(rinside(X7,X2)|~ron(X1,X2)|~rpoint(X3)|X4!=X3|~rpoint(X5)|X1!=X5|~rpoint(X6)|X7!=X6|~rcircle(X8)|X2!=X8|~rR(X7,X4,X1)|~rinside(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_147])).
fof(c_0_162, plain, (![X205]:![X206]:![X207]:![X208]:(((ron(X206,esk12_2(X205,X206))|((X205=X206|(X206!=X207|~rpoint(X207)))|(X205!=X208|~rpoint(X208))))&(rcenter(X205,esk12_2(X205,X206))|((X205=X206|(X206!=X207|~rpoint(X207)))|(X205!=X208|~rpoint(X208)))))&(rcircle(esk12_2(X205,X206))|((X205=X206|(X206!=X207|~rpoint(X207)))|(X205!=X208|~rpoint(X208)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])])])).
cnf(c_0_163,plain,(rR(vd1287,vd1302,vd1329)|rR(vd1302,vd1287,vd1329)|~ron(vd1288,vd1297)|~rR(vd1287,vd1329,vd1302)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_135]), c_0_104]), c_0_69]), c_0_89])])).
cnf(c_0_164,plain,(vf(vd1302,vd1329)=v0|rR(vd1287,vd1302,vd1329)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_149]), c_0_133]), c_0_69]), c_0_104]), c_0_89])])).
cnf(c_0_165,plain,(ron(vd1299,X1)|~rline(X1)|~ron(vd1288,X1)|~ron(vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_111]), c_0_88]), c_0_89])])).
cnf(c_0_166,plain,(vplus(vf(vd1329,X1),vf(vd1287,vd1302))=vf(X1,vd1287)|~rpoint(X1)|~rR(vd1329,X1,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_121]), c_0_69])])).
cnf(c_0_167,plain,(rreal(vf(vd1287,vd1302))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_133]), c_0_69]), c_0_89])])).
cnf(c_0_168,plain,(vf(vd1287,vd1302)!=v0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_133]), c_0_69]), c_0_89])]), c_0_90])).
cnf(c_0_169,plain,(vf(vd1329,vd1302)=vf(vd1302,vd1329)|~rR(vd1287,vd1302,vd1329)|~rR(vd1287,vd1329,vd1302)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_104])])).
cnf(c_0_170,plain,(vf(vd1329,vd1302)=v0|rR(vd1287,vd1329,vd1302)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_69]), c_0_89]), c_0_104])])).
fof(c_0_171, plain, (![X94]:![X95]:((X94=X95&rreal(X95))=>~rless(X94,X94))), inference(fof_simplification,[status(thm)],[c_0_35])).
fof(c_0_172, plain, (![X296]:![X297]:![X298]:![X299]:![X300]:![X301]:(((((((~rless(X296,X298)|X300!=X301)|~rreal(X301))|X298!=X299)|~rreal(X299))|X296!=X297)|~rreal(X297))|rless(vplus(X296,X300),vplus(X298,X300)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])])).
cnf(c_0_173,plain,(vplus(vf(vd1288,X1),vf(vd1287,vd1289))=vf(vd1287,X1)|~rpoint(X1)|~rR(vd1288,vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_88]), c_0_89])])).
cnf(c_0_174,plain,(~ron(X1,X2)|~rpoint(X3)|X1!=X3|~rcircle(X4)|X2!=X4|~rinside(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_160])).
cnf(c_0_175,plain,(rinside(X1,X2)|~rcircle(X2)|~rinside(X3,X2)|~ron(X4,X2)|~rpoint(X1)|~rpoint(X4)|~rpoint(X3)|~rR(X1,X3,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_161])])])])).
fof(c_0_176, plain, (![X282]:![X283]:(~rcenter(X282,X283)|rinside(X282,X283))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])).
cnf(c_0_177,plain,(X2=X4|rcenter(X2,esk12_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_178,plain,(vf(vd1302,vd1329)=v0|rR(vd1287,vd1302,vd1329)|~ron(vd1288,vd1297)|~rR(vd1287,vd1329,vd1302)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_163]), c_0_133]), c_0_69]), c_0_104]), c_0_89])])).
cnf(c_0_179,plain,(vf(vd1302,vd1329)=v0|rR(vd1287,vd1329,vd1302)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_164]), c_0_69]), c_0_89]), c_0_104])])).
cnf(c_0_180,plain,(ron(vd1299,vd1297)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_103]), c_0_116])])).
cnf(c_0_181,plain,(vf(X1,vd1287)!=vf(vd1329,X1)|~rreal(vf(vd1329,X1))|~rpoint(X1)|~rR(vd1329,X1,vd1287)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_166]), c_0_167])]), c_0_168])).
cnf(c_0_182,plain,(vf(vd1329,vd1302)=vf(vd1302,vd1329)|vf(vd1329,vd1302)=v0|~ron(vd1288,vd1297)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170]), c_0_157])).
fof(c_0_183, plain, (![X464]:![X465]:![X466]:![X467]:((((((rR(X465,X464,X467)&rR(X465,X464,X466))&?[X468]:(X467=X468&rpoint(X468)))&?[X469]:(X466=X469&rpoint(X469)))&?[X470]:(X465=X470&rpoint(X470)))&?[X471]:(X464=X471&rpoint(X471)))=>~rR(X465,X466,X467))), inference(fof_simplification,[status(thm)],[c_0_38])).
fof(c_0_184, plain, (![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:((((((~rR(X28,X27,X29)|~rR(X27,X26,X28))|(X29!=X30|~rpoint(X30)))|(X28!=X31|~rpoint(X31)))|(X27!=X32|~rpoint(X32)))|(X26!=X33|~rpoint(X33)))|rR(X27,X26,X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])])).
fof(c_0_185, plain, (![X115]:![X116]:![X117]:![X118]:(((ron(X116,esk8_2(X115,X116))|((X115=X116|(X116!=X117|~rpoint(X117)))|(X115!=X118|~rpoint(X118))))&(ron(X115,esk8_2(X115,X116))|((X115=X116|(X116!=X117|~rpoint(X117)))|(X115!=X118|~rpoint(X118)))))&(rline(esk8_2(X115,X116))|((X115=X116|(X116!=X117|~rpoint(X117)))|(X115!=X118|~rpoint(X118)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])])])).
fof(c_0_186, plain, (![X96]:![X97]:((X96!=X97|~rreal(X97))|~rless(X96,X96))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_171])])])])).
cnf(c_0_187,plain,(rless(vplus(X1,X2),vplus(X3,X2))|~rreal(X4)|X1!=X4|~rreal(X5)|X3!=X5|~rreal(X6)|X2!=X6|~rless(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_172])).
cnf(c_0_188,plain,(vplus(vf(X1,vd1288),vf(vd1287,vd1289))=vf(X1,vd1287)|~rpoint(X1)|~rR(vd1288,X1,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_159]), c_0_89]), c_0_88])])).
cnf(c_0_189,plain,(vplus(vf(X1,vd1288),vf(vd1287,vd1289))=vf(vd1287,X1)|~rpoint(X1)|~rR(vd1288,vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_121]), c_0_88])])).
cnf(c_0_190,plain,(~rcircle(X1)|~rinside(X2,X1)|~ron(X2,X1)|~rpoint(X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_174])])).
cnf(c_0_191,plain,(rinside(vd1329,X1)|~rcircle(X1)|~rinside(vd1287,X1)|~ron(vd1299,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_110]), c_0_69]), c_0_111]), c_0_89])])).
cnf(c_0_192,plain,(rinside(X1,X2)|~rcenter(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_193,plain,(X1=X2|rcenter(X1,esk12_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_177])])).
fof(c_0_194, plain, (![X143]:![X144]:![X145]:![X146]:![X147]:![X148]:(((vf(X145,X143)!=vf(X145,X147)|ron(X143,X146))|(((((~ron(X147,X146)|X147!=X148)|~rpoint(X148))|~rcenter(X145,X146))|X143!=X144)|~rpoint(X144)))&((~ron(X143,X146)|vf(X145,X143)=vf(X145,X147))|(((((~ron(X147,X146)|X147!=X148)|~rpoint(X148))|~rcenter(X145,X146))|X143!=X144)|~rpoint(X144))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_41])])])])])).
cnf(c_0_195,plain,(vf(vd1302,vd1329)=v0|rR(vd1287,vd1302,vd1329)|~ron(vd1288,vd1297)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_179]), c_0_180])).
cnf(c_0_196,plain,(vf(X1,vd1287)!=vf(vd1329,X1)|~rpoint(X1)|~rR(vd1329,X1,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_124]), c_0_69])])).
cnf(c_0_197,plain,(rR(vd1287,vd1302,vd1329)|rR(vd1302,vd1287,vd1329)|rR(vd1329,vd1302,vd1287)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_135]), c_0_104]), c_0_69]), c_0_89])])).
cnf(c_0_198,plain,(vf(vd1329,vd1302)=v0|vf(vd1302,vd1329)!=v0|~ron(vd1288,vd1297)), inference(ef,[status(thm)],[c_0_182])).
fof(c_0_199, plain, (![X472]:![X473]:![X474]:![X475]:![X476]:![X477]:![X478]:![X479]:((((((~rR(X473,X472,X475)|~rR(X473,X472,X474))|(X475!=X476|~rpoint(X476)))|(X474!=X477|~rpoint(X477)))|(X473!=X478|~rpoint(X478)))|(X472!=X479|~rpoint(X479)))|~rR(X473,X474,X475))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_183])])])).
cnf(c_0_200,plain,(rR(X1,X2,X3)|~rpoint(X4)|X2!=X4|~rpoint(X5)|X1!=X5|~rpoint(X6)|X7!=X6|~rpoint(X8)|X3!=X8|~rR(X1,X2,X7)|~rR(X7,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_184])).
fof(c_0_201, plain, (![X704]:![X705]:![X706]:![X707]:![X708]:![X709]:![X710]:![X711]:(((vg(X704,X705,X706)!=v0|ron(X706,X707))|(((((((~ron(X705,X707)|~ron(X704,X707))|X704=X705)|X707!=X708)|~rline(X708))|(X706!=X709|~rpoint(X709)))|(X705!=X710|~rpoint(X710)))|(X704!=X711|~rpoint(X711))))&((~ron(X706,X707)|vg(X704,X705,X706)=v0)|(((((((~ron(X705,X707)|~ron(X704,X707))|X704=X705)|X707!=X708)|~rline(X708))|(X706!=X709|~rpoint(X709)))|(X705!=X710|~rpoint(X710)))|(X704!=X711|~rpoint(X711)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_42])])])])])).
cnf(c_0_202,plain,(rR(vd1288,vd1299,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_151]), c_0_111]), c_0_88]), c_0_89])])).
cnf(c_0_203,plain,(X2=X4|rline(esk8_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_204,plain,(~rless(X1,X1)|~rreal(X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_205,plain,(rless(vplus(X1,X2),vplus(X3,X2))|~rreal(X2)|~rreal(X3)|~rreal(X1)|~rless(X1,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_187])])])).
cnf(c_0_206,plain,(vf(vd1287,X1)=vf(X1,vd1287)|~rpoint(X1)|~rR(vd1288,X1,vd1287)|~rR(vd1288,vd1287,X1)), inference(spm,[status(thm)],[c_0_188, c_0_189])).
fof(c_0_207, plain, (![X306]:![X307]:![X308]:![X309]:((((X308!=X309|~rreal(X309))|X306!=X307)|~rreal(X307))|((rless(X306,X308)|X306=X308)|rless(X308,X306)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_43])])])])).
cnf(c_0_208,plain,(~rcircle(X1)|~rinside(vd1287,X1)|~ron(vd1329,X1)|~ron(vd1299,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_191]), c_0_69])])).
cnf(c_0_209,plain,(X1=X2|rinside(X1,esk12_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(spm,[status(thm)],[c_0_192, c_0_193])).
cnf(c_0_210,plain,(X2=X4|rcircle(esk12_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_211,plain,(ron(X2,X4)|~rpoint(X1)|X2!=X1|~rcenter(X3,X4)|~rpoint(X5)|X6!=X5|~ron(X6,X4)|vf(X3,X2)!=vf(X3,X6)), inference(split_conjunct,[status(thm)],[c_0_194])).
cnf(c_0_212,plain,(vf(vd1302,vd1329)=v0|~ron(vd1288,vd1297)|~rR(vd1302,vd1287,vd1329)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_195]), c_0_69]), c_0_89]), c_0_104])])).
cnf(c_0_213,plain,(rR(vd1302,vd1287,vd1329)|rR(vd1287,vd1302,vd1329)|vf(vd1329,vd1302)!=vf(vd1302,vd1287)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_197]), c_0_104])])).
cnf(c_0_214,plain,(vf(vd1302,vd1329)!=v0|~ron(vd1288,vd1297)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_198]), c_0_104]), c_0_69])]), c_0_127])).
cnf(c_0_215,plain,(~rR(X1,X2,X3)|~rpoint(X4)|X5!=X4|~rpoint(X6)|X1!=X6|~rpoint(X7)|X2!=X7|~rpoint(X8)|X3!=X8|~rR(X1,X5,X2)|~rR(X1,X5,X3)), inference(split_conjunct,[status(thm)],[c_0_199])).
fof(c_0_216, plain, (![X181]:![X182]:![X183]:![X184]:![X185]:![X186]:![X187]:![X188]:((((((~rR(X184,X181,X182)|~rR(X182,X181,X183))|(X184!=X185|~rpoint(X185)))|(X183!=X186|~rpoint(X186)))|(X182!=X187|~rpoint(X187)))|(X181!=X188|~rpoint(X188)))|rR(X184,X181,X183))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_44])])])).
cnf(c_0_217,plain,(rR(X1,X2,X3)|~rpoint(X3)|~rpoint(X4)|~rpoint(X1)|~rpoint(X2)|~rR(X4,X1,X3)|~rR(X1,X2,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_200])])])])).
cnf(c_0_218,plain,(X2=X4|vg(X2,X4,X6)=v0|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rline(X7)|X8!=X7|~ron(X2,X8)|~ron(X4,X8)|~ron(X6,X8)), inference(split_conjunct,[status(thm)],[c_0_201])).
cnf(c_0_219,plain,(X2=X4|ron(X4,esk8_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_220,plain,(ron(vd1287,X1)|~rline(X1)|~ron(vd1288,X1)|~ron(vd1299,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_202]), c_0_89]), c_0_88]), c_0_111])])).
cnf(c_0_221,plain,(X1=X2|rline(esk8_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_203])])).
cnf(c_0_222,plain,(X2=X4|ron(X2,esk8_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_223,plain,(~rreal(X1)|~rless(X1,X1)), inference(er,[status(thm)],[c_0_204])).
cnf(c_0_224,plain,(rless(X1,vplus(X2,v0))|~rreal(X2)|~rreal(X1)|~rless(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_76]), c_0_78])])).
fof(c_0_225, plain, (![X311]:![X312]:![X313]:![X314]:((vplus(X311,X313)=esk16_3(X311,X312,X313)|(((X313!=X314|~rreal(X314))|X311!=X312)|~rreal(X312)))&(rreal(esk16_3(X311,X312,X313))|(((X313!=X314|~rreal(X314))|X311!=X312)|~rreal(X312))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_45])])])])])])).
cnf(c_0_226,plain,(rless(vplus(X1,v0),X2)|~rreal(X2)|~rreal(X1)|~rless(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_76]), c_0_78])])).
cnf(c_0_227,plain,(vf(vd1299,vd1287)=vf(vd1287,vd1299)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_202]), c_0_111]), c_0_151])])).
cnf(c_0_228,plain,(rless(X1,X2)|X2=X1|rless(X2,X1)|~rreal(X3)|X2!=X3|~rreal(X4)|X1!=X4), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_229,plain,(vd1287=X1|~rcircle(esk12_2(vd1287,X1))|~ron(vd1329,esk12_2(vd1287,X1))|~ron(vd1299,esk12_2(vd1287,X1))|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209]), c_0_89])])).
cnf(c_0_230,plain,(X1=X2|rcircle(esk12_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_210])])).
cnf(c_0_231,plain,(X2=X4|ron(X4,esk12_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_232,plain,(ron(X1,X2)|vf(X3,X1)!=vf(X3,X4)|~rcenter(X3,X2)|~ron(X4,X2)|~rpoint(X4)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_211])])).
cnf(c_0_233,plain,(rR(vd1287,vd1302,vd1329)|~ron(vd1288,vd1297)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213]), c_0_195]), c_0_214])).
cnf(c_0_234,plain,(~rpoint(X1)|~rpoint(X2)|~rpoint(X3)|~rpoint(X4)|~rR(X3,X4,X1)|~rR(X3,X4,X2)|~rR(X3,X2,X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_215])])])])).
cnf(c_0_235,plain,(rR(X1,X2,X3)|~rpoint(X4)|X2!=X4|~rpoint(X5)|X6!=X5|~rpoint(X7)|X3!=X7|~rpoint(X8)|X1!=X8|~rR(X6,X2,X3)|~rR(X1,X2,X6)), inference(split_conjunct,[status(thm)],[c_0_216])).
cnf(c_0_236,plain,(rR(vd1287,X1,vd1299)|~rpoint(X1)|~rR(vd1287,X1,vd1329)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_110]), c_0_111]), c_0_69]), c_0_89])])).
cnf(c_0_237,plain,(rpoint(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_238,plain,(vd1289=esk1_0), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_239, plain, (![X675]:![X676]:![X677]:![X678]:![X679]:![X680]:(((((((((((((~ron(X678,X679)&X679=X680)&rline(X680))&X677!=X678)&X676!=X678)&X676!=X677)&X675!=X678)&X675!=X677)&X675!=X676)&?[X681]:(X678=X681&rpoint(X681)))&?[X682]:(X677=X682&rpoint(X682)))&?[X683]:(X676=X683&rpoint(X683)))&?[X684]:(X675=X684&rpoint(X684)))=>(rR(X677,X675,X676)<=>vplus(vg(X675,X677,X678),vg(X678,X677,X676))=vg(X675,X678,X676)))), inference(fof_simplification,[status(thm)],[c_0_46])).
cnf(c_0_240,plain,(vg(X1,X2,X3)=v0|X1=X2|~rline(X4)|~ron(X3,X4)|~ron(X2,X4)|~ron(X1,X4)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_218])])])])).
cnf(c_0_241,plain,(X1=X2|ron(X2,esk8_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_219])])).
cnf(c_0_242,plain,(X1=X2|ron(vd1287,esk8_2(X1,X2))|~ron(vd1288,esk8_2(X1,X2))|~ron(vd1299,esk8_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(spm,[status(thm)],[c_0_220, c_0_221])).
cnf(c_0_243,plain,(X1=X2|ron(X1,esk8_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_222])])).
cnf(c_0_244,plain,(~rreal(vplus(X1,v0))|~rreal(X1)|~rless(vplus(X1,v0),X1)), inference(spm,[status(thm)],[c_0_223, c_0_224])).
cnf(c_0_245,plain,(rreal(esk16_3(X2,X1,X4))|~rreal(X1)|X2!=X1|~rreal(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_246,plain,(vplus(X2,X4)=esk16_3(X2,X1,X4)|~rreal(X1)|X2!=X1|~rreal(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_247,plain,(~rreal(vplus(X1,v0))|~rreal(X1)|~rless(X1,vplus(X1,v0))), inference(spm,[status(thm)],[c_0_223, c_0_226])).
cnf(c_0_248,plain,(vplus(vf(vd1287,vd1302),vf(vd1329,X1))=vf(vd1287,X1)|~rpoint(X1)|~rR(vd1329,vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_133]), c_0_69]), c_0_89])])).
cnf(c_0_249,plain,(vplus(vf(vd1287,vd1299),vf(vd1299,X1))=vf(vd1287,X1)|~rpoint(X1)|~rR(vd1299,vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_227]), c_0_111]), c_0_89])])).
fof(c_0_250, plain, (~rless(vf(vd1287,vd1299),vf(vd1287,vd1302))), inference(fof_simplification,[status(thm)],[c_0_47])).
cnf(c_0_251,plain,(X1=X2|rless(X2,X1)|rless(X1,X2)|~rreal(X1)|~rreal(X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_228])])).
cnf(c_0_252,plain,(vd1287=X1|~ron(vd1329,esk12_2(vd1287,X1))|~ron(vd1299,esk12_2(vd1287,X1))|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230]), c_0_89])])).
cnf(c_0_253,plain,(X1=X2|ron(X2,esk12_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_231])])).
cnf(c_0_254,plain,(X1=X2|ron(X3,esk12_2(X1,X2))|vf(X1,X3)!=vf(X1,X4)|~ron(X4,esk12_2(X1,X2))|~rpoint(X4)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(spm,[status(thm)],[c_0_232, c_0_193])).
cnf(c_0_255,plain,(rR(vd1287,vd1329,vd1302)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_233]), c_0_69]), c_0_89]), c_0_104])])).
cnf(c_0_256,plain,(vf(vd1329,vd1302)=v0|rR(vd1287,vd1302,vd1329)|rR(vd1302,vd1287,vd1329)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_139]), c_0_133]), c_0_104]), c_0_69]), c_0_89])])).
cnf(c_0_257,plain,(~rpoint(X1)|~rR(vd1288,vd1287,X1)|~rR(vd1288,X1,vd1299)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_151]), c_0_111]), c_0_88]), c_0_89])])).
cnf(c_0_258,plain,(vangle(X1,X2,vd1302)=v0|X2=vd1302|X2=X1|rR(X2,X1,vd1302)|~ron(X1,vd1297)|~ron(X2,vd1297)|~rpoint(X1)|~rpoint(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_102]), c_0_103]), c_0_104])])).
cnf(c_0_259,plain,(rR(X1,X2,X3)|~rpoint(X1)|~rpoint(X3)|~rpoint(X4)|~rpoint(X2)|~rR(X4,X2,X3)|~rR(X1,X2,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_235])])])])).
cnf(c_0_260,plain,(rR(vd1287,vd1302,vd1299)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_233]), c_0_104])])).
cnf(c_0_261,plain,(rR(vd1289,vd1287,vd1302)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_262,plain,(rpoint(vd1289)), inference(rw,[status(thm)],[c_0_237, c_0_238])).
fof(c_0_263, plain, (![X685]:![X686]:![X687]:![X688]:![X689]:![X690]:![X691]:![X692]:![X693]:![X694]:(((~rR(X687,X685,X686)|vplus(vg(X685,X687,X688),vg(X688,X687,X686))=vg(X685,X688,X686))|((((((((((((ron(X688,X689)|X689!=X690)|~rline(X690))|X687=X688)|X686=X688)|X686=X687)|X685=X688)|X685=X687)|X685=X686)|(X688!=X691|~rpoint(X691)))|(X687!=X692|~rpoint(X692)))|(X686!=X693|~rpoint(X693)))|(X685!=X694|~rpoint(X694))))&((vplus(vg(X685,X687,X688),vg(X688,X687,X686))!=vg(X685,X688,X686)|rR(X687,X685,X686))|((((((((((((ron(X688,X689)|X689!=X690)|~rline(X690))|X687=X688)|X686=X688)|X686=X687)|X685=X688)|X685=X687)|X685=X686)|(X688!=X691|~rpoint(X691)))|(X687!=X692|~rpoint(X692)))|(X686!=X693|~rpoint(X693)))|(X685!=X694|~rpoint(X694)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_239])])])])])).
fof(c_0_264, plain, (![X745]:![X746]:![X747]:![X748]:![X749]:![X750]:((vg(X745,X746,X747)=vg(X745,X747,X746)|(((X747!=X748|~rpoint(X748))|(X746!=X749|~rpoint(X749)))|(X745!=X750|~rpoint(X750))))&(vg(X745,X746,X747)=vg(X747,X745,X746)|(((X747!=X748|~rpoint(X748))|(X746!=X749|~rpoint(X749)))|(X745!=X750|~rpoint(X750)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_48])])])])).
cnf(c_0_265,plain,(vg(X1,X2,X3)=v0|X4=X3|X1=X2|~ron(X2,esk8_2(X4,X3))|~ron(X1,esk8_2(X4,X3))|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~rpoint(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241]), c_0_221])).
cnf(c_0_266,plain,(vd1288=X1|ron(vd1287,esk8_2(vd1288,X1))|~ron(vd1299,esk8_2(vd1288,X1))|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243]), c_0_88])])).
cnf(c_0_267,plain,(~rreal(vplus(vplus(X1,v0),v0))|~rreal(vplus(X1,v0))|~rreal(X1)|~rless(vplus(vplus(X1,v0),v0),X1)), inference(spm,[status(thm)],[c_0_244, c_0_224])).
cnf(c_0_268,plain,(rreal(esk16_3(X1,X1,X2))|~rreal(X2)|~rreal(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_245])])).
cnf(c_0_269,plain,(esk16_3(X1,X1,X2)=vplus(X1,X2)|~rreal(X2)|~rreal(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_246])])).
cnf(c_0_270,plain,(~rreal(vplus(vplus(X1,v0),v0))|~rreal(vplus(X1,v0))|~rreal(X1)|~rless(X1,vplus(vplus(X1,v0),v0))), inference(spm,[status(thm)],[c_0_247, c_0_226])).
cnf(c_0_271,plain,(rless(vf(vd1287,X1),vplus(X2,vf(vd1329,X1)))|~rreal(vf(vd1329,X1))|~rreal(X2)|~rless(vf(vd1287,vd1302),X2)|~rpoint(X1)|~rR(vd1329,vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_248]), c_0_167])])).
cnf(c_0_272,plain,(vplus(vf(vd1287,vd1299),vf(X1,vd1299))=vf(vd1287,X1)|~rpoint(X1)|~rR(vd1299,vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_121]), c_0_111])])).
cnf(c_0_273,plain,(rreal(vf(vd1287,vd1299))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_227]), c_0_89]), c_0_111])])).
cnf(c_0_274,plain,(~rless(vf(vd1287,vd1299),vf(vd1287,vd1302))), inference(split_conjunct,[status(thm)],[c_0_250])).
cnf(c_0_275,plain,(vf(vd1287,vd1302)=X1|rless(X1,vf(vd1287,vd1302))|rless(vf(vd1287,vd1302),X1)|~rreal(X1)), inference(spm,[status(thm)],[c_0_251, c_0_167])).
cnf(c_0_276,plain,(~ron(vd1299,esk12_2(vd1287,vd1329))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253]), c_0_69]), c_0_89])]), c_0_90])).
cnf(c_0_277,plain,(X1=X2|ron(X3,esk12_2(X1,X2))|vf(X1,X3)!=vf(X1,X2)|~rpoint(X2)|~rpoint(X3)|~rpoint(X1)), inference(spm,[status(thm)],[c_0_254, c_0_253])).
cnf(c_0_278,plain,(rR(vd1329,X1,vd1302)|~ron(vd1288,vd1297)|~rpoint(X1)|~rR(vd1329,X1,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_255]), c_0_104]), c_0_89]), c_0_69])])).
cnf(c_0_279,plain,(vf(vd1329,vd1302)=v0|rR(vd1287,vd1302,vd1329)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_256]), c_0_133]), c_0_69]), c_0_104]), c_0_89])])).
fof(c_0_280, plain, (![X149]:![X150]:![X151]:![X152]:![X153]:![X154]:(((((X149=X151|X149=X150)|(X151!=X152|~rpoint(X152)))|(X150!=X153|~rpoint(X153)))|(X149!=X154|~rpoint(X154)))|vangle(X149,X150,X151)=vangle(X151,X150,X149))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_49])])])).
cnf(c_0_281,plain,(vangle(vd1287,vd1288,vd1302)=v0|vd1288=vd1302|~ron(vd1288,vd1297)|~rR(vd1288,vd1302,vd1299)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_258]), c_0_104]), c_0_116]), c_0_89]), c_0_88])]), c_0_91])).
cnf(c_0_282,plain,(rR(X1,vd1302,vd1299)|~ron(vd1288,vd1297)|~rpoint(X1)|~rR(X1,vd1302,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_260]), c_0_111]), c_0_89]), c_0_104])])).
cnf(c_0_283,plain,(rR(vd1289,vd1302,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_261]), c_0_104]), c_0_262]), c_0_89])])).
cnf(c_0_284,plain,(X2=X4|X2=X6|X2=X8|X4=X6|X4=X8|X6=X8|ron(X8,X10)|rR(X6,X2,X4)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rpoint(X7)|X8!=X7|~rline(X9)|X10!=X9|vplus(vg(X2,X6,X8),vg(X8,X6,X4))!=vg(X2,X8,X4)), inference(split_conjunct,[status(thm)],[c_0_263])).
cnf(c_0_285,plain,(vg(X2,X4,X6)=vg(X2,X6,X4)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5), inference(split_conjunct,[status(thm)],[c_0_264])).
cnf(c_0_286,plain,(vg(X1,X2,X3)=v0|X1=X2|X2=X3|~ron(X1,esk8_2(X2,X3))|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(spm,[status(thm)],[c_0_265, c_0_243])).
cnf(c_0_287,plain,(vd1288=vd1299|ron(vd1287,esk8_2(vd1288,vd1299))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_241]), c_0_111]), c_0_88])])).
cnf(c_0_288,plain,(~rreal(vplus(v0,v0))|~rless(vplus(v0,v0),v0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_98]), c_0_78])])).
cnf(c_0_289,plain,(rreal(vplus(X1,X2))|~rreal(X2)|~rreal(X1)), inference(spm,[status(thm)],[c_0_268, c_0_269])).
cnf(c_0_290,plain,(v0=X1|rless(X1,v0)|rless(v0,X1)|~rreal(X1)), inference(spm,[status(thm)],[c_0_251, c_0_78])).
cnf(c_0_291,plain,(~rreal(vplus(v0,v0))|~rless(v0,vplus(v0,v0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_98]), c_0_78])])).
cnf(c_0_292,plain,(~rreal(vf(vd1329,vd1299))|~rless(vf(vd1287,vd1302),vf(vd1287,vd1299))|~rR(vd1299,vd1287,vd1329)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272]), c_0_133]), c_0_273]), c_0_111]), c_0_110]), c_0_69])]), c_0_274])).
cnf(c_0_293,plain,(vf(vd1287,vd1302)=vf(vd1287,vd1299)|rless(vf(vd1287,vd1302),vf(vd1287,vd1299))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_273]), c_0_274])).
cnf(c_0_294,plain,(vf(vd1287,vd1302)!=vf(vd1287,vd1299)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_277]), c_0_133]), c_0_69]), c_0_111]), c_0_89])]), c_0_90])).
cnf(c_0_295,plain,(~rpoint(X1)|~rR(vd1329,vd1287,X1)|~rR(vd1329,X1,vd1299)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_110]), c_0_111]), c_0_69]), c_0_89])])).
cnf(c_0_296,plain,(rR(vd1329,vd1302,X1)|~ron(vd1288,vd1297)|~rpoint(X1)|~rR(vd1329,X1,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_278]), c_0_104]), c_0_69])])).
cnf(c_0_297,plain,(rR(vd1329,vd1299,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_110]), c_0_111]), c_0_69]), c_0_89])])).
cnf(c_0_298,plain,(ron(vd1288,X1)|~rline(X1)|~ron(vd1299,X1)|~ron(vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_151]), c_0_111]), c_0_88]), c_0_89])])).
cnf(c_0_299,plain,(vf(vd1329,vd1302)=v0|rR(vd1287,vd1329,vd1302)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_279]), c_0_69]), c_0_89]), c_0_104])])).
cnf(c_0_300,plain,(vangle(X1,X2,X3)=vangle(X3,X2,X1)|X1=X2|X1=X3|~rpoint(X4)|X1!=X4|~rpoint(X5)|X2!=X5|~rpoint(X6)|X3!=X6), inference(split_conjunct,[status(thm)],[c_0_280])).
cnf(c_0_301,plain,(vangle(vd1287,vd1288,vd1302)=v0|vd1288=vd1302|~ron(vd1288,vd1297)|~rR(vd1288,vd1302,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_282]), c_0_88])])).
cnf(c_0_302,plain,(vangle(X1,X2,vd1287)=v0|X2=vd1287|X2=X1|rR(X2,X1,vd1287)|~ron(X1,vd1297)|~ron(X2,vd1297)|~rpoint(X1)|~rpoint(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_116]), c_0_103]), c_0_89])])).
cnf(c_0_303,plain,(~rpoint(X1)|~rR(vd1289,vd1302,X1)|~rR(vd1289,X1,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_283]), c_0_89]), c_0_262]), c_0_104])])).
cnf(c_0_304,plain,(rR(X1,vd1302,vd1329)|~ron(vd1288,vd1297)|~rpoint(X1)|~rR(X1,vd1302,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_233]), c_0_69]), c_0_89]), c_0_104])])).
cnf(c_0_305,plain,(X1=X2|X3=X2|X3=X1|X4=X2|X4=X1|X4=X3|ron(X2,X5)|rR(X1,X4,X3)|vplus(vg(X4,X1,X2),vg(X2,X1,X3))!=vg(X4,X2,X3)|~rline(X5)|~rpoint(X2)|~rpoint(X1)|~rpoint(X3)|~rpoint(X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_284])])])])])).
cnf(c_0_306,plain,(vg(X1,X2,vd1329)=v0|X1=X2|~ron(X2,vd1295)|~ron(X1,vd1295)|~rpoint(X2)|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_67]), c_0_68]), c_0_69])])).
cnf(c_0_307,plain,(vg(X1,X2,X3)=vg(X1,X3,X2)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_285])])])).
cnf(c_0_308,plain,(vg(vd1287,vd1288,vd1299)=v0|vd1288=vd1299), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287]), c_0_111]), c_0_88]), c_0_89])]), c_0_91])).
cnf(c_0_309,plain,(~rless(vplus(v0,v0),v0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_289]), c_0_78])])).
cnf(c_0_310,plain,(vplus(X1,X2)=v0|rless(v0,vplus(X1,X2))|rless(vplus(X1,X2),v0)|~rreal(X2)|~rreal(X1)), inference(spm,[status(thm)],[c_0_290, c_0_289])).
cnf(c_0_311,plain,(~rless(v0,vplus(v0,v0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_289]), c_0_78])])).
cnf(c_0_312,plain,(~rless(vf(vd1287,vd1302),vf(vd1287,vd1299))|~rR(vd1299,vd1287,vd1329)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_124]), c_0_111]), c_0_69])])).
cnf(c_0_313,plain,(rless(vf(vd1287,vd1302),vf(vd1287,vd1299))), inference(sr,[status(thm)],[c_0_293, c_0_294])).
cnf(c_0_314,plain,(rinside(vd1289,X1)|~rcircle(X1)|~rinside(vd1287,X1)|~ron(vd1302,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_261]), c_0_262]), c_0_104]), c_0_89])])).
cnf(c_0_315,plain,(rR(vd1287,vd1302,vd1329)|rR(vd1302,vd1287,vd1329)|~ron(vd1299,vd1297)|~rR(vd1329,vd1302,vd1299)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_139]), c_0_104])])).
cnf(c_0_316,plain,(rR(vd1329,vd1302,vd1299)|~ron(vd1288,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_297]), c_0_111])])).
cnf(c_0_317,plain,(ron(vd1288,vd1297)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_103]), c_0_116])])).
cnf(c_0_318,plain,(vf(vd1329,vd1302)=vf(vd1302,vd1329)|vf(vd1329,vd1302)=v0|~ron(vd1299,vd1297)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_299]), c_0_279])).
fof(c_0_319, negated_conjecture, (vangle(vd1287,vd1289,vd1329)!=vangle(vd1287,vd1288,vd1302)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_50])])).
cnf(c_0_320,plain,(vangle(X1,X2,X3)=vangle(X3,X2,X1)|X1=X3|X1=X2|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_300])])])).
cnf(c_0_321,plain,(vangle(vd1302,vd1288,vd1287)=v0|vangle(vd1287,vd1288,vd1302)=v0|vd1288=vd1302|~ron(vd1288,vd1297)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_302]), c_0_102]), c_0_104]), c_0_88])]), c_0_91])).
cnf(c_0_322,plain,(~ron(vd1288,vd1297)|~rR(vd1289,vd1329,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_304]), c_0_69]), c_0_262]), c_0_283])])).
cnf(c_0_323,plain,(ron(vd1289,vd1297)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_324,plain,(vd1287!=vd1289), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_325,plain,(vd1329=X1|vd1329=X2|X3=vd1329|X2=X1|X3=X1|X3=X2|ron(X1,X4)|rR(X2,X3,vd1329)|vplus(vg(X3,X2,X1),v0)!=vg(X3,X1,vd1329)|~rline(X4)|~ron(X2,vd1295)|~ron(X1,vd1295)|~rpoint(X1)|~rpoint(X2)|~rpoint(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306]), c_0_69])])).
cnf(c_0_326,plain,(vg(vd1287,vd1299,vd1288)=v0|vd1288=vd1299), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_308]), c_0_111]), c_0_88]), c_0_89])])).
cnf(c_0_327,plain,(vplus(v0,v0)=v0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_310]), c_0_78])]), c_0_311])).
cnf(c_0_328,plain,(ron(vd1299,vd1295)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_329,plain,(~rR(vd1299,vd1287,vd1329)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_312, c_0_313])])).
cnf(c_0_330,plain,(~rcircle(X1)|~rinside(vd1287,X1)|~ron(vd1289,X1)|~ron(vd1302,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_314]), c_0_262])])).
cnf(c_0_331,plain,(rR(vd1302,vd1287,vd1329)|rR(vd1287,vd1302,vd1329)|~ron(vd1299,vd1297)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_316]), c_0_317])).
cnf(c_0_332,plain,(vf(vd1329,vd1302)=v0|vf(vd1302,vd1329)!=v0|~ron(vd1299,vd1297)), inference(ef,[status(thm)],[c_0_318])).
cnf(c_0_333,negated_conjecture,(vangle(vd1287,vd1289,vd1329)!=vangle(vd1287,vd1288,vd1302)), inference(split_conjunct,[status(thm)],[c_0_319])).
cnf(c_0_334,plain,(vangle(vd1287,vd1288,vd1302)=v0|vd1288=vd1302|~ron(vd1288,vd1297)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_321]), c_0_89]), c_0_88]), c_0_104])]), c_0_117])).
cnf(c_0_335,plain,(vangle(vd1329,vd1289,vd1287)=v0|vd1329=vd1289|~ron(vd1288,vd1297)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_322, c_0_302]), c_0_323]), c_0_69]), c_0_262])]), c_0_324]), c_0_126])).
cnf(c_0_336,plain,(vd1299=vd1287|vd1288=vd1299|vd1329=vd1299|vd1329=vd1288|ron(vd1288,X1)|vg(vd1287,vd1288,vd1329)!=v0|~rline(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_327]), c_0_328]), c_0_86]), c_0_88]), c_0_111]), c_0_89])]), c_0_90]), c_0_91]), c_0_329])).
cnf(c_0_337,plain,(vd1287=X1|~rcircle(esk12_2(vd1287,X1))|~ron(vd1289,esk12_2(vd1287,X1))|~ron(vd1302,esk12_2(vd1287,X1))|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_209]), c_0_89])])).
cnf(c_0_338,plain,(vf(X1,X2)=v0|vf(X2,X3)!=vf(X1,X3)|~rpoint(X3)|~rpoint(X1)|~rpoint(X2)|~rR(X1,X2,X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_143]), c_0_124]), c_0_124])).
cnf(c_0_339,plain,(rR(vd1287,vd1302,vd1329)|rR(vd1302,vd1329,vd1287)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_331]), c_0_69]), c_0_104]), c_0_89])])).
cnf(c_0_340,plain,(vf(vd1302,vd1329)!=v0|~ron(vd1299,vd1297)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_332]), c_0_104]), c_0_69])]), c_0_127])).
cnf(c_0_341,negated_conjecture,(vd1288=vd1302|vangle(vd1287,vd1289,vd1329)!=v0|~ron(vd1288,vd1297)), inference(spm,[status(thm)],[c_0_333, c_0_334])).
cnf(c_0_342,plain,(vangle(vd1287,vd1289,vd1329)=v0|vd1329=vd1289|~ron(vd1288,vd1297)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_335]), c_0_89]), c_0_262]), c_0_69])]), c_0_90])).
cnf(c_0_343,plain,(vd1329=vd1288|vd1329=vd1299|vd1288=vd1299|vd1299=vd1287|ron(vd1288,X1)|~rline(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_306]), c_0_86]), c_0_87]), c_0_88]), c_0_89])]), c_0_91])).
cnf(c_0_344,plain,(vd1287=X1|~ron(vd1289,esk12_2(vd1287,X1))|~ron(vd1302,esk12_2(vd1287,X1))|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_230]), c_0_89])])).
cnf(c_0_345,plain,(rR(vd1287,vd1302,vd1329)|~ron(vd1299,vd1297)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_339]), c_0_89]), c_0_104]), c_0_69])]), c_0_164]), c_0_340])).
cnf(c_0_346,negated_conjecture,(vd1329=vd1289|vd1288=vd1302|~ron(vd1288,vd1297)), inference(spm,[status(thm)],[c_0_341, c_0_342])).
cnf(c_0_347,plain,(vd1299=vd1287|vd1288=vd1299|vd1329=vd1299|vd1329=vd1288|ron(vd1288,vd1297)), inference(spm,[status(thm)],[c_0_343, c_0_103])).
cnf(c_0_348,plain,(~ron(vd1302,esk12_2(vd1287,vd1289))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_253]), c_0_262]), c_0_89])]), c_0_324])).
cnf(c_0_349,plain,(rinside(vd1288,X1)|~rcircle(X1)|~rinside(vd1287,X1)|~ron(vd1299,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_151]), c_0_88]), c_0_111]), c_0_89])])).
cnf(c_0_350,plain,(rR(vd1287,vd1329,vd1302)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_345]), c_0_69]), c_0_89]), c_0_104])])).
cnf(c_0_351,negated_conjecture,(vd1329=vd1288|vd1329=vd1299|vd1288=vd1299|vd1299=vd1287|vd1288=vd1302|vd1329=vd1289), inference(spm,[status(thm)],[c_0_346, c_0_347])).
cnf(c_0_352,plain,(vf(vd1287,vd1289)!=vf(vd1287,vd1302)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_277]), c_0_262]), c_0_104]), c_0_89])]), c_0_324])).
cnf(c_0_353,plain,(~rcircle(X1)|~rinside(vd1287,X1)|~ron(vd1288,X1)|~ron(vd1299,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_349]), c_0_88])])).
cnf(c_0_354,plain,(rR(vd1329,X1,vd1302)|~ron(vd1299,vd1297)|~rpoint(X1)|~rR(vd1329,X1,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_350]), c_0_104]), c_0_89]), c_0_69])])).
cnf(c_0_355,negated_conjecture,(vd1329=vd1289|vd1288=vd1302|vd1299=vd1287|vd1288=vd1299|vd1329=vd1299), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_351]), c_0_159]), c_0_352])).
cnf(c_0_356,plain,(vd1287=X1|~rcircle(esk12_2(vd1287,X1))|~ron(vd1288,esk12_2(vd1287,X1))|~ron(vd1299,esk12_2(vd1287,X1))|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_209]), c_0_89])])).
cnf(c_0_357,plain,(rR(vd1329,vd1302,X1)|~ron(vd1299,vd1297)|~rpoint(X1)|~rR(vd1329,X1,vd1287)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_354]), c_0_104]), c_0_69])])).
cnf(c_0_358,negated_conjecture,(vd1329=vd1299|vd1288=vd1299|vd1299=vd1287|vd1288=vd1302), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_355]), c_0_352])).
cnf(c_0_359,plain,(vd1287=X1|~ron(vd1288,esk12_2(vd1287,X1))|~ron(vd1299,esk12_2(vd1287,X1))|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_230]), c_0_89])])).
cnf(c_0_360,plain,(vf(X1,X2)=v0|vf(X1,X3)!=vf(X3,X2)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~rR(X2,X1,X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_129]), c_0_124]), c_0_124])).
cnf(c_0_361,plain,(rR(vd1329,vd1302,vd1299)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_297]), c_0_111])])).
cnf(c_0_362,negated_conjecture,(vd1288=vd1302|vd1299=vd1287|vd1288=vd1299), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_358]), c_0_294])).
cnf(c_0_363,plain,(~ron(vd1299,esk12_2(vd1287,vd1288))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_253]), c_0_88]), c_0_89])]), c_0_91])).
cnf(c_0_364,plain,(vf(vd1302,vd1299)!=vf(vd1299,vd1329)|~ron(vd1299,vd1297)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_360, c_0_361]), c_0_111]), c_0_69]), c_0_104])]), c_0_340])).
cnf(c_0_365,negated_conjecture,(vd1288=vd1299|vd1299=vd1287), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_362]), c_0_352])).
cnf(c_0_366,plain,(vf(vd1287,vd1289)!=vf(vd1287,vd1299)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_277]), c_0_159]), c_0_88]), c_0_111]), c_0_89])]), c_0_91])).
cnf(c_0_367,plain,(vf(vd1299,vd1329)!=vf(vd1299,vd1302)|~ron(vd1299,vd1297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_364, c_0_121]), c_0_111]), c_0_104])])).
cnf(c_0_368,negated_conjecture,(vd1299=vd1287), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_365]), c_0_366])).
cnf(c_0_369,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_367, c_0_368]), c_0_133]), c_0_368]), c_0_368]), c_0_116])]), ['proof']).
# SZS output end CNFRefutation.
EOF
