% Problem    : SEU406+3 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n091.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:55:47 CDT 2014
% CPUTime    : 
# SinE strategy is gf120_gu_R02_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X1]:![X2]:![X3]:![X4]:~((((((r2_hidden(X1,k2_xboole_0(X3,X4))&r2_hidden(X2,k2_xboole_0(X3,X4)))&~((r2_hidden(X1,k4_xboole_0(X3,X4))&r2_hidden(X2,k4_xboole_0(X3,X4)))))&~((r2_hidden(X1,X4)&r2_hidden(X2,X4))))&~((r2_hidden(X1,k4_xboole_0(X3,X4))&r2_hidden(X2,X4))))&~((r2_hidden(X1,X4)&r2_hidden(X2,k4_xboole_0(X3,X4))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', t1_latsum_1)).
fof(c_0_1, axiom, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+2.ax', d4_xboole_0)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+2.ax', d2_xboole_0)).
fof(c_0_3, axiom, (![X1]:![X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+2.ax', commutativity_k2_xboole_0)).
fof(c_0_4, negated_conjecture, (~(![X1]:![X2]:![X3]:![X4]:~((((((r2_hidden(X1,k2_xboole_0(X3,X4))&r2_hidden(X2,k2_xboole_0(X3,X4)))&~((r2_hidden(X1,k4_xboole_0(X3,X4))&r2_hidden(X2,k4_xboole_0(X3,X4)))))&~((r2_hidden(X1,X4)&r2_hidden(X2,X4))))&~((r2_hidden(X1,k4_xboole_0(X3,X4))&r2_hidden(X2,X4))))&~((r2_hidden(X1,X4)&r2_hidden(X2,k4_xboole_0(X3,X4)))))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_5, plain, (![X1]:![X2]:![X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~r2_hidden(X4,X2))))), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_6, plain, (![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((((~r2_hidden(X8,X7)|(r2_hidden(X8,X5)|r2_hidden(X8,X6)))|X7!=k2_xboole_0(X5,X6))&(((~r2_hidden(X9,X5)|r2_hidden(X9,X7))|X7!=k2_xboole_0(X5,X6))&((~r2_hidden(X9,X6)|r2_hidden(X9,X7))|X7!=k2_xboole_0(X5,X6))))&((((~r2_hidden(esk13_3(X10,X11,X12),X10)|~r2_hidden(esk13_3(X10,X11,X12),X12))|X12=k2_xboole_0(X10,X11))&((~r2_hidden(esk13_3(X10,X11,X12),X11)|~r2_hidden(esk13_3(X10,X11,X12),X12))|X12=k2_xboole_0(X10,X11)))&((r2_hidden(esk13_3(X10,X11,X12),X12)|(r2_hidden(esk13_3(X10,X11,X12),X10)|r2_hidden(esk13_3(X10,X11,X12),X11)))|X12=k2_xboole_0(X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
fof(c_0_7, negated_conjecture, ((((((r2_hidden(esk1_0,k2_xboole_0(esk3_0,esk4_0))&r2_hidden(esk2_0,k2_xboole_0(esk3_0,esk4_0)))&(~r2_hidden(esk1_0,k4_xboole_0(esk3_0,esk4_0))|~r2_hidden(esk2_0,k4_xboole_0(esk3_0,esk4_0))))&(~r2_hidden(esk1_0,esk4_0)|~r2_hidden(esk2_0,esk4_0)))&(~r2_hidden(esk1_0,k4_xboole_0(esk3_0,esk4_0))|~r2_hidden(esk2_0,esk4_0)))&(~r2_hidden(esk1_0,esk4_0)|~r2_hidden(esk2_0,k4_xboole_0(esk3_0,esk4_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
fof(c_0_8, plain, (![X3]:![X4]:k2_xboole_0(X3,X4)=k2_xboole_0(X4,X3)), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((((r2_hidden(X8,X5)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6))&((~r2_hidden(X8,X6)|~r2_hidden(X8,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(X9,X5)|r2_hidden(X9,X6))|r2_hidden(X9,X7))|X7!=k4_xboole_0(X5,X6)))&(((~r2_hidden(esk10_3(X10,X11,X12),X12)|(~r2_hidden(esk10_3(X10,X11,X12),X10)|r2_hidden(esk10_3(X10,X11,X12),X11)))|X12=k4_xboole_0(X10,X11))&(((r2_hidden(esk10_3(X10,X11,X12),X10)|r2_hidden(esk10_3(X10,X11,X12),X12))|X12=k4_xboole_0(X10,X11))&((~r2_hidden(esk10_3(X10,X11,X12),X11)|r2_hidden(esk10_3(X10,X11,X12),X12))|X12=k4_xboole_0(X10,X11)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).
cnf(c_0_10,plain,(r2_hidden(X4,X3)|r2_hidden(X4,X2)|X1!=k2_xboole_0(X2,X3)|~r2_hidden(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,negated_conjecture,(r2_hidden(esk2_0,k2_xboole_0(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,negated_conjecture,(r2_hidden(esk1_0,k2_xboole_0(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14,plain,(r2_hidden(X4,X1)|r2_hidden(X4,X3)|X1!=k4_xboole_0(X2,X3)|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_16,negated_conjecture,(r2_hidden(esk2_0,k2_xboole_0(esk4_0,esk3_0))), inference(rw,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17,negated_conjecture,(r2_hidden(esk1_0,k2_xboole_0(esk4_0,esk3_0))), inference(rw,[status(thm)],[c_0_13, c_0_12])).
cnf(c_0_18,negated_conjecture,(~r2_hidden(esk2_0,k4_xboole_0(esk3_0,esk4_0))|~r2_hidden(esk1_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19,plain,(r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_20,negated_conjecture,(r2_hidden(esk2_0,esk3_0)|r2_hidden(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21,negated_conjecture,(~r2_hidden(esk2_0,esk4_0)|~r2_hidden(esk1_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22,negated_conjecture,(r2_hidden(esk1_0,esk3_0)|r2_hidden(esk1_0,esk4_0)), inference(spm,[status(thm)],[c_0_15, c_0_17])).
cnf(c_0_23,negated_conjecture,(~r2_hidden(esk1_0,esk4_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])).
cnf(c_0_24,negated_conjecture,(~r2_hidden(esk2_0,esk4_0)|~r2_hidden(esk1_0,k4_xboole_0(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_25,negated_conjecture,(r2_hidden(esk1_0,esk3_0)), inference(sr,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26,negated_conjecture,(~r2_hidden(esk2_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_19]), c_0_25])]), c_0_23])).
cnf(c_0_27,negated_conjecture,(~r2_hidden(esk2_0,k4_xboole_0(esk3_0,esk4_0))|~r2_hidden(esk1_0,k4_xboole_0(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_28,negated_conjecture,(r2_hidden(esk2_0,esk3_0)), inference(sr,[status(thm)],[c_0_20, c_0_26])).
cnf(c_0_29,negated_conjecture,(~r2_hidden(esk1_0,k4_xboole_0(esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_19]), c_0_28])]), c_0_26])).
cnf(c_0_30,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_19]), c_0_25])]), c_0_23]), ['proof']).
# SZS output end CNFRefutation.
EOF
