% Problem    : GRP776+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n149.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:23:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_B42_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X2]:(g(X2)=>product(X2,eh)=X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos06)).
fof(c_0_1, axiom, (![X1]:![X2]:(h(X2)=>h(opp(X1)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos10)).
fof(c_0_2, axiom, (![X2]:(g(X2)=>h(f(X2)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos17)).
fof(c_0_3, axiom, (![X1]:![X2]:f(product(X2,X1))=sum(f(X2),f(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos18)).
fof(c_0_4, axiom, (g(eh)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos03)).
fof(c_0_5, axiom, (![X2]:(h(X2)=>sum(X2,opp(X2))=eg)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos15)).
fof(c_0_6, axiom, (![X3]:![X1]:![X2]:(((h(X2)&h(X1))&h(X3))=>sum(sum(X2,X1),X3)=sum(X2,sum(X1,X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos12)).
fof(c_0_7, axiom, (h(eg)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos11)).
fof(c_0_8, axiom, (![X1]:![X2]:((h(X2)&h(X1))=>h(sum(X2,X1)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos09)).
fof(c_0_9, conjecture, (![X4]:(f(eh)=eg&(~(g(X4))|f(inv(X4))=opp(f(X4))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
fof(c_0_10, axiom, (![X3]:![X1]:![X2]:(((g(X2)&g(X1))&g(X3))=>product(product(X2,X1),X3)=product(X2,product(X1,X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos04)).
fof(c_0_11, axiom, (![X2]:(h(X2)=>sum(X2,eg)=X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos14)).
fof(c_0_12, axiom, (![X2]:(g(X2)=>g(inv(X2)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos02)).
fof(c_0_13, axiom, (![X2]:(g(X2)=>product(X2,inv(X2))=eh)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos07)).
fof(c_0_14, axiom, (![X2]:(h(X2)=>sum(opp(X2),X2)=eg)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos16)).
fof(c_0_15, axiom, (![X2]:(h(X2)=>sum(eg,X2)=X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos13)).
fof(c_0_16, plain, (![X3]:(~g(X3)|product(X3,eh)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])).
fof(c_0_17, plain, (![X3]:![X4]:(~h(X3)|h(opp(X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_18, plain, (![X3]:(~g(X3)|h(f(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
fof(c_0_19, plain, (![X3]:![X4]:f(product(X4,X3))=sum(f(X4),f(X3))), inference(variable_rename,[status(thm)],[c_0_3])).
cnf(c_0_20,plain,(product(X1,eh)=X1|~g(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21,plain,(g(eh)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_22, plain, (![X3]:(~h(X3)|sum(X3,opp(X3))=eg)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:(((~h(X6)|~h(X5))|~h(X4))|sum(sum(X6,X5),X4)=sum(X6,sum(X5,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
cnf(c_0_24,plain,(h(opp(X1))|~h(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25,plain,(h(eg)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_26, plain, (![X3]:![X4]:((~h(X4)|~h(X3))|h(sum(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
cnf(c_0_27,plain,(h(f(X1))|~g(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28,plain,(f(product(X1,X2))=sum(f(X1),f(X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29,plain,(product(eh,eh)=eh), inference(spm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_30, negated_conjecture, (~(![X4]:(f(eh)=eg&(~g(X4)|f(inv(X4))=opp(f(X4)))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_9])])).
cnf(c_0_31,plain,(sum(X1,opp(X1))=eg|~h(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32,plain,(sum(sum(X1,X2),X3)=sum(X1,sum(X2,X3))|~h(X3)|~h(X2)|~h(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33,plain,(h(opp(X1))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_34,plain,(h(sum(X1,X2))|~h(X2)|~h(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35,plain,(h(sum(f(X1),f(X2)))|~g(product(X1,X2))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36,plain,(sum(f(eh),f(eh))=f(eh)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_37, negated_conjecture, (((g(esk1_0)|f(eh)!=eg)&(f(inv(esk1_0))!=opp(f(esk1_0))|f(eh)!=eg))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])])).
cnf(c_0_38,plain,(sum(X1,sum(X2,opp(sum(X1,X2))))=eg|~h(X2)|~h(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])]), c_0_34])).
cnf(c_0_39,plain,(h(f(eh))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_29]), c_0_21])]), c_0_36])).
cnf(c_0_40,negated_conjecture,(g(esk1_0)|f(eh)!=eg), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_41, plain, (![X4]:![X5]:![X6]:(((~g(X6)|~g(X5))|~g(X4))|product(product(X6,X5),X4)=product(X6,product(X5,X4)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
fof(c_0_42, plain, (![X3]:(~h(X3)|sum(X3,eg)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).
cnf(c_0_43,plain,(sum(f(eh),sum(f(eh),opp(f(eh))))=eg), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_36]), c_0_39])])).
cnf(c_0_44,negated_conjecture,(product(esk1_0,eh)=esk1_0|f(eh)!=eg), inference(spm,[status(thm)],[c_0_20, c_0_40])).
cnf(c_0_45,plain,(product(product(X1,X2),X3)=product(X1,product(X2,X3))|~g(X3)|~g(X2)|~g(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_46, plain, (![X3]:(~g(X3)|g(inv(X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
cnf(c_0_47,plain,(sum(X1,eg)=X1|~h(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48,plain,(sum(f(eh),eg)=eg), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_31]), c_0_39])])).
cnf(c_0_49,negated_conjecture,(h(sum(f(esk1_0),f(eh)))|f(eh)!=eg), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_44]), c_0_40])).
cnf(c_0_50,negated_conjecture,(sum(f(esk1_0),f(eh))=f(esk1_0)|f(eh)!=eg), inference(spm,[status(thm)],[c_0_28, c_0_44])).
cnf(c_0_51,plain,(product(product(X1,X2),eh)=product(X1,product(X2,eh))|~g(X2)|~g(X1)), inference(spm,[status(thm)],[c_0_45, c_0_21])).
cnf(c_0_52,plain,(g(inv(X1))|~g(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_53,plain,(f(eh)=eg), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_39])])).
cnf(c_0_54,negated_conjecture,(h(f(esk1_0))|f(eh)!=eg), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55,plain,(product(product(X1,inv(X2)),eh)=product(X1,product(inv(X2),eh))|~g(X1)|~g(X2)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56,negated_conjecture,(g(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_53])])).
fof(c_0_57, plain, (![X3]:(~g(X3)|product(X3,inv(X3))=eh)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).
fof(c_0_58, plain, (![X3]:(~h(X3)|sum(opp(X3),X3)=eg)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
cnf(c_0_59,negated_conjecture,(sum(f(esk1_0),eg)=f(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_53]), c_0_53])])).
cnf(c_0_60,negated_conjecture,(h(f(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_53])])).
fof(c_0_61, plain, (![X3]:(~h(X3)|sum(eg,X3)=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
cnf(c_0_62,negated_conjecture,(product(product(esk1_0,inv(X1)),eh)=product(esk1_0,product(inv(X1),eh))|~g(X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_63,plain,(product(X1,inv(X1))=eh|~g(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64,plain,(sum(opp(X1),X1)=eg|~h(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65,negated_conjecture,(sum(f(esk1_0),sum(eg,opp(f(esk1_0))))=eg), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_59]), c_0_25]), c_0_60])])).
cnf(c_0_66,plain,(sum(eg,X1)=X1|~h(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_67,negated_conjecture,(product(esk1_0,product(inv(esk1_0),eh))=eh), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_29]), c_0_56])])).
cnf(c_0_68,plain,(product(inv(X1),eh)=inv(X1)|~g(X1)), inference(spm,[status(thm)],[c_0_20, c_0_52])).
cnf(c_0_69,plain,(sum(opp(X1),sum(X1,X2))=sum(eg,X2)|~h(X2)|~h(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_64]), c_0_33])])).
cnf(c_0_70,negated_conjecture,(sum(f(esk1_0),opp(f(esk1_0)))=eg), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_33])])).
cnf(c_0_71,negated_conjecture,(product(esk1_0,inv(esk1_0))=eh), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_56])])).
cnf(c_0_72,negated_conjecture,(sum(opp(f(esk1_0)),eg)=sum(eg,opp(f(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_33]), c_0_60])])).
cnf(c_0_73,negated_conjecture,(f(eh)!=eg|f(inv(esk1_0))!=opp(f(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_74,negated_conjecture,(sum(f(esk1_0),f(inv(esk1_0)))=eg), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_71]), c_0_53])).
cnf(c_0_75,negated_conjecture,(sum(eg,opp(f(esk1_0)))=opp(f(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_72]), c_0_33])])).
cnf(c_0_76,negated_conjecture,(opp(f(esk1_0))!=f(inv(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_53])])).
cnf(c_0_77,negated_conjecture,(opp(f(esk1_0))=sum(eg,f(inv(esk1_0)))|~h(f(inv(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_74]), c_0_72]), c_0_75]), c_0_60])])).
cnf(c_0_78,plain,(h(sum(f(inv(X1)),f(eh)))|~g(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_68]), c_0_52])).
cnf(c_0_79,plain,(sum(f(inv(X1)),f(eh))=f(inv(X1))|~g(X1)), inference(spm,[status(thm)],[c_0_28, c_0_68])).
cnf(c_0_80,negated_conjecture,(~h(f(inv(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_66])).
cnf(c_0_81,plain,(h(f(inv(X1)))|~g(X1)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_82,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_56])]), ['proof']).
# SZS output end CNFRefutation.
EOF
