% Problem    : NUM346+1 : TPTP v6.1.0. Released v3.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n091.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:29:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
# and selection function SelectCQPrecWNTNp.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:![X2]:![X3]:![X10]:![X11]:![X12]:((((rdn_translate(X1,rdn_pos(X10))&rdn_translate(X2,rdn_pos(X11)))&rdn_add_with_carry(rdnn(n0),X10,X11,X12))&rdn_translate(X3,rdn_pos(X12)))=>sum(X1,X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+2.ax', sum_entry_point_pos_pos)).
fof(c_0_1, axiom, (![X21]:![X22]:![X23]:![X24]:![X25]:((rdn_digit_add(rdnn(X22),rdnn(X23),rdnn(X25),rdnn(n0))&rdn_digit_add(rdnn(X25),rdnn(X21),rdnn(X24),rdnn(n0)))=>rdn_add_with_carry(rdnn(X21),rdnn(X22),rdnn(X23),rdnn(X24)))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+2.ax', add_digit_digit_digit)).
fof(c_0_2, conjecture, (![X1]:(sum(n2,n3,X1)=>less(n4,X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sum_n2_n3_greater_n4)).
fof(c_0_3, axiom, (rdn_translate(n5,rdn_pos(rdnn(n5)))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+0.ax', rdn5)).
fof(c_0_4, axiom, (rdn_digit_add(rdnn(n5),rdnn(n0),rdnn(n5),rdnn(n0))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+2.ax', rdn_digit_add_n5_n0_n5_n0)).
fof(c_0_5, axiom, (![X1]:![X2]:![X15]:![X16]:((sum(X1,X2,X15)&sum(X1,X2,X16))=>X15=X16)), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+2.ax', unique_sum)).
fof(c_0_6, axiom, (rdn_translate(n3,rdn_pos(rdnn(n3)))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+0.ax', rdn3)).
fof(c_0_7, axiom, (rdn_digit_add(rdnn(n2),rdnn(n3),rdnn(n5),rdnn(n0))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+2.ax', rdn_digit_add_n2_n3_n5_n0)).
fof(c_0_8, axiom, (![X1]:![X2]:![X10]:![X11]:(((rdn_translate(X1,rdn_pos(X10))&rdn_translate(X2,rdn_pos(X11)))&rdn_positive_less(X10,X11))=>less(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+1.ax', less_entry_point_pos_pos)).
fof(c_0_9, axiom, (rdn_translate(n2,rdn_pos(rdnn(n2)))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+0.ax', rdn2)).
fof(c_0_10, axiom, (rdn_translate(n4,rdn_pos(rdnn(n4)))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+0.ax', rdn4)).
fof(c_0_11, axiom, (rdn_positive_less(rdnn(n4),rdnn(n5))), file('/export/starexec/sandbox/benchmark/Axioms/NUM005+1.ax', rdn_positive_less45)).
fof(c_0_12, plain, (![X13]:![X14]:![X15]:![X16]:![X17]:![X18]:((((~rdn_translate(X13,rdn_pos(X16))|~rdn_translate(X14,rdn_pos(X17)))|~rdn_add_with_carry(rdnn(n0),X16,X17,X18))|~rdn_translate(X15,rdn_pos(X18)))|sum(X13,X14,X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])).
fof(c_0_13, plain, (![X26]:![X27]:![X28]:![X29]:![X30]:((~rdn_digit_add(rdnn(X27),rdnn(X28),rdnn(X30),rdnn(n0))|~rdn_digit_add(rdnn(X30),rdnn(X26),rdnn(X29),rdnn(n0)))|rdn_add_with_carry(rdnn(X26),rdnn(X27),rdnn(X28),rdnn(X29)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_14, negated_conjecture, (~(![X1]:(sum(n2,n3,X1)=>less(n4,X1)))), inference(assume_negation,[status(cth)],[c_0_2])).
cnf(c_0_15,plain,(sum(X1,X2,X3)|~rdn_translate(X3,rdn_pos(X4))|~rdn_add_with_carry(rdnn(n0),X5,X6,X4)|~rdn_translate(X2,rdn_pos(X6))|~rdn_translate(X1,rdn_pos(X5))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(rdn_translate(n5,rdn_pos(rdnn(n5)))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_17,plain,(rdn_add_with_carry(rdnn(X1),rdnn(X2),rdnn(X3),rdnn(X4))|~rdn_digit_add(rdnn(X5),rdnn(X1),rdnn(X4),rdnn(n0))|~rdn_digit_add(rdnn(X2),rdnn(X3),rdnn(X5),rdnn(n0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(rdn_digit_add(rdnn(n5),rdnn(n0),rdnn(n5),rdnn(n0))), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_19, plain, (![X17]:![X18]:![X19]:![X20]:((~sum(X17,X18,X19)|~sum(X17,X18,X20))|X19=X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
fof(c_0_20, negated_conjecture, ((sum(n2,n3,esk1_0)&~less(n4,esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
cnf(c_0_21,plain,(sum(X1,X2,n5)|~rdn_add_with_carry(rdnn(n0),X3,X4,rdnn(n5))|~rdn_translate(X2,rdn_pos(X4))|~rdn_translate(X1,rdn_pos(X3))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22,plain,(rdn_translate(n3,rdn_pos(rdnn(n3)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23,plain,(rdn_add_with_carry(rdnn(n0),rdnn(X1),rdnn(X2),rdnn(n5))|~rdn_digit_add(rdnn(X1),rdnn(X2),rdnn(n5),rdnn(n0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24,plain,(rdn_digit_add(rdnn(n2),rdnn(n3),rdnn(n5),rdnn(n0))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_25, plain, (![X12]:![X13]:![X14]:![X15]:(((~rdn_translate(X12,rdn_pos(X14))|~rdn_translate(X13,rdn_pos(X15)))|~rdn_positive_less(X14,X15))|less(X12,X13))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
cnf(c_0_26,plain,(X1=X2|~sum(X3,X4,X2)|~sum(X3,X4,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27,negated_conjecture,(sum(n2,n3,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,plain,(sum(X1,n3,n5)|~rdn_add_with_carry(rdnn(n0),X2,rdnn(n3),rdnn(n5))|~rdn_translate(X1,rdn_pos(X2))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29,plain,(rdn_translate(n2,rdn_pos(rdnn(n2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_30,plain,(rdn_add_with_carry(rdnn(n0),rdnn(n2),rdnn(n3),rdnn(n5))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31,plain,(less(X1,X2)|~rdn_positive_less(X3,X4)|~rdn_translate(X2,rdn_pos(X4))|~rdn_translate(X1,rdn_pos(X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32,negated_conjecture,(X1=esk1_0|~sum(n2,n3,X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33,plain,(sum(n2,n3,n5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_34,plain,(less(X1,n5)|~rdn_positive_less(X2,rdnn(n5))|~rdn_translate(X1,rdn_pos(X2))), inference(spm,[status(thm)],[c_0_31, c_0_16])).
cnf(c_0_35,plain,(rdn_translate(n4,rdn_pos(rdnn(n4)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_36,plain,(rdn_positive_less(rdnn(n4),rdnn(n5))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_37,negated_conjecture,(~less(n4,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_38,negated_conjecture,(esk1_0=n5), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39,plain,(less(n4,n5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
cnf(c_0_40,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_39])]), ['proof']).
# SZS output end CNFRefutation.
EOF
