% Problem    : SEU418+4 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n123.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:18:24 CDT 2014
% CPUTime    : 
# SinE strategy is gf120_h_gu_RUU_F100_L01000
# Trying AutoSched0 for 149 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>r1_tarski(k5_relat_1(k3_xboole_0(X1,X2),X3),k3_xboole_0(k5_relat_1(X1,X3),k5_relat_1(X2,X3))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', t8_relset_2)).
fof(c_0_1, axiom, (![X1]:![X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+4.ax', t48_xboole_1)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:((r1_tarski(X1,X2)&r1_tarski(X1,X3))=>r1_tarski(X1,k3_xboole_0(X2,X3)))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+4.ax', t19_xboole_1)).
fof(c_0_3, axiom, (![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k5_relat_1(X1,X3),k5_relat_1(X2,X3))))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+10.ax', t49_relat_1)).
fof(c_0_4, axiom, (![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+3.ax', t2_boole)).
fof(c_0_5, axiom, (![X1]:![X2]:r1_tarski(k4_xboole_0(X1,X2),X1)), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+4.ax', t36_xboole_1)).
fof(c_0_6, axiom, (![X1]:![X2]:![X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+4.ax', t49_xboole_1)).
fof(c_0_7, axiom, (![X1]:k4_xboole_0(X1,k1_xboole_0)=X1), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+3.ax', t3_boole)).
fof(c_0_8, axiom, (![X1]:![X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+4.ax', t37_xboole_1)).
fof(c_0_9, axiom, (![X1]:![X2]:(v1_relat_1(X2)=>v1_relat_1(k4_xboole_0(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+10.ax', t10_relat_1)).
fof(c_0_10, negated_conjecture, (~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>r1_tarski(k5_relat_1(k3_xboole_0(X1,X2),X3),k3_xboole_0(k5_relat_1(X1,X3),k5_relat_1(X2,X3)))))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_11, negated_conjecture, ((v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(v1_relat_1(esk3_0)&~r1_tarski(k5_relat_1(k3_xboole_0(esk1_0,esk2_0),esk3_0),k3_xboole_0(k5_relat_1(esk1_0,esk3_0),k5_relat_1(esk2_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
fof(c_0_12, plain, (![X3]:![X4]:k4_xboole_0(X3,k4_xboole_0(X3,X4))=k3_xboole_0(X3,X4)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_13, plain, (![X4]:![X5]:![X6]:((~r1_tarski(X4,X5)|~r1_tarski(X4,X6))|r1_tarski(X4,k3_xboole_0(X5,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
cnf(c_0_14,negated_conjecture,(~r1_tarski(k5_relat_1(k3_xboole_0(esk1_0,esk2_0),esk3_0),k3_xboole_0(k5_relat_1(esk1_0,esk3_0),k5_relat_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,plain,(k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(r1_tarski(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,negated_conjecture,(~r1_tarski(k5_relat_1(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)),esk3_0),k4_xboole_0(k5_relat_1(esk1_0,esk3_0),k4_xboole_0(k5_relat_1(esk1_0,esk3_0),k5_relat_1(esk2_0,esk3_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_15])).
cnf(c_0_18,plain,(r1_tarski(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_16, c_0_15])).
fof(c_0_19, plain, (![X4]:![X5]:![X6]:(~v1_relat_1(X4)|(~v1_relat_1(X5)|(~v1_relat_1(X6)|(~r1_tarski(X4,X5)|r1_tarski(k5_relat_1(X4,X6),k5_relat_1(X5,X6))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
fof(c_0_20, plain, (![X2]:k3_xboole_0(X2,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_21,negated_conjecture,(~r1_tarski(k5_relat_1(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)),esk3_0),k5_relat_1(esk2_0,esk3_0))|~r1_tarski(k5_relat_1(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)),esk3_0),k5_relat_1(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22,plain,(r1_tarski(k5_relat_1(X1,X2),k5_relat_1(X3,X2))|~r1_tarski(X1,X3)|~v1_relat_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23,negated_conjecture,(v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24,negated_conjecture,(v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_25, plain, (![X3]:![X4]:r1_tarski(k4_xboole_0(X3,X4),X3)), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_26, plain, (![X4]:![X5]:![X6]:k3_xboole_0(X4,k4_xboole_0(X5,X6))=k4_xboole_0(k3_xboole_0(X4,X5),X6)), inference(variable_rename,[status(thm)],[c_0_6])).
cnf(c_0_27,plain,(k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, (![X2]:k4_xboole_0(X2,k1_xboole_0)=X2), inference(variable_rename,[status(thm)],[c_0_7])).
cnf(c_0_29,negated_conjecture,(~v1_relat_1(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)))|~r1_tarski(k5_relat_1(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)),esk3_0),k5_relat_1(esk1_0,esk3_0))|~r1_tarski(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_30,negated_conjecture,(v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_31,plain,(r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, (![X3]:![X4]:![X5]:![X6]:((k4_xboole_0(X3,X4)!=k1_xboole_0|r1_tarski(X3,X4))&(~r1_tarski(X5,X6)|k4_xboole_0(X5,X6)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
cnf(c_0_33,plain,(k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,plain,(k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_27, c_0_15])).
cnf(c_0_35,plain,(k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36,negated_conjecture,(~v1_relat_1(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)))|~r1_tarski(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_22]), c_0_30]), c_0_24]), c_0_31])])).
cnf(c_0_37,plain,(r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38,plain,(k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)=k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_15]), c_0_15])).
cnf(c_0_39,plain,(k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_40, plain, (![X3]:![X4]:(~v1_relat_1(X4)|v1_relat_1(k4_xboole_0(X4,X3)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
cnf(c_0_41,negated_conjecture,(~v1_relat_1(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39]), c_0_35]), c_0_39])])).
cnf(c_0_42,plain,(v1_relat_1(k4_xboole_0(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_43,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_30])]), ['proof']).
# SZS output end CNFRefutation.
EOF
