% Problem    : SWW293+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n141.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:41:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X10]:c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),hAPP(c_COMBK(tc_HOL_Obool,X10),c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__def)).
fof(c_0_1, axiom, (![X19]:![X10]:hAPP(c_Set_OCollect(X10),X19)=X19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__def)).
fof(c_0_2, axiom, (![X335]:![X87]:![X10]:![X94]:hAPP(hAPP(c_COMBK(X94,X10),X87),X335)=X87), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBK__1)).
fof(c_0_3, axiom, (![X10]:c_Orderings_Otop__class_Otop(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),hAPP(c_COMBK(tc_HOL_Obool,X10),c_fTrue))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__def)).
fof(c_0_4, axiom, (![X94]:(class_Lattices_Oboolean__algebra(X94)=>hAPP(c_Groups_Ouminus__class_Ouminus(X94),c_Orderings_Obot__class_Obot(X94))=c_Orderings_Otop__class_Otop(X94))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_compl__bot__eq)).
fof(c_0_5, axiom, (![X16]:![X13]:![X10]:(class_Orderings_Obot(X10)=>hAPP(c_Orderings_Obot__class_Obot(tc_fun(X13,X10)),X16)=c_Orderings_Obot__class_Obot(X10))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_bot__apply)).
fof(c_0_6, axiom, (class_Lattices_Oboolean__algebra(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Lattices_Oboolean__algebra)).
fof(c_0_7, axiom, (class_Orderings_Obot(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Obot)).
fof(c_0_8, axiom, (![X19]:![X10]:(c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),X19)<=>![X3]:~(hBOOL(hAPP(X19,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__Collect__eq)).
fof(c_0_9, axiom, (![X25]:![X16]:![X10]:(class_Lattices_Oboolean__algebra(X10)=>(hAPP(c_Groups_Ouminus__class_Ouminus(X10),X16)=hAPP(c_Groups_Ouminus__class_Ouminus(X10),X25)<=>X16=X25))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_compl__eq__compl__iff)).
fof(c_0_10, axiom, (![X16]:![X13]:![X10]:(class_Orderings_Otop(X10)=>hAPP(c_Orderings_Otop__class_Otop(tc_fun(X13,X10)),X16)=c_Orderings_Otop__class_Otop(X10))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_top__apply)).
fof(c_0_11, axiom, (![X19]:![X10]:hAPP(c_Set_OCollect(X10),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X10),c_fNot),X19))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),hAPP(c_Set_OCollect(X10),X19))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__neg__eq)).
fof(c_0_12, axiom, (class_Orderings_Otop(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Otop)).
fof(c_0_13, axiom, (![X93]:![X94]:(class_Lattices_Oboolean__algebra(X94)=>hAPP(c_Groups_Ouminus__class_Ouminus(X94),hAPP(c_Groups_Ouminus__class_Ouminus(X94),X93))=X93)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__compl)).
fof(c_0_14, axiom, (![X154]:![X10]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),X154))=X154), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__complement)).
fof(c_0_15, axiom, (![X19]:(~(hBOOL(hAPP(c_fNot,X19)))|~(hBOOL(X19)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__1)).
fof(c_0_16, axiom, (![X16]:![X39]:![X125]:![X10]:![X13]:hAPP(c_Sum__Type_Osum_Osum__case(X13,X10,X125,hAPP(c_COMBK(X10,X13),X39),hAPP(c_COMBK(X10,X125),X39)),X16)=X39), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_sum__case__KK)).
fof(c_0_17, axiom, (![X10]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),c_Orderings_Otop__class_Otop(tc_fun(X10,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__UNIV__eq)).
fof(c_0_18, axiom, (![X101]:![X20]:![X19]:(hBOOL(hAPP(X19,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X20,X101)))<=>((X101=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=>hBOOL(hAPP(X19,X20)))&(X101!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=>![X152]:![X123]:(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X123),X101))=>(X20=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Nat_Onat),X101),X152)),X123)=>hBOOL(hAPP(X19,X123)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_split__mod)).
fof(c_0_19, axiom, (hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_nat__0)).
fof(c_0_20, axiom, (c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Pls__def)).
fof(c_0_21, axiom, (![X196]:![X17]:![X19]:![X125]:![X10]:![X13]:hAPP(hAPP(hAPP(c_COMBB(X13,X10,X125),X19),X17),X196)=hAPP(X19,hAPP(X17,X196))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBB__1)).
fof(c_0_22, axiom, (~(hBOOL(c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fFalse__1)).
fof(c_0_23, axiom, (![X84]:![X94]:(class_Divides_Osemiring__div(X94)=>c_Divides_Odiv__class_Omod(X94,X84,c_Groups_Ozero__class_Ozero(X94))=X84)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_mod__by__0)).
fof(c_0_24, axiom, (![X297]:![X94]:hAPP(c_List_Orev(X94),hAPP(c_List_Orev(X94),X297))=X297), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rev__rev__ident)).
fof(c_0_25, axiom, (class_Divides_Osemiring__div(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Nat__Onat__Divides_Osemiring__div)).
fof(c_0_26, axiom, (![X43]:c_Com_Ocom_OSKIP!=hAPP(c_Com_Ocom_OBODY,X43)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_com_Osimps_I18_J)).
fof(c_0_27, axiom, (![X19]:(~(~(hBOOL(X19)))|hBOOL(hAPP(c_fNot,X19)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__2)).
fof(c_0_28, axiom, (![X10]:![X19]:(hBOOL(hAPP(X19,c_Hilbert__Choice_OEps(X10,X19)))<=>?[X147]:hBOOL(hAPP(X19,X147)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_some__eq__ex)).
fof(c_0_29, axiom, (![X14]:![X15]:(hAPP(c_Com_Ocom_OBODY,X15)=hAPP(c_Com_Ocom_OBODY,X14)<=>X15=X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_com_Osimps_I6_J)).
fof(c_0_30, axiom, (![X5]:![X40]:![X29]:![X30]:c_Com_Ocom_OAss(X30,X29)!=c_Com_Ocom_OWhile(X40,X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_com_Osimps_I28_J)).
fof(c_0_31, axiom, (![X5]:![X40]:c_Com_Ocom_OWhile(X40,X5)!=c_Com_Ocom_OSKIP), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_com_Osimps_I17_J)).
fof(c_0_32, axiom, (![X40]:![X41]:c_Com_Ocom_OSKIP!=c_Com_Ocom_OAss(X41,X40)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_com_Osimps_I8_J)).
fof(c_0_33, plain, (![X11]:c_Orderings_Obot__class_Obot(tc_fun(X11,tc_HOL_Obool))=hAPP(c_Set_OCollect(X11),hAPP(c_COMBK(tc_HOL_Obool,X11),c_fFalse))), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_34, plain, (![X20]:![X21]:hAPP(c_Set_OCollect(X21),X20)=X20), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_35, plain, (![X336]:![X337]:![X338]:![X339]:hAPP(hAPP(c_COMBK(X339,X338),X337),X336)=X337), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_36,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_37,plain,(hAPP(c_Set_OCollect(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_38, plain, (![X11]:c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool))=hAPP(c_Set_OCollect(X11),hAPP(c_COMBK(tc_HOL_Obool,X11),c_fTrue))), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_39, plain, (![X95]:(~class_Lattices_Oboolean__algebra(X95)|hAPP(c_Groups_Ouminus__class_Ouminus(X95),c_Orderings_Obot__class_Obot(X95))=c_Orderings_Otop__class_Otop(X95))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
fof(c_0_40, plain, (![X17]:![X18]:![X19]:(~class_Orderings_Obot(X19)|hAPP(c_Orderings_Obot__class_Obot(tc_fun(X18,X19)),X17)=c_Orderings_Obot__class_Obot(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
cnf(c_0_41,plain,(hAPP(hAPP(c_COMBK(X1,X2),X3),X4)=X3), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(X1),c_Orderings_Obot__class_Obot(X1))=c_Orderings_Otop__class_Otop(X1)|~class_Lattices_Oboolean__algebra(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45,plain,(class_Lattices_Oboolean__algebra(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_46,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,X2)),X3)=c_Orderings_Obot__class_Obot(X2)|~class_Orderings_Obot(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47,plain,(class_Orderings_Obot(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_48,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)),X2)=c_fFalse), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_49, plain, (![X19]:![X10]:(c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),X19)<=>![X3]:~hBOOL(hAPP(X19,X3)))), inference(fof_simplification,[status(thm)],[c_0_8])).
fof(c_0_50, plain, (![X26]:![X27]:![X28]:(((hAPP(c_Groups_Ouminus__class_Ouminus(X28),X27)!=hAPP(c_Groups_Ouminus__class_Ouminus(X28),X26)|X27=X26)|~class_Lattices_Oboolean__algebra(X28))&((X27!=X26|hAPP(c_Groups_Ouminus__class_Ouminus(X28),X27)=hAPP(c_Groups_Ouminus__class_Ouminus(X28),X26))|~class_Lattices_Oboolean__algebra(X28)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
fof(c_0_51, plain, (![X17]:![X18]:![X19]:(~class_Orderings_Otop(X19)|hAPP(c_Orderings_Otop__class_Otop(tc_fun(X18,X19)),X17)=c_Orderings_Otop__class_Otop(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
cnf(c_0_52,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue)=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_43, c_0_37])).
cnf(c_0_53,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_Orderings_Obot__class_Obot(tc_HOL_Obool))=c_Orderings_Otop__class_Otop(tc_HOL_Obool)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54,plain,(c_Orderings_Obot__class_Obot(tc_HOL_Obool)=c_fFalse), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
fof(c_0_55, plain, (![X20]:![X21]:hAPP(c_Set_OCollect(X21),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X21),c_fNot),X20))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X21,tc_HOL_Obool)),hAPP(c_Set_OCollect(X21),X20))), inference(variable_rename,[status(thm)],[c_0_11])).
fof(c_0_56, plain, (![X20]:![X21]:![X22]:![X23]:![X25]:((c_Orderings_Obot__class_Obot(tc_fun(X21,tc_HOL_Obool))!=hAPP(c_Set_OCollect(X21),X20)|~hBOOL(hAPP(X20,X22)))&(hBOOL(hAPP(X23,esk309_1(X23)))|c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool))=hAPP(c_Set_OCollect(X25),X23)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_49])])])])])).
cnf(c_0_57,plain,(X2=X3|~class_Lattices_Oboolean__algebra(X1)|hAPP(c_Groups_Ouminus__class_Ouminus(X1),X2)!=hAPP(c_Groups_Ouminus__class_Ouminus(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,X2)),X3)=c_Orderings_Otop__class_Otop(X2)|~class_Orderings_Otop(X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59,plain,(class_Orderings_Otop(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_60,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)),X2)=c_fTrue), inference(spm,[status(thm)],[c_0_41, c_0_52])).
cnf(c_0_61,plain,(c_Orderings_Otop__class_Otop(tc_HOL_Obool)=hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fFalse)), inference(rw,[status(thm)],[c_0_53, c_0_54])).
fof(c_0_62, plain, (![X95]:![X96]:(~class_Lattices_Oboolean__algebra(X96)|hAPP(c_Groups_Ouminus__class_Ouminus(X96),hAPP(c_Groups_Ouminus__class_Ouminus(X96),X95))=X95)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).
fof(c_0_63, plain, (![X155]:![X156]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X156,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X156,tc_HOL_Obool)),X155))=X155), inference(variable_rename,[status(thm)],[c_0_14])).
cnf(c_0_64,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Set_OCollect(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_65, plain, (![X19]:(~hBOOL(hAPP(c_fNot,X19))|~hBOOL(X19))), inference(fof_simplification,[status(thm)],[c_0_15])).
fof(c_0_66, plain, (![X126]:![X127]:![X128]:![X129]:![X130]:hAPP(c_Sum__Type_Osum_Osum__case(X130,X129,X128,hAPP(c_COMBK(X129,X130),X127),hAPP(c_COMBK(X129,X128),X127)),X126)=X127), inference(variable_rename,[status(thm)],[c_0_16])).
cnf(c_0_67,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),X2)|hBOOL(hAPP(X2,esk309_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_68, plain, (![X11]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X11,tc_HOL_Obool)),c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X11,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_17])).
fof(c_0_69, plain, (![X153]:![X154]:![X155]:![X156]:![X157]:![X158]:![X159]:![X160]:((((X153!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|hBOOL(hAPP(X155,X154)))|~hBOOL(hAPP(X155,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X154,X153))))&((X153=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X157),X153))|(X154!=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Nat_Onat),X153),X156)),X157)|hBOOL(hAPP(X155,X157)))))|~hBOOL(hAPP(X155,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X154,X153)))))&((((X158!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|X158=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))|hBOOL(hAPP(X160,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X159,X158))))&(((hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),esk70_3(X158,X159,X160)),X158))|X158=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))|hBOOL(hAPP(X160,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X159,X158))))&(((X159=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Nat_Onat),X158),esk69_3(X158,X159,X160))),esk70_3(X158,X159,X160))|X158=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))|hBOOL(hAPP(X160,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X159,X158))))&((~hBOOL(hAPP(X160,esk70_3(X158,X159,X160)))|X158=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))|hBOOL(hAPP(X160,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X159,X158)))))))&(((X158!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|~hBOOL(hAPP(X160,X159)))|hBOOL(hAPP(X160,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X159,X158))))&(((hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),esk70_3(X158,X159,X160)),X158))|~hBOOL(hAPP(X160,X159)))|hBOOL(hAPP(X160,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X159,X158))))&(((X159=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Nat_Onat),X158),esk69_3(X158,X159,X160))),esk70_3(X158,X159,X160))|~hBOOL(hAPP(X160,X159)))|hBOOL(hAPP(X160,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X159,X158))))&((~hBOOL(hAPP(X160,esk70_3(X158,X159,X160)))|~hBOOL(hAPP(X160,X159)))|hBOOL(hAPP(X160,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X159,X158)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])])])).
cnf(c_0_70,plain,(hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_71,plain,(c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_72,plain,(X1=X2|hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)!=hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X2)), inference(spm,[status(thm)],[c_0_57, c_0_45])).
cnf(c_0_73,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60]), c_0_61])).
cnf(c_0_74,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(X1),hAPP(c_Groups_Ouminus__class_Ouminus(X1),X2))=X2|~class_Lattices_Oboolean__algebra(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_75, plain, (![X197]:![X198]:![X199]:![X200]:![X201]:![X202]:hAPP(hAPP(hAPP(c_COMBB(X202,X201,X200),X199),X198),X197)=hAPP(X199,hAPP(X198,X197))), inference(variable_rename,[status(thm)],[c_0_21])).
cnf(c_0_76,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_77,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2)=hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_37]), c_0_37])).
fof(c_0_78, plain, (![X20]:(~hBOOL(hAPP(c_fNot,X20))|~hBOOL(X20))), inference(variable_rename,[status(thm)],[c_0_65])).
cnf(c_0_79,plain,(hAPP(c_Sum__Type_Osum_Osum__case(X1,X2,X3,hAPP(c_COMBK(X2,X1),X4),hAPP(c_COMBK(X2,X3),X4)),X5)=X4), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_80,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hBOOL(hAPP(X1,esk309_1(X1)))), inference(rw,[status(thm)],[c_0_67, c_0_37])).
cnf(c_0_81,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_82,plain,(hBOOL(hAPP(X1,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X2,X3)))|~hBOOL(hAPP(X1,X2))|X3!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_83,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_84,plain,(c_fFalse=X1|hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)!=c_fTrue), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_85,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1))=X1), inference(spm,[status(thm)],[c_0_74, c_0_45])).
fof(c_0_86, plain, (~hBOOL(c_fFalse)), inference(fof_simplification,[status(thm)],[c_0_22])).
fof(c_0_87, plain, (![X95]:![X96]:(~class_Divides_Osemiring__div(X96)|c_Divides_Odiv__class_Omod(X96,X95,c_Groups_Ozero__class_Ozero(X96))=X95)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])).
cnf(c_0_88,plain,(hAPP(hAPP(hAPP(c_COMBB(X1,X2,X3),X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_89,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_77]), c_0_77])).
fof(c_0_90, plain, (![X298]:![X299]:hAPP(c_List_Orev(X299),hAPP(c_List_Orev(X299),X298))=X298), inference(variable_rename,[status(thm)],[c_0_24])).
cnf(c_0_91,plain,(~hBOOL(X1)|~hBOOL(hAPP(c_fNot,X1))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_92,plain,(c_fFalse=X1|hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_48]), c_0_79])).
cnf(c_0_93,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_81, c_0_77])).
cnf(c_0_94,plain,(hBOOL(hAPP(X1,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X2,X3)))|hAPP(c_Int_Onat,c_Int_OPls)!=X3|~hBOOL(hAPP(X1,X2))), inference(rw,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_95,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)=c_fFalse|X1!=c_fTrue), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_96,plain,(~hBOOL(c_fFalse)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_97,plain,(c_Divides_Odiv__class_Omod(X1,X2,c_Groups_Ozero__class_Ozero(X1))=X2|~class_Divides_Osemiring__div(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_98,plain,(class_Divides_Osemiring__div(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_99,plain,(hAPP(c_fNot,hAPP(c_fNot,hAPP(X1,X2)))=hAPP(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_88])).
cnf(c_0_100,plain,(hAPP(c_List_Orev(X1),hAPP(c_List_Orev(X1),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_101,plain,(hAPP(c_fNot,X1)=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_102,plain,(hAPP(c_fNot,c_fTrue)=c_fFalse), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_93]), c_0_48]), c_0_60])).
cnf(c_0_103,plain,(c_Divides_Odiv__class_Omod(tc_Nat_Onat,X1,X2)!=c_fTrue|hAPP(c_Int_Onat,c_Int_OPls)!=X2|~hBOOL(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])).
cnf(c_0_104,plain,(c_Divides_Odiv__class_Omod(tc_Nat_Onat,X1,hAPP(c_Int_Onat,c_Int_OPls))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_83])).
cnf(c_0_105,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fTrue)=c_fFalse), inference(spm,[status(thm)],[c_0_85, c_0_73])).
fof(c_0_106, plain, (![X44]:c_Com_Ocom_OSKIP!=hAPP(c_Com_Ocom_OBODY,X44)), inference(variable_rename,[status(thm)],[c_0_26])).
cnf(c_0_107,plain,(hAPP(c_fNot,hAPP(c_fNot,X1))=X1), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_108,plain,(hAPP(c_fNot,X1)=c_fFalse|c_fFalse=X1), inference(spm,[status(thm)],[c_0_101, c_0_92])).
cnf(c_0_109,plain,(hAPP(c_fNot,c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_60]), c_0_102])).
cnf(c_0_110,plain,(X1!=c_fTrue|~hBOOL(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_111,plain,(c_fTrue=X1|hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)!=c_fFalse), inference(spm,[status(thm)],[c_0_72, c_0_105])).
fof(c_0_112, plain, (![X19]:(hBOOL(X19)|hBOOL(hAPP(c_fNot,X19)))), inference(fof_simplification,[status(thm)],[c_0_27])).
fof(c_0_113, plain, (![X148]:![X149]:![X151]:![X152]:![X153]:((~hBOOL(hAPP(X149,c_Hilbert__Choice_OEps(X148,X149)))|hBOOL(hAPP(X149,esk351_2(X148,X149))))&(~hBOOL(hAPP(X152,X153))|hBOOL(hAPP(X152,c_Hilbert__Choice_OEps(X151,X152)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])])])).
cnf(c_0_114,plain,(c_Com_Ocom_OSKIP!=hAPP(c_Com_Ocom_OBODY,X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_115,plain,(c_fFalse=X1|c_fTrue=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109])).
cnf(c_0_116,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)!=c_fTrue|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_110, c_0_85])).
cnf(c_0_117,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)=c_fTrue|X1!=c_fFalse), inference(spm,[status(thm)],[c_0_111, c_0_85])).
fof(c_0_118, plain, (![X20]:(hBOOL(X20)|hBOOL(hAPP(c_fNot,X20)))), inference(variable_rename,[status(thm)],[c_0_112])).
fof(c_0_119, plain, (![X16]:![X17]:![X18]:![X19]:((hAPP(c_Com_Ocom_OBODY,X17)!=hAPP(c_Com_Ocom_OBODY,X16)|X17=X16)&(X19!=X18|hAPP(c_Com_Ocom_OBODY,X19)=hAPP(c_Com_Ocom_OBODY,X18)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])).
cnf(c_0_120,plain,(hBOOL(hAPP(X1,esk351_2(X2,X1)))|~hBOOL(hAPP(X1,c_Hilbert__Choice_OEps(X2,X1)))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_121,plain,(hAPP(c_Com_Ocom_OBODY,X1)=c_fFalse|c_fTrue!=c_Com_Ocom_OSKIP), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_122,plain,(X1!=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_123,plain,(hBOOL(hAPP(c_fNot,X1))|hBOOL(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_124,plain,(X1=X2|hAPP(c_Com_Ocom_OBODY,X1)!=hAPP(c_Com_Ocom_OBODY,X2)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_125,plain,(c_fTrue!=c_Com_Ocom_OSKIP|~hBOOL(hAPP(c_Com_Ocom_OBODY,c_Hilbert__Choice_OEps(X1,c_Com_Ocom_OBODY)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_96])).
cnf(c_0_126,plain,(hBOOL(X1)|hAPP(c_fNot,X1)!=c_fFalse), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_127,plain,(X1=X2|c_fTrue!=c_Com_Ocom_OSKIP), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_121]), c_0_121])).
cnf(c_0_128,plain,(c_fTrue!=c_Com_Ocom_OSKIP), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_127])).
cnf(c_0_129,plain,(c_fFalse=X1|X1!=c_Com_Ocom_OSKIP), inference(spm,[status(thm)],[c_0_128, c_0_115])).
fof(c_0_130, plain, (![X41]:![X42]:![X43]:![X44]:c_Com_Ocom_OAss(X44,X43)!=c_Com_Ocom_OWhile(X42,X41)), inference(variable_rename,[status(thm)],[c_0_30])).
cnf(c_0_131,plain,(c_fFalse=c_Com_Ocom_OSKIP), inference(er,[status(thm)],[c_0_129])).
fof(c_0_132, plain, (![X41]:![X42]:c_Com_Ocom_OWhile(X42,X41)!=c_Com_Ocom_OSKIP), inference(variable_rename,[status(thm)],[c_0_31])).
cnf(c_0_133,plain,(c_Com_Ocom_OAss(X1,X2)!=c_Com_Ocom_OWhile(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_134,plain,(c_fTrue=X1|c_Com_Ocom_OSKIP=X1), inference(rw,[status(thm)],[c_0_115, c_0_131])).
cnf(c_0_135,plain,(c_Com_Ocom_OWhile(X1,X2)!=c_Com_Ocom_OSKIP), inference(split_conjunct,[status(thm)],[c_0_132])).
fof(c_0_136, plain, (![X42]:![X43]:c_Com_Ocom_OSKIP!=c_Com_Ocom_OAss(X43,X42)), inference(variable_rename,[status(thm)],[c_0_32])).
cnf(c_0_137,plain,(c_Com_Ocom_OAss(X1,X2)!=c_fTrue), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_135])).
cnf(c_0_138,plain,(c_Com_Ocom_OSKIP!=c_Com_Ocom_OAss(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_139,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_134]), c_0_138]), ['proof']).
# SZS output end CNFRefutation.
EOF
