% Problem    : SWW395+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n135.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:20:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X7]:![X5]:(c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),X7)<=>![X3]:~(hBOOL(hAPP(X7,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__Collect__eq)).
fof(c_0_1, axiom, (![X5]:c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),hAPP(c_COMBK(tc_HOL_Obool,X5),c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__def)).
fof(c_0_2, axiom, (![X7]:![X5]:hAPP(c_Set_OCollect(X5),X7)=X7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__def)).
fof(c_0_3, axiom, (![X7]:![X5]:hAPP(c_Set_OCollect(X5),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X5),c_fNot),X7))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),hAPP(c_Set_OCollect(X5),X7))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__neg__eq)).
fof(c_0_4, axiom, (![X314]:![X313]:![X5]:![X62]:hAPP(hAPP(c_COMBK(X62,X5),X313),X314)=X313), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBK__1)).
fof(c_0_5, axiom, (![X24]:![X5]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),X24))=X24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__complement)).
fof(c_0_6, axiom, (![X7]:(~(hBOOL(hAPP(c_fNot,X7)))|~(hBOOL(X7)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__1)).
fof(c_0_7, axiom, (![X5]:c_Orderings_Otop__class_Otop(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),hAPP(c_COMBK(tc_HOL_Obool,X5),c_fTrue))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__def)).
fof(c_0_8, axiom, (![X128]:c_Groups_Ozero__class_Ozero(tc_Nat_Onat)!=hAPP(c_Nat_OSuc,X128)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Zero__not__Suc)).
fof(c_0_9, axiom, (c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_transfer__nat__int__numerals_I1_J)).
fof(c_0_10, axiom, (c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Pls__def)).
fof(c_0_11, axiom, (![X32]:![X8]:![X7]:![X63]:![X5]:![X22]:hAPP(hAPP(hAPP(c_COMBB(X22,X5,X63),X7),X8),X32)=hAPP(X7,hAPP(X8,X32))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBB__1)).
fof(c_0_12, axiom, (![X5]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X5,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__empty__eq)).
fof(c_0_13, conjecture, (![X27]:![X28]:(hBOOL(hAPP(hAPP(v_P,X27),X28))=>?[X315]:?[X316]:(c_Hoare__Mirabelle_Ohoare__derivs(t_a,v_G,hAPP(hAPP(c_Set_Oinsert(tc_Hoare__Mirabelle_Otriple(t_a)),hAPP(hAPP(hAPP(c_Hoare__Mirabelle_Otriple_Otriple(t_a),X315),v_c),X316)),c_Orderings_Obot__class_Obot(tc_fun(tc_Hoare__Mirabelle_Otriple(t_a),tc_HOL_Obool))))&![X33]:(![X34]:(hBOOL(hAPP(hAPP(X315,X34),X28))=>hBOOL(hAPP(hAPP(X316,X34),X33)))=>hBOOL(hAPP(hAPP(v_Q,X27),X33)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_1)).
fof(c_0_14, axiom, (![X9]:![X26]:![X5]:(class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(X5)=>(X26=hAPP(hAPP(c_Groups_Oplus__class_Oplus(X5),X26),X9)<=>X9=c_Groups_Ozero__class_Ozero(X5)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_add__0__iff)).
fof(c_0_15, axiom, (![X127]:![X128]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X128),hAPP(c_Nat_OSuc,X127))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X128),X127))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_add__Suc__right)).
fof(c_0_16, axiom, (![X79]:![X5]:![X22]:hAPP(c_Relation_Oconverse(X22,X5),hAPP(c_Relation_Oconverse(X5,X22),X79))=X79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_converse__converse)).
fof(c_0_17, axiom, (class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Nat__Onat__Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct)).
fof(c_0_18, axiom, (![X128]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X128),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X128), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Nat_Oadd__0__right)).
fof(c_0_19, axiom, (![X127]:X127!=hAPP(c_Nat_OSuc,X127)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_n__not__Suc__n)).
fof(c_0_20, plain, (![X7]:![X5]:(c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),X7)<=>![X3]:~hBOOL(hAPP(X7,X3)))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_21, plain, (![X6]:c_Orderings_Obot__class_Obot(tc_fun(X6,tc_HOL_Obool))=hAPP(c_Set_OCollect(X6),hAPP(c_COMBK(tc_HOL_Obool,X6),c_fFalse))), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_22, plain, (![X8]:![X9]:hAPP(c_Set_OCollect(X9),X8)=X8), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_23, plain, (![X8]:![X9]:hAPP(c_Set_OCollect(X9),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X9),c_fNot),X8))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X9,tc_HOL_Obool)),hAPP(c_Set_OCollect(X9),X8))), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_24, plain, (![X8]:![X9]:![X10]:![X11]:![X13]:((c_Orderings_Obot__class_Obot(tc_fun(X9,tc_HOL_Obool))!=hAPP(c_Set_OCollect(X9),X8)|~hBOOL(hAPP(X8,X10)))&(hBOOL(hAPP(X11,esk13_1(X11)))|c_Orderings_Obot__class_Obot(tc_fun(X13,tc_HOL_Obool))=hAPP(c_Set_OCollect(X13),X11)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])])).
fof(c_0_25, plain, (![X315]:![X316]:![X317]:![X318]:hAPP(hAPP(c_COMBK(X318,X317),X316),X315)=X316), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_26,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27,plain,(hAPP(c_Set_OCollect(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, (![X25]:![X26]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X26,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X26,tc_HOL_Obool)),X25))=X25), inference(variable_rename,[status(thm)],[c_0_5])).
cnf(c_0_29,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Set_OCollect(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, (![X7]:(~hBOOL(hAPP(c_fNot,X7))|~hBOOL(X7))), inference(fof_simplification,[status(thm)],[c_0_6])).
cnf(c_0_31,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),X2)|hBOOL(hAPP(X2,esk13_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(hAPP(hAPP(c_COMBK(X1,X2),X3),X4)=X3), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_34, plain, (![X6]:c_Orderings_Otop__class_Otop(tc_fun(X6,tc_HOL_Obool))=hAPP(c_Set_OCollect(X6),hAPP(c_COMBK(tc_HOL_Obool,X6),c_fTrue))), inference(variable_rename,[status(thm)],[c_0_7])).
fof(c_0_35, plain, (![X129]:c_Groups_Ozero__class_Ozero(tc_Nat_Onat)!=hAPP(c_Nat_OSuc,X129)), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_36,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37,plain,(c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_38, plain, (![X64]:![X65]:![X66]:![X67]:![X68]:![X69]:hAPP(hAPP(hAPP(c_COMBB(X69,X68,X67),X66),X65),X64)=hAPP(X66,hAPP(X65,X64))), inference(variable_rename,[status(thm)],[c_0_11])).
cnf(c_0_39,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_40,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2)=hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_27]), c_0_27])).
fof(c_0_41, plain, (![X8]:(~hBOOL(hAPP(c_fNot,X8))|~hBOOL(X8))), inference(variable_rename,[status(thm)],[c_0_30])).
cnf(c_0_42,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hBOOL(hAPP(X1,esk13_1(X1)))), inference(rw,[status(thm)],[c_0_31, c_0_27])).
cnf(c_0_43,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)),X2)=c_fFalse), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_44, plain, (![X6]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X6,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X6,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X6,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_12])).
cnf(c_0_45,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_46, negated_conjecture, (~(![X27]:![X28]:(hBOOL(hAPP(hAPP(v_P,X27),X28))=>?[X315]:?[X316]:(c_Hoare__Mirabelle_Ohoare__derivs(t_a,v_G,hAPP(hAPP(c_Set_Oinsert(tc_Hoare__Mirabelle_Otriple(t_a)),hAPP(hAPP(hAPP(c_Hoare__Mirabelle_Otriple_Otriple(t_a),X315),v_c),X316)),c_Orderings_Obot__class_Obot(tc_fun(tc_Hoare__Mirabelle_Otriple(t_a),tc_HOL_Obool))))&![X33]:(![X34]:(hBOOL(hAPP(hAPP(X315,X34),X28))=>hBOOL(hAPP(hAPP(X316,X34),X33)))=>hBOOL(hAPP(hAPP(v_Q,X27),X33))))))), inference(assume_negation,[status(cth)],[c_0_13])).
fof(c_0_47, plain, (![X27]:![X28]:![X29]:(((X28!=hAPP(hAPP(c_Groups_Oplus__class_Oplus(X29),X28),X27)|X27=c_Groups_Ozero__class_Ozero(X29))|~class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(X29))&((X27!=c_Groups_Ozero__class_Ozero(X29)|X28=hAPP(hAPP(c_Groups_Oplus__class_Oplus(X29),X28),X27))|~class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(X29)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
fof(c_0_48, plain, (![X129]:![X130]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X130),hAPP(c_Nat_OSuc,X129))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X130),X129))), inference(variable_rename,[status(thm)],[c_0_15])).
cnf(c_0_49,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)!=hAPP(c_Nat_OSuc,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_50,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_51,plain,(hAPP(hAPP(hAPP(c_COMBB(X1,X2,X3),X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_52,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40]), c_0_40])).
fof(c_0_53, plain, (![X80]:![X81]:![X82]:hAPP(c_Relation_Oconverse(X82,X81),hAPP(c_Relation_Oconverse(X81,X82),X80))=X80), inference(variable_rename,[status(thm)],[c_0_16])).
cnf(c_0_54,plain,(~hBOOL(X1)|~hBOOL(hAPP(c_fNot,X1))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55,plain,(c_fFalse=X1|hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_42]), c_0_43]), c_0_32])).
cnf(c_0_56,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_57,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue)=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_45, c_0_27])).
fof(c_0_58, negated_conjecture, (![X319]:![X320]:![X322]:(hBOOL(hAPP(hAPP(v_P,esk523_0),esk524_0))&(((~hBOOL(hAPP(hAPP(X319,X322),esk524_0))|hBOOL(hAPP(hAPP(X320,X322),esk525_2(X319,X320))))|~c_Hoare__Mirabelle_Ohoare__derivs(t_a,v_G,hAPP(hAPP(c_Set_Oinsert(tc_Hoare__Mirabelle_Otriple(t_a)),hAPP(hAPP(hAPP(c_Hoare__Mirabelle_Otriple_Otriple(t_a),X319),v_c),X320)),c_Orderings_Obot__class_Obot(tc_fun(tc_Hoare__Mirabelle_Otriple(t_a),tc_HOL_Obool)))))&(~hBOOL(hAPP(hAPP(v_Q,esk523_0),esk525_2(X319,X320)))|~c_Hoare__Mirabelle_Ohoare__derivs(t_a,v_G,hAPP(hAPP(c_Set_Oinsert(tc_Hoare__Mirabelle_Otriple(t_a)),hAPP(hAPP(hAPP(c_Hoare__Mirabelle_Otriple_Otriple(t_a),X319),v_c),X320)),c_Orderings_Obot__class_Obot(tc_fun(tc_Hoare__Mirabelle_Otriple(t_a),tc_HOL_Obool)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_46])])])])])).
cnf(c_0_59,plain,(X2=c_Groups_Ozero__class_Ozero(X1)|~class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(X1)|X3!=hAPP(hAPP(c_Groups_Oplus__class_Oplus(X1),X3),X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_60,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),hAPP(c_Nat_OSuc,X2))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_61,plain,(class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_62,plain,(hAPP(c_Nat_OSuc,X1)!=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_49, c_0_50])).
fof(c_0_63, plain, (![X129]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X129),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X129), inference(variable_rename,[status(thm)],[c_0_18])).
cnf(c_0_64,plain,(hAPP(c_fNot,hAPP(c_fNot,hAPP(X1,X2)))=hAPP(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_51])).
cnf(c_0_65,plain,(hAPP(c_Relation_Oconverse(X1,X2),hAPP(c_Relation_Oconverse(X2,X1),X3))=X3), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_66,plain,(hAPP(c_fNot,X1)=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_67,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_56, c_0_40])).
cnf(c_0_68,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)),X2)=c_fTrue), inference(spm,[status(thm)],[c_0_32, c_0_57])).
cnf(c_0_69,negated_conjecture,(hBOOL(hAPP(hAPP(v_P,esk523_0),esk524_0))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_70,plain,(hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2))!=X1), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_50]), c_0_61])]), c_0_62])).
cnf(c_0_71,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X1), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_72, plain, (![X128]:X128!=hAPP(c_Nat_OSuc,X128)), inference(variable_rename,[status(thm)],[c_0_19])).
cnf(c_0_73,plain,(hAPP(c_fNot,hAPP(c_fNot,X1))=X1), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74,plain,(hAPP(c_fNot,X1)=c_fFalse|c_fFalse=X1), inference(spm,[status(thm)],[c_0_66, c_0_55])).
cnf(c_0_75,plain,(hAPP(c_fNot,c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_67]), c_0_68]), c_0_43])).
cnf(c_0_76,negated_conjecture,(hAPP(c_fNot,hAPP(hAPP(v_P,esk523_0),esk524_0))=c_fFalse), inference(spm,[status(thm)],[c_0_66, c_0_69])).
cnf(c_0_77,plain,(hAPP(c_Nat_OSuc,hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2)))!=X1), inference(spm,[status(thm)],[c_0_70, c_0_60])).
cnf(c_0_78,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),hAPP(c_Int_Onat,c_Int_OPls))=X1), inference(rw,[status(thm)],[c_0_71, c_0_50])).
cnf(c_0_79,plain,(X1!=hAPP(c_Nat_OSuc,X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_80,plain,(c_fFalse=X1|c_fTrue=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75])).
cnf(c_0_81,negated_conjecture,(hAPP(hAPP(v_P,esk523_0),esk524_0)=c_fTrue), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_76]), c_0_75])).
cnf(c_0_82,plain,(hAPP(c_Nat_OSuc,hAPP(c_Nat_OSuc,X1))!=X1), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_83,plain,(hAPP(c_Nat_OSuc,X1)=c_fFalse|c_fTrue!=X1), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_84,negated_conjecture,(hAPP(v_P,esk523_0)=c_fFalse|hAPP(c_fTrue,esk524_0)=c_fTrue), inference(spm,[status(thm)],[c_0_81, c_0_80])).
cnf(c_0_85,plain,(hAPP(c_Nat_OSuc,c_fFalse)!=X1|c_fTrue!=X1), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_86,negated_conjecture,(hAPP(c_fTrue,esk524_0)=c_fTrue|hAPP(c_fFalse,esk524_0)=c_fTrue), inference(spm,[status(thm)],[c_0_81, c_0_84])).
cnf(c_0_87,plain,(c_fTrue!=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_80]), c_0_79])).
cnf(c_0_88,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_86, c_0_87]), c_0_87]), ['proof']).
# SZS output end CNFRefutation.
EOF
