% Problem    : GEO286+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n121.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:43:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X49]:![X50]:![X51]:![X52]:![X53]:![X54]:![X55]:![X56]:(((((((((((((ron(X55,X51)&~(ron(X55,X49)))&~(rinside(X55,X49)))&X55=X56)&rpoint(X56))&ron(X53,X51))&rinside(X53,X49))&X53=X54)&rpoint(X54))&X51=X52)&rline(X52))&X49=X50)&rcircle(X50))=>?[X57]:(((rR(X57,X53,X55)&ron(X57,X49))&ron(X57,X51))&rpoint(X57)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(61), 0), imp(cond(axiom(61), 0)))')).
fof(c_0_1, conjecture, (?[X1]:(((rR(X1,vd1123,vd1125)&ron(X1,vd1152))&ron(X1,vd1151))&rpoint(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qe(261)')).
fof(c_0_2, axiom, (((((((((((rless(vf(vd1127,vd1129),vf(vd1123,vd1125))&vd1127!=vd1129)&vd1123!=vd1125)&vd1129=vd1130)&rpoint(vd1130))&vd1127=vd1128)&rpoint(vd1128))&vd1125=vd1126)&rpoint(vd1126))&vd1123=vd1124)&rpoint(vd1124))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(and(and(and(holds(conjunct2(conjunct2(conjunct2(conjunct2(conjunct2(conjunct2(242)))))), 1133, 0), and(holds(conjunct1(conjunct2(conjunct2(conjunct2(conjunct2(conjunct2(242)))))), 1132, 0), and(holds(conjunct1(conjunct2(conjunct2(conjunct2(conjunct2(242))))), 1131, 0), and(pred(conjunct1(conjunct2(conjunct2(conjunct2(242)))), 2), pred(conjunct1(conjunct2(conjunct2(conjunct2(242)))), 1))))), and(pred(conjunct1(conjunct2(conjunct2(242))), 2), pred(conjunct1(conjunct2(conjunct2(242))), 1))), and(pred(conjunct1(conjunct2(242)), 2), pred(conjunct1(conjunct2(242)), 1))), and(pred(conjunct1(242), 2), pred(conjunct1(242), 1)))')).
fof(c_0_3, axiom, (vd1123=vd1129), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(258, 1150, 0)')).
fof(c_0_4, axiom, (![X359]:![X360]:(rcenter(X359,X360)=>rinside(X359,X360))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(77), 0), imp(cond(axiom(77), 0)))')).
fof(c_0_5, axiom, (rline(vd1151)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(259, 0)')).
fof(c_0_6, axiom, (((((ron(vd1127,vd1152)&rcenter(vd1123,vd1152))&rcircle(vd1152))&ron(vd1125,vd1151))&ron(vd1123,vd1151))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(and(pred(comma_conjunct2(comma_conjunct2(comma_conjunct2(259))), 0), and(pred(comma_conjunct1(comma_conjunct2(comma_conjunct2(259))), 0), and(pred(comma_conjunct2(comma_conjunct2(259)), 0), pred(comma_conjunct1(comma_conjunct2(259)), 0)))), pred(comma_conjunct1(259), 0))')).
fof(c_0_7, axiom, (~(rinside(vd1125,vd1152))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'neg(neg(260))')).
fof(c_0_8, axiom, (~(ron(vd1125,vd1152))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'neg(neg(comma_conjunct2(260)))')).
fof(c_0_9, plain, (![X49]:![X50]:![X51]:![X52]:![X53]:![X54]:![X55]:![X56]:(((((((((((((ron(X55,X51)&~ron(X55,X49))&~rinside(X55,X49))&X55=X56)&rpoint(X56))&ron(X53,X51))&rinside(X53,X49))&X53=X54)&rpoint(X54))&X51=X52)&rline(X52))&X49=X50)&rcircle(X50))=>?[X57]:(((rR(X57,X53,X55)&ron(X57,X49))&ron(X57,X51))&rpoint(X57)))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_10, negated_conjecture, (~(?[X1]:(((rR(X1,vd1123,vd1125)&ron(X1,vd1152))&ron(X1,vd1151))&rpoint(X1)))), inference(assume_negation,[status(cth)],[c_0_1])).
fof(c_0_11, plain, (![X58]:![X59]:![X60]:![X61]:![X62]:![X63]:![X64]:![X65]:((((rR(esk4_7(X58,X59,X60,X61,X62,X63,X64),X62,X64)|((((((((((((~ron(X64,X60)|ron(X64,X58))|rinside(X64,X58))|X64!=X65)|~rpoint(X65))|~ron(X62,X60))|~rinside(X62,X58))|X62!=X63)|~rpoint(X63))|X60!=X61)|~rline(X61))|X58!=X59)|~rcircle(X59)))&(ron(esk4_7(X58,X59,X60,X61,X62,X63,X64),X58)|((((((((((((~ron(X64,X60)|ron(X64,X58))|rinside(X64,X58))|X64!=X65)|~rpoint(X65))|~ron(X62,X60))|~rinside(X62,X58))|X62!=X63)|~rpoint(X63))|X60!=X61)|~rline(X61))|X58!=X59)|~rcircle(X59))))&(ron(esk4_7(X58,X59,X60,X61,X62,X63,X64),X60)|((((((((((((~ron(X64,X60)|ron(X64,X58))|rinside(X64,X58))|X64!=X65)|~rpoint(X65))|~ron(X62,X60))|~rinside(X62,X58))|X62!=X63)|~rpoint(X63))|X60!=X61)|~rline(X61))|X58!=X59)|~rcircle(X59))))&(rpoint(esk4_7(X58,X59,X60,X61,X62,X63,X64))|((((((((((((~ron(X64,X60)|ron(X64,X58))|rinside(X64,X58))|X64!=X65)|~rpoint(X65))|~ron(X62,X60))|~rinside(X62,X58))|X62!=X63)|~rpoint(X63))|X60!=X61)|~rline(X61))|X58!=X59)|~rcircle(X59))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])])).
cnf(c_0_12,plain,(vd1129=vd1130), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_13,plain,(vd1123=vd1129), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_14, negated_conjecture, (![X2]:(((~rR(X2,vd1123,vd1125)|~ron(X2,vd1152))|~ron(X2,vd1151))|~rpoint(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
cnf(c_0_15,plain,(rinside(X8,X2)|ron(X8,X2)|rR(esk4_7(X2,X1,X4,X3,X6,X5,X8),X6,X8)|~rcircle(X1)|X2!=X1|~rline(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rinside(X6,X2)|~ron(X6,X4)|~rpoint(X7)|X8!=X7|~ron(X8,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(rpoint(vd1126)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_17,plain,(vd1125=vd1126), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_18,plain,(rpoint(vd1130)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_19,plain,(vd1130=vd1123), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_20,negated_conjecture,(~rpoint(X1)|~ron(X1,vd1151)|~ron(X1,vd1152)|~rR(X1,vd1123,vd1125)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(rinside(X1,X2)|ron(X1,X2)|rR(esk4_7(X2,X2,X3,X3,X4,X4,X1),X4,X1)|~rline(X3)|~rcircle(X2)|~rinside(X4,X2)|~rpoint(X1)|~rpoint(X4)|~ron(X1,X3)|~ron(X4,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_15])])])])).
cnf(c_0_22,plain,(rpoint(vd1125)), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23,plain,(rpoint(vd1123)), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24,plain,(rinside(X8,X2)|ron(X8,X2)|ron(esk4_7(X2,X1,X4,X3,X6,X5,X8),X4)|~rcircle(X1)|X2!=X1|~rline(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rinside(X6,X2)|~ron(X6,X4)|~rpoint(X7)|X8!=X7|~ron(X8,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_25, plain, (![X361]:![X362]:(~rcenter(X361,X362)|rinside(X361,X362))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
cnf(c_0_26,negated_conjecture,(rinside(vd1125,X1)|ron(vd1125,X1)|~rline(X2)|~rcircle(X1)|~rinside(vd1123,X1)|~rpoint(esk4_7(X1,X1,X2,X2,vd1123,vd1123,vd1125))|~ron(esk4_7(X1,X1,X2,X2,vd1123,vd1123,vd1125),vd1152)|~ron(esk4_7(X1,X1,X2,X2,vd1123,vd1123,vd1125),vd1151)|~ron(vd1125,X2)|~ron(vd1123,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])])).
cnf(c_0_27,plain,(rinside(X1,X2)|ron(esk4_7(X2,X2,X3,X3,X4,X4,X1),X3)|ron(X1,X2)|~rline(X3)|~rcircle(X2)|~rinside(X4,X2)|~rpoint(X1)|~rpoint(X4)|~ron(X1,X3)|~ron(X4,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_24])])])])).
cnf(c_0_28,plain,(rline(vd1151)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29,plain,(ron(vd1125,vd1151)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_30,plain,(ron(vd1123,vd1151)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_31,plain,(rinside(X8,X2)|ron(X8,X2)|ron(esk4_7(X2,X1,X4,X3,X6,X5,X8),X2)|~rcircle(X1)|X2!=X1|~rline(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rinside(X6,X2)|~ron(X6,X4)|~rpoint(X7)|X8!=X7|~ron(X8,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_32,plain,(rinside(X1,X2)|~rcenter(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33,plain,(rcenter(vd1123,vd1152)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_34, plain, (~rinside(vd1125,vd1152)), inference(fof_simplification,[status(thm)],[c_0_7])).
fof(c_0_35, plain, (~ron(vd1125,vd1152)), inference(fof_simplification,[status(thm)],[c_0_8])).
cnf(c_0_36,negated_conjecture,(rinside(vd1125,X1)|ron(vd1125,X1)|~rcircle(X1)|~rinside(vd1123,X1)|~rpoint(esk4_7(X1,X1,vd1151,vd1151,vd1123,vd1123,vd1125))|~ron(esk4_7(X1,X1,vd1151,vd1151,vd1123,vd1123,vd1125),vd1152)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_22]), c_0_23])])).
cnf(c_0_37,plain,(rinside(X1,X2)|ron(esk4_7(X2,X2,X3,X3,X4,X4,X1),X2)|ron(X1,X2)|~rline(X3)|~rcircle(X2)|~rinside(X4,X2)|~rpoint(X1)|~rpoint(X4)|~ron(X1,X3)|~ron(X4,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])])])])).
cnf(c_0_38,plain,(rcircle(vd1152)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_39,plain,(rinside(vd1123,vd1152)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40,plain,(~rinside(vd1125,vd1152)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41,plain,(~ron(vd1125,vd1152)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42,plain,(rinside(X8,X2)|ron(X8,X2)|rpoint(esk4_7(X2,X1,X4,X3,X6,X5,X8))|~rcircle(X1)|X2!=X1|~rline(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rinside(X6,X2)|~ron(X6,X4)|~rpoint(X7)|X8!=X7|~ron(X8,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_43,negated_conjecture,(~rpoint(esk4_7(vd1152,vd1152,vd1151,vd1151,vd1123,vd1123,vd1125))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39]), c_0_28]), c_0_22]), c_0_23]), c_0_29]), c_0_30])]), c_0_40]), c_0_41])).
cnf(c_0_44,plain,(rinside(X1,X2)|rpoint(esk4_7(X2,X2,X3,X3,X4,X4,X1))|ron(X1,X2)|~rline(X3)|~rcircle(X2)|~rinside(X4,X2)|~rpoint(X1)|~rpoint(X4)|~ron(X1,X3)|~ron(X4,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_42])])])])).
cnf(c_0_45,negated_conjecture,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_28]), c_0_38]), c_0_39]), c_0_22]), c_0_23]), c_0_29]), c_0_30])]), c_0_40]), c_0_41]), ['proof']).
# SZS output end CNFRefutation.
EOF
