% Problem    : REL029+4 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n146.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:01:01 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4d
# and selection function SelectCQIPrecWNTNp.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:![X2]:converse(composition(X1,X2))=composition(converse(X2),converse(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', converse_multiplicativity)).
fof(c_0_1, axiom, (![X1]:composition(X1,one)=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', composition_identity)).
fof(c_0_2, axiom, (![X1]:converse(converse(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', converse_idempotence)).
fof(c_0_3, axiom, (![X1]:![X2]:join(composition(converse(X1),complement(composition(X1,X2))),complement(X2))=complement(X2)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', converse_cancellativity)).
fof(c_0_4, axiom, (![X1]:![X2]:join(X1,X2)=join(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', maddux1_join_commutativity)).
fof(c_0_5, axiom, (![X1]:zero=meet(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', def_zero)).
fof(c_0_6, axiom, (![X1]:![X2]:meet(X1,X2)=complement(join(complement(X1),complement(X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', maddux4_definiton_of_meet)).
fof(c_0_7, axiom, (![X1]:top=join(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', def_top)).
fof(c_0_8, axiom, (![X1]:![X2]:X1=join(complement(join(complement(X1),complement(X2))),complement(join(complement(X1),X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', maddux3_a_kind_of_de_Morgan)).
fof(c_0_9, axiom, (![X1]:![X2]:![X3]:join(X1,join(X2,X3))=join(join(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', maddux2_join_associativity)).
fof(c_0_10, conjecture, (![X1]:![X2]:![X3]:((join(X1,one)=one&join(X2,one)=one)=>(join(meet(composition(X1,X3),composition(X2,X3)),composition(meet(X1,X2),X3))=composition(meet(X1,X2),X3)&join(composition(meet(X1,X2),X3),meet(composition(X1,X3),composition(X2,X3)))=meet(composition(X1,X3),composition(X2,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
fof(c_0_11, axiom, (![X1]:![X2]:converse(join(X1,X2))=join(converse(X1),converse(X2))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', converse_additivity)).
fof(c_0_12, axiom, (![X1]:![X2]:![X3]:composition(join(X1,X2),X3)=join(composition(X1,X3),composition(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', composition_distributivity)).
fof(c_0_13, axiom, (![X1]:![X2]:![X3]:join(meet(composition(X1,X2),X3),meet(composition(meet(X1,composition(X3,converse(X2))),X2),X3))=meet(composition(meet(X1,composition(X3,converse(X2))),X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+1.ax', modular_law_2)).
fof(c_0_14, axiom, (![X1]:![X2]:![X3]:join(meet(composition(X1,X2),X3),composition(meet(X1,composition(X3,converse(X2))),meet(X2,composition(converse(X1),X3))))=composition(meet(X1,composition(X3,converse(X2))),meet(X2,composition(converse(X1),X3)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+1.ax', dedekind_law)).
fof(c_0_15, axiom, (![X1]:![X2]:![X3]:join(meet(composition(X1,X2),X3),meet(composition(X1,meet(X2,composition(converse(X1),X3))),X3))=meet(composition(X1,meet(X2,composition(converse(X1),X3))),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+1.ax', modular_law_1)).
fof(c_0_16, axiom, (![X1]:![X2]:![X3]:composition(X1,composition(X2,X3))=composition(composition(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', composition_associativity)).
fof(c_0_17, plain, (![X3]:![X4]:converse(composition(X3,X4))=composition(converse(X4),converse(X3))), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_18, plain, (![X2]:composition(X2,one)=X2), inference(variable_rename,[status(thm)],[c_0_1])).
cnf(c_0_19,plain,(converse(composition(X1,X2))=composition(converse(X2),converse(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,plain,(composition(X1,one)=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_21, plain, (![X2]:converse(converse(X2))=X2), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_22,plain,(composition(converse(one),converse(X1))=converse(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_23,plain,(converse(converse(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_24, plain, (![X3]:![X4]:join(composition(converse(X3),complement(composition(X3,X4))),complement(X4))=complement(X4)), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_25, plain, (![X3]:![X4]:join(X3,X4)=join(X4,X3)), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_26,plain,(composition(converse(one),X1)=X1), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_27, plain, (![X2]:zero=meet(X2,complement(X2))), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_28, plain, (![X3]:![X4]:meet(X3,X4)=complement(join(complement(X3),complement(X4)))), inference(variable_rename,[status(thm)],[c_0_6])).
cnf(c_0_29,plain,(join(composition(converse(X1),complement(composition(X1,X2))),complement(X2))=complement(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30,plain,(join(X1,X2)=join(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31,plain,(converse(one)=one), inference(spm,[status(thm)],[c_0_20, c_0_26])).
cnf(c_0_32,plain,(zero=meet(X1,complement(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33,plain,(meet(X1,X2)=complement(join(complement(X1),complement(X2)))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_34, plain, (![X2]:top=join(X2,complement(X2))), inference(variable_rename,[status(thm)],[c_0_7])).
cnf(c_0_35,plain,(join(complement(X1),composition(converse(X2),complement(composition(X2,X1))))=complement(X1)), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36,plain,(composition(one,X1)=X1), inference(rw,[status(thm)],[c_0_26, c_0_31])).
cnf(c_0_37,plain,(complement(join(complement(X1),complement(complement(X1))))=zero), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38,plain,(top=join(X1,complement(X1))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_39, plain, (![X3]:![X4]:X3=join(complement(join(complement(X3),complement(X4))),complement(join(complement(X3),X4)))), inference(variable_rename,[status(thm)],[c_0_8])).
fof(c_0_40, plain, (![X4]:![X5]:![X6]:join(X4,join(X5,X6))=join(join(X4,X5),X6)), inference(variable_rename,[status(thm)],[c_0_9])).
cnf(c_0_41,plain,(join(complement(X1),complement(X1))=complement(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_31]), c_0_36])).
cnf(c_0_42,plain,(complement(top)=zero), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43,plain,(X1=join(complement(join(complement(X1),complement(X2))),complement(join(complement(X1),X2)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44,plain,(join(X1,join(X2,X3))=join(join(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45,plain,(join(zero,zero)=zero), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46,plain,(join(complement(join(complement(X1),X2)),complement(join(complement(X1),complement(X2))))=X1), inference(rw,[status(thm)],[c_0_43, c_0_30])).
cnf(c_0_47,plain,(join(zero,join(zero,X1))=join(zero,X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_48,plain,(join(zero,complement(complement(X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_41]), c_0_38]), c_0_42]), c_0_30])).
fof(c_0_49, negated_conjecture, (~(![X1]:![X2]:![X3]:((join(X1,one)=one&join(X2,one)=one)=>(join(meet(composition(X1,X3),composition(X2,X3)),composition(meet(X1,X2),X3))=composition(meet(X1,X2),X3)&join(composition(meet(X1,X2),X3),meet(composition(X1,X3),composition(X2,X3)))=meet(composition(X1,X3),composition(X2,X3)))))), inference(assume_negation,[status(cth)],[c_0_10])).
cnf(c_0_50,plain,(join(zero,X1)=X1), inference(spm,[status(thm)],[c_0_47, c_0_48])).
fof(c_0_51, negated_conjecture, (((join(esk1_0,one)=one&join(esk2_0,one)=one)&(join(meet(composition(esk1_0,esk3_0),composition(esk2_0,esk3_0)),composition(meet(esk1_0,esk2_0),esk3_0))!=composition(meet(esk1_0,esk2_0),esk3_0)|join(composition(meet(esk1_0,esk2_0),esk4_0),meet(composition(esk1_0,esk4_0),composition(esk2_0,esk4_0)))!=meet(composition(esk1_0,esk4_0),composition(esk2_0,esk4_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_49])])])])).
cnf(c_0_52,plain,(complement(complement(X1))=X1), inference(rw,[status(thm)],[c_0_48, c_0_50])).
cnf(c_0_53,negated_conjecture,(join(esk2_0,one)=one), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_54,plain,(join(X1,X1)=X1), inference(spm,[status(thm)],[c_0_41, c_0_52])).
cnf(c_0_55,negated_conjecture,(join(one,esk2_0)=one), inference(rw,[status(thm)],[c_0_53, c_0_30])).
cnf(c_0_56,plain,(join(top,zero)=top), inference(spm,[status(thm)],[c_0_38, c_0_42])).
cnf(c_0_57,plain,(join(X1,join(complement(X1),X2))=join(top,X2)), inference(spm,[status(thm)],[c_0_44, c_0_38])).
fof(c_0_58, plain, (![X3]:![X4]:converse(join(X3,X4))=join(converse(X3),converse(X4))), inference(variable_rename,[status(thm)],[c_0_11])).
cnf(c_0_59,plain,(join(X1,join(X1,X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_44, c_0_54])).
cnf(c_0_60,negated_conjecture,(join(one,join(esk2_0,X1))=join(one,X1)), inference(spm,[status(thm)],[c_0_44, c_0_55])).
cnf(c_0_61,plain,(join(top,join(zero,X1))=join(top,X1)), inference(spm,[status(thm)],[c_0_44, c_0_56])).
cnf(c_0_62,plain,(join(top,complement(X1))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_41]), c_0_38])).
cnf(c_0_63,plain,(converse(join(X1,X2))=join(converse(X1),converse(X2))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_64,plain,(join(X1,complement(join(complement(X1),X2)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_46]), c_0_30])).
cnf(c_0_65,plain,(join(X1,join(X2,X1))=join(X2,X1)), inference(spm,[status(thm)],[c_0_59, c_0_30])).
fof(c_0_66, plain, (![X4]:![X5]:![X6]:composition(join(X4,X5),X6)=join(composition(X4,X6),composition(X5,X6))), inference(variable_rename,[status(thm)],[c_0_12])).
cnf(c_0_67,negated_conjecture,(join(one,complement(esk2_0))=join(top,one)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_38]), c_0_30])).
cnf(c_0_68,plain,(join(top,X1)=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_48]), c_0_62])).
cnf(c_0_69,negated_conjecture,(join(converse(one),converse(esk2_0))=converse(one)), inference(spm,[status(thm)],[c_0_63, c_0_55])).
cnf(c_0_70,plain,(join(X1,complement(join(X2,complement(X1))))=X1), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_71,plain,(composition(join(X1,X2),X3)=join(composition(X1,X3),composition(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72,negated_conjecture,(join(one,complement(esk2_0))=top), inference(rw,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_73,plain,(join(converse(top),converse(X1))=converse(top)), inference(spm,[status(thm)],[c_0_63, c_0_68])).
cnf(c_0_74,negated_conjecture,(join(one,converse(esk2_0))=one), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_31]), c_0_31])).
cnf(c_0_75,plain,(join(X1,join(X2,X3))=join(X2,join(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_30]), c_0_44])).
cnf(c_0_76,plain,(join(complement(X1),complement(join(X2,X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_70, c_0_52])).
cnf(c_0_77,negated_conjecture,(join(X1,composition(complement(esk2_0),X1))=composition(top,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_36])).
cnf(c_0_78,plain,(join(X1,top)=top), inference(spm,[status(thm)],[c_0_30, c_0_68])).
cnf(c_0_79,plain,(join(converse(top),X1)=converse(top)), inference(spm,[status(thm)],[c_0_73, c_0_23])).
cnf(c_0_80,negated_conjecture,(join(esk1_0,one)=one), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_81,negated_conjecture,(join(X1,composition(converse(esk2_0),X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_74]), c_0_36]), c_0_36])).
cnf(c_0_82,plain,(join(complement(X1),complement(join(complement(X1),complement(X2))))=join(complement(X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_46]), c_0_52]), c_0_75]), c_0_30]), c_0_76])).
cnf(c_0_83,negated_conjecture,(composition(top,top)=top), inference(spm,[status(thm)],[c_0_68, c_0_77])).
cnf(c_0_84,plain,(converse(top)=top), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_85,negated_conjecture,(join(one,esk1_0)=one), inference(rw,[status(thm)],[c_0_80, c_0_30])).
fof(c_0_86, plain, (![X4]:![X5]:![X6]:join(meet(composition(X4,X5),X6),meet(composition(meet(X4,composition(X6,converse(X5))),X5),X6))=meet(composition(meet(X4,composition(X6,converse(X5))),X5),X6)), inference(variable_rename,[status(thm)],[c_0_13])).
cnf(c_0_87,negated_conjecture,(join(converse(X1),composition(converse(X1),esk2_0))=converse(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_81]), c_0_19]), c_0_23])).
fof(c_0_88, plain, (![X4]:![X5]:![X6]:join(meet(composition(X4,X5),X6),composition(meet(X4,composition(X6,converse(X5))),meet(X5,composition(converse(X4),X6))))=composition(meet(X4,composition(X6,converse(X5))),meet(X5,composition(converse(X4),X6)))), inference(variable_rename,[status(thm)],[c_0_14])).
cnf(c_0_89,plain,(join(X1,complement(join(X1,complement(X2))))=join(X1,X2)), inference(spm,[status(thm)],[c_0_82, c_0_52])).
cnf(c_0_90,plain,(join(complement(X1),join(composition(converse(X2),complement(join(composition(X2,X1),composition(X3,X1)))),composition(converse(X3),complement(join(composition(X2,X1),composition(X3,X1))))))=complement(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_71]), c_0_63]), c_0_71])).
cnf(c_0_91,negated_conjecture,(composition(top,zero)=zero), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_83]), c_0_42]), c_0_84]), c_0_42]), c_0_50]), c_0_42])).
cnf(c_0_92,plain,(join(X1,zero)=X1), inference(spm,[status(thm)],[c_0_30, c_0_50])).
cnf(c_0_93,negated_conjecture,(join(converse(one),converse(esk1_0))=converse(one)), inference(spm,[status(thm)],[c_0_63, c_0_85])).
cnf(c_0_94,plain,(join(meet(composition(X1,X2),X3),meet(composition(meet(X1,composition(X3,converse(X2))),X2),X3))=meet(composition(meet(X1,composition(X3,converse(X2))),X2),X3)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_95,negated_conjecture,(join(X1,composition(X1,esk2_0))=X1), inference(spm,[status(thm)],[c_0_87, c_0_23])).
cnf(c_0_96,plain,(join(composition(X1,X2),composition(complement(X1),X2))=composition(top,X2)), inference(spm,[status(thm)],[c_0_71, c_0_38])).
cnf(c_0_97,plain,(join(meet(composition(X1,X2),X3),composition(meet(X1,composition(X3,converse(X2))),meet(X2,composition(converse(X1),X3))))=composition(meet(X1,composition(X3,converse(X2))),meet(X2,composition(converse(X1),X3)))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_98,plain,(join(X1,complement(join(X1,X2)))=join(X1,complement(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_89]), c_0_59])).
cnf(c_0_99,negated_conjecture,(composition(converse(X1),zero)=zero), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_83]), c_0_42]), c_0_78]), c_0_42]), c_0_84]), c_0_78]), c_0_42]), c_0_91]), c_0_92]), c_0_50]), c_0_42])).
cnf(c_0_100,negated_conjecture,(join(one,converse(esk1_0))=one), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_31]), c_0_31])).
cnf(c_0_101,plain,(join(complement(join(complement(composition(X1,X2)),complement(X3))),complement(join(complement(composition(complement(join(complement(X1),complement(composition(X3,converse(X2))))),X2)),complement(X3))))=complement(join(complement(composition(complement(join(complement(X1),complement(composition(X3,converse(X2))))),X2)),complement(X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_102,negated_conjecture,(join(X1,join(X2,composition(X1,esk2_0)))=join(X2,X1)), inference(spm,[status(thm)],[c_0_75, c_0_95])).
cnf(c_0_103,plain,(join(X1,composition(complement(one),X1))=composition(top,X1)), inference(spm,[status(thm)],[c_0_96, c_0_36])).
cnf(c_0_104,plain,(join(complement(join(complement(composition(X1,X2)),complement(X3))),composition(complement(join(complement(X1),complement(composition(X3,converse(X2))))),complement(join(complement(X2),complement(composition(converse(X1),X3))))))=composition(complement(join(complement(X1),complement(composition(X3,converse(X2))))),complement(join(complement(X2),complement(composition(converse(X1),X3)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_105,negated_conjecture,(join(X1,complement(composition(converse(esk2_0),X1)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_81]), c_0_38])).
cnf(c_0_106,negated_conjecture,(composition(X1,zero)=zero), inference(spm,[status(thm)],[c_0_99, c_0_23])).
cnf(c_0_107,negated_conjecture,(join(X1,composition(converse(esk1_0),X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_100]), c_0_36]), c_0_36])).
cnf(c_0_108,negated_conjecture,(join(one,join(converse(esk2_0),X1))=join(one,X1)), inference(spm,[status(thm)],[c_0_44, c_0_74])).
cnf(c_0_109,plain,(join(complement(join(complement(composition(X1,X2)),complement(X3))),complement(join(complement(X3),complement(composition(complement(join(complement(X1),complement(composition(X3,converse(X2))))),X2)))))=complement(join(complement(X3),complement(composition(complement(join(complement(X1),complement(composition(X3,converse(X2))))),X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_30]), c_0_30])).
cnf(c_0_110,negated_conjecture,(join(complement(one),composition(top,esk2_0))=join(esk2_0,complement(one))), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_111,plain,(join(X1,complement(join(X2,X1)))=join(X1,complement(X2))), inference(spm,[status(thm)],[c_0_98, c_0_30])).
cnf(c_0_112,negated_conjecture,(join(complement(one),complement(esk2_0))=complement(esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_55]), c_0_30])).
cnf(c_0_113,negated_conjecture,(complement(join(X1,complement(composition(esk2_0,X1))))=zero), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_52]), c_0_42]), c_0_106]), c_0_92]), c_0_42]), c_0_106]), c_0_30])).
cnf(c_0_114,negated_conjecture,(join(converse(X1),composition(converse(X1),esk1_0))=converse(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_107]), c_0_19]), c_0_23])).
cnf(c_0_115,plain,(join(converse(X1),converse(complement(X1)))=converse(top)), inference(spm,[status(thm)],[c_0_63, c_0_38])).
fof(c_0_116, plain, (![X4]:![X5]:![X6]:join(meet(composition(X4,X5),X6),meet(composition(X4,meet(X5,composition(converse(X4),X6))),X6))=meet(composition(X4,meet(X5,composition(converse(X4),X6))),X6)), inference(variable_rename,[status(thm)],[c_0_15])).
cnf(c_0_117,negated_conjecture,(join(one,composition(converse(esk1_0),converse(esk2_0)))=one), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_107]), c_0_74])).
fof(c_0_118, plain, (![X4]:![X5]:![X6]:composition(X4,composition(X5,X6))=composition(composition(X4,X5),X6)), inference(variable_rename,[status(thm)],[c_0_16])).
cnf(c_0_119,plain,(join(complement(join(complement(composition(X1,converse(X2))),complement(X3))),complement(join(complement(X3),complement(composition(complement(join(complement(X1),complement(composition(X3,X2)))),converse(X2))))))=complement(join(complement(X3),complement(composition(complement(join(complement(X1),complement(composition(X3,X2)))),converse(X2)))))), inference(spm,[status(thm)],[c_0_109, c_0_23])).
cnf(c_0_120,negated_conjecture,(join(complement(one),complement(composition(top,esk2_0)))=complement(esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_110]), c_0_111]), c_0_112])).
cnf(c_0_121,negated_conjecture,(join(X1,composition(esk2_0,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_113]), c_0_92]), c_0_52])).
cnf(c_0_122,negated_conjecture,(join(X1,composition(X1,esk1_0))=X1), inference(spm,[status(thm)],[c_0_114, c_0_23])).
cnf(c_0_123,plain,(join(converse(X1),converse(complement(X1)))=top), inference(rw,[status(thm)],[c_0_115, c_0_84])).
cnf(c_0_124,plain,(join(meet(composition(X1,X2),X3),meet(composition(X1,meet(X2,composition(converse(X1),X3))),X3))=meet(composition(X1,meet(X2,composition(converse(X1),X3))),X3)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_125,negated_conjecture,(join(one,composition(esk2_0,esk1_0))=one), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_117]), c_0_31]), c_0_31]), c_0_19]), c_0_23]), c_0_23])).
cnf(c_0_126,plain,(composition(X1,composition(X2,X3))=composition(composition(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_127,negated_conjecture,(join(X1,complement(composition(converse(esk1_0),X1)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_107]), c_0_38])).
cnf(c_0_128,negated_conjecture,(composition(esk2_0,converse(esk2_0))=converse(esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_36]), c_0_42]), c_0_92]), c_0_52]), c_0_42]), c_0_52]), c_0_50]), c_0_52]), c_0_121]), c_0_42]), c_0_52]), c_0_50]), c_0_52])).
cnf(c_0_129,negated_conjecture,(join(X1,join(X2,composition(X1,esk1_0)))=join(X2,X1)), inference(spm,[status(thm)],[c_0_75, c_0_122])).
cnf(c_0_130,plain,(join(converse(X1),complement(converse(complement(X1))))=converse(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_123]), c_0_42]), c_0_92])).
cnf(c_0_131,plain,(join(complement(join(complement(composition(X1,X2)),complement(X3))),complement(join(complement(composition(X1,complement(join(complement(X2),complement(composition(converse(X1),X3)))))),complement(X3))))=complement(join(complement(composition(X1,complement(join(complement(X2),complement(composition(converse(X1),X3)))))),complement(X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_132,negated_conjecture,(join(X1,composition(esk2_0,composition(esk1_0,X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_125]), c_0_36]), c_0_36]), c_0_126])).
cnf(c_0_133,negated_conjecture,(complement(join(X1,complement(composition(esk1_0,X1))))=zero), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_127]), c_0_52]), c_0_42]), c_0_106]), c_0_92]), c_0_42]), c_0_106]), c_0_30])).
cnf(c_0_134,negated_conjecture,(converse(esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_128]), c_0_23]), c_0_23]), c_0_128])).
cnf(c_0_135,negated_conjecture,(join(complement(one),composition(top,esk1_0))=join(esk1_0,complement(one))), inference(spm,[status(thm)],[c_0_129, c_0_103])).
cnf(c_0_136,negated_conjecture,(join(complement(one),complement(esk1_0))=complement(esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_85]), c_0_30])).
cnf(c_0_137,negated_conjecture,(join(X1,join(composition(X1,esk2_0),X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_44, c_0_95])).
cnf(c_0_138,plain,(join(X1,converse(complement(converse(complement(X1)))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_130]), c_0_23]), c_0_23])).
cnf(c_0_139,plain,(join(X1,converse(complement(converse(X1))))=top), inference(spm,[status(thm)],[c_0_123, c_0_23])).
cnf(c_0_140,plain,(join(complement(join(complement(composition(X1,X2)),complement(X3))),complement(join(complement(X3),complement(composition(X1,complement(join(complement(X2),complement(composition(converse(X1),X3)))))))))=complement(join(complement(X3),complement(composition(X1,complement(join(complement(X2),complement(composition(converse(X1),X3))))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_30]), c_0_30])).
cnf(c_0_141,negated_conjecture,(join(one,composition(esk2_0,composition(esk1_0,esk2_0)))=one), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_132]), c_0_55])).
cnf(c_0_142,negated_conjecture,(join(X1,composition(esk1_0,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_133]), c_0_92]), c_0_52])).
cnf(c_0_143,negated_conjecture,(composition(esk2_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_134]), c_0_134])).
cnf(c_0_144,negated_conjecture,(join(complement(one),complement(composition(top,esk1_0)))=complement(esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_135]), c_0_111]), c_0_136])).
cnf(c_0_145,negated_conjecture,(join(X1,composition(complement(X1),esk2_0))=join(X1,composition(top,esk2_0))), inference(spm,[status(thm)],[c_0_137, c_0_96])).
cnf(c_0_146,plain,(join(complement(X1),converse(complement(converse(X1))))=complement(X1)), inference(spm,[status(thm)],[c_0_138, c_0_52])).
cnf(c_0_147,plain,(join(X1,complement(converse(complement(converse(X1)))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_139]), c_0_42]), c_0_92])).
cnf(c_0_148,negated_conjecture,(join(X1,join(composition(X1,esk1_0),X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_44, c_0_122])).
cnf(c_0_149,plain,(join(complement(join(complement(composition(converse(X1),X2)),complement(X3))),complement(join(complement(X3),complement(composition(converse(X1),complement(join(complement(X2),complement(composition(X1,X3)))))))))=complement(join(complement(X3),complement(composition(converse(X1),complement(join(complement(X2),complement(composition(X1,X3))))))))), inference(spm,[status(thm)],[c_0_140, c_0_23])).
cnf(c_0_150,negated_conjecture,(join(complement(one),complement(composition(esk2_0,composition(esk1_0,esk2_0))))=complement(composition(esk2_0,composition(esk1_0,esk2_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_141]), c_0_30])).
cnf(c_0_151,negated_conjecture,(join(complement(X1),complement(composition(esk1_0,X1)))=complement(composition(esk1_0,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_142]), c_0_30])).
cnf(c_0_152,negated_conjecture,(composition(esk2_0,composition(esk2_0,X1))=composition(esk2_0,X1)), inference(spm,[status(thm)],[c_0_126, c_0_143])).
cnf(c_0_153,negated_conjecture,(join(complement(X1),complement(composition(esk2_0,X1)))=complement(composition(esk2_0,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_121]), c_0_30])).
cnf(c_0_154,negated_conjecture,(composition(esk1_0,converse(esk1_0))=converse(esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_144]), c_0_36]), c_0_42]), c_0_92]), c_0_52]), c_0_42]), c_0_52]), c_0_50]), c_0_52]), c_0_142]), c_0_42]), c_0_52]), c_0_50]), c_0_52])).
cnf(c_0_155,plain,(join(converse(zero),converse(X1))=converse(X1)), inference(spm,[status(thm)],[c_0_63, c_0_50])).
cnf(c_0_156,negated_conjecture,(join(complement(X1),complement(composition(X1,esk2_0)))=complement(composition(X1,esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_95]), c_0_30])).
cnf(c_0_157,negated_conjecture,(join(X1,complement(composition(complement(X1),esk2_0)))=join(X1,complement(composition(top,esk2_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_145]), c_0_98])).
cnf(c_0_158,plain,(complement(converse(complement(converse(X1))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_146]), c_0_52]), c_0_30]), c_0_147])).
cnf(c_0_159,negated_conjecture,(join(X1,composition(complement(X1),esk1_0))=join(X1,composition(top,esk1_0))), inference(spm,[status(thm)],[c_0_148, c_0_96])).
cnf(c_0_160,negated_conjecture,(composition(esk2_0,composition(esk1_0,esk2_0))=composition(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_20]), c_0_134]), c_0_151]), c_0_52]), c_0_134]), c_0_52]), c_0_152]), c_0_153]), c_0_52]), c_0_121]), c_0_134]), c_0_52]), c_0_152]), c_0_153]), c_0_52])).
cnf(c_0_161,negated_conjecture,(converse(esk1_0)=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_154]), c_0_23]), c_0_23]), c_0_154])).
cnf(c_0_162,plain,(join(converse(zero),X1)=X1), inference(spm,[status(thm)],[c_0_155, c_0_23])).
cnf(c_0_163,negated_conjecture,(complement(composition(complement(X1),esk2_0))=join(X1,complement(composition(top,esk2_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_52]), c_0_157])).
cnf(c_0_164,plain,(converse(complement(converse(X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_52, c_0_158])).
cnf(c_0_165,negated_conjecture,(join(complement(X1),complement(composition(X1,esk1_0)))=complement(composition(X1,esk1_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_122]), c_0_30])).
cnf(c_0_166,negated_conjecture,(join(X1,complement(composition(complement(X1),esk1_0)))=join(X1,complement(composition(top,esk1_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_159]), c_0_98])).
cnf(c_0_167,negated_conjecture,(composition(esk2_0,esk1_0)=composition(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_160]), c_0_19]), c_0_134]), c_0_161]), c_0_19]), c_0_134]), c_0_161]), c_0_134]), c_0_126]), c_0_160])).
cnf(c_0_168,plain,(join(converse(complement(X1)),composition(converse(complement(composition(X2,X1))),X2))=converse(complement(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_35]), c_0_19]), c_0_23])).
cnf(c_0_169,plain,(converse(zero)=zero), inference(spm,[status(thm)],[c_0_92, c_0_162])).
cnf(c_0_170,negated_conjecture,(composition(complement(X1),esk2_0)=complement(join(X1,complement(composition(top,esk2_0))))), inference(spm,[status(thm)],[c_0_52, c_0_163])).
cnf(c_0_171,plain,(converse(complement(X1))=complement(converse(X1))), inference(spm,[status(thm)],[c_0_23, c_0_164])).
cnf(c_0_172,negated_conjecture,(complement(composition(complement(X1),esk1_0))=join(X1,complement(composition(top,esk1_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_52]), c_0_166])).
cnf(c_0_173,negated_conjecture,(join(composition(esk1_0,esk2_0),composition(complement(esk2_0),esk1_0))=composition(top,esk1_0)), inference(spm,[status(thm)],[c_0_96, c_0_167])).
cnf(c_0_174,negated_conjecture,(join(complement(esk2_0),composition(top,esk1_0))=join(esk1_0,complement(esk2_0))), inference(spm,[status(thm)],[c_0_129, c_0_77])).
cnf(c_0_175,plain,(composition(converse(complement(composition(X1,top))),X1)=zero), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_42]), c_0_169]), c_0_50]), c_0_169])).
cnf(c_0_176,negated_conjecture,(composition(top,composition(top,X1))=composition(top,X1)), inference(spm,[status(thm)],[c_0_126, c_0_83])).
cnf(c_0_177,negated_conjecture,(join(X1,composition(top,X1))=composition(top,X1)), inference(spm,[status(thm)],[c_0_59, c_0_77])).
cnf(c_0_178,negated_conjecture,(composition(esk2_0,complement(converse(X1)))=complement(join(complement(composition(esk2_0,top)),converse(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_170]), c_0_171]), c_0_63]), c_0_171]), c_0_19]), c_0_134]), c_0_84]), c_0_134]), c_0_171]), c_0_30])).
cnf(c_0_179,negated_conjecture,(composition(complement(X1),esk1_0)=complement(join(X1,complement(composition(top,esk1_0))))), inference(spm,[status(thm)],[c_0_52, c_0_172])).
cnf(c_0_180,negated_conjecture,(join(complement(esk2_0),composition(esk1_0,esk2_0))=join(esk1_0,complement(esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_173]), c_0_174]), c_0_30])).
cnf(c_0_181,plain,(composition(complement(composition(top,converse(X1))),X1)=zero), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_175, c_0_171]), c_0_19]), c_0_84])).
cnf(c_0_182,negated_conjecture,(composition(top,complement(composition(top,X1)))=complement(composition(top,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_176]), c_0_84]), c_0_177])).
cnf(c_0_183,negated_conjecture,(composition(esk2_0,complement(X1))=complement(join(complement(composition(esk2_0,top)),X1))), inference(spm,[status(thm)],[c_0_178, c_0_23])).
cnf(c_0_184,negated_conjecture,(composition(esk1_0,complement(converse(X1)))=complement(join(complement(composition(esk1_0,top)),converse(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_179]), c_0_171]), c_0_63]), c_0_171]), c_0_19]), c_0_161]), c_0_84]), c_0_161]), c_0_171]), c_0_30])).
cnf(c_0_185,negated_conjecture,(join(composition(meet(esk1_0,esk2_0),esk4_0),meet(composition(esk1_0,esk4_0),composition(esk2_0,esk4_0)))!=meet(composition(esk1_0,esk4_0),composition(esk2_0,esk4_0))|join(meet(composition(esk1_0,esk3_0),composition(esk2_0,esk3_0)),composition(meet(esk1_0,esk2_0),esk3_0))!=composition(meet(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_186,negated_conjecture,(complement(composition(esk1_0,esk2_0))=join(complement(esk1_0),complement(esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_180]), c_0_111]), c_0_30]), c_0_151])).
cnf(c_0_187,negated_conjecture,(composition(complement(composition(top,esk2_0)),esk2_0)=zero), inference(spm,[status(thm)],[c_0_181, c_0_134])).
cnf(c_0_188,negated_conjecture,(composition(zero,X1)=zero), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_99]), c_0_169]), c_0_169]), c_0_23])).
cnf(c_0_189,negated_conjecture,(composition(complement(composition(converse(X1),top)),top)=complement(composition(converse(X1),top))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_182]), c_0_84]), c_0_171]), c_0_19]), c_0_84]), c_0_171]), c_0_19]), c_0_84])).
cnf(c_0_190,negated_conjecture,(join(complement(X1),complement(join(complement(composition(esk2_0,top)),composition(esk2_0,X1))))=complement(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_134]), c_0_183])).
cnf(c_0_191,negated_conjecture,(composition(esk1_0,complement(X1))=complement(join(complement(composition(esk1_0,top)),X1))), inference(spm,[status(thm)],[c_0_184, c_0_23])).
cnf(c_0_192,negated_conjecture,(join(complement(join(complement(composition(esk1_0,esk3_0)),complement(composition(esk2_0,esk3_0)))),composition(complement(join(complement(esk1_0),complement(esk2_0))),esk3_0))!=composition(complement(join(complement(esk1_0),complement(esk2_0))),esk3_0)|join(composition(complement(join(complement(esk1_0),complement(esk2_0))),esk4_0),complement(join(complement(composition(esk1_0,esk4_0)),complement(composition(esk2_0,esk4_0)))))!=complement(join(complement(composition(esk1_0,esk4_0)),complement(composition(esk2_0,esk4_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_185, c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_193,negated_conjecture,(composition(esk1_0,esk2_0)=complement(join(complement(esk1_0),complement(esk2_0)))), inference(spm,[status(thm)],[c_0_52, c_0_186])).
cnf(c_0_194,negated_conjecture,(composition(complement(composition(top,esk2_0)),composition(esk2_0,X1))=zero), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_187]), c_0_188])).
cnf(c_0_195,plain,(complement(zero)=top), inference(spm,[status(thm)],[c_0_38, c_0_50])).
cnf(c_0_196,negated_conjecture,(composition(complement(composition(X1,top)),top)=complement(composition(X1,top))), inference(spm,[status(thm)],[c_0_189, c_0_23])).
cnf(c_0_197,negated_conjecture,(join(complement(composition(esk2_0,top)),composition(esk2_0,X1))=join(complement(composition(esk2_0,top)),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_190]), c_0_52]), c_0_52]), c_0_44]), c_0_30]), c_0_121]), c_0_52])).
cnf(c_0_198,negated_conjecture,(composition(complement(composition(top,esk1_0)),esk1_0)=zero), inference(spm,[status(thm)],[c_0_181, c_0_161])).
cnf(c_0_199,negated_conjecture,(join(complement(X1),complement(join(complement(composition(esk1_0,top)),composition(esk1_0,X1))))=complement(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_161]), c_0_191])).
cnf(c_0_200,negated_conjecture,(join(complement(join(complement(composition(esk1_0,esk4_0)),complement(composition(esk2_0,esk4_0)))),composition(complement(join(complement(esk1_0),complement(esk2_0))),esk4_0))!=complement(join(complement(composition(esk1_0,esk4_0)),complement(composition(esk2_0,esk4_0))))|join(complement(join(complement(composition(esk1_0,esk3_0)),complement(composition(esk2_0,esk3_0)))),composition(complement(join(complement(esk1_0),complement(esk2_0))),esk3_0))!=composition(complement(join(complement(esk1_0),complement(esk2_0))),esk3_0)), inference(rw,[status(thm)],[c_0_192, c_0_30])).
cnf(c_0_201,negated_conjecture,(composition(complement(join(complement(esk1_0),complement(esk2_0))),X1)=composition(esk1_0,composition(esk2_0,X1))), inference(spm,[status(thm)],[c_0_126, c_0_193])).
cnf(c_0_202,negated_conjecture,(join(complement(composition(esk2_0,X1)),complement(composition(esk2_0,top)))=complement(composition(esk2_0,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_194]), c_0_171]), c_0_19]), c_0_134]), c_0_84]), c_0_195]), c_0_196])).
cnf(c_0_203,negated_conjecture,(join(complement(composition(esk2_0,top)),complement(composition(esk2_0,X1)))=join(complement(composition(esk2_0,top)),complement(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_197]), c_0_98])).
cnf(c_0_204,negated_conjecture,(composition(complement(composition(top,esk1_0)),composition(esk1_0,X1))=zero), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_198]), c_0_188])).
cnf(c_0_205,negated_conjecture,(join(complement(composition(esk1_0,top)),composition(esk1_0,X1))=join(complement(composition(esk1_0,top)),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_199]), c_0_52]), c_0_52]), c_0_44]), c_0_30]), c_0_142]), c_0_52])).
cnf(c_0_206,negated_conjecture,(join(complement(join(complement(composition(esk1_0,esk4_0)),complement(composition(esk2_0,esk4_0)))),composition(esk1_0,composition(esk2_0,esk4_0)))!=complement(join(complement(composition(esk1_0,esk4_0)),complement(composition(esk2_0,esk4_0))))|join(complement(join(complement(composition(esk1_0,esk3_0)),complement(composition(esk2_0,esk3_0)))),composition(esk1_0,composition(esk2_0,esk3_0)))!=composition(esk1_0,composition(esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_200, c_0_201]), c_0_201]), c_0_201])).
cnf(c_0_207,negated_conjecture,(join(complement(composition(esk2_0,top)),complement(X1))=complement(composition(esk2_0,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_202]), c_0_203])).
cnf(c_0_208,negated_conjecture,(join(complement(composition(esk1_0,X1)),complement(composition(esk1_0,top)))=complement(composition(esk1_0,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_204]), c_0_171]), c_0_19]), c_0_161]), c_0_84]), c_0_195]), c_0_196])).
cnf(c_0_209,negated_conjecture,(join(complement(composition(esk1_0,top)),complement(composition(esk1_0,X1)))=join(complement(composition(esk1_0,top)),complement(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_205]), c_0_98])).
cnf(c_0_210,negated_conjecture,(join(complement(join(complement(esk4_0),join(complement(composition(esk1_0,esk4_0)),complement(composition(esk2_0,top))))),composition(esk1_0,composition(esk2_0,esk4_0)))!=complement(join(complement(esk4_0),join(complement(composition(esk1_0,esk4_0)),complement(composition(esk2_0,top)))))|join(complement(join(complement(esk3_0),join(complement(composition(esk1_0,esk3_0)),complement(composition(esk2_0,top))))),composition(esk1_0,composition(esk2_0,esk3_0)))!=composition(esk1_0,composition(esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_206, c_0_207]), c_0_30]), c_0_75]), c_0_207]), c_0_30]), c_0_75]), c_0_207]), c_0_30]), c_0_75])).
cnf(c_0_211,negated_conjecture,(join(complement(composition(esk1_0,top)),complement(X1))=complement(composition(esk1_0,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_208]), c_0_209])).
cnf(c_0_212,negated_conjecture,(join(complement(join(complement(esk4_0),join(complement(composition(esk1_0,top)),complement(composition(esk2_0,top))))),composition(esk1_0,composition(esk2_0,esk4_0)))!=complement(join(complement(esk4_0),join(complement(composition(esk1_0,top)),complement(composition(esk2_0,top)))))|join(complement(join(complement(esk3_0),join(complement(composition(esk1_0,top)),complement(composition(esk2_0,top))))),composition(esk1_0,composition(esk2_0,esk3_0)))!=composition(esk1_0,composition(esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_210, c_0_211]), c_0_30]), c_0_44]), c_0_59]), c_0_211]), c_0_30]), c_0_44]), c_0_59]), c_0_211]), c_0_30]), c_0_44]), c_0_59])).
cnf(c_0_213,negated_conjecture,(composition(esk2_0,X1)=complement(join(complement(composition(esk2_0,top)),complement(X1)))), inference(spm,[status(thm)],[c_0_183, c_0_52])).
cnf(c_0_214,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_212, c_0_213]), c_0_30]), c_0_191]), c_0_75]), c_0_54]), c_0_213]), c_0_30]), c_0_191]), c_0_75]), c_0_54]), c_0_213]), c_0_30]), c_0_191]), c_0_75])]), ['proof']).
# SZS output end CNFRefutation.
EOF
