% Problem    : SWW318+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n186.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:18:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X18]:![X5]:(hAPP(c_Set_OCollect(X5),X18)=c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))<=>![X3]:~(hBOOL(hAPP(X18,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__empty__eq)).
fof(c_0_1, axiom, (![X5]:c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),hAPP(c_COMBK(tc_HOL_Obool,X5),c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__def)).
fof(c_0_2, axiom, (![X18]:![X5]:hAPP(c_Set_OCollect(X5),X18)=X18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__def)).
fof(c_0_3, axiom, (![X18]:![X5]:hAPP(c_Set_OCollect(X5),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X5),c_fNot),X18))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),hAPP(c_Set_OCollect(X5),X18))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__neg__eq)).
fof(c_0_4, axiom, (![X314]:![X49]:![X5]:![X58]:hAPP(hAPP(c_COMBK(X58,X5),X49),X314)=X49), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBK__1)).
fof(c_0_5, axiom, (![X21]:![X5]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),X21))=X21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__complement)).
fof(c_0_6, axiom, (![X18]:(~(hBOOL(hAPP(c_fNot,X18)))|~(hBOOL(X18)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__1)).
fof(c_0_7, axiom, (![X5]:c_Orderings_Otop__class_Otop(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),hAPP(c_COMBK(tc_HOL_Obool,X5),c_fTrue))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__def)).
fof(c_0_8, axiom, (![X32]:![X16]:![X18]:![X116]:![X5]:![X20]:hAPP(hAPP(hAPP(c_COMBB(X20,X5,X116),X18),X16),X32)=hAPP(X18,hAPP(X16,X32))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBB__1)).
fof(c_0_9, axiom, (![X5]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X5,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__empty__eq)).
fof(c_0_10, axiom, (![X171]:![X5]:![X20]:hAPP(c_Relation_Oconverse(X20,X5),hAPP(c_Relation_Oconverse(X5,X20),X171))=X171), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_converse__converse)).
fof(c_0_11, axiom, (![X25]:![X20]:![X5]:(class_Orderings_Obot(X5)=>hAPP(c_Orderings_Obot__class_Obot(tc_fun(X20,X5)),X25)=c_Orderings_Obot__class_Obot(X5))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_bot__apply)).
fof(c_0_12, axiom, (![X58]:(class_Lattices_Oboolean__algebra(X58)=>hAPP(c_Groups_Ouminus__class_Ouminus(X58),c_Orderings_Obot__class_Obot(X58))=c_Orderings_Otop__class_Otop(X58))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_compl__bot__eq)).
fof(c_0_13, axiom, (class_Orderings_Obot(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Obot)).
fof(c_0_14, axiom, (![X36]:hAPP(c_Com_Ocom_OBODY,X36)!=c_Com_Ocom_OSKIP), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_com_Osimps_I19_J)).
fof(c_0_15, axiom, (![X33]:![X25]:![X5]:(class_Lattices_Oboolean__algebra(X5)=>(hAPP(c_Groups_Ouminus__class_Ouminus(X5),X25)=hAPP(c_Groups_Ouminus__class_Ouminus(X5),X33)<=>X25=X33))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_compl__eq__compl__iff)).
fof(c_0_16, axiom, (![X25]:![X20]:![X5]:(class_Orderings_Otop(X5)=>hAPP(c_Orderings_Otop__class_Otop(tc_fun(X20,X5)),X25)=c_Orderings_Otop__class_Otop(X5))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_top__apply)).
fof(c_0_17, axiom, (class_Lattices_Oboolean__algebra(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Lattices_Oboolean__algebra)).
fof(c_0_18, axiom, (~(hBOOL(c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fFalse__1)).
fof(c_0_19, axiom, (class_Orderings_Otop(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Otop)).
fof(c_0_20, axiom, (![X66]:![X58]:(class_Lattices_Oboolean__algebra(X58)=>hAPP(c_Groups_Ouminus__class_Ouminus(X58),hAPP(c_Groups_Ouminus__class_Ouminus(X58),X66))=X66)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__compl)).
fof(c_0_21, axiom, (![X43]:![X44]:(hAPP(c_Com_Ocom_OBODY,X44)=hAPP(c_Com_Ocom_OBODY,X43)<=>X44=X43)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_com_Osimps_I6_J)).
fof(c_0_22, axiom, (hBOOL(c_fTrue)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fTrue__1)).
fof(c_0_23, axiom, (![X43]:![X11]:![X84]:c_Com_Ocom_OWhile(X84,X11)!=hAPP(c_Com_Ocom_OBODY,X43)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_com_Osimps_I58_J)).
fof(c_0_24, axiom, (![X8]:![X92]:c_Com_Ocom_OSKIP!=c_Com_Ocom_OWhile(X92,X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_com_Osimps_I16_J)).
fof(c_0_25, plain, (![X18]:![X5]:(hAPP(c_Set_OCollect(X5),X18)=c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))<=>![X3]:~hBOOL(hAPP(X18,X3)))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_26, plain, (![X6]:c_Orderings_Obot__class_Obot(tc_fun(X6,tc_HOL_Obool))=hAPP(c_Set_OCollect(X6),hAPP(c_COMBK(tc_HOL_Obool,X6),c_fFalse))), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_27, plain, (![X19]:![X20]:hAPP(c_Set_OCollect(X20),X19)=X19), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_28, plain, (![X19]:![X20]:hAPP(c_Set_OCollect(X20),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X20),c_fNot),X19))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X20,tc_HOL_Obool)),hAPP(c_Set_OCollect(X20),X19))), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_29, plain, (![X19]:![X20]:![X21]:![X22]:![X24]:((hAPP(c_Set_OCollect(X20),X19)!=c_Orderings_Obot__class_Obot(tc_fun(X20,tc_HOL_Obool))|~hBOOL(hAPP(X19,X21)))&(hBOOL(hAPP(X22,esk289_1(X22)))|hAPP(c_Set_OCollect(X24),X22)=c_Orderings_Obot__class_Obot(tc_fun(X24,tc_HOL_Obool))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])])).
fof(c_0_30, plain, (![X315]:![X316]:![X317]:![X318]:hAPP(hAPP(c_COMBK(X318,X317),X316),X315)=X316), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_31,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32,plain,(hAPP(c_Set_OCollect(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_33, plain, (![X22]:![X23]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X23,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X23,tc_HOL_Obool)),X22))=X22), inference(variable_rename,[status(thm)],[c_0_5])).
cnf(c_0_34,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Set_OCollect(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_35, plain, (![X18]:(~hBOOL(hAPP(c_fNot,X18))|~hBOOL(X18))), inference(fof_simplification,[status(thm)],[c_0_6])).
cnf(c_0_36,plain,(hAPP(c_Set_OCollect(X1),X2)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))|hBOOL(hAPP(X2,esk289_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37,plain,(hAPP(hAPP(c_COMBK(X1,X2),X3),X4)=X3), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_39, plain, (![X6]:c_Orderings_Otop__class_Otop(tc_fun(X6,tc_HOL_Obool))=hAPP(c_Set_OCollect(X6),hAPP(c_COMBK(tc_HOL_Obool,X6),c_fTrue))), inference(variable_rename,[status(thm)],[c_0_7])).
fof(c_0_40, plain, (![X117]:![X118]:![X119]:![X120]:![X121]:![X122]:hAPP(hAPP(hAPP(c_COMBB(X122,X121,X120),X119),X118),X117)=hAPP(X119,hAPP(X118,X117))), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_41,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2)=hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_32]), c_0_32])).
fof(c_0_43, plain, (![X19]:(~hBOOL(hAPP(c_fNot,X19))|~hBOOL(X19))), inference(variable_rename,[status(thm)],[c_0_35])).
cnf(c_0_44,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hBOOL(hAPP(X1,esk289_1(X1)))), inference(rw,[status(thm)],[c_0_36, c_0_32])).
cnf(c_0_45,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)),X2)=c_fFalse), inference(spm,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_46, plain, (![X6]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X6,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X6,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X6,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_9])).
cnf(c_0_47,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48,plain,(hAPP(hAPP(hAPP(c_COMBB(X1,X2,X3),X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42]), c_0_42])).
fof(c_0_50, plain, (![X172]:![X173]:![X174]:hAPP(c_Relation_Oconverse(X174,X173),hAPP(c_Relation_Oconverse(X173,X174),X172))=X172), inference(variable_rename,[status(thm)],[c_0_10])).
cnf(c_0_51,plain,(~hBOOL(X1)|~hBOOL(hAPP(c_fNot,X1))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52,plain,(c_fFalse=X1|hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_44]), c_0_45]), c_0_37])).
cnf(c_0_53,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_54,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue)=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_47, c_0_32])).
fof(c_0_55, plain, (![X26]:![X27]:![X28]:(~class_Orderings_Obot(X28)|hAPP(c_Orderings_Obot__class_Obot(tc_fun(X27,X28)),X26)=c_Orderings_Obot__class_Obot(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).
cnf(c_0_56,plain,(hAPP(c_fNot,hAPP(c_fNot,hAPP(X1,X2)))=hAPP(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_48])).
cnf(c_0_57,plain,(hAPP(c_Relation_Oconverse(X1,X2),hAPP(c_Relation_Oconverse(X2,X1),X3))=X3), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58,plain,(hAPP(c_fNot,X1)=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_53, c_0_42])).
cnf(c_0_60,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)),X2)=c_fTrue), inference(spm,[status(thm)],[c_0_37, c_0_54])).
fof(c_0_61, plain, (![X59]:(~class_Lattices_Oboolean__algebra(X59)|hAPP(c_Groups_Ouminus__class_Ouminus(X59),c_Orderings_Obot__class_Obot(X59))=c_Orderings_Otop__class_Otop(X59))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
cnf(c_0_62,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,X2)),X3)=c_Orderings_Obot__class_Obot(X2)|~class_Orderings_Obot(X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63,plain,(class_Orderings_Obot(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_64, plain, (![X37]:hAPP(c_Com_Ocom_OBODY,X37)!=c_Com_Ocom_OSKIP), inference(variable_rename,[status(thm)],[c_0_14])).
cnf(c_0_65,plain,(hAPP(c_fNot,hAPP(c_fNot,X1))=X1), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_66,plain,(hAPP(c_fNot,X1)=c_fFalse|c_fFalse=X1), inference(spm,[status(thm)],[c_0_58, c_0_52])).
cnf(c_0_67,plain,(hAPP(c_fNot,c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_59]), c_0_60]), c_0_45])).
fof(c_0_68, plain, (![X34]:![X35]:![X36]:(((hAPP(c_Groups_Ouminus__class_Ouminus(X36),X35)!=hAPP(c_Groups_Ouminus__class_Ouminus(X36),X34)|X35=X34)|~class_Lattices_Oboolean__algebra(X36))&((X35!=X34|hAPP(c_Groups_Ouminus__class_Ouminus(X36),X35)=hAPP(c_Groups_Ouminus__class_Ouminus(X36),X34))|~class_Lattices_Oboolean__algebra(X36)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
fof(c_0_69, plain, (![X26]:![X27]:![X28]:(~class_Orderings_Otop(X28)|hAPP(c_Orderings_Otop__class_Otop(tc_fun(X27,X28)),X26)=c_Orderings_Otop__class_Otop(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).
cnf(c_0_70,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(X1),c_Orderings_Obot__class_Obot(X1))=c_Orderings_Otop__class_Otop(X1)|~class_Lattices_Oboolean__algebra(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71,plain,(class_Lattices_Oboolean__algebra(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_72,plain,(c_Orderings_Obot__class_Obot(tc_HOL_Obool)=c_fFalse), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_45])).
cnf(c_0_73,plain,(hAPP(c_Com_Ocom_OBODY,X1)!=c_Com_Ocom_OSKIP), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_74,plain,(c_fFalse=X1|c_fTrue=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
fof(c_0_75, plain, (~hBOOL(c_fFalse)), inference(fof_simplification,[status(thm)],[c_0_18])).
cnf(c_0_76,plain,(X2=X3|~class_Lattices_Oboolean__algebra(X1)|hAPP(c_Groups_Ouminus__class_Ouminus(X1),X2)!=hAPP(c_Groups_Ouminus__class_Ouminus(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_77,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,X2)),X3)=c_Orderings_Otop__class_Otop(X2)|~class_Orderings_Otop(X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_78,plain,(class_Orderings_Otop(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_79,plain,(c_Orderings_Otop__class_Otop(tc_HOL_Obool)=hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fFalse)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])).
fof(c_0_80, plain, (![X67]:![X68]:(~class_Lattices_Oboolean__algebra(X68)|hAPP(c_Groups_Ouminus__class_Ouminus(X68),hAPP(c_Groups_Ouminus__class_Ouminus(X68),X67))=X67)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).
cnf(c_0_81,plain,(~hBOOL(hAPP(X1,X2))|hAPP(c_Set_OCollect(X3),X1)!=c_Orderings_Obot__class_Obot(tc_fun(X3,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_82,plain,(hAPP(c_Com_Ocom_OBODY,X1)=c_fFalse|c_fTrue!=c_Com_Ocom_OSKIP), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_83,plain,(~hBOOL(c_fFalse)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_84,plain,(X1=X2|hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)!=hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X2)), inference(spm,[status(thm)],[c_0_76, c_0_71])).
cnf(c_0_85,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_60]), c_0_79])).
cnf(c_0_86,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(X1),hAPP(c_Groups_Ouminus__class_Ouminus(X1),X2))=X2|~class_Lattices_Oboolean__algebra(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
fof(c_0_87, plain, (![X45]:![X46]:![X47]:![X48]:((hAPP(c_Com_Ocom_OBODY,X46)!=hAPP(c_Com_Ocom_OBODY,X45)|X46=X45)&(X48!=X47|hAPP(c_Com_Ocom_OBODY,X48)=hAPP(c_Com_Ocom_OBODY,X47)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])).
cnf(c_0_88,plain,(X1!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|~hBOOL(hAPP(X1,X3))), inference(rw,[status(thm)],[c_0_81, c_0_32])).
cnf(c_0_89,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=c_Com_Ocom_OBODY|c_fTrue!=c_Com_Ocom_OSKIP), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_82]), c_0_83])).
cnf(c_0_90,plain,(c_fFalse=X1|hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)!=c_fTrue), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_91,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1))=X1), inference(spm,[status(thm)],[c_0_86, c_0_71])).
cnf(c_0_92,plain,(X1=X2|hAPP(c_Com_Ocom_OBODY,X1)!=hAPP(c_Com_Ocom_OBODY,X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_93,plain,(c_fTrue!=c_Com_Ocom_OSKIP|X1!=c_Com_Ocom_OBODY|~hBOOL(hAPP(X1,X2))), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_94,plain,(hBOOL(c_fTrue)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_95,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)=c_fFalse|X1!=c_fTrue), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_96,plain,(X1=X2|c_fTrue!=c_Com_Ocom_OSKIP), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_82]), c_0_82])).
cnf(c_0_97,plain,(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool)!=c_Com_Ocom_OBODY|c_fTrue!=c_Com_Ocom_OSKIP), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_85]), c_0_94])])).
cnf(c_0_98,plain,(c_fTrue!=c_Com_Ocom_OSKIP), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_95]), c_0_96]), c_0_97]), c_0_96])).
cnf(c_0_99,plain,(c_fFalse=X1|X1!=c_Com_Ocom_OSKIP), inference(spm,[status(thm)],[c_0_98, c_0_74])).
fof(c_0_100, plain, (![X85]:![X86]:![X87]:c_Com_Ocom_OWhile(X87,X86)!=hAPP(c_Com_Ocom_OBODY,X85)), inference(variable_rename,[status(thm)],[c_0_23])).
cnf(c_0_101,plain,(c_fFalse=c_Com_Ocom_OSKIP), inference(er,[status(thm)],[c_0_99])).
fof(c_0_102, plain, (![X93]:![X94]:c_Com_Ocom_OSKIP!=c_Com_Ocom_OWhile(X94,X93)), inference(variable_rename,[status(thm)],[c_0_24])).
cnf(c_0_103,plain,(c_Com_Ocom_OWhile(X1,X2)!=hAPP(c_Com_Ocom_OBODY,X3)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_104,plain,(c_fTrue=X1|c_Com_Ocom_OSKIP=X1), inference(rw,[status(thm)],[c_0_74, c_0_101])).
cnf(c_0_105,plain,(c_Com_Ocom_OSKIP!=c_Com_Ocom_OWhile(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_106,plain,(hAPP(c_Com_Ocom_OBODY,X1)!=c_fTrue), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_105])).
cnf(c_0_107,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_104]), c_0_73]), ['proof']).
# SZS output end CNFRefutation.
EOF
