% Problem    : REL042+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n112.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:46:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:![X2]:converse(composition(X1,X2))=composition(converse(X2),converse(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', converse_multiplicativity)).
fof(c_0_1, axiom, (![X1]:converse(converse(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', converse_idempotence)).
fof(c_0_2, axiom, (![X1]:composition(X1,one)=X1), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', composition_identity)).
fof(c_0_3, axiom, (![X1]:![X2]:join(composition(converse(X1),complement(composition(X1,X2))),complement(X2))=complement(X2)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', converse_cancellativity)).
fof(c_0_4, axiom, (![X1]:![X2]:join(X1,X2)=join(X2,X1)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', maddux1_join_commutativity)).
fof(c_0_5, axiom, (![X1]:zero=meet(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', def_zero)).
fof(c_0_6, axiom, (![X1]:![X2]:meet(X1,X2)=complement(join(complement(X1),complement(X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', maddux4_definiton_of_meet)).
fof(c_0_7, axiom, (![X1]:top=join(X1,complement(X1))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', def_top)).
fof(c_0_8, axiom, (![X1]:![X2]:X1=join(complement(join(complement(X1),complement(X2))),complement(join(complement(X1),X2)))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', maddux3_a_kind_of_de_Morgan)).
fof(c_0_9, axiom, (![X1]:![X2]:![X3]:join(X1,join(X2,X3))=join(join(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', maddux2_join_associativity)).
fof(c_0_10, axiom, (![X1]:![X2]:converse(join(X1,X2))=join(converse(X1),converse(X2))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', converse_additivity)).
fof(c_0_11, conjecture, (![X1]:(![X2]:meet(composition(X1,X2),composition(X1,complement(X2)))=zero=>join(composition(converse(X1),X1),one)=one)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
fof(c_0_12, axiom, (![X1]:![X2]:![X3]:composition(join(X1,X2),X3)=join(composition(X1,X3),composition(X2,X3))), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', composition_distributivity)).
fof(c_0_13, axiom, (![X1]:![X2]:![X3]:composition(X1,composition(X2,X3))=composition(composition(X1,X2),X3)), file('/export/starexec/sandbox/benchmark/Axioms/REL001+0.ax', composition_associativity)).
fof(c_0_14, plain, (![X3]:![X4]:converse(composition(X3,X4))=composition(converse(X4),converse(X3))), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_15, plain, (![X2]:converse(converse(X2))=X2), inference(variable_rename,[status(thm)],[c_0_1])).
cnf(c_0_16,plain,(converse(composition(X1,X2))=composition(converse(X2),converse(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17,plain,(converse(converse(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_18, plain, (![X2]:composition(X2,one)=X2), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_19,plain,(converse(composition(converse(X1),X2))=composition(converse(X2),X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,plain,(composition(X1,one)=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_21, plain, (![X3]:![X4]:join(composition(converse(X3),complement(composition(X3,X4))),complement(X4))=complement(X4)), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_22, plain, (![X3]:![X4]:join(X3,X4)=join(X4,X3)), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_23,plain,(composition(converse(one),X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_17])).
fof(c_0_24, plain, (![X2]:zero=meet(X2,complement(X2))), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_25, plain, (![X3]:![X4]:meet(X3,X4)=complement(join(complement(X3),complement(X4)))), inference(variable_rename,[status(thm)],[c_0_6])).
cnf(c_0_26,plain,(join(composition(converse(X1),complement(composition(X1,X2))),complement(X2))=complement(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27,plain,(join(X1,X2)=join(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28,plain,(converse(one)=one), inference(spm,[status(thm)],[c_0_20, c_0_23])).
cnf(c_0_29,plain,(zero=meet(X1,complement(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30,plain,(meet(X1,X2)=complement(join(complement(X1),complement(X2)))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_31, plain, (![X2]:top=join(X2,complement(X2))), inference(variable_rename,[status(thm)],[c_0_7])).
cnf(c_0_32,plain,(join(complement(X1),composition(converse(X2),complement(composition(X2,X1))))=complement(X1)), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33,plain,(composition(one,X1)=X1), inference(rw,[status(thm)],[c_0_23, c_0_28])).
cnf(c_0_34,plain,(complement(join(complement(X1),complement(complement(X1))))=zero), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35,plain,(top=join(X1,complement(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_36, plain, (![X3]:![X4]:X3=join(complement(join(complement(X3),complement(X4))),complement(join(complement(X3),X4)))), inference(variable_rename,[status(thm)],[c_0_8])).
fof(c_0_37, plain, (![X4]:![X5]:![X6]:join(X4,join(X5,X6))=join(join(X4,X5),X6)), inference(variable_rename,[status(thm)],[c_0_9])).
cnf(c_0_38,plain,(join(complement(X1),complement(X1))=complement(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_28]), c_0_33])).
cnf(c_0_39,plain,(complement(top)=zero), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40,plain,(X1=join(complement(join(complement(X1),complement(X2))),complement(join(complement(X1),X2)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41,plain,(join(X1,join(X2,X3))=join(join(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_42,plain,(join(zero,zero)=zero), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43,plain,(join(complement(join(complement(X1),X2)),complement(join(complement(X1),complement(X2))))=X1), inference(rw,[status(thm)],[c_0_40, c_0_27])).
cnf(c_0_44,plain,(join(zero,join(zero,X1))=join(zero,X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45,plain,(join(zero,complement(complement(X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_38]), c_0_35]), c_0_39]), c_0_27])).
cnf(c_0_46,plain,(join(zero,X1)=X1), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_47,plain,(complement(complement(X1))=X1), inference(rw,[status(thm)],[c_0_45, c_0_46])).
fof(c_0_48, plain, (![X3]:![X4]:converse(join(X3,X4))=join(converse(X3),converse(X4))), inference(variable_rename,[status(thm)],[c_0_10])).
cnf(c_0_49,plain,(join(X1,X1)=X1), inference(spm,[status(thm)],[c_0_38, c_0_47])).
cnf(c_0_50,plain,(converse(join(X1,X2))=join(converse(X1),converse(X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_51,plain,(join(X1,join(X1,X2))=join(X1,X2)), inference(spm,[status(thm)],[c_0_41, c_0_49])).
cnf(c_0_52,plain,(converse(join(converse(X1),X2))=join(X1,converse(X2))), inference(spm,[status(thm)],[c_0_50, c_0_17])).
cnf(c_0_53,plain,(join(X1,top)=top), inference(spm,[status(thm)],[c_0_51, c_0_35])).
cnf(c_0_54,plain,(join(X1,converse(complement(converse(X1))))=converse(top)), inference(spm,[status(thm)],[c_0_52, c_0_35])).
cnf(c_0_55,plain,(join(top,X1)=top), inference(spm,[status(thm)],[c_0_27, c_0_53])).
cnf(c_0_56,plain,(converse(top)=top), inference(spm,[status(thm)],[c_0_54, c_0_55])).
fof(c_0_57, negated_conjecture, (~(![X1]:(![X2]:meet(composition(X1,X2),composition(X1,complement(X2)))=zero=>join(composition(converse(X1),X1),one)=one))), inference(assume_negation,[status(cth)],[c_0_11])).
cnf(c_0_58,plain,(join(X1,converse(complement(converse(X1))))=top), inference(rw,[status(thm)],[c_0_54, c_0_56])).
fof(c_0_59, negated_conjecture, (![X4]:(meet(composition(esk1_0,X4),composition(esk1_0,complement(X4)))=zero&join(composition(converse(esk1_0),esk1_0),one)!=one)), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_57])])])])).
cnf(c_0_60,plain,(complement(join(complement(X1),complement(converse(complement(converse(complement(X1)))))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_58]), c_0_39]), c_0_46])).
cnf(c_0_61,plain,(join(X1,complement(join(complement(X1),X2)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_43]), c_0_27])).
cnf(c_0_62,plain,(join(X1,join(X2,X1))=join(X2,X1)), inference(spm,[status(thm)],[c_0_51, c_0_27])).
fof(c_0_63, plain, (![X4]:![X5]:![X6]:composition(join(X4,X5),X6)=join(composition(X4,X6),composition(X5,X6))), inference(variable_rename,[status(thm)],[c_0_12])).
cnf(c_0_64,plain,(join(complement(one),composition(converse(X1),complement(X1)))=complement(one)), inference(spm,[status(thm)],[c_0_32, c_0_20])).
cnf(c_0_65,negated_conjecture,(meet(composition(esk1_0,X1),composition(esk1_0,complement(X1)))=zero), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_66,plain,(complement(join(X1,complement(converse(complement(converse(X1))))))=complement(X1)), inference(spm,[status(thm)],[c_0_60, c_0_47])).
cnf(c_0_67,plain,(join(X1,complement(join(X2,complement(X1))))=X1), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68,plain,(composition(join(X1,X2),X3)=join(composition(X1,X3),composition(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_69,plain,(join(complement(one),composition(converse(complement(X1)),X1))=complement(one)), inference(spm,[status(thm)],[c_0_64, c_0_47])).
cnf(c_0_70,negated_conjecture,(complement(join(complement(composition(esk1_0,X1)),complement(composition(esk1_0,complement(X1)))))=zero), inference(rw,[status(thm)],[c_0_65, c_0_30])).
cnf(c_0_71,plain,(join(X1,complement(converse(complement(converse(X1)))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_66]), c_0_43])).
cnf(c_0_72,plain,(join(complement(join(complement(X1),X2)),complement(join(complement(X2),complement(X1))))=X1), inference(spm,[status(thm)],[c_0_43, c_0_27])).
cnf(c_0_73,plain,(join(complement(join(X1,X2)),complement(join(X1,complement(X2))))=complement(X1)), inference(spm,[status(thm)],[c_0_43, c_0_47])).
cnf(c_0_74,plain,(join(complement(X1),complement(join(X2,X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_67, c_0_47])).
cnf(c_0_75,plain,(converse(join(composition(converse(X1),X2),X3))=join(composition(converse(X2),X1),converse(X3))), inference(spm,[status(thm)],[c_0_52, c_0_19])).
cnf(c_0_76,plain,(join(composition(X1,converse(X2)),converse(composition(X2,X3)))=composition(join(X1,converse(X3)),converse(X2))), inference(spm,[status(thm)],[c_0_68, c_0_16])).
cnf(c_0_77,plain,(converse(join(X1,converse(X2)))=join(converse(X1),X2)), inference(spm,[status(thm)],[c_0_50, c_0_17])).
cnf(c_0_78,plain,(join(complement(one),converse(complement(one)))=complement(one)), inference(spm,[status(thm)],[c_0_69, c_0_20])).
cnf(c_0_79,plain,(join(X1,join(X2,X3))=join(X3,join(X1,X2))), inference(spm,[status(thm)],[c_0_27, c_0_41])).
cnf(c_0_80,negated_conjecture,(complement(join(complement(esk1_0),complement(composition(esk1_0,complement(one)))))=zero), inference(spm,[status(thm)],[c_0_70, c_0_20])).
cnf(c_0_81,plain,(join(X1,converse(complement(converse(complement(X1)))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_71]), c_0_17]), c_0_17])).
cnf(c_0_82,plain,(join(X1,complement(join(X1,complement(X2))))=join(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_47]), c_0_47]), c_0_41]), c_0_74])).
cnf(c_0_83,plain,(join(X1,zero)=X1), inference(spm,[status(thm)],[c_0_27, c_0_46])).
cnf(c_0_84,plain,(join(composition(X1,X2),composition(X1,X3))=composition(X1,join(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_50]), c_0_16]), c_0_17]), c_0_17]), c_0_17])).
cnf(c_0_85,plain,(join(complement(X1),composition(X2,complement(composition(converse(X2),X1))))=complement(X1)), inference(spm,[status(thm)],[c_0_32, c_0_17])).
cnf(c_0_86,plain,(join(X1,composition(X2,X1))=composition(join(one,X2),X1)), inference(spm,[status(thm)],[c_0_68, c_0_33])).
cnf(c_0_87,plain,(converse(complement(one))=complement(one)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_27]), c_0_78])).
cnf(c_0_88,plain,(join(X1,join(complement(X1),X2))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_35]), c_0_53])).
cnf(c_0_89,negated_conjecture,(complement(join(complement(esk1_0),composition(esk1_0,complement(one))))=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_80]), c_0_47]), c_0_46])).
cnf(c_0_90,plain,(join(complement(X1),converse(complement(converse(X1))))=complement(X1)), inference(spm,[status(thm)],[c_0_81, c_0_47])).
cnf(c_0_91,plain,(join(complement(join(X1,complement(X2))),complement(join(complement(X2),complement(X1))))=X2), inference(spm,[status(thm)],[c_0_43, c_0_27])).
cnf(c_0_92,plain,(join(X1,complement(join(X1,X2)))=join(X1,complement(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_82]), c_0_51])).
cnf(c_0_93,negated_conjecture,(join(complement(composition(esk1_0,X1)),complement(composition(esk1_0,complement(X1))))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_70]), c_0_41]), c_0_83])).
cnf(c_0_94,plain,(join(X1,composition(X1,X2))=composition(X1,join(X2,one))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_20]), c_0_27])).
cnf(c_0_95,plain,(join(X1,composition(X2,complement(composition(converse(X2),complement(X1)))))=X1), inference(spm,[status(thm)],[c_0_85, c_0_47])).
cnf(c_0_96,plain,(join(X1,composition(complement(one),X1))=composition(top,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_58]), c_0_28]), c_0_87])).
cnf(c_0_97,plain,(converse(composition(X1,converse(X2)))=composition(X2,converse(X1))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_98,plain,(join(X1,join(X2,complement(X1)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_47]), c_0_79])).
cnf(c_0_99,plain,(join(converse(X1),join(converse(X2),X3))=join(converse(join(X1,X2)),X3)), inference(spm,[status(thm)],[c_0_41, c_0_50])).
cnf(c_0_100,negated_conjecture,(join(complement(esk1_0),composition(esk1_0,complement(one)))=complement(esk1_0)), inference(spm,[status(thm)],[c_0_47, c_0_89])).
cnf(c_0_101,plain,(complement(converse(complement(converse(X1))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_90]), c_0_47]), c_0_27]), c_0_71])).
cnf(c_0_102,plain,(join(complement(join(X1,X2)),complement(join(X2,complement(X1))))=complement(X2)), inference(spm,[status(thm)],[c_0_91, c_0_47])).
cnf(c_0_103,negated_conjecture,(join(complement(composition(esk1_0,X1)),composition(esk1_0,complement(X1)))=complement(composition(esk1_0,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_39]), c_0_83]), c_0_47])).
cnf(c_0_104,plain,(composition(X1,join(one,complement(composition(converse(X1),complement(X1)))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_27])).
cnf(c_0_105,plain,(join(one,complement(composition(converse(X1),complement(X1))))=complement(composition(converse(X1),complement(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_64]), c_0_47]), c_0_27])).
cnf(c_0_106,plain,(join(complement(join(X1,X2)),complement(join(complement(X2),X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_72, c_0_47])).
cnf(c_0_107,plain,(join(X1,composition(X1,complement(one)))=composition(X1,top)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_96]), c_0_97]), c_0_56]), c_0_97]), c_0_87])).
cnf(c_0_108,plain,(join(complement(converse(X1)),converse(join(X1,X2)))=top), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_27])).
cnf(c_0_109,negated_conjecture,(join(esk1_0,complement(composition(esk1_0,complement(one))))=complement(composition(esk1_0,complement(one)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_100]), c_0_47]), c_0_27])).
cnf(c_0_110,plain,(converse(complement(converse(X1)))=complement(X1)), inference(spm,[status(thm)],[c_0_47, c_0_101])).
fof(c_0_111, plain, (![X4]:![X5]:![X6]:composition(X4,composition(X5,X6))=composition(composition(X4,X5),X6)), inference(variable_rename,[status(thm)],[c_0_13])).
cnf(c_0_112,negated_conjecture,(join(composition(esk1_0,X1),complement(composition(esk1_0,top)))=complement(composition(esk1_0,complement(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_47]), c_0_47]), c_0_84]), c_0_27]), c_0_35])).
cnf(c_0_113,plain,(composition(X1,complement(composition(converse(X1),complement(X1))))=X1), inference(rw,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_114,negated_conjecture,(complement(composition(esk1_0,complement(one)))=join(esk1_0,complement(composition(esk1_0,top)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_100]), c_0_27]), c_0_107]), c_0_47]), c_0_27])).
cnf(c_0_115,negated_conjecture,(join(complement(converse(esk1_0)),converse(complement(composition(esk1_0,complement(one)))))=top), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_116,plain,(converse(complement(X1))=complement(converse(X1))), inference(spm,[status(thm)],[c_0_17, c_0_110])).
cnf(c_0_117,plain,(converse(composition(X1,complement(one)))=composition(complement(one),converse(X1))), inference(spm,[status(thm)],[c_0_16, c_0_87])).
cnf(c_0_118,plain,(composition(X1,composition(X2,X3))=composition(composition(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_119,negated_conjecture,(complement(composition(esk1_0,composition(converse(esk1_0),complement(esk1_0))))=join(esk1_0,complement(composition(esk1_0,top)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_47])).
cnf(c_0_120,negated_conjecture,(complement(join(esk1_0,complement(composition(esk1_0,top))))=composition(esk1_0,complement(one))), inference(spm,[status(thm)],[c_0_47, c_0_114])).
cnf(c_0_121,negated_conjecture,(join(complement(converse(esk1_0)),complement(composition(complement(one),converse(esk1_0))))=top), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_116]), c_0_117])).
cnf(c_0_122,plain,(converse(composition(X1,composition(converse(X2),X3)))=composition(converse(X3),composition(X2,converse(X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_19]), c_0_118])).
cnf(c_0_123,negated_conjecture,(composition(esk1_0,composition(converse(esk1_0),complement(esk1_0)))=composition(esk1_0,complement(one))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_119]), c_0_120])).
cnf(c_0_124,negated_conjecture,(complement(join(complement(converse(esk1_0)),composition(complement(one),converse(esk1_0))))=converse(esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_121]), c_0_39]), c_0_47]), c_0_46])).
cnf(c_0_125,negated_conjecture,(composition(complement(converse(esk1_0)),composition(esk1_0,converse(esk1_0)))=composition(complement(one),converse(esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_117]), c_0_116])).
cnf(c_0_126,negated_conjecture,(join(complement(converse(esk1_0)),composition(complement(one),converse(esk1_0)))=complement(converse(esk1_0))), inference(spm,[status(thm)],[c_0_47, c_0_124])).
cnf(c_0_127,plain,(join(complement(converse(X1)),composition(X2,complement(converse(composition(X1,X2)))))=complement(converse(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_16]), c_0_17])).
cnf(c_0_128,negated_conjecture,(composition(complement(converse(esk1_0)),join(one,composition(esk1_0,converse(esk1_0))))=complement(converse(esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_125]), c_0_126]), c_0_27])).
cnf(c_0_129,plain,(join(X1,composition(join(one,X2),X1))=composition(join(one,X2),X1)), inference(spm,[status(thm)],[c_0_51, c_0_86])).
cnf(c_0_130,negated_conjecture,(join(esk1_0,composition(esk1_0,composition(converse(esk1_0),esk1_0)))=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_110]), c_0_47]), c_0_110]), c_0_47]), c_0_129]), c_0_110]), c_0_47]), c_0_86]), c_0_118])).
cnf(c_0_131,negated_conjecture,(join(esk1_0,join(composition(esk1_0,composition(converse(esk1_0),esk1_0)),X1))=join(esk1_0,X1)), inference(spm,[status(thm)],[c_0_41, c_0_130])).
cnf(c_0_132,plain,(join(X1,complement(composition(X2,complement(composition(converse(X2),X1)))))=complement(composition(X2,complement(composition(converse(X2),X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_85]), c_0_47]), c_0_27])).
cnf(c_0_133,negated_conjecture,(complement(composition(esk1_0,complement(composition(converse(esk1_0),esk1_0))))=join(esk1_0,complement(composition(esk1_0,top)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_112]), c_0_132])).
cnf(c_0_134,plain,(converse(join(complement(one),X1))=join(complement(one),converse(X1))), inference(spm,[status(thm)],[c_0_52, c_0_87])).
cnf(c_0_135,plain,(join(complement(one),composition(converse(X1),composition(X2,complement(composition(converse(X2),X1)))))=complement(one)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_19]), c_0_118])).
cnf(c_0_136,negated_conjecture,(composition(esk1_0,complement(composition(converse(esk1_0),esk1_0)))=composition(esk1_0,complement(one))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_133]), c_0_120])).
cnf(c_0_137,plain,(join(X1,composition(X1,X2))=composition(X1,join(one,X2))), inference(spm,[status(thm)],[c_0_84, c_0_20])).
cnf(c_0_138,plain,(join(complement(one),composition(join(complement(one),converse(X1)),complement(join(complement(one),X1))))=complement(one)), inference(spm,[status(thm)],[c_0_64, c_0_134])).
cnf(c_0_139,negated_conjecture,(join(complement(one),composition(converse(esk1_0),composition(esk1_0,complement(one))))=complement(one)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_140,plain,(composition(join(X1,composition(X1,X2)),X3)=composition(X1,composition(join(X2,one),X3))), inference(spm,[status(thm)],[c_0_118, c_0_94])).
cnf(c_0_141,plain,(join(X1,composition(X1,join(one,X2)))=composition(X1,join(one,X2))), inference(spm,[status(thm)],[c_0_51, c_0_137])).
cnf(c_0_142,negated_conjecture,(join(composition(converse(esk1_0),esk1_0),one)!=one), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_143,negated_conjecture,(composition(complement(one),join(one,composition(converse(esk1_0),esk1_0)))=complement(one)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_19]), c_0_117]), c_0_118]), c_0_47]), c_0_140]), c_0_27]), c_0_20]), c_0_141])).
cnf(c_0_144,negated_conjecture,(join(one,composition(converse(esk1_0),esk1_0))!=one), inference(rw,[status(thm)],[c_0_142, c_0_27])).
cnf(c_0_145,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_143]), c_0_116]), c_0_28]), c_0_47]), c_0_116]), c_0_28]), c_0_47]), c_0_20]), c_0_51]), c_0_116]), c_0_28]), c_0_47]), c_0_144]), ['proof']).
# SZS output end CNFRefutation.
EOF
