% Problem    : GEO285+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n179.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:34:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X67]:![X68]:![X69]:![X70]:(((((rless(X67,X69)&X69=X70)&rreal(X70))&X67=X68)&rreal(X68))=>(~(rless(X69,X67))|X67=X69))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(147), 0), imp(cond(axiom(147), 0)))')).
fof(c_0_1, axiom, (![X15]:![X16]:![X17]:![X18]:((((ron(X16,X18)&rcenter(X17,X18))&?[X19]:(X16=X19&rpoint(X19)))&?[X20]:(X15=X20&rpoint(X20)))=>(rless(vf(X17,X15),vf(X17,X16))<=>rinside(X15,X18)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(186), 0), imp(cond(axiom(186), 0)))')).
fof(c_0_2, axiom, (![X670]:![X671]:((?[X672]:(X671=X672&rpoint(X672))&?[X673]:(X670=X673&rpoint(X673)))=>vf(X670,X671)=vf(X671,X670))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(166), 0), imp(cond(axiom(166), 0)))')).
fof(c_0_3, axiom, (((((((((((rless(vf(vd1127,vd1129),vf(vd1123,vd1125))&vd1127!=vd1129)&vd1123!=vd1125)&vd1129=vd1130)&rpoint(vd1130))&vd1127=vd1128)&rpoint(vd1128))&vd1125=vd1126)&rpoint(vd1126))&vd1123=vd1124)&rpoint(vd1124))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(and(and(and(holds(conjunct2(conjunct2(conjunct2(conjunct2(conjunct2(conjunct2(242)))))), 1133, 0), and(holds(conjunct1(conjunct2(conjunct2(conjunct2(conjunct2(conjunct2(242)))))), 1132, 0), and(holds(conjunct1(conjunct2(conjunct2(conjunct2(conjunct2(242))))), 1131, 0), and(pred(conjunct1(conjunct2(conjunct2(conjunct2(242)))), 2), pred(conjunct1(conjunct2(conjunct2(conjunct2(242)))), 1))))), and(pred(conjunct1(conjunct2(conjunct2(242))), 2), pred(conjunct1(conjunct2(conjunct2(242))), 1))), and(pred(conjunct1(conjunct2(242)), 2), pred(conjunct1(conjunct2(242)), 1))), and(pred(conjunct1(242), 2), pred(conjunct1(242), 1)))')).
fof(c_0_4, axiom, (vd1123=vd1129), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(258, 1150, 0)')).
fof(c_0_5, axiom, (((((ron(vd1127,vd1152)&rcenter(vd1123,vd1152))&rcircle(vd1152))&ron(vd1125,vd1151))&ron(vd1123,vd1151))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(and(pred(comma_conjunct2(comma_conjunct2(comma_conjunct2(259))), 0), and(pred(comma_conjunct1(comma_conjunct2(comma_conjunct2(259))), 0), and(pred(comma_conjunct2(comma_conjunct2(259)), 0), pred(comma_conjunct1(comma_conjunct2(259)), 0)))), pred(comma_conjunct1(259), 0))')).
fof(c_0_6, axiom, (![X60]:![X61]:((?[X62]:(X61=X62&rpoint(X62))&?[X63]:(X60=X63&rpoint(X63)))=>?[X64]:(vf(X60,X61)=X64&rreal(X64)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(160), 0), imp(cond(axiom(160), 0)))')).
fof(c_0_7, conjecture, (~(rinside(vd1125,vd1152))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'neg(neg(260))')).
fof(c_0_8, axiom, (![X65]:![X66]:((X65=X66&rreal(X66))=>~(rless(X65,X65)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(149), 0), imp(cond(axiom(149), 0)))')).
fof(c_0_9, plain, (![X67]:![X68]:![X69]:![X70]:(((((rless(X67,X69)&X69=X70)&rreal(X70))&X67=X68)&rreal(X68))=>(~rless(X69,X67)|X67=X69))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_10, plain, (![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:(((~rless(vf(X23,X21),vf(X23,X22))|rinside(X21,X24))|(((~ron(X22,X24)|~rcenter(X23,X24))|(X22!=X25|~rpoint(X25)))|(X21!=X26|~rpoint(X26))))&((~rinside(X21,X24)|rless(vf(X23,X21),vf(X23,X22)))|(((~ron(X22,X24)|~rcenter(X23,X24))|(X22!=X25|~rpoint(X25)))|(X21!=X26|~rpoint(X26)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_11, plain, (![X71]:![X72]:![X73]:![X74]:(((((~rless(X71,X73)|X73!=X74)|~rreal(X74))|X71!=X72)|~rreal(X72))|(~rless(X73,X71)|X71=X73))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
cnf(c_0_12,plain,(rless(vf(X5,X2),vf(X5,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rcenter(X5,X6)|~ron(X4,X6)|~rinside(X2,X6)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, (![X674]:![X675]:![X676]:![X677]:(((X675!=X676|~rpoint(X676))|(X674!=X677|~rpoint(X677)))|vf(X674,X675)=vf(X675,X674))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
cnf(c_0_14,plain,(vd1129=vd1130), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_15,plain,(vd1123=vd1129), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16,plain,(X1=X2|~rless(X2,X1)|~rreal(X3)|X1!=X3|~rreal(X4)|X2!=X4|~rless(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(rless(vf(vd1127,vd1129),vf(vd1123,vd1125))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_18,plain,(rless(vf(X1,X2),vf(X1,X3))|~rpoint(X3)|~rpoint(X2)|~rcenter(X1,X4)|~ron(X3,X4)|~rinside(X2,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_12])])).
cnf(c_0_19,plain,(rcenter(vd1123,vd1152)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_20,plain,(vf(X1,X2)=vf(X2,X1)|~rpoint(X3)|X1!=X3|~rpoint(X4)|X2!=X4), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,plain,(rpoint(vd1130)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_22,plain,(vd1130=vd1123), inference(rw,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_23, plain, (![X65]:![X66]:![X67]:![X68]:((vf(X65,X66)=esk8_2(X65,X66)|((X66!=X67|~rpoint(X67))|(X65!=X68|~rpoint(X68))))&(rreal(esk8_2(X65,X66))|((X66!=X67|~rpoint(X67))|(X65!=X68|~rpoint(X68)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
cnf(c_0_24,plain,(X1=X2|~rreal(X2)|~rreal(X1)|~rless(X2,X1)|~rless(X1,X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_16])])).
cnf(c_0_25,plain,(rless(vf(vd1127,vd1123),vf(vd1123,vd1125))), inference(rw,[status(thm)],[c_0_17, c_0_15])).
cnf(c_0_26,plain,(rless(vf(vd1123,X1),vf(vd1123,X2))|~rpoint(X2)|~rpoint(X1)|~ron(X2,vd1152)|~rinside(X1,vd1152)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_27,plain,(vf(X1,X2)=vf(X2,X1)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_20])])).
cnf(c_0_28,plain,(rpoint(vd1123)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29,plain,(rpoint(vd1128)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_30,plain,(vd1127=vd1128), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_31,plain,(rpoint(vd1126)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_32,plain,(vd1125=vd1126), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_33, negated_conjecture, (rinside(vd1125,vd1152)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_7])])).
cnf(c_0_34,plain,(rreal(esk8_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35,plain,(vf(X2,X4)=esk8_2(X2,X4)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36,plain,(vf(vd1127,vd1123)=vf(vd1123,vd1125)|~rreal(vf(vd1127,vd1123))|~rreal(vf(vd1123,vd1125))|~rless(vf(vd1123,vd1125),vf(vd1127,vd1123))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_37,plain,(rless(vf(vd1123,X1),vf(X2,vd1123))|~rpoint(X2)|~rpoint(X1)|~ron(X2,vd1152)|~rinside(X1,vd1152)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_38,plain,(rpoint(vd1127)), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_39,plain,(rpoint(vd1125)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40,plain,(ron(vd1127,vd1152)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_41,negated_conjecture,(rinside(vd1125,vd1152)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42,plain,(rreal(esk8_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_34])])).
cnf(c_0_43,plain,(esk8_2(X1,X2)=vf(X1,X2)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_35])])).
fof(c_0_44, plain, (![X65]:![X66]:((X65=X66&rreal(X66))=>~rless(X65,X65))), inference(fof_simplification,[status(thm)],[c_0_8])).
cnf(c_0_45,plain,(vf(vd1127,vd1123)=vf(vd1123,vd1125)|~rreal(vf(vd1127,vd1123))|~rreal(vf(vd1123,vd1125))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41])])).
cnf(c_0_46,plain,(rreal(vf(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_47, plain, (![X67]:![X68]:((X67!=X68|~rreal(X68))|~rless(X67,X67))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_44])])])])).
cnf(c_0_48,plain,(vf(vd1127,vd1123)=vf(vd1123,vd1125)|~rreal(vf(vd1123,vd1125))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_28]), c_0_38])])).
cnf(c_0_49,plain,(~rless(X1,X1)|~rreal(X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_50,plain,(vf(vd1127,vd1123)=vf(vd1123,vd1125)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_46]), c_0_39]), c_0_28])])).
cnf(c_0_51,plain,(~rreal(X1)|~rless(X1,X1)), inference(er,[status(thm)],[c_0_49])).
cnf(c_0_52,plain,(rless(vf(vd1123,vd1125),vf(vd1123,vd1125))), inference(rw,[status(thm)],[c_0_25, c_0_50])).
cnf(c_0_53,plain,(rreal(vf(vd1123,vd1125))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_50]), c_0_28]), c_0_38])])).
cnf(c_0_54,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])]), ['proof']).
# SZS output end CNFRefutation.
EOF
