% Problem    : SWW330+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n126.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:21:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X6]:(~(~(hBOOL(X6)))|hBOOL(hAPP(c_fNot,X6)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__2)).
fof(c_0_1, axiom, (![X65]:![X19]:![X7]:(class_Lattices_Oboolean__algebra(X7)=>(hAPP(c_Groups_Ouminus__class_Ouminus(X7),X19)=hAPP(c_Groups_Ouminus__class_Ouminus(X7),X65)<=>X19=X65))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_compl__eq__compl__iff)).
fof(c_0_2, axiom, (![X28]:(class_Lattices_Oboolean__algebra(X28)=>hAPP(c_Groups_Ouminus__class_Ouminus(X28),c_Orderings_Obot__class_Obot(X28))=c_Orderings_Otop__class_Otop(X28))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_compl__bot__eq)).
fof(c_0_3, axiom, (![X7]:c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))=hAPP(c_Set_OCollect(X7),hAPP(c_COMBK(tc_HOL_Obool,X7),c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__def)).
fof(c_0_4, axiom, (![X6]:![X7]:hAPP(c_Set_OCollect(X7),X6)=X6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__def)).
fof(c_0_5, axiom, (![X7]:c_Orderings_Otop__class_Otop(tc_fun(X7,tc_HOL_Obool))=hAPP(c_Set_OCollect(X7),hAPP(c_COMBK(tc_HOL_Obool,X7),c_fTrue))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__def)).
fof(c_0_6, axiom, (![X21]:![X7]:(class_Finite__Set_Ofinite(X7)=>hBOOL(hAPP(c_Finite__Set_Ofinite(X7),X21)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite__code)).
fof(c_0_7, axiom, (![X311]:(class_Finite__Set_Ofinite(X311)=>class_Finite__Set_Ofinite(tc_Option_Ooption(X311)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Option__Ooption__Finite__Set_Ofinite)).
fof(c_0_8, axiom, (![X6]:![X7]:(class_Enum_Oenum(X7)=>(?[X97]:hBOOL(hAPP(X6,X97))<=>hBOOL(hAPP(c_Enum_Oenum__class_Oenum__ex(X7),X6))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_exists__code)).
fof(c_0_9, axiom, (class_Lattices_Oboolean__algebra(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Lattices_Oboolean__algebra)).
fof(c_0_10, axiom, (![X31]:![X28]:(class_Lattices_Oboolean__algebra(X28)=>hAPP(c_Groups_Ouminus__class_Ouminus(X28),hAPP(c_Groups_Ouminus__class_Ouminus(X28),X31))=X31)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__compl)).
fof(c_0_11, axiom, (![X312]:![X58]:![X7]:![X28]:hAPP(hAPP(c_COMBK(X28,X7),X58),X312)=X58), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBK__1)).
fof(c_0_12, axiom, (![X6]:![X7]:(?[X3]:(hBOOL(hAPP(hAPP(c_member(X7),X3),c_Orderings_Otop__class_Otop(tc_fun(X7,tc_HOL_Obool))))&hBOOL(hAPP(X6,X3)))<=>?[X97]:hBOOL(hAPP(X6,X97)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_bex__UNIV)).
fof(c_0_13, axiom, (![X19]:![X7]:hBOOL(hAPP(hAPP(c_member(X7),X19),c_Orderings_Otop__class_Otop(tc_fun(X7,tc_HOL_Obool))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__I)).
fof(c_0_14, axiom, (![X6]:(~(hBOOL(hAPP(c_fNot,X6)))|~(hBOOL(X6)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__1)).
fof(c_0_15, axiom, (![X6]:![X7]:(c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))=hAPP(c_Set_OCollect(X7),X6)<=>![X3]:~(hBOOL(hAPP(X6,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__Collect__eq)).
fof(c_0_16, axiom, (![X19]:![X22]:![X7]:(class_Orderings_Obot(X7)=>hAPP(c_Orderings_Obot__class_Obot(tc_fun(X22,X7)),X19)=c_Orderings_Obot__class_Obot(X7))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_bot__fun__def)).
fof(c_0_17, axiom, (![X19]:![X22]:![X7]:(class_Orderings_Otop(X7)=>hAPP(c_Orderings_Otop__class_Otop(tc_fun(X22,X7)),X19)=c_Orderings_Otop__class_Otop(X7))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_top__apply)).
fof(c_0_18, axiom, (![X249]:![X6]:![X7]:(hAPP(c_List_Ofilter(X7,X6),X249)=X249<=>![X3]:(hBOOL(hAPP(hAPP(c_member(X7),X3),hAPP(c_List_Oset(X7),X249)))=>hBOOL(hAPP(X6,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_filter__id__conv)).
fof(c_0_19, axiom, (class_Finite__Set_Ofinite(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Finite__Set_Ofinite)).
fof(c_0_20, axiom, (![X40]:![X7]:hAPP(c_Set_Oimage(X7,X7,c_COMBI(X7)),X40)=X40), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_image__ident)).
fof(c_0_21, axiom, (class_Enum_Oenum(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Enum_Oenum)).
fof(c_0_22, axiom, (![X6]:![X7]:hAPP(c_Set_OCollect(X7),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X7),c_fNot),X6))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X7,tc_HOL_Obool)),hAPP(c_Set_OCollect(X7),X6))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__neg__eq)).
fof(c_0_23, axiom, (![X72]:![X24]:![X6]:(hBOOL(hAPP(X6,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X24,X72)))<=>((X72=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=>hBOOL(hAPP(X6,X24)))&(X72!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=>![X167]:![X137]:(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X137),X72))=>(X24=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Nat_Onat),X72),X167)),X137)=>hBOOL(hAPP(X6,X137)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_split__mod)).
fof(c_0_24, axiom, (c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_transfer__nat__int__numerals_I1_J)).
fof(c_0_25, axiom, (c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Pls__def)).
fof(c_0_26, axiom, (class_Orderings_Obot(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Obot)).
fof(c_0_27, axiom, (class_Orderings_Otop(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Otop)).
fof(c_0_28, axiom, (![X21]:![X7]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X7,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X7,tc_HOL_Obool)),X21))=X21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__complement)).
fof(c_0_29, axiom, (~(hBOOL(c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fFalse__1)).
fof(c_0_30, axiom, (![X34]:![X28]:(class_Divides_Osemiring__div(X28)=>c_Divides_Odiv__class_Omod(X28,X34,c_Groups_Ozero__class_Ozero(X28))=X34)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_mod__by__0)).
fof(c_0_31, axiom, (![X233]:![X28]:hAPP(c_Fun_Oid(X28),X233)=X233), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_id__def)).
fof(c_0_32, conjecture, ((![X3]:(hBOOL(hAPP(hAPP(c_member(tc_Hoare__Mirabelle_Otriple(t_a)),X3),v_G))=>c_Hoare__Mirabelle_Otriple__valid(t_a,c_Groups_Ozero__class_Ozero(tc_Nat_Onat),X3))=>![X3]:(hBOOL(hAPP(hAPP(c_member(tc_Hoare__Mirabelle_Otriple(t_a)),X3),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(t_a),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a)),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),hAPP(hAPP(c_COMBB(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_fun(tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),tc_Com_Opname),c_Hoare__Mirabelle_Otriple_Otriple(t_a)),v_P)),c_Com_Ocom_OBODY)),v_Q)),v_Procs)))=>c_Hoare__Mirabelle_Otriple__valid(t_a,c_Groups_Ozero__class_Ozero(tc_Nat_Onat),X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_1)).
fof(c_0_33, axiom, (![X68]:![X4]:![X6]:![X36]:![X7]:![X22]:hAPP(hAPP(hAPP(c_COMBB(X22,X7,X36),X6),X4),X68)=hAPP(X6,hAPP(X4,X68))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBB__1)).
fof(c_0_34, axiom, (![X7]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X7,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X7,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__empty__eq)).
fof(c_0_35, axiom, (class_Divides_Osemiring__div(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Nat__Onat__Divides_Osemiring__div)).
fof(c_0_36, axiom, (![X27]:hAPP(c_Nat_Osize__class_Osize(tc_Nat_Onat),X27)=X27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_nat__size)).
fof(c_0_37, axiom, (![X21]:![X7]:hAPP(c_Set_OCollect(X7),hAPP(hAPP(c_COMBC(X7,tc_fun(X7,tc_HOL_Obool),tc_HOL_Obool),c_member(X7)),X21))=X21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__mem__eq)).
fof(c_0_38, axiom, (![X21]:![X7]:(?[X3]:hBOOL(hAPP(hAPP(c_member(X7),X3),X21))<=>X21!=c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_ex__in__conv)).
fof(c_0_39, axiom, (![X251]:![X28]:hAPP(c_List_Orev(X28),hAPP(c_List_Orev(X28),X251))=X251), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rev__rev__ident)).
fof(c_0_40, axiom, (![X105]:hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),X105)=X105), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_number__of__is__id)).
fof(c_0_41, axiom, (![X68]:![X4]:![X6]:![X7]:![X36]:![X22]:hAPP(hAPP(hAPP(c_COMBC(X22,X36,X7),X6),X4),X68)=hAPP(hAPP(X6,X68),X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBC__1)).
fof(c_0_42, axiom, (![X21]:![X2]:![X22]:![X7]:(c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))=hAPP(c_Set_Oimage(X22,X7,X2),X21)<=>X21=c_Orderings_Obot__class_Obot(tc_fun(X22,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__is__image)).
fof(c_0_43, axiom, (hBOOL(c_fTrue)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fTrue__1)).
fof(c_0_44, axiom, (![X27]:~(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X27),X27)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_less__not__refl)).
fof(c_0_45, axiom, (![X126]:(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X126))<=>?[X138]:![X3]:(hBOOL(hAPP(hAPP(c_member(tc_Nat_Onat),X3),X126))=>hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X3),X138))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite__nat__set__iff__bounded)).
fof(c_0_46, axiom, (![X21]:(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X21))<=>hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Int_Oint),hAPP(c_Set_Oimage(tc_Nat_Onat,tc_Int_Oint,c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint)),X21))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_transfer__nat__int__set__relations_I1_J)).
fof(c_0_47, axiom, (![X7]:hBOOL(hAPP(c_Finite__Set_Ofinite(X7),c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite_OemptyI)).
fof(c_0_48, axiom, (class_Finite__Set_Ofinite(tc_Product__Type_Ounit)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Product____Type__Ounit__Finite__Set_Ofinite)).
fof(c_0_49, plain, (![X6]:(hBOOL(X6)|hBOOL(hAPP(c_fNot,X6)))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_50, plain, (![X66]:![X67]:![X68]:(((hAPP(c_Groups_Ouminus__class_Ouminus(X68),X67)!=hAPP(c_Groups_Ouminus__class_Ouminus(X68),X66)|X67=X66)|~class_Lattices_Oboolean__algebra(X68))&((X67!=X66|hAPP(c_Groups_Ouminus__class_Ouminus(X68),X67)=hAPP(c_Groups_Ouminus__class_Ouminus(X68),X66))|~class_Lattices_Oboolean__algebra(X68)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
fof(c_0_51, plain, (![X29]:(~class_Lattices_Oboolean__algebra(X29)|hAPP(c_Groups_Ouminus__class_Ouminus(X29),c_Orderings_Obot__class_Obot(X29))=c_Orderings_Otop__class_Otop(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
fof(c_0_52, plain, (![X8]:c_Orderings_Obot__class_Obot(tc_fun(X8,tc_HOL_Obool))=hAPP(c_Set_OCollect(X8),hAPP(c_COMBK(tc_HOL_Obool,X8),c_fFalse))), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_53, plain, (![X8]:![X9]:hAPP(c_Set_OCollect(X9),X8)=X8), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_54, plain, (![X8]:c_Orderings_Otop__class_Otop(tc_fun(X8,tc_HOL_Obool))=hAPP(c_Set_OCollect(X8),hAPP(c_COMBK(tc_HOL_Obool,X8),c_fTrue))), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_55, plain, (![X22]:![X23]:(~class_Finite__Set_Ofinite(X23)|hBOOL(hAPP(c_Finite__Set_Ofinite(X23),X22)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
fof(c_0_56, plain, (![X312]:(~class_Finite__Set_Ofinite(X312)|class_Finite__Set_Ofinite(tc_Option_Ooption(X312)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
fof(c_0_57, plain, (![X98]:![X99]:![X100]:(((~hBOOL(hAPP(X98,X100))|hBOOL(hAPP(c_Enum_Oenum__class_Oenum__ex(X99),X98)))|~class_Enum_Oenum(X99))&((~hBOOL(hAPP(c_Enum_Oenum__class_Oenum__ex(X99),X98))|hBOOL(hAPP(X98,esk479_2(X98,X99))))|~class_Enum_Oenum(X99)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
fof(c_0_58, plain, (![X7]:(hBOOL(X7)|hBOOL(hAPP(c_fNot,X7)))), inference(variable_rename,[status(thm)],[c_0_49])).
cnf(c_0_59,plain,(X2=X3|~class_Lattices_Oboolean__algebra(X1)|hAPP(c_Groups_Ouminus__class_Ouminus(X1),X2)!=hAPP(c_Groups_Ouminus__class_Ouminus(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_60,plain,(class_Lattices_Oboolean__algebra(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_61,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(X1),c_Orderings_Obot__class_Obot(X1))=c_Orderings_Otop__class_Otop(X1)|~class_Lattices_Oboolean__algebra(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_62, plain, (![X32]:![X33]:(~class_Lattices_Oboolean__algebra(X33)|hAPP(c_Groups_Ouminus__class_Ouminus(X33),hAPP(c_Groups_Ouminus__class_Ouminus(X33),X32))=X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
fof(c_0_63, plain, (![X313]:![X314]:![X315]:![X316]:hAPP(hAPP(c_COMBK(X316,X315),X314),X313)=X314), inference(variable_rename,[status(thm)],[c_0_11])).
cnf(c_0_64,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_65,plain,(hAPP(c_Set_OCollect(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_66,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_67,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(X1),X2))|~class_Finite__Set_Ofinite(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_68,plain,(class_Finite__Set_Ofinite(tc_Option_Ooption(X1))|~class_Finite__Set_Ofinite(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_69, plain, (![X98]:![X99]:![X100]:![X102]:![X103]:![X104]:(((~hBOOL(hAPP(hAPP(c_member(X99),X100),c_Orderings_Otop__class_Otop(tc_fun(X99,tc_HOL_Obool))))|~hBOOL(hAPP(X98,X100)))|hBOOL(hAPP(X98,esk536_1(X98))))&((hBOOL(hAPP(hAPP(c_member(X104),esk537_2(X102,X104)),c_Orderings_Otop__class_Otop(tc_fun(X104,tc_HOL_Obool))))|~hBOOL(hAPP(X102,X103)))&(hBOOL(hAPP(X102,esk537_2(X102,X104)))|~hBOOL(hAPP(X102,X103)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])])).
fof(c_0_70, plain, (![X20]:![X21]:hBOOL(hAPP(hAPP(c_member(X21),X20),c_Orderings_Otop__class_Otop(tc_fun(X21,tc_HOL_Obool))))), inference(variable_rename,[status(thm)],[c_0_13])).
fof(c_0_71, plain, (![X6]:(~hBOOL(hAPP(c_fNot,X6))|~hBOOL(X6))), inference(fof_simplification,[status(thm)],[c_0_14])).
cnf(c_0_72,plain,(hBOOL(hAPP(c_Enum_Oenum__class_Oenum__ex(X1),X2))|~class_Enum_Oenum(X1)|~hBOOL(hAPP(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_73,plain,(hBOOL(hAPP(c_fNot,X1))|hBOOL(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_74, plain, (![X6]:![X7]:(c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))=hAPP(c_Set_OCollect(X7),X6)<=>![X3]:~hBOOL(hAPP(X6,X3)))), inference(fof_simplification,[status(thm)],[c_0_15])).
cnf(c_0_75,plain,(X1=X2|hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)!=hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X2)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_76,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_Orderings_Obot__class_Obot(tc_HOL_Obool))=c_Orderings_Otop__class_Otop(tc_HOL_Obool)), inference(spm,[status(thm)],[c_0_61, c_0_60])).
cnf(c_0_77,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(X1),hAPP(c_Groups_Ouminus__class_Ouminus(X1),X2))=X2|~class_Lattices_Oboolean__algebra(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_78, plain, (![X23]:![X24]:![X25]:(~class_Orderings_Obot(X25)|hAPP(c_Orderings_Obot__class_Obot(tc_fun(X24,X25)),X23)=c_Orderings_Obot__class_Obot(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).
cnf(c_0_79,plain,(hAPP(hAPP(c_COMBK(X1,X2),X3),X4)=X3), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_80,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_64, c_0_65])).
fof(c_0_81, plain, (![X23]:![X24]:![X25]:(~class_Orderings_Otop(X25)|hAPP(c_Orderings_Otop__class_Otop(tc_fun(X24,X25)),X23)=c_Orderings_Otop__class_Otop(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
cnf(c_0_82,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue)=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_66, c_0_65])).
fof(c_0_83, plain, (![X250]:![X251]:![X252]:![X253]:![X254]:![X255]:![X256]:((hAPP(c_List_Ofilter(X252,X251),X250)!=X250|(~hBOOL(hAPP(hAPP(c_member(X252),X253),hAPP(c_List_Oset(X252),X250)))|hBOOL(hAPP(X251,X253))))&((hBOOL(hAPP(hAPP(c_member(X256),esk454_3(X254,X255,X256)),hAPP(c_List_Oset(X256),X254)))|hAPP(c_List_Ofilter(X256,X255),X254)=X254)&(~hBOOL(hAPP(X255,esk454_3(X254,X255,X256)))|hAPP(c_List_Ofilter(X256,X255),X254)=X254)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])])])).
cnf(c_0_84,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Option_Ooption(X1)),X2))|~class_Finite__Set_Ofinite(X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_85,plain,(class_Finite__Set_Ofinite(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_86,plain,(hBOOL(hAPP(X1,esk536_1(X1)))|~hBOOL(hAPP(X1,X2))|~hBOOL(hAPP(hAPP(c_member(X3),X2),c_Orderings_Otop__class_Otop(tc_fun(X3,tc_HOL_Obool))))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_87,plain,(hBOOL(hAPP(hAPP(c_member(X1),X2),c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))))), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_88, plain, (![X41]:![X42]:hAPP(c_Set_Oimage(X42,X42,c_COMBI(X42)),X41)=X41), inference(variable_rename,[status(thm)],[c_0_20])).
fof(c_0_89, plain, (![X7]:(~hBOOL(hAPP(c_fNot,X7))|~hBOOL(X7))), inference(variable_rename,[status(thm)],[c_0_71])).
cnf(c_0_90,plain,(hBOOL(hAPP(c_Enum_Oenum__class_Oenum__ex(X1),c_fNot))|hBOOL(X2)|~class_Enum_Oenum(X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_91,plain,(class_Enum_Oenum(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_92, plain, (![X8]:![X9]:hAPP(c_Set_OCollect(X9),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X9),c_fNot),X8))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X9,tc_HOL_Obool)),hAPP(c_Set_OCollect(X9),X8))), inference(variable_rename,[status(thm)],[c_0_22])).
fof(c_0_93, plain, (![X8]:![X9]:![X10]:![X11]:![X13]:((c_Orderings_Obot__class_Obot(tc_fun(X9,tc_HOL_Obool))!=hAPP(c_Set_OCollect(X9),X8)|~hBOOL(hAPP(X8,X10)))&(hBOOL(hAPP(X11,esk157_1(X11)))|c_Orderings_Obot__class_Obot(tc_fun(X13,tc_HOL_Obool))=hAPP(c_Set_OCollect(X13),X11)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_74])])])])])).
fof(c_0_94, plain, (![X168]:![X169]:![X170]:![X171]:![X172]:![X173]:![X174]:![X175]:((((X168!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|hBOOL(hAPP(X170,X169)))|~hBOOL(hAPP(X170,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X169,X168))))&((X168=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X172),X168))|(X169!=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Nat_Onat),X168),X171)),X172)|hBOOL(hAPP(X170,X172)))))|~hBOOL(hAPP(X170,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X169,X168)))))&((((X173!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|X173=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))|hBOOL(hAPP(X175,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X174,X173))))&(((hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),esk236_3(X173,X174,X175)),X173))|X173=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))|hBOOL(hAPP(X175,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X174,X173))))&(((X174=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Nat_Onat),X173),esk235_3(X173,X174,X175))),esk236_3(X173,X174,X175))|X173=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))|hBOOL(hAPP(X175,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X174,X173))))&((~hBOOL(hAPP(X175,esk236_3(X173,X174,X175)))|X173=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))|hBOOL(hAPP(X175,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X174,X173)))))))&(((X173!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|~hBOOL(hAPP(X175,X174)))|hBOOL(hAPP(X175,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X174,X173))))&(((hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),esk236_3(X173,X174,X175)),X173))|~hBOOL(hAPP(X175,X174)))|hBOOL(hAPP(X175,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X174,X173))))&(((X174=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Nat_Onat),X173),esk235_3(X173,X174,X175))),esk236_3(X173,X174,X175))|~hBOOL(hAPP(X175,X174)))|hBOOL(hAPP(X175,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X174,X173))))&((~hBOOL(hAPP(X175,esk236_3(X173,X174,X175)))|~hBOOL(hAPP(X175,X174)))|hBOOL(hAPP(X175,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X174,X173)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])])])).
cnf(c_0_95,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_96,plain,(c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_97,plain,(c_Orderings_Obot__class_Obot(tc_HOL_Obool)=X1|c_Orderings_Otop__class_Otop(tc_HOL_Obool)!=hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_98,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1))=X1), inference(spm,[status(thm)],[c_0_77, c_0_60])).
cnf(c_0_99,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,X2)),X3)=c_Orderings_Obot__class_Obot(X2)|~class_Orderings_Obot(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_100,plain,(class_Orderings_Obot(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_101,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)),X2)=c_fFalse), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_102,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,X2)),X3)=c_Orderings_Otop__class_Otop(X2)|~class_Orderings_Otop(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_103,plain,(class_Orderings_Otop(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_104,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)),X2)=c_fTrue), inference(spm,[status(thm)],[c_0_79, c_0_82])).
cnf(c_0_105,plain,(hAPP(c_List_Ofilter(X1,X2),X3)=X3|~hBOOL(hAPP(X2,esk454_3(X3,X2,X1)))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_106,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Option_Ooption(tc_HOL_Obool)),X1))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_107,plain,(hBOOL(hAPP(X1,esk536_1(X1)))|~hBOOL(hAPP(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_108,plain,(hAPP(c_Set_Oimage(X1,X1,c_COMBI(X1)),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_109,plain,(~hBOOL(X1)|~hBOOL(hAPP(c_fNot,X1))), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_110,plain,(hBOOL(hAPP(c_Enum_Oenum__class_Oenum__ex(tc_HOL_Obool),c_fNot))|hBOOL(X1)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
fof(c_0_111, plain, (![X22]:![X23]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X23,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X23,tc_HOL_Obool)),X22))=X22), inference(variable_rename,[status(thm)],[c_0_28])).
cnf(c_0_112,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Set_OCollect(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_113,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),X2)|hBOOL(hAPP(X2,esk157_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_114,plain,(hBOOL(hAPP(X1,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X2,X3)))|~hBOOL(hAPP(X1,X2))|X3!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_115,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_116,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)=c_Orderings_Obot__class_Obot(tc_HOL_Obool)|c_Orderings_Otop__class_Otop(tc_HOL_Obool)!=X1), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_117,plain,(c_Orderings_Obot__class_Obot(tc_HOL_Obool)=c_fFalse), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])).
cnf(c_0_118,plain,(c_Orderings_Otop__class_Otop(tc_HOL_Obool)=c_fTrue), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104])).
fof(c_0_119, plain, (~hBOOL(c_fFalse)), inference(fof_simplification,[status(thm)],[c_0_29])).
fof(c_0_120, plain, (![X35]:![X36]:(~class_Divides_Osemiring__div(X36)|c_Divides_Odiv__class_Omod(X36,X35,c_Groups_Ozero__class_Ozero(X36))=X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])).
fof(c_0_121, plain, (![X234]:![X235]:hAPP(c_Fun_Oid(X235),X234)=X234), inference(variable_rename,[status(thm)],[c_0_31])).
cnf(c_0_122,plain,(hAPP(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_Option_Ooption(tc_HOL_Obool))),X2)=X2), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_123,plain,(hBOOL(esk536_1(c_Set_Oimage(X1,X1,c_COMBI(X1))))|~hBOOL(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_108])).
cnf(c_0_124,plain,(hBOOL(hAPP(c_Enum_Oenum__class_Oenum__ex(tc_HOL_Obool),c_fNot))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_110])).
fof(c_0_125, negated_conjecture, (~((![X3]:(hBOOL(hAPP(hAPP(c_member(tc_Hoare__Mirabelle_Otriple(t_a)),X3),v_G))=>c_Hoare__Mirabelle_Otriple__valid(t_a,c_Groups_Ozero__class_Ozero(tc_Nat_Onat),X3))=>![X3]:(hBOOL(hAPP(hAPP(c_member(tc_Hoare__Mirabelle_Otriple(t_a)),X3),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(t_a),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a)),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),hAPP(hAPP(c_COMBB(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_fun(tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),tc_Com_Opname),c_Hoare__Mirabelle_Otriple_Otriple(t_a)),v_P)),c_Com_Ocom_OBODY)),v_Q)),v_Procs)))=>c_Hoare__Mirabelle_Otriple__valid(t_a,c_Groups_Ozero__class_Ozero(tc_Nat_Onat),X3))))), inference(assume_negation,[status(cth)],[c_0_32])).
fof(c_0_126, plain, (![X69]:![X70]:![X71]:![X72]:![X73]:![X74]:hAPP(hAPP(hAPP(c_COMBB(X74,X73,X72),X71),X70),X69)=hAPP(X71,hAPP(X70,X69))), inference(variable_rename,[status(thm)],[c_0_33])).
cnf(c_0_127,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_128,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2)=hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_65]), c_0_65])).
cnf(c_0_129,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hBOOL(hAPP(X1,esk157_1(X1)))), inference(rw,[status(thm)],[c_0_113, c_0_65])).
fof(c_0_130, plain, (![X8]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X8,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X8,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X8,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_34])).
cnf(c_0_131,plain,(hBOOL(hAPP(X1,c_Divides_Odiv__class_Omod(tc_Nat_Onat,X2,X3)))|hAPP(c_Int_Onat,c_Int_OPls)!=X3|~hBOOL(hAPP(X1,X2))), inference(rw,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_132,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)=c_fFalse|c_fTrue!=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_117]), c_0_118])).
cnf(c_0_133,plain,(~hBOOL(c_fFalse)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_134,plain,(c_Divides_Odiv__class_Omod(X1,X2,c_Groups_Ozero__class_Ozero(X1))=X2|~class_Divides_Osemiring__div(X1)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_135,plain,(class_Divides_Osemiring__div(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_136, plain, (![X28]:hAPP(c_Nat_Osize__class_Osize(tc_Nat_Onat),X28)=X28), inference(variable_rename,[status(thm)],[c_0_36])).
cnf(c_0_137,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_HOL_Obool),X1))), inference(spm,[status(thm)],[c_0_67, c_0_85])).
cnf(c_0_138,plain,(hBOOL(hAPP(X1,esk537_2(X1,X3)))|~hBOOL(hAPP(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_139,plain,(hAPP(c_Fun_Oid(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_140,plain,(hBOOL(esk536_1(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_Option_Ooption(tc_HOL_Obool)))))|~hBOOL(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_122]), c_0_122])).
cnf(c_0_141,plain,(hBOOL(esk536_1(c_Set_Oimage(X1,X1,c_COMBI(X1))))), inference(spm,[status(thm)],[c_0_123, c_0_124])).
fof(c_0_142, plain, (![X22]:![X23]:hAPP(c_Set_OCollect(X23),hAPP(hAPP(c_COMBC(X23,tc_fun(X23,tc_HOL_Obool),tc_HOL_Obool),c_member(X23)),X22))=X22), inference(variable_rename,[status(thm)],[c_0_37])).
fof(c_0_143, plain, (![X22]:![X23]:![X24]:![X25]:![X26]:((~hBOOL(hAPP(hAPP(c_member(X23),X24),X22))|X22!=c_Orderings_Obot__class_Obot(tc_fun(X23,tc_HOL_Obool)))&(X25=c_Orderings_Obot__class_Obot(tc_fun(X26,tc_HOL_Obool))|hBOOL(hAPP(hAPP(c_member(X26),esk17_2(X25,X26)),X25))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])])])).
fof(c_0_144, negated_conjecture, (![X4]:((~hBOOL(hAPP(hAPP(c_member(tc_Hoare__Mirabelle_Otriple(t_a)),X4),v_G))|c_Hoare__Mirabelle_Otriple__valid(t_a,c_Groups_Ozero__class_Ozero(tc_Nat_Onat),X4))&(hBOOL(hAPP(hAPP(c_member(tc_Hoare__Mirabelle_Otriple(t_a)),esk563_0),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(t_a),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a)),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),hAPP(hAPP(c_COMBB(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_fun(tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),tc_Com_Opname),c_Hoare__Mirabelle_Otriple_Otriple(t_a)),v_P)),c_Com_Ocom_OBODY)),v_Q)),v_Procs)))&~c_Hoare__Mirabelle_Otriple__valid(t_a,c_Groups_Ozero__class_Ozero(tc_Nat_Onat),esk563_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_125])])])])).
cnf(c_0_145,plain,(hAPP(hAPP(hAPP(c_COMBB(X1,X2,X3),X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_146,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_128]), c_0_128])).
fof(c_0_147, plain, (![X252]:![X253]:hAPP(c_List_Orev(X253),hAPP(c_List_Orev(X253),X252))=X252), inference(variable_rename,[status(thm)],[c_0_39])).
cnf(c_0_148,plain,(c_fFalse=X1|hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_129]), c_0_101]), c_0_79])).
cnf(c_0_149,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_150,plain,(c_Divides_Odiv__class_Omod(tc_Nat_Onat,X1,X2)!=c_fTrue|hAPP(c_Int_Onat,c_Int_OPls)!=X2|~hBOOL(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_133])).
cnf(c_0_151,plain,(c_Divides_Odiv__class_Omod(tc_Nat_Onat,X1,hAPP(c_Int_Onat,c_Int_OPls))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_115])).
cnf(c_0_152,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_117]), c_0_118])).
fof(c_0_153, plain, (![X106]:hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),X106)=X106), inference(variable_rename,[status(thm)],[c_0_40])).
cnf(c_0_154,plain,(hAPP(c_Nat_Osize__class_Osize(tc_Nat_Onat),X1)=X1), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_155,plain,(hAPP(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_HOL_Obool)),X2)=X2), inference(spm,[status(thm)],[c_0_105, c_0_137])).
cnf(c_0_156,plain,(hBOOL(esk537_2(c_Fun_Oid(X1),X2))|~hBOOL(X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_139])).
cnf(c_0_157,plain,(hBOOL(esk536_1(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_Option_Ooption(tc_HOL_Obool)))))), inference(spm,[status(thm)],[c_0_140, c_0_141])).
fof(c_0_158, plain, (![X69]:![X70]:![X71]:![X72]:![X73]:![X74]:hAPP(hAPP(hAPP(c_COMBC(X74,X73,X72),X71),X70),X69)=hAPP(hAPP(X71,X69),X70)), inference(variable_rename,[status(thm)],[c_0_41])).
cnf(c_0_159,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBC(X1,tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool),c_member(X1)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_160,plain,(X1!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|~hBOOL(hAPP(hAPP(c_member(X2),X3),X1))), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_161,negated_conjecture,(hBOOL(hAPP(hAPP(c_member(tc_Hoare__Mirabelle_Otriple(t_a)),esk563_0),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(t_a),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a)),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),hAPP(hAPP(c_COMBB(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_fun(tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),tc_Com_Opname),c_Hoare__Mirabelle_Otriple_Otriple(t_a)),v_P)),c_Com_Ocom_OBODY)),v_Q)),v_Procs)))), inference(split_conjunct,[status(thm)],[c_0_144])).
fof(c_0_162, plain, (![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:((c_Orderings_Obot__class_Obot(tc_fun(X26,tc_HOL_Obool))!=hAPP(c_Set_Oimage(X25,X26,X24),X23)|X23=c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool)))&(X27!=c_Orderings_Obot__class_Obot(tc_fun(X29,tc_HOL_Obool))|c_Orderings_Obot__class_Obot(tc_fun(X30,tc_HOL_Obool))=hAPP(c_Set_Oimage(X29,X30,X28),X27)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_42])])])])).
cnf(c_0_163,plain,(hAPP(c_fNot,hAPP(c_fNot,hAPP(X1,X2)))=hAPP(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_145])).
cnf(c_0_164,plain,(hAPP(c_List_Orev(X1),hAPP(c_List_Orev(X1),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_165,plain,(hAPP(c_fNot,X1)=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_109, c_0_148])).
cnf(c_0_166,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_149, c_0_128])).
cnf(c_0_167,plain,(X1!=c_fTrue|~hBOOL(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1))), inference(spm,[status(thm)],[c_0_150, c_0_151])).
cnf(c_0_168,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)=X2|X1!=hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X2)), inference(spm,[status(thm)],[c_0_75, c_0_98])).
cnf(c_0_169,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fTrue)=c_fFalse), inference(spm,[status(thm)],[c_0_98, c_0_152])).
cnf(c_0_170,plain,(hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),X1)=X1), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_171,plain,(hBOOL(esk536_1(c_Nat_Osize__class_Osize(tc_Nat_Onat)))|~hBOOL(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_154]), c_0_154])).
cnf(c_0_172,plain,(hBOOL(c_fTrue)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_173,plain,(hBOOL(esk537_2(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_HOL_Obool)),X2))|~hBOOL(X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_155]), c_0_155])).
cnf(c_0_174,plain,(hBOOL(esk537_2(c_Fun_Oid(X1),X2))), inference(spm,[status(thm)],[c_0_156, c_0_157])).
fof(c_0_175, plain, (![X27]:~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X27),X27))), inference(fof_simplification,[status(thm)],[c_0_44])).
fof(c_0_176, plain, (![X139]:![X141]:![X142]:![X143]:((~hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X139))|(~hBOOL(hAPP(hAPP(c_member(tc_Nat_Onat),X141),X139))|hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X141),esk85_1(X139)))))&((hBOOL(hAPP(hAPP(c_member(tc_Nat_Onat),esk86_2(X142,X143)),X142))|hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X142)))&(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),esk86_2(X142,X143)),X143))|hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X142)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_45])])])])])])).
cnf(c_0_177,plain,(hAPP(hAPP(hAPP(c_COMBC(X1,X2,X3),X4),X5),X6)=hAPP(hAPP(X4,X6),X5)), inference(split_conjunct,[status(thm)],[c_0_158])).
cnf(c_0_178,plain,(hAPP(hAPP(c_COMBC(X1,tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool),c_member(X1)),X2)=X2), inference(rw,[status(thm)],[c_0_159, c_0_65])).
cnf(c_0_179,negated_conjecture,(hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(t_a),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a)),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),hAPP(hAPP(c_COMBB(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_fun(tc_Com_Ocom,tc_fun(tc_fun(t_a,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(t_a))),tc_Com_Opname),c_Hoare__Mirabelle_Otriple_Otriple(t_a)),v_P)),c_Com_Ocom_OBODY)),v_Q)),v_Procs)!=c_Orderings_Obot__class_Obot(tc_fun(tc_Hoare__Mirabelle_Otriple(t_a),tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_180,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_Oimage(X2,X1,X3),X4)|X4!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_181,plain,(hAPP(c_fNot,hAPP(c_fNot,X1))=X1), inference(spm,[status(thm)],[c_0_163, c_0_164])).
cnf(c_0_182,plain,(hAPP(c_fNot,X1)=c_fFalse|c_fFalse=X1), inference(spm,[status(thm)],[c_0_165, c_0_148])).
cnf(c_0_183,plain,(hAPP(c_fNot,c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_166]), c_0_104]), c_0_101])).
cnf(c_0_184,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)!=c_fTrue|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_167, c_0_98])).
cnf(c_0_185,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1)=c_fTrue|X1!=c_fFalse), inference(spm,[status(thm)],[c_0_168, c_0_169])).
cnf(c_0_186,plain,(hBOOL(esk536_1(c_Int_Onumber__class_Onumber__of(tc_Int_Oint)))|~hBOOL(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_170]), c_0_170])).
cnf(c_0_187,plain,(hBOOL(esk536_1(c_Nat_Osize__class_Osize(tc_Nat_Onat)))), inference(spm,[status(thm)],[c_0_171, c_0_172])).
cnf(c_0_188,plain,(hBOOL(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),esk536_1(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool))))|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_107, c_0_98])).
cnf(c_0_189,plain,(hBOOL(esk537_2(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_HOL_Obool)),X2))), inference(spm,[status(thm)],[c_0_173, c_0_174])).
fof(c_0_190, plain, (![X28]:~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X28),X28))), inference(variable_rename,[status(thm)],[c_0_175])).
cnf(c_0_191,plain,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X1),esk85_1(X2)))|~hBOOL(hAPP(hAPP(c_member(tc_Nat_Onat),X1),X2))|~hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X2))), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_192,plain,(hAPP(hAPP(c_member(X1),X2),X3)=hAPP(X3,X2)), inference(spm,[status(thm)],[c_0_177, c_0_178])).
cnf(c_0_193,plain,(~hBOOL(hAPP(X1,X2))|c_Orderings_Obot__class_Obot(tc_fun(X3,tc_HOL_Obool))!=hAPP(c_Set_OCollect(X3),X1)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_194,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(tc_Com_Opname,tc_HOL_Obool))!=v_Procs), inference(spm,[status(thm)],[c_0_179, c_0_180])).
cnf(c_0_195,plain,(c_fFalse=X1|c_fTrue=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182]), c_0_183])).
cnf(c_0_196,plain,(X1!=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_184, c_0_185])).
cnf(c_0_197,plain,(hBOOL(esk536_1(c_Int_Onumber__class_Onumber__of(tc_Int_Oint)))), inference(spm,[status(thm)],[c_0_186, c_0_187])).
cnf(c_0_198,plain,(hBOOL(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),esk536_1(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool))))), inference(spm,[status(thm)],[c_0_188, c_0_189])).
cnf(c_0_199,plain,(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X1),X1))), inference(split_conjunct,[status(thm)],[c_0_190])).
cnf(c_0_200,plain,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X1),esk85_1(X2)))|~hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X2))|~hBOOL(hAPP(X2,X1))), inference(rw,[status(thm)],[c_0_191, c_0_192])).
fof(c_0_201, plain, (![X22]:![X23]:((~hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X22))|hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Int_Oint),hAPP(c_Set_Oimage(tc_Nat_Onat,tc_Int_Oint,c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint)),X22))))&(~hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Int_Oint),hAPP(c_Set_Oimage(tc_Nat_Onat,tc_Int_Oint,c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint)),X23)))|hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X23))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_46])])])])).
fof(c_0_202, plain, (![X8]:hBOOL(hAPP(c_Finite__Set_Ofinite(X8),c_Orderings_Obot__class_Obot(tc_fun(X8,tc_HOL_Obool))))), inference(variable_rename,[status(thm)],[c_0_47])).
cnf(c_0_203,plain,(X1!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|~hBOOL(hAPP(X1,X3))), inference(rw,[status(thm)],[c_0_193, c_0_65])).
cnf(c_0_204,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(tc_Com_Opname,tc_HOL_Obool))=c_fFalse|v_Procs!=c_fTrue), inference(spm,[status(thm)],[c_0_194, c_0_195])).
cnf(c_0_205,plain,(esk536_1(c_Int_Onumber__class_Onumber__of(tc_Int_Oint))!=c_fFalse), inference(spm,[status(thm)],[c_0_196, c_0_197])).
cnf(c_0_206,plain,(esk536_1(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool))=c_fFalse), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_195]), c_0_169]), c_0_133])).
cnf(c_0_207,plain,(class_Finite__Set_Ofinite(tc_Product__Type_Ounit)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_208,plain,(~hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X1))|~hBOOL(hAPP(X1,esk85_1(X1)))), inference(spm,[status(thm)],[c_0_199, c_0_200])).
cnf(c_0_209,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X1))|~hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Int_Oint),hAPP(c_Set_Oimage(tc_Nat_Onat,tc_Int_Oint,c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint)),X1)))), inference(split_conjunct,[status(thm)],[c_0_201])).
cnf(c_0_210,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(X1),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))))), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_211,negated_conjecture,(v_Procs!=c_fTrue|X1!=c_fFalse|~hBOOL(hAPP(X1,X2))), inference(spm,[status(thm)],[c_0_203, c_0_204])).
cnf(c_0_212,plain,(c_Int_Onumber__class_Onumber__of(tc_Int_Oint)=c_fFalse|esk536_1(c_fTrue)!=c_fFalse), inference(spm,[status(thm)],[c_0_205, c_0_195])).
cnf(c_0_213,plain,(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool)=c_fFalse|esk536_1(c_fTrue)=c_fFalse), inference(spm,[status(thm)],[c_0_206, c_0_195])).
cnf(c_0_214,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Product__Type_Ounit),X1))), inference(spm,[status(thm)],[c_0_67, c_0_207])).
cnf(c_0_215,plain,(~hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),c_Finite__Set_Ofinite(tc_HOL_Obool)))), inference(spm,[status(thm)],[c_0_208, c_0_137])).
cnf(c_0_216,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),X1))|c_Orderings_Obot__class_Obot(tc_fun(tc_Nat_Onat,tc_HOL_Obool))!=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_180]), c_0_210])])).
cnf(c_0_217,negated_conjecture,(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool)!=c_fFalse|v_Procs!=c_fTrue), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_152]), c_0_172])])).
cnf(c_0_218,plain,(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool)=c_fFalse|c_Int_Onumber__class_Onumber__of(tc_Int_Oint)=c_fFalse), inference(spm,[status(thm)],[c_0_212, c_0_213])).
cnf(c_0_219,plain,(hAPP(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_Product__Type_Ounit)),X2)=X2), inference(spm,[status(thm)],[c_0_105, c_0_214])).
cnf(c_0_220,plain,(c_Orderings_Obot__class_Obot(tc_fun(tc_Nat_Onat,tc_HOL_Obool))!=c_Finite__Set_Ofinite(tc_HOL_Obool)), inference(spm,[status(thm)],[c_0_215, c_0_216])).
cnf(c_0_221,negated_conjecture,(c_Int_Onumber__class_Onumber__of(tc_Int_Oint)!=c_fFalse|v_Procs!=c_fTrue|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_211, c_0_170])).
cnf(c_0_222,negated_conjecture,(c_Int_Onumber__class_Onumber__of(tc_Int_Oint)=c_fFalse|v_Procs!=c_fTrue), inference(spm,[status(thm)],[c_0_217, c_0_218])).
cnf(c_0_223,plain,(hBOOL(esk537_2(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_Product__Type_Ounit)),X2))|~hBOOL(X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_219]), c_0_219])).
cnf(c_0_224,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))!=v_Procs), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_129]), c_0_101]), c_0_133])).
cnf(c_0_225,plain,(c_Orderings_Obot__class_Obot(tc_fun(tc_Nat_Onat,tc_HOL_Obool))=c_fFalse|c_Finite__Set_Ofinite(tc_HOL_Obool)!=c_fTrue), inference(spm,[status(thm)],[c_0_220, c_0_195])).
cnf(c_0_226,negated_conjecture,(v_Procs!=c_fTrue|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_221, c_0_222])).
cnf(c_0_227,plain,(hBOOL(esk537_2(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_Product__Type_Ounit)),X2))), inference(spm,[status(thm)],[c_0_223, c_0_189])).
cnf(c_0_228,negated_conjecture,(c_Finite__Set_Ofinite(tc_HOL_Obool)!=c_fTrue|v_Procs!=c_fFalse), inference(spm,[status(thm)],[c_0_224, c_0_225])).
cnf(c_0_229,negated_conjecture,(v_Procs!=c_fTrue), inference(spm,[status(thm)],[c_0_226, c_0_227])).
cnf(c_0_230,plain,(hAPP(c_Finite__Set_Ofinite(tc_HOL_Obool),X1)!=c_fFalse), inference(spm,[status(thm)],[c_0_196, c_0_137])).
cnf(c_0_231,negated_conjecture,(c_Finite__Set_Ofinite(tc_HOL_Obool)=c_fFalse|v_Procs!=c_fFalse), inference(spm,[status(thm)],[c_0_228, c_0_195])).
cnf(c_0_232,negated_conjecture,(v_Procs=c_fFalse), inference(spm,[status(thm)],[c_0_229, c_0_195])).
cnf(c_0_233,negated_conjecture,(hAPP(c_fFalse,X1)!=c_fFalse|v_Procs!=c_fFalse), inference(spm,[status(thm)],[c_0_230, c_0_231])).
cnf(c_0_234,plain,(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),c_Finite__Set_Ofinite(tc_HOL_Obool))=c_fFalse), inference(spm,[status(thm)],[c_0_215, c_0_148])).
cnf(c_0_235,negated_conjecture,(c_Finite__Set_Ofinite(tc_HOL_Obool)=c_fFalse), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_231, c_0_232])])).
cnf(c_0_236,plain,(hAPP(c_fTrue,c_fFalse)=c_fTrue|c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool)=c_fFalse), inference(spm,[status(thm)],[c_0_152, c_0_195])).
cnf(c_0_237,negated_conjecture,(hAPP(c_fFalse,X1)!=c_fFalse), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_233, c_0_232])])).
cnf(c_0_238,plain,(hAPP(c_Finite__Set_Ofinite(tc_Nat_Onat),c_fFalse)=c_fFalse), inference(rw,[status(thm)],[c_0_234, c_0_235])).
cnf(c_0_239,plain,(hAPP(c_fTrue,c_fFalse)=c_fTrue), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_236]), c_0_237])).
cnf(c_0_240,plain,(c_fTrue!=c_fFalse), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_152]), c_0_172])])).
cnf(c_0_241,plain,(c_Finite__Set_Ofinite(tc_Nat_Onat)=c_fFalse), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_195]), c_0_239]), c_0_240])).
cnf(c_0_242,plain,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_238, c_0_241]), c_0_237]), ['proof']).
# SZS output end CNFRefutation.
EOF
