% Problem    : NUM925+2 : TPTP v6.1.0. Released v5.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n153.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:28:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___042_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X21]:number_number_of_int(X21)=X21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_176_number__of__is__id)).
fof(c_0_1, conjecture, (hAPP_nat_int(power_power_int(hAPP_int_int(plus_plus_int(one_one_int),hAPP_nat_int(semiri1621563631at_int,n))),number_number_of_nat(bit0(bit1(pls))))!=zero_zero_int), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_2, axiom, (![X21]:succ(bit0(X21))=bit1(X21)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_568_succ__Bit0)).
fof(c_0_3, axiom, (![X21]:bit0(X21)=hAPP_int_int(plus_plus_int(X21),X21)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_103_Bit0__def)).
fof(c_0_4, axiom, (![X21]:succ(X21)=hAPP_int_int(plus_plus_int(X21),one_one_int)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_573_succ__def)).
fof(c_0_5, axiom, (![X17]:hAPP_int_int(plus_plus_int(zero_zero_int),X17)=X17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_105_zadd__0)).
fof(c_0_6, axiom, (zero_zero_int=number_number_of_int(pls)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_35_zero__is__num__zero)).
fof(c_0_7, axiom, (![X78]:![X79]:(hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,zero_zero_int),X79))=>hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,zero_zero_int),hAPP_nat_int(power_power_int(X79),X78))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_267_zero__less__power)).
fof(c_0_8, axiom, (~(hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,pls),zero_zero_int)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_79_bin__less__0__simps_I1_J)).
fof(c_0_9, axiom, (hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,zero_zero_int),hAPP_int_int(plus_plus_int(one_one_int),hAPP_nat_int(semiri1621563631at_int,n))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_0_n1pos)).
fof(c_0_10, plain, (![X22]:number_number_of_int(X22)=X22), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_11, negated_conjecture, (hAPP_nat_int(power_power_int(hAPP_int_int(plus_plus_int(one_one_int),hAPP_nat_int(semiri1621563631at_int,n))),number_number_of_nat(bit0(bit1(pls))))=zero_zero_int), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_1])])).
fof(c_0_12, plain, (![X22]:succ(bit0(X22))=bit1(X22)), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_13, plain, (![X22]:bit0(X22)=hAPP_int_int(plus_plus_int(X22),X22)), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_14, plain, (![X22]:succ(X22)=hAPP_int_int(plus_plus_int(X22),one_one_int)), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_15, plain, (![X18]:hAPP_int_int(plus_plus_int(zero_zero_int),X18)=X18), inference(variable_rename,[status(thm)],[c_0_5])).
cnf(c_0_16,plain,(zero_zero_int=number_number_of_int(pls)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17,plain,(number_number_of_int(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_18, plain, (![X80]:![X81]:(~hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,zero_zero_int),X81))|hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,zero_zero_int),hAPP_nat_int(power_power_int(X81),X80))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
cnf(c_0_19,negated_conjecture,(hAPP_nat_int(power_power_int(hAPP_int_int(plus_plus_int(one_one_int),hAPP_nat_int(semiri1621563631at_int,n))),number_number_of_nat(bit0(bit1(pls))))=zero_zero_int), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20,plain,(succ(bit0(X1))=bit1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21,plain,(bit0(X1)=hAPP_int_int(plus_plus_int(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,plain,(succ(X1)=hAPP_int_int(plus_plus_int(X1),one_one_int)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23,plain,(hAPP_int_int(plus_plus_int(zero_zero_int),X1)=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24,plain,(pls=zero_zero_int), inference(rw,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_25, plain, (~hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,pls),zero_zero_int))), inference(fof_simplification,[status(thm)],[c_0_8])).
cnf(c_0_26,plain,(hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,zero_zero_int),hAPP_nat_int(power_power_int(X1),X2)))|~hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,zero_zero_int),X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,negated_conjecture,(hAPP_nat_int(power_power_int(hAPP_int_int(plus_plus_int(one_one_int),hAPP_nat_int(semiri1621563631at_int,n))),number_number_of_nat(hAPP_int_int(plus_plus_int(hAPP_int_int(plus_plus_int(hAPP_int_int(plus_plus_int(pls),pls)),one_one_int)),hAPP_int_int(plus_plus_int(hAPP_int_int(plus_plus_int(pls),pls)),one_one_int))))=zero_zero_int), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_21]), c_0_22]), c_0_22])).
cnf(c_0_28,plain,(hAPP_int_int(plus_plus_int(pls),X1)=X1), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29,plain,(hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,zero_zero_int),hAPP_int_int(plus_plus_int(one_one_int),hAPP_nat_int(semiri1621563631at_int,n))))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_30,plain,(~hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,pls),zero_zero_int))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31,plain,(hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,pls),hAPP_nat_int(power_power_int(X1),X2)))|~hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,pls),X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_24]), c_0_24])).
cnf(c_0_32,negated_conjecture,(hAPP_nat_int(power_power_int(hAPP_int_int(plus_plus_int(one_one_int),hAPP_nat_int(semiri1621563631at_int,n))),number_number_of_nat(hAPP_int_int(plus_plus_int(one_one_int),one_one_int)))=pls), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_28]), c_0_28]), c_0_28]), c_0_24])).
cnf(c_0_33,plain,(hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,pls),hAPP_int_int(plus_plus_int(one_one_int),hAPP_nat_int(semiri1621563631at_int,n))))), inference(rw,[status(thm)],[c_0_29, c_0_24])).
cnf(c_0_34,plain,(~hBOOL(hAPP_int_bool(hAPP_i1948725293t_bool(ord_less_int,pls),pls))), inference(rw,[status(thm)],[c_0_30, c_0_24])).
cnf(c_0_35,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])]), c_0_34]), ['proof']).
# SZS output end CNFRefutation.
EOF
