% Problem    : SWV398+1 : TPTP v6.1.0. Released v3.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n131.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:55:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_C45_F1_SE_CS_SP_PI_CO_S0V
# and selection function PSelectComplexExceptRRHorn.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:![X2]:![X3]:![X4]:(((X2!=create_slb&contains_slb(X2,findmin_pqp_res(X1)))&strictly_less_than(findmin_pqp_res(X1),lookup_slb(X2,findmin_pqp_res(X1))))=>findmin_cpq_eff(triple(X1,X2,X3))=triple(X1,update_slb(X2,findmin_pqp_res(X1)),bad))), file('/export/starexec/sandbox/benchmark/Axioms/SWV007+3.ax', ax48)).
fof(c_0_1, axiom, (![X1]:![X2]:![X3]:![X4]:((X2!=create_slb&~(contains_slb(X2,findmin_pqp_res(X1))))=>findmin_cpq_eff(triple(X1,X2,X3))=triple(X1,update_slb(X2,findmin_pqp_res(X1)),bad))), file('/export/starexec/sandbox/benchmark/Axioms/SWV007+3.ax', ax47)).
fof(c_0_2, axiom, (![X1]:![X2]:(strictly_less_than(X1,X2)<=>(less_than(X1,X2)&~(less_than(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV007+0.ax', stricly_smaller_definition)).
fof(c_0_3, axiom, (![X1]:![X2]:(less_than(X1,X2)|less_than(X2,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV007+0.ax', totality)).
fof(c_0_4, conjecture, (![X1]:![X2]:![X3]:(~(check_cpq(triple(X1,X2,X3)))=>~(check_cpq(findmin_cpq_eff(triple(X1,X2,X3)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', l24_co)).
fof(c_0_5, lemma, (![X1]:![X2]:![X3]:(check_cpq(triple(X1,X2,X3))<=>![X4]:![X5]:(pair_in_list(X2,X4,X5)=>less_than(X5,X4)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', l26_li4142)).
fof(c_0_6, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:(less_than(X5,X4)=>(check_cpq(triple(X1,insert_slb(X2,pair(X4,X5)),X3))<=>check_cpq(triple(X1,X2,X3))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV007+3.ax', ax37)).
fof(c_0_7, axiom, (![X1]:![X2]:![X3]:![X4]:(((X2!=create_slb&contains_slb(X2,findmin_pqp_res(X1)))&less_than(lookup_slb(X2,findmin_pqp_res(X1)),findmin_pqp_res(X1)))=>findmin_cpq_eff(triple(X1,X2,X3))=triple(X1,update_slb(X2,findmin_pqp_res(X1)),X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV007+3.ax', ax49)).
fof(c_0_8, axiom, (![X1]:![X2]:![X3]:![X4]:insert_cpq(triple(X1,X2,X3),X4)=triple(insert_pqp(X1,X4),insert_slb(X2,pair(X4,bottom)),X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWV007+3.ax', ax42)).
fof(c_0_9, axiom, (![X1]:less_than(bottom,X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWV007+0.ax', bottom_smallest)).
fof(c_0_10, lemma, (![X1]:(?[X2]:?[X3]:(pair_in_list(X1,X2,X3)&strictly_less_than(X2,X3))=>![X4]:?[X5]:?[X6]:(pair_in_list(update_slb(X1,X4),X5,X6)&strictly_less_than(X5,X6)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', l34_l35)).
fof(c_0_11, axiom, (![X1]:![X2]:check_cpq(triple(X1,create_slb,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWV007+3.ax', ax36)).
fof(c_0_12, plain, (![X1]:![X2]:![X3]:(((X2!=create_slb&contains_slb(X2,findmin_pqp_res(X1)))&strictly_less_than(findmin_pqp_res(X1),lookup_slb(X2,findmin_pqp_res(X1))))=>findmin_cpq_eff(triple(X1,X2,X3))=triple(X1,update_slb(X2,findmin_pqp_res(X1)),bad))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_13, plain, (![X1]:![X2]:![X3]:((X2!=create_slb&~contains_slb(X2,findmin_pqp_res(X1)))=>findmin_cpq_eff(triple(X1,X2,X3))=triple(X1,update_slb(X2,findmin_pqp_res(X1)),bad))), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_14, plain, (![X1]:![X2]:(strictly_less_than(X1,X2)<=>(less_than(X1,X2)&~less_than(X2,X1)))), inference(fof_simplification,[status(thm)],[c_0_2])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:(((X5=create_slb|~contains_slb(X5,findmin_pqp_res(X4)))|~strictly_less_than(findmin_pqp_res(X4),lookup_slb(X5,findmin_pqp_res(X4))))|findmin_cpq_eff(triple(X4,X5,X6))=triple(X4,update_slb(X5,findmin_pqp_res(X4)),bad))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:((X5=create_slb|contains_slb(X5,findmin_pqp_res(X4)))|findmin_cpq_eff(triple(X4,X5,X6))=triple(X4,update_slb(X5,findmin_pqp_res(X4)),bad))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
fof(c_0_17, plain, (![X3]:![X4]:![X5]:![X6]:(((less_than(X3,X4)|~strictly_less_than(X3,X4))&(~less_than(X4,X3)|~strictly_less_than(X3,X4)))&((~less_than(X5,X6)|less_than(X6,X5))|strictly_less_than(X5,X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])])).
fof(c_0_18, plain, (![X3]:![X4]:(less_than(X3,X4)|less_than(X4,X3))), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_19, negated_conjecture, (~(![X1]:![X2]:![X3]:(~check_cpq(triple(X1,X2,X3))=>~check_cpq(findmin_cpq_eff(triple(X1,X2,X3)))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_4])])).
cnf(c_0_20,plain,(findmin_cpq_eff(triple(X1,X2,X3))=triple(X1,update_slb(X2,findmin_pqp_res(X1)),bad)|X2=create_slb|~strictly_less_than(findmin_pqp_res(X1),lookup_slb(X2,findmin_pqp_res(X1)))|~contains_slb(X2,findmin_pqp_res(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(findmin_cpq_eff(triple(X1,X2,X3))=triple(X1,update_slb(X2,findmin_pqp_res(X1)),bad)|contains_slb(X2,findmin_pqp_res(X1))|X2=create_slb), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(strictly_less_than(X1,X2)|less_than(X2,X1)|~less_than(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,plain,(less_than(X1,X2)|less_than(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, negated_conjecture, ((~check_cpq(triple(esk5_0,esk6_0,esk7_0))&check_cpq(findmin_cpq_eff(triple(esk5_0,esk6_0,esk7_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
fof(c_0_25, lemma, (![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:![X15]:((~check_cpq(triple(X6,X7,X8))|(~pair_in_list(X7,X9,X10)|less_than(X10,X9)))&((pair_in_list(X12,esk1_2(X11,X12),esk2_2(X11,X12))|check_cpq(triple(X11,X12,X15)))&(~less_than(esk2_2(X11,X12),esk1_2(X11,X12))|check_cpq(triple(X11,X12,X15)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).
fof(c_0_26, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((~check_cpq(triple(X6,insert_slb(X7,pair(X9,X10)),X8))|check_cpq(triple(X6,X7,X8)))|~less_than(X10,X9))&((~check_cpq(triple(X6,X7,X8))|check_cpq(triple(X6,insert_slb(X7,pair(X9,X10)),X8)))|~less_than(X10,X9)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_27,plain,(triple(X1,update_slb(X2,findmin_pqp_res(X1)),bad)=findmin_cpq_eff(triple(X1,X2,X3))|create_slb=X2|~strictly_less_than(findmin_pqp_res(X1),lookup_slb(X2,findmin_pqp_res(X1)))), inference(csr,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28,plain,(strictly_less_than(X1,X2)|less_than(X2,X1)), inference(csr,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29,negated_conjecture,(check_cpq(findmin_cpq_eff(triple(esk5_0,esk6_0,esk7_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30,plain,(findmin_cpq_eff(triple(X1,X2,X3))=findmin_cpq_eff(triple(X1,X2,X4))|create_slb=X2|contains_slb(X2,findmin_pqp_res(X1))), inference(spm,[status(thm)],[c_0_21, c_0_21])).
cnf(c_0_31,lemma,(less_than(X1,X2)|~pair_in_list(X3,X2,X1)|~check_cpq(triple(X4,X3,X5))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32,plain,(check_cpq(triple(X3,insert_slb(X4,pair(X2,X1)),X5))|~less_than(X1,X2)|~check_cpq(triple(X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_33, plain, (![X1]:![X2]:![X3]:(((X2!=create_slb&contains_slb(X2,findmin_pqp_res(X1)))&less_than(lookup_slb(X2,findmin_pqp_res(X1)),findmin_pqp_res(X1)))=>findmin_cpq_eff(triple(X1,X2,X3))=triple(X1,update_slb(X2,findmin_pqp_res(X1)),X3))), inference(fof_simplification,[status(thm)],[c_0_7])).
cnf(c_0_34,plain,(triple(X1,update_slb(X2,findmin_pqp_res(X1)),bad)=findmin_cpq_eff(triple(X1,X2,X3))|create_slb=X2|less_than(lookup_slb(X2,findmin_pqp_res(X1)),findmin_pqp_res(X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35,negated_conjecture,(esk6_0=create_slb|check_cpq(findmin_cpq_eff(triple(esk5_0,esk6_0,X1)))|contains_slb(esk6_0,findmin_pqp_res(esk5_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36,lemma,(less_than(X1,X2)|~check_cpq(triple(X3,X4,X5))|~pair_in_list(insert_slb(X4,pair(X6,X7)),X2,X1)|~less_than(X7,X6)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37,lemma,(check_cpq(triple(X1,X2,X3))|pair_in_list(X2,esk1_2(X1,X2),esk2_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38,lemma,(check_cpq(triple(X1,X2,X3))|~less_than(esk2_2(X1,X2),esk1_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_39, plain, (![X5]:![X6]:![X7]:![X8]:insert_cpq(triple(X5,X6,X7),X8)=triple(insert_pqp(X5,X8),insert_slb(X6,pair(X8,bottom)),X7)), inference(variable_rename,[status(thm)],[c_0_8])).
fof(c_0_40, plain, (![X2]:less_than(bottom,X2)), inference(variable_rename,[status(thm)],[c_0_9])).
fof(c_0_41, plain, (![X4]:![X5]:![X6]:(((X5=create_slb|~contains_slb(X5,findmin_pqp_res(X4)))|~less_than(lookup_slb(X5,findmin_pqp_res(X4)),findmin_pqp_res(X4)))|findmin_cpq_eff(triple(X4,X5,X6))=triple(X4,update_slb(X5,findmin_pqp_res(X4)),X6))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])])).
cnf(c_0_42,plain,(findmin_cpq_eff(triple(X1,X2,X3))=findmin_cpq_eff(triple(X1,X2,X4))|create_slb=X2|less_than(lookup_slb(X2,findmin_pqp_res(X1)),findmin_pqp_res(X1))), inference(spm,[status(thm)],[c_0_34, c_0_34])).
cnf(c_0_43,negated_conjecture,(esk6_0=create_slb|check_cpq(triple(esk5_0,update_slb(esk6_0,findmin_pqp_res(esk5_0)),bad))|contains_slb(esk6_0,findmin_pqp_res(esk5_0))), inference(spm,[status(thm)],[c_0_35, c_0_21])).
cnf(c_0_44,lemma,(check_cpq(triple(X1,insert_slb(X2,pair(X3,X4)),X5))|~check_cpq(triple(X6,X2,X7))|~less_than(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_45,plain,(insert_cpq(triple(X1,X2,X3),X4)=triple(insert_pqp(X1,X4),insert_slb(X2,pair(X4,bottom)),X3)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46,plain,(less_than(bottom,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47,plain,(findmin_cpq_eff(triple(X1,X2,X3))=triple(X1,update_slb(X2,findmin_pqp_res(X1)),X3)|X2=create_slb|~less_than(lookup_slb(X2,findmin_pqp_res(X1)),findmin_pqp_res(X1))|~contains_slb(X2,findmin_pqp_res(X1))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48,negated_conjecture,(esk6_0=create_slb|check_cpq(findmin_cpq_eff(triple(esk5_0,esk6_0,X1)))|less_than(lookup_slb(esk6_0,findmin_pqp_res(esk5_0)),findmin_pqp_res(esk5_0))), inference(spm,[status(thm)],[c_0_29, c_0_42])).
cnf(c_0_49,lemma,(esk6_0=create_slb|contains_slb(esk6_0,findmin_pqp_res(esk5_0))|less_than(X1,X2)|~pair_in_list(update_slb(esk6_0,findmin_pqp_res(esk5_0)),X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_43])).
cnf(c_0_50,lemma,(check_cpq(insert_cpq(triple(X1,X2,X3),X4))|~check_cpq(triple(X5,X2,X6))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])])).
cnf(c_0_51,negated_conjecture,(triple(esk5_0,update_slb(esk6_0,findmin_pqp_res(esk5_0)),X1)=findmin_cpq_eff(triple(esk5_0,esk6_0,X1))|esk6_0=create_slb|check_cpq(findmin_cpq_eff(triple(esk5_0,esk6_0,X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_35])).
fof(c_0_52, lemma, (![X7]:![X8]:![X9]:![X10]:((pair_in_list(update_slb(X7,X10),esk3_2(X7,X10),esk4_2(X7,X10))|(~pair_in_list(X7,X8,X9)|~strictly_less_than(X8,X9)))&(strictly_less_than(esk3_2(X7,X10),esk4_2(X7,X10))|(~pair_in_list(X7,X8,X9)|~strictly_less_than(X8,X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
cnf(c_0_53,plain,(check_cpq(triple(X3,X4,X5))|~less_than(X1,X2)|~check_cpq(triple(X3,insert_slb(X4,pair(X2,X1)),X5))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_54,negated_conjecture,(esk6_0=create_slb|check_cpq(triple(esk5_0,update_slb(esk6_0,findmin_pqp_res(esk5_0)),bad))|less_than(lookup_slb(esk6_0,findmin_pqp_res(esk5_0)),findmin_pqp_res(esk5_0))), inference(spm,[status(thm)],[c_0_48, c_0_34])).
cnf(c_0_55,lemma,(esk6_0=create_slb|check_cpq(triple(X1,update_slb(esk6_0,findmin_pqp_res(esk5_0)),X2))|contains_slb(esk6_0,findmin_pqp_res(esk5_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_37]), c_0_38])).
cnf(c_0_56,lemma,(esk6_0=create_slb|check_cpq(insert_cpq(triple(X1,update_slb(esk6_0,findmin_pqp_res(esk5_0)),X2),X3))|check_cpq(findmin_cpq_eff(triple(esk5_0,esk6_0,X4)))|~check_cpq(findmin_cpq_eff(triple(esk5_0,esk6_0,X5)))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_57,lemma,(strictly_less_than(esk3_2(X3,X4),esk4_2(X3,X4))|~strictly_less_than(X1,X2)|~pair_in_list(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_58,plain,(check_cpq(triple(insert_pqp(X1,X2),X3,X4))|~check_cpq(insert_cpq(triple(X1,X3,X4),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_45]), c_0_46])])).
cnf(c_0_59,negated_conjecture,(triple(esk5_0,update_slb(esk6_0,findmin_pqp_res(esk5_0)),X1)=findmin_cpq_eff(triple(esk5_0,esk6_0,X1))|esk6_0=create_slb|check_cpq(triple(esk5_0,update_slb(esk6_0,findmin_pqp_res(esk5_0)),bad))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_54]), c_0_55])).
cnf(c_0_60,negated_conjecture,(esk6_0=create_slb|check_cpq(insert_cpq(triple(X1,update_slb(esk6_0,findmin_pqp_res(esk5_0)),X2),X3))|check_cpq(findmin_cpq_eff(triple(esk5_0,esk6_0,X4)))), inference(spm,[status(thm)],[c_0_56, c_0_29])).
cnf(c_0_61,lemma,(pair_in_list(update_slb(X3,X4),esk3_2(X3,X4),esk4_2(X3,X4))|~strictly_less_than(X1,X2)|~pair_in_list(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62,lemma,(check_cpq(triple(X1,X2,X3))|strictly_less_than(esk3_2(X2,X4),esk4_2(X2,X4))|~strictly_less_than(esk1_2(X1,X2),esk2_2(X1,X2))), inference(spm,[status(thm)],[c_0_57, c_0_37])).
cnf(c_0_63,lemma,(less_than(X1,X2)|~check_cpq(insert_cpq(triple(X3,X4,X5),X6))|~pair_in_list(X4,X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_58])).
cnf(c_0_64,lemma,(esk6_0=create_slb|check_cpq(insert_cpq(triple(X1,update_slb(esk6_0,findmin_pqp_res(esk5_0)),X2),X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_59]), c_0_60]), c_0_50])).
cnf(c_0_65,lemma,(check_cpq(triple(X1,X2,X3))|pair_in_list(update_slb(X2,X4),esk3_2(X2,X4),esk4_2(X2,X4))|~strictly_less_than(esk1_2(X1,X2),esk2_2(X1,X2))), inference(spm,[status(thm)],[c_0_61, c_0_37])).
cnf(c_0_66,plain,(~strictly_less_than(X1,X2)|~less_than(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_67,lemma,(check_cpq(triple(X1,X2,X3))|strictly_less_than(esk3_2(X2,X4),esk4_2(X2,X4))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_28]), c_0_38])).
cnf(c_0_68,lemma,(esk6_0=create_slb|less_than(X1,X2)|~pair_in_list(update_slb(esk6_0,findmin_pqp_res(esk5_0)),X2,X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_69,lemma,(check_cpq(triple(X1,X2,X3))|pair_in_list(update_slb(X2,X4),esk3_2(X2,X4),esk4_2(X2,X4))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_28]), c_0_38])).
cnf(c_0_70,lemma,(check_cpq(triple(X1,X2,X3))|~less_than(esk4_2(X2,X4),esk3_2(X2,X4))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_71,negated_conjecture,(~check_cpq(triple(esk5_0,esk6_0,esk7_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_72,lemma,(esk6_0=create_slb|check_cpq(triple(X1,esk6_0,X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])).
fof(c_0_73, plain, (![X3]:![X4]:check_cpq(triple(X3,create_slb,X4))), inference(variable_rename,[status(thm)],[c_0_11])).
cnf(c_0_74,negated_conjecture,(esk6_0=create_slb), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_75,plain,(check_cpq(triple(X1,create_slb,X2))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_76,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_74]), c_0_75])]), ['proof']).
# SZS output end CNFRefutation.
EOF
