% Problem    : GEO344+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n179.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:56:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X16]:![X17]:![X18]:![X19]:(((((((((ron(X17,X19)&ron(X17,X18))&ron(X16,X19))&ron(X16,X18))&?[X20]:(X19=X20&rline(X20)))&?[X21]:(X18=X21&rline(X21)))&X16!=X17)&?[X22]:(X17=X22&rpoint(X22)))&?[X23]:(X16=X23&rpoint(X23)))=>X18=X19)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(73), 0), imp(cond(axiom(73), 0)))')).
fof(c_0_1, conjecture, (![X1]:(((ron(vd1409,X1)&ron(vd1401,X1))&rline(X1))=>X1=vd1419)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(theu(the(409), 1), imp(the(409)))')).
fof(c_0_2, axiom, ((((((((((vd1406=vd1407&ron(vd1400,vd1406))&ron(vd1399,vd1406))&rline(vd1407))&vd1400!=vd1401)&vd1399!=vd1401)&vd1399!=vd1400)&?[X2]:(vd1401=X2&rpoint(X2)))&?[X3]:(vd1400=X3&rpoint(X3)))&?[X4]:(vd1399=X4&rpoint(X4)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(pred(conjunct2(405), 5), and(pred(s2(plural(conjunct2(405))), 0), and(pred(s1(plural(conjunct2(405))), 0), and(pred(conjunct2(405), 1), and(pred(conjunct1(405), 9), and(pred(conjunct1(405), 8), and(pred(conjunct1(405), 7), and(qe(s3(plural(405))), and(qe(s2(plural(405))), qe(s1(plural(405))))))))))))')).
fof(c_0_3, axiom, (rpoint(vd1409)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(406, 0)')).
fof(c_0_4, axiom, (vd1409!=vd1401), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct2(conjunct2(conjunct2(406))), 1413, 0)')).
fof(c_0_5, axiom, (((ron(vd1409,vd1419)&ron(vd1401,vd1419))&rline(vd1419))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(pred(s2(plural(the(409))), 0), and(pred(s1(plural(the(409))), 0), pred(the(409), 0)))')).
fof(c_0_6, plain, (![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:(((((((((~ron(X25,X27)|~ron(X25,X26))|~ron(X24,X27))|~ron(X24,X26))|(X27!=X28|~rline(X28)))|(X26!=X29|~rline(X29)))|X24=X25)|(X25!=X30|~rpoint(X30)))|(X24!=X31|~rpoint(X31)))|X26=X27)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])).
fof(c_0_7, negated_conjecture, (~(![X1]:(((ron(vd1409,X1)&ron(vd1401,X1))&rline(X1))=>X1=vd1419))), inference(assume_negation,[status(cth)],[c_0_1])).
fof(c_0_8, plain, ((((((((((vd1406=vd1407&ron(vd1400,vd1406))&ron(vd1399,vd1406))&rline(vd1407))&vd1400!=vd1401)&vd1399!=vd1401)&vd1399!=vd1400)&(vd1401=esk2_0&rpoint(esk2_0)))&(vd1400=esk3_0&rpoint(esk3_0)))&(vd1399=esk4_0&rpoint(esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_2])])).
cnf(c_0_9,plain,(X1=X2|X4=X6|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rline(X7)|X1!=X7|~rline(X8)|X2!=X8|~ron(X4,X1)|~ron(X4,X2)|~ron(X6,X1)|~ron(X6,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, negated_conjecture, ((((ron(vd1409,esk1_0)&ron(vd1401,esk1_0))&rline(esk1_0))&esk1_0!=vd1419)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_11,plain,(rpoint(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,plain,(vd1401=esk2_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(X1=X2|X3=X4|~rpoint(X2)|~rpoint(X1)|~rline(X4)|~rline(X3)|~ron(X2,X4)|~ron(X2,X3)|~ron(X1,X4)|~ron(X1,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_9])])])])).
cnf(c_0_14,negated_conjecture,(ron(vd1401,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(rpoint(vd1401)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16,negated_conjecture,(rline(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,negated_conjecture,(X1=vd1401|X2=esk1_0|~rpoint(X1)|~rline(X2)|~ron(vd1401,X2)|~ron(X1,esk1_0)|~ron(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_18,negated_conjecture,(ron(vd1409,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19,plain,(rpoint(vd1409)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_20,plain,(vd1409!=vd1401), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_21,negated_conjecture,(X1=esk1_0|~rline(X1)|~ron(vd1401,X1)|~ron(vd1409,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_22,plain,(rline(vd1419)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_23,plain,(ron(vd1401,vd1419)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_24,plain,(ron(vd1409,vd1419)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25,negated_conjecture,(esk1_0!=vd1419), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])]), c_0_25]), ['proof']).
# SZS output end CNFRefutation.
EOF
