% Problem    : SWW360+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n127.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:13:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X9]:![X7]:hAPP(c_Set_OCollect(X7),hAPP(hAPP(c_COMBC(X7,tc_fun(X7,tc_HOL_Obool),tc_HOL_Obool),c_member(X7)),X9))=X9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__mem__eq)).
fof(c_0_1, axiom, (![X36]:![X7]:hAPP(c_Set_OCollect(X7),X36)=X36), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__def)).
fof(c_0_2, axiom, (![X143]:![X35]:![X36]:![X7]:![X142]:![X5]:hAPP(hAPP(hAPP(c_COMBC(X5,X142,X7),X36),X35),X143)=hAPP(hAPP(X36,X143),X35)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBC__1)).
fof(c_0_3, axiom, (![X12]:![X8]:![X2]:![X7]:![X5]:hAPP(c_Set_Oimage(X5,X7,X2),hAPP(hAPP(c_Set_Oinsert(X5),X8),X12))=hAPP(hAPP(c_Set_Oinsert(X7),hAPP(X2,X8)),hAPP(c_Set_Oimage(X5,X7,X2),X12))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_image__insert)).
fof(c_0_4, axiom, (![X9]:![X8]:![X7]:(hBOOL(hAPP(hAPP(c_member(X7),X8),X9))=>hAPP(hAPP(c_Set_Oinsert(X7),X8),X9)=X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_insert__absorb)).
fof(c_0_5, axiom, (![X9]:![X8]:![X7]:c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))!=hAPP(hAPP(c_Set_Oinsert(X7),X8),X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__not__insert)).
fof(c_0_6, axiom, (![X12]:![X8]:![X9]:![X7]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X7,tc_HOL_Obool)),X9),hAPP(hAPP(c_Set_Oinsert(X7),X8),X12))=hAPP(hAPP(c_Set_Oinsert(X7),X8),hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X7,tc_HOL_Obool)),X9),X12))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Un__insert__right)).
fof(c_0_7, hypothesis, (v_G=hAPP(c_Set_Oimage(tc_Com_Opname,t_a,v_mgt__call),v_U)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_1)).
fof(c_0_8, axiom, (![X36]:![X7]:(?[X3]:(hBOOL(hAPP(hAPP(c_member(X7),X3),c_Orderings_Otop__class_Otop(tc_fun(X7,tc_HOL_Obool))))&hBOOL(hAPP(X36,X3)))<=>?[X120]:hBOOL(hAPP(X36,X120)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_bex__UNIV)).
fof(c_0_9, axiom, (![X18]:![X7]:hBOOL(hAPP(hAPP(c_member(X7),X18),c_Orderings_Otop__class_Otop(tc_fun(X7,tc_HOL_Obool))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__I)).
fof(c_0_10, axiom, (![X36]:(~(~(hBOOL(X36)))|hBOOL(hAPP(c_fNot,X36)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__2)).
fof(c_0_11, axiom, (![X36]:![X7]:(hAPP(c_Set_OCollect(X7),X36)=c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))<=>![X3]:~(hBOOL(hAPP(X36,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__empty__eq)).
fof(c_0_12, axiom, (![X7]:c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))=hAPP(c_Set_OCollect(X7),hAPP(c_COMBK(tc_HOL_Obool,X7),c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__def)).
fof(c_0_13, axiom, (![X9]:![X7]:(class_Finite__Set_Ofinite(X7)=>hBOOL(hAPP(c_Finite__Set_Ofinite(X7),X9)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite__code)).
fof(c_0_14, axiom, (![X36]:(~(hBOOL(hAPP(c_fNot,X36)))|~(hBOOL(X36)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__1)).
fof(c_0_15, axiom, (![X12]:![X9]:![X7]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X7,tc_HOL_Obool)),X9),X12)=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X7,tc_HOL_Obool)),X12),X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Un__commute)).
fof(c_0_16, conjecture, (![X3]:(hBOOL(hAPP(hAPP(c_member(tc_Com_Opname),X3),v_U))=>v_P(hAPP(c_Set_Oimage(tc_Com_Opname,t_a,v_mgt__call),v_U),hAPP(hAPP(c_Set_Oinsert(t_a),hAPP(v_mgt__call,X3)),c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_2)).
fof(c_0_17, axiom, (![X36]:![X7]:hAPP(c_Set_OCollect(X7),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X7),c_fNot),X36))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X7,tc_HOL_Obool)),hAPP(c_Set_OCollect(X7),X36))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__neg__eq)).
fof(c_0_18, axiom, (![X281]:![X280]:![X7]:![X23]:hAPP(hAPP(c_COMBK(X23,X7),X280),X281)=X280), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBK__1)).
fof(c_0_19, axiom, (![X204]:![X36]:![X7]:(hAPP(c_List_Ofilter(X7,X36),X204)=X204<=>![X3]:(hBOOL(hAPP(hAPP(c_member(X7),X3),hAPP(c_List_Oset(X7),X204)))=>hBOOL(hAPP(X36,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_filter__id__conv)).
fof(c_0_20, axiom, (class_Finite__Set_Ofinite(tc_Product__Type_Ounit)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Product____Type__Ounit__Finite__Set_Ofinite)).
fof(c_0_21, axiom, (![X144]:![X7]:hAPP(c_Set_Oimage(X7,X7,c_COMBI(X7)),X144)=X144), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_image__ident)).
fof(c_0_22, axiom, (![X12]:![X8]:![X7]:c_Complete__Lattice_OSup__class_OSup(tc_fun(X7,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X7,tc_HOL_Obool)),X8),X12))=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X7,tc_HOL_Obool)),X8),c_Complete__Lattice_OSup__class_OSup(tc_fun(X7,tc_HOL_Obool),X12))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Union__insert)).
fof(c_0_23, axiom, (![X9]:![X18]:![X7]:(hBOOL(hAPP(hAPP(c_member(X7),X18),X9))<=>hBOOL(hAPP(X9,X18)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_mem__def)).
fof(c_0_24, axiom, (![X9]:![X7]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X7,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X7,tc_HOL_Obool)),X9))=X9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__complement)).
fof(c_0_25, axiom, (![X7]:c_Orderings_Otop__class_Otop(tc_fun(X7,tc_HOL_Obool))=hAPP(c_Set_OCollect(X7),hAPP(c_COMBK(tc_HOL_Obool,X7),c_fTrue))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__def)).
fof(c_0_26, axiom, (![X7]:![X36]:(hBOOL(hAPP(X36,c_Hilbert__Choice_OEps(X7,X36)))<=>?[X120]:hBOOL(hAPP(X36,X120)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_some__eq__ex)).
fof(c_0_27, axiom, (![X143]:![X35]:![X36]:![X142]:![X7]:![X5]:hAPP(hAPP(hAPP(c_COMBB(X5,X7,X142),X36),X35),X143)=hAPP(X36,hAPP(X35,X143))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBB__1)).
fof(c_0_28, axiom, (![X7]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X7,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X7,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__empty__eq)).
fof(c_0_29, axiom, (![X22]:![X23]:(class_Lattices_Oboolean__algebra(X23)=>hAPP(c_Groups_Ouminus__class_Ouminus(X23),hAPP(c_Groups_Ouminus__class_Ouminus(X23),X22))=X22)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__compl)).
fof(c_0_30, axiom, (![X280]:![X23]:hAPP(c_COMBI(X23),X280)=X280), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBI__1)).
fof(c_0_31, axiom, (![X1]:![X2]:(![X3]:hAPP(X2,X3)=hAPP(X1,X3)=>X2=X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_ext)).
fof(c_0_32, axiom, (![X9]:![X7]:hAPP(c_Set_OPow(X7),X9)=hAPP(c_Set_OCollect(tc_fun(X7,tc_HOL_Obool)),hAPP(hAPP(c_COMBC(tc_fun(X7,tc_HOL_Obool),tc_fun(X7,tc_HOL_Obool),tc_HOL_Obool),c_Orderings_Oord__class_Oless__eq(tc_fun(X7,tc_HOL_Obool))),X9))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Pow__def)).
fof(c_0_33, axiom, (![X7]:hBOOL(hAPP(c_Finite__Set_Ofinite(X7),c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite_OemptyI)).
fof(c_0_34, axiom, (![X202]:![X23]:hAPP(c_List_Orev(X23),hAPP(c_List_Orev(X23),X202))=X202), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rev__rev__ident)).
fof(c_0_35, axiom, (class_Lattices_Oboolean__algebra(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Lattices_Oboolean__algebra)).
fof(c_0_36, axiom, (![X9]:![X7]:(?[X3]:hBOOL(hAPP(hAPP(c_member(X7),X3),X9))<=>X9!=c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_ex__in__conv)).
fof(c_0_37, axiom, (![X9]:![X7]:c_Complete__Lattice_OSup__class_OSup(tc_fun(X7,tc_HOL_Obool),hAPP(c_Set_OPow(X7),X9))=X9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Union__Pow__eq)).
fof(c_0_38, axiom, (![X12]:![X9]:![X7]:hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(X7,tc_HOL_Obool)),X9),hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X7,tc_HOL_Obool)),X9),X12)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Un__upper1)).
fof(c_0_39, axiom, (![X23]:(class_Lattices_Oboolean__algebra(X23)=>hAPP(c_Groups_Ouminus__class_Ouminus(X23),c_Orderings_Obot__class_Obot(X23))=c_Orderings_Otop__class_Otop(X23))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_compl__bot__eq)).
fof(c_0_40, axiom, (![X18]:![X5]:![X7]:(class_Orderings_Obot(X7)=>hAPP(c_Orderings_Obot__class_Obot(tc_fun(X5,X7)),X18)=c_Orderings_Obot__class_Obot(X7))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_bot__apply)).
fof(c_0_41, axiom, (![X9]:![X19]:![X18]:![X7]:hAPP(hAPP(c_Set_Oinsert(X7),X18),hAPP(hAPP(c_Set_Oinsert(X7),X19),X9))=hAPP(hAPP(c_Set_Oinsert(X7),X19),hAPP(hAPP(c_Set_Oinsert(X7),X18),X9))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_insert__commute)).
fof(c_0_42, axiom, (![X8]:![X7]:hAPP(c_Set_OCollect(X7),hAPP(c_fequal,X8))=hAPP(hAPP(c_Set_Oinsert(X7),X8),c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_singleton__conv2)).
fof(c_0_43, axiom, (![X18]:~(![X31]:(X18=hAPP(c_Int_OAbs__Integ,X31)=>~(hBOOL(hAPP(hAPP(c_member(tc_fun(tc_prod(tc_Nat_Onat,tc_Nat_Onat),tc_HOL_Obool)),X31),c_Int_OInteg)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Abs__Integ__cases)).
fof(c_0_44, axiom, (![X22]:![X23]:hAPP(c_Fun_Oid(X23),X22)=X22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_id__apply)).
fof(c_0_45, hypothesis, (hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Com_Opname),v_U))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_46, axiom, (class_Orderings_Obot(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Obot)).
fof(c_0_47, axiom, (![X20]:![X11]:![X16]:![X8]:![X7]:(hAPP(hAPP(c_Set_Oinsert(X7),X8),hAPP(hAPP(c_Set_Oinsert(X7),X16),c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))))=hAPP(hAPP(c_Set_Oinsert(X7),X11),hAPP(hAPP(c_Set_Oinsert(X7),X20),c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))))<=>((X8=X11&X16=X20)|(X8=X20&X16=X11)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_doubleton__eq__iff)).
fof(c_0_48, axiom, (![X12]:![X9]:![X7]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X7,tc_HOL_Obool)),X9),X12)=c_Complete__Lattice_OSup__class_OSup(tc_fun(X7,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X7,tc_HOL_Obool)),X9),hAPP(hAPP(c_Set_Oinsert(tc_fun(X7,tc_HOL_Obool)),X12),c_Orderings_Obot__class_Obot(tc_fun(tc_fun(X7,tc_HOL_Obool),tc_HOL_Obool)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Un__eq__Union)).
fof(c_0_49, axiom, (![X9]:![X2]:![X5]:![X7]:(c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))=hAPP(c_Set_Oimage(X5,X7,X2),X9)<=>X9=c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__is__image)).
fof(c_0_50, axiom, (![X4]:![X5]:![X6]:![X7]:(hBOOL(hAPP(c_Finite__Set_Ofinite(X7),X6))=>hBOOL(hAPP(c_Finite__Set_Ofinite(X5),hAPP(c_Set_Oimage(X7,X5,X4),X6))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite__imageI)).
fof(c_0_51, axiom, (![X18]:![X5]:![X7]:(class_Orderings_Otop(X7)=>hAPP(c_Orderings_Otop__class_Otop(tc_fun(X5,X7)),X18)=c_Orderings_Otop__class_Otop(X7))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_top__apply)).
fof(c_0_52, axiom, (hBOOL(c_fTrue)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fTrue__1)).
fof(c_0_53, axiom, (class_Orderings_Otop(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Otop)).
fof(c_0_54, axiom, (![X129]:![X7]:hAPP(c_Relation_OField(X7),X129)=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X7,tc_HOL_Obool)),hAPP(c_Relation_ODomain(X7,X7),X129)),hAPP(c_Relation_ORange(X7,X7),X129))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Field__def)).
fof(c_0_55, axiom, (![X7]:hAPP(c_Relation_OField(X7),c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X7,X7),tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Field__empty)).
fof(c_0_56, axiom, (![X5]:![X7]:hAPP(c_Relation_ODomain(X7,X5),c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X7,X5),tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Domain__empty)).
fof(c_0_57, axiom, (![X9]:hAPP(c_Set_Oimage(tc_Int_Oint,tc_Nat_Onat,c_Int_Onat),hAPP(c_Set_Oimage(tc_Nat_Onat,tc_Int_Oint,c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint)),X9))=X9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_transfer__int__nat__set__return__embed)).
fof(c_0_58, axiom, (![X9]:![X1]:![X142]:![X2]:![X7]:![X5]:hAPP(c_Set_Oimage(X5,X7,X2),hAPP(c_Set_Oimage(X142,X5,X1),X9))=hAPP(c_Set_Oimage(X142,X7,hAPP(hAPP(c_COMBB(X5,X7,X142),X2),X1)),X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_image__image)).
fof(c_0_59, axiom, (![X19]:![X18]:(X18!=X19|hBOOL(hAPP(hAPP(c_fequal,X18),X19)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fequal__2)).
fof(c_0_60, axiom, (~(hBOOL(c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fFalse__1)).
fof(c_0_61, axiom, (![X25]:![X22]:![X23]:(class_Orderings_Opreorder(X23)=>(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(X23),X22),X25))=>~(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(X23),X25),X22)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_order__less__not__sym)).
fof(c_0_62, axiom, (![X18]:hBOOL(hAPP(hAPP(c_member(tc_fun(tc_prod(tc_Nat_Onat,tc_Nat_Onat),tc_HOL_Obool)),hAPP(c_Int_ORep__Integ,X18)),c_Int_OInteg))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Rep__Integ)).
fof(c_0_63, axiom, (![X54]:hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X54),hAPP(c_Nat_OSuc,X54)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_lessI)).
fof(c_0_64, axiom, (![X18]:hAPP(c_Int_OAbs__Integ,hAPP(c_Int_ORep__Integ,X18))=X18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Rep__Integ__inverse)).
fof(c_0_65, axiom, (class_Orderings_Opreorder(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Nat__Onat__Orderings_Opreorder)).
fof(c_0_66, plain, (![X10]:![X11]:hAPP(c_Set_OCollect(X11),hAPP(hAPP(c_COMBC(X11,tc_fun(X11,tc_HOL_Obool),tc_HOL_Obool),c_member(X11)),X10))=X10), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_67, plain, (![X37]:![X38]:hAPP(c_Set_OCollect(X38),X37)=X37), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_68, plain, (![X144]:![X145]:![X146]:![X147]:![X148]:![X149]:hAPP(hAPP(hAPP(c_COMBC(X149,X148,X147),X146),X145),X144)=hAPP(hAPP(X146,X144),X145)), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_69,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBC(X1,tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool),c_member(X1)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_70,plain,(hAPP(c_Set_OCollect(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_71, plain, (![X13]:![X14]:![X15]:![X16]:![X17]:hAPP(c_Set_Oimage(X17,X16,X15),hAPP(hAPP(c_Set_Oinsert(X17),X14),X13))=hAPP(hAPP(c_Set_Oinsert(X16),hAPP(X15,X14)),hAPP(c_Set_Oimage(X17,X16,X15),X13))), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_72, plain, (![X10]:![X11]:![X12]:(~hBOOL(hAPP(hAPP(c_member(X12),X11),X10))|hAPP(hAPP(c_Set_Oinsert(X12),X11),X10)=X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
cnf(c_0_73,plain,(hAPP(hAPP(hAPP(c_COMBC(X1,X2,X3),X4),X5),X6)=hAPP(hAPP(X4,X6),X5)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_74,plain,(hAPP(hAPP(c_COMBC(X1,tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool),c_member(X1)),X2)=X2), inference(rw,[status(thm)],[c_0_69, c_0_70])).
fof(c_0_75, plain, (![X10]:![X11]:![X12]:c_Orderings_Obot__class_Obot(tc_fun(X12,tc_HOL_Obool))!=hAPP(hAPP(c_Set_Oinsert(X12),X11),X10)), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_76, plain, (![X13]:![X14]:![X15]:![X16]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X16,tc_HOL_Obool)),X15),hAPP(hAPP(c_Set_Oinsert(X16),X14),X13))=hAPP(hAPP(c_Set_Oinsert(X16),X14),hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X16,tc_HOL_Obool)),X15),X13))), inference(variable_rename,[status(thm)],[c_0_6])).
cnf(c_0_77,plain,(hAPP(c_Set_Oimage(X1,X2,X3),hAPP(hAPP(c_Set_Oinsert(X1),X4),X5))=hAPP(hAPP(c_Set_Oinsert(X2),hAPP(X3,X4)),hAPP(c_Set_Oimage(X1,X2,X3),X5))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_78,hypothesis,(v_G=hAPP(c_Set_Oimage(tc_Com_Opname,t_a,v_mgt__call),v_U)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_79,plain,(hAPP(hAPP(c_Set_Oinsert(X1),X2),X3)=X3|~hBOOL(hAPP(hAPP(c_member(X1),X2),X3))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_80,plain,(hAPP(hAPP(c_member(X1),X2),X3)=hAPP(X3,X2)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
fof(c_0_81, plain, (![X121]:![X122]:![X123]:![X125]:![X126]:![X127]:(((~hBOOL(hAPP(hAPP(c_member(X122),X123),c_Orderings_Otop__class_Otop(tc_fun(X122,tc_HOL_Obool))))|~hBOOL(hAPP(X121,X123)))|hBOOL(hAPP(X121,esk507_1(X121))))&((hBOOL(hAPP(hAPP(c_member(X127),esk508_2(X125,X127)),c_Orderings_Otop__class_Otop(tc_fun(X127,tc_HOL_Obool))))|~hBOOL(hAPP(X125,X126)))&(hBOOL(hAPP(X125,esk508_2(X125,X127)))|~hBOOL(hAPP(X125,X126)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])])).
fof(c_0_82, plain, (![X19]:![X20]:hBOOL(hAPP(hAPP(c_member(X20),X19),c_Orderings_Otop__class_Otop(tc_fun(X20,tc_HOL_Obool))))), inference(variable_rename,[status(thm)],[c_0_9])).
fof(c_0_83, plain, (![X36]:(hBOOL(X36)|hBOOL(hAPP(c_fNot,X36)))), inference(fof_simplification,[status(thm)],[c_0_10])).
cnf(c_0_84,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))!=hAPP(hAPP(c_Set_Oinsert(X1),X2),X3)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_85,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(hAPP(c_Set_Oinsert(X1),X3),X4))=hAPP(hAPP(c_Set_Oinsert(X1),X3),hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X4))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_86,hypothesis,(hAPP(c_Set_Oimage(tc_Com_Opname,t_a,v_mgt__call),hAPP(hAPP(c_Set_Oinsert(tc_Com_Opname),X1),v_U))=hAPP(hAPP(c_Set_Oinsert(t_a),hAPP(v_mgt__call,X1)),v_G)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_87,plain,(hAPP(hAPP(c_Set_Oinsert(X1),X2),X3)=X3|~hBOOL(hAPP(X3,X2))), inference(rw,[status(thm)],[c_0_79, c_0_80])).
fof(c_0_88, plain, (![X36]:![X7]:(hAPP(c_Set_OCollect(X7),X36)=c_Orderings_Obot__class_Obot(tc_fun(X7,tc_HOL_Obool))<=>![X3]:~hBOOL(hAPP(X36,X3)))), inference(fof_simplification,[status(thm)],[c_0_11])).
fof(c_0_89, plain, (![X8]:c_Orderings_Obot__class_Obot(tc_fun(X8,tc_HOL_Obool))=hAPP(c_Set_OCollect(X8),hAPP(c_COMBK(tc_HOL_Obool,X8),c_fFalse))), inference(variable_rename,[status(thm)],[c_0_12])).
fof(c_0_90, plain, (![X10]:![X11]:(~class_Finite__Set_Ofinite(X11)|hBOOL(hAPP(c_Finite__Set_Ofinite(X11),X10)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).
fof(c_0_91, plain, (![X36]:(~hBOOL(hAPP(c_fNot,X36))|~hBOOL(X36))), inference(fof_simplification,[status(thm)],[c_0_14])).
cnf(c_0_92,plain,(hBOOL(hAPP(X1,esk507_1(X1)))|~hBOOL(hAPP(X1,X2))|~hBOOL(hAPP(hAPP(c_member(X3),X2),c_Orderings_Otop__class_Otop(tc_fun(X3,tc_HOL_Obool))))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_93,plain,(hBOOL(hAPP(hAPP(c_member(X1),X2),c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))))), inference(split_conjunct,[status(thm)],[c_0_82])).
fof(c_0_94, plain, (![X37]:(hBOOL(X37)|hBOOL(hAPP(c_fNot,X37)))), inference(variable_rename,[status(thm)],[c_0_83])).
cnf(c_0_95,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(hAPP(c_Set_Oinsert(X1),X3),X4))!=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_96,hypothesis,(hAPP(hAPP(c_Set_Oinsert(t_a),hAPP(v_mgt__call,X1)),v_G)=v_G|~hBOOL(hAPP(v_U,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_78])).
fof(c_0_97, plain, (![X13]:![X14]:![X15]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X15,tc_HOL_Obool)),X14),X13)=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X15,tc_HOL_Obool)),X13),X14)), inference(variable_rename,[status(thm)],[c_0_15])).
fof(c_0_98, negated_conjecture, (~(![X3]:(hBOOL(hAPP(hAPP(c_member(tc_Com_Opname),X3),v_U))=>v_P(hAPP(c_Set_Oimage(tc_Com_Opname,t_a,v_mgt__call),v_U),hAPP(hAPP(c_Set_Oinsert(t_a),hAPP(v_mgt__call,X3)),c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))))))), inference(assume_negation,[status(cth)],[c_0_16])).
fof(c_0_99, plain, (![X37]:![X38]:hAPP(c_Set_OCollect(X38),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X38),c_fNot),X37))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X38,tc_HOL_Obool)),hAPP(c_Set_OCollect(X38),X37))), inference(variable_rename,[status(thm)],[c_0_17])).
fof(c_0_100, plain, (![X37]:![X38]:![X39]:![X40]:![X42]:((hAPP(c_Set_OCollect(X38),X37)!=c_Orderings_Obot__class_Obot(tc_fun(X38,tc_HOL_Obool))|~hBOOL(hAPP(X37,X39)))&(hBOOL(hAPP(X40,esk282_1(X40)))|hAPP(c_Set_OCollect(X42),X40)=c_Orderings_Obot__class_Obot(tc_fun(X42,tc_HOL_Obool))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_88])])])])])).
fof(c_0_101, plain, (![X282]:![X283]:![X284]:![X285]:hAPP(hAPP(c_COMBK(X285,X284),X283),X282)=X283), inference(variable_rename,[status(thm)],[c_0_18])).
cnf(c_0_102,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse))), inference(split_conjunct,[status(thm)],[c_0_89])).
fof(c_0_103, plain, (![X205]:![X206]:![X207]:![X208]:![X209]:![X210]:![X211]:((hAPP(c_List_Ofilter(X207,X206),X205)!=X205|(~hBOOL(hAPP(hAPP(c_member(X207),X208),hAPP(c_List_Oset(X207),X205)))|hBOOL(hAPP(X206,X208))))&((hBOOL(hAPP(hAPP(c_member(X211),esk428_3(X209,X210,X211)),hAPP(c_List_Oset(X211),X209)))|hAPP(c_List_Ofilter(X211,X210),X209)=X209)&(~hBOOL(hAPP(X210,esk428_3(X209,X210,X211)))|hAPP(c_List_Ofilter(X211,X210),X209)=X209)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])])])).
cnf(c_0_104,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(X1),X2))|~class_Finite__Set_Ofinite(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_105,plain,(class_Finite__Set_Ofinite(tc_Product__Type_Ounit)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_106, plain, (![X145]:![X146]:hAPP(c_Set_Oimage(X146,X146,c_COMBI(X146)),X145)=X145), inference(variable_rename,[status(thm)],[c_0_21])).
fof(c_0_107, plain, (![X37]:(~hBOOL(hAPP(c_fNot,X37))|~hBOOL(X37))), inference(variable_rename,[status(thm)],[c_0_91])).
cnf(c_0_108,plain,(hBOOL(hAPP(X1,esk507_1(X1)))|~hBOOL(hAPP(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_109,plain,(hBOOL(hAPP(c_fNot,X1))|hBOOL(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_110,hypothesis,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(t_a,tc_HOL_Obool)),X1),v_G)!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))|~hBOOL(hAPP(v_U,X2))), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_111,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X3)=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X3),X2)), inference(split_conjunct,[status(thm)],[c_0_97])).
fof(c_0_112, plain, (![X13]:![X14]:![X15]:c_Complete__Lattice_OSup__class_OSup(tc_fun(X15,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X15,tc_HOL_Obool)),X14),X13))=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X15,tc_HOL_Obool)),X14),c_Complete__Lattice_OSup__class_OSup(tc_fun(X15,tc_HOL_Obool),X13))), inference(variable_rename,[status(thm)],[c_0_22])).
fof(c_0_113, plain, (![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:((~hBOOL(hAPP(hAPP(c_member(X21),X20),X19))|hBOOL(hAPP(X19,X20)))&(~hBOOL(hAPP(X22,X23))|hBOOL(hAPP(hAPP(c_member(X24),X23),X22))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).
fof(c_0_114, negated_conjecture, ((hBOOL(hAPP(hAPP(c_member(tc_Com_Opname),esk568_0),v_U))&~v_P(hAPP(c_Set_Oimage(tc_Com_Opname,t_a,v_mgt__call),v_U),hAPP(hAPP(c_Set_Oinsert(t_a),hAPP(v_mgt__call,esk568_0)),c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_98])])])).
fof(c_0_115, plain, (![X10]:![X11]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X11,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X11,tc_HOL_Obool)),X10))=X10), inference(variable_rename,[status(thm)],[c_0_24])).
cnf(c_0_116,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Set_OCollect(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_117,plain,(hAPP(c_Set_OCollect(X1),X2)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))|hBOOL(hAPP(X2,esk282_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_118,plain,(hAPP(hAPP(c_COMBK(X1,X2),X3),X4)=X3), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_119,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_102, c_0_70])).
fof(c_0_120, plain, (![X8]:c_Orderings_Otop__class_Otop(tc_fun(X8,tc_HOL_Obool))=hAPP(c_Set_OCollect(X8),hAPP(c_COMBK(tc_HOL_Obool,X8),c_fTrue))), inference(variable_rename,[status(thm)],[c_0_25])).
cnf(c_0_121,plain,(hAPP(c_List_Ofilter(X1,X2),X3)=X3|~hBOOL(hAPP(X2,esk428_3(X3,X2,X1)))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_122,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Product__Type_Ounit),X1))), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_123,plain,(hAPP(c_Set_Oimage(X1,X1,c_COMBI(X1)),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_124,plain,(~hBOOL(X1)|~hBOOL(hAPP(c_fNot,X1))), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_125,plain,(hBOOL(hAPP(c_fNot,esk507_1(c_fNot)))|hBOOL(X1)), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_126,hypothesis,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(t_a,tc_HOL_Obool)),v_G),X1)!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))|~hBOOL(hAPP(v_U,X2))), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_127,plain,(c_Complete__Lattice_OSup__class_OSup(tc_fun(X1,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X1,tc_HOL_Obool)),X2),X3))=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),c_Complete__Lattice_OSup__class_OSup(tc_fun(X1,tc_HOL_Obool),X3))), inference(split_conjunct,[status(thm)],[c_0_112])).
fof(c_0_128, plain, (![X121]:![X122]:![X124]:![X125]:![X126]:((~hBOOL(hAPP(X122,c_Hilbert__Choice_OEps(X121,X122)))|hBOOL(hAPP(X122,esk330_2(X121,X122))))&(~hBOOL(hAPP(X125,X126))|hBOOL(hAPP(X125,c_Hilbert__Choice_OEps(X124,X125)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])])])).
cnf(c_0_129,plain,(hBOOL(hAPP(X1,X2))|~hBOOL(hAPP(hAPP(c_member(X3),X2),X1))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_130,negated_conjecture,(hBOOL(hAPP(hAPP(c_member(tc_Com_Opname),esk568_0),v_U))), inference(split_conjunct,[status(thm)],[c_0_114])).
fof(c_0_131, plain, (![X144]:![X145]:![X146]:![X147]:![X148]:![X149]:hAPP(hAPP(hAPP(c_COMBB(X149,X148,X147),X146),X145),X144)=hAPP(X146,hAPP(X145,X144))), inference(variable_rename,[status(thm)],[c_0_27])).
cnf(c_0_132,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_133,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2)=hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_70]), c_0_70])).
cnf(c_0_134,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hBOOL(hAPP(X1,esk282_1(X1)))), inference(rw,[status(thm)],[c_0_117, c_0_70])).
cnf(c_0_135,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)),X2)=c_fFalse), inference(spm,[status(thm)],[c_0_118, c_0_119])).
fof(c_0_136, plain, (![X8]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X8,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X8,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X8,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_28])).
cnf(c_0_137,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue))), inference(split_conjunct,[status(thm)],[c_0_120])).
fof(c_0_138, plain, (![X24]:![X25]:(~class_Lattices_Oboolean__algebra(X25)|hAPP(c_Groups_Ouminus__class_Ouminus(X25),hAPP(c_Groups_Ouminus__class_Ouminus(X25),X24))=X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])).
fof(c_0_139, plain, (![X281]:![X282]:hAPP(c_COMBI(X282),X281)=X281), inference(variable_rename,[status(thm)],[c_0_30])).
cnf(c_0_140,plain,(hAPP(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_Product__Type_Ounit)),X2)=X2), inference(spm,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_141,plain,(hBOOL(esk507_1(c_Set_Oimage(X1,X1,c_COMBI(X1))))|~hBOOL(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_123]), c_0_123])).
cnf(c_0_142,plain,(hBOOL(hAPP(c_fNot,esk507_1(c_fNot)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_125])).
fof(c_0_143, plain, (![X4]:![X5]:(hAPP(X5,esk1_2(X4,X5))!=hAPP(X4,esk1_2(X4,X5))|X5=X4)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])).
cnf(c_0_144,hypothesis,(c_Complete__Lattice_OSup__class_OSup(tc_fun(t_a,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(t_a,tc_HOL_Obool)),v_G),X1))!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))|~hBOOL(hAPP(v_U,X2))), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_145,plain,(hBOOL(hAPP(X1,c_Hilbert__Choice_OEps(X2,X1)))|~hBOOL(hAPP(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_146,negated_conjecture,(hBOOL(hAPP(v_U,esk568_0))), inference(spm,[status(thm)],[c_0_129, c_0_130])).
fof(c_0_147, plain, (![X10]:![X11]:hAPP(c_Set_OPow(X11),X10)=hAPP(c_Set_OCollect(tc_fun(X11,tc_HOL_Obool)),hAPP(hAPP(c_COMBC(tc_fun(X11,tc_HOL_Obool),tc_fun(X11,tc_HOL_Obool),tc_HOL_Obool),c_Orderings_Oord__class_Oless__eq(tc_fun(X11,tc_HOL_Obool))),X10))), inference(variable_rename,[status(thm)],[c_0_32])).
fof(c_0_148, plain, (![X8]:hBOOL(hAPP(c_Finite__Set_Ofinite(X8),c_Orderings_Obot__class_Obot(tc_fun(X8,tc_HOL_Obool))))), inference(variable_rename,[status(thm)],[c_0_33])).
cnf(c_0_149,plain,(hAPP(hAPP(hAPP(c_COMBB(X1,X2,X3),X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_150,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_132, c_0_133]), c_0_133])).
fof(c_0_151, plain, (![X203]:![X204]:hAPP(c_List_Orev(X204),hAPP(c_List_Orev(X204),X203))=X203), inference(variable_rename,[status(thm)],[c_0_34])).
cnf(c_0_152,plain,(c_fFalse=X1|hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_134]), c_0_135]), c_0_118])).
cnf(c_0_153,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_154,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue)=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_137, c_0_70])).
cnf(c_0_155,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(X1),hAPP(c_Groups_Ouminus__class_Ouminus(X1),X2))=X2|~class_Lattices_Oboolean__algebra(X1)), inference(split_conjunct,[status(thm)],[c_0_138])).
cnf(c_0_156,plain,(class_Lattices_Oboolean__algebra(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_157,plain,(hBOOL(hAPP(X1,esk508_2(X1,X3)))|~hBOOL(hAPP(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_158,plain,(hAPP(c_COMBI(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_139])).
cnf(c_0_159,plain,(hBOOL(esk507_1(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_Product__Type_Ounit))))|~hBOOL(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_140]), c_0_140])).
cnf(c_0_160,plain,(hBOOL(esk507_1(c_Set_Oimage(X1,X1,c_COMBI(X1))))), inference(spm,[status(thm)],[c_0_141, c_0_142])).
fof(c_0_161, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:((~hBOOL(hAPP(hAPP(c_member(X11),X12),X10))|X10!=c_Orderings_Obot__class_Obot(tc_fun(X11,tc_HOL_Obool)))&(X13=c_Orderings_Obot__class_Obot(tc_fun(X14,tc_HOL_Obool))|hBOOL(hAPP(hAPP(c_member(X14),esk16_2(X13,X14)),X13))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])])])).
cnf(c_0_162,plain,(X1=X2|hAPP(X1,esk1_2(X2,X1))!=hAPP(X2,esk1_2(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_163,hypothesis,(c_Complete__Lattice_OSup__class_OSup(tc_fun(t_a,tc_HOL_Obool),X1)!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))|~hBOOL(hAPP(v_U,X2))|~hBOOL(hAPP(X1,v_G))), inference(spm,[status(thm)],[c_0_144, c_0_87])).
cnf(c_0_164,plain,(hBOOL(hAPP(X1,esk330_2(X2,X1)))|~hBOOL(hAPP(X1,c_Hilbert__Choice_OEps(X2,X1)))), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_165,negated_conjecture,(hBOOL(hAPP(v_U,c_Hilbert__Choice_OEps(X1,v_U)))), inference(spm,[status(thm)],[c_0_145, c_0_146])).
fof(c_0_166, plain, (![X10]:![X11]:c_Complete__Lattice_OSup__class_OSup(tc_fun(X11,tc_HOL_Obool),hAPP(c_Set_OPow(X11),X10))=X10), inference(variable_rename,[status(thm)],[c_0_37])).
cnf(c_0_167,plain,(hAPP(c_Set_OPow(X1),X2)=hAPP(c_Set_OCollect(tc_fun(X1,tc_HOL_Obool)),hAPP(hAPP(c_COMBC(tc_fun(X1,tc_HOL_Obool),tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool),c_Orderings_Oord__class_Oless__eq(tc_fun(X1,tc_HOL_Obool))),X2))), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_168,plain,(hAPP(c_Set_Oimage(X1,X2,X3),hAPP(hAPP(c_Set_Oinsert(X1),X4),X5))!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_84, c_0_77])).
cnf(c_0_169,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(X1),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))))), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_170,plain,(hAPP(c_fNot,hAPP(c_fNot,hAPP(X1,X2)))=hAPP(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_149])).
cnf(c_0_171,plain,(hAPP(c_List_Orev(X1),hAPP(c_List_Orev(X1),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_151])).
cnf(c_0_172,plain,(hAPP(c_fNot,X1)=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_124, c_0_152])).
cnf(c_0_173,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_153, c_0_133])).
cnf(c_0_174,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)),X2)=c_fTrue), inference(spm,[status(thm)],[c_0_118, c_0_154])).
cnf(c_0_175,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),X1))=X1), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_176,plain,(hBOOL(esk508_2(c_COMBI(X1),X2))|~hBOOL(X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158]), c_0_158])).
cnf(c_0_177,plain,(hBOOL(esk507_1(c_List_Ofilter(X1,c_Finite__Set_Ofinite(tc_Product__Type_Ounit))))), inference(spm,[status(thm)],[c_0_159, c_0_160])).
cnf(c_0_178,plain,(X1!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|~hBOOL(hAPP(hAPP(c_member(X2),X3),X1))), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_179,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hAPP(X1,esk1_2(c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool)),X1))!=c_fFalse), inference(spm,[status(thm)],[c_0_162, c_0_135])).
cnf(c_0_180,hypothesis,(c_Complete__Lattice_OSup__class_OSup(tc_fun(t_a,tc_HOL_Obool),X1)!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))|~hBOOL(hAPP(X1,v_G))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_165])])).
cnf(c_0_181,plain,(c_Complete__Lattice_OSup__class_OSup(tc_fun(X1,tc_HOL_Obool),hAPP(c_Set_OPow(X1),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_182,plain,(hAPP(hAPP(c_COMBC(tc_fun(X1,tc_HOL_Obool),tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool),c_Orderings_Oord__class_Oless__eq(tc_fun(X1,tc_HOL_Obool))),X2)=hAPP(c_Set_OPow(X1),X2)), inference(rw,[status(thm)],[c_0_167, c_0_70])).
fof(c_0_183, plain, (![X13]:![X14]:![X15]:hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(X15,tc_HOL_Obool)),X14),hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X15,tc_HOL_Obool)),X14),X13)))), inference(variable_rename,[status(thm)],[c_0_38])).
cnf(c_0_184,plain,(hAPP(c_Set_Oimage(X1,X2,X3),X4)!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|~hBOOL(hAPP(X4,X5))), inference(spm,[status(thm)],[c_0_168, c_0_87])).
cnf(c_0_185,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(X1),esk507_1(c_Finite__Set_Ofinite(X1))))), inference(spm,[status(thm)],[c_0_108, c_0_169])).
cnf(c_0_186,plain,(hAPP(c_fNot,hAPP(c_fNot,X1))=X1), inference(spm,[status(thm)],[c_0_170, c_0_171])).
cnf(c_0_187,plain,(hAPP(c_fNot,X1)=c_fFalse|c_fFalse=X1), inference(spm,[status(thm)],[c_0_172, c_0_152])).
cnf(c_0_188,plain,(hAPP(c_fNot,c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_173]), c_0_174]), c_0_135])).
cnf(c_0_189,plain,(hBOOL(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),esk507_1(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool))))|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_108, c_0_175])).
cnf(c_0_190,plain,(hBOOL(esk508_2(c_COMBI(X1),X2))), inference(spm,[status(thm)],[c_0_176, c_0_177])).
cnf(c_0_191,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(tc_Com_Opname,tc_HOL_Obool))!=v_U), inference(spm,[status(thm)],[c_0_178, c_0_130])).
cnf(c_0_192,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_179, c_0_135])).
cnf(c_0_193,hypothesis,(X1!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))|~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(t_a,tc_HOL_Obool)),v_G),X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181]), c_0_182]), c_0_73])).
cnf(c_0_194,plain,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X3)))), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_195,plain,(hAPP(c_Set_Oimage(X1,X2,X3),c_Finite__Set_Ofinite(X4))!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_184, c_0_185])).
cnf(c_0_196,plain,(c_fFalse=X1|c_fTrue=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187]), c_0_188])).
cnf(c_0_197,plain,(hBOOL(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),esk507_1(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool))))), inference(spm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_198,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))!=v_U), inference(spm,[status(thm)],[c_0_191, c_0_192])).
fof(c_0_199, plain, (![X24]:(~class_Lattices_Oboolean__algebra(X24)|hAPP(c_Groups_Ouminus__class_Ouminus(X24),c_Orderings_Obot__class_Obot(X24))=c_Orderings_Otop__class_Otop(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])).
fof(c_0_200, plain, (![X19]:![X20]:![X21]:(~class_Orderings_Obot(X21)|hAPP(c_Orderings_Obot__class_Obot(tc_fun(X20,X21)),X19)=c_Orderings_Obot__class_Obot(X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])).
cnf(c_0_201,hypothesis,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(t_a,tc_HOL_Obool)),v_G),X1)!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_193, c_0_194])).
fof(c_0_202, plain, (![X20]:![X21]:![X22]:![X23]:hAPP(hAPP(c_Set_Oinsert(X23),X22),hAPP(hAPP(c_Set_Oinsert(X23),X21),X20))=hAPP(hAPP(c_Set_Oinsert(X23),X21),hAPP(hAPP(c_Set_Oinsert(X23),X22),X20))), inference(variable_rename,[status(thm)],[c_0_41])).
fof(c_0_203, plain, (![X9]:![X10]:hAPP(c_Set_OCollect(X10),hAPP(c_fequal,X9))=hAPP(hAPP(c_Set_Oinsert(X10),X9),c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool)))), inference(variable_rename,[status(thm)],[c_0_42])).
fof(c_0_204, plain, (![X18]:~(![X31]:(X18=hAPP(c_Int_OAbs__Integ,X31)=>~hBOOL(hAPP(hAPP(c_member(tc_fun(tc_prod(tc_Nat_Onat,tc_Nat_Onat),tc_HOL_Obool)),X31),c_Int_OInteg))))), inference(fof_simplification,[status(thm)],[c_0_43])).
fof(c_0_205, plain, (![X24]:![X25]:hAPP(c_Fun_Oid(X25),X24)=X24), inference(variable_rename,[status(thm)],[c_0_44])).
cnf(c_0_206,plain,(c_Finite__Set_Ofinite(X1)!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_195, c_0_123])).
cnf(c_0_207,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(tc_Com_Opname,tc_HOL_Obool))=c_fFalse|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_191, c_0_196])).
cnf(c_0_208,hypothesis,(hBOOL(hAPP(c_Finite__Set_Ofinite(tc_Com_Opname),v_U))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_209,plain,(hAPP(c_Set_Oimage(X1,X2,X3),c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool))!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_184, c_0_197])).
cnf(c_0_210,hypothesis,(c_Set_Oimage(tc_Com_Opname,t_a,v_mgt__call)=c_fFalse|hAPP(c_fTrue,v_U)=v_G), inference(spm,[status(thm)],[c_0_78, c_0_196])).
cnf(c_0_211,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=c_fFalse|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_198, c_0_196])).
cnf(c_0_212,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(X1),c_Orderings_Obot__class_Obot(X1))=c_Orderings_Otop__class_Otop(X1)|~class_Lattices_Oboolean__algebra(X1)), inference(split_conjunct,[status(thm)],[c_0_199])).
cnf(c_0_213,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,X2)),X3)=c_Orderings_Obot__class_Obot(X2)|~class_Orderings_Obot(X2)), inference(split_conjunct,[status(thm)],[c_0_200])).
cnf(c_0_214,plain,(class_Orderings_Obot(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_215,hypothesis,(c_Complete__Lattice_OSup__class_OSup(tc_fun(t_a,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(t_a,tc_HOL_Obool)),v_G),X1))!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_201, c_0_127])).
cnf(c_0_216,plain,(hAPP(hAPP(c_Set_Oinsert(X1),X2),hAPP(hAPP(c_Set_Oinsert(X1),X3),X4))=hAPP(hAPP(c_Set_Oinsert(X1),X3),hAPP(hAPP(c_Set_Oinsert(X1),X2),X4))), inference(split_conjunct,[status(thm)],[c_0_202])).
fof(c_0_217, plain, (![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:(((((X24=X21|X24=X22)|hAPP(hAPP(c_Set_Oinsert(X25),X24),hAPP(hAPP(c_Set_Oinsert(X25),X23),c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool))))!=hAPP(hAPP(c_Set_Oinsert(X25),X22),hAPP(hAPP(c_Set_Oinsert(X25),X21),c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool)))))&((X23=X22|X24=X22)|hAPP(hAPP(c_Set_Oinsert(X25),X24),hAPP(hAPP(c_Set_Oinsert(X25),X23),c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool))))!=hAPP(hAPP(c_Set_Oinsert(X25),X22),hAPP(hAPP(c_Set_Oinsert(X25),X21),c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool))))))&(((X24=X21|X23=X21)|hAPP(hAPP(c_Set_Oinsert(X25),X24),hAPP(hAPP(c_Set_Oinsert(X25),X23),c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool))))!=hAPP(hAPP(c_Set_Oinsert(X25),X22),hAPP(hAPP(c_Set_Oinsert(X25),X21),c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool)))))&((X23=X22|X23=X21)|hAPP(hAPP(c_Set_Oinsert(X25),X24),hAPP(hAPP(c_Set_Oinsert(X25),X23),c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool))))!=hAPP(hAPP(c_Set_Oinsert(X25),X22),hAPP(hAPP(c_Set_Oinsert(X25),X21),c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool)))))))&(((X29!=X27|X28!=X26)|hAPP(hAPP(c_Set_Oinsert(X30),X29),hAPP(hAPP(c_Set_Oinsert(X30),X28),c_Orderings_Obot__class_Obot(tc_fun(X30,tc_HOL_Obool))))=hAPP(hAPP(c_Set_Oinsert(X30),X27),hAPP(hAPP(c_Set_Oinsert(X30),X26),c_Orderings_Obot__class_Obot(tc_fun(X30,tc_HOL_Obool)))))&((X29!=X26|X28!=X27)|hAPP(hAPP(c_Set_Oinsert(X30),X29),hAPP(hAPP(c_Set_Oinsert(X30),X28),c_Orderings_Obot__class_Obot(tc_fun(X30,tc_HOL_Obool))))=hAPP(hAPP(c_Set_Oinsert(X30),X27),hAPP(hAPP(c_Set_Oinsert(X30),X26),c_Orderings_Obot__class_Obot(tc_fun(X30,tc_HOL_Obool)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_47])])])])])).
cnf(c_0_218,plain,(hAPP(c_Set_OCollect(X1),hAPP(c_fequal,X2))=hAPP(hAPP(c_Set_Oinsert(X1),X2),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))), inference(split_conjunct,[status(thm)],[c_0_203])).
fof(c_0_219, plain, (![X13]:![X14]:![X15]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X15,tc_HOL_Obool)),X14),X13)=c_Complete__Lattice_OSup__class_OSup(tc_fun(X15,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X15,tc_HOL_Obool)),X14),hAPP(hAPP(c_Set_Oinsert(tc_fun(X15,tc_HOL_Obool)),X13),c_Orderings_Obot__class_Obot(tc_fun(tc_fun(X15,tc_HOL_Obool),tc_HOL_Obool)))))), inference(variable_rename,[status(thm)],[c_0_48])).
fof(c_0_220, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:((c_Orderings_Obot__class_Obot(tc_fun(X13,tc_HOL_Obool))!=hAPP(c_Set_Oimage(X12,X13,X11),X10)|X10=c_Orderings_Obot__class_Obot(tc_fun(X12,tc_HOL_Obool)))&(X14!=c_Orderings_Obot__class_Obot(tc_fun(X16,tc_HOL_Obool))|c_Orderings_Obot__class_Obot(tc_fun(X17,tc_HOL_Obool))=hAPP(c_Set_Oimage(X16,X17,X15),X14)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_49])])])])).
fof(c_0_221, plain, (![X32]:(X32=hAPP(c_Int_OAbs__Integ,esk320_1(X32))&hBOOL(hAPP(hAPP(c_member(tc_fun(tc_prod(tc_Nat_Onat,tc_Nat_Onat),tc_HOL_Obool)),esk320_1(X32)),c_Int_OInteg)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_204])])])).
cnf(c_0_222,plain,(hAPP(c_Fun_Oid(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_205])).
fof(c_0_223, plain, (![X8]:![X9]:![X10]:![X11]:(~hBOOL(hAPP(c_Finite__Set_Ofinite(X11),X10))|hBOOL(hAPP(c_Finite__Set_Ofinite(X9),hAPP(c_Set_Oimage(X11,X9,X8),X10))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_50])])).
cnf(c_0_224,negated_conjecture,(c_Finite__Set_Ofinite(X1)!=c_fFalse|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_206, c_0_207])).
cnf(c_0_225,hypothesis,(c_Finite__Set_Ofinite(tc_Com_Opname)=c_fFalse|hBOOL(hAPP(c_fTrue,v_U))), inference(spm,[status(thm)],[c_0_208, c_0_196])).
cnf(c_0_226,hypothesis,(hAPP(c_fTrue,v_U)=v_G|hAPP(c_fFalse,c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool))!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_209, c_0_210])).
cnf(c_0_227,negated_conjecture,(hAPP(c_fFalse,X1)=c_fFalse|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_135, c_0_211])).
fof(c_0_228, plain, (![X19]:![X20]:![X21]:(~class_Orderings_Otop(X21)|hAPP(c_Orderings_Otop__class_Otop(tc_fun(X20,X21)),X19)=c_Orderings_Otop__class_Otop(X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_51])])).
cnf(c_0_229,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_Orderings_Obot__class_Obot(tc_HOL_Obool))=c_Orderings_Otop__class_Otop(tc_HOL_Obool)), inference(spm,[status(thm)],[c_0_212, c_0_156])).
cnf(c_0_230,plain,(c_Orderings_Obot__class_Obot(tc_HOL_Obool)=c_fFalse), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214]), c_0_135])).
cnf(c_0_231,hypothesis,(c_Complete__Lattice_OSup__class_OSup(tc_fun(t_a,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(t_a,tc_HOL_Obool)),X1),hAPP(hAPP(c_Set_Oinsert(tc_fun(t_a,tc_HOL_Obool)),v_G),X2)))!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_215, c_0_216])).
cnf(c_0_232,plain,(hAPP(hAPP(c_Set_Oinsert(X1),X2),hAPP(hAPP(c_Set_Oinsert(X1),X3),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))))=hAPP(hAPP(c_Set_Oinsert(X1),X4),hAPP(hAPP(c_Set_Oinsert(X1),X5),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))))|X3!=X4|X2!=X5), inference(split_conjunct,[status(thm)],[c_0_217])).
cnf(c_0_233,plain,(hAPP(hAPP(c_Set_Oinsert(X1),X2),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=hAPP(c_fequal,X2)), inference(rw,[status(thm)],[c_0_218, c_0_70])).
cnf(c_0_234,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X3)=c_Complete__Lattice_OSup__class_OSup(tc_fun(X1,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(hAPP(c_Set_Oinsert(tc_fun(X1,tc_HOL_Obool)),X3),c_Orderings_Obot__class_Obot(tc_fun(tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool)))))), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_235,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|c_Orderings_Obot__class_Obot(tc_fun(X3,tc_HOL_Obool))!=hAPP(c_Set_Oimage(X2,X3,X4),X1)), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_236,plain,(X1=hAPP(c_Int_OAbs__Integ,esk320_1(X1))), inference(split_conjunct,[status(thm)],[c_0_221])).
cnf(c_0_237,plain,(hBOOL(esk507_1(c_Fun_Oid(X1)))|~hBOOL(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_222]), c_0_222])).
cnf(c_0_238,plain,(hBOOL(c_fTrue)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_239,plain,(hBOOL(hAPP(c_Finite__Set_Ofinite(X1),hAPP(c_Set_Oimage(X2,X1,X3),X4)))|~hBOOL(hAPP(c_Finite__Set_Ofinite(X2),X4))), inference(split_conjunct,[status(thm)],[c_0_223])).
cnf(c_0_240,negated_conjecture,(hBOOL(hAPP(c_fTrue,v_U))|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_224, c_0_225])).
cnf(c_0_241,negated_conjecture,(hAPP(c_fTrue,v_U)=v_G|c_fTrue!=v_U), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227]), c_0_211])).
cnf(c_0_242,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,X2)),X3)=c_Orderings_Otop__class_Otop(X2)|~class_Orderings_Otop(X2)), inference(split_conjunct,[status(thm)],[c_0_228])).
cnf(c_0_243,plain,(class_Orderings_Otop(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_244,plain,(c_Orderings_Otop__class_Otop(tc_HOL_Obool)=hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fFalse)), inference(rw,[status(thm)],[c_0_229, c_0_230])).
cnf(c_0_245,hypothesis,(c_Complete__Lattice_OSup__class_OSup(tc_fun(t_a,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(t_a,tc_HOL_Obool)),X1),X2))!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))|~hBOOL(hAPP(X2,v_G))), inference(spm,[status(thm)],[c_0_231, c_0_87])).
cnf(c_0_246,plain,(hAPP(hAPP(c_Set_Oinsert(X1),X2),hAPP(c_fequal,X3))=hAPP(hAPP(c_Set_Oinsert(X1),X3),hAPP(c_fequal,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_232, c_0_233]), c_0_233])])])).
cnf(c_0_247,plain,(c_Complete__Lattice_OSup__class_OSup(tc_fun(X1,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(c_fequal,X3)))=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X3)), inference(rw,[status(thm)],[c_0_234, c_0_233])).
fof(c_0_248, plain, (![X130]:![X131]:hAPP(c_Relation_OField(X131),X130)=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X131,tc_HOL_Obool)),hAPP(c_Relation_ODomain(X131,X131),X130)),hAPP(c_Relation_ORange(X131,X131),X130))), inference(variable_rename,[status(thm)],[c_0_54])).
fof(c_0_249, plain, (![X8]:hAPP(c_Relation_OField(X8),c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X8,X8),tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X8,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_55])).
fof(c_0_250, plain, (![X8]:![X9]:hAPP(c_Relation_ODomain(X9,X8),c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X9,X8),tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X9,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_56])).
cnf(c_0_251,hypothesis,(c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))!=v_G), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_78]), c_0_191])).
cnf(c_0_252,plain,(hBOOL(hAPP(c_Int_OAbs__Integ,esk507_1(c_Int_OAbs__Integ)))|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_108, c_0_236])).
cnf(c_0_253,plain,(hBOOL(esk507_1(c_Fun_Oid(X1)))), inference(spm,[status(thm)],[c_0_237, c_0_238])).
fof(c_0_254, plain, (![X10]:hAPP(c_Set_Oimage(tc_Int_Oint,tc_Nat_Onat,c_Int_Onat),hAPP(c_Set_Oimage(tc_Nat_Onat,tc_Int_Oint,c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint)),X10))=X10), inference(variable_rename,[status(thm)],[c_0_57])).
fof(c_0_255, plain, (![X143]:![X144]:![X145]:![X146]:![X147]:![X148]:hAPP(c_Set_Oimage(X148,X147,X146),hAPP(c_Set_Oimage(X145,X148,X144),X143))=hAPP(c_Set_Oimage(X145,X147,hAPP(hAPP(c_COMBB(X148,X147,X145),X146),X144)),X143)), inference(variable_rename,[status(thm)],[c_0_58])).
cnf(c_0_256,hypothesis,(hBOOL(hAPP(c_Finite__Set_Ofinite(t_a),v_G))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_78]), c_0_208])])).
cnf(c_0_257,negated_conjecture,(hBOOL(v_G)|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_240, c_0_241])).
cnf(c_0_258,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243]), c_0_174]), c_0_244])).
cnf(c_0_259,hypothesis,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(t_a,tc_HOL_Obool)),X1),X2)!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))|~hBOOL(hAPP(hAPP(c_fequal,X1),v_G))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246]), c_0_247])).
cnf(c_0_260,plain,(hAPP(c_Relation_OField(X1),X2)=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Relation_ODomain(X1,X1),X2)),hAPP(c_Relation_ORange(X1,X1),X2))), inference(split_conjunct,[status(thm)],[c_0_248])).
cnf(c_0_261,plain,(hAPP(c_Relation_OField(X1),c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X1,X1),tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_249])).
cnf(c_0_262,plain,(hAPP(c_Relation_ODomain(X1,X2),c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X1,X2),tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_250])).
cnf(c_0_263,hypothesis,(c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))=c_fFalse|v_G!=c_fTrue), inference(spm,[status(thm)],[c_0_251, c_0_196])).
cnf(c_0_264,plain,(hBOOL(hAPP(hAPP(c_member(tc_fun(tc_prod(tc_Nat_Onat,tc_Nat_Onat),tc_HOL_Obool)),esk320_1(X1)),c_Int_OInteg))), inference(split_conjunct,[status(thm)],[c_0_221])).
cnf(c_0_265,plain,(hBOOL(hAPP(c_Int_OAbs__Integ,esk507_1(c_Int_OAbs__Integ)))), inference(spm,[status(thm)],[c_0_252, c_0_253])).
cnf(c_0_266,plain,(hAPP(c_Set_Oimage(tc_Int_Oint,tc_Nat_Onat,c_Int_Onat),hAPP(c_Set_Oimage(tc_Nat_Onat,tc_Int_Oint,c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint)),X1))=X1), inference(split_conjunct,[status(thm)],[c_0_254])).
cnf(c_0_267,plain,(hAPP(c_Set_Oimage(X1,X2,X3),hAPP(c_Set_Oimage(X4,X1,X5),X6))=hAPP(c_Set_Oimage(X4,X2,hAPP(hAPP(c_COMBB(X1,X2,X4),X3),X5)),X6)), inference(split_conjunct,[status(thm)],[c_0_255])).
cnf(c_0_268,hypothesis,(c_Finite__Set_Ofinite(t_a)=c_fFalse|hBOOL(hAPP(c_fTrue,v_G))), inference(spm,[status(thm)],[c_0_256, c_0_196])).
cnf(c_0_269,negated_conjecture,(hAPP(c_fNot,v_G)=c_fFalse|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_172, c_0_257])).
cnf(c_0_270,plain,(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool)!=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_209, c_0_123])).
cnf(c_0_271,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fTrue)=c_fFalse), inference(spm,[status(thm)],[c_0_175, c_0_258])).
cnf(c_0_272,hypothesis,(hAPP(c_Relation_OField(t_a),X1)!=c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))|~hBOOL(hAPP(hAPP(c_fequal,hAPP(c_Relation_ODomain(t_a,t_a),X1)),v_G))), inference(spm,[status(thm)],[c_0_259, c_0_260])).
cnf(c_0_273,plain,(hAPP(c_Relation_OField(X1),c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_261, c_0_192])).
cnf(c_0_274,plain,(hAPP(c_Relation_ODomain(X1,X2),c_Orderings_Obot__class_Obot(tc_fun(X3,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_262, c_0_192])).
cnf(c_0_275,negated_conjecture,(v_U!=c_fFalse|v_G!=c_fTrue), inference(spm,[status(thm)],[c_0_198, c_0_263])).
cnf(c_0_276,plain,(c_Orderings_Obot__class_Obot(tc_fun(tc_fun(tc_prod(tc_Nat_Onat,tc_Nat_Onat),tc_HOL_Obool),tc_HOL_Obool))!=c_Int_OInteg), inference(spm,[status(thm)],[c_0_178, c_0_264])).
fof(c_0_277, plain, (![X20]:![X21]:(X21!=X20|hBOOL(hAPP(hAPP(c_fequal,X21),X20)))), inference(variable_rename,[status(thm)],[c_0_59])).
cnf(c_0_278,plain,(hAPP(c_Set_Oimage(X1,X2,X3),c_Int_OAbs__Integ)!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_184, c_0_265])).
cnf(c_0_279,plain,(hAPP(c_Set_Oimage(tc_Nat_Onat,tc_Nat_Onat,hAPP(hAPP(c_COMBB(tc_Int_Oint,tc_Nat_Onat,tc_Nat_Onat),c_Int_Onat),c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint))),X1)=X1), inference(rw,[status(thm)],[c_0_266, c_0_267])).
cnf(c_0_280,negated_conjecture,(hBOOL(hAPP(c_fTrue,v_G))|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_224, c_0_268])).
cnf(c_0_281,negated_conjecture,(v_G=c_fTrue|c_fTrue!=v_U), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_269]), c_0_188])).
cnf(c_0_282,negated_conjecture,(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool)!=c_fFalse|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_270, c_0_211])).
cnf(c_0_283,plain,(hAPP(c_fTrue,c_fTrue)=c_fFalse|c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool)=c_fFalse), inference(spm,[status(thm)],[c_0_271, c_0_196])).
fof(c_0_284, plain, (~hBOOL(c_fFalse)), inference(fof_simplification,[status(thm)],[c_0_60])).
cnf(c_0_285,hypothesis,(~hBOOL(hAPP(hAPP(c_fequal,c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))),v_G))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_273]), c_0_274])).
cnf(c_0_286,negated_conjecture,(v_G=c_fFalse|v_U!=c_fFalse), inference(spm,[status(thm)],[c_0_275, c_0_196])).
cnf(c_0_287,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))!=c_Int_OInteg), inference(spm,[status(thm)],[c_0_276, c_0_192])).
cnf(c_0_288,plain,(hBOOL(hAPP(hAPP(c_fequal,X1),X2))|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_277])).
cnf(c_0_289,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))!=c_Int_OAbs__Integ), inference(spm,[status(thm)],[c_0_278, c_0_123])).
cnf(c_0_290,plain,(c_Orderings_Obot__class_Obot(tc_fun(tc_Nat_Onat,tc_HOL_Obool))!=c_Int_OAbs__Integ), inference(spm,[status(thm)],[c_0_278, c_0_279])).
cnf(c_0_291,negated_conjecture,(hBOOL(hAPP(c_fTrue,c_fTrue))|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_280, c_0_281])).
cnf(c_0_292,negated_conjecture,(hAPP(c_fTrue,c_fTrue)=c_fFalse|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_282, c_0_283])).
cnf(c_0_293,plain,(~hBOOL(c_fFalse)), inference(split_conjunct,[status(thm)],[c_0_284])).
cnf(c_0_294,hypothesis,(v_U!=c_fFalse|~hBOOL(hAPP(hAPP(c_fequal,c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_HOL_Obool))),c_fFalse))), inference(spm,[status(thm)],[c_0_285, c_0_286])).
cnf(c_0_295,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=c_fFalse|c_fTrue!=c_Int_OInteg), inference(spm,[status(thm)],[c_0_287, c_0_196])).
cnf(c_0_296,plain,(hBOOL(hAPP(hAPP(c_fequal,X1),X1))), inference(er,[status(thm)],[c_0_288])).
cnf(c_0_297,negated_conjecture,(c_fFalse!=c_Int_OAbs__Integ|c_fTrue!=v_U), inference(spm,[status(thm)],[c_0_289, c_0_207])).
cnf(c_0_298,plain,(c_Orderings_Obot__class_Obot(tc_fun(tc_Nat_Onat,tc_HOL_Obool))=c_fFalse|c_fTrue!=c_Int_OAbs__Integ), inference(spm,[status(thm)],[c_0_290, c_0_196])).
cnf(c_0_299,negated_conjecture,(c_fTrue!=v_U), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_292]), c_0_293])).
cnf(c_0_300,hypothesis,(v_U!=c_fFalse|c_fTrue!=c_Int_OInteg), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_295]), c_0_296])])).
cnf(c_0_301,negated_conjecture,(c_fFalse=X1|c_fFalse!=c_Int_OAbs__Integ|X1!=v_U), inference(spm,[status(thm)],[c_0_297, c_0_196])).
cnf(c_0_302,negated_conjecture,(v_U!=c_fFalse|c_fTrue!=c_Int_OAbs__Integ), inference(spm,[status(thm)],[c_0_198, c_0_298])).
cnf(c_0_303,negated_conjecture,(c_fFalse=X1|X1!=v_U), inference(spm,[status(thm)],[c_0_299, c_0_196])).
cnf(c_0_304,hypothesis,(c_fFalse=X1|v_U!=c_fFalse|X1!=c_Int_OInteg), inference(spm,[status(thm)],[c_0_300, c_0_196])).
cnf(c_0_305,negated_conjecture,(v_U=c_fFalse|c_fFalse!=c_Int_OAbs__Integ), inference(er,[status(thm)],[c_0_301])).
cnf(c_0_306,negated_conjecture,(c_fFalse=X1|v_U!=c_fFalse|X1!=c_Int_OAbs__Integ), inference(spm,[status(thm)],[c_0_302, c_0_196])).
cnf(c_0_307,negated_conjecture,(v_U=c_fFalse), inference(er,[status(thm)],[c_0_303])).
fof(c_0_308, plain, (![X25]:![X22]:![X23]:(class_Orderings_Opreorder(X23)=>(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(X23),X22),X25))=>~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(X23),X25),X22))))), inference(fof_simplification,[status(thm)],[c_0_61])).
fof(c_0_309, plain, (![X19]:hBOOL(hAPP(hAPP(c_member(tc_fun(tc_prod(tc_Nat_Onat,tc_Nat_Onat),tc_HOL_Obool)),hAPP(c_Int_ORep__Integ,X19)),c_Int_OInteg))), inference(variable_rename,[status(thm)],[c_0_62])).
cnf(c_0_310,negated_conjecture,(c_fFalse=X1|c_fFalse!=c_Int_OAbs__Integ|X1!=c_Int_OInteg), inference(spm,[status(thm)],[c_0_304, c_0_305])).
cnf(c_0_311,negated_conjecture,(c_fFalse=X1|X1!=c_Int_OAbs__Integ), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_306, c_0_307])])).
fof(c_0_312, plain, (![X26]:![X27]:![X28]:(~class_Orderings_Opreorder(X28)|(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(X28),X27),X26))|~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(X28),X26),X27))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_308])])).
fof(c_0_313, plain, (![X55]:hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X55),hAPP(c_Nat_OSuc,X55)))), inference(variable_rename,[status(thm)],[c_0_63])).
cnf(c_0_314,plain,(hBOOL(hAPP(hAPP(c_member(tc_fun(tc_prod(tc_Nat_Onat,tc_Nat_Onat),tc_HOL_Obool)),hAPP(c_Int_ORep__Integ,X1)),c_Int_OInteg))), inference(split_conjunct,[status(thm)],[c_0_309])).
cnf(c_0_315,negated_conjecture,(c_Int_OInteg=c_fFalse|c_fFalse!=c_Int_OAbs__Integ), inference(er,[status(thm)],[c_0_310])).
cnf(c_0_316,negated_conjecture,(c_fFalse=c_Int_OAbs__Integ), inference(er,[status(thm)],[c_0_311])).
fof(c_0_317, plain, (![X19]:hAPP(c_Int_OAbs__Integ,hAPP(c_Int_ORep__Integ,X19))=X19), inference(variable_rename,[status(thm)],[c_0_64])).
cnf(c_0_318,plain,(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(X1),X2),X3))|~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(X1),X3),X2))|~class_Orderings_Opreorder(X1)), inference(split_conjunct,[status(thm)],[c_0_312])).
cnf(c_0_319,plain,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X1),hAPP(c_Nat_OSuc,X1)))), inference(split_conjunct,[status(thm)],[c_0_313])).
cnf(c_0_320,plain,(class_Orderings_Opreorder(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_321,plain,(hBOOL(hAPP(c_Int_OInteg,hAPP(c_Int_ORep__Integ,X1)))), inference(spm,[status(thm)],[c_0_129, c_0_314])).
cnf(c_0_322,negated_conjecture,(c_Int_OInteg=c_Int_OAbs__Integ), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_315, c_0_316]), c_0_316])])).
cnf(c_0_323,plain,(hAPP(c_Int_OAbs__Integ,hAPP(c_Int_ORep__Integ,X1))=X1), inference(split_conjunct,[status(thm)],[c_0_317])).
cnf(c_0_324,plain,(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),hAPP(c_Nat_OSuc,X1)),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_319]), c_0_320])])).
cnf(c_0_325,plain,(hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_321, c_0_322]), c_0_323])).
cnf(c_0_326,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_324, c_0_325])]), ['proof']).
# SZS output end CNFRefutation.
EOF
