% Problem    : NUM926+5 : TPTP v6.1.0. Released v5.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n188.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:16:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X4]:?[X5]:plus_plus(int,power_power(int,X4,number_number_of(nat,bit0(bit1(pls)))),power_power(int,X5,number_number_of(nat,bit0(bit1(pls)))))=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_1, axiom, (![X18]:number_number_of(int,X18)=X18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_86_number__of__is__id)).
fof(c_0_2, axiom, (![X15]:![X11]:plus_plus(int,X15,X11)=plus_plus(int,X11,X15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_89_zadd__commute)).
fof(c_0_3, axiom, ((t=one_one(int)=>?[X4]:?[X5]:plus_plus(int,power_power(int,X4,number_number_of(nat,bit0(bit1(pls)))),power_power(int,X5,number_number_of(nat,bit0(bit1(pls)))))=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1__096t_A_061_A1_A_061_061_062_AEX_Ax_Ay_O_Ax_A_094_A2_A_L_Ay_A_094_A2_A_06)).
fof(c_0_4, axiom, ((ord_less(int,one_one(int),t)=>?[X4]:?[X5]:plus_plus(int,power_power(int,X4,number_number_of(nat,bit0(bit1(pls)))),power_power(int,X5,number_number_of(nat,bit0(bit1(pls)))))=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_2__0961_A_060_At_A_061_061_062_AEX_Ax_Ay_O_Ax_A_094_A2_A_L_Ay_A_094_A2_A_06)).
fof(c_0_5, axiom, (![X16]:![X17]:(ord_less(int,X16,X17)<=>(ord_less_eq(int,X16,X17)&X16!=X17))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_22_zless__le)).
fof(c_0_6, axiom, (ord_less_eq(int,one_one(int),t)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_0_tpos)).
fof(c_0_7, negated_conjecture, (~(?[X4]:?[X5]:plus_plus(int,power_power(int,X4,number_number_of(nat,bit0(bit1(pls)))),power_power(int,X5,number_number_of(nat,bit0(bit1(pls)))))=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_8, negated_conjecture, (![X6]:![X7]:plus_plus(int,power_power(int,X6,number_number_of(nat,bit0(bit1(pls)))),power_power(int,X7,number_number_of(nat,bit0(bit1(pls)))))!=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
fof(c_0_9, plain, (![X19]:number_number_of(int,X19)=X19), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_10, plain, (![X16]:![X17]:plus_plus(int,X16,X17)=plus_plus(int,X17,X16)), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_11, plain, ((t!=one_one(int)|plus_plus(int,power_power(int,esk1_0,number_number_of(nat,bit0(bit1(pls)))),power_power(int,esk2_0,number_number_of(nat,bit0(bit1(pls)))))=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
cnf(c_0_12,negated_conjecture,(plus_plus(int,power_power(int,X1,number_number_of(nat,bit0(bit1(pls)))),power_power(int,X2,number_number_of(nat,bit0(bit1(pls)))))!=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13,plain,(number_number_of(int,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(plus_plus(int,X1,X2)=plus_plus(int,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ((~ord_less(int,one_one(int),t)|plus_plus(int,power_power(int,esk3_0,number_number_of(nat,bit0(bit1(pls)))),power_power(int,esk4_0,number_number_of(nat,bit0(bit1(pls)))))=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
fof(c_0_16, plain, (![X18]:![X19]:![X20]:![X21]:(((ord_less_eq(int,X18,X19)|~ord_less(int,X18,X19))&(X18!=X19|~ord_less(int,X18,X19)))&((~ord_less_eq(int,X20,X21)|X20=X21)|ord_less(int,X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
cnf(c_0_17,plain,(plus_plus(int,power_power(int,esk1_0,number_number_of(nat,bit0(bit1(pls)))),power_power(int,esk2_0,number_number_of(nat,bit0(bit1(pls)))))=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int))|t!=one_one(int)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,negated_conjecture,(plus_plus(int,power_power(int,X1,number_number_of(nat,bit0(bit1(pls)))),power_power(int,X2,number_number_of(nat,bit0(bit1(pls)))))!=plus_plus(int,one_one(int),times_times(int,bit0(bit0(bit1(pls))),m))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
cnf(c_0_19,plain,(plus_plus(int,power_power(int,esk3_0,number_number_of(nat,bit0(bit1(pls)))),power_power(int,esk4_0,number_number_of(nat,bit0(bit1(pls)))))=plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int))|~ord_less(int,one_one(int),t)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20,plain,(ord_less(int,X1,X2)|X1=X2|~ord_less_eq(int,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21,plain,(ord_less_eq(int,one_one(int),t)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22,plain,(t!=one_one(int)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_13]), c_0_14]), c_0_18])).
cnf(c_0_23,plain,(~ord_less(int,one_one(int),t)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_13]), c_0_14]), c_0_18])).
cnf(c_0_24,plain,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23]), ['proof']).
# SZS output end CNFRefutation.
EOF
