% Problem    : SCT163+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n146.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:28:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X11]:![X9]:(class_Finite__Set_Ofinite(X9)=>c_Finite__Set_Ofinite(X9,X11))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite)).
fof(c_0_1, axiom, (![X24]:hAPP(c_Nat_OSuc,X24)!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Suc__neq__Zero)).
fof(c_0_2, axiom, (![X6]:![X2]:![X9]:![X13]:(c_Set_Oimage(X13,X9,X2,c_Orderings_Otop__class_Otop(tc_fun(X13,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X9,tc_HOL_Obool))=>hAPP(X2,hAPP(c_Hilbert__Choice_Oinv__into(X13,X9,c_Orderings_Otop__class_Otop(tc_fun(X13,tc_HOL_Obool)),X2),X6))=X6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_surj__f__inv__f)).
fof(c_0_3, axiom, (![X11]:![X2]:![X13]:![X9]:(c_Finite__Set_Ofinite(X9,c_Set_Oimage(X13,X9,X2,X11))=>(c_Fun_Oinj__on(X13,X9,X2,X11)=>c_Finite__Set_Ofinite(X13,X11)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite__imageD)).
fof(c_0_4, axiom, (class_Finite__Set_Ofinite(tc_Product__Type_Ounit)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Product____Type__Ounit__Finite__Set_Ofinite)).
fof(c_0_5, axiom, (![X24]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X24),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Nat_Oadd__0__right)).
fof(c_0_6, axiom, (![X11]:![X2]:![X13]:![X9]:(c_Fun_Oinj__on(X9,X13,X2,X11)<=>![X3]:(hBOOL(hAPP(hAPP(c_member(X9),X3),X11))=>![X43]:(hBOOL(hAPP(hAPP(c_member(X9),X43),X11))=>(hAPP(X2,X3)=hAPP(X2,X43)=>X3=X43))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_inj__on__def)).
fof(c_0_7, axiom, (![X2]:![X9]:(c_Finite__Set_Ofinite(X9,c_Orderings_Otop__class_Otop(tc_fun(X9,tc_HOL_Obool)))=>(c_Fun_Oinj__on(X9,X9,X2,c_Orderings_Otop__class_Otop(tc_fun(X9,tc_HOL_Obool)))=>c_Set_Oimage(X9,X9,X2,c_Orderings_Otop__class_Otop(tc_fun(X9,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X9,tc_HOL_Obool))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite__UNIV__inj__surj)).
fof(c_0_8, axiom, (![X62]:c_Fun_Oinj__on(tc_Nat_Onat,tc_Nat_Onat,c_Nat_OSuc,X62)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_inj__Suc)).
fof(c_0_9, plain, (![X12]:![X13]:(~class_Finite__Set_Ofinite(X13)|c_Finite__Set_Ofinite(X13,X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])).
fof(c_0_10, plain, (![X25]:hAPP(c_Nat_OSuc,X25)!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_11, plain, (![X14]:![X15]:![X16]:![X17]:(c_Set_Oimage(X17,X16,X15,c_Orderings_Otop__class_Otop(tc_fun(X17,tc_HOL_Obool)))!=c_Orderings_Otop__class_Otop(tc_fun(X16,tc_HOL_Obool))|hAPP(X15,hAPP(c_Hilbert__Choice_Oinv__into(X17,X16,c_Orderings_Otop__class_Otop(tc_fun(X17,tc_HOL_Obool)),X15),X14))=X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
fof(c_0_12, plain, (![X14]:![X15]:![X16]:![X17]:(~c_Finite__Set_Ofinite(X17,c_Set_Oimage(X16,X17,X15,X14))|(~c_Fun_Oinj__on(X16,X17,X15,X14)|c_Finite__Set_Ofinite(X16,X14)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
cnf(c_0_13,plain,(c_Finite__Set_Ofinite(X1,X2)|~class_Finite__Set_Ofinite(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(class_Finite__Set_Ofinite(tc_Product__Type_Ounit)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_15, plain, (![X25]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X25),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X25), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_16, plain, (![X44]:![X45]:![X46]:![X47]:![X48]:![X49]:![X50]:![X51]:![X52]:![X53]:((~c_Fun_Oinj__on(X47,X46,X45,X44)|(~hBOOL(hAPP(hAPP(c_member(X47),X48),X44))|(~hBOOL(hAPP(hAPP(c_member(X47),X49),X44))|(hAPP(X45,X48)!=hAPP(X45,X49)|X48=X49))))&((hBOOL(hAPP(hAPP(c_member(X53),esk5_4(X50,X51,X52,X53)),X50))|c_Fun_Oinj__on(X53,X52,X51,X50))&((hBOOL(hAPP(hAPP(c_member(X53),esk6_4(X50,X51,X52,X53)),X50))|c_Fun_Oinj__on(X53,X52,X51,X50))&((hAPP(X51,esk5_4(X50,X51,X52,X53))=hAPP(X51,esk6_4(X50,X51,X52,X53))|c_Fun_Oinj__on(X53,X52,X51,X50))&(esk5_4(X50,X51,X52,X53)!=esk6_4(X50,X51,X52,X53)|c_Fun_Oinj__on(X53,X52,X51,X50))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])])).
fof(c_0_17, plain, (![X10]:![X11]:(~c_Finite__Set_Ofinite(X11,c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool)))|(~c_Fun_Oinj__on(X11,X11,X10,c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool)))|c_Set_Oimage(X11,X11,X10,c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
fof(c_0_18, plain, (![X63]:c_Fun_Oinj__on(tc_Nat_Onat,tc_Nat_Onat,c_Nat_OSuc,X63)), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_19,plain,(hAPP(c_Nat_OSuc,X1)!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20,plain,(hAPP(X1,hAPP(c_Hilbert__Choice_Oinv__into(X2,X3,c_Orderings_Otop__class_Otop(tc_fun(X2,tc_HOL_Obool)),X1),X4))=X4|c_Set_Oimage(X2,X3,X1,c_Orderings_Otop__class_Otop(tc_fun(X2,tc_HOL_Obool)))!=c_Orderings_Otop__class_Otop(tc_fun(X3,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21,plain,(c_Finite__Set_Ofinite(X1,X2)|~c_Fun_Oinj__on(X1,X3,X4,X2)|~c_Finite__Set_Ofinite(X3,c_Set_Oimage(X1,X3,X4,X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,plain,(c_Finite__Set_Ofinite(tc_Product__Type_Ounit,X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_23,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24,plain,(c_Fun_Oinj__on(X1,X2,X3,X4)|hAPP(X3,esk5_4(X4,X3,X2,X1))=hAPP(X3,esk6_4(X4,X3,X2,X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,plain,(c_Fun_Oinj__on(X1,X2,X3,X4)|esk5_4(X4,X3,X2,X1)!=esk6_4(X4,X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26,plain,(c_Set_Oimage(X1,X1,X2,c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))|~c_Fun_Oinj__on(X1,X1,X2,c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))|~c_Finite__Set_Ofinite(X1,c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27,plain,(c_Fun_Oinj__on(tc_Nat_Onat,tc_Nat_Onat,c_Nat_OSuc,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28,plain,(c_Set_Oimage(X1,X2,c_Nat_OSuc,c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))!=c_Orderings_Otop__class_Otop(tc_fun(X2,tc_HOL_Obool))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20])])).
cnf(c_0_29,plain,(c_Finite__Set_Ofinite(X1,X2)|~c_Fun_Oinj__on(X1,tc_Product__Type_Ounit,X3,X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_30,plain,(c_Fun_Oinj__on(X1,X2,c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X3)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_23]), c_0_25])).
cnf(c_0_31,plain,(~c_Finite__Set_Ofinite(tc_Nat_Onat,c_Orderings_Otop__class_Otop(tc_fun(tc_Nat_Onat,tc_HOL_Obool)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_32,plain,(c_Finite__Set_Ofinite(X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32])]), ['proof']).
# SZS output end CNFRefutation.
EOF
