% Problem    : SWV094+1 : TPTP v6.1.0. Bugfixed v3.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n183.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:55:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (((((((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)=use&a_select2(rho_defuse,n1)=use)&a_select2(rho_defuse,n2)=use)&a_select2(sigma_defuse,n0)=use)&a_select2(sigma_defuse,n1)=use)&a_select2(sigma_defuse,n2)=use)&a_select2(sigma_defuse,n3)=use)&a_select2(sigma_defuse,n4)=use)&a_select2(sigma_defuse,n5)=use)&a_select3(u_defuse,n0,n0)=use)&a_select3(u_defuse,n1,n0)=use)&a_select3(u_defuse,n2,n0)=use)&a_select2(xinit_defuse,n3)=use)&a_select2(xinit_defuse,n4)=use)&a_select2(xinit_defuse,n5)=use)&a_select2(xinit_mean_defuse,n0)=use)&a_select2(xinit_mean_defuse,n1)=use)&a_select2(xinit_mean_defuse,n2)=use)&a_select2(xinit_mean_defuse,n3)=use)&a_select2(xinit_mean_defuse,n4)=use)&a_select2(xinit_mean_defuse,n5)=use)&a_select2(xinit_noise_defuse,n0)=use)&a_select2(xinit_noise_defuse,n1)=use)&a_select2(xinit_noise_defuse,n2)=use)&a_select2(xinit_noise_defuse,n3)=use)&a_select2(xinit_noise_defuse,n4)=use)&a_select2(xinit_noise_defuse,n5)=use)&leq(n0,pv5))&leq(n0,pv51))&leq(n0,pv52))&leq(pv5,minus(n999,n1)))&leq(pv51,minus(n6,n1)))&leq(pv52,minus(n6,n1)))&![X14]:![X18]:((((leq(n0,X14)&leq(n0,X18))&leq(X14,n2))&leq(X18,minus(pv5,n1)))=>(a_select3(u_defuse,X14,X18)=use&a_select3(z_defuse,X14,X18)=use)))=>((pv51!=pv52=>(((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)=use&a_select2(rho_defuse,n1)=use)&a_select2(rho_defuse,n2)=use)&a_select2(sigma_defuse,n0)=use)&a_select2(sigma_defuse,n1)=use)&a_select2(sigma_defuse,n2)=use)&a_select2(sigma_defuse,n3)=use)&a_select2(sigma_defuse,n4)=use)&a_select2(sigma_defuse,n5)=use)&a_select3(u_defuse,n0,n0)=use)&a_select3(u_defuse,n1,n0)=use)&a_select3(u_defuse,n2,n0)=use)&a_select2(xinit_defuse,n3)=use)&a_select2(xinit_defuse,n4)=use)&a_select2(xinit_defuse,n5)=use)&a_select2(xinit_mean_defuse,n0)=use)&a_select2(xinit_mean_defuse,n1)=use)&a_select2(xinit_mean_defuse,n2)=use)&a_select2(xinit_mean_defuse,n3)=use)&a_select2(xinit_mean_defuse,n4)=use)&a_select2(xinit_mean_defuse,n5)=use)&a_select2(xinit_noise_defuse,n0)=use)&a_select2(xinit_noise_defuse,n1)=use)&a_select2(xinit_noise_defuse,n2)=use)&a_select2(xinit_noise_defuse,n3)=use)&a_select2(xinit_noise_defuse,n4)=use)&a_select2(xinit_noise_defuse,n5)=use)&leq(n0,pv5))&leq(n0,pv51))&leq(pv5,minus(n999,n1)))&leq(pv51,minus(n6,n1)))&![X4]:![X20]:((((leq(n0,X4)&leq(n0,X20))&leq(X4,n2))&leq(X20,minus(pv5,n1)))=>(a_select3(u_defuse,X4,X20)=use&a_select3(z_defuse,X4,X20)=use))))&(pv51=pv52=>(((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)=use&a_select2(rho_defuse,n1)=use)&a_select2(rho_defuse,n2)=use)&a_select2(sigma_defuse,n0)=use)&a_select2(sigma_defuse,n1)=use)&a_select2(sigma_defuse,n2)=use)&a_select2(sigma_defuse,n3)=use)&a_select2(sigma_defuse,n4)=use)&a_select2(sigma_defuse,n5)=use)&a_select3(u_defuse,n0,n0)=use)&a_select3(u_defuse,n1,n0)=use)&a_select3(u_defuse,n2,n0)=use)&a_select2(xinit_defuse,n3)=use)&a_select2(xinit_defuse,n4)=use)&a_select2(xinit_defuse,n5)=use)&a_select2(xinit_mean_defuse,n0)=use)&a_select2(xinit_mean_defuse,n1)=use)&a_select2(xinit_mean_defuse,n2)=use)&a_select2(xinit_mean_defuse,n3)=use)&a_select2(xinit_mean_defuse,n4)=use)&a_select2(xinit_mean_defuse,n5)=use)&a_select2(xinit_noise_defuse,n0)=use)&a_select2(xinit_noise_defuse,n1)=use)&a_select2(xinit_noise_defuse,n2)=use)&a_select2(xinit_noise_defuse,n3)=use)&a_select2(xinit_noise_defuse,n4)=use)&a_select2(xinit_noise_defuse,n5)=use)&leq(n0,pv5))&leq(n0,pv51))&leq(pv5,minus(n999,n1)))&leq(pv51,minus(n6,n1)))&![X21]:![X22]:((((leq(n0,X21)&leq(n0,X22))&leq(X21,n2))&leq(X22,minus(pv5,n1)))=>(a_select3(u_defuse,X21,X22)=use&a_select3(z_defuse,X21,X22)=use))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', quaternion_ds1_inuse_0006)).
fof(c_0_1, axiom, (![X1]:![X2]:(leq(X1,pred(X2))<=>gt(X2,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', leq_gt_pred)).
fof(c_0_2, axiom, (![X1]:minus(X1,n1)=pred(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', pred_minus_1)).
fof(c_0_3, axiom, (![X1]:![X2]:(gt(X2,X1)=>leq(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', leq_gt1)).
fof(c_0_4, axiom, (![X1]:~(gt(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', irreflexivity_gt)).
fof(c_0_5, axiom, (![X1]:![X2]:(leq(X1,X2)<=>gt(succ(X2),X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', leq_succ_gt_equiv)).
fof(c_0_6, axiom, (![X1]:plus(X1,n1)=succ(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', succ_plus_1_r)).
fof(c_0_7, axiom, (![X1]:![X2]:((leq(X1,X2)&X1!=X2)=>gt(X2,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', leq_gt2)).
fof(c_0_8, axiom, (![X1]:plus(n1,X1)=succ(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', succ_plus_1_l)).
fof(c_0_9, axiom, (![X1]:succ(pred(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', succ_pred)).
fof(c_0_10, axiom, (![X1]:![X2]:![X3]:((leq(X1,X2)&leq(X2,X3))=>leq(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', transitivity_leq)).
fof(c_0_11, axiom, (![X1]:leq(X1,X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', reflexivity_leq)).
fof(c_0_12, negated_conjecture, (~(((((((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)=use&a_select2(rho_defuse,n1)=use)&a_select2(rho_defuse,n2)=use)&a_select2(sigma_defuse,n0)=use)&a_select2(sigma_defuse,n1)=use)&a_select2(sigma_defuse,n2)=use)&a_select2(sigma_defuse,n3)=use)&a_select2(sigma_defuse,n4)=use)&a_select2(sigma_defuse,n5)=use)&a_select3(u_defuse,n0,n0)=use)&a_select3(u_defuse,n1,n0)=use)&a_select3(u_defuse,n2,n0)=use)&a_select2(xinit_defuse,n3)=use)&a_select2(xinit_defuse,n4)=use)&a_select2(xinit_defuse,n5)=use)&a_select2(xinit_mean_defuse,n0)=use)&a_select2(xinit_mean_defuse,n1)=use)&a_select2(xinit_mean_defuse,n2)=use)&a_select2(xinit_mean_defuse,n3)=use)&a_select2(xinit_mean_defuse,n4)=use)&a_select2(xinit_mean_defuse,n5)=use)&a_select2(xinit_noise_defuse,n0)=use)&a_select2(xinit_noise_defuse,n1)=use)&a_select2(xinit_noise_defuse,n2)=use)&a_select2(xinit_noise_defuse,n3)=use)&a_select2(xinit_noise_defuse,n4)=use)&a_select2(xinit_noise_defuse,n5)=use)&leq(n0,pv5))&leq(n0,pv51))&leq(n0,pv52))&leq(pv5,minus(n999,n1)))&leq(pv51,minus(n6,n1)))&leq(pv52,minus(n6,n1)))&![X14]:![X18]:((((leq(n0,X14)&leq(n0,X18))&leq(X14,n2))&leq(X18,minus(pv5,n1)))=>(a_select3(u_defuse,X14,X18)=use&a_select3(z_defuse,X14,X18)=use)))=>((pv51!=pv52=>(((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)=use&a_select2(rho_defuse,n1)=use)&a_select2(rho_defuse,n2)=use)&a_select2(sigma_defuse,n0)=use)&a_select2(sigma_defuse,n1)=use)&a_select2(sigma_defuse,n2)=use)&a_select2(sigma_defuse,n3)=use)&a_select2(sigma_defuse,n4)=use)&a_select2(sigma_defuse,n5)=use)&a_select3(u_defuse,n0,n0)=use)&a_select3(u_defuse,n1,n0)=use)&a_select3(u_defuse,n2,n0)=use)&a_select2(xinit_defuse,n3)=use)&a_select2(xinit_defuse,n4)=use)&a_select2(xinit_defuse,n5)=use)&a_select2(xinit_mean_defuse,n0)=use)&a_select2(xinit_mean_defuse,n1)=use)&a_select2(xinit_mean_defuse,n2)=use)&a_select2(xinit_mean_defuse,n3)=use)&a_select2(xinit_mean_defuse,n4)=use)&a_select2(xinit_mean_defuse,n5)=use)&a_select2(xinit_noise_defuse,n0)=use)&a_select2(xinit_noise_defuse,n1)=use)&a_select2(xinit_noise_defuse,n2)=use)&a_select2(xinit_noise_defuse,n3)=use)&a_select2(xinit_noise_defuse,n4)=use)&a_select2(xinit_noise_defuse,n5)=use)&leq(n0,pv5))&leq(n0,pv51))&leq(pv5,minus(n999,n1)))&leq(pv51,minus(n6,n1)))&![X4]:![X20]:((((leq(n0,X4)&leq(n0,X20))&leq(X4,n2))&leq(X20,minus(pv5,n1)))=>(a_select3(u_defuse,X4,X20)=use&a_select3(z_defuse,X4,X20)=use))))&(pv51=pv52=>(((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)=use&a_select2(rho_defuse,n1)=use)&a_select2(rho_defuse,n2)=use)&a_select2(sigma_defuse,n0)=use)&a_select2(sigma_defuse,n1)=use)&a_select2(sigma_defuse,n2)=use)&a_select2(sigma_defuse,n3)=use)&a_select2(sigma_defuse,n4)=use)&a_select2(sigma_defuse,n5)=use)&a_select3(u_defuse,n0,n0)=use)&a_select3(u_defuse,n1,n0)=use)&a_select3(u_defuse,n2,n0)=use)&a_select2(xinit_defuse,n3)=use)&a_select2(xinit_defuse,n4)=use)&a_select2(xinit_defuse,n5)=use)&a_select2(xinit_mean_defuse,n0)=use)&a_select2(xinit_mean_defuse,n1)=use)&a_select2(xinit_mean_defuse,n2)=use)&a_select2(xinit_mean_defuse,n3)=use)&a_select2(xinit_mean_defuse,n4)=use)&a_select2(xinit_mean_defuse,n5)=use)&a_select2(xinit_noise_defuse,n0)=use)&a_select2(xinit_noise_defuse,n1)=use)&a_select2(xinit_noise_defuse,n2)=use)&a_select2(xinit_noise_defuse,n3)=use)&a_select2(xinit_noise_defuse,n4)=use)&a_select2(xinit_noise_defuse,n5)=use)&leq(n0,pv5))&leq(n0,pv51))&leq(pv5,minus(n999,n1)))&leq(pv51,minus(n6,n1)))&![X21]:![X22]:((((leq(n0,X21)&leq(n0,X22))&leq(X21,n2))&leq(X22,minus(pv5,n1)))=>(a_select3(u_defuse,X21,X22)=use&a_select3(z_defuse,X21,X22)=use)))))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_13, negated_conjecture, (![X23]:![X24]:((((((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)=use&a_select2(rho_defuse,n1)=use)&a_select2(rho_defuse,n2)=use)&a_select2(sigma_defuse,n0)=use)&a_select2(sigma_defuse,n1)=use)&a_select2(sigma_defuse,n2)=use)&a_select2(sigma_defuse,n3)=use)&a_select2(sigma_defuse,n4)=use)&a_select2(sigma_defuse,n5)=use)&a_select3(u_defuse,n0,n0)=use)&a_select3(u_defuse,n1,n0)=use)&a_select3(u_defuse,n2,n0)=use)&a_select2(xinit_defuse,n3)=use)&a_select2(xinit_defuse,n4)=use)&a_select2(xinit_defuse,n5)=use)&a_select2(xinit_mean_defuse,n0)=use)&a_select2(xinit_mean_defuse,n1)=use)&a_select2(xinit_mean_defuse,n2)=use)&a_select2(xinit_mean_defuse,n3)=use)&a_select2(xinit_mean_defuse,n4)=use)&a_select2(xinit_mean_defuse,n5)=use)&a_select2(xinit_noise_defuse,n0)=use)&a_select2(xinit_noise_defuse,n1)=use)&a_select2(xinit_noise_defuse,n2)=use)&a_select2(xinit_noise_defuse,n3)=use)&a_select2(xinit_noise_defuse,n4)=use)&a_select2(xinit_noise_defuse,n5)=use)&leq(n0,pv5))&leq(n0,pv51))&leq(n0,pv52))&leq(pv5,minus(n999,n1)))&leq(pv51,minus(n6,n1)))&leq(pv52,minus(n6,n1)))&((a_select3(u_defuse,X23,X24)=use|(((~leq(n0,X23)|~leq(n0,X24))|~leq(X23,n2))|~leq(X24,minus(pv5,n1))))&(a_select3(z_defuse,X23,X24)=use|(((~leq(n0,X23)|~leq(n0,X24))|~leq(X23,n2))|~leq(X24,minus(pv5,n1))))))&(((pv51=pv52|pv51!=pv52)&((((((leq(n0,esk26_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|pv51!=pv52)&((leq(n0,esk27_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|pv51!=pv52))&((leq(esk26_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|pv51!=pv52))&((leq(esk27_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|pv51!=pv52))&(((a_select3(u_defuse,esk26_0,esk27_0)!=use|a_select3(z_defuse,esk26_0,esk27_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|pv51!=pv52)))&((((((pv51=pv52|(leq(n0,esk24_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((((((leq(n0,esk26_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk24_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((leq(n0,esk27_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk24_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk26_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk24_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk27_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk24_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&(((a_select3(u_defuse,esk26_0,esk27_0)!=use|a_select3(z_defuse,esk26_0,esk27_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk24_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))))&((pv51=pv52|(leq(n0,esk25_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((((((leq(n0,esk26_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk25_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((leq(n0,esk27_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk25_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk26_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk25_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk27_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk25_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&(((a_select3(u_defuse,esk26_0,esk27_0)!=use|a_select3(z_defuse,esk26_0,esk27_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(n0,esk25_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))))&((pv51=pv52|(leq(esk24_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((((((leq(n0,esk26_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk24_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((leq(n0,esk27_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk24_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk26_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk24_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk27_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk24_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&(((a_select3(u_defuse,esk26_0,esk27_0)!=use|a_select3(z_defuse,esk26_0,esk27_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk24_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))))&((pv51=pv52|(leq(esk25_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((((((leq(n0,esk26_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk25_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((leq(n0,esk27_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk25_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk26_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk25_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk27_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk25_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&(((a_select3(u_defuse,esk26_0,esk27_0)!=use|a_select3(z_defuse,esk26_0,esk27_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|(leq(esk25_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))))&((pv51=pv52|((a_select3(u_defuse,esk24_0,esk25_0)!=use|a_select3(z_defuse,esk24_0,esk25_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((((((leq(n0,esk26_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|((a_select3(u_defuse,esk24_0,esk25_0)!=use|a_select3(z_defuse,esk24_0,esk25_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1)))))&((leq(n0,esk27_0)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|((a_select3(u_defuse,esk24_0,esk25_0)!=use|a_select3(z_defuse,esk24_0,esk25_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk26_0,n2)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|((a_select3(u_defuse,esk24_0,esk25_0)!=use|a_select3(z_defuse,esk24_0,esk25_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&((leq(esk27_0,minus(pv5,n1))|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|((a_select3(u_defuse,esk24_0,esk25_0)!=use|a_select3(z_defuse,esk24_0,esk25_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))&(((a_select3(u_defuse,esk26_0,esk27_0)!=use|a_select3(z_defuse,esk26_0,esk27_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))|((a_select3(u_defuse,esk24_0,esk25_0)!=use|a_select3(z_defuse,esk24_0,esk25_0)!=use)|((((((((((((((((((((((((((((((a_select2(rho_defuse,n0)!=use|a_select2(rho_defuse,n1)!=use)|a_select2(rho_defuse,n2)!=use)|a_select2(sigma_defuse,n0)!=use)|a_select2(sigma_defuse,n1)!=use)|a_select2(sigma_defuse,n2)!=use)|a_select2(sigma_defuse,n3)!=use)|a_select2(sigma_defuse,n4)!=use)|a_select2(sigma_defuse,n5)!=use)|a_select3(u_defuse,n0,n0)!=use)|a_select3(u_defuse,n1,n0)!=use)|a_select3(u_defuse,n2,n0)!=use)|a_select2(xinit_defuse,n3)!=use)|a_select2(xinit_defuse,n4)!=use)|a_select2(xinit_defuse,n5)!=use)|a_select2(xinit_mean_defuse,n0)!=use)|a_select2(xinit_mean_defuse,n1)!=use)|a_select2(xinit_mean_defuse,n2)!=use)|a_select2(xinit_mean_defuse,n3)!=use)|a_select2(xinit_mean_defuse,n4)!=use)|a_select2(xinit_mean_defuse,n5)!=use)|a_select2(xinit_noise_defuse,n0)!=use)|a_select2(xinit_noise_defuse,n1)!=use)|a_select2(xinit_noise_defuse,n2)!=use)|a_select2(xinit_noise_defuse,n3)!=use)|a_select2(xinit_noise_defuse,n4)!=use)|a_select2(xinit_noise_defuse,n5)!=use)|~leq(n0,pv5))|~leq(n0,pv51))|~leq(pv5,minus(n999,n1)))|~leq(pv51,minus(n6,n1))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).
cnf(c_0_14,negated_conjecture,(leq(esk27_0,minus(pv5,n1))|pv51!=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_15,negated_conjecture,(a_select2(rho_defuse,n0)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16,negated_conjecture,(a_select2(rho_defuse,n1)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,negated_conjecture,(a_select2(rho_defuse,n2)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,negated_conjecture,(a_select2(sigma_defuse,n0)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,negated_conjecture,(a_select2(sigma_defuse,n1)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,negated_conjecture,(a_select2(sigma_defuse,n2)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,negated_conjecture,(a_select2(sigma_defuse,n3)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,negated_conjecture,(a_select2(sigma_defuse,n4)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23,negated_conjecture,(a_select2(sigma_defuse,n5)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24,negated_conjecture,(a_select2(xinit_defuse,n3)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,negated_conjecture,(a_select2(xinit_defuse,n4)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_26,negated_conjecture,(a_select2(xinit_defuse,n5)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_27,negated_conjecture,(a_select2(xinit_mean_defuse,n0)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_28,negated_conjecture,(a_select2(xinit_mean_defuse,n1)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_29,negated_conjecture,(a_select2(xinit_mean_defuse,n2)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_30,negated_conjecture,(a_select2(xinit_mean_defuse,n3)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31,negated_conjecture,(a_select2(xinit_mean_defuse,n4)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_32,negated_conjecture,(a_select2(xinit_mean_defuse,n5)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_33,negated_conjecture,(a_select2(xinit_noise_defuse,n0)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_34,negated_conjecture,(a_select2(xinit_noise_defuse,n1)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35,negated_conjecture,(a_select2(xinit_noise_defuse,n2)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_36,negated_conjecture,(a_select2(xinit_noise_defuse,n3)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_37,negated_conjecture,(a_select2(xinit_noise_defuse,n4)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_38,negated_conjecture,(a_select2(xinit_noise_defuse,n5)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_39,negated_conjecture,(a_select3(u_defuse,n0,n0)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_40,negated_conjecture,(a_select3(u_defuse,n1,n0)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_41,negated_conjecture,(a_select3(u_defuse,n2,n0)=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_42,negated_conjecture,(leq(n0,pv5)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_43,negated_conjecture,(leq(n0,pv51)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_44,negated_conjecture,(leq(pv5,minus(n999,n1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_45,negated_conjecture,(leq(pv51,minus(n6,n1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_46,negated_conjecture,(leq(n0,esk27_0)|pv51!=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_47, plain, (![X3]:![X4]:![X5]:![X6]:((~leq(X3,pred(X4))|gt(X4,X3))&(~gt(X6,X5)|leq(X5,pred(X6))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_48, plain, (![X2]:minus(X2,n1)=pred(X2)), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_49,negated_conjecture,(pv51!=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use|a_select3(z_defuse,esk26_0,esk27_0)!=use|a_select3(u_defuse,esk26_0,esk27_0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_50,negated_conjecture,(a_select3(z_defuse,X2,X1)=use|~leq(X1,minus(pv5,n1))|~leq(X2,n2)|~leq(n0,X1)|~leq(n0,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_51,negated_conjecture,(leq(esk27_0,minus(pv5,n1))|pv52!=pv51), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_52,negated_conjecture,(leq(n0,esk27_0)|pv52!=pv51), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_53,negated_conjecture,(leq(n0,esk26_0)|pv51!=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54,negated_conjecture,(leq(esk26_0,n2)|pv51!=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_55, plain, (![X3]:![X4]:(~gt(X4,X3)|leq(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
cnf(c_0_56,plain,(gt(X1,X2)|~leq(X2,pred(X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57,plain,(minus(X1,n1)=pred(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_58, plain, (![X1]:~gt(X1,X1)), inference(fof_simplification,[status(thm)],[c_0_4])).
fof(c_0_59, plain, (![X3]:![X4]:![X5]:![X6]:((~leq(X3,X4)|gt(succ(X4),X3))&(~gt(succ(X6),X5)|leq(X5,X6)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
fof(c_0_60, plain, (![X2]:plus(X2,n1)=succ(X2)), inference(variable_rename,[status(thm)],[c_0_6])).
cnf(c_0_61,plain,(leq(X1,pred(X2))|~gt(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_62, plain, (![X3]:![X4]:((~leq(X3,X4)|X3=X4)|gt(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
cnf(c_0_63,negated_conjecture,(a_select3(u_defuse,esk26_0,esk27_0)!=use|a_select3(z_defuse,esk26_0,esk27_0)!=use|pv52!=pv51), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_64,negated_conjecture,(a_select3(z_defuse,X1,esk27_0)=use|pv52!=pv51|~leq(X1,n2)|~leq(n0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_65,negated_conjecture,(leq(n0,esk26_0)|pv52!=pv51), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_66,negated_conjecture,(leq(esk26_0,n2)|pv52!=pv51), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_67,negated_conjecture,(a_select3(u_defuse,X2,X1)=use|~leq(X1,minus(pv5,n1))|~leq(X2,n2)|~leq(n0,X1)|~leq(n0,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_68,plain,(leq(X1,X2)|~gt(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_69,plain,(gt(X1,X2)|~leq(X2,minus(X1,n1))), inference(rw,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_70,negated_conjecture,(leq(esk25_0,minus(pv5,n1))|pv51=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_71, plain, (![X2]:~gt(X2,X2)), inference(variable_rename,[status(thm)],[c_0_58])).
cnf(c_0_72,plain,(gt(succ(X1),X2)|~leq(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_73,plain,(plus(X1,n1)=succ(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_74, plain, (![X2]:plus(n1,X2)=succ(X2)), inference(variable_rename,[status(thm)],[c_0_8])).
fof(c_0_75, plain, (![X2]:succ(pred(X2))=X2), inference(variable_rename,[status(thm)],[c_0_9])).
cnf(c_0_76,plain,(leq(X1,minus(X2,n1))|~gt(X2,X1)), inference(rw,[status(thm)],[c_0_61, c_0_57])).
cnf(c_0_77,plain,(gt(X1,X2)|X2=X1|~leq(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_78,negated_conjecture,(leq(esk24_0,n2)|pv51=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_79,negated_conjecture,(a_select3(u_defuse,esk26_0,esk27_0)!=use|pv52!=pv51), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_80,negated_conjecture,(a_select3(u_defuse,X1,esk27_0)=use|pv52!=pv51|~leq(X1,n2)|~leq(n0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_51]), c_0_52])).
cnf(c_0_81,negated_conjecture,(leq(n0,esk24_0)|pv51=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_82,negated_conjecture,(leq(n0,esk25_0)|pv51=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_83,plain,(leq(X1,X2)|~leq(X1,minus(X2,n1))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_84,negated_conjecture,(pv52=pv51|leq(esk25_0,minus(pv5,n1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_85,plain,(~gt(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_86,plain,(gt(plus(X1,n1),X2)|~leq(X2,X1)), inference(rw,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_87,plain,(plus(n1,X1)=succ(X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
fof(c_0_88, plain, (![X4]:![X5]:![X6]:((~leq(X4,X5)|~leq(X5,X6))|leq(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
cnf(c_0_89,plain,(succ(pred(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_90,negated_conjecture,(pv51=pv52|~leq(pv51,minus(n6,n1))|~leq(pv5,minus(n999,n1))|~leq(n0,pv51)|~leq(n0,pv5)|a_select2(xinit_noise_defuse,n5)!=use|a_select2(xinit_noise_defuse,n4)!=use|a_select2(xinit_noise_defuse,n3)!=use|a_select2(xinit_noise_defuse,n2)!=use|a_select2(xinit_noise_defuse,n1)!=use|a_select2(xinit_noise_defuse,n0)!=use|a_select2(xinit_mean_defuse,n5)!=use|a_select2(xinit_mean_defuse,n4)!=use|a_select2(xinit_mean_defuse,n3)!=use|a_select2(xinit_mean_defuse,n2)!=use|a_select2(xinit_mean_defuse,n1)!=use|a_select2(xinit_mean_defuse,n0)!=use|a_select2(xinit_defuse,n5)!=use|a_select2(xinit_defuse,n4)!=use|a_select2(xinit_defuse,n3)!=use|a_select3(u_defuse,n2,n0)!=use|a_select3(u_defuse,n1,n0)!=use|a_select3(u_defuse,n0,n0)!=use|a_select2(sigma_defuse,n5)!=use|a_select2(sigma_defuse,n4)!=use|a_select2(sigma_defuse,n3)!=use|a_select2(sigma_defuse,n2)!=use|a_select2(sigma_defuse,n1)!=use|a_select2(sigma_defuse,n0)!=use|a_select2(rho_defuse,n2)!=use|a_select2(rho_defuse,n1)!=use|a_select2(rho_defuse,n0)!=use|a_select3(z_defuse,esk24_0,esk25_0)!=use|a_select3(u_defuse,esk24_0,esk25_0)!=use), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_91,plain,(X1=X2|leq(X2,minus(X1,n1))|~leq(X2,X1)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_92,negated_conjecture,(pv52=pv51|leq(esk24_0,n2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_93,negated_conjecture,(pv52!=pv51), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_65]), c_0_66])).
cnf(c_0_94,negated_conjecture,(pv52=pv51|leq(n0,esk24_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_95,negated_conjecture,(pv52=pv51|leq(n0,esk25_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_96,negated_conjecture,(pv52=pv51|leq(esk25_0,pv5)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_97,plain,(~leq(plus(X1,n1),X1)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_98,plain,(plus(n1,X1)=plus(X1,n1)), inference(rw,[status(thm)],[c_0_87, c_0_73])).
cnf(c_0_99,plain,(leq(X1,X2)|~leq(X3,X2)|~leq(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_100,plain,(plus(minus(X1,n1),n1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_73]), c_0_57])).
cnf(c_0_101,negated_conjecture,(pv52=pv51|a_select3(u_defuse,esk24_0,esk25_0)!=use|a_select3(z_defuse,esk24_0,esk25_0)!=use), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_102,negated_conjecture,(a_select3(z_defuse,X1,X2)=use|pv5=X2|~leq(X1,n2)|~leq(n0,X1)|~leq(n0,X2)|~leq(X2,pv5)), inference(spm,[status(thm)],[c_0_50, c_0_91])).
cnf(c_0_103,negated_conjecture,(leq(esk24_0,n2)), inference(sr,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_104,negated_conjecture,(leq(n0,esk24_0)), inference(sr,[status(thm)],[c_0_94, c_0_93])).
cnf(c_0_105,negated_conjecture,(leq(n0,esk25_0)), inference(sr,[status(thm)],[c_0_95, c_0_93])).
cnf(c_0_106,negated_conjecture,(leq(esk25_0,pv5)), inference(sr,[status(thm)],[c_0_96, c_0_93])).
cnf(c_0_107,plain,(~leq(plus(n1,X1),X1)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_108,negated_conjecture,(pv52=pv51|leq(X1,minus(pv5,n1))|~leq(X1,esk25_0)), inference(spm,[status(thm)],[c_0_99, c_0_84])).
cnf(c_0_109,plain,(plus(n1,minus(X1,n1))=X1), inference(rw,[status(thm)],[c_0_100, c_0_98])).
cnf(c_0_110,negated_conjecture,(esk25_0=pv5|a_select3(u_defuse,esk24_0,esk25_0)!=use), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_93])).
cnf(c_0_111,negated_conjecture,(a_select3(u_defuse,X1,X2)=use|pv5=X2|~leq(X1,n2)|~leq(n0,X1)|~leq(n0,X2)|~leq(X2,pv5)), inference(spm,[status(thm)],[c_0_67, c_0_91])).
fof(c_0_112, plain, (![X2]:leq(X2,X2)), inference(variable_rename,[status(thm)],[c_0_11])).
cnf(c_0_113,negated_conjecture,(pv52=pv51|~leq(pv5,esk25_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109])).
cnf(c_0_114,negated_conjecture,(esk25_0=pv5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])).
cnf(c_0_115,plain,(leq(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_116,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_114]), c_0_115])]), c_0_93]), ['proof']).
# SZS output end CNFRefutation.
EOF
