% Problem    : SWW470+3 : TPTP v6.1.0. Released v5.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n151.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:57:36 CDT 2014
% CPUTime    : 
# SinE strategy is gf120_h_gu_R02_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X15]:![X3]:![X14]:![X13]:(![X17]:![X18]:(hBOOL(hAPP_state_bool(hAPP_a2036067514e_bool(X13,X17),X18))=>?[X38]:?[X39]:(hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X3),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(X38,X14,X39)),bot_bo797238721a_bool)))&![X21]:(![X22]:(hBOOL(hAPP_state_bool(hAPP_a2036067514e_bool(X38,X22),X18))=>hBOOL(hAPP_state_bool(hAPP_a2036067514e_bool(X39,X22),X21)))=>hBOOL(hAPP_state_bool(hAPP_a2036067514e_bool(X15,X17),X21)))))=>hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X3),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(X13,X14,X15)),bot_bo797238721a_bool))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_123_conseq)).
fof(c_0_1, axiom, (![X291]:![X292]:hAPP_a2036067514e_bool(hAPP_f762886889e_bool(cOMBK_1458035955bool_a,X291),X292)=X291), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_COMBK_1_1_COMBK_000tc__fun_Itc__Com__Ostate_Mtc__HOL__Obool_J_000t__a_U)).
fof(c_0_2, axiom, (![X291]:![X292]:(is_bool(X291)=>hAPP_state_bool(hAPP_b2019457360e_bool(cOMBK_bool_state,X291),X292)=X291)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_COMBK_1_1_COMBK_000tc__HOL__Obool_000tc__Com__Ostate_U)).
fof(c_0_3, hypothesis, (is_bool(fFalse)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', gsy_c_fFalse)).
fof(c_0_4, axiom, (~(hBOOL(fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_fFalse_1_1_U)).
fof(c_0_5, conjecture, (hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(g),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(hAPP_f762886889e_bool(cOMBK_1458035955bool_a,hAPP_b2019457360e_bool(cOMBK_bool_state,fFalse)),c,hAPP_f762886889e_bool(hAPP_f1261923407e_bool(cOMBC_892787026e_bool,hAPP_f963367678e_bool(hAPP_f375255701e_bool(cOMBB_145932198bool_a,cOMBS_1378840469l_bool),hAPP_f1509969235l_bool(hAPP_f1178339559l_bool(cOMBB_1355796797bool_a,hAPP_f1561913689l_bool(cOMBB_188601460_state,fconj)),p))),hAPP_f1759915619e_bool(hAPP_f2073279419e_bool(cOMBB_160679318_state,fNot),b)))),bot_bo797238721a_bool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_6, plain, (![X40]:![X41]:![X42]:![X43]:![X46]:![X47]:![X49]:((hBOOL(hAPP_state_bool(hAPP_a2036067514e_bool(X43,esk25_4(X40,X41,X42,X43)),esk26_4(X40,X41,X42,X43)))|hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X41),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(X43,X42,X40)),bot_bo797238721a_bool))))&((((~hBOOL(hAPP_state_bool(hAPP_a2036067514e_bool(X46,X49),esk26_4(X40,X41,X42,X43)))|hBOOL(hAPP_state_bool(hAPP_a2036067514e_bool(X47,X49),esk27_6(X40,X41,X42,X43,X46,X47))))|~hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X41),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(X46,X42,X47)),bot_bo797238721a_bool))))|hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X41),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(X43,X42,X40)),bot_bo797238721a_bool))))&((~hBOOL(hAPP_state_bool(hAPP_a2036067514e_bool(X40,esk25_4(X40,X41,X42,X43)),esk27_6(X40,X41,X42,X43,X46,X47)))|~hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X41),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(X46,X42,X47)),bot_bo797238721a_bool))))|hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X41),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(X43,X42,X40)),bot_bo797238721a_bool))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])).
fof(c_0_7, plain, (![X293]:![X294]:hAPP_a2036067514e_bool(hAPP_f762886889e_bool(cOMBK_1458035955bool_a,X293),X294)=X293), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_8, plain, (![X293]:![X294]:(~is_bool(X293)|hAPP_state_bool(hAPP_b2019457360e_bool(cOMBK_bool_state,X293),X294)=X293)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_9,plain,(hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X1),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(X2,X3,X4)),bot_bo797238721a_bool)))|hBOOL(hAPP_state_bool(hAPP_a2036067514e_bool(X2,esk25_4(X4,X1,X3,X2)),esk26_4(X4,X1,X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(hAPP_a2036067514e_bool(hAPP_f762886889e_bool(cOMBK_1458035955bool_a,X1),X2)=X1), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,plain,(hAPP_state_bool(hAPP_b2019457360e_bool(cOMBK_bool_state,X1),X2)=X1|~is_bool(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12,hypothesis,(is_bool(fFalse)), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_13, plain, (~hBOOL(fFalse)), inference(fof_simplification,[status(thm)],[c_0_4])).
fof(c_0_14, negated_conjecture, (~hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(g),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(hAPP_f762886889e_bool(cOMBK_1458035955bool_a,hAPP_b2019457360e_bool(cOMBK_bool_state,fFalse)),c,hAPP_f762886889e_bool(hAPP_f1261923407e_bool(cOMBC_892787026e_bool,hAPP_f963367678e_bool(hAPP_f375255701e_bool(cOMBB_145932198bool_a,cOMBS_1378840469l_bool),hAPP_f1509969235l_bool(hAPP_f1178339559l_bool(cOMBB_1355796797bool_a,hAPP_f1561913689l_bool(cOMBB_188601460_state,fconj)),p))),hAPP_f1759915619e_bool(hAPP_f2073279419e_bool(cOMBB_160679318_state,fNot),b)))),bot_bo797238721a_bool)))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_5])])).
cnf(c_0_15,plain,(hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X1),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(hAPP_f762886889e_bool(cOMBK_1458035955bool_a,X2),X3,X4)),bot_bo797238721a_bool)))|hBOOL(hAPP_state_bool(X2,esk26_4(X4,X1,X3,hAPP_f762886889e_bool(cOMBK_1458035955bool_a,X2))))), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_16,hypothesis,(hAPP_state_bool(hAPP_b2019457360e_bool(cOMBK_bool_state,fFalse),X1)=fFalse), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17,plain,(~hBOOL(fFalse)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,negated_conjecture,(~hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(g),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(hAPP_f762886889e_bool(cOMBK_1458035955bool_a,hAPP_b2019457360e_bool(cOMBK_bool_state,fFalse)),c,hAPP_f762886889e_bool(hAPP_f1261923407e_bool(cOMBC_892787026e_bool,hAPP_f963367678e_bool(hAPP_f375255701e_bool(cOMBB_145932198bool_a,cOMBS_1378840469l_bool),hAPP_f1509969235l_bool(hAPP_f1178339559l_bool(cOMBB_1355796797bool_a,hAPP_f1561913689l_bool(cOMBB_188601460_state,fconj)),p))),hAPP_f1759915619e_bool(hAPP_f2073279419e_bool(cOMBB_160679318_state,fNot),b)))),bot_bo797238721a_bool)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,hypothesis,(hBOOL(hAPP_f1695230391l_bool(hoare_2102800559rivs_a(X1),hAPP_f20753329a_bool(hAPP_H1743777351a_bool(insert956547291iple_a,hoare_1916936827iple_a(hAPP_f762886889e_bool(cOMBK_1458035955bool_a,hAPP_b2019457360e_bool(cOMBK_bool_state,fFalse)),X2,X3)),bot_bo797238721a_bool)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
cnf(c_0_20,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19])]), ['proof']).
# SZS output end CNFRefutation.
EOF
