% Problem    : GEO127+1 : TPTP v6.1.0. Released v2.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n154.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:19:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_B42_F1_PI_SE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X9]:![X3]:(incident_o(X3,X9)<=>?[X5]:(ordered_by(X9,X3,X5)|ordered_by(X9,X5,X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', theorem_4_12)).
fof(c_0_1, axiom, (![X9]:![X3]:![X5]:(ordered_by(X9,X3,X5)=>(incident_o(X3,X9)&incident_o(X5,X9)))), file('/export/starexec/sandbox/benchmark/Axioms/GEO004+2.ax', o1)).
fof(c_0_2, axiom, (![X3]:![X9]:(start_point(X3,X9)<=>(incident_o(X3,X9)&![X5]:((X3!=X5&incident_o(X5,X9))=>ordered_by(X9,X3,X5))))), file('/export/starexec/sandbox/benchmark/Axioms/GEO004+2.ax', start_point_defn)).
fof(c_0_3, axiom, (![X9]:?[X3]:start_point(X3,X9)), file('/export/starexec/sandbox/benchmark/Axioms/GEO004+2.ax', o4)).
fof(c_0_4, axiom, (![X3]:![X1]:(inner_point(X3,X1)<=>(incident_c(X3,X1)&~(end_point(X3,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/GEO004+0.ax', inner_point_defn)).
fof(c_0_5, axiom, (![X9]:?[X1]:(open(X1)&![X3]:(incident_o(X3,X9)<=>incident_c(X3,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/GEO004+2.ax', o2)).
fof(c_0_6, axiom, (![X1]:?[X3]:inner_point(X3,X1)), file('/export/starexec/sandbox/benchmark/Axioms/GEO004+0.ax', c3)).
fof(c_0_7, axiom, (![X3]:![X1]:(end_point(X3,X1)<=>(incident_c(X3,X1)&![X2]:![X4]:((((part_of(X2,X1)&part_of(X4,X1))&incident_c(X3,X2))&incident_c(X3,X4))=>(part_of(X2,X4)|part_of(X4,X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/GEO004+0.ax', end_point_defn)).
fof(c_0_8, axiom, (![X1]:(closed(X1)<=>~(?[X3]:end_point(X3,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/GEO004+0.ax', closed_defn)).
fof(c_0_9, axiom, (![X1]:(open(X1)<=>?[X3]:end_point(X3,X1))), file('/export/starexec/sandbox/benchmark/Axioms/GEO004+0.ax', open_defn)).
fof(c_0_10, negated_conjecture, (~(![X9]:![X3]:(incident_o(X3,X9)<=>?[X5]:(ordered_by(X9,X3,X5)|ordered_by(X9,X5,X3))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_11, negated_conjecture, (![X12]:![X13]:(((~ordered_by(esk25_0,esk26_0,X12)|~incident_o(esk26_0,esk25_0))&(~ordered_by(esk25_0,X13,esk26_0)|~incident_o(esk26_0,esk25_0)))&(incident_o(esk26_0,esk25_0)|(ordered_by(esk25_0,esk26_0,esk27_0)|ordered_by(esk25_0,esk28_0,esk26_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])])).
fof(c_0_12, plain, (![X10]:![X11]:![X12]:((incident_o(X11,X10)|~ordered_by(X10,X11,X12))&(incident_o(X12,X10)|~ordered_by(X10,X11,X12)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
cnf(c_0_13,negated_conjecture,(~incident_o(esk26_0,esk25_0)|~ordered_by(esk25_0,X1,esk26_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14,plain,(incident_o(X3,X1)|~ordered_by(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_15, plain, (![X10]:![X11]:![X12]:![X13]:![X14]:(((incident_o(X10,X11)|~start_point(X10,X11))&(((X10=X12|~incident_o(X12,X11))|ordered_by(X11,X10,X12))|~start_point(X10,X11)))&((((X13!=esk15_2(X13,X14)|~incident_o(X13,X14))|start_point(X13,X14))&((incident_o(esk15_2(X13,X14),X14)|~incident_o(X13,X14))|start_point(X13,X14)))&((~ordered_by(X14,X13,esk15_2(X13,X14))|~incident_o(X13,X14))|start_point(X13,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
cnf(c_0_16,negated_conjecture,(ordered_by(esk25_0,esk28_0,esk26_0)|ordered_by(esk25_0,esk26_0,esk27_0)|incident_o(esk26_0,esk25_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(~ordered_by(esk25_0,X1,esk26_0)), inference(csr,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18,plain,(incident_o(X2,X1)|~ordered_by(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(ordered_by(X2,X1,X3)|X1=X3|~start_point(X1,X2)|~incident_o(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20,negated_conjecture,(incident_o(esk26_0,esk25_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_16, c_0_17]), c_0_18])).
fof(c_0_21, plain, (![X10]:start_point(esk20_1(X10),X10)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_3])])).
cnf(c_0_22,negated_conjecture,(X1=esk26_0|~start_point(X1,esk25_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_19]), c_0_20])])).
cnf(c_0_23,plain,(start_point(esk20_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_24,negated_conjecture,(~incident_o(esk26_0,esk25_0)|~ordered_by(esk25_0,esk26_0,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25,negated_conjecture,(esk20_1(esk25_0)=esk26_0), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_26, plain, (![X3]:![X1]:(inner_point(X3,X1)<=>(incident_c(X3,X1)&~end_point(X3,X1)))), inference(fof_simplification,[status(thm)],[c_0_4])).
cnf(c_0_27,negated_conjecture,(~ordered_by(esk25_0,esk26_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_20])])).
cnf(c_0_28,negated_conjecture,(start_point(esk26_0,esk25_0)), inference(spm,[status(thm)],[c_0_23, c_0_25])).
fof(c_0_29, plain, (![X10]:![X12]:![X13]:(open(esk17_1(X10))&((~incident_o(X12,X10)|incident_c(X12,esk17_1(X10)))&(~incident_c(X13,esk17_1(X10))|incident_o(X13,X10))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
fof(c_0_30, plain, (![X4]:![X5]:![X6]:![X7]:(((incident_c(X4,X5)|~inner_point(X4,X5))&(~end_point(X4,X5)|~inner_point(X4,X5)))&((~incident_c(X6,X7)|end_point(X6,X7))|inner_point(X6,X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])])])).
fof(c_0_31, plain, (![X4]:inner_point(esk8_1(X4),X4)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_6])])).
fof(c_0_32, plain, (![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:(((incident_c(X5,X6)|~end_point(X5,X6))&(((((~part_of(X7,X6)|~part_of(X8,X6))|~incident_c(X5,X7))|~incident_c(X5,X8))|(part_of(X7,X8)|part_of(X8,X7)))|~end_point(X5,X6)))&((((((part_of(esk3_2(X9,X10),X10)|~incident_c(X9,X10))|end_point(X9,X10))&((part_of(esk4_2(X9,X10),X10)|~incident_c(X9,X10))|end_point(X9,X10)))&((incident_c(X9,esk3_2(X9,X10))|~incident_c(X9,X10))|end_point(X9,X10)))&((incident_c(X9,esk4_2(X9,X10))|~incident_c(X9,X10))|end_point(X9,X10)))&(((~part_of(esk3_2(X9,X10),esk4_2(X9,X10))|~incident_c(X9,X10))|end_point(X9,X10))&((~part_of(esk4_2(X9,X10),esk3_2(X9,X10))|~incident_c(X9,X10))|end_point(X9,X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])])).
fof(c_0_33, plain, (![X4]:![X5]:![X6]:((~closed(X4)|~end_point(X5,X4))&(end_point(esk6_1(X6),X6)|closed(X6)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
cnf(c_0_34,negated_conjecture,(esk26_0=X1|~incident_o(X1,esk25_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_19]), c_0_28])])).
cnf(c_0_35,plain,(incident_o(X1,X2)|~incident_c(X1,esk17_1(X2))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36,plain,(incident_c(X1,X2)|~inner_point(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37,plain,(inner_point(esk8_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38,plain,(incident_c(X1,X2)|~end_point(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39,plain,(closed(X1)|end_point(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40,plain,(~inner_point(X1,X2)|~end_point(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41,negated_conjecture,(esk26_0=X1|~incident_c(X1,esk17_1(esk25_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42,plain,(incident_c(esk8_1(X1),X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_43, plain, (![X4]:![X6]:![X7]:((~open(X4)|end_point(esk7_1(X4),X4))&(~end_point(X7,X6)|open(X6)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
cnf(c_0_44,plain,(closed(X1)|incident_c(esk6_1(X1),X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45,plain,(~end_point(esk8_1(X1),X1)), inference(spm,[status(thm)],[c_0_40, c_0_37])).
cnf(c_0_46,negated_conjecture,(esk8_1(esk17_1(esk25_0))=esk26_0), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47,plain,(~end_point(X1,X2)|~closed(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_48,plain,(end_point(esk7_1(X1),X1)|~open(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49,negated_conjecture,(esk6_1(esk17_1(esk25_0))=esk26_0|closed(esk17_1(esk25_0))), inference(spm,[status(thm)],[c_0_41, c_0_44])).
cnf(c_0_50,negated_conjecture,(~end_point(esk26_0,esk17_1(esk25_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51,plain,(~open(X1)|~closed(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52,negated_conjecture,(closed(esk17_1(esk25_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_49]), c_0_50])).
cnf(c_0_53,plain,(open(esk17_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_54,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])]), ['proof']).
# SZS output end CNFRefutation.
EOF
