% Problem    : SWW309+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n115.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:53:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X5]:c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),hAPP(c_COMBK(tc_HOL_Obool,X5),c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__def)).
fof(c_0_1, axiom, (![X10]:![X5]:hAPP(c_Set_OCollect(X5),X10)=X10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__def)).
fof(c_0_2, axiom, (![X10]:![X5]:(c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),X10)<=>![X3]:~(hBOOL(hAPP(X10,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__Collect__eq)).
fof(c_0_3, axiom, (![X311]:![X310]:![X5]:![X66]:hAPP(hAPP(c_COMBK(X66,X5),X310),X311)=X310), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBK__1)).
fof(c_0_4, axiom, (![X10]:![X5]:hAPP(c_Set_OCollect(X5),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X5),c_fNot),X10))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),hAPP(c_Set_OCollect(X5),X10))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__neg__eq)).
fof(c_0_5, axiom, (![X1]:![X2]:(![X3]:hAPP(X2,X3)=hAPP(X1,X3)=>X2=X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_ext)).
fof(c_0_6, axiom, (![X23]:![X5]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),X23))=X23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__complement)).
fof(c_0_7, axiom, (![X10]:(~(hBOOL(hAPP(c_fNot,X10)))|~(hBOOL(X10)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__1)).
fof(c_0_8, axiom, (![X5]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X5,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__empty__eq)).
fof(c_0_9, axiom, (![X5]:c_Orderings_Otop__class_Otop(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),hAPP(c_COMBK(tc_HOL_Obool,X5),c_fTrue))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__def)).
fof(c_0_10, axiom, (![X115]:![X116]:(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X116),X115)=X116=>X115=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_add__eq__self__zero)).
fof(c_0_11, axiom, (c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_transfer__nat__int__numerals_I1_J)).
fof(c_0_12, axiom, (c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Pls__def)).
fof(c_0_13, axiom, (![X116]:c_Groups_Ozero__class_Ozero(tc_Nat_Onat)!=hAPP(c_Nat_OSuc,X116)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Zero__not__Suc)).
fof(c_0_14, axiom, (![X35]:![X27]:![X10]:![X51]:![X5]:![X21]:hAPP(hAPP(hAPP(c_COMBB(X21,X5,X51),X10),X27),X35)=hAPP(X10,hAPP(X27,X35))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBB__1)).
fof(c_0_15, axiom, (![X115]:![X116]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X116),hAPP(c_Nat_OSuc,X115))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X116),X115))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_add__Suc__right)).
fof(c_0_16, axiom, (![X233]:![X66]:hAPP(c_List_Orev(X66),hAPP(c_List_Orev(X66),X233))=X233), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rev__rev__ident)).
fof(c_0_17, axiom, (![X116]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X116),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X116), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Nat_Oadd__0__right)).
fof(c_0_18, axiom, (![X115]:X115!=hAPP(c_Nat_OSuc,X115)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_n__not__Suc__n)).
fof(c_0_19, plain, (![X6]:c_Orderings_Obot__class_Obot(tc_fun(X6,tc_HOL_Obool))=hAPP(c_Set_OCollect(X6),hAPP(c_COMBK(tc_HOL_Obool,X6),c_fFalse))), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_20, plain, (![X11]:![X12]:hAPP(c_Set_OCollect(X12),X11)=X11), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_21, plain, (![X10]:![X5]:(c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),X10)<=>![X3]:~hBOOL(hAPP(X10,X3)))), inference(fof_simplification,[status(thm)],[c_0_2])).
fof(c_0_22, plain, (![X312]:![X313]:![X314]:![X315]:hAPP(hAPP(c_COMBK(X315,X314),X313),X312)=X313), inference(variable_rename,[status(thm)],[c_0_3])).
cnf(c_0_23,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24,plain,(hAPP(c_Set_OCollect(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, (![X11]:![X12]:hAPP(c_Set_OCollect(X12),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X12),c_fNot),X11))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X12,tc_HOL_Obool)),hAPP(c_Set_OCollect(X12),X11))), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_26, plain, (![X11]:![X12]:![X13]:![X14]:![X16]:((c_Orderings_Obot__class_Obot(tc_fun(X12,tc_HOL_Obool))!=hAPP(c_Set_OCollect(X12),X11)|~hBOOL(hAPP(X11,X13)))&(hBOOL(hAPP(X14,esk11_1(X14)))|c_Orderings_Obot__class_Obot(tc_fun(X16,tc_HOL_Obool))=hAPP(c_Set_OCollect(X16),X14)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).
fof(c_0_27, plain, (![X4]:![X5]:(hAPP(X5,esk1_2(X4,X5))!=hAPP(X4,esk1_2(X4,X5))|X5=X4)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_28,plain,(hAPP(hAPP(c_COMBK(X1,X2),X3),X4)=X3), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_30, plain, (![X24]:![X25]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X25,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X25,tc_HOL_Obool)),X24))=X24), inference(variable_rename,[status(thm)],[c_0_6])).
cnf(c_0_31,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Set_OCollect(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, (![X10]:(~hBOOL(hAPP(c_fNot,X10))|~hBOOL(X10))), inference(fof_simplification,[status(thm)],[c_0_7])).
cnf(c_0_33,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),X2)|hBOOL(hAPP(X2,esk11_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_34, plain, (![X6]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X6,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X6,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X6,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_35,plain,(X1=X2|hAPP(X1,esk1_2(X2,X1))!=hAPP(X2,esk1_2(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)),X2)=c_fFalse), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_37, plain, (![X6]:c_Orderings_Otop__class_Otop(tc_fun(X6,tc_HOL_Obool))=hAPP(c_Set_OCollect(X6),hAPP(c_COMBK(tc_HOL_Obool,X6),c_fTrue))), inference(variable_rename,[status(thm)],[c_0_9])).
fof(c_0_38, plain, (![X117]:![X118]:(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X118),X117)!=X118|X117=c_Groups_Ozero__class_Ozero(tc_Nat_Onat))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
cnf(c_0_39,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_40,plain,(c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_41, plain, (![X117]:c_Groups_Ozero__class_Ozero(tc_Nat_Onat)!=hAPP(c_Nat_OSuc,X117)), inference(variable_rename,[status(thm)],[c_0_13])).
fof(c_0_42, plain, (![X52]:![X53]:![X54]:![X55]:![X56]:![X57]:hAPP(hAPP(hAPP(c_COMBB(X57,X56,X55),X54),X53),X52)=hAPP(X54,hAPP(X53,X52))), inference(variable_rename,[status(thm)],[c_0_14])).
cnf(c_0_43,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_44,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2)=hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_24]), c_0_24])).
fof(c_0_45, plain, (![X11]:(~hBOOL(hAPP(c_fNot,X11))|~hBOOL(X11))), inference(variable_rename,[status(thm)],[c_0_32])).
cnf(c_0_46,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hBOOL(hAPP(X1,esk11_1(X1)))), inference(rw,[status(thm)],[c_0_33, c_0_24])).
cnf(c_0_47,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hAPP(X1,esk1_2(c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool)),X1))!=c_fFalse), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_49,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_50,plain,(X1=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X2),X1)!=X2), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_51,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_52, plain, (![X117]:![X118]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X118),hAPP(c_Nat_OSuc,X117))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X118),X117))), inference(variable_rename,[status(thm)],[c_0_15])).
cnf(c_0_53,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)!=hAPP(c_Nat_OSuc,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_54,plain,(hAPP(hAPP(hAPP(c_COMBB(X1,X2,X3),X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_55,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_44])).
fof(c_0_56, plain, (![X234]:![X235]:hAPP(c_List_Orev(X235),hAPP(c_List_Orev(X235),X234))=X234), inference(variable_rename,[status(thm)],[c_0_16])).
cnf(c_0_57,plain,(~hBOOL(X1)|~hBOOL(hAPP(c_fNot,X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_58,plain,(c_fFalse=X1|hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_46]), c_0_36]), c_0_28])).
cnf(c_0_59,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_47, c_0_44])).
cnf(c_0_60,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_48, c_0_36])).
cnf(c_0_61,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue)=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_49, c_0_24])).
cnf(c_0_62,plain,(hAPP(c_Int_Onat,c_Int_OPls)=X1|hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X2),X1)!=X2), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_63,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),hAPP(c_Nat_OSuc,X2))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64,plain,(hAPP(c_Nat_OSuc,X1)!=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_53, c_0_51])).
fof(c_0_65, plain, (![X117]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X117),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X117), inference(variable_rename,[status(thm)],[c_0_17])).
cnf(c_0_66,plain,(hAPP(c_fNot,hAPP(c_fNot,hAPP(X1,X2)))=hAPP(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_54])).
cnf(c_0_67,plain,(hAPP(c_List_Orev(X1),hAPP(c_List_Orev(X1),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68,plain,(hAPP(c_fNot,X1)=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_69,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_70,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)),X2)=c_fTrue), inference(spm,[status(thm)],[c_0_28, c_0_61])).
cnf(c_0_71,plain,(hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2))!=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])).
cnf(c_0_72,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X1), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_73, plain, (![X116]:X116!=hAPP(c_Nat_OSuc,X116)), inference(variable_rename,[status(thm)],[c_0_18])).
cnf(c_0_74,plain,(hAPP(c_fNot,hAPP(c_fNot,X1))=X1), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_75,plain,(hAPP(c_fNot,X1)=c_fFalse|c_fFalse=X1), inference(spm,[status(thm)],[c_0_68, c_0_58])).
cnf(c_0_76,plain,(hAPP(c_fNot,c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_69]), c_0_70]), c_0_36])).
cnf(c_0_77,plain,(hAPP(c_Nat_OSuc,hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2)))!=X1), inference(spm,[status(thm)],[c_0_71, c_0_63])).
cnf(c_0_78,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),hAPP(c_Int_Onat,c_Int_OPls))=X1), inference(rw,[status(thm)],[c_0_72, c_0_51])).
cnf(c_0_79,plain,(X1!=hAPP(c_Nat_OSuc,X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_80,plain,(c_fFalse=X1|c_fTrue=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
cnf(c_0_81,plain,(hAPP(c_Nat_OSuc,hAPP(c_Nat_OSuc,X1))!=X1), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_82,plain,(hAPP(c_Nat_OSuc,X1)=c_fFalse|c_fTrue!=X1), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_83,plain,(hAPP(c_Nat_OSuc,c_fFalse)!=X1|c_fTrue!=X1), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_84,plain,(c_fTrue!=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_80]), c_0_79])).
cnf(c_0_85,plain,($false), inference(sr,[status(thm)],[c_0_70, c_0_84]), ['proof']).
# SZS output end CNFRefutation.
EOF
