% Problem    : SWW097+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n101.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:06:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_C45_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (?[X5]:(false2=phi(f7(X5))&~(?[X10]:forallprefers(f7(X10),f7(X5))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_false2)).
fof(c_0_1, axiom, (![X5]:f7(X5)=lazy_impl(prop(X5),X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_f7)).
fof(c_0_2, axiom, (![X1]:(prop(X1)=false<=>~(bool(X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', prop_false)).
fof(c_0_3, axiom, (![X1]:((d(X1)&phi(X1)=X1)|(~(d(X1))&phi(X1)=err))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_phi)).
fof(c_0_4, axiom, (![X4]:lazy_impl(false,X4)=true), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', lazy_impl_axiom2)).
fof(c_0_5, axiom, (![X1]:![X2]:(forallprefers(X1,X2)<=>(((~(d(X1))&d(X2))|(((d(X1)&d(X2))&~(bool(X1)))&bool(X2)))|(X1=false&X2=true)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_forallprefers)).
fof(c_0_6, axiom, (![X1]:(prop(X1)=true<=>bool(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', prop_true)).
fof(c_0_7, axiom, (![X4]:lazy_impl(true,X4)=phi(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', lazy_impl_axiom3)).
fof(c_0_8, axiom, (![X1]:(bool(X1)<=>(X1=false|X1=true))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_bool)).
fof(c_0_9, axiom, (![X3]:![X4]:(~(bool(X3))=>lazy_impl(X3,X4)=phi(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', lazy_impl_axiom1)).
fof(c_0_10, axiom, (((d(true)&d(false))&d(err))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', false_true_err_in_d)).
fof(c_0_11, conjecture, (![X5]:![X6]:lazy_and1(X5,X6)=lazy_and2(X5,X6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', lazy_and1_lazy_and2)).
fof(c_0_12, axiom, (![X5]:![X6]:?[X7]:(lazy_and2(X5,X6)=phi(f2(X5,X6,X7))&~(?[X8]:forallprefers(f2(X5,X6,X8),f2(X5,X6,X7))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_lazy_and2)).
fof(c_0_13, axiom, (![X3]:![X4]:(~(bool(X3))=>impl(X3,X4)=phi(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', impl_axiom1)).
fof(c_0_14, axiom, (![X5]:![X6]:![X7]:f2(X5,X6,X7)=lazy_impl(prop(X7),impl(lazy_impl(X5,impl(X6,X7)),X7))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_f2)).
fof(c_0_15, axiom, (((true!=false&true!=err)&false!=err)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', distinct_false_true_err)).
fof(c_0_16, axiom, (![X3]:![X4]:(~(bool(X3))=>lazy_and1(X3,X4)=phi(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', lazy_and1_axiom1)).
fof(c_0_17, axiom, (![X4]:(bool(X4)=>impl(true,X4)=X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', impl_axiom4)).
fof(c_0_18, axiom, (![X4]:lazy_and1(true,X4)=phi(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', lazy_and1_axiom3)).
fof(c_0_19, axiom, (![X4]:(bool(X4)=>impl(false,X4)=true)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', impl_axiom3)).
fof(c_0_20, axiom, (![X4]:lazy_and1(false,X4)=false), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', lazy_and1_axiom2)).
fof(c_0_21, plain, (![X12]:(false2=phi(f7(esk7_0))&~forallprefers(f7(X12),f7(esk7_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])).
fof(c_0_22, plain, (![X6]:f7(X6)=lazy_impl(prop(X6),X6)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_23, plain, (![X1]:(prop(X1)=false<=>~bool(X1))), inference(fof_simplification,[status(thm)],[c_0_2])).
fof(c_0_24, plain, (![X1]:((d(X1)&phi(X1)=X1)|(~d(X1)&phi(X1)=err))), inference(fof_simplification,[status(thm)],[c_0_3])).
cnf(c_0_25,plain,(~forallprefers(f7(X1),f7(esk7_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26,plain,(f7(X1)=lazy_impl(prop(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, (![X2]:![X3]:((prop(X2)!=false|~bool(X2))&(bool(X3)|prop(X3)=false))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).
fof(c_0_28, plain, (![X5]:lazy_impl(false,X5)=true), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_29, plain, (![X1]:![X2]:(forallprefers(X1,X2)<=>(((~d(X1)&d(X2))|(((d(X1)&d(X2))&~bool(X1))&bool(X2)))|(X1=false&X2=true)))), inference(fof_simplification,[status(thm)],[c_0_5])).
fof(c_0_30, plain, (![X2]:![X3]:((prop(X2)!=true|bool(X2))&(~bool(X3)|prop(X3)=true))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
fof(c_0_31, plain, (![X2]:(((~d(X2)|d(X2))&(phi(X2)=err|d(X2)))&((~d(X2)|phi(X2)=X2)&(phi(X2)=err|phi(X2)=X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[c_0_24])])).
fof(c_0_32, plain, (![X5]:lazy_impl(true,X5)=phi(X5)), inference(variable_rename,[status(thm)],[c_0_7])).
cnf(c_0_33,plain,(~forallprefers(lazy_impl(prop(X1),X1),lazy_impl(prop(esk7_0),esk7_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26]), c_0_26])).
cnf(c_0_34,plain,(prop(X1)=false|bool(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35,plain,(lazy_impl(false,X1)=true), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_36, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=false|(d(X3)|~d(X3)))|~forallprefers(X3,X4))&((X4=true|(d(X3)|~d(X3)))|~forallprefers(X3,X4)))&(((X3=false|(d(X4)|~d(X3)))|~forallprefers(X3,X4))&((X4=true|(d(X4)|~d(X3)))|~forallprefers(X3,X4))))&(((X3=false|(~bool(X3)|~d(X3)))|~forallprefers(X3,X4))&((X4=true|(~bool(X3)|~d(X3)))|~forallprefers(X3,X4))))&(((X3=false|(bool(X4)|~d(X3)))|~forallprefers(X3,X4))&((X4=true|(bool(X4)|~d(X3)))|~forallprefers(X3,X4))))&((((((X3=false|(d(X3)|d(X4)))|~forallprefers(X3,X4))&((X4=true|(d(X3)|d(X4)))|~forallprefers(X3,X4)))&(((X3=false|(d(X4)|d(X4)))|~forallprefers(X3,X4))&((X4=true|(d(X4)|d(X4)))|~forallprefers(X3,X4))))&(((X3=false|(~bool(X3)|d(X4)))|~forallprefers(X3,X4))&((X4=true|(~bool(X3)|d(X4)))|~forallprefers(X3,X4))))&(((X3=false|(bool(X4)|d(X4)))|~forallprefers(X3,X4))&((X4=true|(bool(X4)|d(X4)))|~forallprefers(X3,X4)))))&((((d(X5)|~d(X6))|forallprefers(X5,X6))&((((~d(X5)|~d(X6))|bool(X5))|~bool(X6))|forallprefers(X5,X6)))&((X5!=false|X6!=true)|forallprefers(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])])).
cnf(c_0_37,plain,(prop(X1)=true|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38,plain,(phi(X1)=X1|~d(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39,plain,(lazy_impl(true,X1)=phi(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40,plain,(bool(esk7_0)|~forallprefers(lazy_impl(prop(X1),X1),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])).
cnf(c_0_41,plain,(forallprefers(X1,X2)|X2!=true|X1!=false), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_42, plain, (![X2]:![X3]:((~bool(X2)|(X2=false|X2=true))&((X3!=false|bool(X3))&(X3!=true|bool(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
cnf(c_0_43,plain,(~forallprefers(lazy_impl(true,X1),lazy_impl(prop(esk7_0),esk7_0))|~bool(X1)), inference(spm,[status(thm)],[c_0_33, c_0_37])).
cnf(c_0_44,plain,(lazy_impl(true,X1)=X1|~d(X1)), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45,plain,(bool(esk7_0)|~forallprefers(lazy_impl(true,X1),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_40, c_0_37])).
cnf(c_0_46,plain,(forallprefers(false,X1)|true!=X1), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_47,plain,(bool(X1)|X1!=false), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_48, plain, (![X3]:![X4]:(~bool(X3)=>lazy_impl(X3,X4)=phi(X3))), inference(fof_simplification,[status(thm)],[c_0_9])).
cnf(c_0_49,plain,(~forallprefers(X1,lazy_impl(prop(esk7_0),esk7_0))|~d(X1)|~bool(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50,plain,(bool(esk7_0)|~forallprefers(X1,true)|~d(X1)|~bool(X1)), inference(spm,[status(thm)],[c_0_45, c_0_44])).
cnf(c_0_51,plain,(forallprefers(false,true)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_52,plain,(d(false)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_53,plain,(bool(false)), inference(er,[status(thm)],[c_0_47])).
fof(c_0_54, plain, (![X5]:![X6]:(bool(X5)|lazy_impl(X5,X6)=phi(X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_48])])])])).
cnf(c_0_55,plain,(~forallprefers(X1,lazy_impl(true,esk7_0))|~d(X1)|~bool(esk7_0)|~bool(X1)), inference(spm,[status(thm)],[c_0_49, c_0_37])).
cnf(c_0_56,plain,(bool(esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_53])])).
cnf(c_0_57,plain,(X1=true|X1=false|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_58,plain,(lazy_impl(X1,X2)=phi(X1)|bool(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_59,plain,(~forallprefers(X1,lazy_impl(true,esk7_0))|~d(X1)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56])])).
cnf(c_0_60,plain,(d(true)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_61,plain,(esk7_0=false|true=esk7_0), inference(spm,[status(thm)],[c_0_57, c_0_56])).
cnf(c_0_62,plain,(lazy_impl(true,X1)=lazy_impl(X1,X2)|bool(X1)), inference(rw,[status(thm)],[c_0_58, c_0_39])).
cnf(c_0_63,plain,(~forallprefers(lazy_impl(true,X1),lazy_impl(true,esk7_0))|~bool(esk7_0)|~bool(X1)), inference(spm,[status(thm)],[c_0_43, c_0_37])).
cnf(c_0_64,plain,(~forallprefers(X1,esk7_0)|~d(esk7_0)|~d(X1)|~bool(X1)), inference(spm,[status(thm)],[c_0_59, c_0_44])).
cnf(c_0_65,plain,(esk7_0=false|d(esk7_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_66,plain,(~forallprefers(lazy_impl(true,prop(X1)),lazy_impl(prop(esk7_0),esk7_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_62]), c_0_43])).
fof(c_0_67, negated_conjecture, (~(![X5]:![X6]:lazy_and1(X5,X6)=lazy_and2(X5,X6))), inference(assume_negation,[status(cth)],[c_0_11])).
fof(c_0_68, plain, (![X9]:![X10]:![X12]:(lazy_and2(X9,X10)=phi(f2(X9,X10,esk2_2(X9,X10)))&~forallprefers(f2(X9,X10,X12),f2(X9,X10,esk2_2(X9,X10))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
fof(c_0_69, plain, (![X3]:![X4]:(~bool(X3)=>impl(X3,X4)=phi(X3))), inference(fof_simplification,[status(thm)],[c_0_13])).
cnf(c_0_70,plain,(~forallprefers(lazy_impl(true,X1),lazy_impl(true,esk7_0))|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_56])])).
cnf(c_0_71,plain,(esk7_0=false|~forallprefers(X1,esk7_0)|~d(X1)|~bool(X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_72,plain,(esk7_0=false|forallprefers(false,esk7_0)), inference(spm,[status(thm)],[c_0_51, c_0_61])).
cnf(c_0_73,plain,(~forallprefers(lazy_impl(true,prop(X1)),lazy_impl(true,esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_37]), c_0_56])])).
fof(c_0_74, negated_conjecture, (lazy_and1(esk8_0,esk9_0)!=lazy_and2(esk8_0,esk9_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_67])])])).
cnf(c_0_75,plain,(lazy_and2(X1,X2)=phi(f2(X1,X2,esk2_2(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_76, plain, (![X8]:![X9]:![X10]:f2(X8,X9,X10)=lazy_impl(prop(X10),impl(lazy_impl(X8,impl(X9,X10)),X10))), inference(variable_rename,[status(thm)],[c_0_14])).
fof(c_0_77, plain, (![X5]:![X6]:(bool(X5)|impl(X5,X6)=phi(X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_69])])])])).
cnf(c_0_78,plain,(phi(X1)=X1|phi(X1)=err), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_79,plain,(~forallprefers(lazy_impl(true,X1),esk7_0)|~d(esk7_0)|~bool(X1)), inference(spm,[status(thm)],[c_0_70, c_0_44])).
cnf(c_0_80,plain,(esk7_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_52]), c_0_53])])).
cnf(c_0_81,plain,(forallprefers(X1,X2)|bool(X1)|~bool(X2)|~d(X2)|~d(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_82,plain,(forallprefers(X1,X2)|d(X1)|~d(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_83,plain,(~forallprefers(lazy_impl(true,prop(X1)),esk7_0)|~d(esk7_0)), inference(spm,[status(thm)],[c_0_73, c_0_44])).
cnf(c_0_84,negated_conjecture,(lazy_and1(esk8_0,esk9_0)!=lazy_and2(esk8_0,esk9_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_85,plain,(lazy_impl(true,f2(X1,X2,esk2_2(X1,X2)))=lazy_and2(X1,X2)), inference(rw,[status(thm)],[c_0_75, c_0_39])).
cnf(c_0_86,plain,(f2(X1,X2,X3)=lazy_impl(prop(X3),impl(lazy_impl(X1,impl(X2,X3)),X3))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_87,plain,(impl(X1,X2)=phi(X1)|bool(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_88,plain,(lazy_impl(true,X1)=X1|lazy_impl(true,X1)=err), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_39]), c_0_39])).
cnf(c_0_89,plain,(~forallprefers(lazy_impl(true,X1),false)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_80]), c_0_80]), c_0_52])])).
cnf(c_0_90,plain,(forallprefers(X1,X2)|bool(X1)|~d(X2)|~bool(X2)), inference(csr,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_91,plain,(~forallprefers(lazy_impl(true,prop(X1)),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_80]), c_0_80]), c_0_52])])).
cnf(c_0_92,negated_conjecture,(lazy_impl(true,lazy_impl(prop(esk2_2(esk8_0,esk9_0)),impl(lazy_impl(esk8_0,impl(esk9_0,esk2_2(esk8_0,esk9_0))),esk2_2(esk8_0,esk9_0))))!=lazy_and1(esk8_0,esk9_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_85]), c_0_86])).
cnf(c_0_93,plain,(lazy_impl(true,X1)=impl(X1,X2)|bool(X1)), inference(rw,[status(thm)],[c_0_87, c_0_39])).
cnf(c_0_94,plain,(lazy_impl(true,X1)=X1|err!=X1), inference(ef,[status(thm)],[c_0_88])).
cnf(c_0_95,plain,(bool(lazy_impl(true,X1))|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_52]), c_0_53])])).
cnf(c_0_96,plain,(bool(lazy_impl(true,prop(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_90]), c_0_52]), c_0_53])])).
cnf(c_0_97,plain,(d(X1)|phi(X1)=err), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_98,negated_conjecture,(bool(lazy_impl(esk8_0,impl(esk9_0,esk2_2(esk8_0,esk9_0))))|lazy_impl(true,lazy_impl(prop(esk2_2(esk8_0,esk9_0)),lazy_impl(true,lazy_impl(esk8_0,impl(esk9_0,esk2_2(esk8_0,esk9_0))))))!=lazy_and1(esk8_0,esk9_0)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_99,plain,(lazy_impl(true,err)=err), inference(er,[status(thm)],[c_0_94])).
cnf(c_0_100,plain,(lazy_impl(true,X1)=false|lazy_impl(true,X1)=true|~bool(X1)), inference(spm,[status(thm)],[c_0_57, c_0_95])).
cnf(c_0_101,plain,(true!=err), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_102,plain,(false!=err), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_103,plain,(lazy_impl(true,prop(X1))=false|lazy_impl(true,prop(X1))=true), inference(spm,[status(thm)],[c_0_57, c_0_96])).
fof(c_0_104, plain, (![X3]:![X4]:(~bool(X3)=>lazy_and1(X3,X4)=phi(X3))), inference(fof_simplification,[status(thm)],[c_0_16])).
cnf(c_0_105,plain,(lazy_impl(true,X1)=err|d(X1)), inference(rw,[status(thm)],[c_0_97, c_0_39])).
cnf(c_0_106,negated_conjecture,(bool(lazy_impl(true,esk8_0))|lazy_impl(true,lazy_impl(prop(esk2_2(esk8_0,esk9_0)),lazy_impl(true,lazy_impl(true,esk8_0))))!=lazy_and1(esk8_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_62]), c_0_95])).
cnf(c_0_107,plain,(~bool(err)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_99]), c_0_101]), c_0_102])).
cnf(c_0_108,plain,(lazy_impl(true,prop(X1))=false|prop(X1)=true), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_103]), c_0_101])).
fof(c_0_109, plain, (![X5]:![X6]:(bool(X5)|lazy_and1(X5,X6)=phi(X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_104])])])])).
cnf(c_0_110,plain,(d(X1)|~forallprefers(err,false)|~bool(X1)), inference(spm,[status(thm)],[c_0_89, c_0_105])).
cnf(c_0_111,plain,(~forallprefers(f2(X1,X2,X3),f2(X1,X2,esk2_2(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_112,negated_conjecture,(d(esk8_0)|lazy_impl(true,lazy_impl(prop(esk2_2(esk8_0,esk9_0)),err))!=lazy_and1(esk8_0,esk9_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_105]), c_0_99]), c_0_107])).
cnf(c_0_113,plain,(prop(X1)=true|prop(X1)=false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_108]), c_0_102])).
cnf(c_0_114,plain,(lazy_and1(X1,X2)=phi(X1)|bool(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_115,plain,(d(X1)|bool(err)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_90]), c_0_52]), c_0_53])])).
cnf(c_0_116,plain,(~forallprefers(lazy_impl(prop(X3),impl(lazy_impl(X1,impl(X2,X3)),X3)),lazy_impl(prop(esk2_2(X1,X2)),impl(lazy_impl(X1,impl(X2,esk2_2(X1,X2))),esk2_2(X1,X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_86]), c_0_86])).
cnf(c_0_117,plain,(lazy_impl(err,X1)=err|bool(err)), inference(spm,[status(thm)],[c_0_62, c_0_99])).
cnf(c_0_118,negated_conjecture,(prop(esk2_2(esk8_0,esk9_0))=false|d(esk8_0)|lazy_and1(esk8_0,esk9_0)!=err), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_99]), c_0_99])).
cnf(c_0_119,plain,(lazy_impl(true,X1)=lazy_and1(X1,X2)|bool(X1)), inference(rw,[status(thm)],[c_0_114, c_0_39])).
cnf(c_0_120,plain,(d(X1)|~bool(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_115]), c_0_101]), c_0_102])).
cnf(c_0_121,plain,(bool(lazy_impl(X1,impl(X2,esk2_2(X1,X2))))|~forallprefers(lazy_impl(prop(X3),impl(lazy_impl(X1,impl(X2,X3)),X3)),lazy_impl(prop(esk2_2(X1,X2)),lazy_impl(true,lazy_impl(X1,impl(X2,esk2_2(X1,X2))))))), inference(spm,[status(thm)],[c_0_116, c_0_93])).
cnf(c_0_122,plain,(lazy_impl(err,X1)=err), inference(sr,[status(thm)],[c_0_117, c_0_107])).
cnf(c_0_123,negated_conjecture,(prop(esk2_2(esk8_0,esk9_0))=false|d(esk8_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_105]), c_0_120])).
cnf(c_0_124,plain,(~forallprefers(lazy_impl(prop(X1),impl(err,X1)),lazy_impl(prop(esk2_2(err,X2)),err))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_122]), c_0_99]), c_0_107])).
cnf(c_0_125,plain,(~forallprefers(lazy_impl(prop(X1),X1),lazy_impl(true,esk7_0))|~bool(esk7_0)), inference(spm,[status(thm)],[c_0_33, c_0_37])).
cnf(c_0_126,negated_conjecture,(d(esk8_0)|~forallprefers(lazy_impl(prop(X1),impl(lazy_impl(esk8_0,impl(esk9_0,X1)),X1)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_123]), c_0_35])).
cnf(c_0_127,plain,(lazy_impl(X1,X2)=err|d(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_62]), c_0_120])).
cnf(c_0_128,plain,(~forallprefers(lazy_impl(prop(X1),err),lazy_impl(prop(esk2_2(err,X2)),err))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_93]), c_0_99]), c_0_107])).
cnf(c_0_129,plain,(~forallprefers(lazy_impl(prop(X1),X1),lazy_impl(true,esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_56])])).
cnf(c_0_130,negated_conjecture,(d(esk8_0)|~forallprefers(lazy_impl(prop(X1),impl(err,X1)),true)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_131,plain,(~forallprefers(err,lazy_impl(prop(esk2_2(err,X1)),err))|~bool(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_37]), c_0_99])).
cnf(c_0_132,plain,(bool(X1)|X1!=true), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_133,plain,(~forallprefers(lazy_impl(prop(X1),X1),esk7_0)|~d(esk7_0)), inference(spm,[status(thm)],[c_0_129, c_0_44])).
cnf(c_0_134,negated_conjecture,(d(esk8_0)|~forallprefers(lazy_impl(prop(X1),err),true)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_93]), c_0_99]), c_0_107])).
cnf(c_0_135,plain,(bool(esk2_2(err,X1))|~forallprefers(err,true)|~bool(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_34]), c_0_35])).
cnf(c_0_136,plain,(bool(true)), inference(er,[status(thm)],[c_0_132])).
cnf(c_0_137,plain,(~forallprefers(lazy_impl(prop(X1),X1),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_80]), c_0_80]), c_0_52])])).
cnf(c_0_138,negated_conjecture,(d(esk8_0)|~forallprefers(err,true)|~bool(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_37]), c_0_99])).
cnf(c_0_139,plain,(bool(esk2_2(err,X1))|~bool(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_90]), c_0_60]), c_0_136])]), c_0_107])).
cnf(c_0_140,plain,(bool(lazy_impl(prop(X1),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_90]), c_0_52]), c_0_53])])).
cnf(c_0_141,negated_conjecture,(bool(lazy_impl(true,esk8_0))|lazy_impl(true,lazy_impl(prop(esk2_2(esk8_0,esk9_0)),lazy_impl(true,esk8_0)))!=lazy_and1(esk8_0,esk9_0)|~d(lazy_impl(true,esk8_0))), inference(spm,[status(thm)],[c_0_106, c_0_44])).
cnf(c_0_142,negated_conjecture,(d(esk8_0)|~bool(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_90]), c_0_60]), c_0_136])]), c_0_107])).
cnf(c_0_143,plain,(bool(esk2_2(err,X1))), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_144,negated_conjecture,(bool(esk8_0)|lazy_impl(true,lazy_impl(prop(esk2_2(esk8_0,esk9_0)),esk8_0))!=lazy_and1(esk8_0,esk9_0)|~d(esk8_0)), inference(spm,[status(thm)],[c_0_141, c_0_44])).
cnf(c_0_145,negated_conjecture,(d(esk8_0)), inference(spm,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_146,negated_conjecture,(bool(esk8_0)|lazy_impl(true,lazy_impl(prop(esk2_2(esk8_0,esk9_0)),esk8_0))!=lazy_and1(esk8_0,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_144, c_0_145])])).
cnf(c_0_147,negated_conjecture,(prop(esk2_2(esk8_0,esk9_0))=false|bool(esk8_0)|lazy_and1(esk8_0,esk9_0)!=lazy_impl(true,lazy_impl(true,esk8_0))), inference(spm,[status(thm)],[c_0_146, c_0_113])).
cnf(c_0_148,negated_conjecture,(prop(esk2_2(esk8_0,esk9_0))=false|bool(esk8_0)|lazy_impl(true,lazy_impl(true,esk8_0))!=lazy_impl(true,esk8_0)), inference(spm,[status(thm)],[c_0_147, c_0_119])).
cnf(c_0_149,negated_conjecture,(prop(esk2_2(esk8_0,esk9_0))=false|bool(esk8_0)|lazy_impl(true,esk8_0)!=esk8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_44]), c_0_145])])).
cnf(c_0_150,plain,(bool(X1)|~forallprefers(lazy_impl(prop(X2),impl(lazy_impl(true,X1),X2)),lazy_impl(prop(esk2_2(X1,X3)),impl(lazy_impl(X1,impl(X3,esk2_2(X1,X3))),esk2_2(X1,X3))))), inference(spm,[status(thm)],[c_0_116, c_0_62])).
cnf(c_0_151,negated_conjecture,(prop(esk2_2(esk8_0,esk9_0))=false|bool(esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_44]), c_0_145])])).
cnf(c_0_152,plain,(false2=phi(f7(esk7_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_153,negated_conjecture,(bool(esk8_0)|~forallprefers(lazy_impl(prop(X1),impl(lazy_impl(true,esk8_0),X1)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_35])).
cnf(c_0_154,plain,(lazy_impl(true,lazy_impl(prop(esk7_0),esk7_0))=false2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_39]), c_0_26])).
cnf(c_0_155,negated_conjecture,(bool(esk8_0)|~forallprefers(lazy_impl(prop(X1),impl(esk8_0,X1)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_44]), c_0_145])])).
cnf(c_0_156,plain,(lazy_impl(true,lazy_impl(true,esk7_0))=false2|~bool(esk7_0)), inference(spm,[status(thm)],[c_0_154, c_0_37])).
cnf(c_0_157,negated_conjecture,(bool(esk8_0)|~forallprefers(lazy_impl(prop(X1),lazy_impl(true,esk8_0)),true)), inference(spm,[status(thm)],[c_0_155, c_0_93])).
cnf(c_0_158,plain,(lazy_impl(true,lazy_impl(true,esk7_0))=false2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_156, c_0_56])])).
cnf(c_0_159,negated_conjecture,(bool(esk8_0)|~forallprefers(lazy_impl(prop(X1),esk8_0),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_44]), c_0_145])])).
cnf(c_0_160,plain,(lazy_impl(true,esk7_0)=false2|~d(lazy_impl(true,esk7_0))), inference(spm,[status(thm)],[c_0_44, c_0_158])).
cnf(c_0_161,negated_conjecture,(bool(esk8_0)|~forallprefers(lazy_impl(true,esk8_0),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_159, c_0_37])).
cnf(c_0_162,plain,(false2=esk7_0|~d(esk7_0)), inference(spm,[status(thm)],[c_0_160, c_0_44])).
cnf(c_0_163,negated_conjecture,(bool(esk8_0)|~forallprefers(esk8_0,true)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_44]), c_0_145])])).
cnf(c_0_164,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(prop(esk2_2(false,X2)),impl(true,esk2_2(false,X2))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_35]), c_0_35])).
fof(c_0_165, plain, (![X5]:(~bool(X5)|impl(true,X5)=X5)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
cnf(c_0_166,plain,(false2=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_162, c_0_80]), c_0_80]), c_0_52])])).
cnf(c_0_167,negated_conjecture,(bool(esk8_0)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_90]), c_0_60]), c_0_136])])).
fof(c_0_168, plain, (![X5]:lazy_and1(true,X5)=phi(X5)), inference(variable_rename,[status(thm)],[c_0_18])).
cnf(c_0_169,plain,(bool(esk2_2(false,X1))|~forallprefers(lazy_impl(prop(X2),impl(true,X2)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_34]), c_0_35])).
cnf(c_0_170,plain,(impl(true,X1)=X1|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_171,plain,(lazy_impl(true,lazy_impl(prop(false),false))=false), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_154, c_0_80]), c_0_80]), c_0_166])).
cnf(c_0_172,negated_conjecture,(bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk2_2(esk8_0,esk9_0)),impl(lazy_impl(esk8_0,lazy_impl(true,esk9_0)),esk2_2(esk8_0,esk9_0))))!=lazy_and1(esk8_0,esk9_0)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_173,negated_conjecture,(bool(esk8_0)), inference(spm,[status(thm)],[c_0_167, c_0_143])).
cnf(c_0_174,plain,(lazy_and1(true,X1)=phi(X1)), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_175,plain,(bool(esk2_2(false,X1))|~forallprefers(lazy_impl(prop(X2),X2),true)|~bool(X2)), inference(spm,[status(thm)],[c_0_169, c_0_170])).
cnf(c_0_176,plain,(lazy_impl(prop(false),false)=false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_171]), c_0_102])).
cnf(c_0_177,negated_conjecture,(d(esk9_0)|lazy_impl(true,lazy_impl(prop(esk2_2(esk8_0,esk9_0)),impl(lazy_impl(esk8_0,err),esk2_2(esk8_0,esk9_0))))!=lazy_and1(esk8_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_105]), c_0_120])).
cnf(c_0_178,negated_conjecture,(esk8_0=false|esk8_0=true), inference(spm,[status(thm)],[c_0_57, c_0_173])).
cnf(c_0_179,plain,(lazy_and1(true,X1)=lazy_impl(true,X1)), inference(rw,[status(thm)],[c_0_174, c_0_39])).
cnf(c_0_180,plain,(bool(esk2_2(false,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176]), c_0_51]), c_0_53])])).
cnf(c_0_181,negated_conjecture,(esk8_0=false|d(esk9_0)|lazy_impl(true,lazy_impl(prop(esk2_2(true,esk9_0)),impl(err,esk2_2(true,esk9_0))))!=lazy_impl(true,esk9_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178]), c_0_99]), c_0_179])).
cnf(c_0_182,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(true,impl(true,esk2_2(false,X2))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_37]), c_0_180])])).
cnf(c_0_183,plain,(bool(X1)|~forallprefers(lazy_impl(prop(X2),impl(lazy_impl(X3,lazy_impl(true,X1)),X2)),lazy_impl(prop(esk2_2(X3,X1)),impl(lazy_impl(X3,impl(X1,esk2_2(X3,X1))),esk2_2(X3,X1))))), inference(spm,[status(thm)],[c_0_116, c_0_93])).
cnf(c_0_184,negated_conjecture,(esk8_0=false|d(esk9_0)|lazy_impl(true,lazy_impl(prop(esk2_2(true,esk9_0)),err))!=lazy_impl(true,esk9_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_93]), c_0_99]), c_0_107])).
cnf(c_0_185,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(true,esk2_2(false,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_170]), c_0_180])])).
cnf(c_0_186,plain,(d(esk2_2(false,X1))), inference(spm,[status(thm)],[c_0_120, c_0_180])).
cnf(c_0_187,plain,(d(X1)|~forallprefers(lazy_impl(prop(X2),impl(lazy_impl(X3,err),X2)),lazy_impl(prop(esk2_2(X3,X1)),impl(lazy_impl(X3,impl(X1,esk2_2(X3,X1))),esk2_2(X3,X1))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_105]), c_0_120])).
cnf(c_0_188,negated_conjecture,(prop(esk2_2(true,esk9_0))=false|esk8_0=false|d(esk9_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_113]), c_0_99]), c_0_99]), c_0_105])).
cnf(c_0_189,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),esk2_2(false,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_44]), c_0_186])])).
cnf(c_0_190,negated_conjecture,(esk8_0=false|d(esk9_0)|~forallprefers(lazy_impl(prop(X1),impl(err,X1)),true)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_99]), c_0_35])).
cnf(c_0_191,plain,(~forallprefers(lazy_impl(prop(X1),X1),esk2_2(false,X2))|~bool(X1)), inference(spm,[status(thm)],[c_0_189, c_0_170])).
cnf(c_0_192,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(prop(esk2_2(true,esk9_0)),impl(lazy_impl(true,impl(esk9_0,esk2_2(true,esk9_0))),esk2_2(true,esk9_0))))!=lazy_impl(true,esk9_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_178]), c_0_179])).
cnf(c_0_193,negated_conjecture,(esk8_0=false|d(esk9_0)|~forallprefers(lazy_impl(prop(X1),err),true)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_93]), c_0_99]), c_0_107])).
cnf(c_0_194,plain,(~forallprefers(false2,lazy_impl(prop(esk7_0),esk7_0))|~bool(lazy_impl(true,esk7_0))), inference(spm,[status(thm)],[c_0_43, c_0_158])).
cnf(c_0_195,plain,(~forallprefers(false,esk2_2(false,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_176]), c_0_53])])).
cnf(c_0_196,plain,(esk2_2(false,X1)=false|esk2_2(false,X1)=true), inference(spm,[status(thm)],[c_0_57, c_0_180])).
cnf(c_0_197,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(prop(esk2_2(true,esk9_0)),impl(impl(esk9_0,esk2_2(true,esk9_0)),esk2_2(true,esk9_0))))!=lazy_impl(true,esk9_0)|~d(impl(esk9_0,esk2_2(true,esk9_0)))), inference(spm,[status(thm)],[c_0_192, c_0_44])).
cnf(c_0_198,negated_conjecture,(esk8_0=false|d(esk9_0)|~forallprefers(err,true)|~bool(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_37]), c_0_99])).
cnf(c_0_199,plain,(prop(X1)=false|~forallprefers(lazy_impl(true,X1),false)), inference(spm,[status(thm)],[c_0_137, c_0_113])).
cnf(c_0_200,plain,(~forallprefers(false2,lazy_impl(prop(esk7_0),esk7_0))|~d(esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_44]), c_0_56])])).
cnf(c_0_201,plain,(esk2_2(false,X1)=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196]), c_0_51])])).
cnf(c_0_202,negated_conjecture,(esk8_0=false|bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk2_2(true,esk9_0)),impl(lazy_impl(true,esk9_0),esk2_2(true,esk9_0))))!=lazy_impl(true,esk9_0)|~d(lazy_impl(true,esk9_0))), inference(spm,[status(thm)],[c_0_197, c_0_93])).
cnf(c_0_203,negated_conjecture,(esk8_0=false|d(esk9_0)|~bool(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_90]), c_0_60]), c_0_136])]), c_0_107])).
cnf(c_0_204,plain,(lazy_impl(true,X1)=X1|prop(X1)=false|~forallprefers(err,false)), inference(spm,[status(thm)],[c_0_199, c_0_88])).
cnf(c_0_205,plain,(~forallprefers(false,lazy_impl(prop(false),false))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_200, c_0_80]), c_0_80]), c_0_80]), c_0_52])]), c_0_166])).
cnf(c_0_206,plain,(lazy_impl(true,true)=true|lazy_impl(true,true)=false|prop(X1)=false), inference(spm,[status(thm)],[c_0_103, c_0_113])).
cnf(c_0_207,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),false)), inference(rw,[status(thm)],[c_0_189, c_0_201])).
cnf(c_0_208,negated_conjecture,(esk8_0=false|bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk2_2(true,esk9_0)),impl(esk9_0,esk2_2(true,esk9_0))))!=esk9_0|~d(esk9_0)), inference(spm,[status(thm)],[c_0_202, c_0_44])).
cnf(c_0_209,negated_conjecture,(esk8_0=false|d(esk9_0)), inference(spm,[status(thm)],[c_0_203, c_0_173])).
cnf(c_0_210,plain,(lazy_impl(true,X1)=X1|prop(X1)=false|bool(err)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_90]), c_0_52]), c_0_53])])).
cnf(c_0_211,plain,(~forallprefers(lazy_impl(prop(X1),impl(lazy_impl(X2,X1),X1)),lazy_impl(prop(esk2_2(X2,true)),impl(lazy_impl(X2,impl(true,esk2_2(X2,true))),esk2_2(X2,true))))|~bool(X1)), inference(spm,[status(thm)],[c_0_116, c_0_170])).
cnf(c_0_212,plain,(lazy_impl(true,true)=false|lazy_impl(true,true)=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206]), c_0_35]), c_0_51])])).
cnf(c_0_213,plain,(~forallprefers(lazy_impl(true,impl(true,X1)),false)|~bool(X1)), inference(spm,[status(thm)],[c_0_207, c_0_37])).
cnf(c_0_214,negated_conjecture,(esk8_0=false|bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk2_2(true,esk9_0)),impl(esk9_0,esk2_2(true,esk9_0))))!=esk9_0), inference(spm,[status(thm)],[c_0_208, c_0_209])).
cnf(c_0_215,plain,(lazy_impl(true,X1)=X1|prop(X1)=false), inference(sr,[status(thm)],[c_0_210, c_0_107])).
cnf(c_0_216,plain,(true!=false), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_217,plain,(lazy_impl(true,true)=false|~forallprefers(lazy_impl(prop(true),impl(true,true)),lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212]), c_0_136])])).
cnf(c_0_218,plain,(lazy_impl(prop(X1),X1)=false|lazy_impl(prop(X1),X1)=true), inference(spm,[status(thm)],[c_0_57, c_0_140])).
cnf(c_0_219,plain,(d(impl(true,X1))|~forallprefers(err,false)|~bool(X1)), inference(spm,[status(thm)],[c_0_213, c_0_105])).
cnf(c_0_220,negated_conjecture,(esk8_0=false|bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk2_2(true,esk9_0)),lazy_impl(true,esk9_0)))!=esk9_0), inference(spm,[status(thm)],[c_0_214, c_0_93])).
cnf(c_0_221,plain,(bool(lazy_impl(prop(X1),impl(true,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_90]), c_0_52]), c_0_53])])).
cnf(c_0_222,plain,(lazy_impl(X1,X2)=err|lazy_impl(X1,X2)=X1|bool(X1)), inference(spm,[status(thm)],[c_0_88, c_0_62])).
cnf(c_0_223,plain,(lazy_impl(true,X1)=X1|~bool(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_215]), c_0_216])).
cnf(c_0_224,plain,(lazy_impl(true,true)=false|~forallprefers(lazy_impl(true,impl(true,true)),lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_37]), c_0_136])])).
cnf(c_0_225,plain,(lazy_impl(true,X1)=true|lazy_impl(true,X1)=false|prop(X1)=false), inference(spm,[status(thm)],[c_0_218, c_0_113])).
cnf(c_0_226,plain,(~forallprefers(impl(true,X1),false)|~d(impl(true,X1))|~bool(X1)), inference(spm,[status(thm)],[c_0_213, c_0_44])).
cnf(c_0_227,plain,(d(impl(true,X1))|~bool(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_90]), c_0_52]), c_0_53])]), c_0_107])).
cnf(c_0_228,negated_conjecture,(esk8_0=false|bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk2_2(true,esk9_0)),esk9_0))!=esk9_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_44]), c_0_209])).
cnf(c_0_229,plain,(bool(esk2_2(X1,X2))|~forallprefers(lazy_impl(prop(X3),impl(lazy_impl(X1,impl(X2,X3)),X3)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_34]), c_0_35])).
cnf(c_0_230,plain,(prop(X1)=false|bool(lazy_impl(true,impl(true,X1)))), inference(spm,[status(thm)],[c_0_221, c_0_113])).
cnf(c_0_231,plain,(lazy_impl(X1,X2)=err|lazy_impl(true,X1)=X1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_222]), c_0_223])).
cnf(c_0_232,plain,(lazy_impl(true,esk2_2(true,true))=false|lazy_impl(true,esk2_2(true,true))=true|lazy_impl(true,true)=false|~forallprefers(lazy_impl(true,impl(true,true)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_225]), c_0_35])).
cnf(c_0_233,plain,(lazy_impl(true,X1)=false|true=X1|~bool(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_100]), c_0_101])).
cnf(c_0_234,plain,(~forallprefers(impl(true,X1),false)|~bool(X1)), inference(spm,[status(thm)],[c_0_226, c_0_227])).
cnf(c_0_235,negated_conjecture,(prop(esk2_2(true,esk9_0))=false|esk8_0=false|bool(esk9_0)|lazy_impl(true,lazy_impl(true,esk9_0))!=esk9_0), inference(spm,[status(thm)],[c_0_228, c_0_113])).
cnf(c_0_236,plain,(bool(esk2_2(true,X1))|~forallprefers(lazy_impl(prop(X2),impl(impl(X1,X2),X2)),true)|~bool(impl(X1,X2))), inference(spm,[status(thm)],[c_0_229, c_0_223])).
cnf(c_0_237,plain,(prop(X1)=false|~forallprefers(lazy_impl(true,impl(true,X1)),false)), inference(spm,[status(thm)],[c_0_207, c_0_113])).
cnf(c_0_238,plain,(prop(X1)=false|d(impl(true,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_105]), c_0_107])).
cnf(c_0_239,plain,(lazy_impl(true,true)=true), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_231]), c_0_102]), c_0_101])).
cnf(c_0_240,plain,(lazy_impl(true,esk2_2(true,true))=true|lazy_impl(true,esk2_2(true,true))=false|impl(true,true)=true|lazy_impl(true,true)=false|~bool(impl(true,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233]), c_0_51])])).
cnf(c_0_241,plain,(bool(impl(true,X1))|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_90]), c_0_52]), c_0_53])])).
cnf(c_0_242,negated_conjecture,(prop(esk2_2(true,esk9_0))=false|esk8_0=false|bool(esk9_0)|lazy_impl(true,esk9_0)!=esk9_0|~d(lazy_impl(true,esk9_0))), inference(spm,[status(thm)],[c_0_235, c_0_44])).
cnf(c_0_243,plain,(bool(esk2_2(true,true))|~forallprefers(lazy_impl(prop(X1),impl(X1,X1)),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_236, c_0_170])).
cnf(c_0_244,plain,(prop(X1)=false|~forallprefers(impl(true,X1),false)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_44]), c_0_238])).
cnf(c_0_245,plain,(~forallprefers(lazy_impl(true,impl(true,true)),lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_224, c_0_239]), c_0_216])).
cnf(c_0_246,plain,(lazy_impl(true,esk2_2(true,true))=false|lazy_impl(true,esk2_2(true,true))=true|lazy_impl(true,true)=false|impl(true,true)=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241]), c_0_136])])).
cnf(c_0_247,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(true,impl(true,false)))), inference(rw,[status(thm)],[c_0_182, c_0_201])).
cnf(c_0_248,plain,(bool(X1)|~forallprefers(lazy_impl(prop(X2),impl(lazy_impl(X3,X1),X2)),lazy_impl(prop(esk2_2(X3,X1)),impl(lazy_impl(X3,impl(X1,esk2_2(X3,X1))),esk2_2(X3,X1))))|~d(X1)), inference(spm,[status(thm)],[c_0_183, c_0_44])).
cnf(c_0_249,negated_conjecture,(prop(esk2_2(true,esk9_0))=false|esk8_0=false|bool(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_44]), c_0_209])).
cnf(c_0_250,plain,(bool(esk2_2(true,true))|~forallprefers(lazy_impl(true,impl(X1,X1)),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_243, c_0_37])).
cnf(c_0_251,plain,(prop(X1)=false|bool(impl(true,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_90]), c_0_52]), c_0_53])])).
cnf(c_0_252,plain,(~forallprefers(impl(true,true),lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))|~d(impl(true,true))), inference(spm,[status(thm)],[c_0_245, c_0_44])).
cnf(c_0_253,plain,(lazy_impl(true,esk2_2(true,true))=false|impl(true,true)=true|lazy_impl(true,true)=false|esk2_2(true,true)=true), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_246]), c_0_101])).
cnf(c_0_254,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),impl(true,false))|~d(impl(true,false))), inference(spm,[status(thm)],[c_0_247, c_0_44])).
cnf(c_0_255,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(prop(X1),impl(lazy_impl(true,esk9_0),X1)),true)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_249]), c_0_35]), c_0_209])).
cnf(c_0_256,plain,(bool(esk2_2(true,true))|~forallprefers(impl(X1,X1),true)|~d(impl(X1,X1))|~bool(X1)), inference(spm,[status(thm)],[c_0_250, c_0_44])).
cnf(c_0_257,plain,(impl(true,X1)=false|impl(true,X1)=true|prop(X1)=false), inference(spm,[status(thm)],[c_0_57, c_0_251])).
cnf(c_0_258,plain,(~forallprefers(impl(true,true),lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_227]), c_0_136])])).
cnf(c_0_259,plain,(esk2_2(true,true)=true|lazy_impl(true,true)=false|impl(true,true)=true|esk2_2(true,true)=false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_253]), c_0_102])).
cnf(c_0_260,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),impl(true,false))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_227]), c_0_53])])).
cnf(c_0_261,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(prop(X1),impl(esk9_0,X1)),true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_44]), c_0_209])).
cnf(c_0_262,plain,(lazy_impl(true,esk2_2(true,true))=esk2_2(true,true)|lazy_impl(true,true)=false|~forallprefers(lazy_impl(prop(true),impl(true,true)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_215]), c_0_35])).
cnf(c_0_263,plain,(lazy_impl(true,true)=false|bool(esk2_2(true,true))|~forallprefers(lazy_impl(true,impl(true,true)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_34]), c_0_35])).
cnf(c_0_264,plain,(bool(esk2_2(true,true))|~forallprefers(impl(true,true),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_227]), c_0_136])])).
cnf(c_0_265,plain,(impl(true,esk2_2(true,true))=true|impl(true,esk2_2(true,true))=false|lazy_impl(true,true)=false|~forallprefers(lazy_impl(true,impl(true,true)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_257]), c_0_35])).
cnf(c_0_266,plain,(~forallprefers(impl(true,true),lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,esk2_2(true,true)),esk2_2(true,true))))|~bool(esk2_2(true,true))), inference(spm,[status(thm)],[c_0_258, c_0_170])).
cnf(c_0_267,plain,(esk2_2(true,true)=false|impl(true,true)=true|esk2_2(true,true)=true), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_259, c_0_239]), c_0_216])).
cnf(c_0_268,plain,(~forallprefers(lazy_impl(prop(X1),X1),impl(true,false))|~bool(X1)), inference(spm,[status(thm)],[c_0_260, c_0_170])).
cnf(c_0_269,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(prop(X1),lazy_impl(true,esk9_0)),true)), inference(spm,[status(thm)],[c_0_261, c_0_93])).
cnf(c_0_270,plain,(lazy_impl(true,esk2_2(true,true))=esk2_2(true,true)|lazy_impl(true,true)=false|~forallprefers(lazy_impl(true,impl(true,true)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_37]), c_0_136])])).
cnf(c_0_271,plain,(bool(esk2_2(true,true))|~forallprefers(lazy_impl(true,impl(true,true)),true)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_263, c_0_239]), c_0_216])).
cnf(c_0_272,plain,(bool(impl(true,true))|bool(esk2_2(true,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_90]), c_0_60]), c_0_136])])).
cnf(c_0_273,plain,(impl(true,esk2_2(true,true))=false|impl(true,esk2_2(true,true))=true|impl(true,true)=true|lazy_impl(true,true)=false|~bool(impl(true,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_233]), c_0_51])])).
cnf(c_0_274,plain,(impl(true,true)=true|esk2_2(true,true)=false|~forallprefers(impl(true,true),lazy_impl(prop(true),impl(true,true)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267]), c_0_239]), c_0_136])])).
cnf(c_0_275,plain,(d(lazy_impl(prop(X1),impl(true,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_82]), c_0_52])])).
cnf(c_0_276,plain,(~forallprefers(false,impl(true,false))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_176]), c_0_53])])).
cnf(c_0_277,plain,(impl(true,X1)=false|impl(true,X1)=true|~bool(X1)), inference(spm,[status(thm)],[c_0_57, c_0_241])).
cnf(c_0_278,plain,(lazy_impl(true,esk7_0)=false2|~d(esk7_0)), inference(spm,[status(thm)],[c_0_158, c_0_44])).
cnf(c_0_279,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(prop(X1),esk9_0),true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_44]), c_0_209])).
cnf(c_0_280,plain,(lazy_impl(true,esk2_2(true,true))=esk2_2(true,true)|impl(true,true)=true|lazy_impl(true,true)=false|~bool(impl(true,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_233]), c_0_51])])).
cnf(c_0_281,plain,(~forallprefers(lazy_impl(true,impl(true,true)),lazy_impl(true,impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))|~bool(esk2_2(true,true))), inference(spm,[status(thm)],[c_0_245, c_0_37])).
cnf(c_0_282,plain,(impl(true,true)=true|bool(esk2_2(true,true))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_233]), c_0_51])]), c_0_272])).
cnf(c_0_283,plain,(impl(true,esk2_2(true,true))=true|impl(true,esk2_2(true,true))=false|lazy_impl(true,true)=false|impl(true,true)=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_241]), c_0_136])])).
cnf(c_0_284,plain,(~forallprefers(impl(true,true),lazy_impl(true,impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))|~bool(esk2_2(true,true))), inference(spm,[status(thm)],[c_0_258, c_0_37])).
cnf(c_0_285,plain,(esk2_2(true,true)=false|bool(impl(true,true))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_90]), c_0_275]), c_0_221])]), c_0_132])).
cnf(c_0_286,plain,(impl(true,false)=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_277]), c_0_51]), c_0_53])])).
cnf(c_0_287,plain,(lazy_impl(true,false)=false), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_278, c_0_80]), c_0_80]), c_0_52])]), c_0_166])).
fof(c_0_288, plain, (![X5]:(~bool(X5)|impl(false,X5)=true)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).
cnf(c_0_289,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(true,esk9_0),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_279, c_0_37])).
cnf(c_0_290,plain,(~forallprefers(lazy_impl(true,impl(true,true)),lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,esk2_2(true,true)),esk2_2(true,true))))|~bool(esk2_2(true,true))), inference(spm,[status(thm)],[c_0_245, c_0_170])).
cnf(c_0_291,plain,(lazy_impl(true,esk2_2(true,true))=esk2_2(true,true)|lazy_impl(true,true)=false|impl(true,true)=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_241]), c_0_136])])).
cnf(c_0_292,plain,(impl(true,true)=true|~forallprefers(lazy_impl(true,impl(true,true)),lazy_impl(true,impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))), inference(spm,[status(thm)],[c_0_281, c_0_282])).
cnf(c_0_293,plain,(impl(true,esk2_2(true,true))=false|impl(true,esk2_2(true,true))=true|impl(true,true)=true), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_283, c_0_239]), c_0_216])).
cnf(c_0_294,plain,(bool(impl(true,true))|~forallprefers(impl(true,true),lazy_impl(true,impl(false,false)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_285]), c_0_286]), c_0_287]), c_0_53])])).
cnf(c_0_295,plain,(impl(false,X1)=true|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_288])).
cnf(c_0_296,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(esk9_0,true)|~bool(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_44]), c_0_209])).
cnf(c_0_297,plain,(impl(true,true)=true|~forallprefers(lazy_impl(true,impl(true,true)),lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,esk2_2(true,true)),esk2_2(true,true))))), inference(spm,[status(thm)],[c_0_290, c_0_282])).
cnf(c_0_298,plain,(lazy_impl(true,esk2_2(true,true))=esk2_2(true,true)|impl(true,true)=true), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_291, c_0_239]), c_0_216])).
cnf(c_0_299,plain,(impl(true,esk2_2(true,true))=false|impl(true,true)=true|~forallprefers(lazy_impl(true,impl(true,true)),lazy_impl(true,impl(true,esk2_2(true,true))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293]), c_0_239])).
cnf(c_0_300,plain,(bool(impl(true,true))|~forallprefers(impl(true,true),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_295]), c_0_239]), c_0_53])])).
cnf(c_0_301,negated_conjecture,(lazy_impl(prop(esk2_2(esk8_0,esk9_0)),impl(lazy_impl(esk8_0,impl(esk9_0,esk2_2(esk8_0,esk9_0))),esk2_2(esk8_0,esk9_0)))!=lazy_and1(esk8_0,esk9_0)|~d(lazy_impl(prop(esk2_2(esk8_0,esk9_0)),impl(lazy_impl(esk8_0,impl(esk9_0,esk2_2(esk8_0,esk9_0))),esk2_2(esk8_0,esk9_0))))), inference(spm,[status(thm)],[c_0_92, c_0_44])).
cnf(c_0_302,negated_conjecture,(esk8_0=false|bool(esk9_0)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_90]), c_0_60]), c_0_136])])).
cnf(c_0_303,plain,(impl(true,true)=true|~forallprefers(lazy_impl(true,impl(true,true)),lazy_impl(prop(esk2_2(true,true)),impl(esk2_2(true,true),esk2_2(true,true))))), inference(spm,[status(thm)],[c_0_297, c_0_298])).
cnf(c_0_304,plain,(impl(true,esk2_2(true,true))=false|impl(true,true)=true|~forallprefers(lazy_impl(true,impl(true,true)),true)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_293]), c_0_239]), c_0_216])).
cnf(c_0_305,plain,(bool(impl(true,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_90]), c_0_60]), c_0_136])])).
cnf(c_0_306,negated_conjecture,(bool(esk2_2(esk8_0,esk9_0))|lazy_and1(esk8_0,esk9_0)!=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_34]), c_0_35]), c_0_35]), c_0_60])])).
cnf(c_0_307,negated_conjecture,(esk8_0=false|bool(esk9_0)), inference(spm,[status(thm)],[c_0_302, c_0_143])).
cnf(c_0_308,plain,(impl(true,true)=true|~forallprefers(lazy_impl(true,impl(true,true)),lazy_impl(true,impl(esk2_2(true,true),esk2_2(true,true))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_37]), c_0_282])).
cnf(c_0_309,plain,(impl(true,esk2_2(true,true))=false|impl(true,true)=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_233]), c_0_51]), c_0_305])])).
cnf(c_0_310,negated_conjecture,(esk8_0=false|bool(esk2_2(true,esk9_0))|lazy_impl(true,esk9_0)!=true), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_178]), c_0_179])).
cnf(c_0_311,negated_conjecture,(esk8_0=false|esk9_0=false|esk9_0=true), inference(spm,[status(thm)],[c_0_57, c_0_307])).
cnf(c_0_312,plain,(impl(true,true)=true|~forallprefers(false,lazy_impl(true,impl(esk2_2(true,true),esk2_2(true,true))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_233]), c_0_305])])).
cnf(c_0_313,plain,(impl(true,true)=true|esk2_2(true,true)=false), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_309]), c_0_282])).
cnf(c_0_314,negated_conjecture,(lazy_impl(true,lazy_impl(true,impl(lazy_impl(esk8_0,impl(esk9_0,esk2_2(esk8_0,esk9_0))),esk2_2(esk8_0,esk9_0))))!=lazy_and1(esk8_0,esk9_0)|~bool(esk2_2(esk8_0,esk9_0))), inference(spm,[status(thm)],[c_0_92, c_0_37])).
cnf(c_0_315,negated_conjecture,(lazy_impl(true,esk9_0)=false|esk8_0=false|bool(esk2_2(true,esk9_0))|~bool(esk9_0)), inference(spm,[status(thm)],[c_0_310, c_0_100])).
cnf(c_0_316,negated_conjecture,(esk9_0=false|esk8_0=false|bool(esk2_2(true,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311]), c_0_239])])).
cnf(c_0_317,plain,(impl(true,true)=true|~forallprefers(false,lazy_impl(true,impl(false,false)))), inference(spm,[status(thm)],[c_0_312, c_0_313])).
cnf(c_0_318,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(true,impl(lazy_impl(true,impl(esk9_0,esk2_2(true,esk9_0))),esk2_2(true,esk9_0))))!=lazy_impl(true,esk9_0)|~bool(esk2_2(true,esk9_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_178]), c_0_179])).
cnf(c_0_319,negated_conjecture,(lazy_impl(true,esk9_0)=false|esk8_0=false|bool(esk2_2(true,esk9_0))), inference(spm,[status(thm)],[c_0_315, c_0_307])).
cnf(c_0_320,negated_conjecture,(esk9_0=false|esk8_0=false|lazy_impl(true,lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))!=true), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_311]), c_0_239])).
cnf(c_0_321,negated_conjecture,(esk2_2(true,true)=false|esk2_2(true,true)=true|esk8_0=false|esk9_0=false), inference(spm,[status(thm)],[c_0_57, c_0_316])).
cnf(c_0_322,plain,(impl(true,true)=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_295]), c_0_239]), c_0_51]), c_0_53])])).
cnf(c_0_323,negated_conjecture,(lazy_impl(true,esk9_0)=false|esk8_0=false|lazy_impl(true,lazy_impl(true,impl(lazy_impl(true,impl(esk9_0,esk2_2(true,esk9_0))),esk2_2(true,esk9_0))))!=lazy_impl(true,esk9_0)), inference(spm,[status(thm)],[c_0_318, c_0_319])).
cnf(c_0_324,negated_conjecture,(esk2_2(true,true)=false|esk8_0=false|esk9_0=false|lazy_impl(true,lazy_impl(prop(true),true))!=true), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_321]), c_0_322]), c_0_239]), c_0_322])).
cnf(c_0_325,negated_conjecture,(esk9_0=false|esk8_0=false|lazy_impl(true,lazy_impl(true,impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))!=true), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_311]), c_0_239]), c_0_239]), c_0_216])).
cnf(c_0_326,negated_conjecture,(esk2_2(true,true)=false|esk9_0=false|esk8_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_37]), c_0_239]), c_0_239]), c_0_136])])).
cnf(c_0_327,plain,(bool(esk2_2(true,X1))|~forallprefers(lazy_impl(prop(X2),impl(impl(X1,X2),X2)),true)|~d(impl(X1,X2))), inference(spm,[status(thm)],[c_0_229, c_0_44])).
cnf(c_0_328,negated_conjecture,(esk8_0=false|esk9_0=false|lazy_impl(true,lazy_impl(true,impl(false,false)))!=true), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_286]), c_0_287])).
cnf(c_0_329,plain,(bool(esk2_2(true,false))|~forallprefers(lazy_impl(prop(X1),impl(true,X1)),true)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_295]), c_0_60])])).
cnf(c_0_330,negated_conjecture,(esk9_0=false|esk8_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_295]), c_0_239]), c_0_239]), c_0_53])])).
cnf(c_0_331,plain,(bool(esk2_2(true,false))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_286]), c_0_176]), c_0_51]), c_0_53])])).
cnf(c_0_332,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(true,impl(lazy_impl(true,impl(false,esk2_2(true,false))),esk2_2(true,false))))!=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_330]), c_0_287]), c_0_331])])).
cnf(c_0_333,plain,(~forallprefers(lazy_impl(prop(X1),impl(lazy_impl(X2,true),X1)),lazy_impl(prop(esk2_2(X2,false)),impl(lazy_impl(X2,impl(false,esk2_2(X2,false))),esk2_2(X2,false))))|~bool(X1)), inference(spm,[status(thm)],[c_0_116, c_0_295])).
cnf(c_0_334,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(true,impl(true,esk2_2(true,false))))!=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_295]), c_0_239]), c_0_331])])).
cnf(c_0_335,plain,(impl(true,X1)=false|true=X1|~bool(X1)), inference(spm,[status(thm)],[c_0_170, c_0_277])).
cnf(c_0_336,plain,(~forallprefers(lazy_impl(prop(X1),impl(lazy_impl(X2,true),X1)),lazy_impl(prop(esk2_2(X2,false)),impl(lazy_impl(X2,true),esk2_2(X2,false))))|~bool(esk2_2(X2,false))|~bool(X1)), inference(spm,[status(thm)],[c_0_333, c_0_295])).
cnf(c_0_337,negated_conjecture,(esk2_2(true,false)=true|esk8_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_335]), c_0_287]), c_0_287]), c_0_331])])).
cnf(c_0_338,negated_conjecture,(esk8_0=false|~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(prop(true),true))|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_337]), c_0_239]), c_0_239]), c_0_322]), c_0_136])])).
cnf(c_0_339,negated_conjecture,(esk8_0=false|~forallprefers(false,lazy_impl(prop(true),true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_286]), c_0_176]), c_0_53])])).
fof(c_0_340, plain, (![X5]:lazy_and1(false,X5)=false), inference(variable_rename,[status(thm)],[c_0_20])).
cnf(c_0_341,negated_conjecture,(lazy_impl(prop(esk2_2(esk8_0,esk9_0)),impl(lazy_impl(esk8_0,impl(esk9_0,esk2_2(esk8_0,esk9_0))),esk2_2(esk8_0,esk9_0)))=true|lazy_and1(esk8_0,esk9_0)!=false|~bool(lazy_impl(prop(esk2_2(esk8_0,esk9_0)),impl(lazy_impl(esk8_0,impl(esk9_0,esk2_2(esk8_0,esk9_0))),esk2_2(esk8_0,esk9_0))))), inference(spm,[status(thm)],[c_0_92, c_0_233])).
cnf(c_0_342,negated_conjecture,(esk8_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_37]), c_0_239]), c_0_51]), c_0_136])])).
cnf(c_0_343,plain,(lazy_and1(false,X1)=false), inference(split_conjunct,[status(thm)],[c_0_340])).
cnf(c_0_344,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_341, c_0_342]), c_0_201]), c_0_342]), c_0_342]), c_0_201]), c_0_35]), c_0_342]), c_0_201]), c_0_286]), c_0_176]), c_0_342]), c_0_343]), c_0_342]), c_0_201]), c_0_342]), c_0_342]), c_0_201]), c_0_35]), c_0_342]), c_0_201]), c_0_286]), c_0_176]), c_0_53])]), c_0_216]), ['proof']).
# SZS output end CNFRefutation.
EOF
