% Problem    : PUZ133+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n098.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:48:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C41_F1_PI_AE_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (((queens_p&![X1]:q(X1)=p(perm(X1)))=>queens_q)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', queens_sym)).
fof(c_0_1, axiom, ((![X1]:![X2]:(((((le(s(n0),X1)&le(X1,n))&le(s(X1),X2))&le(X2,n))&(le(s(X1),X2)<=>le(s(perm(X2)),perm(X1))))=>((q(X1)!=q(X2)&plus(q(X1),X1)!=plus(q(X2),X2))&minus(q(X1),X1)!=minus(q(X2),X2)))=>queens_q)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', queens_q)).
fof(c_0_2, axiom, (![X3]:![X4]:![X5]:((le(X3,X4)&le(X4,X5))=>le(X3,X5))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', le_trans)).
fof(c_0_3, axiom, (![X2]:![X1]:minus(X1,X2)=minus(perm(X2),perm(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', permutation_another_one)).
fof(c_0_4, axiom, (![X1]:perm(X1)=minus(s(n),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', permutation)).
fof(c_0_5, axiom, (![X3]:le(X3,s(X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', succ_le)).
fof(c_0_6, axiom, (![X1]:![X2]:![X6]:![X7]:(minus(X1,X2)=minus(X6,X7)<=>minus(X1,X6)=minus(X2,X7))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', minus1)).
fof(c_0_7, axiom, (![X1]:![X2]:![X6]:![X7]:(plus(X1,X2)=plus(X6,X7)<=>minus(X1,X6)=minus(X7,X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', plus1)).
fof(c_0_8, axiom, ((queens_p=>![X1]:![X2]:((((le(s(n0),X1)&le(X1,n))&le(s(X1),X2))&le(X2,n))=>((p(X1)!=p(X2)&plus(p(X1),X1)!=plus(p(X2),X2))&minus(p(X1),X1)!=minus(p(X2),X2))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', queens_p)).
fof(c_0_9, axiom, (![X1]:((le(s(n0),X1)&le(X1,n))=>(le(s(n0),perm(X1))&le(perm(X1),n)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', permutation_range)).
fof(c_0_10, negated_conjecture, (~(((queens_p&![X1]:q(X1)=p(perm(X1)))=>queens_q))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_11, plain, (((((((le(s(n0),esk1_0)|queens_q)&(le(esk1_0,n)|queens_q))&(le(s(esk1_0),esk2_0)|queens_q))&(le(esk2_0,n)|queens_q))&(((~le(s(esk1_0),esk2_0)|le(s(perm(esk2_0)),perm(esk1_0)))|queens_q)&((~le(s(perm(esk2_0)),perm(esk1_0))|le(s(esk1_0),esk2_0))|queens_q)))&(((q(esk1_0)=q(esk2_0)|plus(q(esk1_0),esk1_0)=plus(q(esk2_0),esk2_0))|minus(q(esk1_0),esk1_0)=minus(q(esk2_0),esk2_0))|queens_q))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_12, negated_conjecture, (![X2]:((queens_p&q(X2)=p(perm(X2)))&~queens_q)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
fof(c_0_13, plain, (![X6]:![X7]:![X8]:((~le(X6,X7)|~le(X7,X8))|le(X6,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
cnf(c_0_14,plain,(queens_q|le(s(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,negated_conjecture,(~queens_q), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, (![X3]:![X4]:minus(X4,X3)=minus(perm(X3),perm(X4))), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_17, plain, (![X2]:perm(X2)=minus(s(n),X2)), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_18,plain,(le(X1,X2)|~le(X3,X2)|~le(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(le(s(esk1_0),esk2_0)), inference(sr,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_20, plain, (![X4]:le(X4,s(X4))), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_21, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:((minus(X8,X9)!=minus(X10,X11)|minus(X8,X10)=minus(X9,X11))&(minus(X12,X14)!=minus(X13,X15)|minus(X12,X13)=minus(X14,X15)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_22,plain,(minus(X1,X2)=minus(perm(X2),perm(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(perm(X1)=minus(s(n),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,plain,(queens_q|minus(q(esk1_0),esk1_0)=minus(q(esk2_0),esk2_0)|plus(q(esk1_0),esk1_0)=plus(q(esk2_0),esk2_0)|q(esk1_0)=q(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25,negated_conjecture,(q(X1)=p(perm(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_26, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:((plus(X8,X9)!=plus(X10,X11)|minus(X8,X10)=minus(X11,X9))&(minus(X12,X14)!=minus(X15,X13)|plus(X12,X13)=plus(X14,X15)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
cnf(c_0_27,plain,(le(X1,esk2_0)|~le(X1,s(esk1_0))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_28,plain,(le(X1,s(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29,plain,(minus(X1,X2)=minus(X3,X4)|minus(X1,X3)!=minus(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30,plain,(minus(minus(s(n),X2),minus(s(n),X1))=minus(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_23])).
cnf(c_0_31,plain,(p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))|plus(p(minus(s(n),esk2_0)),esk2_0)=plus(p(minus(s(n),esk1_0)),esk1_0)|minus(p(minus(s(n),esk2_0)),esk2_0)=minus(p(minus(s(n),esk1_0)),esk1_0)|queens_q), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_23]), c_0_23]), c_0_23]), c_0_23]), c_0_23]), c_0_23])).
cnf(c_0_32,plain,(plus(X1,X2)=plus(X3,X4)|minus(X1,X3)!=minus(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_33, plain, (![X3]:![X4]:((((p(X3)!=p(X4)|(((~le(s(n0),X3)|~le(X3,n))|~le(s(X3),X4))|~le(X4,n)))|~queens_p)&((plus(p(X3),X3)!=plus(p(X4),X4)|(((~le(s(n0),X3)|~le(X3,n))|~le(s(X3),X4))|~le(X4,n)))|~queens_p))&((minus(p(X3),X3)!=minus(p(X4),X4)|(((~le(s(n0),X3)|~le(X3,n))|~le(s(X3),X4))|~le(X4,n)))|~queens_p))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
fof(c_0_34, plain, (![X2]:((le(s(n0),perm(X2))|(~le(s(n0),X2)|~le(X2,n)))&(le(perm(X2),n)|(~le(s(n0),X2)|~le(X2,n))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
cnf(c_0_35,plain,(le(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36,plain,(queens_q|le(s(n0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_37,plain,(minus(minus(s(n),X1),X2)=minus(minus(s(n),X3),X4)|minus(X3,X1)!=minus(X2,X4)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38,plain,(minus(p(minus(s(n),esk2_0)),esk2_0)=minus(p(minus(s(n),esk1_0)),esk1_0)|plus(p(minus(s(n),esk2_0)),esk2_0)=plus(p(minus(s(n),esk1_0)),esk1_0)|p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))), inference(sr,[status(thm)],[c_0_31, c_0_15])).
cnf(c_0_39,plain,(plus(X1,X2)=plus(X2,X1)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_40,plain,(~queens_p|~le(X1,n)|~le(s(X2),X1)|~le(X2,n)|~le(s(n0),X2)|minus(p(X2),X2)!=minus(p(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41,negated_conjecture,(queens_p), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42,plain,(minus(X1,minus(s(n),X2))=minus(X3,minus(s(n),X4))|minus(X1,X3)!=minus(X4,X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_43,plain,(le(perm(X1),n)|~le(X1,n)|~le(s(n0),X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44,plain,(queens_q|le(esk2_0,n)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_45,plain,(le(X1,esk2_0)|~le(X1,esk1_0)), inference(spm,[status(thm)],[c_0_18, c_0_35])).
cnf(c_0_46,plain,(le(s(n0),esk1_0)), inference(sr,[status(thm)],[c_0_36, c_0_15])).
cnf(c_0_47,plain,(le(s(n0),perm(X1))|~le(X1,n)|~le(s(n0),X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48,plain,(minus(minus(s(n),X1),X2)=minus(minus(s(n),X2),X1)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_49,plain,(minus(X1,X2)=minus(X3,X4)|plus(X1,X4)!=plus(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_50,plain,(minus(p(minus(s(n),esk2_0)),esk2_0)=minus(p(minus(s(n),esk1_0)),esk1_0)|plus(esk2_0,p(minus(s(n),esk2_0)))=plus(esk1_0,p(minus(s(n),esk1_0)))|p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39]), c_0_39])).
cnf(c_0_51,plain,(minus(p(X1),X1)!=minus(p(X2),X2)|~le(s(n0),X2)|~le(s(X2),X1)|~le(X2,n)|~le(X1,n)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41])])).
cnf(c_0_52,plain,(minus(X1,minus(s(n),X2))=minus(X2,minus(s(n),X1))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_53,plain,(le(minus(s(n),X1),n)|~le(X1,n)|~le(s(n0),X1)), inference(rw,[status(thm)],[c_0_43, c_0_23])).
cnf(c_0_54,plain,(le(esk2_0,n)), inference(sr,[status(thm)],[c_0_44, c_0_15])).
cnf(c_0_55,plain,(le(s(n0),esk2_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_56,plain,(le(s(n0),minus(s(n),X1))|~le(X1,n)|~le(s(n0),X1)), inference(rw,[status(thm)],[c_0_47, c_0_23])).
cnf(c_0_57,plain,(queens_q|le(esk1_0,n)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_58,plain,(queens_q|le(s(perm(esk2_0)),perm(esk1_0))|~le(s(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_59,plain,(minus(minus(s(n),minus(s(n),X1)),X2)=minus(X1,X2)), inference(spm,[status(thm)],[c_0_30, c_0_48])).
cnf(c_0_60,plain,(minus(p(minus(s(n),esk2_0)),esk2_0)=minus(p(minus(s(n),esk1_0)),esk1_0)|p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))|minus(esk2_0,X1)=minus(X2,p(minus(s(n),esk2_0)))|plus(esk1_0,p(minus(s(n),esk1_0)))!=plus(X1,X2)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_61,plain,(minus(X1,minus(s(n),p(minus(s(n),X1))))!=minus(p(X2),X2)|~le(s(X2),minus(s(n),X1))|~le(minus(s(n),X1),n)|~le(s(n0),X2)|~le(X2,n)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_62,plain,(le(minus(s(n),esk2_0),n)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])])).
cnf(c_0_63,plain,(le(s(n0),minus(s(n),esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_54]), c_0_55])])).
cnf(c_0_64,plain,(le(esk1_0,n)), inference(sr,[status(thm)],[c_0_57, c_0_15])).
cnf(c_0_65,plain,(queens_q|le(s(minus(s(n),esk2_0)),minus(s(n),esk1_0))|~le(s(esk1_0),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_23]), c_0_23])).
cnf(c_0_66,plain,(minus(X1,minus(s(n),X2))=minus(X3,minus(s(n),X4))|minus(X2,X3)!=minus(X4,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_59]), c_0_30])).
cnf(c_0_67,plain,(minus(p(minus(s(n),esk1_0)),p(minus(s(n),esk2_0)))=minus(esk2_0,esk1_0)|minus(p(minus(s(n),esk2_0)),esk2_0)=minus(p(minus(s(n),esk1_0)),esk1_0)|p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_68,plain,(minus(X1,minus(s(n),p(minus(s(n),X1))))!=minus(esk2_0,minus(s(n),p(minus(s(n),esk2_0))))|~le(s(minus(s(n),esk2_0)),minus(s(n),X1))|~le(minus(s(n),X1),n)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])]), c_0_52])).
cnf(c_0_69,plain,(le(minus(s(n),esk1_0),n)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_64]), c_0_46])])).
cnf(c_0_70,plain,(le(s(minus(s(n),esk2_0)),minus(s(n),esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_19])]), c_0_15])).
cnf(c_0_71,plain,(minus(p(minus(s(n),esk2_0)),minus(s(n),X1))=minus(X2,minus(s(n),p(minus(s(n),esk1_0))))|minus(p(minus(s(n),esk2_0)),esk2_0)=minus(p(minus(s(n),esk1_0)),esk1_0)|p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))|minus(X1,X2)!=minus(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_72,plain,(minus(esk2_0,minus(s(n),p(minus(s(n),esk2_0))))!=minus(esk1_0,minus(s(n),p(minus(s(n),esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])])).
cnf(c_0_73,plain,(minus(p(minus(s(n),esk2_0)),esk2_0)=minus(p(minus(s(n),esk1_0)),esk1_0)|p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_71]), c_0_52]), c_0_72])).
cnf(c_0_74,plain,(p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))|minus(p(minus(s(n),esk2_0)),X1)=minus(esk2_0,X2)|minus(p(minus(s(n),esk1_0)),esk1_0)!=minus(X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_73])).
cnf(c_0_75,plain,(~queens_p|~le(X1,n)|~le(s(X2),X1)|~le(X2,n)|~le(s(n0),X2)|plus(p(X2),X2)!=plus(p(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_76,plain,(plus(X1,minus(s(n),X2))=plus(X3,minus(s(n),X4))|minus(X1,X3)!=minus(X2,X4)), inference(spm,[status(thm)],[c_0_32, c_0_30])).
cnf(c_0_77,plain,(minus(p(minus(s(n),esk2_0)),p(minus(s(n),esk1_0)))=minus(esk2_0,esk1_0)|p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))), inference(er,[status(thm)],[c_0_74])).
cnf(c_0_78,plain,(plus(p(X1),X1)!=plus(p(X2),X2)|~le(s(n0),X2)|~le(s(X2),X1)|~le(X2,n)|~le(X1,n)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_41])])).
cnf(c_0_79,plain,(plus(p(minus(s(n),esk2_0)),minus(s(n),X1))=plus(p(minus(s(n),esk1_0)),minus(s(n),X2))|p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))|minus(esk2_0,esk1_0)!=minus(X1,X2)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_80,plain,(plus(X1,p(X1))!=plus(X2,p(X2))|~le(s(n0),X2)|~le(s(X2),X1)|~le(X2,n)|~le(X1,n)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_39]), c_0_39])).
cnf(c_0_81,plain,(plus(minus(s(n),esk2_0),p(minus(s(n),esk2_0)))=plus(minus(s(n),esk1_0),p(minus(s(n),esk1_0)))|p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_79]), c_0_39]), c_0_39])).
cnf(c_0_82,plain,(~queens_p|~le(X1,n)|~le(s(X2),X1)|~le(X2,n)|~le(s(n0),X2)|p(X2)!=p(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_83,plain,(p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))|plus(X1,p(X1))!=plus(minus(s(n),esk1_0),p(minus(s(n),esk1_0)))|~le(s(minus(s(n),esk2_0)),X1)|~le(X1,n)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_63]), c_0_62])])).
cnf(c_0_84,plain,(p(X1)!=p(X2)|~le(s(n0),X2)|~le(s(X2),X1)|~le(X2,n)|~le(X1,n)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_41])])).
cnf(c_0_85,plain,(p(minus(s(n),esk2_0))=p(minus(s(n),esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_83]), c_0_70]), c_0_69])])).
cnf(c_0_86,plain,(p(X1)!=p(minus(s(n),esk1_0))|~le(s(minus(s(n),esk2_0)),X1)|~le(X1,n)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_63]), c_0_62])])).
cnf(c_0_87,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_86]), c_0_70]), c_0_69])]), ['proof']).
# SZS output end CNFRefutation.
EOF
