% Problem    : LCL899+1 : TPTP v6.1.0. Released v5.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n094.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:00:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C45_F1_AE_R8_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:'+'(X1,'0')=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_03)).
fof(c_0_1, axiom, (![X1]:![X2]:'+'(X1,X2)='+'(X2,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_02)).
fof(c_0_2, axiom, (![X7]:![X8]:(('>='(X7,X8)&'>='(X8,X7))=>X7=X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_06)).
fof(c_0_3, axiom, (![X1]:'>='(X1,'0')), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_08)).
fof(c_0_4, axiom, (![X9]:![X10]:![X11]:('>='('+'(X9,X10),X11)<=>'>='(X10,'==>'(X9,X11)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_07)).
fof(c_0_5, axiom, (![X1]:'+'(X1,'1')='1'), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_12)).
fof(c_0_6, axiom, (![X12]:![X13]:![X14]:('>='(X12,X13)=>'>='('+'(X12,X14),'+'(X13,X14)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_09)).
fof(c_0_7, axiom, (![X1]:![X2]:'==>'('==>'(X1,X2),X2)='==>'('==>'(X2,X1),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_13)).
fof(c_0_8, conjecture, (![X21]:'==>'('==>'(X21,'1'),'1')=X21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_14)).
fof(c_0_9, plain, (![X2]:'+'(X2,'0')=X2), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_10, plain, (![X3]:![X4]:'+'(X3,X4)='+'(X4,X3)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_11, plain, (![X9]:![X10]:((~'>='(X9,X10)|~'>='(X10,X9))|X9=X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
fof(c_0_12, plain, (![X2]:'>='(X2,'0')), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_13, plain, (![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:((~'>='('+'(X12,X13),X14)|'>='(X13,'==>'(X12,X14)))&(~'>='(X16,'==>'(X15,X17))|'>='('+'(X15,X16),X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
fof(c_0_14, plain, (![X2]:'+'(X2,'1')='1'), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_15, plain, (![X15]:![X16]:![X17]:(~'>='(X15,X16)|'>='('+'(X15,X17),'+'(X16,X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_16,plain,('+'(X1,'0')=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17,plain,('+'(X1,X2)='+'(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,plain,(X1=X2|~'>='(X2,X1)|~'>='(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19,plain,('>='(X1,'0')), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,('>='(X1,'==>'(X2,X3))|~'>='('+'(X2,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,plain,('+'(X1,'1')='1'), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22,plain,('>='('+'(X1,X2),'+'(X3,X2))|~'>='(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,plain,('+'('0',X1)=X1), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24,plain,('0'=X1|~'>='('0',X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25,plain,('>='(X1,'==>'(X2,'0'))), inference(spm,[status(thm)],[c_0_20, c_0_19])).
fof(c_0_26, plain, (![X3]:![X4]:'==>'('==>'(X3,X4),X4)='==>'('==>'(X4,X3),X3)), inference(variable_rename,[status(thm)],[c_0_7])).
fof(c_0_27, negated_conjecture, (~(![X21]:'==>'('==>'(X21,'1'),'1')=X21)), inference(assume_negation,[status(cth)],[c_0_8])).
cnf(c_0_28,plain,('+'('1',X1)='1'), inference(spm,[status(thm)],[c_0_21, c_0_17])).
cnf(c_0_29,plain,('>='('+'(X1,X2),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_19]), c_0_23])).
cnf(c_0_30,plain,('==>'(X1,'0')='0'), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31,plain,('==>'('==>'(X1,X2),X2)='==>'('==>'(X2,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_32, negated_conjecture, ('==>'('==>'(esk1_0,'1'),'1')!=esk1_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])).
cnf(c_0_33,plain,('>='(X1,'==>'('1',X2))|~'>='('1',X2)), inference(spm,[status(thm)],[c_0_20, c_0_28])).
cnf(c_0_34,plain,('>='('1',X1)), inference(spm,[status(thm)],[c_0_29, c_0_28])).
cnf(c_0_35,plain,('>='('+'(X1,X2),X3)|~'>='(X2,'==>'(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_36,plain,('==>'('==>'('0',X1),X1)='0'), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37,negated_conjecture,('==>'('==>'(esk1_0,'1'),'1')!=esk1_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38,plain,('>='(X1,'==>'('1',X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34])])).
cnf(c_0_39,plain,('>='('+'('==>'('0',X1),X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_19])])).
cnf(c_0_40,negated_conjecture,('==>'('==>'('1',esk1_0),esk1_0)!=esk1_0), inference(rw,[status(thm)],[c_0_37, c_0_31])).
cnf(c_0_41,plain,('==>'('1',X1)='0'), inference(spm,[status(thm)],[c_0_24, c_0_38])).
cnf(c_0_42,plain,('>='('==>'('0',X1),X1)), inference(spm,[status(thm)],[c_0_39, c_0_16])).
cnf(c_0_43,plain,('>='(X1,'==>'(X2,X1))), inference(spm,[status(thm)],[c_0_20, c_0_29])).
cnf(c_0_44,negated_conjecture,('==>'('0',esk1_0)!=esk1_0), inference(rw,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45,plain,('==>'('0',X1)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_42]), c_0_43])])).
cnf(c_0_46,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45])]), ['proof']).
# SZS output end CNFRefutation.
EOF
