% Problem    : SWB025+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n138.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:45:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C45_F1_AE_R8_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X2]:![X12]:![X33]:![X14]:![X34]:((((iext(uri_rdf_first,X12,X33)&iext(uri_rdf_rest,X12,X14))&iext(uri_rdf_first,X14,X34))&iext(uri_rdf_rest,X14,uri_rdf_nil))=>(iext(uri_owl_propertyChainAxiom,X2,X12)<=>(((ip(X2)&ip(X33))&ip(X34))&![X36]:![X27]:![X28]:((iext(X33,X36,X27)&iext(X34,X27,X28))=>iext(X2,X36,X28)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_chain_002)).
fof(c_0_1, axiom, (?[X42]:?[X43]:?[X44]:?[X45]:?[X46]:((((((((((((((iext(uri_owl_propertyChainAxiom,uri_ex_hasUncle,X42)&iext(uri_rdf_first,X42,uri_ex_hasCousin))&iext(uri_rdf_rest,X42,X43))&iext(uri_rdf_first,X43,uri_ex_hasFather))&iext(uri_rdf_rest,X43,uri_rdf_nil))&iext(uri_owl_propertyChainAxiom,uri_ex_hasCousin,X44))&iext(uri_rdf_first,X44,uri_ex_hasUncle))&iext(uri_rdf_rest,X44,X45))&iext(uri_rdf_first,X45,X46))&iext(uri_rdf_rest,X45,uri_rdf_nil))&iext(uri_owl_inverseOf,X46,uri_ex_hasFather))&iext(uri_ex_hasFather,uri_ex_alice,uri_ex_dave))&iext(uri_ex_hasCousin,uri_ex_alice,uri_ex_bob))&iext(uri_ex_hasFather,uri_ex_bob,uri_ex_charly))&iext(uri_ex_hasUncle,uri_ex_bob,uri_ex_dave))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_premise_fullish_025_Cyclic_Dependencies_between_Complex_Properties)).
fof(c_0_2, axiom, (![X33]:![X34]:(iext(uri_owl_inverseOf,X33,X34)<=>((ip(X33)&ip(X34))&![X4]:![X7]:(iext(X33,X4,X7)<=>iext(X34,X7,X4))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_inv)).
fof(c_0_3, conjecture, ((iext(uri_ex_hasUncle,uri_ex_alice,uri_ex_charly)&iext(uri_ex_hasCousin,uri_ex_bob,uri_ex_alice))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_conclusion_fullish_025_Cyclic_Dependencies_between_Complex_Properties)).
fof(c_0_4, plain, (![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:![X43]:![X44]:((((((ip(X37)|~iext(uri_owl_propertyChainAxiom,X37,X38))|(((~iext(uri_rdf_first,X38,X39)|~iext(uri_rdf_rest,X38,X40))|~iext(uri_rdf_first,X40,X41))|~iext(uri_rdf_rest,X40,uri_rdf_nil)))&((ip(X39)|~iext(uri_owl_propertyChainAxiom,X37,X38))|(((~iext(uri_rdf_first,X38,X39)|~iext(uri_rdf_rest,X38,X40))|~iext(uri_rdf_first,X40,X41))|~iext(uri_rdf_rest,X40,uri_rdf_nil))))&((ip(X41)|~iext(uri_owl_propertyChainAxiom,X37,X38))|(((~iext(uri_rdf_first,X38,X39)|~iext(uri_rdf_rest,X38,X40))|~iext(uri_rdf_first,X40,X41))|~iext(uri_rdf_rest,X40,uri_rdf_nil))))&((((~iext(X39,X42,X43)|~iext(X41,X43,X44))|iext(X37,X42,X44))|~iext(uri_owl_propertyChainAxiom,X37,X38))|(((~iext(uri_rdf_first,X38,X39)|~iext(uri_rdf_rest,X38,X40))|~iext(uri_rdf_first,X40,X41))|~iext(uri_rdf_rest,X40,uri_rdf_nil))))&(((((iext(X39,esk152_5(X37,X38,X39,X40,X41),esk153_5(X37,X38,X39,X40,X41))|((~ip(X37)|~ip(X39))|~ip(X41)))|iext(uri_owl_propertyChainAxiom,X37,X38))|(((~iext(uri_rdf_first,X38,X39)|~iext(uri_rdf_rest,X38,X40))|~iext(uri_rdf_first,X40,X41))|~iext(uri_rdf_rest,X40,uri_rdf_nil)))&(((iext(X41,esk153_5(X37,X38,X39,X40,X41),esk154_5(X37,X38,X39,X40,X41))|((~ip(X37)|~ip(X39))|~ip(X41)))|iext(uri_owl_propertyChainAxiom,X37,X38))|(((~iext(uri_rdf_first,X38,X39)|~iext(uri_rdf_rest,X38,X40))|~iext(uri_rdf_first,X40,X41))|~iext(uri_rdf_rest,X40,uri_rdf_nil))))&(((~iext(X37,esk152_5(X37,X38,X39,X40,X41),esk154_5(X37,X38,X39,X40,X41))|((~ip(X37)|~ip(X39))|~ip(X41)))|iext(uri_owl_propertyChainAxiom,X37,X38))|(((~iext(uri_rdf_first,X38,X39)|~iext(uri_rdf_rest,X38,X40))|~iext(uri_rdf_first,X40,X41))|~iext(uri_rdf_rest,X40,uri_rdf_nil)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])).
fof(c_0_5, plain, (((((((((((((((iext(uri_owl_propertyChainAxiom,uri_ex_hasUncle,esk192_0)&iext(uri_rdf_first,esk192_0,uri_ex_hasCousin))&iext(uri_rdf_rest,esk192_0,esk193_0))&iext(uri_rdf_first,esk193_0,uri_ex_hasFather))&iext(uri_rdf_rest,esk193_0,uri_rdf_nil))&iext(uri_owl_propertyChainAxiom,uri_ex_hasCousin,esk194_0))&iext(uri_rdf_first,esk194_0,uri_ex_hasUncle))&iext(uri_rdf_rest,esk194_0,esk195_0))&iext(uri_rdf_first,esk195_0,esk196_0))&iext(uri_rdf_rest,esk195_0,uri_rdf_nil))&iext(uri_owl_inverseOf,esk196_0,uri_ex_hasFather))&iext(uri_ex_hasFather,uri_ex_alice,uri_ex_dave))&iext(uri_ex_hasCousin,uri_ex_alice,uri_ex_bob))&iext(uri_ex_hasFather,uri_ex_bob,uri_ex_charly))&iext(uri_ex_hasUncle,uri_ex_bob,uri_ex_dave))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[c_0_1])])])).
cnf(c_0_6,plain,(iext(X5,X6,X7)|~iext(uri_rdf_rest,X1,uri_rdf_nil)|~iext(uri_rdf_first,X1,X2)|~iext(uri_rdf_rest,X3,X1)|~iext(uri_rdf_first,X3,X4)|~iext(uri_owl_propertyChainAxiom,X5,X3)|~iext(X2,X8,X7)|~iext(X4,X6,X8)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(iext(uri_rdf_rest,esk195_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,plain,(iext(X1,X2,X3)|~iext(uri_rdf_first,esk195_0,X4)|~iext(uri_rdf_rest,X5,esk195_0)|~iext(uri_rdf_first,X5,X6)|~iext(uri_owl_propertyChainAxiom,X1,X5)|~iext(X6,X2,X7)|~iext(X4,X7,X3)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_9,plain,(iext(uri_rdf_first,esk195_0,esk196_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,plain,(iext(uri_rdf_rest,esk193_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11,plain,(iext(X1,X2,X3)|~iext(uri_rdf_rest,X4,esk195_0)|~iext(uri_rdf_first,X4,X5)|~iext(uri_owl_propertyChainAxiom,X1,X4)|~iext(esk196_0,X6,X3)|~iext(X5,X2,X6)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12,plain,(iext(uri_rdf_rest,esk194_0,esk195_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_13, plain, (![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:((((ip(X35)|~iext(uri_owl_inverseOf,X35,X36))&(ip(X36)|~iext(uri_owl_inverseOf,X35,X36)))&(((~iext(X35,X37,X38)|iext(X36,X38,X37))|~iext(uri_owl_inverseOf,X35,X36))&((~iext(X36,X40,X39)|iext(X35,X39,X40))|~iext(uri_owl_inverseOf,X35,X36))))&((((~iext(X41,esk159_2(X41,X42),esk160_2(X41,X42))|~iext(X42,esk160_2(X41,X42),esk159_2(X41,X42)))|(~ip(X41)|~ip(X42)))|iext(uri_owl_inverseOf,X41,X42))&(((iext(X41,esk159_2(X41,X42),esk160_2(X41,X42))|iext(X42,esk160_2(X41,X42),esk159_2(X41,X42)))|(~ip(X41)|~ip(X42)))|iext(uri_owl_inverseOf,X41,X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
cnf(c_0_14,plain,(iext(X1,X2,X3)|~iext(uri_rdf_first,esk193_0,X4)|~iext(uri_rdf_rest,X5,esk193_0)|~iext(uri_rdf_first,X5,X6)|~iext(uri_owl_propertyChainAxiom,X1,X5)|~iext(X6,X2,X7)|~iext(X4,X7,X3)), inference(spm,[status(thm)],[c_0_6, c_0_10])).
cnf(c_0_15,plain,(iext(uri_rdf_first,esk193_0,uri_ex_hasFather)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16,plain,(iext(X1,X2,X3)|~iext(uri_rdf_first,esk194_0,X4)|~iext(uri_owl_propertyChainAxiom,X1,esk194_0)|~iext(esk196_0,X5,X3)|~iext(X4,X2,X5)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17,plain,(iext(uri_rdf_first,esk194_0,uri_ex_hasUncle)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18,plain,(iext(X1,X3,X4)|~iext(uri_owl_inverseOf,X1,X2)|~iext(X2,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(iext(uri_owl_inverseOf,esk196_0,uri_ex_hasFather)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_20,plain,(iext(X1,X2,X3)|~iext(uri_rdf_rest,X4,esk193_0)|~iext(uri_rdf_first,X4,X5)|~iext(uri_owl_propertyChainAxiom,X1,X4)|~iext(uri_ex_hasFather,X6,X3)|~iext(X5,X2,X6)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21,plain,(iext(uri_rdf_rest,esk192_0,esk193_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_22,plain,(iext(X1,X2,X3)|~iext(uri_owl_propertyChainAxiom,X1,esk194_0)|~iext(esk196_0,X4,X3)|~iext(uri_ex_hasUncle,X2,X4)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23,plain,(iext(uri_owl_propertyChainAxiom,uri_ex_hasCousin,esk194_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_24,plain,(iext(esk196_0,X1,X2)|~iext(uri_ex_hasFather,X2,X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25,plain,(iext(uri_ex_hasFather,uri_ex_alice,uri_ex_dave)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_26,plain,(iext(X1,X2,X3)|~iext(uri_rdf_first,esk192_0,X4)|~iext(uri_owl_propertyChainAxiom,X1,esk192_0)|~iext(uri_ex_hasFather,X5,X3)|~iext(X4,X2,X5)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,plain,(iext(uri_rdf_first,esk192_0,uri_ex_hasCousin)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_28, negated_conjecture, (~((iext(uri_ex_hasUncle,uri_ex_alice,uri_ex_charly)&iext(uri_ex_hasCousin,uri_ex_bob,uri_ex_alice)))), inference(assume_negation,[status(cth)],[c_0_3])).
cnf(c_0_29,plain,(iext(uri_ex_hasCousin,X1,X2)|~iext(esk196_0,X3,X2)|~iext(uri_ex_hasUncle,X1,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30,plain,(iext(esk196_0,uri_ex_dave,uri_ex_alice)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31,plain,(iext(X1,X2,X3)|~iext(uri_owl_propertyChainAxiom,X1,esk192_0)|~iext(uri_ex_hasFather,X4,X3)|~iext(uri_ex_hasCousin,X2,X4)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32,plain,(iext(uri_owl_propertyChainAxiom,uri_ex_hasUncle,esk192_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_33, negated_conjecture, ((~iext(uri_ex_hasUncle,uri_ex_alice,uri_ex_charly)|~iext(uri_ex_hasCousin,uri_ex_bob,uri_ex_alice))), inference(fof_nnf,[status(thm)],[c_0_28])).
cnf(c_0_34,plain,(iext(uri_ex_hasCousin,X1,uri_ex_alice)|~iext(uri_ex_hasUncle,X1,uri_ex_dave)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35,plain,(iext(uri_ex_hasUncle,uri_ex_bob,uri_ex_dave)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_36,plain,(iext(uri_ex_hasUncle,X1,X2)|~iext(uri_ex_hasFather,X3,X2)|~iext(uri_ex_hasCousin,X1,X3)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37,plain,(iext(uri_ex_hasFather,uri_ex_bob,uri_ex_charly)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_38,negated_conjecture,(~iext(uri_ex_hasCousin,uri_ex_bob,uri_ex_alice)|~iext(uri_ex_hasUncle,uri_ex_alice,uri_ex_charly)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39,plain,(iext(uri_ex_hasCousin,uri_ex_bob,uri_ex_alice)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40,plain,(iext(uri_ex_hasUncle,X1,uri_ex_charly)|~iext(uri_ex_hasCousin,X1,uri_ex_bob)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41,plain,(iext(uri_ex_hasCousin,uri_ex_alice,uri_ex_bob)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_42,negated_conjecture,(~iext(uri_ex_hasUncle,uri_ex_alice,uri_ex_charly)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39])])).
cnf(c_0_43,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), ['proof']).
# SZS output end CNFRefutation.
EOF
