% Problem    : NUM926+7 : TPTP v6.1.0. Released v5.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n097.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:49:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X19]:hAPP(int,int,plus_plus(int,pls),X19)=ti(int,X19)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_126_add__Pls)).
fof(c_0_1, axiom, (![X19]:number_number_of(int,X19)=ti(int,X19)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_88_number__of__is__id)).
fof(c_0_2, conjecture, (?[X6]:?[X7]:hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,X6),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,X7),number_number_of(nat,bit0(bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_3, hypothesis, (![X4]:bit0(ti(int,X4))=bit0(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tsy_c_Int_OBit0_arg1)).
fof(c_0_4, axiom, (![X19]:bit0(X19)=hAPP(int,int,plus_plus(int,X19),X19)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_128_Bit0__def)).
fof(c_0_5, axiom, (![X14]:![X12]:hAPP(int,int,plus_plus(int,number_number_of(int,X14)),number_number_of(int,X12))=number_number_of(int,hAPP(int,int,plus_plus(int,X14),X12))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_134_plus__numeral__code_I9_J)).
fof(c_0_6, hypothesis, (![X4]:ti(int,bit1(X4))=bit1(X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tsy_c_Int_OBit1_res)).
fof(c_0_7, axiom, ((t=one_one(int)=>?[X6]:?[X7]:hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,X6),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,X7),number_number_of(nat,bit0(bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1__096t_A_061_A1_A_061_061_062_AEX_Ax_Ay_O_Ax_A_094_A2_A_L_Ay_A_094_A2_A_06)).
fof(c_0_8, axiom, (![X19]:![X50]:hAPP(int,int,plus_plus(int,bit1(X19)),bit0(X50))=bit1(hAPP(int,int,plus_plus(int,X19),X50))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_151_add__Bit1__Bit0)).
fof(c_0_9, axiom, (hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat))=number_number_of(nat,bit0(bit1(pls)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_37_nat__1__add__1)).
fof(c_0_10, axiom, (![X16]:![X12]:hAPP(int,int,times_times(int,X16),X12)=hAPP(int,int,times_times(int,X12),X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_87_zmult__commute)).
fof(c_0_11, axiom, (![X16]:![X12]:hAPP(int,int,plus_plus(int,X16),X12)=hAPP(int,int,plus_plus(int,X12),X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_91_zadd__commute)).
fof(c_0_12, axiom, ((hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),one_one(int)),t))=>?[X6]:?[X7]:hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,X6),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,X7),number_number_of(nat,bit0(bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_2__0961_A_060_At_A_061_061_062_AEX_Ax_Ay_O_Ax_A_094_A2_A_L_Ay_A_094_A2_A_06)).
fof(c_0_13, axiom, (![X42]:![X43]:![X44]:hAPP(int,int,plus_plus(int,hAPP(int,int,plus_plus(int,X42),X43)),X44)=hAPP(int,int,plus_plus(int,X42),hAPP(int,int,plus_plus(int,X43),X44))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_89_zadd__assoc)).
fof(c_0_14, axiom, (one_one(int)=number_number_of(int,bit1(pls))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_160_one__is__num__one)).
fof(c_0_15, axiom, (![X3]:(order(X3)=>![X10]:![X11]:(hBOOL(hAPP(X3,bool,hAPP(X3,fun(X3,bool),ord_less_eq(X3),X10),X11))=>(ti(X3,X10)!=ti(X3,X11)=>hBOOL(hAPP(X3,bool,hAPP(X3,fun(X3,bool),ord_less(X3),X10),X11)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_332_order__le__neq__implies__less)).
fof(c_0_16, axiom, (hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less_eq(int),one_one(int)),t))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_0_tpos)).
fof(c_0_17, axiom, (ti(int,t)=t), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tsy_v_t_____res)).
fof(c_0_18, axiom, (order(int)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Int_Oint___Orderings_Oorder)).
fof(c_0_19, plain, (![X20]:hAPP(int,int,plus_plus(int,pls),X20)=ti(int,X20)), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_20, plain, (![X20]:number_number_of(int,X20)=ti(int,X20)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_21, negated_conjecture, (~(?[X6]:?[X7]:hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,X6),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,X7),number_number_of(nat,bit0(bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int)))), inference(assume_negation,[status(cth)],[c_0_2])).
fof(c_0_22, hypothesis, (![X5]:bit0(ti(int,X5))=bit0(X5)), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_23, plain, (![X20]:bit0(X20)=hAPP(int,int,plus_plus(int,X20),X20)), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_24, plain, (![X15]:![X16]:hAPP(int,int,plus_plus(int,number_number_of(int,X15)),number_number_of(int,X16))=number_number_of(int,hAPP(int,int,plus_plus(int,X15),X16))), inference(variable_rename,[status(thm)],[c_0_5])).
cnf(c_0_25,plain,(hAPP(int,int,plus_plus(int,pls),X1)=ti(int,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(number_number_of(int,X1)=ti(int,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_27, negated_conjecture, (![X8]:![X9]:hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,X8),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,X9),number_number_of(nat,bit0(bit1(pls)))))!=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).
fof(c_0_28, hypothesis, (![X5]:ti(int,bit1(X5))=bit1(X5)), inference(variable_rename,[status(thm)],[c_0_6])).
fof(c_0_29, plain, ((t!=one_one(int)|hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,esk1_0),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,esk2_0),number_number_of(nat,bit0(bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_30,hypothesis,(bit0(ti(int,X1))=bit0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31,plain,(bit0(X1)=hAPP(int,int,plus_plus(int,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,plain,(hAPP(int,int,plus_plus(int,number_number_of(int,X1)),number_number_of(int,X2))=number_number_of(int,hAPP(int,int,plus_plus(int,X1),X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,plain,(number_number_of(int,X1)=hAPP(int,int,plus_plus(int,pls),X1)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_34, plain, (![X51]:![X52]:hAPP(int,int,plus_plus(int,bit1(X51)),bit0(X52))=bit1(hAPP(int,int,plus_plus(int,X51),X52))), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_35,negated_conjecture,(hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,X1),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,X2),number_number_of(nat,bit0(bit1(pls)))))!=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36,plain,(hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat))=number_number_of(nat,bit0(bit1(pls)))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_37, plain, (![X17]:![X18]:hAPP(int,int,times_times(int,X17),X18)=hAPP(int,int,times_times(int,X18),X17)), inference(variable_rename,[status(thm)],[c_0_10])).
fof(c_0_38, plain, (![X17]:![X18]:hAPP(int,int,plus_plus(int,X17),X18)=hAPP(int,int,plus_plus(int,X18),X17)), inference(variable_rename,[status(thm)],[c_0_11])).
cnf(c_0_39,hypothesis,(ti(int,bit1(X1))=bit1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_40, plain, ((~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),one_one(int)),t))|hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,esk3_0),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,esk4_0),number_number_of(nat,bit0(bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_41,plain,(hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,esk1_0),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,esk2_0),number_number_of(nat,bit0(bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int))|t!=one_one(int)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42,hypothesis,(hAPP(int,int,plus_plus(int,ti(int,X1)),ti(int,X1))=hAPP(int,int,plus_plus(int,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_31])).
cnf(c_0_43,plain,(hAPP(int,int,plus_plus(int,hAPP(int,int,plus_plus(int,pls),X1)),hAPP(int,int,plus_plus(int,pls),X2))=hAPP(int,int,plus_plus(int,pls),hAPP(int,int,plus_plus(int,X1),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_33]), c_0_33])).
fof(c_0_44, plain, (![X45]:![X46]:![X47]:hAPP(int,int,plus_plus(int,hAPP(int,int,plus_plus(int,X45),X46)),X47)=hAPP(int,int,plus_plus(int,X45),hAPP(int,int,plus_plus(int,X46),X47))), inference(variable_rename,[status(thm)],[c_0_13])).
cnf(c_0_45,plain,(hAPP(int,int,plus_plus(int,bit1(X1)),bit0(X2))=bit1(hAPP(int,int,plus_plus(int,X1),X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46,negated_conjecture,(hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,hAPP(int,int,plus_plus(int,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))),hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))))),m)),one_one(int))!=hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,X1),number_number_of(nat,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))))),hAPP(nat,int,power_power(int,X2),number_number_of(nat,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_31]), c_0_31]), c_0_31]), c_0_31])).
cnf(c_0_47,plain,(number_number_of(nat,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls)))=hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat))), inference(rw,[status(thm)],[c_0_36, c_0_31])).
cnf(c_0_48,plain,(hAPP(int,int,times_times(int,X1),X2)=hAPP(int,int,times_times(int,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49,plain,(hAPP(int,int,plus_plus(int,X1),X2)=hAPP(int,int,plus_plus(int,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50,plain,(one_one(int)=number_number_of(int,bit1(pls))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_51,hypothesis,(number_number_of(int,bit1(X1))=bit1(X1)), inference(rw,[status(thm)],[c_0_39, c_0_26])).
cnf(c_0_52,plain,(hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,esk3_0),number_number_of(nat,bit0(bit1(pls))))),hAPP(nat,int,power_power(int,esk4_0),number_number_of(nat,bit0(bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls))))),m)),one_one(int))|~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),one_one(int)),t))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_53,plain,(hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,esk1_0),number_number_of(nat,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))))),hAPP(nat,int,power_power(int,esk2_0),number_number_of(nat,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,hAPP(int,int,plus_plus(int,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))),hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))))),m)),one_one(int))|one_one(int)!=t), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_31]), c_0_31]), c_0_31]), c_0_31])).
cnf(c_0_54,hypothesis,(hAPP(int,int,plus_plus(int,pls),hAPP(int,int,plus_plus(int,X1),X1))=hAPP(int,int,plus_plus(int,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_26]), c_0_33]), c_0_26]), c_0_33]), c_0_43])).
cnf(c_0_55,plain,(hAPP(int,int,plus_plus(int,hAPP(int,int,plus_plus(int,X1),X2)),X3)=hAPP(int,int,plus_plus(int,X1),hAPP(int,int,plus_plus(int,X2),X3))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56,plain,(hAPP(int,int,plus_plus(int,bit1(X1)),hAPP(int,int,plus_plus(int,X2),X2))=bit1(hAPP(int,int,plus_plus(int,X1),X2))), inference(rw,[status(thm)],[c_0_45, c_0_31])).
cnf(c_0_57,negated_conjecture,(hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,X1),hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat)))),hAPP(nat,int,power_power(int,X2),hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat))))!=hAPP(int,int,plus_plus(int,one_one(int)),hAPP(int,int,times_times(int,m),hAPP(int,int,plus_plus(int,pls),hAPP(int,int,plus_plus(int,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))),hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_47]), c_0_33]), c_0_48]), c_0_49])).
cnf(c_0_58,plain,(one_one(int)=bit1(pls)), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_59,hypothesis,(hAPP(int,int,plus_plus(int,pls),bit1(X1))=bit1(X1)), inference(rw,[status(thm)],[c_0_51, c_0_33])).
cnf(c_0_60,plain,(hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,esk3_0),number_number_of(nat,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))))),hAPP(nat,int,power_power(int,esk4_0),number_number_of(nat,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls)))))=hAPP(int,int,plus_plus(int,hAPP(int,int,times_times(int,number_number_of(int,hAPP(int,int,plus_plus(int,hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))),hAPP(int,int,plus_plus(int,bit1(pls)),bit1(pls))))),m)),one_one(int))|~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),one_one(int)),t))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_31]), c_0_31]), c_0_31]), c_0_31])).
fof(c_0_61, plain, (![X12]:![X13]:![X14]:(~order(X12)|(~hBOOL(hAPP(X12,bool,hAPP(X12,fun(X12,bool),ord_less_eq(X12),X13),X14))|(ti(X12,X13)=ti(X12,X14)|hBOOL(hAPP(X12,bool,hAPP(X12,fun(X12,bool),ord_less(X12),X13),X14)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
cnf(c_0_62,plain,(hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less_eq(int),one_one(int)),t))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_63,plain,(ti(int,t)=t), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_64,plain,(hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,esk1_0),hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat)))),hAPP(nat,int,power_power(int,esk2_0),hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat))))=hAPP(int,int,plus_plus(int,one_one(int)),hAPP(int,int,times_times(int,m),hAPP(int,int,plus_plus(int,bit1(pls)),bit1(hAPP(int,int,plus_plus(int,pls),bit1(pls))))))|one_one(int)!=t), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_47]), c_0_47]), c_0_33]), c_0_54]), c_0_55]), c_0_56]), c_0_48]), c_0_49])).
cnf(c_0_65,negated_conjecture,(hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,X1),hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat)))),hAPP(nat,int,power_power(int,X2),hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat))))!=hAPP(int,int,plus_plus(int,bit1(pls)),hAPP(int,int,times_times(int,m),hAPP(int,int,plus_plus(int,bit1(pls)),bit1(bit1(pls)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_54]), c_0_58]), c_0_55]), c_0_56]), c_0_59])).
cnf(c_0_66,plain,(hAPP(int,int,plus_plus(int,hAPP(nat,int,power_power(int,esk3_0),hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat)))),hAPP(nat,int,power_power(int,esk4_0),hAPP(nat,nat,plus_plus(nat,one_one(nat)),one_one(nat))))=hAPP(int,int,plus_plus(int,one_one(int)),hAPP(int,int,times_times(int,m),hAPP(int,int,plus_plus(int,bit1(pls)),bit1(hAPP(int,int,plus_plus(int,pls),bit1(pls))))))|~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),one_one(int)),t))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_47]), c_0_47]), c_0_33]), c_0_54]), c_0_55]), c_0_56]), c_0_48]), c_0_49])).
cnf(c_0_67,plain,(hBOOL(hAPP(X1,bool,hAPP(X1,fun(X1,bool),ord_less(X1),X2),X3))|ti(X1,X2)=ti(X1,X3)|~hBOOL(hAPP(X1,bool,hAPP(X1,fun(X1,bool),ord_less_eq(X1),X2),X3))|~order(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_68,plain,(hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less_eq(int),bit1(pls)),t))), inference(rw,[status(thm)],[c_0_62, c_0_58])).
cnf(c_0_69,plain,(ti(int,X1)=hAPP(int,int,plus_plus(int,pls),X1)), inference(rw,[status(thm)],[c_0_26, c_0_33])).
cnf(c_0_70,plain,(hAPP(int,int,plus_plus(int,pls),t)=t), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_26]), c_0_33])).
cnf(c_0_71,plain,(order(int)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_72,plain,(bit1(pls)!=t), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_58]), c_0_58]), c_0_59]), c_0_65])).
cnf(c_0_73,plain,(~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),bit1(pls)),t))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_58]), c_0_58]), c_0_59]), c_0_65])).
cnf(c_0_74,plain,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_59]), c_0_26]), c_0_33]), c_0_70]), c_0_71])]), c_0_72]), c_0_73]), ['proof']).
# SZS output end CNFRefutation.
EOF
