% Problem    : SWB066+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n187.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:59:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C45_F1_AE_R8_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Presaturation interreduction done
# No success with AutoSched1
# Trying AutoSched2 for 23 seconds
# AutoSched2-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X11]:![X12]:![X21]:((iext(uri_rdf_first,X12,X21)&iext(uri_rdf_rest,X12,uri_rdf_nil))=>(iext(uri_owl_oneOf,X11,X12)<=>(ic(X11)&![X4]:(icext(X11,X4)<=>X4=X21))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_enum_class_001)).
fof(c_0_1, axiom, (?[X42]:?[X43]:?[X44]:?[X45]:?[X46]:?[X47]:?[X37]:?[X38]:(((((((((((((((((iext(uri_owl_oneOf,X44,X37)&iext(uri_rdfs_domain,uri_ex_p,X38))&iext(uri_rdfs_range,uri_ex_p,X42))&iext(uri_owl_oneOf,X42,X46))&iext(uri_rdf_first,X47,uri_ex_s))&iext(uri_rdf_rest,X47,uri_rdf_nil))&iext(uri_ex_p,uri_ex_s,uri_ex_o))&iext(uri_rdf_first,X37,uri_ex_o))&iext(uri_rdf_rest,X37,uri_rdf_nil))&iext(uri_owl_oneOf,X45,X47))&iext(uri_rdfs_domain,uri_ex_q,X44))&iext(uri_rdfs_range,uri_ex_q,X45))&iext(uri_ex_q,uri_ex_o,uri_ex_s))&iext(uri_rdf_first,X43,uri_ex_s))&iext(uri_rdf_rest,X43,uri_rdf_nil))&iext(uri_rdf_first,X46,uri_ex_o))&iext(uri_rdf_rest,X46,uri_rdf_nil))&iext(uri_owl_oneOf,X38,X43))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', premise_rdfbased_sem_inv_ext)).
fof(c_0_2, axiom, (![X2]:![X5]:(iext(uri_rdfs_domain,X2,X5)<=>((ip(X2)&ic(X5))&![X4]:![X7]:(iext(X2,X4,X7)=>icext(X5,X4))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_rdfsext_domain)).
fof(c_0_3, axiom, (![X1]:![X2]:![X3]:(iext(X2,X1,X3)=>ip(X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', simple_iext_property)).
fof(c_0_4, axiom, (![X2]:![X5]:(iext(uri_rdfs_range,X2,X5)<=>((ip(X2)&ip(X5))&![X4]:![X7]:(iext(X2,X4,X7)=>icext(X5,X7))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_rdfsext_range)).
fof(c_0_5, axiom, (![X2]:(icext(uri_owl_InverseFunctionalProperty,X2)<=>(ip(X2)&![X37]:![X38]:![X7]:((iext(X2,X37,X7)&iext(X2,X38,X7))=>X37=X38)))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_char_inversefunctional)).
fof(c_0_6, axiom, (![X4]:![X7]:(iext(uri_owl_oneOf,X4,X7)=>(ic(X4)&icext(uri_rdf_List,X7)))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_prop_oneof_ext)).
fof(c_0_7, axiom, (![X2]:(icext(uri_owl_FunctionalProperty,X2)<=>(ip(X2)&![X4]:![X27]:![X28]:((iext(X2,X4,X27)&iext(X2,X4,X28))=>X27=X28)))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_char_functional)).
fof(c_0_8, axiom, (![X33]:![X34]:(iext(uri_owl_inverseOf,X33,X34)<=>((ip(X33)&ip(X34))&![X4]:![X7]:(iext(X33,X4,X7)<=>iext(X34,X7,X4))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_inv)).
fof(c_0_9, conjecture, (iext(uri_owl_inverseOf,uri_ex_q,uri_ex_p)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conclusion_rdfbased_sem_inv_ext)).
fof(c_0_10, plain, (![X22]:![X23]:![X24]:![X25]:![X26]:((((ic(X22)|~iext(uri_owl_oneOf,X22,X23))|(~iext(uri_rdf_first,X23,X24)|~iext(uri_rdf_rest,X23,uri_rdf_nil)))&((((~icext(X22,X25)|X25=X24)|~iext(uri_owl_oneOf,X22,X23))|(~iext(uri_rdf_first,X23,X24)|~iext(uri_rdf_rest,X23,uri_rdf_nil)))&(((X26!=X24|icext(X22,X26))|~iext(uri_owl_oneOf,X22,X23))|(~iext(uri_rdf_first,X23,X24)|~iext(uri_rdf_rest,X23,uri_rdf_nil)))))&(((((~icext(X22,esk11_3(X22,X23,X24))|esk11_3(X22,X23,X24)!=X24)|~ic(X22))|iext(uri_owl_oneOf,X22,X23))|(~iext(uri_rdf_first,X23,X24)|~iext(uri_rdf_rest,X23,uri_rdf_nil)))&((((icext(X22,esk11_3(X22,X23,X24))|esk11_3(X22,X23,X24)=X24)|~ic(X22))|iext(uri_owl_oneOf,X22,X23))|(~iext(uri_rdf_first,X23,X24)|~iext(uri_rdf_rest,X23,uri_rdf_nil)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])])).
fof(c_0_11, plain, ((((((((((((((((((iext(uri_owl_oneOf,esk194_0,esk198_0)&iext(uri_rdfs_domain,uri_ex_p,esk199_0))&iext(uri_rdfs_range,uri_ex_p,esk192_0))&iext(uri_owl_oneOf,esk192_0,esk196_0))&iext(uri_rdf_first,esk197_0,uri_ex_s))&iext(uri_rdf_rest,esk197_0,uri_rdf_nil))&iext(uri_ex_p,uri_ex_s,uri_ex_o))&iext(uri_rdf_first,esk198_0,uri_ex_o))&iext(uri_rdf_rest,esk198_0,uri_rdf_nil))&iext(uri_owl_oneOf,esk195_0,esk197_0))&iext(uri_rdfs_domain,uri_ex_q,esk194_0))&iext(uri_rdfs_range,uri_ex_q,esk195_0))&iext(uri_ex_q,uri_ex_o,uri_ex_s))&iext(uri_rdf_first,esk193_0,uri_ex_s))&iext(uri_rdf_rest,esk193_0,uri_rdf_nil))&iext(uri_rdf_first,esk196_0,uri_ex_o))&iext(uri_rdf_rest,esk196_0,uri_rdf_nil))&iext(uri_owl_oneOf,esk199_0,esk193_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_1])])).
cnf(c_0_12,plain,(X4=X2|~iext(uri_rdf_rest,X1,uri_rdf_nil)|~iext(uri_rdf_first,X1,X2)|~iext(uri_owl_oneOf,X3,X1)|~icext(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13,plain,(iext(uri_rdf_rest,esk198_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:((((ip(X8)|~iext(uri_rdfs_domain,X8,X9))&(ic(X9)|~iext(uri_rdfs_domain,X8,X9)))&((~iext(X8,X10,X11)|icext(X9,X10))|~iext(uri_rdfs_domain,X8,X9)))&(((iext(X12,esk105_2(X12,X13),esk106_2(X12,X13))|(~ip(X12)|~ic(X13)))|iext(uri_rdfs_domain,X12,X13))&((~icext(X13,esk105_2(X12,X13))|(~ip(X12)|~ic(X13)))|iext(uri_rdfs_domain,X12,X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
fof(c_0_15, plain, (![X4]:![X5]:![X6]:(~iext(X5,X4,X6)|ip(X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_16,plain,(X1=X2|~icext(X3,X2)|~iext(uri_rdf_first,esk198_0,X1)|~iext(uri_owl_oneOf,X3,esk198_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17,plain,(iext(uri_rdf_first,esk198_0,uri_ex_o)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,plain,(icext(X2,X3)|~iext(uri_rdfs_domain,X1,X2)|~iext(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,plain,(iext(uri_rdfs_domain,uri_ex_q,esk194_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20,plain,(ip(X1)|~iext(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(iext(uri_ex_q,uri_ex_o,uri_ex_s)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22,plain,(iext(uri_rdf_rest,esk193_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_23,plain,(uri_ex_o=X1|~icext(X2,X1)|~iext(uri_owl_oneOf,X2,esk198_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24,plain,(iext(uri_owl_oneOf,esk194_0,esk198_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25,plain,(icext(esk194_0,X1)|~iext(uri_ex_q,X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26,plain,(iext(uri_rdfs_domain,X1,X2)|iext(X1,esk105_2(X1,X2),esk106_2(X1,X2))|~ic(X2)|~ip(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_27,plain,(ip(uri_ex_q)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_28, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:((((ip(X8)|~iext(uri_rdfs_range,X8,X9))&(ip(X9)|~iext(uri_rdfs_range,X8,X9)))&((~iext(X8,X10,X11)|icext(X9,X11))|~iext(uri_rdfs_range,X8,X9)))&(((iext(X12,esk107_2(X12,X13),esk108_2(X12,X13))|(~ip(X12)|~ip(X13)))|iext(uri_rdfs_range,X12,X13))&((~icext(X13,esk108_2(X12,X13))|(~ip(X12)|~ip(X13)))|iext(uri_rdfs_range,X12,X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
cnf(c_0_29,plain,(iext(uri_rdf_rest,esk197_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_30,plain,(X1=X2|~icext(X3,X2)|~iext(uri_rdf_first,esk193_0,X1)|~iext(uri_owl_oneOf,X3,esk193_0)), inference(spm,[status(thm)],[c_0_12, c_0_22])).
cnf(c_0_31,plain,(iext(uri_rdf_first,esk193_0,uri_ex_s)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_32,plain,(iext(uri_rdfs_domain,uri_ex_p,esk199_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_33, plain, (![X39]:![X40]:![X41]:![X42]:![X43]:(((ip(X39)|~icext(uri_owl_InverseFunctionalProperty,X39))&(((~iext(X39,X40,X42)|~iext(X39,X41,X42))|X40=X41)|~icext(uri_owl_InverseFunctionalProperty,X39)))&((((iext(X43,esk166_1(X43),esk168_1(X43))|~ip(X43))|icext(uri_owl_InverseFunctionalProperty,X43))&((iext(X43,esk167_1(X43),esk168_1(X43))|~ip(X43))|icext(uri_owl_InverseFunctionalProperty,X43)))&((esk166_1(X43)!=esk167_1(X43)|~ip(X43))|icext(uri_owl_InverseFunctionalProperty,X43))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).
cnf(c_0_34,plain,(iext(uri_ex_p,uri_ex_s,uri_ex_o)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_35,plain,(uri_ex_o=X1|~icext(esk194_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_36,plain,(icext(esk194_0,esk105_2(uri_ex_q,X1))|iext(uri_rdfs_domain,uri_ex_q,X1)|~ic(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_37,plain,(icext(X2,X3)|~iext(uri_rdfs_range,X1,X2)|~iext(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38,plain,(iext(uri_rdfs_range,uri_ex_p,esk192_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_39, plain, (![X8]:![X9]:((ic(X8)|~iext(uri_owl_oneOf,X8,X9))&(icext(uri_rdf_List,X9)|~iext(uri_owl_oneOf,X8,X9)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_40,plain,(X1=X2|~icext(X3,X2)|~iext(uri_rdf_first,esk197_0,X1)|~iext(uri_owl_oneOf,X3,esk197_0)), inference(spm,[status(thm)],[c_0_12, c_0_29])).
cnf(c_0_41,plain,(iext(uri_rdf_first,esk197_0,uri_ex_s)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_42,plain,(iext(uri_rdfs_range,uri_ex_q,esk195_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_43, plain, (![X29]:![X30]:![X31]:![X32]:![X33]:(((ip(X29)|~icext(uri_owl_FunctionalProperty,X29))&(((~iext(X29,X30,X31)|~iext(X29,X30,X32))|X31=X32)|~icext(uri_owl_FunctionalProperty,X29)))&((((iext(X33,esk163_1(X33),esk164_1(X33))|~ip(X33))|icext(uri_owl_FunctionalProperty,X33))&((iext(X33,esk163_1(X33),esk165_1(X33))|~ip(X33))|icext(uri_owl_FunctionalProperty,X33)))&((esk164_1(X33)!=esk165_1(X33)|~ip(X33))|icext(uri_owl_FunctionalProperty,X33))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])])).
cnf(c_0_44,plain,(uri_ex_s=X1|~icext(X2,X1)|~iext(uri_owl_oneOf,X2,esk193_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_45,plain,(iext(uri_owl_oneOf,esk199_0,esk193_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_46,plain,(icext(esk199_0,X1)|~iext(uri_ex_p,X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_32])).
cnf(c_0_47,plain,(icext(uri_owl_InverseFunctionalProperty,X1)|iext(X1,esk167_1(X1),esk168_1(X1))|~ip(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_48,plain,(ip(uri_ex_p)), inference(spm,[status(thm)],[c_0_20, c_0_34])).
cnf(c_0_49,plain,(iext(uri_rdfs_domain,X1,X2)|~ic(X2)|~ip(X1)|~icext(X2,esk105_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_50,plain,(esk105_2(uri_ex_q,X1)=uri_ex_o|iext(uri_rdfs_domain,uri_ex_q,X1)|~ic(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_51,plain,(icext(esk192_0,X1)|~iext(uri_ex_p,X2,X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_52,plain,(ic(X1)|~iext(uri_owl_oneOf,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53,plain,(iext(uri_owl_oneOf,esk192_0,esk196_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_54,plain,(uri_ex_s=X1|~icext(X2,X1)|~iext(uri_owl_oneOf,X2,esk197_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_55,plain,(iext(uri_owl_oneOf,esk195_0,esk197_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_56,plain,(icext(esk195_0,X1)|~iext(uri_ex_q,X2,X1)), inference(spm,[status(thm)],[c_0_37, c_0_42])).
cnf(c_0_57,plain,(icext(uri_owl_FunctionalProperty,X1)|iext(X1,esk163_1(X1),esk165_1(X1))|~ip(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_58,plain,(uri_ex_s=X1|~icext(esk199_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_59,plain,(icext(esk199_0,esk167_1(uri_ex_p))|icext(uri_owl_InverseFunctionalProperty,uri_ex_p)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])])).
cnf(c_0_60,plain,(iext(uri_rdf_rest,esk196_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_61,plain,(iext(uri_rdfs_domain,uri_ex_q,X1)|~ic(X1)|~icext(X1,uri_ex_o)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_27])])).
cnf(c_0_62,plain,(icext(esk192_0,uri_ex_o)), inference(spm,[status(thm)],[c_0_51, c_0_34])).
cnf(c_0_63,plain,(ic(esk192_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
fof(c_0_64, plain, (![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:((((ip(X35)|~iext(uri_owl_inverseOf,X35,X36))&(ip(X36)|~iext(uri_owl_inverseOf,X35,X36)))&(((~iext(X35,X37,X38)|iext(X36,X38,X37))|~iext(uri_owl_inverseOf,X35,X36))&((~iext(X36,X40,X39)|iext(X35,X39,X40))|~iext(uri_owl_inverseOf,X35,X36))))&((((~iext(X41,esk159_2(X41,X42),esk160_2(X41,X42))|~iext(X42,esk160_2(X41,X42),esk159_2(X41,X42)))|(~ip(X41)|~ip(X42)))|iext(uri_owl_inverseOf,X41,X42))&(((iext(X41,esk159_2(X41,X42),esk160_2(X41,X42))|iext(X42,esk160_2(X41,X42),esk159_2(X41,X42)))|(~ip(X41)|~ip(X42)))|iext(uri_owl_inverseOf,X41,X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])])).
cnf(c_0_65,plain,(uri_ex_s=X1|~icext(esk195_0,X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_66,plain,(icext(esk195_0,esk165_1(uri_ex_q))|icext(uri_owl_FunctionalProperty,uri_ex_q)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_27])])).
cnf(c_0_67,plain,(icext(uri_owl_InverseFunctionalProperty,X1)|iext(X1,esk166_1(X1),esk168_1(X1))|~ip(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_68,plain,(icext(uri_owl_InverseFunctionalProperty,X1)|~ip(X1)|esk166_1(X1)!=esk167_1(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_69,plain,(esk167_1(uri_ex_p)=uri_ex_s|icext(uri_owl_InverseFunctionalProperty,uri_ex_p)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_70,plain,(X1=X2|~icext(X3,X2)|~iext(uri_rdf_first,esk196_0,X1)|~iext(uri_owl_oneOf,X3,esk196_0)), inference(spm,[status(thm)],[c_0_12, c_0_60])).
cnf(c_0_71,plain,(iext(uri_rdf_first,esk196_0,uri_ex_o)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_72,plain,(iext(uri_rdfs_domain,uri_ex_q,esk192_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])])).
cnf(c_0_73,plain,(iext(uri_owl_inverseOf,X1,X2)|iext(X2,esk160_2(X1,X2),esk159_2(X1,X2))|iext(X1,esk159_2(X1,X2),esk160_2(X1,X2))|~ip(X2)|~ip(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_74,plain,(icext(uri_owl_FunctionalProperty,X1)|iext(X1,esk163_1(X1),esk164_1(X1))|~ip(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_75,plain,(icext(uri_owl_FunctionalProperty,X1)|~ip(X1)|esk164_1(X1)!=esk165_1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_76,plain,(esk165_1(uri_ex_q)=uri_ex_s|icext(uri_owl_FunctionalProperty,uri_ex_q)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_77,plain,(icext(esk199_0,esk166_1(uri_ex_p))|icext(uri_owl_InverseFunctionalProperty,uri_ex_p)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_67]), c_0_48])])).
cnf(c_0_78,plain,(icext(uri_owl_InverseFunctionalProperty,uri_ex_p)|esk166_1(uri_ex_p)!=uri_ex_s), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_48])])).
cnf(c_0_79,plain,(uri_ex_o=X1|~icext(X2,X1)|~iext(uri_owl_oneOf,X2,esk196_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_80,plain,(icext(esk192_0,X1)|~iext(uri_ex_q,X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_72])).
cnf(c_0_81,plain,(icext(esk192_0,esk160_2(uri_ex_p,X1))|iext(X1,esk160_2(uri_ex_p,X1),esk159_2(uri_ex_p,X1))|iext(uri_owl_inverseOf,uri_ex_p,X1)|~ip(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_73]), c_0_48])])).
cnf(c_0_82,plain,(icext(esk195_0,esk164_1(uri_ex_q))|icext(uri_owl_FunctionalProperty,uri_ex_q)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_74]), c_0_27])])).
cnf(c_0_83,plain,(icext(uri_owl_FunctionalProperty,uri_ex_q)|esk164_1(uri_ex_q)!=uri_ex_s), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_27])])).
cnf(c_0_84,plain,(X2=X3|~icext(uri_owl_InverseFunctionalProperty,X1)|~iext(X1,X3,X4)|~iext(X1,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_85,plain,(icext(uri_owl_InverseFunctionalProperty,uri_ex_p)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_77]), c_0_78])).
cnf(c_0_86,plain,(uri_ex_o=X1|~icext(esk192_0,X1)), inference(spm,[status(thm)],[c_0_79, c_0_53])).
cnf(c_0_87,plain,(icext(esk192_0,esk160_2(uri_ex_p,uri_ex_q))|iext(uri_owl_inverseOf,uri_ex_p,uri_ex_q)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_27])])).
cnf(c_0_88,plain,(X2=X3|~icext(uri_owl_FunctionalProperty,X1)|~iext(X1,X4,X3)|~iext(X1,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_89,plain,(icext(uri_owl_FunctionalProperty,uri_ex_q)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_82]), c_0_83])).
cnf(c_0_90,plain,(iext(uri_owl_inverseOf,X1,X2)|~ip(X2)|~ip(X1)|~iext(X2,esk160_2(X1,X2),esk159_2(X1,X2))|~iext(X1,esk159_2(X1,X2),esk160_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_91,plain,(X1=X2|~iext(uri_ex_p,X2,X3)|~iext(uri_ex_p,X1,X3)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92,plain,(esk160_2(uri_ex_p,uri_ex_q)=uri_ex_o|iext(uri_owl_inverseOf,uri_ex_p,uri_ex_q)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_93,plain,(X1=X2|~iext(uri_ex_q,X3,X2)|~iext(uri_ex_q,X3,X1)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
fof(c_0_94, negated_conjecture, (~iext(uri_owl_inverseOf,uri_ex_q,uri_ex_p)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_9])])).
cnf(c_0_95,plain,(iext(uri_owl_inverseOf,X1,X2)|~iext(X2,esk160_2(X1,X2),esk159_2(X1,X2))|~iext(X1,esk159_2(X1,X2),esk160_2(X1,X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_90, c_0_20]), c_0_20])).
cnf(c_0_96,plain,(X1=uri_ex_s|~iext(uri_ex_p,X1,uri_ex_o)), inference(spm,[status(thm)],[c_0_91, c_0_34])).
cnf(c_0_97,plain,(iext(uri_ex_p,esk159_2(uri_ex_p,uri_ex_q),uri_ex_o)|iext(uri_ex_q,uri_ex_o,esk159_2(uri_ex_p,uri_ex_q))|iext(uri_owl_inverseOf,uri_ex_p,uri_ex_q)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_92]), c_0_27]), c_0_48])])).
cnf(c_0_98,plain,(X1=uri_ex_s|~iext(uri_ex_q,uri_ex_o,X1)), inference(spm,[status(thm)],[c_0_93, c_0_21])).
cnf(c_0_99,plain,(icext(esk192_0,esk159_2(X1,uri_ex_p))|iext(X1,esk159_2(X1,uri_ex_p),esk160_2(X1,uri_ex_p))|iext(uri_owl_inverseOf,X1,uri_ex_p)|~ip(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_73]), c_0_48])])).
cnf(c_0_100,negated_conjecture,(~iext(uri_owl_inverseOf,uri_ex_q,uri_ex_p)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_101,plain,(iext(uri_owl_inverseOf,uri_ex_p,uri_ex_q)|~iext(uri_ex_q,uri_ex_o,esk159_2(uri_ex_p,uri_ex_q))|~iext(uri_ex_p,esk159_2(uri_ex_p,uri_ex_q),uri_ex_o)), inference(spm,[status(thm)],[c_0_95, c_0_92])).
cnf(c_0_102,plain,(esk159_2(uri_ex_p,uri_ex_q)=uri_ex_s|iext(uri_owl_inverseOf,uri_ex_p,uri_ex_q)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98])).
cnf(c_0_103,plain,(icext(esk192_0,esk159_2(uri_ex_q,uri_ex_p))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_99]), c_0_27])]), c_0_100])).
cnf(c_0_104,plain,(iext(X2,X3,X4)|~iext(uri_owl_inverseOf,X1,X2)|~iext(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_105,plain,(iext(uri_owl_inverseOf,uri_ex_p,uri_ex_q)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_21]), c_0_34])])).
cnf(c_0_106,plain,(iext(X1,X3,X4)|~iext(uri_owl_inverseOf,X1,X2)|~iext(X2,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_107,plain,(esk159_2(uri_ex_q,uri_ex_p)=uri_ex_o), inference(spm,[status(thm)],[c_0_86, c_0_103])).
cnf(c_0_108,plain,(iext(uri_ex_q,X1,X2)|~iext(uri_ex_p,X2,X1)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_109,plain,(iext(uri_ex_p,X1,X2)|~iext(uri_ex_q,X2,X1)), inference(spm,[status(thm)],[c_0_106, c_0_105])).
cnf(c_0_110,plain,(icext(esk199_0,esk160_2(X1,uri_ex_p))|iext(X1,esk159_2(X1,uri_ex_p),esk160_2(X1,uri_ex_p))|iext(uri_owl_inverseOf,X1,uri_ex_p)|~ip(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_73]), c_0_48])])).
cnf(c_0_111,plain,(~iext(uri_ex_p,esk160_2(uri_ex_q,uri_ex_p),uri_ex_o)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_107]), c_0_100]), c_0_108])).
cnf(c_0_112,plain,(icext(esk199_0,esk160_2(uri_ex_q,uri_ex_p))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_107]), c_0_27])]), c_0_111]), c_0_100])).
cnf(c_0_113,plain,(esk160_2(uri_ex_q,uri_ex_p)=uri_ex_s), inference(spm,[status(thm)],[c_0_58, c_0_112])).
cnf(c_0_114,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_113]), c_0_34])]), ['proof']).
# SZS output end CNFRefutation.
EOF
