% Problem    : SCT166+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n119.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:47:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X46]:~(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X46),X46)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_less__irrefl__nat)).
fof(c_0_1, axiom, (![X46]:![X48]:(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X48),X46))=>hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X48),hAPP(c_Nat_OSuc,X46))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_less__SucI)).
fof(c_0_2, axiom, (![X56]:(c_Finite__Set_Ofinite(tc_Nat_Onat,X56)<=>?[X17]:![X3]:(hBOOL(hAPP(hAPP(c_member(tc_Nat_Onat),X3),X56))=>hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X3),X17))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite__nat__set__iff__bounded)).
fof(c_0_3, axiom, (![X14]:![X8]:![X2]:![X10]:hBOOL(hAPP(hAPP(c_member(X10),hAPP(X2,X8)),c_Set_Oimage(X14,X10,X2,c_Orderings_Otop__class_Otop(tc_fun(X14,tc_HOL_Obool)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rangeI)).
fof(c_0_4, axiom, (![X65]:![X14]:![X40]:![X10]:(c_Finite__Set_Ofinite(X10,X40)=>c_Finite__Set_Ofinite(X14,c_Set_Oimage(X10,X14,X65,X40)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite__imageI)).
fof(c_0_5, axiom, (c_Finite__Set_Ofinite(tc_Arrow__Order__Mirabelle_Oindi,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oindi,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_finite__indi)).
fof(c_0_6, plain, (![X46]:~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X46),X46))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_7, plain, (![X47]:~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X47),X47))), inference(variable_rename,[status(thm)],[c_0_6])).
fof(c_0_8, plain, (![X49]:![X50]:(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X50),X49))|hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X50),hAPP(c_Nat_OSuc,X49))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
cnf(c_0_9,plain,(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X1),X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,plain,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X1),hAPP(c_Nat_OSuc,X2)))|~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X1),X2))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, plain, (![X57]:![X59]:![X60]:![X61]:((~c_Finite__Set_Ofinite(tc_Nat_Onat,X57)|(~hBOOL(hAPP(hAPP(c_member(tc_Nat_Onat),X59),X57))|hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X59),esk41_1(X57)))))&((hBOOL(hAPP(hAPP(c_member(tc_Nat_Onat),esk42_2(X60,X61)),X60))|c_Finite__Set_Ofinite(tc_Nat_Onat,X60))&(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),esk42_2(X60,X61)),X61))|c_Finite__Set_Ofinite(tc_Nat_Onat,X60))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
cnf(c_0_12,plain,(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),hAPP(c_Nat_OSuc,X1)),X1))), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13,plain,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless(tc_Nat_Onat),X1),esk41_1(X2)))|~hBOOL(hAPP(hAPP(c_member(tc_Nat_Onat),X1),X2))|~c_Finite__Set_Ofinite(tc_Nat_Onat,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, (![X15]:![X16]:![X17]:![X18]:hBOOL(hAPP(hAPP(c_member(X18),hAPP(X17,X16)),c_Set_Oimage(X15,X18,X17,c_Orderings_Otop__class_Otop(tc_fun(X15,tc_HOL_Obool)))))), inference(variable_rename,[status(thm)],[c_0_3])).
cnf(c_0_15,plain,(~c_Finite__Set_Ofinite(tc_Nat_Onat,X1)|~hBOOL(hAPP(hAPP(c_member(tc_Nat_Onat),hAPP(c_Nat_OSuc,esk41_1(X1))),X1))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16,plain,(hBOOL(hAPP(hAPP(c_member(X1),hAPP(X2,X3)),c_Set_Oimage(X4,X1,X2,c_Orderings_Otop__class_Otop(tc_fun(X4,tc_HOL_Obool)))))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, (![X66]:![X67]:![X68]:![X69]:(~c_Finite__Set_Ofinite(X69,X68)|c_Finite__Set_Ofinite(X67,c_Set_Oimage(X69,X67,X66,X68)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
cnf(c_0_18,plain,(~c_Finite__Set_Ofinite(tc_Nat_Onat,c_Set_Oimage(X1,tc_Nat_Onat,c_Nat_OSuc,c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19,plain,(c_Finite__Set_Ofinite(X1,c_Set_Oimage(X2,X1,X3,X4))|~c_Finite__Set_Ofinite(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20,plain,(c_Finite__Set_Ofinite(tc_Arrow__Order__Mirabelle_Oindi,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oindi,tc_HOL_Obool)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_21,plain,(~c_Finite__Set_Ofinite(X1,c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22,plain,($false), inference(sr,[status(thm)],[c_0_20, c_0_21]), ['proof']).
# SZS output end CNFRefutation.
EOF
