% Problem    : SWB014+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n131.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:49:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C45_F1_AE_R8_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X11]:![X12]:![X13]:![X14]:![X15]:((((iext(uri_rdf_first,X12,X13)&iext(uri_rdf_rest,X12,X14))&iext(uri_rdf_first,X14,X15))&iext(uri_rdf_rest,X14,uri_rdf_nil))=>(iext(uri_owl_unionOf,X11,X12)<=>(((ic(X11)&ic(X13))&ic(X15))&![X4]:(icext(X11,X4)<=>(icext(X13,X4)|icext(X15,X4))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_bool_unionof_class_002)).
fof(c_0_1, axiom, (?[X43]:?[X44]:?[X45]:(((((((iext(uri_rdf_type,uri_ex_Eagle,uri_ex_Species)&iext(uri_rdf_type,uri_ex_Falcon,uri_ex_Species))&iext(uri_rdf_type,uri_ex_harry,X43))&iext(uri_owl_unionOf,X43,X44))&iext(uri_rdf_first,X44,uri_ex_Eagle))&iext(uri_rdf_rest,X44,X45))&iext(uri_rdf_first,X45,uri_ex_Falcon))&iext(uri_rdf_rest,X45,uri_rdf_nil))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_premise_fullish_014_Harry_belongs_to_some_Species)).
fof(c_0_2, axiom, (![X4]:![X5]:(iext(uri_rdf_type,X4,X5)<=>icext(X5,X4))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', rdfs_cext_def)).
fof(c_0_3, conjecture, (?[X42]:(iext(uri_rdf_type,uri_ex_harry,X42)&iext(uri_rdf_type,X42,uri_ex_Species))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_conclusion_fullish_014_Harry_belongs_to_some_Species)).
fof(c_0_4, plain, (![X16]:![X17]:![X18]:![X19]:![X20]:![X21]:![X22]:((((((ic(X16)|~iext(uri_owl_unionOf,X16,X17))|(((~iext(uri_rdf_first,X17,X18)|~iext(uri_rdf_rest,X17,X19))|~iext(uri_rdf_first,X19,X20))|~iext(uri_rdf_rest,X19,uri_rdf_nil)))&((ic(X18)|~iext(uri_owl_unionOf,X16,X17))|(((~iext(uri_rdf_first,X17,X18)|~iext(uri_rdf_rest,X17,X19))|~iext(uri_rdf_first,X19,X20))|~iext(uri_rdf_rest,X19,uri_rdf_nil))))&((ic(X20)|~iext(uri_owl_unionOf,X16,X17))|(((~iext(uri_rdf_first,X17,X18)|~iext(uri_rdf_rest,X17,X19))|~iext(uri_rdf_first,X19,X20))|~iext(uri_rdf_rest,X19,uri_rdf_nil))))&((((~icext(X16,X21)|(icext(X18,X21)|icext(X20,X21)))|~iext(uri_owl_unionOf,X16,X17))|(((~iext(uri_rdf_first,X17,X18)|~iext(uri_rdf_rest,X17,X19))|~iext(uri_rdf_first,X19,X20))|~iext(uri_rdf_rest,X19,uri_rdf_nil)))&((((~icext(X18,X22)|icext(X16,X22))|~iext(uri_owl_unionOf,X16,X17))|(((~iext(uri_rdf_first,X17,X18)|~iext(uri_rdf_rest,X17,X19))|~iext(uri_rdf_first,X19,X20))|~iext(uri_rdf_rest,X19,uri_rdf_nil)))&(((~icext(X20,X22)|icext(X16,X22))|~iext(uri_owl_unionOf,X16,X17))|(((~iext(uri_rdf_first,X17,X18)|~iext(uri_rdf_rest,X17,X19))|~iext(uri_rdf_first,X19,X20))|~iext(uri_rdf_rest,X19,uri_rdf_nil))))))&((((((~icext(X18,esk8_5(X16,X17,X18,X19,X20))|~icext(X16,esk8_5(X16,X17,X18,X19,X20)))|((~ic(X16)|~ic(X18))|~ic(X20)))|iext(uri_owl_unionOf,X16,X17))|(((~iext(uri_rdf_first,X17,X18)|~iext(uri_rdf_rest,X17,X19))|~iext(uri_rdf_first,X19,X20))|~iext(uri_rdf_rest,X19,uri_rdf_nil)))&((((~icext(X20,esk8_5(X16,X17,X18,X19,X20))|~icext(X16,esk8_5(X16,X17,X18,X19,X20)))|((~ic(X16)|~ic(X18))|~ic(X20)))|iext(uri_owl_unionOf,X16,X17))|(((~iext(uri_rdf_first,X17,X18)|~iext(uri_rdf_rest,X17,X19))|~iext(uri_rdf_first,X19,X20))|~iext(uri_rdf_rest,X19,uri_rdf_nil))))&((((icext(X16,esk8_5(X16,X17,X18,X19,X20))|(icext(X18,esk8_5(X16,X17,X18,X19,X20))|icext(X20,esk8_5(X16,X17,X18,X19,X20))))|((~ic(X16)|~ic(X18))|~ic(X20)))|iext(uri_owl_unionOf,X16,X17))|(((~iext(uri_rdf_first,X17,X18)|~iext(uri_rdf_rest,X17,X19))|~iext(uri_rdf_first,X19,X20))|~iext(uri_rdf_rest,X19,uri_rdf_nil)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])])).
fof(c_0_5, plain, ((((((((iext(uri_rdf_type,uri_ex_Eagle,uri_ex_Species)&iext(uri_rdf_type,uri_ex_Falcon,uri_ex_Species))&iext(uri_rdf_type,uri_ex_harry,esk192_0))&iext(uri_owl_unionOf,esk192_0,esk193_0))&iext(uri_rdf_first,esk193_0,uri_ex_Eagle))&iext(uri_rdf_rest,esk193_0,esk194_0))&iext(uri_rdf_first,esk194_0,uri_ex_Falcon))&iext(uri_rdf_rest,esk194_0,uri_rdf_nil))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_1])])).
cnf(c_0_6,plain,(icext(X2,X6)|icext(X4,X6)|~iext(uri_rdf_rest,X1,uri_rdf_nil)|~iext(uri_rdf_first,X1,X2)|~iext(uri_rdf_rest,X3,X1)|~iext(uri_rdf_first,X3,X4)|~iext(uri_owl_unionOf,X5,X3)|~icext(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7,plain,(iext(uri_rdf_rest,esk194_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,plain,(icext(X1,X2)|icext(X3,X2)|~icext(X4,X2)|~iext(uri_rdf_first,esk194_0,X3)|~iext(uri_rdf_rest,X5,esk194_0)|~iext(uri_rdf_first,X5,X1)|~iext(uri_owl_unionOf,X4,X5)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_9,plain,(iext(uri_rdf_first,esk194_0,uri_ex_Falcon)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,plain,(icext(uri_ex_Falcon,X1)|icext(X2,X1)|~icext(X3,X1)|~iext(uri_rdf_rest,X4,esk194_0)|~iext(uri_rdf_first,X4,X2)|~iext(uri_owl_unionOf,X3,X4)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_11,plain,(iext(uri_rdf_rest,esk193_0,esk194_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12,plain,(icext(uri_ex_Falcon,X1)|icext(X2,X1)|~icext(X3,X1)|~iext(uri_rdf_first,esk193_0,X2)|~iext(uri_owl_unionOf,X3,esk193_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_13,plain,(iext(uri_rdf_first,esk193_0,uri_ex_Eagle)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_14, plain, (![X6]:![X7]:![X8]:![X9]:((~iext(uri_rdf_type,X6,X7)|icext(X7,X6))&(~icext(X9,X8)|iext(uri_rdf_type,X8,X9)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
fof(c_0_15, negated_conjecture, (~(?[X42]:(iext(uri_rdf_type,uri_ex_harry,X42)&iext(uri_rdf_type,X42,uri_ex_Species)))), inference(assume_negation,[status(cth)],[c_0_3])).
cnf(c_0_16,plain,(icext(uri_ex_Eagle,X1)|icext(uri_ex_Falcon,X1)|~icext(X2,X1)|~iext(uri_owl_unionOf,X2,esk193_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17,plain,(iext(uri_owl_unionOf,esk192_0,esk193_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18,plain,(icext(X1,X2)|~iext(uri_rdf_type,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,plain,(iext(uri_rdf_type,uri_ex_harry,esk192_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_20, negated_conjecture, (![X43]:(~iext(uri_rdf_type,uri_ex_harry,X43)|~iext(uri_rdf_type,X43,uri_ex_Species))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
cnf(c_0_21,plain,(icext(uri_ex_Falcon,X1)|icext(uri_ex_Eagle,X1)|~icext(esk192_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22,plain,(icext(esk192_0,uri_ex_harry)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23,negated_conjecture,(~iext(uri_rdf_type,X1,uri_ex_Species)|~iext(uri_rdf_type,uri_ex_harry,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24,plain,(iext(uri_rdf_type,uri_ex_Falcon,uri_ex_Species)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25,plain,(iext(uri_rdf_type,X1,X2)|~icext(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26,plain,(icext(uri_ex_Eagle,uri_ex_harry)|icext(uri_ex_Falcon,uri_ex_harry)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27,negated_conjecture,(~iext(uri_rdf_type,uri_ex_harry,uri_ex_Falcon)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28,plain,(iext(uri_rdf_type,uri_ex_Eagle,uri_ex_Species)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29,plain,(icext(uri_ex_Eagle,uri_ex_harry)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
cnf(c_0_30,negated_conjecture,(~iext(uri_rdf_type,uri_ex_harry,uri_ex_Eagle)), inference(spm,[status(thm)],[c_0_23, c_0_28])).
cnf(c_0_31,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_29]), c_0_30]), ['proof']).
# SZS output end CNFRefutation.
EOF
