% Problem    : SWW098+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n100.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:33:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_C45_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (?[X5]:(false2=phi(f7(X5))&~(?[X10]:forallprefers(f7(X10),f7(X5))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_false2)).
fof(c_0_1, axiom, (![X5]:f7(X5)=lazy_impl(prop(X5),X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_f7)).
fof(c_0_2, axiom, (![X1]:(prop(X1)=false<=>~(bool(X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', prop_false)).
fof(c_0_3, axiom, (![X1]:((d(X1)&phi(X1)=X1)|(~(d(X1))&phi(X1)=err))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_phi)).
fof(c_0_4, axiom, (![X4]:lazy_impl(false,X4)=true), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', lazy_impl_axiom2)).
fof(c_0_5, axiom, (![X1]:![X2]:(forallprefers(X1,X2)<=>(((~(d(X1))&d(X2))|(((d(X1)&d(X2))&~(bool(X1)))&bool(X2)))|(X1=false&X2=true)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_forallprefers)).
fof(c_0_6, axiom, (![X1]:(prop(X1)=true<=>bool(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', prop_true)).
fof(c_0_7, axiom, (![X4]:lazy_impl(true,X4)=phi(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', lazy_impl_axiom3)).
fof(c_0_8, axiom, (![X1]:(bool(X1)<=>(X1=false|X1=true))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_bool)).
fof(c_0_9, axiom, (((d(true)&d(false))&d(err))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', false_true_err_in_d)).
fof(c_0_10, axiom, (![X3]:![X4]:(~(bool(X3))=>lazy_impl(X3,X4)=phi(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', lazy_impl_axiom1)).
fof(c_0_11, axiom, (![X5]:![X6]:?[X7]:(lazy_and2(X5,X6)=phi(f2(X5,X6,X7))&~(?[X8]:forallprefers(f2(X5,X6,X8),f2(X5,X6,X7))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_lazy_and2)).
fof(c_0_12, axiom, (![X5]:![X6]:![X7]:f2(X5,X6,X7)=lazy_impl(prop(X7),impl(lazy_impl(X5,impl(X6,X7)),X7))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_f2)).
fof(c_0_13, axiom, (![X4]:(bool(X4)=>impl(true,X4)=X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', impl_axiom4)).
fof(c_0_14, axiom, (((true!=false&true!=err)&false!=err)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', distinct_false_true_err)).
fof(c_0_15, axiom, (![X3]:![X4]:(~(bool(X3))=>impl(X3,X4)=phi(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', impl_axiom1)).
fof(c_0_16, axiom, (![X4]:(bool(X4)=>impl(false,X4)=true)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', impl_axiom3)).
fof(c_0_17, conjecture, (![X5]:![X6]:or1(X5,X6)=or2(X5,X6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', or1_or2)).
fof(c_0_18, axiom, (![X5]:![X6]:?[X7]:(or2(X5,X6)=phi(f3(X5,X6,X7))&~(?[X8]:forallprefers(f3(X5,X6,X8),f3(X5,X6,X7))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_or2)).
fof(c_0_19, axiom, (![X5]:![X6]:![X7]:f3(X5,X6,X7)=lazy_impl(prop(X7),impl(impl(X5,X7),impl(impl(X6,X7),X7)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_f3)).
fof(c_0_20, axiom, (![X3]:![X4]:(~(bool(X3))=>or1(X3,X4)=phi(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', or1_axiom1)).
fof(c_0_21, axiom, (![X3]:![X4]:((bool(X3)&~(bool(X4)))=>impl(X3,X4)=phi(X4))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', impl_axiom2)).
fof(c_0_22, axiom, (![X3]:![X4]:((bool(X3)&~(bool(X4)))=>or1(X3,X4)=phi(X4))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', or1_axiom2)).
fof(c_0_23, axiom, (![X5]:![X6]:?[X7]:(and2(X5,X6)=phi(f1(X5,X6,X7))&~(?[X8]:forallprefers(f1(X5,X6,X8),f1(X5,X6,X7))))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_and2)).
fof(c_0_24, axiom, (![X5]:![X6]:![X7]:f1(X5,X6,X7)=lazy_impl(prop(X7),impl(impl(X5,impl(X6,X7)),X7))), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', def_f1)).
fof(c_0_25, axiom, (![X4]:(bool(X4)=>or1(true,X4)=true)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', or1_axiom3)).
fof(c_0_26, axiom, (![X4]:(bool(X4)=>or1(false,X4)=X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWV012+0.ax', or1_axiom4)).
fof(c_0_27, plain, (![X12]:(false2=phi(f7(esk7_0))&~forallprefers(f7(X12),f7(esk7_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])).
fof(c_0_28, plain, (![X6]:f7(X6)=lazy_impl(prop(X6),X6)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_29, plain, (![X1]:(prop(X1)=false<=>~bool(X1))), inference(fof_simplification,[status(thm)],[c_0_2])).
fof(c_0_30, plain, (![X1]:((d(X1)&phi(X1)=X1)|(~d(X1)&phi(X1)=err))), inference(fof_simplification,[status(thm)],[c_0_3])).
cnf(c_0_31,plain,(~forallprefers(f7(X1),f7(esk7_0))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32,plain,(f7(X1)=lazy_impl(prop(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_33, plain, (![X2]:![X3]:((prop(X2)!=false|~bool(X2))&(bool(X3)|prop(X3)=false))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])).
fof(c_0_34, plain, (![X5]:lazy_impl(false,X5)=true), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_35, plain, (![X1]:![X2]:(forallprefers(X1,X2)<=>(((~d(X1)&d(X2))|(((d(X1)&d(X2))&~bool(X1))&bool(X2)))|(X1=false&X2=true)))), inference(fof_simplification,[status(thm)],[c_0_5])).
fof(c_0_36, plain, (![X2]:![X3]:((prop(X2)!=true|bool(X2))&(~bool(X3)|prop(X3)=true))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
fof(c_0_37, plain, (![X2]:(((~d(X2)|d(X2))&(phi(X2)=err|d(X2)))&((~d(X2)|phi(X2)=X2)&(phi(X2)=err|phi(X2)=X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[c_0_30])])).
fof(c_0_38, plain, (![X5]:lazy_impl(true,X5)=phi(X5)), inference(variable_rename,[status(thm)],[c_0_7])).
cnf(c_0_39,plain,(~forallprefers(lazy_impl(prop(X1),X1),lazy_impl(prop(esk7_0),esk7_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32]), c_0_32])).
cnf(c_0_40,plain,(prop(X1)=false|bool(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41,plain,(lazy_impl(false,X1)=true), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_42, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=false|(d(X3)|~d(X3)))|~forallprefers(X3,X4))&((X4=true|(d(X3)|~d(X3)))|~forallprefers(X3,X4)))&(((X3=false|(d(X4)|~d(X3)))|~forallprefers(X3,X4))&((X4=true|(d(X4)|~d(X3)))|~forallprefers(X3,X4))))&(((X3=false|(~bool(X3)|~d(X3)))|~forallprefers(X3,X4))&((X4=true|(~bool(X3)|~d(X3)))|~forallprefers(X3,X4))))&(((X3=false|(bool(X4)|~d(X3)))|~forallprefers(X3,X4))&((X4=true|(bool(X4)|~d(X3)))|~forallprefers(X3,X4))))&((((((X3=false|(d(X3)|d(X4)))|~forallprefers(X3,X4))&((X4=true|(d(X3)|d(X4)))|~forallprefers(X3,X4)))&(((X3=false|(d(X4)|d(X4)))|~forallprefers(X3,X4))&((X4=true|(d(X4)|d(X4)))|~forallprefers(X3,X4))))&(((X3=false|(~bool(X3)|d(X4)))|~forallprefers(X3,X4))&((X4=true|(~bool(X3)|d(X4)))|~forallprefers(X3,X4))))&(((X3=false|(bool(X4)|d(X4)))|~forallprefers(X3,X4))&((X4=true|(bool(X4)|d(X4)))|~forallprefers(X3,X4)))))&((((d(X5)|~d(X6))|forallprefers(X5,X6))&((((~d(X5)|~d(X6))|bool(X5))|~bool(X6))|forallprefers(X5,X6)))&((X5!=false|X6!=true)|forallprefers(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])])])).
cnf(c_0_43,plain,(prop(X1)=true|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44,plain,(phi(X1)=X1|~d(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45,plain,(lazy_impl(true,X1)=phi(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46,plain,(bool(esk7_0)|~forallprefers(lazy_impl(prop(X1),X1),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_47,plain,(forallprefers(X1,X2)|X2!=true|X1!=false), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_48, plain, (![X2]:![X3]:((~bool(X2)|(X2=false|X2=true))&((X3!=false|bool(X3))&(X3!=true|bool(X3))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
cnf(c_0_49,plain,(~forallprefers(lazy_impl(true,X1),lazy_impl(prop(esk7_0),esk7_0))|~bool(X1)), inference(spm,[status(thm)],[c_0_39, c_0_43])).
cnf(c_0_50,plain,(lazy_impl(true,X1)=X1|~d(X1)), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_51,plain,(bool(esk7_0)|~forallprefers(lazy_impl(true,X1),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_46, c_0_43])).
cnf(c_0_52,plain,(forallprefers(false,X1)|true!=X1), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_53,plain,(bool(X1)|X1!=false), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54,plain,(false2=phi(f7(esk7_0))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_55,plain,(~forallprefers(X1,lazy_impl(prop(esk7_0),esk7_0))|~d(X1)|~bool(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_56,plain,(bool(esk7_0)|~forallprefers(X1,true)|~d(X1)|~bool(X1)), inference(spm,[status(thm)],[c_0_51, c_0_50])).
cnf(c_0_57,plain,(forallprefers(false,true)), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_58,plain,(d(false)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_59,plain,(bool(false)), inference(er,[status(thm)],[c_0_53])).
fof(c_0_60, plain, (![X3]:![X4]:(~bool(X3)=>lazy_impl(X3,X4)=phi(X3))), inference(fof_simplification,[status(thm)],[c_0_10])).
cnf(c_0_61,plain,(lazy_impl(true,lazy_impl(prop(esk7_0),esk7_0))=false2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_45]), c_0_32])).
cnf(c_0_62,plain,(~forallprefers(X1,lazy_impl(true,esk7_0))|~d(X1)|~bool(esk7_0)|~bool(X1)), inference(spm,[status(thm)],[c_0_55, c_0_43])).
cnf(c_0_63,plain,(bool(esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), c_0_59])])).
cnf(c_0_64,plain,(X1=true|X1=false|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_65, plain, (![X5]:![X6]:(bool(X5)|lazy_impl(X5,X6)=phi(X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_60])])])])).
cnf(c_0_66,plain,(lazy_impl(true,lazy_impl(true,esk7_0))=false2|~bool(esk7_0)), inference(spm,[status(thm)],[c_0_61, c_0_43])).
cnf(c_0_67,plain,(~forallprefers(X1,lazy_impl(true,esk7_0))|~d(X1)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63])])).
cnf(c_0_68,plain,(d(true)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_69,plain,(esk7_0=false|true=esk7_0), inference(spm,[status(thm)],[c_0_64, c_0_63])).
cnf(c_0_70,plain,(lazy_impl(X1,X2)=phi(X1)|bool(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_71, plain, (![X9]:![X10]:![X12]:(lazy_and2(X9,X10)=phi(f2(X9,X10,esk2_2(X9,X10)))&~forallprefers(f2(X9,X10,X12),f2(X9,X10,esk2_2(X9,X10))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
fof(c_0_72, plain, (![X8]:![X9]:![X10]:f2(X8,X9,X10)=lazy_impl(prop(X10),impl(lazy_impl(X8,impl(X9,X10)),X10))), inference(variable_rename,[status(thm)],[c_0_12])).
cnf(c_0_73,plain,(lazy_impl(true,lazy_impl(true,esk7_0))=false2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_63])])).
cnf(c_0_74,plain,(~forallprefers(X1,esk7_0)|~d(esk7_0)|~d(X1)|~bool(X1)), inference(spm,[status(thm)],[c_0_67, c_0_50])).
cnf(c_0_75,plain,(esk7_0=false|d(esk7_0)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_76,plain,(lazy_impl(true,X1)=lazy_impl(X1,X2)|bool(X1)), inference(rw,[status(thm)],[c_0_70, c_0_45])).
cnf(c_0_77,plain,(~forallprefers(f2(X1,X2,X3),f2(X1,X2,esk2_2(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_78,plain,(f2(X1,X2,X3)=lazy_impl(prop(X3),impl(lazy_impl(X1,impl(X2,X3)),X3))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_79,plain,(lazy_impl(true,esk7_0)=false2|~d(lazy_impl(true,esk7_0))), inference(spm,[status(thm)],[c_0_50, c_0_73])).
cnf(c_0_80,plain,(esk7_0=false|~forallprefers(X1,esk7_0)|~d(X1)|~bool(X1)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_81,plain,(esk7_0=false|forallprefers(false,esk7_0)), inference(spm,[status(thm)],[c_0_57, c_0_69])).
cnf(c_0_82,plain,(~forallprefers(lazy_impl(true,X1),lazy_impl(true,esk7_0))|~bool(esk7_0)|~bool(X1)), inference(spm,[status(thm)],[c_0_49, c_0_43])).
cnf(c_0_83,plain,(~forallprefers(lazy_impl(true,prop(X1)),lazy_impl(prop(esk7_0),esk7_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_76]), c_0_49])).
cnf(c_0_84,plain,(~forallprefers(lazy_impl(prop(X3),impl(lazy_impl(X1,impl(X2,X3)),X3)),lazy_impl(prop(esk2_2(X1,X2)),impl(lazy_impl(X1,impl(X2,esk2_2(X1,X2))),esk2_2(X1,X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_78]), c_0_78])).
cnf(c_0_85,plain,(false2=esk7_0|~d(esk7_0)), inference(spm,[status(thm)],[c_0_79, c_0_50])).
cnf(c_0_86,plain,(esk7_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_58]), c_0_59])])).
cnf(c_0_87,plain,(~forallprefers(lazy_impl(true,X1),lazy_impl(true,esk7_0))|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_63])])).
cnf(c_0_88,plain,(~forallprefers(lazy_impl(true,prop(X1)),lazy_impl(true,esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_43]), c_0_63])])).
cnf(c_0_89,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(prop(esk2_2(false,X2)),impl(true,esk2_2(false,X2))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_41]), c_0_41])).
fof(c_0_90, plain, (![X5]:(~bool(X5)|impl(true,X5)=X5)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).
cnf(c_0_91,plain,(phi(X1)=X1|phi(X1)=err), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_92,plain,(false2=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86]), c_0_86]), c_0_58])])).
cnf(c_0_93,plain,(~forallprefers(lazy_impl(true,X1),esk7_0)|~d(esk7_0)|~bool(X1)), inference(spm,[status(thm)],[c_0_87, c_0_50])).
cnf(c_0_94,plain,(d(X1)|phi(X1)=err), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_95,plain,(forallprefers(X1,X2)|bool(X1)|~bool(X2)|~d(X2)|~d(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_96,plain,(forallprefers(X1,X2)|d(X1)|~d(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_97,plain,(~forallprefers(lazy_impl(true,prop(X1)),esk7_0)|~d(esk7_0)), inference(spm,[status(thm)],[c_0_88, c_0_50])).
cnf(c_0_98,plain,(bool(esk2_2(false,X1))|~forallprefers(lazy_impl(prop(X2),impl(true,X2)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_40]), c_0_41])).
cnf(c_0_99,plain,(impl(true,X1)=X1|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_100,plain,(lazy_impl(true,X1)=X1|lazy_impl(true,X1)=err), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_45]), c_0_45])).
cnf(c_0_101,plain,(lazy_impl(true,lazy_impl(prop(false),false))=false), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_86]), c_0_86]), c_0_92])).
cnf(c_0_102,plain,(false!=err), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_103,plain,(~forallprefers(lazy_impl(true,X1),false)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_86]), c_0_86]), c_0_58])])).
cnf(c_0_104,plain,(lazy_impl(true,X1)=err|d(X1)), inference(rw,[status(thm)],[c_0_94, c_0_45])).
fof(c_0_105, plain, (![X3]:![X4]:(~bool(X3)=>impl(X3,X4)=phi(X3))), inference(fof_simplification,[status(thm)],[c_0_15])).
cnf(c_0_106,plain,(forallprefers(X1,X2)|bool(X1)|~d(X2)|~bool(X2)), inference(csr,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_107,plain,(~forallprefers(lazy_impl(true,prop(X1)),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_86]), c_0_86]), c_0_58])])).
cnf(c_0_108,plain,(bool(esk2_2(false,X1))|~forallprefers(lazy_impl(prop(X2),X2),true)|~bool(X2)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_109,plain,(lazy_impl(prop(false),false)=false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102])).
cnf(c_0_110,plain,(d(X1)|~forallprefers(err,false)|~bool(X1)), inference(spm,[status(thm)],[c_0_103, c_0_104])).
fof(c_0_111, plain, (![X5]:![X6]:(bool(X5)|impl(X5,X6)=phi(X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_105])])])])).
cnf(c_0_112,plain,(lazy_impl(true,X1)=X1|err!=X1), inference(ef,[status(thm)],[c_0_100])).
cnf(c_0_113,plain,(bool(lazy_impl(true,X1))|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_114,plain,(~forallprefers(lazy_impl(prop(X1),X1),lazy_impl(true,esk7_0))|~bool(esk7_0)), inference(spm,[status(thm)],[c_0_39, c_0_43])).
cnf(c_0_115,plain,(bool(lazy_impl(true,prop(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_116,plain,(bool(esk2_2(false,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_57]), c_0_59])])).
cnf(c_0_117,plain,(d(X1)|bool(err)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_118,plain,(true!=err), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_119,plain,(impl(X1,X2)=phi(X1)|bool(X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_120,plain,(lazy_impl(true,err)=err), inference(er,[status(thm)],[c_0_112])).
cnf(c_0_121,plain,(lazy_impl(true,X1)=false|lazy_impl(true,X1)=true|~bool(X1)), inference(spm,[status(thm)],[c_0_64, c_0_113])).
cnf(c_0_122,plain,(~forallprefers(lazy_impl(prop(X1),X1),lazy_impl(true,esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_63])])).
cnf(c_0_123,plain,(lazy_impl(true,prop(X1))=false|lazy_impl(true,prop(X1))=true), inference(spm,[status(thm)],[c_0_64, c_0_115])).
cnf(c_0_124,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(true,impl(true,esk2_2(false,X2))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_43]), c_0_116])])).
cnf(c_0_125,plain,(d(X1)|~bool(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_117]), c_0_118]), c_0_102])).
cnf(c_0_126,plain,(lazy_impl(true,X1)=impl(X1,X2)|bool(X1)), inference(rw,[status(thm)],[c_0_119, c_0_45])).
cnf(c_0_127,plain,(lazy_impl(err,X1)=err|bool(err)), inference(spm,[status(thm)],[c_0_76, c_0_120])).
cnf(c_0_128,plain,(~bool(err)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_120]), c_0_118]), c_0_102])).
cnf(c_0_129,plain,(~forallprefers(lazy_impl(prop(X1),X1),esk7_0)|~d(esk7_0)), inference(spm,[status(thm)],[c_0_122, c_0_50])).
cnf(c_0_130,plain,(lazy_impl(true,prop(X1))=false|prop(X1)=true), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_123]), c_0_118])).
cnf(c_0_131,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(true,esk2_2(false,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_99]), c_0_116])])).
cnf(c_0_132,plain,(d(esk2_2(false,X1))), inference(spm,[status(thm)],[c_0_125, c_0_116])).
cnf(c_0_133,plain,(bool(lazy_impl(X1,impl(X2,esk2_2(X1,X2))))|~forallprefers(lazy_impl(prop(X3),impl(lazy_impl(X1,impl(X2,X3)),X3)),lazy_impl(prop(esk2_2(X1,X2)),lazy_impl(true,lazy_impl(X1,impl(X2,esk2_2(X1,X2))))))), inference(spm,[status(thm)],[c_0_84, c_0_126])).
cnf(c_0_134,plain,(lazy_impl(err,X1)=err), inference(sr,[status(thm)],[c_0_127, c_0_128])).
cnf(c_0_135,plain,(~forallprefers(lazy_impl(prop(X1),X1),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_86]), c_0_86]), c_0_58])])).
cnf(c_0_136,plain,(prop(X1)=true|prop(X1)=false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_130]), c_0_102])).
cnf(c_0_137,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),esk2_2(false,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_50]), c_0_132])])).
cnf(c_0_138,plain,(~forallprefers(lazy_impl(prop(X1),impl(err,X1)),lazy_impl(prop(esk2_2(err,X2)),err))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_134]), c_0_120]), c_0_128])).
cnf(c_0_139,plain,(prop(X1)=false|~forallprefers(lazy_impl(true,X1),false)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_140,plain,(~forallprefers(lazy_impl(prop(X1),X1),esk2_2(false,X2))|~bool(X1)), inference(spm,[status(thm)],[c_0_137, c_0_99])).
cnf(c_0_141,plain,(~forallprefers(lazy_impl(prop(X1),err),lazy_impl(prop(esk2_2(err,X2)),err))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_126]), c_0_120]), c_0_128])).
cnf(c_0_142,plain,(lazy_impl(true,X1)=X1|prop(X1)=false|~forallprefers(err,false)), inference(spm,[status(thm)],[c_0_139, c_0_100])).
cnf(c_0_143,plain,(~forallprefers(false,esk2_2(false,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_109]), c_0_59])])).
cnf(c_0_144,plain,(esk2_2(false,X1)=false|esk2_2(false,X1)=true), inference(spm,[status(thm)],[c_0_64, c_0_116])).
cnf(c_0_145,plain,(~forallprefers(err,lazy_impl(prop(esk2_2(err,X1)),err))|~bool(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_43]), c_0_120])).
cnf(c_0_146,plain,(bool(X1)|X1!=true), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_147,plain,(lazy_impl(true,X1)=X1|prop(X1)=false|bool(err)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_148,plain,(esk2_2(false,X1)=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_57])])).
cnf(c_0_149,plain,(prop(X1)=false|bool(lazy_impl(true,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_150,plain,(bool(esk2_2(err,X1))|~forallprefers(err,true)|~bool(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_40]), c_0_41])).
cnf(c_0_151,plain,(bool(true)), inference(er,[status(thm)],[c_0_146])).
cnf(c_0_152,plain,(lazy_impl(true,X1)=X1|prop(X1)=false), inference(sr,[status(thm)],[c_0_147, c_0_128])).
cnf(c_0_153,plain,(true!=false), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_154,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),false)), inference(rw,[status(thm)],[c_0_137, c_0_148])).
cnf(c_0_155,plain,(bool(lazy_impl(true,esk2_2(err,X1)))|~forallprefers(err,true)|~bool(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_149]), c_0_41])).
cnf(c_0_156,plain,(bool(esk2_2(err,X1))|~bool(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_106]), c_0_68]), c_0_151])]), c_0_128])).
cnf(c_0_157,plain,(bool(esk2_2(X1,X2))|~forallprefers(lazy_impl(prop(X3),impl(lazy_impl(X1,impl(X2,X3)),X3)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_40]), c_0_41])).
cnf(c_0_158,plain,(lazy_impl(true,X1)=X1|~bool(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_152]), c_0_153])).
cnf(c_0_159,plain,(~forallprefers(lazy_impl(true,impl(true,X1)),false)|~bool(X1)), inference(spm,[status(thm)],[c_0_154, c_0_43])).
cnf(c_0_160,plain,(bool(lazy_impl(true,esk2_2(err,X1)))|~bool(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_106]), c_0_68]), c_0_151])]), c_0_128])).
cnf(c_0_161,plain,(bool(esk2_2(err,X1))), inference(spm,[status(thm)],[c_0_156, c_0_116])).
cnf(c_0_162,plain,(bool(esk2_2(true,X1))|~forallprefers(lazy_impl(prop(X2),impl(impl(X1,X2),X2)),true)|~bool(impl(X1,X2))), inference(spm,[status(thm)],[c_0_157, c_0_158])).
cnf(c_0_163,plain,(bool(lazy_impl(true,impl(true,X1)))|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_164,plain,(bool(lazy_impl(true,esk2_2(err,X1)))), inference(spm,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_165,plain,(lazy_impl(X1,X2)=err|lazy_impl(X1,X2)=X1|bool(X1)), inference(spm,[status(thm)],[c_0_100, c_0_76])).
cnf(c_0_166,plain,(bool(esk2_2(true,true))|~forallprefers(lazy_impl(prop(X1),impl(X1,X1)),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_162, c_0_99])).
cnf(c_0_167,plain,(lazy_impl(true,impl(true,X1))=false|lazy_impl(true,impl(true,X1))=true|~bool(X1)), inference(spm,[status(thm)],[c_0_64, c_0_163])).
cnf(c_0_168,plain,(lazy_impl(true,esk2_2(err,X1))=false|lazy_impl(true,esk2_2(err,X1))=true), inference(spm,[status(thm)],[c_0_64, c_0_164])).
cnf(c_0_169,plain,(lazy_impl(X1,X2)=err|lazy_impl(true,X1)=X1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_165]), c_0_158])).
cnf(c_0_170,plain,(bool(esk2_2(true,true))|~forallprefers(lazy_impl(true,impl(X1,X1)),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_166, c_0_43])).
cnf(c_0_171,plain,(lazy_impl(true,impl(true,X1))=false|impl(true,X1)=true|~bool(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_167]), c_0_118])).
cnf(c_0_172,plain,(bool(lazy_impl(prop(X1),impl(true,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_173,plain,(~forallprefers(lazy_impl(prop(X1),impl(lazy_impl(X2,X1),X1)),lazy_impl(prop(esk2_2(X2,true)),impl(lazy_impl(X2,impl(true,esk2_2(X2,true))),esk2_2(X2,true))))|~bool(X1)), inference(spm,[status(thm)],[c_0_84, c_0_99])).
cnf(c_0_174,plain,(lazy_impl(true,true)=true), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_102]), c_0_118])).
cnf(c_0_175,plain,(impl(true,true)=true|bool(esk2_2(true,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171]), c_0_57]), c_0_151])])).
cnf(c_0_176,plain,(prop(X1)=false|bool(lazy_impl(true,impl(true,X1)))), inference(spm,[status(thm)],[c_0_172, c_0_136])).
cnf(c_0_177,plain,(~forallprefers(lazy_impl(prop(true),impl(true,true)),lazy_impl(prop(esk2_2(true,true)),impl(lazy_impl(true,impl(true,esk2_2(true,true))),esk2_2(true,true))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_151])])).
cnf(c_0_178,plain,(impl(true,true)=true|esk2_2(true,true)=false|esk2_2(true,true)=true), inference(spm,[status(thm)],[c_0_64, c_0_175])).
cnf(c_0_179,plain,(prop(X1)=false|d(impl(true,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_104]), c_0_128])).
cnf(c_0_180,plain,(d(impl(true,X1))|~forallprefers(err,false)|~bool(X1)), inference(spm,[status(thm)],[c_0_159, c_0_104])).
cnf(c_0_181,plain,(esk2_2(true,true)=false|impl(true,true)=true|~forallprefers(lazy_impl(prop(true),impl(true,true)),lazy_impl(prop(true),impl(lazy_impl(true,impl(true,true)),true)))), inference(spm,[status(thm)],[c_0_177, c_0_178])).
fof(c_0_182, plain, (![X5]:(~bool(X5)|impl(false,X5)=true)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).
cnf(c_0_183,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(true,impl(true,false)))), inference(rw,[status(thm)],[c_0_124, c_0_148])).
cnf(c_0_184,plain,(d(impl(true,false))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_179]), c_0_41]), c_0_153])).
cnf(c_0_185,plain,(~forallprefers(impl(true,X1),false)|~d(impl(true,X1))|~bool(X1)), inference(spm,[status(thm)],[c_0_159, c_0_50])).
cnf(c_0_186,plain,(d(impl(true,X1))|~bool(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_106]), c_0_58]), c_0_59])]), c_0_128])).
fof(c_0_187, negated_conjecture, (~(![X5]:![X6]:or1(X5,X6)=or2(X5,X6))), inference(assume_negation,[status(cth)],[c_0_17])).
fof(c_0_188, plain, (![X9]:![X10]:![X12]:(or2(X9,X10)=phi(f3(X9,X10,esk3_2(X9,X10)))&~forallprefers(f3(X9,X10,X12),f3(X9,X10,esk3_2(X9,X10))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).
cnf(c_0_189,plain,(impl(true,true)=true|esk2_2(true,true)=false|~forallprefers(lazy_impl(prop(true),impl(true,true)),lazy_impl(prop(true),impl(false,true)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_171]), c_0_151])])).
cnf(c_0_190,plain,(impl(false,X1)=true|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_191,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),impl(true,false))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_50]), c_0_184])])).
cnf(c_0_192,plain,(~forallprefers(impl(true,X1),false)|~bool(X1)), inference(spm,[status(thm)],[c_0_185, c_0_186])).
fof(c_0_193, negated_conjecture, (or1(esk8_0,esk9_0)!=or2(esk8_0,esk9_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_187])])])).
cnf(c_0_194,plain,(or2(X1,X2)=phi(f3(X1,X2,esk3_2(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_188])).
fof(c_0_195, plain, (![X8]:![X9]:![X10]:f3(X8,X9,X10)=lazy_impl(prop(X10),impl(impl(X8,X10),impl(impl(X9,X10),X10)))), inference(variable_rename,[status(thm)],[c_0_19])).
cnf(c_0_196,plain,(esk2_2(true,true)=false|impl(true,true)=true|~forallprefers(lazy_impl(prop(true),impl(true,true)),lazy_impl(prop(true),true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_151])])).
cnf(c_0_197,plain,(~forallprefers(lazy_impl(prop(X1),X1),impl(true,false))|~bool(X1)), inference(spm,[status(thm)],[c_0_191, c_0_99])).
cnf(c_0_198,plain,(bool(impl(true,X1))|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_199,negated_conjecture,(or1(esk8_0,esk9_0)!=or2(esk8_0,esk9_0)), inference(split_conjunct,[status(thm)],[c_0_193])).
cnf(c_0_200,plain,(lazy_impl(true,f3(X1,X2,esk3_2(X1,X2)))=or2(X1,X2)), inference(rw,[status(thm)],[c_0_194, c_0_45])).
cnf(c_0_201,plain,(f3(X1,X2,X3)=lazy_impl(prop(X3),impl(impl(X1,X3),impl(impl(X2,X3),X3)))), inference(split_conjunct,[status(thm)],[c_0_195])).
cnf(c_0_202,plain,(impl(true,true)=true|esk2_2(true,true)=false|~forallprefers(lazy_impl(true,impl(true,true)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_43]), c_0_174]), c_0_151])])).
cnf(c_0_203,plain,(~forallprefers(false,impl(true,false))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_109]), c_0_59])])).
cnf(c_0_204,plain,(impl(true,X1)=false|impl(true,X1)=true|~bool(X1)), inference(spm,[status(thm)],[c_0_64, c_0_198])).
cnf(c_0_205,plain,(lazy_impl(true,esk7_0)=false2|~d(esk7_0)), inference(spm,[status(thm)],[c_0_73, c_0_50])).
cnf(c_0_206,negated_conjecture,(lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),impl(impl(esk9_0,esk3_2(esk8_0,esk9_0)),esk3_2(esk8_0,esk9_0)))))!=or1(esk8_0,esk9_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_199, c_0_200]), c_0_201])).
cnf(c_0_207,plain,(esk2_2(true,true)=false|impl(true,true)=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_171]), c_0_57]), c_0_151])])).
cnf(c_0_208,plain,(impl(true,false)=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_204]), c_0_57]), c_0_59])])).
cnf(c_0_209,plain,(lazy_impl(true,false)=false), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_205, c_0_86]), c_0_86]), c_0_58])]), c_0_92])).
cnf(c_0_210,negated_conjecture,(bool(esk8_0)|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(lazy_impl(true,esk8_0),impl(impl(esk9_0,esk3_2(esk8_0,esk9_0)),esk3_2(esk8_0,esk9_0)))))!=or1(esk8_0,esk9_0)), inference(spm,[status(thm)],[c_0_206, c_0_126])).
fof(c_0_211, plain, (![X3]:![X4]:(~bool(X3)=>or1(X3,X4)=phi(X3))), inference(fof_simplification,[status(thm)],[c_0_20])).
cnf(c_0_212,plain,(impl(true,true)=true|~forallprefers(lazy_impl(prop(true),impl(true,true)),lazy_impl(prop(false),impl(false,false)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_207]), c_0_208]), c_0_209])).
cnf(c_0_213,negated_conjecture,(bool(lazy_impl(true,esk8_0))|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),lazy_impl(true,lazy_impl(true,esk8_0))))!=or1(esk8_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_126]), c_0_113])).
fof(c_0_214, plain, (![X5]:![X6]:(bool(X5)|or1(X5,X6)=phi(X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_211])])])])).
cnf(c_0_215,plain,(impl(true,true)=true|~forallprefers(lazy_impl(prop(true),impl(true,true)),lazy_impl(prop(false),true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_190]), c_0_59])])).
cnf(c_0_216,negated_conjecture,(d(esk8_0)|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),err))!=or1(esk8_0,esk9_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_104]), c_0_120]), c_0_128])).
cnf(c_0_217,plain,(or1(X1,X2)=phi(X1)|bool(X1)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_218,plain,(~forallprefers(f3(X1,X2,X3),f3(X1,X2,esk3_2(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_219,plain,(impl(true,true)=true|~forallprefers(lazy_impl(prop(true),impl(true,true)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_43]), c_0_174]), c_0_59])])).
cnf(c_0_220,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|bool(lazy_impl(true,esk8_0))|lazy_impl(true,lazy_impl(true,lazy_impl(true,lazy_impl(true,esk8_0))))!=or1(esk8_0,esk9_0)), inference(spm,[status(thm)],[c_0_213, c_0_136])).
cnf(c_0_221,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|d(esk8_0)|or1(esk8_0,esk9_0)!=err), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_136]), c_0_120]), c_0_120])).
cnf(c_0_222,plain,(lazy_impl(true,X1)=or1(X1,X2)|bool(X1)), inference(rw,[status(thm)],[c_0_217, c_0_45])).
cnf(c_0_223,plain,(~forallprefers(lazy_impl(prop(X3),impl(impl(X1,X3),impl(impl(X2,X3),X3))),lazy_impl(prop(esk3_2(X1,X2)),impl(impl(X1,esk3_2(X1,X2)),impl(impl(X2,esk3_2(X1,X2)),esk3_2(X1,X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_218, c_0_201]), c_0_201])).
cnf(c_0_224,plain,(impl(true,true)=true|~forallprefers(lazy_impl(true,impl(true,true)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_43]), c_0_151])])).
cnf(c_0_225,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|bool(lazy_impl(true,esk8_0))|lazy_impl(true,lazy_impl(true,lazy_impl(true,esk8_0)))!=or1(esk8_0,esk9_0)|~d(lazy_impl(true,esk8_0))), inference(spm,[status(thm)],[c_0_220, c_0_50])).
cnf(c_0_226,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|d(esk8_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_222]), c_0_104]), c_0_125])).
cnf(c_0_227,plain,(bool(esk3_2(X1,X2))|~forallprefers(lazy_impl(prop(X3),impl(impl(X1,X3),impl(impl(X2,X3),X3))),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_40]), c_0_41])).
cnf(c_0_228,plain,(impl(true,true)=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_171]), c_0_57]), c_0_151])])).
cnf(c_0_229,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|bool(esk8_0)|or1(esk8_0,esk9_0)!=lazy_impl(true,lazy_impl(true,esk8_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_50]), c_0_226])).
cnf(c_0_230,plain,(bool(esk3_2(X1,true))|~forallprefers(lazy_impl(prop(true),impl(impl(X1,true),true)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_228]), c_0_228])).
cnf(c_0_231,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|bool(esk8_0)|lazy_impl(true,lazy_impl(true,esk8_0))!=lazy_impl(true,esk8_0)), inference(spm,[status(thm)],[c_0_229, c_0_222])).
cnf(c_0_232,plain,(bool(esk3_2(X1,true))|~forallprefers(lazy_impl(true,impl(impl(X1,true),true)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_43]), c_0_151])])).
cnf(c_0_233,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|bool(esk8_0)|~d(lazy_impl(true,esk8_0))), inference(spm,[status(thm)],[c_0_231, c_0_50])).
cnf(c_0_234,plain,(bool(lazy_impl(true,impl(impl(X1,true),true)))|bool(esk3_2(X1,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_106]), c_0_68]), c_0_151])])).
cnf(c_0_235,plain,(bool(X1)|~forallprefers(lazy_impl(prop(X2),impl(lazy_impl(true,X1),impl(impl(X3,X2),X2))),lazy_impl(prop(esk3_2(X1,X3)),impl(impl(X1,esk3_2(X1,X3)),impl(impl(X3,esk3_2(X1,X3)),esk3_2(X1,X3)))))), inference(spm,[status(thm)],[c_0_223, c_0_126])).
cnf(c_0_236,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|bool(esk8_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_50]), c_0_226])).
cnf(c_0_237,plain,(d(impl(impl(X1,true),true))|bool(esk3_2(X1,true))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_104]), c_0_128])).
cnf(c_0_238,negated_conjecture,(bool(esk8_0)|~forallprefers(lazy_impl(prop(X1),impl(lazy_impl(true,esk8_0),impl(impl(esk9_0,X1),X1))),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_236]), c_0_41])).
cnf(c_0_239,plain,(bool(esk3_2(X1,true))|~forallprefers(impl(impl(X1,true),true),true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_50]), c_0_237])).
cnf(c_0_240,negated_conjecture,(bool(lazy_impl(true,esk8_0))|~forallprefers(lazy_impl(prop(X1),lazy_impl(true,lazy_impl(true,esk8_0))),true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_126]), c_0_113])).
cnf(c_0_241,plain,(bool(esk3_2(X1,true))|bool(X1)|~forallprefers(impl(lazy_impl(true,X1),true),true)), inference(spm,[status(thm)],[c_0_239, c_0_126])).
cnf(c_0_242,negated_conjecture,(d(esk8_0)|~forallprefers(lazy_impl(prop(X1),err),true)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_104]), c_0_120]), c_0_128])).
cnf(c_0_243,plain,(bool(esk3_2(err,true))|~forallprefers(impl(err,true),true)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_120]), c_0_128])).
cnf(c_0_244,negated_conjecture,(d(esk8_0)|~forallprefers(err,true)|~bool(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_43]), c_0_120])).
cnf(c_0_245,plain,(bool(esk3_2(err,true))|~forallprefers(err,true)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_126]), c_0_120]), c_0_128])).
cnf(c_0_246,negated_conjecture,(d(esk8_0)|~bool(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_106]), c_0_68]), c_0_151])]), c_0_128])).
cnf(c_0_247,plain,(bool(esk3_2(err,true))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_106]), c_0_68]), c_0_151])]), c_0_128])).
cnf(c_0_248,negated_conjecture,(d(esk8_0)), inference(spm,[status(thm)],[c_0_246, c_0_247])).
cnf(c_0_249,negated_conjecture,(bool(esk8_0)|~forallprefers(lazy_impl(prop(X1),lazy_impl(true,esk8_0)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_50]), c_0_248])])).
cnf(c_0_250,negated_conjecture,(bool(esk8_0)|~forallprefers(lazy_impl(prop(X1),esk8_0),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_50]), c_0_248])])).
fof(c_0_251, plain, (![X3]:![X4]:((bool(X3)&~bool(X4))=>impl(X3,X4)=phi(X4))), inference(fof_simplification,[status(thm)],[c_0_21])).
cnf(c_0_252,negated_conjecture,(bool(esk8_0)|~forallprefers(lazy_impl(true,esk8_0),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_250, c_0_43])).
cnf(c_0_253,negated_conjecture,(bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),impl(lazy_impl(true,esk9_0),esk3_2(esk8_0,esk9_0)))))!=or1(esk8_0,esk9_0)), inference(spm,[status(thm)],[c_0_206, c_0_126])).
fof(c_0_254, plain, (![X5]:![X6]:((~bool(X5)|bool(X6))|impl(X5,X6)=phi(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_251])])).
cnf(c_0_255,negated_conjecture,(bool(esk8_0)|~forallprefers(esk8_0,true)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_50]), c_0_248])])).
cnf(c_0_256,negated_conjecture,(d(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),impl(err,esk3_2(esk8_0,esk9_0)))))!=or1(esk8_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_104]), c_0_125])).
cnf(c_0_257,plain,(impl(X1,X2)=phi(X2)|bool(X2)|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_254])).
cnf(c_0_258,negated_conjecture,(bool(esk8_0)|~bool(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_106]), c_0_68]), c_0_151])])).
cnf(c_0_259,negated_conjecture,(d(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),err)))!=or1(esk8_0,esk9_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_126]), c_0_120]), c_0_128])).
cnf(c_0_260,plain,(lazy_impl(true,X2)=impl(X1,X2)|bool(X2)|~bool(X1)), inference(rw,[status(thm)],[c_0_257, c_0_45])).
cnf(c_0_261,negated_conjecture,(bool(esk8_0)), inference(spm,[status(thm)],[c_0_258, c_0_247])).
cnf(c_0_262,plain,(prop(X1)=false|~forallprefers(lazy_impl(true,impl(true,X1)),false)), inference(spm,[status(thm)],[c_0_154, c_0_136])).
cnf(c_0_263,negated_conjecture,(bool(lazy_impl(true,esk9_0))|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),lazy_impl(true,lazy_impl(true,esk9_0)))))!=or1(esk8_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_126]), c_0_113])).
cnf(c_0_264,negated_conjecture,(d(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),err))!=or1(esk8_0,esk9_0)|~bool(impl(esk8_0,esk3_2(esk8_0,esk9_0)))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_260]), c_0_120]), c_0_128])).
cnf(c_0_265,negated_conjecture,(esk8_0=false|esk8_0=true), inference(spm,[status(thm)],[c_0_64, c_0_261])).
cnf(c_0_266,plain,(prop(X1)=false|~forallprefers(impl(true,X1),false)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_50]), c_0_179])).
fof(c_0_267, plain, (![X3]:![X4]:((bool(X3)&~bool(X4))=>or1(X3,X4)=phi(X4))), inference(fof_simplification,[status(thm)],[c_0_22])).
cnf(c_0_268,negated_conjecture,(bool(lazy_impl(true,lazy_impl(true,esk9_0)))|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),lazy_impl(true,lazy_impl(true,lazy_impl(true,esk9_0)))))!=or1(esk8_0,esk9_0)|~bool(impl(esk8_0,esk3_2(esk8_0,esk9_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_260]), c_0_113])).
cnf(c_0_269,negated_conjecture,(esk8_0=false|d(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(true,esk9_0)),err))!=or1(true,esk9_0)|~bool(impl(true,esk3_2(true,esk9_0)))), inference(spm,[status(thm)],[c_0_264, c_0_265])).
cnf(c_0_270,plain,(prop(X1)=false|bool(impl(true,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_106]), c_0_58]), c_0_59])])).
fof(c_0_271, plain, (![X5]:![X6]:((~bool(X5)|bool(X6))|or1(X5,X6)=phi(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_267])])).
fof(c_0_272, plain, (![X9]:![X10]:![X12]:(and2(X9,X10)=phi(f1(X9,X10,esk1_2(X9,X10)))&~forallprefers(f1(X9,X10,X12),f1(X9,X10,esk1_2(X9,X10))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).
fof(c_0_273, plain, (![X8]:![X9]:![X10]:f1(X8,X9,X10)=lazy_impl(prop(X10),impl(impl(X8,impl(X9,X10)),X10))), inference(variable_rename,[status(thm)],[c_0_24])).
cnf(c_0_274,negated_conjecture,(esk8_0=false|bool(lazy_impl(true,lazy_impl(true,esk9_0)))|lazy_impl(true,lazy_impl(prop(esk3_2(true,esk9_0)),lazy_impl(true,lazy_impl(true,lazy_impl(true,esk9_0)))))!=or1(true,esk9_0)|~bool(impl(true,esk3_2(true,esk9_0)))), inference(spm,[status(thm)],[c_0_268, c_0_265])).
cnf(c_0_275,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|d(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(true,esk9_0)),err))!=or1(true,esk9_0)), inference(spm,[status(thm)],[c_0_269, c_0_270])).
cnf(c_0_276,plain,(or1(X1,X2)=phi(X2)|bool(X2)|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_277,plain,(~forallprefers(f1(X1,X2,X3),f1(X1,X2,esk1_2(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_272])).
cnf(c_0_278,plain,(f1(X1,X2,X3)=lazy_impl(prop(X3),impl(impl(X1,impl(X2,X3)),X3))), inference(split_conjunct,[status(thm)],[c_0_273])).
cnf(c_0_279,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|bool(lazy_impl(true,lazy_impl(true,esk9_0)))|lazy_impl(true,lazy_impl(prop(esk3_2(true,esk9_0)),lazy_impl(true,lazy_impl(true,lazy_impl(true,esk9_0)))))!=or1(true,esk9_0)), inference(spm,[status(thm)],[c_0_274, c_0_270])).
cnf(c_0_280,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|d(esk9_0)|or1(true,esk9_0)!=err), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_136]), c_0_120]), c_0_120]), c_0_153])).
cnf(c_0_281,plain,(lazy_impl(true,X2)=or1(X1,X2)|bool(X2)|~bool(X1)), inference(rw,[status(thm)],[c_0_276, c_0_45])).
cnf(c_0_282,plain,(~forallprefers(lazy_impl(prop(X3),impl(impl(X1,impl(X2,X3)),X3)),lazy_impl(prop(esk1_2(X1,X2)),impl(impl(X1,impl(X2,esk1_2(X1,X2))),esk1_2(X1,X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_277, c_0_278]), c_0_278])).
cnf(c_0_283,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|bool(lazy_impl(true,lazy_impl(true,esk9_0)))|lazy_impl(true,lazy_impl(true,lazy_impl(true,lazy_impl(true,lazy_impl(true,esk9_0)))))!=or1(true,esk9_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_136]), c_0_153])).
cnf(c_0_284,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|d(esk9_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281]), c_0_151])]), c_0_104]), c_0_125])).
cnf(c_0_285,plain,(bool(esk1_2(X1,X2))|~forallprefers(lazy_impl(prop(X3),impl(impl(X1,impl(X2,X3)),X3)),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_40]), c_0_41])).
cnf(c_0_286,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|bool(lazy_impl(true,esk9_0))|lazy_impl(true,lazy_impl(true,lazy_impl(true,lazy_impl(true,esk9_0))))!=or1(true,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_50]), c_0_284])).
cnf(c_0_287,plain,(bool(esk1_2(false,X1))|~forallprefers(lazy_impl(prop(X2),impl(true,X2)),true)|~bool(impl(X1,X2))), inference(spm,[status(thm)],[c_0_285, c_0_190])).
cnf(c_0_288,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|bool(esk9_0)|lazy_impl(true,lazy_impl(true,lazy_impl(true,esk9_0)))!=or1(true,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_50]), c_0_284])).
cnf(c_0_289,plain,(bool(esk1_2(false,true))|~forallprefers(lazy_impl(prop(X1),impl(true,X1)),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_287, c_0_99])).
cnf(c_0_290,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|bool(esk9_0)|or1(true,esk9_0)!=lazy_impl(true,lazy_impl(true,esk9_0))|~d(lazy_impl(true,lazy_impl(true,esk9_0)))), inference(spm,[status(thm)],[c_0_288, c_0_50])).
cnf(c_0_291,plain,(bool(esk1_2(false,true))|~forallprefers(lazy_impl(prop(X1),X1),true)|~bool(X1)), inference(spm,[status(thm)],[c_0_289, c_0_99])).
cnf(c_0_292,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|bool(esk9_0)|lazy_impl(true,lazy_impl(true,esk9_0))!=lazy_impl(true,esk9_0)|~d(lazy_impl(true,lazy_impl(true,esk9_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_281]), c_0_151])])).
cnf(c_0_293,plain,(~forallprefers(lazy_impl(prop(X1),impl(impl(false,impl(X2,X1)),X1)),lazy_impl(prop(esk1_2(false,X2)),impl(true,esk1_2(false,X2))))|~bool(impl(X2,esk1_2(false,X2)))), inference(spm,[status(thm)],[c_0_282, c_0_190])).
cnf(c_0_294,plain,(bool(esk1_2(false,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_109]), c_0_57]), c_0_59])])).
cnf(c_0_295,plain,(~forallprefers(lazy_impl(true,impl(impl(X1,X2),impl(impl(X3,X2),X2))),lazy_impl(prop(esk3_2(X1,X3)),impl(impl(X1,esk3_2(X1,X3)),impl(impl(X3,esk3_2(X1,X3)),esk3_2(X1,X3)))))|~bool(X2)), inference(spm,[status(thm)],[c_0_223, c_0_43])).
cnf(c_0_296,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|bool(esk9_0)|~d(lazy_impl(true,esk9_0))), inference(spm,[status(thm)],[c_0_292, c_0_50])).
cnf(c_0_297,plain,(~forallprefers(lazy_impl(prop(X1),impl(impl(false,impl(true,X1)),X1)),lazy_impl(prop(esk1_2(false,true)),impl(true,esk1_2(false,true))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_99]), c_0_294])])).
cnf(c_0_298,plain,(~forallprefers(lazy_impl(true,impl(false,impl(impl(X1,false),false))),lazy_impl(prop(esk3_2(true,X1)),impl(impl(true,esk3_2(true,X1)),impl(impl(X1,esk3_2(true,X1)),esk3_2(true,X1)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_208]), c_0_59])])).
cnf(c_0_299,negated_conjecture,(prop(esk3_2(true,esk9_0))=false|esk8_0=false|bool(esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_50]), c_0_284])).
cnf(c_0_300,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(prop(esk1_2(false,true)),impl(true,esk1_2(false,true))))|~bool(impl(true,X1))), inference(spm,[status(thm)],[c_0_297, c_0_190])).
cnf(c_0_301,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(true,impl(false,impl(impl(esk9_0,false),false))),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299]), c_0_41])).
cnf(c_0_302,plain,(~forallprefers(false,lazy_impl(prop(esk1_2(false,true)),impl(true,esk1_2(false,true))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_208]), c_0_109]), c_0_59])])).
cnf(c_0_303,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(true,impl(false,impl(lazy_impl(true,esk9_0),false))),true)), inference(spm,[status(thm)],[c_0_301, c_0_126])).
cnf(c_0_304,plain,(~forallprefers(false,lazy_impl(true,impl(true,esk1_2(false,true))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_43]), c_0_294])])).
cnf(c_0_305,plain,(esk1_2(false,true)=false|esk1_2(false,true)=true), inference(spm,[status(thm)],[c_0_64, c_0_294])).
cnf(c_0_306,negated_conjecture,(esk8_0=false|d(esk9_0)|~forallprefers(lazy_impl(true,impl(false,impl(err,false))),true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_104]), c_0_125])).
cnf(c_0_307,plain,(esk1_2(false,true)=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_305]), c_0_228]), c_0_174]), c_0_57])])).
cnf(c_0_308,negated_conjecture,(esk8_0=false|d(esk9_0)|~forallprefers(lazy_impl(true,impl(false,err)),true)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_126]), c_0_120]), c_0_128])).
cnf(c_0_309,plain,(~forallprefers(lazy_impl(prop(X1),impl(impl(false,impl(true,X1)),X1)),false)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_297, c_0_307]), c_0_307]), c_0_208]), c_0_109])).
cnf(c_0_310,negated_conjecture,(esk8_0=false|d(esk9_0)|~forallprefers(err,true)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_260]), c_0_120]), c_0_120]), c_0_59])]), c_0_128])).
cnf(c_0_311,plain,(~forallprefers(lazy_impl(prop(true),impl(impl(false,true),true)),false)), inference(spm,[status(thm)],[c_0_309, c_0_228])).
cnf(c_0_312,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(true,impl(false,impl(esk9_0,false))),true)|~d(esk9_0)), inference(spm,[status(thm)],[c_0_303, c_0_50])).
cnf(c_0_313,negated_conjecture,(esk8_0=false|d(esk9_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_106]), c_0_68]), c_0_151])]), c_0_128])).
cnf(c_0_314,plain,(~forallprefers(lazy_impl(true,impl(impl(false,true),true)),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_43]), c_0_151])])).
cnf(c_0_315,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(true,impl(false,impl(esk9_0,false))),true)), inference(spm,[status(thm)],[c_0_312, c_0_313])).
cnf(c_0_316,plain,(bool(lazy_impl(true,impl(impl(false,true),true)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_317,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(true,impl(false,lazy_impl(true,esk9_0))),true)), inference(spm,[status(thm)],[c_0_315, c_0_126])).
cnf(c_0_318,plain,(d(impl(impl(false,true),true))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_104]), c_0_128])).
cnf(c_0_319,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(true,impl(false,esk9_0)),true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_50]), c_0_313])).
cnf(c_0_320,plain,(~forallprefers(impl(impl(false,true),true),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_50]), c_0_318])])).
cnf(c_0_321,negated_conjecture,(lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),impl(impl(esk9_0,esk3_2(esk8_0,esk9_0)),esk3_2(esk8_0,esk9_0))))!=or1(esk8_0,esk9_0)|~d(lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),impl(impl(esk9_0,esk3_2(esk8_0,esk9_0)),esk3_2(esk8_0,esk9_0)))))), inference(spm,[status(thm)],[c_0_206, c_0_50])).
cnf(c_0_322,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(true,lazy_impl(true,esk9_0)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_260]), c_0_59])])).
cnf(c_0_323,plain,(~forallprefers(lazy_impl(true,impl(false,true)),false)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_126]), c_0_103])).
cnf(c_0_324,plain,(~forallprefers(lazy_impl(prop(false),impl(impl(false,false),false)),false)), inference(spm,[status(thm)],[c_0_309, c_0_208])).
cnf(c_0_325,negated_conjecture,(bool(esk3_2(esk8_0,esk9_0))|or1(esk8_0,esk9_0)!=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_40]), c_0_41]), c_0_41]), c_0_68])])).
fof(c_0_326, plain, (![X5]:(~bool(X5)|or1(true,X5)=true)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])).
cnf(c_0_327,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(lazy_impl(true,esk9_0),true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_322, c_0_50]), c_0_313])).
cnf(c_0_328,plain,(impl(true,X1)=false|impl(true,X1)=true|prop(X1)=false), inference(spm,[status(thm)],[c_0_64, c_0_270])).
cnf(c_0_329,plain,(bool(esk3_2(X1,true))|~forallprefers(lazy_impl(prop(X2),impl(impl(X1,X2),impl(X2,X2))),true)|~bool(X2)), inference(spm,[status(thm)],[c_0_227, c_0_99])).
cnf(c_0_330,plain,(bool(lazy_impl(true,impl(false,true)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_331,plain,(~forallprefers(lazy_impl(true,impl(impl(false,false),false)),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_43]), c_0_59])])).
cnf(c_0_332,negated_conjecture,(esk8_0=false|bool(esk3_2(true,esk9_0))|or1(true,esk9_0)!=true), inference(spm,[status(thm)],[c_0_325, c_0_265])).
cnf(c_0_333,plain,(or1(true,X1)=true|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_326])).
cnf(c_0_334,negated_conjecture,(esk8_0=false|bool(esk9_0)|~forallprefers(esk9_0,true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_50]), c_0_313])).
cnf(c_0_335,negated_conjecture,(impl(true,esk3_2(esk8_0,esk9_0))=true|impl(true,esk3_2(esk8_0,esk9_0))=false|or1(esk8_0,esk9_0)!=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_328]), c_0_41]), c_0_41]), c_0_68])])).
cnf(c_0_336,plain,(~forallprefers(lazy_impl(prop(X1),impl(impl(X2,X1),impl(X1,X1))),lazy_impl(prop(esk3_2(X2,true)),impl(impl(X2,esk3_2(X2,true)),impl(impl(true,esk3_2(X2,true)),esk3_2(X2,true)))))|~bool(X1)), inference(spm,[status(thm)],[c_0_223, c_0_99])).
cnf(c_0_337,plain,(bool(esk3_2(X1,true))|~forallprefers(lazy_impl(prop(false),impl(impl(X1,false),true)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_190]), c_0_59])])).
cnf(c_0_338,plain,(d(impl(false,true))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_104]), c_0_128])).
cnf(c_0_339,plain,(d(impl(impl(false,false),false))|~forallprefers(err,false)), inference(spm,[status(thm)],[c_0_331, c_0_104])).
cnf(c_0_340,negated_conjecture,(lazy_impl(true,lazy_impl(true,impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),impl(impl(esk9_0,esk3_2(esk8_0,esk9_0)),esk3_2(esk8_0,esk9_0)))))!=or1(esk8_0,esk9_0)|~bool(esk3_2(esk8_0,esk9_0))), inference(spm,[status(thm)],[c_0_206, c_0_43])).
cnf(c_0_341,negated_conjecture,(esk8_0=false|bool(esk3_2(true,esk9_0))|~bool(esk9_0)), inference(spm,[status(thm)],[c_0_332, c_0_333])).
cnf(c_0_342,negated_conjecture,(esk8_0=false|bool(esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_106]), c_0_68]), c_0_151])])).
cnf(c_0_343,negated_conjecture,(impl(true,esk3_2(true,esk9_0))=false|impl(true,esk3_2(true,esk9_0))=true|esk8_0=false|or1(true,esk9_0)!=true), inference(spm,[status(thm)],[c_0_335, c_0_265])).
cnf(c_0_344,plain,(~forallprefers(lazy_impl(prop(false),impl(impl(X1,false),true)),lazy_impl(prop(esk3_2(X1,true)),impl(impl(X1,esk3_2(X1,true)),impl(impl(true,esk3_2(X1,true)),esk3_2(X1,true)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_190]), c_0_59])])).
cnf(c_0_345,plain,(bool(esk3_2(true,true))|~forallprefers(lazy_impl(prop(false),impl(false,true)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_99]), c_0_59])])).
cnf(c_0_346,plain,(~forallprefers(impl(false,true),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_50]), c_0_338])])).
cnf(c_0_347,plain,(~forallprefers(impl(impl(false,false),false),false)|~d(impl(impl(false,false),false))), inference(spm,[status(thm)],[c_0_331, c_0_50])).
cnf(c_0_348,plain,(d(impl(impl(false,false),false))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_106]), c_0_58]), c_0_59])]), c_0_128])).
cnf(c_0_349,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(true,impl(impl(true,esk3_2(true,esk9_0)),impl(impl(esk9_0,esk3_2(true,esk9_0)),esk3_2(true,esk9_0)))))!=or1(true,esk9_0)|~bool(esk3_2(true,esk9_0))), inference(spm,[status(thm)],[c_0_340, c_0_265])).
cnf(c_0_350,negated_conjecture,(esk8_0=false|bool(esk3_2(true,esk9_0))), inference(spm,[status(thm)],[c_0_341, c_0_342])).
cnf(c_0_351,negated_conjecture,(impl(true,esk3_2(true,esk9_0))=true|impl(true,esk3_2(true,esk9_0))=false|esk8_0=false|~bool(esk9_0)), inference(spm,[status(thm)],[c_0_343, c_0_333])).
cnf(c_0_352,plain,(~forallprefers(lazy_impl(prop(false),impl(false,true)),lazy_impl(prop(esk3_2(true,true)),impl(impl(true,esk3_2(true,true)),impl(impl(true,esk3_2(true,true)),esk3_2(true,true)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_99]), c_0_59])])).
cnf(c_0_353,plain,(bool(esk3_2(true,true))|~forallprefers(lazy_impl(true,impl(false,true)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_43]), c_0_59])])).
cnf(c_0_354,plain,(lazy_impl(true,X1)=false|true=X1|~bool(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_121]), c_0_118])).
cnf(c_0_355,plain,(bool(impl(false,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_356,plain,(~forallprefers(impl(impl(false,false),false),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_347, c_0_348])])).
cnf(c_0_357,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(true,impl(impl(true,esk3_2(true,esk9_0)),impl(impl(esk9_0,esk3_2(true,esk9_0)),esk3_2(true,esk9_0)))))!=or1(true,esk9_0)), inference(spm,[status(thm)],[c_0_349, c_0_350])).
cnf(c_0_358,negated_conjecture,(esk8_0=false|esk9_0=false|esk9_0=true), inference(spm,[status(thm)],[c_0_64, c_0_342])).
cnf(c_0_359,negated_conjecture,(impl(true,esk3_2(true,esk9_0))=false|impl(true,esk3_2(true,esk9_0))=true|esk8_0=false), inference(spm,[status(thm)],[c_0_351, c_0_342])).
cnf(c_0_360,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(prop(esk3_2(true,esk9_0)),impl(impl(true,esk3_2(true,esk9_0)),impl(impl(esk9_0,esk3_2(true,esk9_0)),esk3_2(true,esk9_0)))))!=or1(true,esk9_0)), inference(spm,[status(thm)],[c_0_206, c_0_265])).
cnf(c_0_361,plain,(~forallprefers(lazy_impl(true,impl(false,true)),lazy_impl(prop(esk3_2(true,true)),impl(impl(true,esk3_2(true,true)),impl(impl(true,esk3_2(true,true)),esk3_2(true,true)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_43]), c_0_59])])).
cnf(c_0_362,plain,(impl(false,true)=true|bool(esk3_2(true,true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_354]), c_0_57]), c_0_355])])).
cnf(c_0_363,plain,(~forallprefers(lazy_impl(true,impl(false,false)),false)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_126]), c_0_103])).
cnf(c_0_364,negated_conjecture,(esk9_0=false|esk8_0=false|lazy_impl(true,lazy_impl(true,impl(impl(true,esk3_2(true,true)),impl(impl(true,esk3_2(true,true)),esk3_2(true,true)))))!=or1(true,true)), inference(spm,[status(thm)],[c_0_357, c_0_358])).
cnf(c_0_365,negated_conjecture,(impl(true,esk3_2(true,true))=true|impl(true,esk3_2(true,true))=false|esk9_0=false|esk8_0=false), inference(spm,[status(thm)],[c_0_359, c_0_358])).
cnf(c_0_366,negated_conjecture,(esk8_0=false|bool(esk3_2(true,esk9_0))|lazy_impl(true,lazy_impl(prop(esk3_2(true,esk9_0)),impl(lazy_impl(true,esk3_2(true,esk9_0)),impl(impl(esk9_0,esk3_2(true,esk9_0)),esk3_2(true,esk9_0)))))!=or1(true,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_360, c_0_260]), c_0_151])])).
cnf(c_0_367,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(prop(esk3_2(true,esk9_0)),impl(esk3_2(true,esk9_0),impl(impl(esk9_0,esk3_2(true,esk9_0)),esk3_2(true,esk9_0)))))!=or1(true,esk9_0)|~bool(esk3_2(true,esk9_0))), inference(spm,[status(thm)],[c_0_360, c_0_99])).
cnf(c_0_368,plain,(~forallprefers(lazy_impl(true,impl(false,true)),lazy_impl(true,impl(impl(true,esk3_2(true,true)),impl(impl(true,esk3_2(true,true)),esk3_2(true,true)))))|~bool(esk3_2(true,true))), inference(spm,[status(thm)],[c_0_361, c_0_43])).
cnf(c_0_369,plain,(impl(false,true)=true|esk3_2(true,true)=false|esk3_2(true,true)=true), inference(spm,[status(thm)],[c_0_64, c_0_362])).
cnf(c_0_370,plain,(bool(lazy_impl(true,impl(false,false)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_371,negated_conjecture,(impl(true,esk3_2(true,true))=false|esk8_0=false|esk9_0=false|lazy_impl(true,lazy_impl(true,impl(true,impl(true,esk3_2(true,true)))))!=or1(true,true)), inference(spm,[status(thm)],[c_0_364, c_0_365])).
cnf(c_0_372,negated_conjecture,(lazy_impl(true,esk3_2(true,esk9_0))=err|esk8_0=false|lazy_impl(true,lazy_impl(prop(esk3_2(true,esk9_0)),impl(esk3_2(true,esk9_0),impl(impl(esk9_0,esk3_2(true,esk9_0)),esk3_2(true,esk9_0)))))!=or1(true,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_100]), c_0_367])).
cnf(c_0_373,plain,(esk3_2(true,true)=false|impl(false,true)=true|~forallprefers(lazy_impl(true,impl(false,true)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_369]), c_0_228]), c_0_228]), c_0_228]), c_0_228]), c_0_174]), c_0_151])])).
cnf(c_0_374,plain,(d(impl(false,false))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_370, c_0_104]), c_0_128])).
cnf(c_0_375,plain,(~forallprefers(lazy_impl(prop(X1),impl(impl(X2,X1),X1)),lazy_impl(prop(esk1_2(X2,true)),impl(impl(X2,impl(true,esk1_2(X2,true))),esk1_2(X2,true))))|~bool(X1)), inference(spm,[status(thm)],[c_0_282, c_0_99])).
cnf(c_0_376,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(true,impl(esk3_2(true,esk9_0),impl(impl(esk9_0,esk3_2(true,esk9_0)),esk3_2(true,esk9_0)))))!=or1(true,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_99]), c_0_350])).
cnf(c_0_377,negated_conjecture,(impl(true,esk3_2(true,true))=false|esk9_0=false|esk8_0=false|or1(true,true)!=true), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_365]), c_0_228]), c_0_174]), c_0_174]), c_0_153])).
cnf(c_0_378,negated_conjecture,(lazy_impl(true,esk3_2(true,esk9_0))=err|esk8_0=false|lazy_impl(true,lazy_impl(true,impl(esk3_2(true,esk9_0),impl(impl(esk9_0,esk3_2(true,esk9_0)),esk3_2(true,esk9_0)))))!=or1(true,esk9_0)|~bool(esk3_2(true,esk9_0))), inference(spm,[status(thm)],[c_0_372, c_0_43])).
cnf(c_0_379,negated_conjecture,(esk3_2(true,esk9_0)=false|esk3_2(true,esk9_0)=true|esk8_0=false), inference(spm,[status(thm)],[c_0_64, c_0_350])).
cnf(c_0_380,plain,(impl(false,true)=true|~forallprefers(false,lazy_impl(prop(esk3_2(true,true)),impl(impl(true,esk3_2(true,true)),impl(impl(true,esk3_2(true,true)),esk3_2(true,true)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_354]), c_0_355])])).
cnf(c_0_381,plain,(impl(false,true)=true|esk3_2(true,true)=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_354]), c_0_57]), c_0_355])])).
cnf(c_0_382,plain,(~forallprefers(impl(false,false),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_50]), c_0_374])])).
cnf(c_0_383,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(prop(false),impl(impl(false,false),false)))|~bool(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_375, c_0_190]), c_0_307]), c_0_307]), c_0_208]), c_0_307])).
cnf(c_0_384,negated_conjecture,(esk9_0=false|esk8_0=false|lazy_impl(true,lazy_impl(true,impl(esk3_2(true,true),impl(impl(true,esk3_2(true,true)),esk3_2(true,true)))))!=or1(true,true)), inference(spm,[status(thm)],[c_0_376, c_0_358])).
cnf(c_0_385,negated_conjecture,(impl(true,esk3_2(true,true))=false|esk8_0=false|esk9_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_333]), c_0_151])])).
cnf(c_0_386,negated_conjecture,(esk3_2(true,esk9_0)=false|esk8_0=false|lazy_impl(true,lazy_impl(true,impl(true,impl(impl(esk9_0,true),true))))!=or1(true,esk9_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_379]), c_0_174]), c_0_151])]), c_0_118])).
cnf(c_0_387,plain,(impl(false,true)=true|~forallprefers(false,lazy_impl(prop(false),impl(false,impl(false,false))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_381]), c_0_208]), c_0_208])).
cnf(c_0_388,plain,(bool(impl(false,false))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_389,plain,(~forallprefers(false,lazy_impl(prop(false),impl(impl(false,false),false)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_208]), c_0_109]), c_0_59])])).
cnf(c_0_390,negated_conjecture,(esk8_0=false|esk9_0=false|lazy_impl(true,lazy_impl(true,impl(esk3_2(true,true),impl(false,esk3_2(true,true)))))!=or1(true,true)), inference(spm,[status(thm)],[c_0_384, c_0_385])).
cnf(c_0_391,negated_conjecture,(esk9_0=false|esk8_0=false|bool(esk3_2(true,true))), inference(spm,[status(thm)],[c_0_350, c_0_358])).
cnf(c_0_392,negated_conjecture,(esk3_2(true,true)=false|esk9_0=false|esk8_0=false|or1(true,true)!=true), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_358]), c_0_228]), c_0_228]), c_0_228]), c_0_174]), c_0_174])).
cnf(c_0_393,plain,(impl(false,true)=true|~forallprefers(false,lazy_impl(prop(false),true))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_190]), c_0_388])])).
cnf(c_0_394,plain,(~forallprefers(false,lazy_impl(true,impl(impl(false,false),false)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_43]), c_0_59])])).
cnf(c_0_395,plain,(bool(impl(impl(false,false),false))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_396,negated_conjecture,(esk9_0=false|esk8_0=false|lazy_impl(true,lazy_impl(true,impl(esk3_2(true,true),true)))!=or1(true,true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_190]), c_0_391])).
cnf(c_0_397,negated_conjecture,(esk3_2(true,true)=false|esk8_0=false|esk9_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_333]), c_0_151])])).
cnf(c_0_398,plain,(impl(false,true)=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_43]), c_0_174]), c_0_57]), c_0_59])])).
cnf(c_0_399,plain,(~forallprefers(lazy_impl(prop(X1),impl(impl(false,X1),X1)),false)|~bool(X1)), inference(spm,[status(thm)],[c_0_309, c_0_99])).
cnf(c_0_400,plain,(bool(impl(X1,impl(X2,X3)))|~forallprefers(lazy_impl(prop(X3),lazy_impl(true,impl(X1,impl(X2,X3)))),lazy_impl(prop(esk1_2(X1,X2)),impl(impl(X1,impl(X2,esk1_2(X1,X2))),esk1_2(X1,X2))))), inference(spm,[status(thm)],[c_0_282, c_0_126])).
cnf(c_0_401,plain,(lazy_impl(true,impl(impl(false,false),false))=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_394, c_0_121]), c_0_57]), c_0_395])])).
cnf(c_0_402,negated_conjecture,(esk8_0=false|esk9_0=false|or1(true,true)!=true), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_397]), c_0_398]), c_0_174]), c_0_174])).
cnf(c_0_403,plain,(~forallprefers(lazy_impl(true,impl(impl(false,X1),X1)),false)|~bool(X1)), inference(spm,[status(thm)],[c_0_399, c_0_43])).
cnf(c_0_404,plain,(d(impl(X1,impl(X2,X3)))|~forallprefers(lazy_impl(prop(X3),err),lazy_impl(prop(esk1_2(X1,X2)),impl(impl(X1,impl(X2,esk1_2(X1,X2))),esk1_2(X1,X2))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_104]), c_0_125])).
cnf(c_0_405,plain,(impl(impl(false,false),false)=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_401]), c_0_348])])).
cnf(c_0_406,negated_conjecture,(esk9_0=false|esk8_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_333]), c_0_151])])).
cnf(c_0_407,plain,(prop(X1)=false|~forallprefers(lazy_impl(true,impl(impl(false,X1),X1)),false)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_136]), c_0_40])).
cnf(c_0_408,plain,(d(impl(impl(false,X1),X1))|~forallprefers(err,false)|~bool(X1)), inference(spm,[status(thm)],[c_0_403, c_0_104])).
cnf(c_0_409,plain,(d(impl(false,impl(true,X1)))|~forallprefers(lazy_impl(prop(X1),err),false)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_307]), c_0_208]), c_0_405]), c_0_109])).
cnf(c_0_410,negated_conjecture,(esk3_2(true,false)=false|esk8_0=false|or1(true,false)!=true), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_406]), c_0_398]), c_0_228]), c_0_228]), c_0_174]), c_0_174])).
cnf(c_0_411,plain,(prop(X1)=false|~forallprefers(impl(impl(false,X1),X1),false)|~d(impl(impl(false,X1),X1))), inference(spm,[status(thm)],[c_0_407, c_0_50])).
cnf(c_0_412,plain,(d(impl(impl(false,X1),X1))|~bool(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_106]), c_0_58]), c_0_59])]), c_0_128])).
cnf(c_0_413,plain,(prop(X1)=false|d(impl(false,impl(true,X1)))|~forallprefers(err,false)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_136]), c_0_120])).
cnf(c_0_414,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(true,impl(esk3_2(true,false),impl(impl(false,esk3_2(true,false)),esk3_2(true,false)))))!=or1(true,false)), inference(spm,[status(thm)],[c_0_376, c_0_406])).
cnf(c_0_415,negated_conjecture,(esk3_2(true,false)=false|esk8_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_333]), c_0_59])])).
cnf(c_0_416,plain,(prop(X1)=false|~forallprefers(impl(impl(false,X1),X1),false)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_411, c_0_412]), c_0_40])).
cnf(c_0_417,plain,(prop(X1)=false|d(impl(false,impl(true,X1)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_106]), c_0_58]), c_0_59])]), c_0_128])).
cnf(c_0_418,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|d(esk9_0)|lazy_impl(true,lazy_impl(true,impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),err)))!=or1(esk8_0,esk9_0)), inference(spm,[status(thm)],[c_0_259, c_0_136])).
cnf(c_0_419,negated_conjecture,(esk8_0=false|lazy_impl(true,lazy_impl(true,impl(false,false)))!=or1(true,false)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_415]), c_0_405])).
cnf(c_0_420,plain,(prop(X1)=false|~forallprefers(lazy_impl(true,impl(false,X1)),false)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_126]), c_0_103])).
cnf(c_0_421,plain,(prop(X1)=false|d(impl(false,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_99]), c_0_40])).
cnf(c_0_422,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|d(esk9_0)|or1(esk8_0,esk9_0)!=err|~bool(impl(esk8_0,esk3_2(esk8_0,esk9_0)))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_260]), c_0_120]), c_0_120]), c_0_120]), c_0_128])).
cnf(c_0_423,negated_conjecture,(esk8_0=false|or1(true,false)!=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_190]), c_0_174]), c_0_174]), c_0_59])])).
cnf(c_0_424,plain,(prop(X1)=false|~forallprefers(impl(false,X1),false)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_50]), c_0_421])).
cnf(c_0_425,plain,(bool(X1)|~forallprefers(lazy_impl(prop(X2),impl(impl(X3,X2),impl(lazy_impl(true,X1),X2))),lazy_impl(prop(esk3_2(X3,X1)),impl(impl(X3,esk3_2(X3,X1)),impl(impl(X1,esk3_2(X3,X1)),esk3_2(X3,X1)))))), inference(spm,[status(thm)],[c_0_223, c_0_126])).
cnf(c_0_426,negated_conjecture,(prop(esk3_2(esk8_0,esk9_0))=false|d(esk9_0)|~bool(impl(esk8_0,esk3_2(esk8_0,esk9_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_281]), c_0_261])]), c_0_104]), c_0_125])).
cnf(c_0_427,negated_conjecture,(esk8_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_423, c_0_333]), c_0_59])])).
cnf(c_0_428,plain,(prop(X1)=false|bool(impl(false,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_424, c_0_106]), c_0_58]), c_0_59])])).
cnf(c_0_429,plain,(d(X1)|~forallprefers(lazy_impl(prop(X2),impl(impl(X3,X2),impl(err,X2))),lazy_impl(prop(esk3_2(X3,X1)),impl(impl(X3,esk3_2(X3,X1)),impl(impl(X1,esk3_2(X3,X1)),esk3_2(X3,X1)))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_425, c_0_104]), c_0_125])).
cnf(c_0_430,negated_conjecture,(prop(esk3_2(false,esk9_0))=false|d(esk9_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_426, c_0_427]), c_0_427]), c_0_427]), c_0_428])).
cnf(c_0_431,negated_conjecture,(d(esk9_0)|~forallprefers(lazy_impl(prop(X1),impl(impl(false,X1),impl(err,X1))),true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_430]), c_0_41])).
cnf(c_0_432,negated_conjecture,(d(esk9_0)|~forallprefers(lazy_impl(prop(true),impl(true,impl(err,true))),true)), inference(spm,[status(thm)],[c_0_431, c_0_398])).
cnf(c_0_433,negated_conjecture,(d(esk9_0)|~forallprefers(lazy_impl(prop(true),impl(true,err)),true)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_432, c_0_126]), c_0_120]), c_0_128])).
cnf(c_0_434,negated_conjecture,(bool(lazy_impl(true,esk9_0))|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),lazy_impl(true,esk9_0))))!=or1(esk8_0,esk9_0)|~d(lazy_impl(true,esk9_0))), inference(spm,[status(thm)],[c_0_263, c_0_50])).
cnf(c_0_435,negated_conjecture,(d(esk9_0)|~forallprefers(lazy_impl(prop(true),err),true)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_260]), c_0_120]), c_0_151])]), c_0_128])).
cnf(c_0_436,negated_conjecture,(bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(esk8_0,esk9_0)),impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),esk9_0)))!=or1(esk8_0,esk9_0)|~d(esk9_0)), inference(spm,[status(thm)],[c_0_434, c_0_50])).
cnf(c_0_437,negated_conjecture,(d(esk9_0)|~forallprefers(err,true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_43]), c_0_120]), c_0_151])])).
cnf(c_0_438,negated_conjecture,(bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(false,esk9_0)),impl(impl(false,esk3_2(false,esk9_0)),esk9_0)))!=or1(false,esk9_0)|~d(esk9_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_436, c_0_427]), c_0_427]), c_0_427]), c_0_427])).
cnf(c_0_439,negated_conjecture,(d(esk9_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_437, c_0_106]), c_0_68]), c_0_151])]), c_0_128])).
cnf(c_0_440,negated_conjecture,(bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(false,esk9_0)),impl(impl(false,esk3_2(false,esk9_0)),esk9_0)))!=or1(false,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_438, c_0_439])])).
cnf(c_0_441,negated_conjecture,(bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(false,esk9_0)),lazy_impl(true,esk9_0)))!=or1(false,esk9_0)|~bool(impl(false,esk3_2(false,esk9_0)))), inference(spm,[status(thm)],[c_0_440, c_0_260])).
cnf(c_0_442,negated_conjecture,(prop(esk3_2(false,esk9_0))=false|bool(esk9_0)|lazy_impl(true,lazy_impl(prop(esk3_2(false,esk9_0)),lazy_impl(true,esk9_0)))!=or1(false,esk9_0)), inference(spm,[status(thm)],[c_0_441, c_0_428])).
cnf(c_0_443,negated_conjecture,(prop(esk3_2(false,esk9_0))=false|bool(esk9_0)|lazy_impl(true,lazy_impl(true,lazy_impl(true,esk9_0)))!=or1(false,esk9_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_136]), c_0_153])).
cnf(c_0_444,negated_conjecture,(prop(esk3_2(false,esk9_0))=false|bool(esk9_0)|or1(false,esk9_0)!=lazy_impl(true,lazy_impl(true,esk9_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_443, c_0_50]), c_0_439])])).
cnf(c_0_445,negated_conjecture,(prop(esk3_2(false,esk9_0))=false|bool(esk9_0)|lazy_impl(true,lazy_impl(true,esk9_0))!=lazy_impl(true,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_281]), c_0_59])])).
cnf(c_0_446,negated_conjecture,(prop(esk3_2(false,esk9_0))=false|bool(esk9_0)|lazy_impl(true,esk9_0)!=esk9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_445, c_0_50]), c_0_439])])).
cnf(c_0_447,plain,(bool(X1)|~forallprefers(lazy_impl(prop(X2),impl(impl(X3,X2),impl(X1,X2))),lazy_impl(prop(esk3_2(X3,X1)),impl(impl(X3,esk3_2(X3,X1)),impl(impl(X1,esk3_2(X3,X1)),esk3_2(X3,X1)))))|~d(X1)), inference(spm,[status(thm)],[c_0_425, c_0_50])).
cnf(c_0_448,negated_conjecture,(prop(esk3_2(false,esk9_0))=false|bool(esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_446, c_0_50]), c_0_439])])).
cnf(c_0_449,negated_conjecture,(bool(esk9_0)|~forallprefers(lazy_impl(prop(X1),impl(impl(false,X1),impl(esk9_0,X1))),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_447, c_0_448]), c_0_41]), c_0_439])])).
cnf(c_0_450,negated_conjecture,(bool(esk9_0)|~forallprefers(lazy_impl(prop(true),impl(true,impl(esk9_0,true))),true)), inference(spm,[status(thm)],[c_0_449, c_0_398])).
cnf(c_0_451,negated_conjecture,(bool(esk9_0)|~forallprefers(lazy_impl(true,impl(true,impl(esk9_0,true))),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_43]), c_0_151])])).
cnf(c_0_452,negated_conjecture,(bool(esk9_0)|~forallprefers(lazy_impl(true,impl(true,lazy_impl(true,esk9_0))),true)), inference(spm,[status(thm)],[c_0_451, c_0_126])).
cnf(c_0_453,negated_conjecture,(bool(esk9_0)|~forallprefers(lazy_impl(true,impl(true,esk9_0)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_452, c_0_50]), c_0_439])])).
cnf(c_0_454,negated_conjecture,(bool(esk9_0)|~forallprefers(lazy_impl(true,lazy_impl(true,esk9_0)),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_453, c_0_260]), c_0_151])])).
cnf(c_0_455,negated_conjecture,(bool(esk9_0)|~forallprefers(lazy_impl(true,esk9_0),true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_454, c_0_50]), c_0_439])])).
fof(c_0_456, plain, (![X5]:(~bool(X5)|or1(false,X5)=X5)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])).
cnf(c_0_457,negated_conjecture,(bool(esk9_0)|~forallprefers(esk9_0,true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_455, c_0_50]), c_0_439])])).
cnf(c_0_458,negated_conjecture,(bool(esk3_2(false,esk9_0))|or1(false,esk9_0)!=true), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_325, c_0_427]), c_0_427])).
cnf(c_0_459,plain,(or1(false,X1)=X1|~bool(X1)), inference(split_conjunct,[status(thm)],[c_0_456])).
cnf(c_0_460,negated_conjecture,(bool(esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_457, c_0_106]), c_0_68]), c_0_151])])).
cnf(c_0_461,negated_conjecture,(lazy_impl(true,lazy_impl(prop(esk3_2(false,esk9_0)),impl(impl(false,esk3_2(false,esk9_0)),impl(impl(esk9_0,esk3_2(false,esk9_0)),esk3_2(false,esk9_0)))))!=or1(false,esk9_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_206, c_0_427]), c_0_427]), c_0_427]), c_0_427]), c_0_427]), c_0_427])).
cnf(c_0_462,negated_conjecture,(bool(esk3_2(false,esk9_0))|esk9_0!=true), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_458, c_0_459]), c_0_146])).
cnf(c_0_463,negated_conjecture,(esk9_0=false|esk9_0=true), inference(spm,[status(thm)],[c_0_64, c_0_460])).
cnf(c_0_464,negated_conjecture,(lazy_impl(true,lazy_impl(prop(esk3_2(false,esk9_0)),impl(true,impl(impl(esk9_0,esk3_2(false,esk9_0)),esk3_2(false,esk9_0)))))!=or1(false,esk9_0)|~bool(esk3_2(false,esk9_0))), inference(spm,[status(thm)],[c_0_461, c_0_190])).
cnf(c_0_465,negated_conjecture,(esk9_0=false|bool(esk3_2(false,true))), inference(spm,[status(thm)],[c_0_462, c_0_463])).
cnf(c_0_466,negated_conjecture,(esk9_0=false|lazy_impl(true,lazy_impl(prop(esk3_2(false,true)),impl(true,impl(impl(true,esk3_2(false,true)),esk3_2(false,true)))))!=or1(false,true)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_463]), c_0_465])).
cnf(c_0_467,negated_conjecture,(esk3_2(false,true)=false|esk3_2(false,true)=true|esk9_0=false), inference(spm,[status(thm)],[c_0_64, c_0_465])).
cnf(c_0_468,negated_conjecture,(esk3_2(false,true)=false|esk9_0=false|lazy_impl(true,lazy_impl(prop(true),true))!=or1(false,true)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_466, c_0_467]), c_0_228]), c_0_228]), c_0_228])).
cnf(c_0_469,negated_conjecture,(esk3_2(false,true)=false|esk9_0=false|or1(false,true)!=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_43]), c_0_174]), c_0_174]), c_0_151])])).
cnf(c_0_470,negated_conjecture,(esk3_2(false,true)=false|esk9_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_469, c_0_459]), c_0_151])])).
cnf(c_0_471,plain,(~forallprefers(lazy_impl(prop(X1),impl(true,impl(impl(X2,X1),X1))),lazy_impl(prop(esk3_2(false,X2)),impl(impl(false,esk3_2(false,X2)),impl(impl(X2,esk3_2(false,X2)),esk3_2(false,X2)))))|~bool(X1)), inference(spm,[status(thm)],[c_0_223, c_0_190])).
cnf(c_0_472,negated_conjecture,(esk9_0=false|lazy_impl(true,lazy_impl(prop(false),impl(true,impl(false,false))))!=or1(false,true)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_466, c_0_470]), c_0_208])).
cnf(c_0_473,plain,(~forallprefers(lazy_impl(true,impl(impl(false,impl(true,X1)),X1)),false)|~bool(X1)), inference(spm,[status(thm)],[c_0_309, c_0_43])).
cnf(c_0_474,plain,(impl(true,esk3_2(false,X1))=true|impl(true,esk3_2(false,X1))=false|~forallprefers(lazy_impl(prop(X2),impl(true,impl(impl(X1,X2),X2))),true)|~bool(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_471, c_0_328]), c_0_41])).
cnf(c_0_475,plain,(d(impl(true,esk3_2(false,X1)))|~forallprefers(lazy_impl(prop(X2),impl(true,impl(impl(X1,X2),X2))),true)|~bool(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_471, c_0_179]), c_0_41])).
cnf(c_0_476,plain,(bool(esk3_2(X1,false))|~forallprefers(lazy_impl(prop(false),impl(impl(X1,false),false)),true)), inference(spm,[status(thm)],[c_0_227, c_0_405])).
cnf(c_0_477,negated_conjecture,(esk9_0=false|lazy_impl(true,lazy_impl(prop(false),true))!=or1(false,true)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_472, c_0_190]), c_0_228]), c_0_59])])).
cnf(c_0_478,plain,(~forallprefers(impl(impl(false,impl(true,X1)),X1),false)|~d(impl(impl(false,impl(true,X1)),X1))|~bool(X1)), inference(spm,[status(thm)],[c_0_473, c_0_50])).
cnf(c_0_479,plain,(impl(true,esk3_2(false,false))=false|impl(true,esk3_2(false,false))=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_474, c_0_405]), c_0_208]), c_0_109]), c_0_57]), c_0_59])])).
cnf(c_0_480,plain,(d(impl(true,esk3_2(false,false)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_475, c_0_405]), c_0_208]), c_0_109]), c_0_57]), c_0_59])])).
cnf(c_0_481,plain,(bool(esk3_2(false,false))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_190]), c_0_208]), c_0_109]), c_0_57]), c_0_59])])).
cnf(c_0_482,negated_conjecture,(lazy_impl(true,impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),impl(impl(esk9_0,esk3_2(esk8_0,esk9_0)),esk3_2(esk8_0,esk9_0))))!=or1(esk8_0,esk9_0)|~d(lazy_impl(true,impl(impl(esk8_0,esk3_2(esk8_0,esk9_0)),impl(impl(esk9_0,esk3_2(esk8_0,esk9_0)),esk3_2(esk8_0,esk9_0)))))|~bool(esk3_2(esk8_0,esk9_0))), inference(spm,[status(thm)],[c_0_321, c_0_43])).
cnf(c_0_483,negated_conjecture,(esk9_0=false|or1(false,true)!=true), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_477, c_0_43]), c_0_174]), c_0_174]), c_0_59])])).
cnf(c_0_484,plain,(~forallprefers(lazy_impl(true,impl(impl(X1,impl(X2,X3)),X3)),lazy_impl(prop(esk1_2(X1,X2)),impl(impl(X1,impl(X2,esk1_2(X1,X2))),esk1_2(X1,X2))))|~bool(X3)), inference(spm,[status(thm)],[c_0_282, c_0_43])).
cnf(c_0_485,plain,(impl(true,esk3_2(false,false))=false|~forallprefers(impl(true,esk3_2(false,false)),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_478, c_0_479]), c_0_398]), c_0_398]), c_0_480]), c_0_481])])).
cnf(c_0_486,negated_conjecture,(lazy_impl(true,impl(impl(false,esk3_2(false,esk9_0)),impl(impl(esk9_0,esk3_2(false,esk9_0)),esk3_2(false,esk9_0))))!=or1(false,esk9_0)|~d(lazy_impl(true,impl(impl(false,esk3_2(false,esk9_0)),impl(impl(esk9_0,esk3_2(false,esk9_0)),esk3_2(false,esk9_0)))))|~bool(esk3_2(false,esk9_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_482, c_0_427]), c_0_427]), c_0_427]), c_0_427]), c_0_427]), c_0_427]), c_0_427]), c_0_427]), c_0_427]), c_0_427])).
cnf(c_0_487,negated_conjecture,(esk9_0=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_483, c_0_459]), c_0_151])])).
cnf(c_0_488,plain,(~forallprefers(lazy_impl(true,impl(true,X1)),lazy_impl(prop(esk1_2(false,X2)),impl(impl(false,impl(X2,esk1_2(false,X2))),esk1_2(false,X2))))|~bool(impl(X2,X1))|~bool(X1)), inference(spm,[status(thm)],[c_0_484, c_0_190])).
cnf(c_0_489,plain,(bool(impl(true,esk3_2(false,false)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_106]), c_0_58]), c_0_59])]), c_0_53])).
cnf(c_0_490,negated_conjecture,(lazy_impl(true,impl(impl(false,esk3_2(false,false)),impl(impl(false,esk3_2(false,false)),esk3_2(false,false))))!=or1(false,false)|~d(lazy_impl(true,impl(impl(false,esk3_2(false,false)),impl(impl(false,esk3_2(false,false)),esk3_2(false,false)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_486, c_0_487]), c_0_487]), c_0_487]), c_0_487]), c_0_487]), c_0_487]), c_0_487]), c_0_487]), c_0_487]), c_0_487]), c_0_481])])).
cnf(c_0_491,plain,(~forallprefers(lazy_impl(true,impl(true,esk3_2(false,false))),false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_488, c_0_489]), c_0_307]), c_0_307]), c_0_208]), c_0_307]), c_0_405]), c_0_109]), c_0_481])])).
cnf(c_0_492,negated_conjecture,(lazy_impl(true,impl(true,impl(true,esk3_2(false,false))))!=or1(false,false)|~d(lazy_impl(true,impl(true,impl(true,esk3_2(false,false)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_490, c_0_190]), c_0_481])])).
cnf(c_0_493,plain,(d(lazy_impl(true,impl(true,esk3_2(false,false))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_96]), c_0_58])])).
cnf(c_0_494,plain,(lazy_impl(true,esk3_2(false,X1))=esk3_2(false,X1)|~forallprefers(lazy_impl(prop(X2),impl(true,impl(impl(X1,X2),X2))),true)|~bool(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_471, c_0_152]), c_0_41])).
cnf(c_0_495,negated_conjecture,(lazy_impl(true,impl(true,esk3_2(false,false)))!=or1(false,false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_99]), c_0_493]), c_0_481])])).
cnf(c_0_496,plain,(lazy_impl(true,esk3_2(false,false))=esk3_2(false,false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_405]), c_0_208]), c_0_109]), c_0_57]), c_0_59])])).
cnf(c_0_497,negated_conjecture,(or1(false,false)!=esk3_2(false,false)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_495, c_0_99]), c_0_496]), c_0_481])])).
cnf(c_0_498,plain,(~forallprefers(lazy_impl(prop(X1),impl(impl(X2,X1),impl(true,X1))),lazy_impl(prop(esk3_2(X2,false)),impl(impl(X2,esk3_2(X2,false)),impl(impl(false,esk3_2(X2,false)),esk3_2(X2,false)))))|~bool(X1)), inference(spm,[status(thm)],[c_0_223, c_0_190])).
cnf(c_0_499,plain,(esk3_2(false,false)=false|esk3_2(false,false)=true), inference(spm,[status(thm)],[c_0_64, c_0_481])).
cnf(c_0_500,negated_conjecture,(esk3_2(false,false)!=false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_497, c_0_459]), c_0_59])])).
cnf(c_0_501,plain,(~forallprefers(lazy_impl(prop(false),impl(impl(X1,false),false)),lazy_impl(prop(esk3_2(X1,false)),impl(impl(X1,esk3_2(X1,false)),impl(impl(false,esk3_2(X1,false)),esk3_2(X1,false)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_498, c_0_208]), c_0_59])])).
cnf(c_0_502,plain,(esk3_2(false,false)=true), inference(sr,[status(thm)],[c_0_499, c_0_500])).
cnf(c_0_503,plain,(~forallprefers(false,lazy_impl(prop(true),true))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_501, c_0_502]), c_0_405]), c_0_109]), c_0_398]), c_0_398]), c_0_228]), c_0_228])).
cnf(c_0_504,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_503, c_0_43]), c_0_174]), c_0_57]), c_0_151])]), ['proof']).
# SZS output end CNFRefutation.
EOF
