% Problem    : LCL896+1 : TPTP v6.1.0. Released v5.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n113.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:37:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____102_C18_F1_PI_AE_CS_SP_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:'+'(X1,'0')=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_03)).
fof(c_0_1, axiom, (![X1]:![X2]:'+'(X1,X2)='+'(X2,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_02)).
fof(c_0_2, axiom, (![X9]:![X10]:![X11]:('>='('+'(X9,X10),X11)<=>'>='(X10,'==>'(X9,X11)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_07)).
fof(c_0_3, axiom, (![X12]:![X13]:![X14]:('>='(X12,X13)=>'>='('+'(X12,X14),'+'(X13,X14)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_09)).
fof(c_0_4, axiom, (![X1]:'>='(X1,'0')), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_08)).
fof(c_0_5, axiom, (![X1]:'>='(X1,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_04)).
fof(c_0_6, axiom, (![X7]:![X8]:(('>='(X7,X8)&'>='(X8,X7))=>X7=X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_06)).
fof(c_0_7, axiom, (![X18]:![X19]:![X20]:('>='(X18,X19)=>'>='('==>'(X20,X18),'==>'(X20,X19)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_11)).
fof(c_0_8, axiom, (![X1]:![X2]:![X3]:'+'('+'(X1,'==>'(X1,X2)),'==>'('+'(X1,'==>'(X1,X2)),X3))='+'(X1,'==>'(X1,'+'(X2,'==>'(X2,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_12)).
fof(c_0_9, axiom, (![X1]:![X2]:![X3]:'+'('+'(X1,X2),X3)='+'(X1,'+'(X2,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_01)).
fof(c_0_10, conjecture, (![X21]:![X22]:'+'(X21,'==>'(X21,X22))='+'(X22,'==>'(X22,X21))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_13)).
fof(c_0_11, plain, (![X2]:'+'(X2,'0')=X2), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_12, plain, (![X3]:![X4]:'+'(X3,X4)='+'(X4,X3)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_13, plain, (![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:((~'>='('+'(X12,X13),X14)|'>='(X13,'==>'(X12,X14)))&(~'>='(X16,'==>'(X15,X17))|'>='('+'(X15,X16),X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
fof(c_0_14, plain, (![X15]:![X16]:![X17]:(~'>='(X15,X16)|'>='('+'(X15,X17),'+'(X16,X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
fof(c_0_15, plain, (![X2]:'>='(X2,'0')), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_16,plain,('+'(X1,'0')=X1), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,('+'(X1,X2)='+'(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,('>='(X1,'==>'(X2,X3))|~'>='('+'(X2,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,('>='('+'(X1,X2),'+'(X3,X2))|~'>='(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,('>='(X1,'0')), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,('+'('0',X1)=X1), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_22, plain, (![X2]:'>='(X2,X2)), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_23, plain, (![X9]:![X10]:((~'>='(X9,X10)|~'>='(X10,X9))|X9=X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
cnf(c_0_24,plain,('>='('0','==>'(X1,X2))|~'>='(X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_16])).
cnf(c_0_25,plain,('>='('+'(X1,X2),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_26,plain,('>='(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, (![X21]:![X22]:![X23]:(~'>='(X21,X22)|'>='('==>'(X23,X21),'==>'(X23,X22)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
cnf(c_0_28,plain,('>='('+'(X1,X2),X3)|~'>='(X2,'==>'(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_29, plain, (![X4]:![X5]:![X6]:'+'('+'(X4,'==>'(X4,X5)),'==>'('+'(X4,'==>'(X4,X5)),X6))='+'(X4,'==>'(X4,'+'(X5,'==>'(X5,X6))))), inference(variable_rename,[status(thm)],[c_0_8])).
fof(c_0_30, plain, (![X4]:![X5]:![X6]:'+'('+'(X4,X5),X6)='+'(X4,'+'(X5,X6))), inference(variable_rename,[status(thm)],[c_0_9])).
cnf(c_0_31,plain,(X1=X2|~'>='(X2,X1)|~'>='(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,plain,('>='('0','==>'('+'(X1,X2),X2))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33,plain,('>='(X1,'==>'(X2,'+'(X2,X1)))), inference(spm,[status(thm)],[c_0_18, c_0_26])).
cnf(c_0_34,plain,('>='('==>'(X1,X2),'==>'(X1,X3))|~'>='(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35,plain,('>='('+'(X1,'==>'(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_28, c_0_26])).
cnf(c_0_36,plain,('+'('+'(X1,'==>'(X1,X2)),'==>'('+'(X1,'==>'(X1,X2)),X3))='+'(X1,'==>'(X1,'+'(X2,'==>'(X2,X3))))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37,plain,('+'('+'(X1,X2),X3)='+'(X1,'+'(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38,plain,('==>'('+'(X1,X2),X2)='0'), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_20])])).
cnf(c_0_39,plain,('==>'(X1,'+'(X1,X2))=X2|~'>='('==>'(X1,'+'(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_31, c_0_33])).
cnf(c_0_40,plain,('>='('==>'(X1,'+'(X2,'==>'(X2,X3))),'==>'(X1,X3))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41,plain,('+'(X1,'+'('==>'(X1,X2),'==>'('+'(X1,'==>'(X1,X2)),X3)))='+'(X1,'==>'(X1,'+'(X2,'==>'(X2,X3))))), inference(rw,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42,plain,('==>'('+'(X1,X2),X1)='0'), inference(spm,[status(thm)],[c_0_38, c_0_17])).
cnf(c_0_43,plain,('==>'(X1,'+'(X1,'==>'(X1,X2)))='==>'(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44,plain,('+'(X1,'==>'(X1,'+'(X2,'==>'(X2,X1))))='+'(X1,'==>'(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_16])).
fof(c_0_45, negated_conjecture, (~(![X21]:![X22]:'+'(X21,'==>'(X21,X22))='+'(X22,'==>'(X22,X21)))), inference(assume_negation,[status(cth)],[c_0_10])).
cnf(c_0_46,plain,('==>'(X1,'+'(X2,'==>'(X2,X1)))='==>'(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_43])).
fof(c_0_47, negated_conjecture, ('+'(esk1_0,'==>'(esk1_0,esk2_0))!='+'(esk2_0,'==>'(esk2_0,esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_45])])])).
cnf(c_0_48,plain,('>='('+'(X1,'==>'(X1,X2)),'+'(X2,'==>'(X2,X1)))), inference(spm,[status(thm)],[c_0_35, c_0_46])).
cnf(c_0_49,negated_conjecture,('+'(esk1_0,'==>'(esk1_0,esk2_0))!='+'(esk2_0,'==>'(esk2_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_50,plain,('+'(X1,'==>'(X1,X2))='+'(X2,'==>'(X2,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_48]), c_0_48])])).
cnf(c_0_51,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50])]), ['proof']).
# SZS output end CNFRefutation.
EOF
