% Problem    : LCL496+1 : TPTP v6.1.0. Released v3.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n190.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:21:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____011_C18_F1_PI_AE_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, ((modus_ponens<=>![X1]:![X2]:((is_a_theorem(X1)&is_a_theorem(implies(X1,X2)))=>is_a_theorem(X2)))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', modus_ponens)).
fof(c_0_1, axiom, ((r5<=>![X4]:![X5]:![X6]:is_a_theorem(implies(implies(X5,X6),implies(or(X4,X5),or(X4,X6)))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', r5)).
fof(c_0_2, axiom, ((r3<=>![X4]:![X5]:is_a_theorem(implies(or(X4,X5),or(X5,X4))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', r3)).
fof(c_0_3, axiom, (modus_ponens), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', principia_modus_ponens)).
fof(c_0_4, axiom, (r5), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', principia_r5)).
fof(c_0_5, axiom, ((r2<=>![X4]:![X5]:is_a_theorem(implies(X5,or(X4,X5))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', r2)).
fof(c_0_6, axiom, (r3), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', principia_r3)).
fof(c_0_7, axiom, ((op_implies_or=>![X1]:![X2]:implies(X1,X2)=or(not(X1),X2))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+1.ax', op_implies_or)).
fof(c_0_8, axiom, (r2), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', principia_r2)).
fof(c_0_9, axiom, (op_implies_or), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', principia_op_implies_or)).
fof(c_0_10, axiom, ((r4<=>![X4]:![X5]:![X6]:is_a_theorem(implies(or(X4,or(X5,X6)),or(X5,or(X4,X6)))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', r4)).
fof(c_0_11, axiom, ((r1<=>![X4]:is_a_theorem(implies(or(X4,X4),X4)))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', r1)).
fof(c_0_12, axiom, (r4), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', principia_r4)).
fof(c_0_13, axiom, (r1), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', principia_r1)).
fof(c_0_14, axiom, ((op_and=>![X1]:![X2]:and(X1,X2)=not(or(not(X1),not(X2))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+1.ax', op_and)).
fof(c_0_15, axiom, (op_and), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', principia_op_and)).
fof(c_0_16, axiom, ((op_equiv=>![X1]:![X2]:equiv(X1,X2)=and(implies(X1,X2),implies(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+1.ax', op_equiv)).
fof(c_0_17, axiom, (op_equiv), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', principia_op_equiv)).
fof(c_0_18, axiom, ((substitution_of_equivalents<=>![X1]:![X2]:(is_a_theorem(equiv(X1,X2))=>X1=X2))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', substitution_of_equivalents)).
fof(c_0_19, axiom, (substitution_of_equivalents), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+4.ax', substitution_of_equivalents)).
fof(c_0_20, axiom, ((cn1<=>![X4]:![X5]:![X6]:is_a_theorem(implies(implies(X4,X5),implies(implies(X5,X6),implies(X4,X6)))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', cn1)).
fof(c_0_21, conjecture, (cn1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', luka_cn1)).
fof(c_0_22, plain, (![X3]:![X4]:((~modus_ponens|((~is_a_theorem(X3)|~is_a_theorem(implies(X3,X4)))|is_a_theorem(X4)))&(((is_a_theorem(esk1_0)|modus_ponens)&(is_a_theorem(implies(esk1_0,esk2_0))|modus_ponens))&(~is_a_theorem(esk2_0)|modus_ponens)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])).
fof(c_0_23, plain, (![X7]:![X8]:![X9]:((~r5|is_a_theorem(implies(implies(X8,X9),implies(or(X7,X8),or(X7,X9)))))&(~is_a_theorem(implies(implies(esk54_0,esk55_0),implies(or(esk53_0,esk54_0),or(esk53_0,esk55_0))))|r5))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_24, plain, (![X6]:![X7]:((~r3|is_a_theorem(implies(or(X6,X7),or(X7,X6))))&(~is_a_theorem(implies(or(esk48_0,esk49_0),or(esk49_0,esk48_0)))|r3))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_25,plain,(is_a_theorem(X1)|~is_a_theorem(implies(X2,X1))|~is_a_theorem(X2)|~modus_ponens), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26,plain,(modus_ponens), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_27,plain,(is_a_theorem(implies(implies(X1,X2),implies(or(X3,X1),or(X3,X2))))|~r5), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,plain,(r5), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_29, plain, (![X6]:![X7]:((~r2|is_a_theorem(implies(X7,or(X6,X7))))&(~is_a_theorem(implies(esk47_0,or(esk46_0,esk47_0)))|r2))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
cnf(c_0_30,plain,(is_a_theorem(implies(or(X1,X2),or(X2,X1)))|~r3), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,plain,(r3), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_32, plain, (![X3]:![X4]:(~op_implies_or|implies(X3,X4)=or(not(X3),X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_33,plain,(is_a_theorem(X1)|~is_a_theorem(implies(X2,X1))|~is_a_theorem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26])])).
cnf(c_0_34,plain,(is_a_theorem(implies(implies(X1,X2),implies(or(X3,X1),or(X3,X2))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28])])).
cnf(c_0_35,plain,(is_a_theorem(implies(X1,or(X2,X1)))|~r2), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36,plain,(r2), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_37,plain,(is_a_theorem(implies(or(X1,X2),or(X2,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31])])).
cnf(c_0_38,plain,(implies(X1,X2)=or(not(X1),X2)|~op_implies_or), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39,plain,(op_implies_or), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40,plain,(is_a_theorem(implies(or(X1,X2),or(X1,X3)))|~is_a_theorem(implies(X2,X3))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41,plain,(is_a_theorem(implies(X1,or(X2,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
cnf(c_0_42,plain,(is_a_theorem(or(X1,X2))|~is_a_theorem(or(X2,X1))), inference(spm,[status(thm)],[c_0_33, c_0_37])).
cnf(c_0_43,plain,(or(not(X1),X2)=implies(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39])])).
fof(c_0_44, plain, (![X7]:![X8]:![X9]:((~r4|is_a_theorem(implies(or(X7,or(X8,X9)),or(X8,or(X7,X9)))))&(~is_a_theorem(implies(or(esk50_0,or(esk51_0,esk52_0)),or(esk51_0,or(esk50_0,esk52_0))))|r4))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
cnf(c_0_45,plain,(is_a_theorem(implies(or(X1,X2),or(X1,or(X3,X2))))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46,plain,(is_a_theorem(or(X1,not(X2)))|~is_a_theorem(implies(X2,X1))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_47, plain, (![X5]:((~r1|is_a_theorem(implies(or(X5,X5),X5)))&(~is_a_theorem(implies(or(esk45_0,esk45_0),esk45_0))|r1))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
cnf(c_0_48,plain,(is_a_theorem(implies(or(X1,or(X2,X3)),or(X2,or(X1,X3))))|~r4), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_49,plain,(r4), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_50,plain,(is_a_theorem(or(X1,or(X2,X3)))|~is_a_theorem(or(X1,X3))), inference(spm,[status(thm)],[c_0_33, c_0_45])).
cnf(c_0_51,plain,(is_a_theorem(or(or(X1,X2),not(X2)))), inference(spm,[status(thm)],[c_0_46, c_0_41])).
cnf(c_0_52,plain,(is_a_theorem(implies(or(X1,X1),X1))|~r1), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53,plain,(r1), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54,plain,(is_a_theorem(implies(or(X1,or(X2,X3)),or(X2,or(X1,X3))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49])])).
cnf(c_0_55,plain,(is_a_theorem(or(or(X1,X2),or(X3,not(X2))))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_56,plain,(is_a_theorem(implies(or(X1,X1),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53])])).
cnf(c_0_57,plain,(is_a_theorem(or(X1,or(X2,X3)))|~is_a_theorem(or(X2,or(X1,X3)))), inference(spm,[status(thm)],[c_0_33, c_0_54])).
cnf(c_0_58,plain,(is_a_theorem(or(or(X1,not(X2)),or(X3,X2)))), inference(spm,[status(thm)],[c_0_42, c_0_55])).
cnf(c_0_59,plain,(is_a_theorem(X1)|~is_a_theorem(or(X1,X1))), inference(spm,[status(thm)],[c_0_33, c_0_56])).
cnf(c_0_60,plain,(is_a_theorem(or(X1,or(or(X2,not(X3)),X3)))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_61,plain,(is_a_theorem(or(or(X1,not(X2)),X2))), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_62,plain,(is_a_theorem(or(X1,or(X2,not(X1))))), inference(spm,[status(thm)],[c_0_42, c_0_61])).
cnf(c_0_63,plain,(is_a_theorem(or(X1,or(X2,not(X2))))), inference(spm,[status(thm)],[c_0_57, c_0_62])).
cnf(c_0_64,plain,(is_a_theorem(or(X1,not(X1)))), inference(spm,[status(thm)],[c_0_59, c_0_63])).
cnf(c_0_65,plain,(is_a_theorem(or(X1,implies(X2,X3)))|~is_a_theorem(implies(X2,or(X1,X3)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_43]), c_0_43])).
cnf(c_0_66,plain,(is_a_theorem(implies(X1,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_64]), c_0_43])).
cnf(c_0_67,plain,(is_a_theorem(implies(or(X1,not(X2)),implies(X2,X1)))), inference(spm,[status(thm)],[c_0_37, c_0_43])).
cnf(c_0_68,plain,(is_a_theorem(or(X1,implies(or(X1,X2),X2)))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
fof(c_0_69, plain, (![X3]:![X4]:(~op_and|and(X3,X4)=not(or(not(X3),not(X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
cnf(c_0_70,plain,(is_a_theorem(implies(X1,X2))|~is_a_theorem(or(X2,not(X1)))), inference(spm,[status(thm)],[c_0_33, c_0_67])).
cnf(c_0_71,plain,(is_a_theorem(implies(X1,implies(implies(X1,X2),X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_43]), c_0_43])).
cnf(c_0_72,plain,(and(X1,X2)=not(or(not(X1),not(X2)))|~op_and), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_73,plain,(op_and), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_74,plain,(is_a_theorem(implies(X1,not(not(X1))))), inference(spm,[status(thm)],[c_0_64, c_0_43])).
cnf(c_0_75,plain,(is_a_theorem(implies(X1,not(X2)))|~is_a_theorem(implies(X2,not(X1)))), inference(spm,[status(thm)],[c_0_70, c_0_43])).
cnf(c_0_76,plain,(is_a_theorem(implies(implies(X1,X2),X2))|~is_a_theorem(X1)), inference(spm,[status(thm)],[c_0_33, c_0_71])).
cnf(c_0_77,plain,(not(implies(X1,not(X2)))=and(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_43]), c_0_73])])).
fof(c_0_78, plain, (![X3]:![X4]:(~op_equiv|equiv(X3,X4)=and(implies(X3,X4),implies(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
cnf(c_0_79,plain,(is_a_theorem(implies(or(X1,X2),or(X1,not(not(X2)))))), inference(spm,[status(thm)],[c_0_40, c_0_74])).
cnf(c_0_80,plain,(is_a_theorem(implies(X1,and(X2,X1)))|~is_a_theorem(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])).
cnf(c_0_81,plain,(equiv(X1,X2)=and(implies(X1,X2),implies(X2,X1))|~op_equiv), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_82,plain,(op_equiv), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_83,plain,(is_a_theorem(or(X1,not(not(X2))))|~is_a_theorem(or(X1,X2))), inference(spm,[status(thm)],[c_0_33, c_0_79])).
fof(c_0_84, plain, (![X3]:![X4]:((~substitution_of_equivalents|(~is_a_theorem(equiv(X3,X4))|X3=X4))&((is_a_theorem(equiv(esk3_0,esk4_0))|substitution_of_equivalents)&(esk3_0!=esk4_0|substitution_of_equivalents)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])])).
cnf(c_0_85,plain,(is_a_theorem(and(X1,X2))|~is_a_theorem(X2)|~is_a_theorem(X1)), inference(spm,[status(thm)],[c_0_33, c_0_80])).
cnf(c_0_86,plain,(and(implies(X1,X2),implies(X2,X1))=equiv(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_82])])).
cnf(c_0_87,plain,(is_a_theorem(or(X1,not(not(not(X1)))))), inference(spm,[status(thm)],[c_0_83, c_0_64])).
cnf(c_0_88,plain,(X1=X2|~is_a_theorem(equiv(X1,X2))|~substitution_of_equivalents), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_89,plain,(substitution_of_equivalents), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_90,plain,(is_a_theorem(equiv(X1,X2))|~is_a_theorem(implies(X2,X1))|~is_a_theorem(implies(X1,X2))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_91,plain,(is_a_theorem(implies(not(not(X1)),X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_87]), c_0_43])).
cnf(c_0_92,plain,(X1=X2|~is_a_theorem(equiv(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_89])])).
cnf(c_0_93,plain,(is_a_theorem(equiv(or(X1,X2),or(X2,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_37]), c_0_37])])).
cnf(c_0_94,plain,(is_a_theorem(equiv(X1,not(not(X1))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_74])])).
cnf(c_0_95,plain,(or(X1,X2)=or(X2,X1)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_96,plain,(not(not(X1))=X1), inference(spm,[status(thm)],[c_0_92, c_0_94])).
cnf(c_0_97,plain,(is_a_theorem(implies(or(X1,implies(X2,X3)),implies(X2,or(X1,X3))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_43]), c_0_43])).
cnf(c_0_98,plain,(is_a_theorem(implies(implies(X1,or(X2,X3)),or(X2,implies(X1,X3))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_43]), c_0_43])).
cnf(c_0_99,plain,(or(X1,not(X2))=implies(X2,X1)), inference(spm,[status(thm)],[c_0_43, c_0_95])).
cnf(c_0_100,plain,(not(and(X1,X2))=implies(X1,not(X2))), inference(spm,[status(thm)],[c_0_96, c_0_77])).
cnf(c_0_101,plain,(is_a_theorem(equiv(implies(X1,or(X2,X3)),or(X2,implies(X1,X3))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_97]), c_0_98])])).
fof(c_0_102, plain, (![X7]:![X8]:![X9]:((~cn1|is_a_theorem(implies(implies(X7,X8),implies(implies(X8,X9),implies(X7,X9)))))&(~is_a_theorem(implies(implies(esk39_0,esk40_0),implies(implies(esk40_0,esk41_0),implies(esk39_0,esk41_0))))|cn1))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).
fof(c_0_103, negated_conjecture, (~cn1), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_21])])).
cnf(c_0_104,plain,(or(X1,implies(X2,not(X3)))=implies(and(X2,X3),X1)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_105,plain,(or(X1,implies(X2,X3))=implies(X2,or(X1,X3))), inference(spm,[status(thm)],[c_0_92, c_0_101])).
cnf(c_0_106,plain,(is_a_theorem(implies(implies(X1,X2),implies(implies(X3,X1),implies(X3,X2))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_43]), c_0_43])).
cnf(c_0_107,plain,(cn1|~is_a_theorem(implies(implies(esk39_0,esk40_0),implies(implies(esk40_0,esk41_0),implies(esk39_0,esk41_0))))), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_108,negated_conjecture,(~cn1), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_109,plain,(implies(X1,not(X2))=implies(X2,not(X1))), inference(spm,[status(thm)],[c_0_43, c_0_99])).
cnf(c_0_110,plain,(implies(and(X1,X2),X3)=implies(X1,implies(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_105]), c_0_99])).
cnf(c_0_111,plain,(is_a_theorem(implies(implies(X1,X2),implies(X1,X3)))|~is_a_theorem(implies(X2,X3))), inference(spm,[status(thm)],[c_0_33, c_0_106])).
cnf(c_0_112,plain,(~is_a_theorem(implies(implies(esk39_0,esk40_0),implies(implies(esk40_0,esk41_0),implies(esk39_0,esk41_0))))), inference(sr,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_113,plain,(implies(X1,implies(X2,X3))=implies(X2,implies(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_109]), c_0_104]), c_0_99]), c_0_110])).
cnf(c_0_114,plain,(is_a_theorem(implies(implies(X1,X2),implies(X1,implies(implies(X2,X3),X3))))), inference(spm,[status(thm)],[c_0_111, c_0_71])).
cnf(c_0_115,plain,(~is_a_theorem(implies(esk39_0,implies(implies(esk39_0,esk40_0),implies(implies(esk40_0,esk41_0),esk41_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_113]), c_0_113])).
cnf(c_0_116,plain,(is_a_theorem(implies(X1,implies(implies(X1,X2),implies(implies(X2,X3),X3))))), inference(rw,[status(thm)],[c_0_114, c_0_113])).
cnf(c_0_117,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_116])]), ['proof']).
# SZS output end CNFRefutation.
EOF
