% Problem    : NUM924+8 : TPTP v6.1.0. Released v5.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n191.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:44:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X6]:hAPP(int,int,succ,hAPP(int,int,pred,X6))=ti(int,X6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_2899_succ__pred)).
fof(c_0_1, axiom, (![X1]:![X228]:hAPP(X1,X1,combi(X1),X228)=ti(X1,X228)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_COMBI_1_1_U)).
fof(c_0_2, axiom, (![X15]:hAPP(int,int,number_number_of(int),X15)=ti(int,X15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_22_number__of__is__id)).
fof(c_0_3, axiom, (zero_zero(nat)=hAPP(int,nat,nat_1,zero_zero(int))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1102_transfer__nat__int__numerals_I1_J)).
fof(c_0_4, axiom, (pls=zero_zero(int)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_123_Pls__def)).
fof(c_0_5, conjecture, (hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),one_one(int))),zero_zero(int)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_6, axiom, (hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls)))=hAPP(nat,nat,suc,hAPP(nat,nat,suc,zero_zero(nat)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1426_numeral__2__eq__2)).
fof(c_0_7, axiom, (one_one(nat)=hAPP(int,nat,nat_1,one_one(int))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1115_transfer__nat__int__numerals_I2_J)).
fof(c_0_8, axiom, (one_one(nat)=hAPP(nat,nat,suc,zero_zero(nat))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1330_One__nat__def)).
fof(c_0_9, axiom, (![X15]:hAPP(int,int,pred,hAPP(int,int,bit1,X15))=hAPP(int,int,bit0,X15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_2900_pred__Bit1)).
fof(c_0_10, axiom, (![X15]:hAPP(int,int,succ,X15)=hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X15),one_one(int))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1528_succ__def)).
fof(c_0_11, axiom, (one_one(int)=hAPP(int,int,number_number_of(int),hAPP(int,int,bit1,pls))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_112_one__is__num__one)).
fof(c_0_12, axiom, (hAPP(int,int,bit0,pls)=pls), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_122_Bit0__Pls)).
fof(c_0_13, axiom, (hAPP(int,int,succ,pls)=hAPP(int,int,bit1,pls)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1525_succ__Pls)).
fof(c_0_14, axiom, (hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls)))=hAPP(int,nat,nat_1,hAPP(int,int,number_number_of(int),hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1148_transfer__nat__int__numerals_I3_J)).
fof(c_0_15, axiom, (![X16]:hAPP(int,int,hAPP(int,fun(int,int),times_times(int),one_one(int)),X16)=ti(int,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_66_zmult__1)).
fof(c_0_16, axiom, (ti(int,m)=m), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tsy_v_m_res)).
fof(c_0_17, axiom, (![X16]:hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X16),zero_zero(int))=ti(int,X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_134_zadd__0__right)).
fof(c_0_18, axiom, (![X15]:hAPP(int,int,pred,hAPP(int,int,bit0,X15))=hAPP(int,int,bit1,hAPP(int,int,pred,X15))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_2901_pred__Bit0)).
fof(c_0_19, axiom, (![X15]:hAPP(int,int,succ,hAPP(int,int,bit1,X15))=hAPP(int,int,bit0,hAPP(int,int,succ,X15))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1526_succ__Bit1)).
fof(c_0_20, axiom, (![X13]:hAPP(int,nat,number_number_of(nat),X13)=hAPP(int,nat,nat_1,hAPP(int,int,number_number_of(int),X13))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_1112_nat__number__of__def)).
fof(c_0_21, axiom, (![X15]:![X32]:hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,bit0,X15)),X32)=hAPP(int,int,bit0,hAPP(int,int,hAPP(int,fun(int,int),times_times(int),X15),X32))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_63_mult__Bit0)).
fof(c_0_22, axiom, (![X15]:![X32]:hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(int,int,bit0,X15)),hAPP(int,int,bit1,X32))=hAPP(int,int,bit1,hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X15),X32))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_163_add__Bit0__Bit1)).
fof(c_0_23, axiom, (![X16]:![X12]:hAPP(int,int,hAPP(int,fun(int,int),times_times(int),X16),X12)=hAPP(int,int,hAPP(int,fun(int,int),times_times(int),X12),X16)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_23_zmult__commute)).
fof(c_0_24, axiom, (hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),one_one(int))=hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,number_number_of(int),hAPP(int,int,bit0,hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),m)),one_one(int))),t)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_3_t)).
fof(c_0_25, axiom, (![X12]:hAPP(int,int,hAPP(int,fun(int,int),times_times(int),pls),X12)=pls), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_62_mult__Pls)).
fof(c_0_26, axiom, (hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,number_number_of(int),hAPP(int,int,bit0,hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),m)),one_one(int))),t)),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,number_number_of(int),hAPP(int,int,bit0,hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),m)),one_one(int))),zero_zero(int))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_2__096_I4_A_K_Am_A_L_A1_J_A_K_At_A_060_A_I4_A_K_Am_A_L_A1_J_A_K_A0_096)).
fof(c_0_27, plain, (![X7]:hAPP(int,int,succ,hAPP(int,int,pred,X7))=ti(int,X7)), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_28, plain, (![X229]:![X230]:hAPP(X229,X229,combi(X229),X230)=ti(X229,X230)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_29, plain, (![X16]:hAPP(int,int,number_number_of(int),X16)=ti(int,X16)), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_30,plain,(zero_zero(nat)=hAPP(int,nat,nat_1,zero_zero(int))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_31,plain,(pls=zero_zero(int)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_32,plain,(hAPP(int,int,succ,hAPP(int,int,pred,X1))=ti(int,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33,plain,(hAPP(X1,X1,combi(X1),X2)=ti(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34,plain,(hAPP(int,int,number_number_of(int),X1)=ti(int,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, negated_conjecture, (~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),one_one(int))),zero_zero(int)))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_5])])).
cnf(c_0_36,plain,(hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls)))=hAPP(nat,nat,suc,hAPP(nat,nat,suc,zero_zero(nat)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_37,plain,(zero_zero(nat)=hAPP(int,nat,nat_1,pls)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38,plain,(one_one(nat)=hAPP(int,nat,nat_1,one_one(int))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_39,plain,(one_one(nat)=hAPP(nat,nat,suc,zero_zero(nat))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_40,plain,(hAPP(int,int,succ,hAPP(int,int,pred,X1))=hAPP(int,int,combi(int),X1)), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41,plain,(hAPP(int,int,number_number_of(int),X1)=hAPP(int,int,combi(int),X1)), inference(rw,[status(thm)],[c_0_34, c_0_33])).
fof(c_0_42, plain, (![X16]:hAPP(int,int,pred,hAPP(int,int,bit1,X16))=hAPP(int,int,bit0,X16)), inference(variable_rename,[status(thm)],[c_0_9])).
cnf(c_0_43,negated_conjecture,(~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),one_one(int))),zero_zero(int)))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44,plain,(hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls)))=hAPP(nat,nat,suc,hAPP(nat,nat,suc,hAPP(int,nat,nat_1,pls)))), inference(rw,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_45, plain, (![X16]:hAPP(int,int,succ,X16)=hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X16),one_one(int))), inference(variable_rename,[status(thm)],[c_0_10])).
cnf(c_0_46,plain,(hAPP(int,nat,nat_1,one_one(int))=hAPP(nat,nat,suc,zero_zero(nat))), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_47,plain,(one_one(int)=hAPP(int,int,number_number_of(int),hAPP(int,int,bit1,pls))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_48,plain,(hAPP(int,int,number_number_of(int),X1)=hAPP(int,int,succ,hAPP(int,int,pred,X1))), inference(rw,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_49,plain,(hAPP(int,int,pred,hAPP(int,int,bit1,X1))=hAPP(int,int,bit0,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50,plain,(hAPP(int,int,bit0,pls)=pls), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_51,plain,(hAPP(int,int,succ,pls)=hAPP(int,int,bit1,pls)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_52,plain,(hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls)))=hAPP(int,nat,nat_1,hAPP(int,int,number_number_of(int),hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_53, plain, (![X17]:hAPP(int,int,hAPP(int,fun(int,int),times_times(int),one_one(int)),X17)=ti(int,X17)), inference(variable_rename,[status(thm)],[c_0_15])).
cnf(c_0_54,plain,(ti(int,m)=m), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_55, plain, (![X17]:hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X17),zero_zero(int))=ti(int,X17)), inference(variable_rename,[status(thm)],[c_0_17])).
cnf(c_0_56,negated_conjecture,(~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(nat,nat,suc,hAPP(nat,nat,suc,hAPP(int,nat,nat_1,pls))))),one_one(int))),pls))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_31])).
cnf(c_0_57,plain,(hAPP(int,int,succ,X1)=hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X1),one_one(int))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_58,plain,(hAPP(int,nat,nat_1,one_one(int))=hAPP(nat,nat,suc,hAPP(int,nat,nat_1,pls))), inference(rw,[status(thm)],[c_0_46, c_0_37])).
cnf(c_0_59,plain,(one_one(int)=hAPP(int,int,bit1,pls)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_60,plain,(hAPP(int,nat,nat_1,hAPP(int,int,number_number_of(int),hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))=hAPP(nat,nat,suc,hAPP(nat,nat,suc,hAPP(int,nat,nat_1,pls)))), inference(rw,[status(thm)],[c_0_52, c_0_44])).
fof(c_0_61, plain, (![X16]:hAPP(int,int,pred,hAPP(int,int,bit0,X16))=hAPP(int,int,bit1,hAPP(int,int,pred,X16))), inference(variable_rename,[status(thm)],[c_0_18])).
fof(c_0_62, plain, (![X16]:hAPP(int,int,succ,hAPP(int,int,bit1,X16))=hAPP(int,int,bit0,hAPP(int,int,succ,X16))), inference(variable_rename,[status(thm)],[c_0_19])).
cnf(c_0_63,plain,(hAPP(int,int,hAPP(int,fun(int,int),times_times(int),one_one(int)),X1)=ti(int,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_64,plain,(hAPP(int,int,combi(int),m)=m), inference(rw,[status(thm)],[c_0_54, c_0_33])).
cnf(c_0_65,plain,(hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X1),zero_zero(int))=ti(int,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_66, plain, (![X14]:hAPP(int,nat,number_number_of(nat),X14)=hAPP(int,nat,nat_1,hAPP(int,int,number_number_of(int),X14))), inference(variable_rename,[status(thm)],[c_0_20])).
cnf(c_0_67,negated_conjecture,(~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,succ,hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(nat,nat,suc,hAPP(nat,nat,suc,hAPP(int,nat,nat_1,pls)))))),pls))), inference(rw,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_68,plain,(hAPP(nat,nat,suc,hAPP(int,nat,nat_1,pls))=hAPP(int,nat,nat_1,hAPP(int,int,bit1,pls))), inference(rw,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_69,plain,(hAPP(int,nat,nat_1,hAPP(int,int,succ,hAPP(int,int,pred,hAPP(int,int,bit0,hAPP(int,int,bit1,pls)))))=hAPP(nat,nat,suc,hAPP(nat,nat,suc,hAPP(int,nat,nat_1,pls)))), inference(rw,[status(thm)],[c_0_60, c_0_48])).
cnf(c_0_70,plain,(hAPP(int,int,pred,hAPP(int,int,bit0,X1))=hAPP(int,int,bit1,hAPP(int,int,pred,X1))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71,plain,(hAPP(int,int,succ,hAPP(int,int,bit1,X1))=hAPP(int,int,bit0,hAPP(int,int,succ,X1))), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_72, plain, (![X33]:![X34]:hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,bit0,X33)),X34)=hAPP(int,int,bit0,hAPP(int,int,hAPP(int,fun(int,int),times_times(int),X33),X34))), inference(variable_rename,[status(thm)],[c_0_21])).
cnf(c_0_73,plain,(hAPP(int,int,hAPP(int,fun(int,int),times_times(int),one_one(int)),X1)=hAPP(int,int,combi(int),X1)), inference(rw,[status(thm)],[c_0_63, c_0_33])).
cnf(c_0_74,plain,(hAPP(int,int,number_number_of(int),m)=m), inference(rw,[status(thm)],[c_0_64, c_0_41])).
fof(c_0_75, plain, (![X33]:![X34]:hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(int,int,bit0,X33)),hAPP(int,int,bit1,X34))=hAPP(int,int,bit1,hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X33),X34))), inference(variable_rename,[status(thm)],[c_0_22])).
cnf(c_0_76,plain,(hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X1),zero_zero(int))=hAPP(int,int,combi(int),X1)), inference(rw,[status(thm)],[c_0_65, c_0_33])).
fof(c_0_77, plain, (![X17]:![X18]:hAPP(int,int,hAPP(int,fun(int,int),times_times(int),X17),X18)=hAPP(int,int,hAPP(int,fun(int,int),times_times(int),X18),X17)), inference(variable_rename,[status(thm)],[c_0_23])).
cnf(c_0_78,plain,(hAPP(int,nat,number_number_of(nat),X1)=hAPP(int,nat,nat_1,hAPP(int,int,number_number_of(int),X1))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_79,negated_conjecture,(~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,succ,hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(nat,nat,suc,hAPP(int,nat,nat_1,hAPP(int,int,bit1,pls)))))),pls))), inference(rw,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_80,plain,(hAPP(nat,nat,suc,hAPP(int,nat,nat_1,hAPP(int,int,bit1,pls)))=hAPP(int,nat,nat_1,hAPP(int,int,bit0,hAPP(int,int,bit1,pls)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_68]), c_0_70]), c_0_49]), c_0_50]), c_0_71]), c_0_51])).
cnf(c_0_81,plain,(hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(int,nat,number_number_of(nat),hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),one_one(int))=hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,number_number_of(int),hAPP(int,int,bit0,hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),m)),one_one(int))),t)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_82,plain,(hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,bit0,X1)),X2)=hAPP(int,int,bit0,hAPP(int,int,hAPP(int,fun(int,int),times_times(int),X1),X2))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_83,plain,(hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,bit1,pls)),X1)=hAPP(int,int,succ,hAPP(int,int,pred,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_59]), c_0_41]), c_0_48])).
cnf(c_0_84,plain,(hAPP(int,int,succ,hAPP(int,int,pred,m))=m), inference(rw,[status(thm)],[c_0_74, c_0_48])).
cnf(c_0_85,plain,(hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(int,int,bit0,X1)),hAPP(int,int,bit1,X2))=hAPP(int,int,bit1,hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X1),X2))), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_86,plain,(hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X1),pls)=hAPP(int,int,succ,hAPP(int,int,pred,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_31]), c_0_41]), c_0_48])).
cnf(c_0_87,plain,(hAPP(int,int,hAPP(int,fun(int,int),times_times(int),X1),X2)=hAPP(int,int,hAPP(int,fun(int,int),times_times(int),X2),X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_88,plain,(hAPP(int,nat,number_number_of(nat),X1)=hAPP(int,nat,nat_1,hAPP(int,int,succ,hAPP(int,int,pred,X1)))), inference(rw,[status(thm)],[c_0_78, c_0_48])).
cnf(c_0_89,plain,(hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),X1),hAPP(int,int,bit1,pls))=hAPP(int,int,succ,X1)), inference(rw,[status(thm)],[c_0_57, c_0_59])).
fof(c_0_90, plain, (![X13]:hAPP(int,int,hAPP(int,fun(int,int),times_times(int),pls),X13)=pls), inference(variable_rename,[status(thm)],[c_0_25])).
cnf(c_0_91,negated_conjecture,(~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,succ,hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(int,nat,nat_1,hAPP(int,int,bit0,hAPP(int,int,bit1,pls)))))),pls))), inference(rw,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_92,plain,(hAPP(int,int,succ,hAPP(nat,int,hAPP(int,fun(nat,int),power_power(int),s),hAPP(int,nat,nat_1,hAPP(int,int,bit0,hAPP(int,int,bit1,pls)))))=hAPP(int,int,hAPP(int,fun(int,int),times_times(int),t),hAPP(int,int,bit1,hAPP(int,int,bit0,m)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_48]), c_0_70]), c_0_70]), c_0_49]), c_0_50]), c_0_71]), c_0_71]), c_0_51]), c_0_82]), c_0_82]), c_0_83]), c_0_84]), c_0_59]), c_0_85]), c_0_86]), c_0_70]), c_0_71]), c_0_84]), c_0_87]), c_0_88]), c_0_70]), c_0_49]), c_0_50]), c_0_71]), c_0_51]), c_0_59]), c_0_89])).
cnf(c_0_93,plain,(hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,number_number_of(int),hAPP(int,int,bit0,hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),m)),one_one(int))),t)),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,hAPP(int,fun(int,int),plus_plus(int),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),hAPP(int,int,number_number_of(int),hAPP(int,int,bit0,hAPP(int,int,bit0,hAPP(int,int,bit1,pls))))),m)),one_one(int))),zero_zero(int))))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_94,plain,(hAPP(int,int,hAPP(int,fun(int,int),times_times(int),pls),X1)=pls), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_95,negated_conjecture,(~hBOOL(hAPP(int,bool,hAPP(int,fun(int,bool),ord_less(int),hAPP(int,int,hAPP(int,fun(int,int),times_times(int),t),hAPP(int,int,bit1,hAPP(int,int,bit0,m)))),pls))), inference(rw,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_96,plain,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_48]), c_0_70]), c_0_70]), c_0_49]), c_0_50]), c_0_71]), c_0_71]), c_0_51]), c_0_82]), c_0_82]), c_0_83]), c_0_84]), c_0_59]), c_0_85]), c_0_86]), c_0_70]), c_0_71]), c_0_84]), c_0_87]), c_0_48]), c_0_70]), c_0_70]), c_0_49]), c_0_50]), c_0_71]), c_0_71]), c_0_51]), c_0_82]), c_0_82]), c_0_83]), c_0_84]), c_0_59]), c_0_85]), c_0_86]), c_0_70]), c_0_71]), c_0_84]), c_0_31]), c_0_87]), c_0_94]), c_0_95]), ['proof']).
# SZS output end CNFRefutation.
EOF
