% Problem    : PRO006+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n184.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:58:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2n
# and selection function SelectCQIArNTNpEqFirst.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (~(?[X112]:occurrence_of(X112,tptp0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals)).
fof(c_0_1, axiom, (![X106]:(occurrence_of(X106,tptp0)=>?[X107]:?[X108]:((((occurrence_of(X107,tptp4)&root_occ(X107,X106))&occurrence_of(X108,tptp3))&leaf_occ(X108,X106))&next_subocc(X107,X108,tptp0)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_35)).
fof(c_0_2, axiom, (![X5]:![X6]:![X7]:((occurrence_of(X5,X6)&occurrence_of(X5,X7))=>X6=X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_02)).
fof(c_0_3, axiom, (![X100]:![X101]:(root_occ(X100,X101)<=>?[X102]:((occurrence_of(X101,X102)&subactivity_occurrence(X100,X101))&root(X100,X102)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_33)).
fof(c_0_4, axiom, (![X74]:![X75]:((root(X75,X74)&~(atomic(X74)))=>?[X76]:(occurrence_of(X76,X74)&subactivity_occurrence(X75,X76)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_26)).
fof(c_0_5, axiom, (![X109]:![X110]:((((occurrence_of(X110,tptp0)&subactivity_occurrence(X109,X110))&arboreal(X109))&~(leaf_occ(X109,X110)))=>?[X111]:(occurrence_of(X111,tptp1)&next_subocc(X109,X111,tptp0)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_48)).
fof(c_0_6, axiom, (~(atomic(tptp0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_37)).
fof(c_0_7, axiom, (![X17]:![X18]:(occurrence_of(X17,X18)=>(arboreal(X17)<=>atomic(X18)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_07)).
fof(c_0_8, axiom, (atomic(tptp4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_38)).
fof(c_0_9, axiom, (![X61]:![X62]:![X63]:(next_subocc(X61,X62,X63)<=>(min_precedes(X61,X62,X63)&~(?[X64]:(min_precedes(X61,X64,X63)&min_precedes(X64,X62,X63)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_22)).
fof(c_0_10, axiom, (![X1]:![X2]:(occurrence_of(X2,X1)=>(activity(X1)&activity_occurrence(X2)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos)).
fof(c_0_11, axiom, (![X103]:![X104]:(leaf_occ(X103,X104)<=>?[X105]:((occurrence_of(X104,X105)&subactivity_occurrence(X103,X104))&leaf(X103,X105)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_34)).
fof(c_0_12, axiom, (![X3]:(activity_occurrence(X3)=>?[X4]:(activity(X4)&occurrence_of(X3,X4)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_01)).
fof(c_0_13, axiom, (![X57]:![X58]:(leaf(X57,X58)<=>((root(X57,X58)|?[X59]:min_precedes(X59,X57,X58))&~(?[X60]:min_precedes(X57,X60,X58))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_21)).
fof(c_0_14, axiom, (![X70]:![X71]:![X72]:(min_precedes(X71,X72,X70)=>?[X73]:((occurrence_of(X73,X70)&subactivity_occurrence(X71,X73))&subactivity_occurrence(X72,X73)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_25)).
fof(c_0_15, axiom, (![X36]:![X37]:![X38]:(min_precedes(X36,X37,X38)=>~(root(X37,X38)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_14)).
fof(c_0_16, axiom, (![X80]:![X81]:![X82]:![X83]:(((((occurrence_of(X81,X80)&arboreal(X82))&arboreal(X83))&subactivity_occurrence(X82,X81))&subactivity_occurrence(X83,X81))=>((min_precedes(X82,X83,X80)|min_precedes(X83,X82,X80))|X82=X83))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_28)).
fof(c_0_17, axiom, (![X68]:![X69]:(subactivity_occurrence(X68,X69)=>(activity_occurrence(X68)&activity_occurrence(X69)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_24)).
fof(c_0_18, axiom, (atomic(tptp3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_39)).
fof(c_0_19, axiom, (atomic(tptp1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_41)).
fof(c_0_20, axiom, (tptp1!=tptp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_45)).
fof(c_0_21, negated_conjecture, (~(~(?[X112]:occurrence_of(X112,tptp0)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_22, plain, (![X109]:(((((occurrence_of(esk16_1(X109),tptp4)|~occurrence_of(X109,tptp0))&(root_occ(esk16_1(X109),X109)|~occurrence_of(X109,tptp0)))&(occurrence_of(esk17_1(X109),tptp3)|~occurrence_of(X109,tptp0)))&(leaf_occ(esk17_1(X109),X109)|~occurrence_of(X109,tptp0)))&(next_subocc(esk16_1(X109),esk17_1(X109),tptp0)|~occurrence_of(X109,tptp0)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_23, negated_conjecture, (occurrence_of(esk19_0,tptp0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
fof(c_0_24, plain, (![X8]:![X9]:![X10]:((~occurrence_of(X8,X9)|~occurrence_of(X8,X10))|X9=X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
fof(c_0_25, plain, (![X103]:![X104]:![X106]:![X107]:![X108]:((((occurrence_of(X104,esk14_2(X103,X104))|~root_occ(X103,X104))&(subactivity_occurrence(X103,X104)|~root_occ(X103,X104)))&(root(X103,esk14_2(X103,X104))|~root_occ(X103,X104)))&(((~occurrence_of(X107,X108)|~subactivity_occurrence(X106,X107))|~root(X106,X108))|root_occ(X106,X107)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])])).
cnf(c_0_26,plain,(root_occ(esk16_1(X1),X1)|~occurrence_of(X1,tptp0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27,negated_conjecture,(occurrence_of(esk19_0,tptp0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28,plain,(X1=X2|~occurrence_of(X3,X2)|~occurrence_of(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29,plain,(occurrence_of(X2,esk14_2(X1,X2))|~root_occ(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30,negated_conjecture,(root_occ(esk16_1(esk19_0),esk19_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_31, plain, (![X74]:![X75]:((root(X75,X74)&~atomic(X74))=>?[X76]:(occurrence_of(X76,X74)&subactivity_occurrence(X75,X76)))), inference(fof_simplification,[status(thm)],[c_0_4])).
cnf(c_0_32,negated_conjecture,(X1=tptp0|~occurrence_of(esk19_0,X1)), inference(spm,[status(thm)],[c_0_28, c_0_27])).
cnf(c_0_33,negated_conjecture,(occurrence_of(esk19_0,esk14_2(esk16_1(esk19_0),esk19_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_34, plain, (![X109]:![X110]:((((occurrence_of(X110,tptp0)&subactivity_occurrence(X109,X110))&arboreal(X109))&~leaf_occ(X109,X110))=>?[X111]:(occurrence_of(X111,tptp1)&next_subocc(X109,X111,tptp0)))), inference(fof_simplification,[status(thm)],[c_0_5])).
fof(c_0_35, plain, (![X77]:![X78]:((occurrence_of(esk11_2(X77,X78),X77)|(~root(X78,X77)|atomic(X77)))&(subactivity_occurrence(X78,esk11_2(X77,X78))|(~root(X78,X77)|atomic(X77))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])])).
cnf(c_0_36,plain,(root(X1,esk14_2(X1,X2))|~root_occ(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37,negated_conjecture,(esk14_2(esk16_1(esk19_0),esk19_0)=tptp0), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_38, plain, (~atomic(tptp0)), inference(fof_simplification,[status(thm)],[c_0_6])).
fof(c_0_39, plain, (![X112]:![X113]:((occurrence_of(esk18_1(X112),tptp1)|(((~occurrence_of(X113,tptp0)|~subactivity_occurrence(X112,X113))|~arboreal(X112))|leaf_occ(X112,X113)))&(next_subocc(X112,esk18_1(X112),tptp0)|(((~occurrence_of(X113,tptp0)|~subactivity_occurrence(X112,X113))|~arboreal(X112))|leaf_occ(X112,X113))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])])])])).
cnf(c_0_40,plain,(atomic(X1)|occurrence_of(esk11_2(X1,X2),X1)|~root(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41,negated_conjecture,(root(esk16_1(esk19_0),tptp0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_30]), c_0_37])).
cnf(c_0_42,plain,(~atomic(tptp0)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_43, plain, (![X19]:![X20]:(((~arboreal(X19)|atomic(X20))|~occurrence_of(X19,X20))&((~atomic(X20)|arboreal(X19))|~occurrence_of(X19,X20)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_44,plain,(occurrence_of(esk16_1(X1),tptp4)|~occurrence_of(X1,tptp0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_45,plain,(leaf_occ(X1,X2)|next_subocc(X1,esk18_1(X1),tptp0)|~arboreal(X1)|~subactivity_occurrence(X1,X2)|~occurrence_of(X2,tptp0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46,negated_conjecture,(occurrence_of(esk11_2(tptp0,esk16_1(esk19_0)),tptp0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_47,plain,(atomic(X1)|subactivity_occurrence(X2,esk11_2(X1,X2))|~root(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48,plain,(arboreal(X1)|~occurrence_of(X1,X2)|~atomic(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49,negated_conjecture,(occurrence_of(esk16_1(esk19_0),tptp4)), inference(spm,[status(thm)],[c_0_44, c_0_27])).
cnf(c_0_50,plain,(atomic(tptp4)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_51, plain, (![X65]:![X66]:![X67]:![X68]:![X69]:![X70]:![X71]:(((min_precedes(X65,X66,X67)|~next_subocc(X65,X66,X67))&((~min_precedes(X65,X68,X67)|~min_precedes(X68,X66,X67))|~next_subocc(X65,X66,X67)))&(((min_precedes(X69,esk8_3(X69,X70,X71),X71)|~min_precedes(X69,X70,X71))|next_subocc(X69,X70,X71))&((min_precedes(esk8_3(X69,X70,X71),X70,X71)|~min_precedes(X69,X70,X71))|next_subocc(X69,X70,X71))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])])).
cnf(c_0_52,negated_conjecture,(leaf_occ(X1,esk11_2(tptp0,esk16_1(esk19_0)))|next_subocc(X1,esk18_1(X1),tptp0)|~subactivity_occurrence(X1,esk11_2(tptp0,esk16_1(esk19_0)))|~arboreal(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53,negated_conjecture,(subactivity_occurrence(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_41]), c_0_42])).
cnf(c_0_54,negated_conjecture,(arboreal(esk16_1(esk19_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
fof(c_0_55, plain, (![X3]:![X4]:((activity(X3)|~occurrence_of(X4,X3))&(activity_occurrence(X4)|~occurrence_of(X4,X3)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
fof(c_0_56, plain, (![X106]:![X107]:![X109]:![X110]:![X111]:((((occurrence_of(X107,esk15_2(X106,X107))|~leaf_occ(X106,X107))&(subactivity_occurrence(X106,X107)|~leaf_occ(X106,X107)))&(leaf(X106,esk15_2(X106,X107))|~leaf_occ(X106,X107)))&(((~occurrence_of(X110,X111)|~subactivity_occurrence(X109,X110))|~leaf(X109,X111))|leaf_occ(X109,X110)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])])).
cnf(c_0_57,plain,(min_precedes(X1,X2,X3)|~next_subocc(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_58,negated_conjecture,(leaf_occ(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0)))|next_subocc(esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)),tptp0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])])).
fof(c_0_59, plain, (![X5]:((activity(esk1_1(X5))|~activity_occurrence(X5))&(occurrence_of(X5,esk1_1(X5))|~activity_occurrence(X5)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
cnf(c_0_60,plain,(activity_occurrence(X1)|~occurrence_of(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_61, plain, (![X61]:![X62]:![X64]:![X65]:![X66]:![X67]:((((root(X61,X62)|min_precedes(esk6_2(X61,X62),X61,X62))|~leaf(X61,X62))&(~min_precedes(X61,X64,X62)|~leaf(X61,X62)))&(((~root(X65,X66)|min_precedes(X65,esk7_2(X65,X66),X66))|leaf(X65,X66))&((~min_precedes(X67,X65,X66)|min_precedes(X65,esk7_2(X65,X66),X66))|leaf(X65,X66))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])])).
cnf(c_0_62,plain,(leaf(X1,esk15_2(X1,X2))|~leaf_occ(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63,negated_conjecture,(leaf_occ(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0)))|min_precedes(esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)),tptp0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_64,plain,(occurrence_of(X1,esk1_1(X1))|~activity_occurrence(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65,plain,(occurrence_of(X2,esk15_2(X1,X2))|~leaf_occ(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_66,negated_conjecture,(X1=tptp0|~occurrence_of(esk11_2(tptp0,esk16_1(esk19_0)),X1)), inference(spm,[status(thm)],[c_0_28, c_0_46])).
cnf(c_0_67,negated_conjecture,(activity_occurrence(esk11_2(tptp0,esk16_1(esk19_0)))), inference(spm,[status(thm)],[c_0_60, c_0_46])).
cnf(c_0_68,plain,(~leaf(X1,X2)|~min_precedes(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_69,negated_conjecture,(leaf(esk16_1(esk19_0),esk15_2(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0))))|min_precedes(esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)),tptp0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_70,plain,(X1=esk1_1(X2)|~occurrence_of(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_64]), c_0_60])).
cnf(c_0_71,negated_conjecture,(min_precedes(esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)),tptp0)|occurrence_of(esk11_2(tptp0,esk16_1(esk19_0)),esk15_2(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0))))), inference(spm,[status(thm)],[c_0_65, c_0_63])).
cnf(c_0_72,negated_conjecture,(esk1_1(esk11_2(tptp0,esk16_1(esk19_0)))=tptp0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_64]), c_0_67])])).
cnf(c_0_73,plain,(next_subocc(esk16_1(X1),esk17_1(X1),tptp0)|~occurrence_of(X1,tptp0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_74,negated_conjecture,(min_precedes(esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)),tptp0)|~min_precedes(esk16_1(esk19_0),X1,esk15_2(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0))))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75,negated_conjecture,(esk15_2(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0)))=tptp0|min_precedes(esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)),tptp0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])).
cnf(c_0_76,negated_conjecture,(next_subocc(esk16_1(esk19_0),esk17_1(esk19_0),tptp0)), inference(spm,[status(thm)],[c_0_73, c_0_27])).
fof(c_0_77, plain, (![X74]:![X75]:![X76]:(((occurrence_of(esk10_3(X74,X75,X76),X74)|~min_precedes(X75,X76,X74))&(subactivity_occurrence(X75,esk10_3(X74,X75,X76))|~min_precedes(X75,X76,X74)))&(subactivity_occurrence(X76,esk10_3(X74,X75,X76))|~min_precedes(X75,X76,X74)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).
cnf(c_0_78,negated_conjecture,(min_precedes(esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)),tptp0)|~min_precedes(esk16_1(esk19_0),X1,tptp0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_79,negated_conjecture,(min_precedes(esk16_1(esk19_0),esk17_1(esk19_0),tptp0)), inference(spm,[status(thm)],[c_0_57, c_0_76])).
cnf(c_0_80,plain,(leaf_occ(X1,X2)|occurrence_of(esk18_1(X1),tptp1)|~arboreal(X1)|~subactivity_occurrence(X1,X2)|~occurrence_of(X2,tptp0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_81,plain,(occurrence_of(esk10_3(X3,X1,X2),X3)|~min_precedes(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_82,negated_conjecture,(min_precedes(esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)),tptp0)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_83,negated_conjecture,(leaf_occ(X1,esk11_2(tptp0,esk16_1(esk19_0)))|occurrence_of(esk18_1(X1),tptp1)|~subactivity_occurrence(X1,esk11_2(tptp0,esk16_1(esk19_0)))|~arboreal(X1)), inference(spm,[status(thm)],[c_0_80, c_0_46])).
cnf(c_0_84,plain,(leaf_occ(esk17_1(X1),X1)|~occurrence_of(X1,tptp0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_85,negated_conjecture,(occurrence_of(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))),tptp0)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_86,negated_conjecture,(leaf_occ(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0)))|occurrence_of(esk18_1(esk16_1(esk19_0)),tptp1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_53]), c_0_54])])).
fof(c_0_87, plain, (![X36]:![X37]:![X38]:(min_precedes(X36,X37,X38)=>~root(X37,X38))), inference(fof_simplification,[status(thm)],[c_0_15])).
cnf(c_0_88,negated_conjecture,(leaf_occ(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_89,negated_conjecture,(leaf(esk16_1(esk19_0),esk15_2(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0))))|occurrence_of(esk18_1(esk16_1(esk19_0)),tptp1)), inference(spm,[status(thm)],[c_0_62, c_0_86])).
cnf(c_0_90,negated_conjecture,(occurrence_of(esk11_2(tptp0,esk16_1(esk19_0)),esk15_2(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0))))|occurrence_of(esk18_1(esk16_1(esk19_0)),tptp1)), inference(spm,[status(thm)],[c_0_65, c_0_86])).
fof(c_0_91, plain, (![X84]:![X85]:![X86]:![X87]:(((((~occurrence_of(X85,X84)|~arboreal(X86))|~arboreal(X87))|~subactivity_occurrence(X86,X85))|~subactivity_occurrence(X87,X85))|((min_precedes(X86,X87,X84)|min_precedes(X87,X86,X84))|X86=X87))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).
fof(c_0_92, plain, (![X70]:![X71]:((activity_occurrence(X70)|~subactivity_occurrence(X70,X71))&(activity_occurrence(X71)|~subactivity_occurrence(X70,X71)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
fof(c_0_93, plain, (![X39]:![X40]:![X41]:(~min_precedes(X39,X40,X41)|~root(X40,X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_87])])).
cnf(c_0_94,plain,(occurrence_of(esk17_1(X1),tptp3)|~occurrence_of(X1,tptp0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_95,negated_conjecture,(leaf(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk15_2(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))), inference(spm,[status(thm)],[c_0_62, c_0_88])).
cnf(c_0_96,negated_conjecture,(occurrence_of(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))),esk15_2(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))), inference(spm,[status(thm)],[c_0_65, c_0_88])).
cnf(c_0_97,negated_conjecture,(esk1_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))=tptp0), inference(spm,[status(thm)],[c_0_70, c_0_85])).
cnf(c_0_98,negated_conjecture,(occurrence_of(esk18_1(esk16_1(esk19_0)),tptp1)|~min_precedes(esk16_1(esk19_0),X1,esk15_2(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0))))), inference(spm,[status(thm)],[c_0_68, c_0_89])).
cnf(c_0_99,negated_conjecture,(esk15_2(esk16_1(esk19_0),esk11_2(tptp0,esk16_1(esk19_0)))=tptp0|occurrence_of(esk18_1(esk16_1(esk19_0)),tptp1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_90]), c_0_72])).
cnf(c_0_100,plain,(X1=X2|min_precedes(X2,X1,X3)|min_precedes(X1,X2,X3)|~subactivity_occurrence(X2,X4)|~subactivity_occurrence(X1,X4)|~arboreal(X2)|~arboreal(X1)|~occurrence_of(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_101,plain,(activity_occurrence(X2)|~subactivity_occurrence(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_102,plain,(subactivity_occurrence(X1,esk10_3(X3,X1,X2))|~min_precedes(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_103,plain,(~root(X1,X2)|~min_precedes(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_104,negated_conjecture,(root_occ(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))), inference(spm,[status(thm)],[c_0_26, c_0_85])).
cnf(c_0_105,negated_conjecture,(next_subocc(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp0)), inference(spm,[status(thm)],[c_0_73, c_0_85])).
cnf(c_0_106,plain,(subactivity_occurrence(X1,X2)|~leaf_occ(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_107,negated_conjecture,(occurrence_of(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp3)), inference(spm,[status(thm)],[c_0_94, c_0_85])).
cnf(c_0_108,plain,(atomic(tptp3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_109,negated_conjecture,(~min_precedes(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),X1,esk15_2(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))), inference(spm,[status(thm)],[c_0_68, c_0_95])).
cnf(c_0_110,negated_conjecture,(esk15_2(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))=tptp0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_96]), c_0_97])).
cnf(c_0_111,negated_conjecture,(occurrence_of(esk18_1(esk16_1(esk19_0)),tptp1)|~min_precedes(esk16_1(esk19_0),X1,tptp0)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_112,plain,(X1=X2|min_precedes(X1,X2,esk1_1(X3))|min_precedes(X2,X1,esk1_1(X3))|~subactivity_occurrence(X2,X3)|~subactivity_occurrence(X1,X3)|~arboreal(X2)|~arboreal(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_64]), c_0_101])).
cnf(c_0_113,negated_conjecture,(subactivity_occurrence(esk16_1(esk19_0),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))), inference(spm,[status(thm)],[c_0_102, c_0_82])).
cnf(c_0_114,negated_conjecture,(~min_precedes(X1,esk16_1(esk19_0),tptp0)), inference(spm,[status(thm)],[c_0_103, c_0_41])).
cnf(c_0_115,plain,(subactivity_occurrence(X1,X2)|~root_occ(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_116,negated_conjecture,(occurrence_of(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp4)), inference(spm,[status(thm)],[c_0_44, c_0_85])).
cnf(c_0_117,negated_conjecture,(root(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk14_2(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))), inference(spm,[status(thm)],[c_0_36, c_0_104])).
cnf(c_0_118,negated_conjecture,(occurrence_of(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))),esk14_2(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))), inference(spm,[status(thm)],[c_0_29, c_0_104])).
cnf(c_0_119,negated_conjecture,(min_precedes(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp0)), inference(spm,[status(thm)],[c_0_57, c_0_105])).
cnf(c_0_120,negated_conjecture,(subactivity_occurrence(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))), inference(spm,[status(thm)],[c_0_106, c_0_88])).
cnf(c_0_121,negated_conjecture,(arboreal(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_107]), c_0_108])])).
cnf(c_0_122,negated_conjecture,(~min_precedes(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),X1,tptp0)), inference(rw,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_123,plain,(subactivity_occurrence(X2,esk10_3(X3,X1,X2))|~min_precedes(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_124,negated_conjecture,(occurrence_of(esk18_1(esk16_1(esk19_0)),tptp1)), inference(spm,[status(thm)],[c_0_111, c_0_79])).
cnf(c_0_125,plain,(atomic(tptp1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_126,negated_conjecture,(X1=esk16_1(esk19_0)|min_precedes(esk16_1(esk19_0),X1,tptp0)|~subactivity_occurrence(X1,esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))|~arboreal(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_97]), c_0_97]), c_0_54])]), c_0_114])).
cnf(c_0_127,negated_conjecture,(subactivity_occurrence(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))), inference(spm,[status(thm)],[c_0_115, c_0_104])).
cnf(c_0_128,negated_conjecture,(arboreal(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_116]), c_0_50])])).
cnf(c_0_129,negated_conjecture,(~min_precedes(X1,esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk14_2(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))), inference(spm,[status(thm)],[c_0_103, c_0_117])).
cnf(c_0_130,negated_conjecture,(esk14_2(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))=tptp0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_118]), c_0_97])).
cnf(c_0_131,negated_conjecture,(occurrence_of(esk10_3(tptp0,esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))),tptp0)), inference(spm,[status(thm)],[c_0_81, c_0_119])).
cnf(c_0_132,negated_conjecture,(X1=esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))|min_precedes(X1,esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp0)|~subactivity_occurrence(X1,esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))|~arboreal(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_120]), c_0_97]), c_0_97]), c_0_121])]), c_0_122])).
cnf(c_0_133,negated_conjecture,(subactivity_occurrence(esk18_1(esk16_1(esk19_0)),esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))), inference(spm,[status(thm)],[c_0_123, c_0_82])).
cnf(c_0_134,negated_conjecture,(arboreal(esk18_1(esk16_1(esk19_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_124]), c_0_125])])).
cnf(c_0_135,negated_conjecture,(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))=esk16_1(esk19_0)|min_precedes(esk16_1(esk19_0),esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_128])])).
cnf(c_0_136,negated_conjecture,(~min_precedes(X1,esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp0)), inference(rw,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_137,negated_conjecture,(occurrence_of(esk17_1(esk10_3(tptp0,esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))),tptp3)), inference(spm,[status(thm)],[c_0_94, c_0_131])).
cnf(c_0_138,plain,(~next_subocc(X1,X2,X3)|~min_precedes(X4,X2,X3)|~min_precedes(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_139,negated_conjecture,(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))=esk18_1(esk16_1(esk19_0))|min_precedes(esk18_1(esk16_1(esk19_0)),esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_134])])).
cnf(c_0_140,negated_conjecture,(esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))=esk16_1(esk19_0)), inference(sr,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_141,negated_conjecture,(esk1_1(esk17_1(esk10_3(tptp0,esk16_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))))=tptp3), inference(spm,[status(thm)],[c_0_70, c_0_137])).
cnf(c_0_142,negated_conjecture,(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))=esk18_1(esk16_1(esk19_0))|~next_subocc(X1,esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp0)|~min_precedes(X1,esk18_1(esk16_1(esk19_0)),tptp0)), inference(spm,[status(thm)],[c_0_138, c_0_139])).
cnf(c_0_143,negated_conjecture,(next_subocc(esk16_1(esk19_0),esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))),tptp0)), inference(rw,[status(thm)],[c_0_105, c_0_140])).
cnf(c_0_144,negated_conjecture,(esk1_1(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0)))))))=tptp3), inference(rw,[status(thm)],[c_0_141, c_0_140])).
cnf(c_0_145,negated_conjecture,(esk17_1(esk10_3(tptp0,esk16_1(esk19_0),esk18_1(esk16_1(esk19_0))))=esk18_1(esk16_1(esk19_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_82]), c_0_143])])).
cnf(c_0_146,negated_conjecture,(esk1_1(esk18_1(esk16_1(esk19_0)))=tptp1), inference(spm,[status(thm)],[c_0_70, c_0_124])).
cnf(c_0_147,plain,(tptp1!=tptp3), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_148,negated_conjecture,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_144, c_0_145]), c_0_145]), c_0_146]), c_0_147]), ['proof']).
# SZS output end CNFRefutation.
EOF
