% Problem    : SET776+4 : TPTP v6.1.0. Released v2.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n114.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:16:56 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X4]:![X8]:![X7]:((pre_order(X8,X4)&![X1]:![X2]:((member(X1,X4)&member(X2,X4))=>(apply(X7,X1,X2)<=>(apply(X8,X1,X2)&apply(X8,X2,X1)))))=>![X9]:![X10]:![X11]:![X12]:((((member(X9,X4)&member(X10,X4))&member(X11,X4))&member(X12,X4))=>(((apply(X7,X9,X11)&apply(X7,X10,X12))&apply(X8,X9,X10))=>apply(X8,X11,X12))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', thIII12)).
fof(c_0_1, axiom, (![X7]:![X4]:(pre_order(X7,X4)<=>(![X3]:(member(X3,X4)=>apply(X7,X3,X3))&![X3]:![X5]:![X6]:(((member(X3,X4)&member(X5,X4))&member(X6,X4))=>((apply(X7,X3,X5)&apply(X7,X5,X6))=>apply(X7,X3,X6)))))), file('/export/starexec/sandbox/benchmark/Axioms/SET006+2.ax', pre_order)).
fof(c_0_2, negated_conjecture, (~(![X4]:![X8]:![X7]:((pre_order(X8,X4)&![X1]:![X2]:((member(X1,X4)&member(X2,X4))=>(apply(X7,X1,X2)<=>(apply(X8,X1,X2)&apply(X8,X2,X1)))))=>![X9]:![X10]:![X11]:![X12]:((((member(X9,X4)&member(X10,X4))&member(X11,X4))&member(X12,X4))=>(((apply(X7,X9,X11)&apply(X7,X10,X12))&apply(X8,X9,X10))=>apply(X8,X11,X12)))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_3, negated_conjecture, (![X16]:![X17]:((pre_order(esk16_0,esk15_0)&((((apply(esk16_0,X16,X17)|~apply(esk17_0,X16,X17))|(~member(X16,esk15_0)|~member(X17,esk15_0)))&((apply(esk16_0,X17,X16)|~apply(esk17_0,X16,X17))|(~member(X16,esk15_0)|~member(X17,esk15_0))))&(((~apply(esk16_0,X16,X17)|~apply(esk16_0,X17,X16))|apply(esk17_0,X16,X17))|(~member(X16,esk15_0)|~member(X17,esk15_0)))))&((((member(esk18_0,esk15_0)&member(esk19_0,esk15_0))&member(esk20_0,esk15_0))&member(esk21_0,esk15_0))&(((apply(esk17_0,esk18_0,esk20_0)&apply(esk17_0,esk19_0,esk21_0))&apply(esk16_0,esk18_0,esk19_0))&~apply(esk16_0,esk20_0,esk21_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
fof(c_0_4, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:((((~member(X10,X9)|apply(X8,X10,X10))|~pre_order(X8,X9))&((((~member(X11,X9)|~member(X12,X9))|~member(X13,X9))|((~apply(X8,X11,X12)|~apply(X8,X12,X13))|apply(X8,X11,X13)))|~pre_order(X8,X9)))&((((((member(esk12_2(X14,X15),X15)|member(esk11_2(X14,X15),X15))|pre_order(X14,X15))&((member(esk13_2(X14,X15),X15)|member(esk11_2(X14,X15),X15))|pre_order(X14,X15)))&((member(esk14_2(X14,X15),X15)|member(esk11_2(X14,X15),X15))|pre_order(X14,X15)))&((((apply(X14,esk12_2(X14,X15),esk13_2(X14,X15))|member(esk11_2(X14,X15),X15))|pre_order(X14,X15))&((apply(X14,esk13_2(X14,X15),esk14_2(X14,X15))|member(esk11_2(X14,X15),X15))|pre_order(X14,X15)))&((~apply(X14,esk12_2(X14,X15),esk14_2(X14,X15))|member(esk11_2(X14,X15),X15))|pre_order(X14,X15))))&(((((member(esk12_2(X14,X15),X15)|~apply(X14,esk11_2(X14,X15),esk11_2(X14,X15)))|pre_order(X14,X15))&((member(esk13_2(X14,X15),X15)|~apply(X14,esk11_2(X14,X15),esk11_2(X14,X15)))|pre_order(X14,X15)))&((member(esk14_2(X14,X15),X15)|~apply(X14,esk11_2(X14,X15),esk11_2(X14,X15)))|pre_order(X14,X15)))&((((apply(X14,esk12_2(X14,X15),esk13_2(X14,X15))|~apply(X14,esk11_2(X14,X15),esk11_2(X14,X15)))|pre_order(X14,X15))&((apply(X14,esk13_2(X14,X15),esk14_2(X14,X15))|~apply(X14,esk11_2(X14,X15),esk11_2(X14,X15)))|pre_order(X14,X15)))&((~apply(X14,esk12_2(X14,X15),esk14_2(X14,X15))|~apply(X14,esk11_2(X14,X15),esk11_2(X14,X15)))|pre_order(X14,X15))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])])).
cnf(c_0_5,negated_conjecture,(apply(esk16_0,X2,X1)|~member(X1,esk15_0)|~member(X2,esk15_0)|~apply(esk17_0,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6,negated_conjecture,(apply(esk17_0,esk19_0,esk21_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_7,negated_conjecture,(member(esk19_0,esk15_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_8,negated_conjecture,(member(esk21_0,esk15_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_9,plain,(apply(X1,X3,X4)|~pre_order(X1,X2)|~apply(X1,X5,X4)|~apply(X1,X3,X5)|~member(X4,X2)|~member(X5,X2)|~member(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10,negated_conjecture,(apply(esk16_0,esk19_0,esk21_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7]), c_0_8])])).
cnf(c_0_11,negated_conjecture,(apply(esk16_0,X1,esk21_0)|~pre_order(esk16_0,X2)|~apply(esk16_0,X1,esk19_0)|~member(esk19_0,X2)|~member(esk21_0,X2)|~member(X1,X2)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12,negated_conjecture,(pre_order(esk16_0,esk15_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_13,negated_conjecture,(apply(esk16_0,esk18_0,esk19_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_14,negated_conjecture,(apply(esk16_0,X1,X2)|~member(X1,esk15_0)|~member(X2,esk15_0)|~apply(esk17_0,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_15,negated_conjecture,(apply(esk17_0,esk18_0,esk20_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_16,negated_conjecture,(member(esk18_0,esk15_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_17,negated_conjecture,(member(esk20_0,esk15_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_18,negated_conjecture,(~apply(esk16_0,esk20_0,esk21_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_19,negated_conjecture,(apply(esk16_0,X1,esk21_0)|~apply(esk16_0,X1,esk19_0)|~member(X1,esk15_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_7]), c_0_8])])).
cnf(c_0_20,negated_conjecture,(apply(esk16_0,X1,esk19_0)|~pre_order(esk16_0,X2)|~apply(esk16_0,X1,esk18_0)|~member(esk18_0,X2)|~member(esk19_0,X2)|~member(X1,X2)), inference(spm,[status(thm)],[c_0_9, c_0_13])).
cnf(c_0_21,negated_conjecture,(apply(esk16_0,esk20_0,esk18_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])])).
cnf(c_0_22,negated_conjecture,(~apply(esk16_0,esk20_0,esk19_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_17])])).
cnf(c_0_23,negated_conjecture,(~pre_order(esk16_0,X1)|~member(esk18_0,X1)|~member(esk19_0,X1)|~member(esk20_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
cnf(c_0_24,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_12]), c_0_16]), c_0_7]), c_0_17])]), ['proof']).
# SZS output end CNFRefutation.
EOF
