% Problem    : GEO288+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n110.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:43:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X141]:![X142]:((?[X143]:(X142=X143&rpoint(X143))&?[X144]:(X141=X144&rpoint(X144)))=>vf(X141,X142)=vf(X142,X141))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(166), 0), imp(cond(axiom(166), 0)))')).
fof(c_0_1, axiom, (((((((((((((~(?[X1]:(((ron(vd1165,X1)&ron(vd1164,X1))&ron(vd1163,X1))&rline(X1)))&vangle(vd1157,vd1156,vd1158)=vangle(vd1164,vd1163,vd1165))&vf(vd1156,vd1158)=vf(vd1163,vd1165))&vf(vd1156,vd1157)=vf(vd1163,vd1164))&?[X2]:(vd1165=X2&rpoint(X2)))&?[X3]:(vd1164=X3&rpoint(X3)))&?[X4]:(vd1163=X4&rpoint(X4)))&vd1157!=vd1158)&vd1156!=vd1158)&vd1156!=vd1157)&?[X5]:(vd1158=X5&rpoint(X5)))&?[X6]:(vd1157=X6&rpoint(X6)))&?[X7]:(vd1156=X7&rpoint(X7)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(neg(neg(conjunct2(conjunct2(conjunct2(plural(comma_conjunct2(265))))))), and(holds(conjunct1(conjunct2(conjunct2(plural(comma_conjunct2(265))))), 1172, 0), and(holds(conjunct1(conjunct2(plural(comma_conjunct2(265)))), 1171, 0), and(holds(conjunct1(plural(comma_conjunct2(265))), 1170, 0), and(qe(s3(plural(comma_conjunct2(265)))), and(qe(s2(plural(comma_conjunct2(265)))), and(qe(s1(plural(comma_conjunct2(265)))), and(pred(comma_conjunct1(265), 9), and(pred(comma_conjunct1(265), 8), and(pred(comma_conjunct1(265), 7), and(qe(s3(plural(265))), and(qe(s2(plural(265))), qe(s1(plural(265)))))))))))))))')).
fof(c_0_2, axiom, (![X196]:![X197]:![X198]:![X199]:((((ron(X197,X199)&rcenter(X198,X199))&?[X200]:(X197=X200&rpoint(X200)))&?[X201]:(X196=X201&rpoint(X201)))=>(rless(vf(X198,X196),vf(X198,X197))<=>rinside(X196,X199)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(186), 0), imp(cond(axiom(186), 0)))')).
fof(c_0_3, axiom, (![X145]:![X146]:((?[X147]:(X146=X147&rpoint(X147))&?[X148]:(X145=X148&rpoint(X148)))=>?[X149]:(vf(X145,X146)=X149&rreal(X149)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(160), 0), imp(cond(axiom(160), 0)))')).
fof(c_0_4, axiom, (![X226]:![X227]:![X228]:![X229]:((((X228=X229&rreal(X229))&X226=X227)&rreal(X227))=>((rless(X226,X228)|X226=X228)|rless(X228,X226)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(145), 0), imp(cond(axiom(145), 0)))')).
fof(c_0_5, axiom, (![X517]:![X518]:(((X517!=X518&?[X519]:(X518=X519&rpoint(X519)))&?[X520]:(X517=X520&rpoint(X520)))=>?[X521]:((ron(X518,X521)&rcenter(X517,X521))&rcircle(X521)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(53), 0), imp(cond(axiom(53), 0)))')).
fof(c_0_6, axiom, (![X442]:![X443]:![X444]:![X445]:(((((rinside(X442,X444)&X444=X445)&rcircle(X445))&X442=X443)&rpoint(X443))=>~(ron(X442,X444)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(79), 0), imp(cond(axiom(79), 0)))')).
fof(c_0_7, axiom, (![X51]:![X52]:(((X51!=X52&?[X53]:(X52=X53&rpoint(X53)))&?[X54]:(X51=X54&rpoint(X54)))=>?[X55]:((ron(X52,X55)&ron(X51,X55))&rline(X55)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(51), 0), imp(cond(axiom(51), 0)))')).
fof(c_0_8, axiom, (![X47]:![X48]:((?[X49]:(X48=X49&rpoint(X49))&?[X50]:(X47=X50&rpoint(X50)))=>(vf(X47,X48)=v0<=>X47=X48))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(162), 0), imp(cond(axiom(162), 0)))')).
fof(c_0_9, axiom, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:(((((((((((((~(?[X14]:(((ron(X13,X14)&ron(X12,X14))&ron(X11,X14))&rline(X14)))&vangle(X9,X8,X10)=vangle(X12,X11,X13))&vf(X8,X10)=vf(X11,X13))&vf(X8,X9)=vf(X11,X12))&?[X15]:(X13=X15&rpoint(X15)))&?[X16]:(X12=X16&rpoint(X16)))&?[X17]:(X11=X17&rpoint(X17)))&X9!=X10)&X8!=X10)&X8!=X9)&?[X18]:(X10=X18&rpoint(X18)))&?[X19]:(X9=X19&rpoint(X19)))&?[X20]:(X8=X20&rpoint(X20)))=>vangle(X11,X12,X13)=vangle(X8,X9,X10))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(proof(267), 0), 1)')).
fof(c_0_10, axiom, (rpoint(vd14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(axiom(5), 0)')).
fof(c_0_11, conjecture, (vangle(vd1156,vd1157,vd1158)=vangle(vd1163,vd1164,vd1165)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct1(conjunct2(266)), 1176, 0)')).
fof(c_0_12, plain, (![X145]:![X146]:![X147]:![X148]:(((X146!=X147|~rpoint(X147))|(X145!=X148|~rpoint(X148)))|vf(X145,X146)=vf(X146,X145))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])).
fof(c_0_13, plain, (![X8]:(((((((((((((((~ron(vd1165,X8)|~ron(vd1164,X8))|~ron(vd1163,X8))|~rline(X8))&vangle(vd1157,vd1156,vd1158)=vangle(vd1164,vd1163,vd1165))&vf(vd1156,vd1158)=vf(vd1163,vd1165))&vf(vd1156,vd1157)=vf(vd1163,vd1164))&(vd1165=esk1_0&rpoint(esk1_0)))&(vd1164=esk2_0&rpoint(esk2_0)))&(vd1163=esk3_0&rpoint(esk3_0)))&vd1157!=vd1158)&vd1156!=vd1158)&vd1156!=vd1157)&(vd1158=esk4_0&rpoint(esk4_0)))&(vd1157=esk5_0&rpoint(esk5_0)))&(vd1156=esk6_0&rpoint(esk6_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_14, plain, (![X202]:![X203]:![X204]:![X205]:![X206]:![X207]:(((~rless(vf(X204,X202),vf(X204,X203))|rinside(X202,X205))|(((~ron(X203,X205)|~rcenter(X204,X205))|(X203!=X206|~rpoint(X206)))|(X202!=X207|~rpoint(X207))))&((~rinside(X202,X205)|rless(vf(X204,X202),vf(X204,X203)))|(((~ron(X203,X205)|~rcenter(X204,X205))|(X203!=X206|~rpoint(X206)))|(X202!=X207|~rpoint(X207)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_15,plain,(vf(X1,X2)=vf(X2,X1)|~rpoint(X3)|X1!=X3|~rpoint(X4)|X2!=X4), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(rpoint(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17,plain,(vd1163=esk3_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(rpoint(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19,plain,(vd1165=esk1_0), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_20, plain, (![X150]:![X151]:![X152]:![X153]:((vf(X150,X151)=esk54_2(X150,X151)|((X151!=X152|~rpoint(X152))|(X150!=X153|~rpoint(X153))))&(rreal(esk54_2(X150,X151))|((X151!=X152|~rpoint(X152))|(X150!=X153|~rpoint(X153)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_21,plain,(rpoint(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,plain,(vd1164=esk2_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23,plain,(rinside(X2,X6)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rcenter(X5,X6)|~ron(X4,X6)|~rless(vf(X5,X2),vf(X5,X4))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24,plain,(vf(vd1156,vd1158)=vf(vd1163,vd1165)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25,plain,(vf(X1,X2)=vf(X2,X1)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_15])])).
cnf(c_0_26,plain,(rpoint(vd1163)), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_27,plain,(rpoint(vd1165)), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_28,plain,(rreal(esk54_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29,plain,(vf(X2,X4)=esk54_2(X2,X4)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30,plain,(vf(vd1156,vd1157)=vf(vd1163,vd1164)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31,plain,(rpoint(vd1164)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_32,plain,(rinside(X1,X2)|~rcenter(X3,X2)|~rless(vf(X3,X1),vf(X3,X4))|~rpoint(X4)|~rpoint(X1)|~ron(X4,X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_23])])).
cnf(c_0_33,plain,(vf(vd1165,vd1163)=vf(vd1156,vd1158)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27])])).
cnf(c_0_34,plain,(rpoint(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35,plain,(vd1158=esk4_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_36,plain,(rpoint(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_37,plain,(vd1156=esk6_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_38,plain,(rreal(esk54_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_28])])).
cnf(c_0_39,plain,(esk54_2(X1,X2)=vf(X1,X2)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_29])])).
cnf(c_0_40,plain,(vf(vd1164,vd1163)=vf(vd1156,vd1157)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_25]), c_0_26]), c_0_31])])).
cnf(c_0_41,plain,(rpoint(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_42,plain,(vd1157=esk5_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_43,plain,(rinside(X1,X2)|~rcenter(vd1165,X2)|~rless(vf(vd1165,X1),vf(vd1156,vd1158))|~rpoint(X1)|~ron(vd1163,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_26])])).
cnf(c_0_44,plain,(rpoint(vd1158)), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_45,plain,(rpoint(vd1156)), inference(rw,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_46, plain, (![X230]:![X231]:![X232]:![X233]:((((X232!=X233|~rreal(X233))|X230!=X231)|~rreal(X231))|((rless(X230,X232)|X230=X232)|rless(X232,X230)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_47,plain,(rreal(vf(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_48,plain,(rinside(vd1163,X1)|~rcenter(vd1165,X1)|~rless(vf(vd1156,vd1158),vf(vd1165,X2))|~rpoint(X2)|~ron(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_26])])).
cnf(c_0_49,plain,(rinside(X1,X2)|~rcenter(vd1164,X2)|~rless(vf(vd1164,X1),vf(vd1156,vd1157))|~rpoint(X1)|~ron(vd1163,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_40]), c_0_26])])).
cnf(c_0_50,plain,(rpoint(vd1157)), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51,plain,(rinside(vd1163,X1)|~rcenter(vd1164,X1)|~rless(vf(vd1156,vd1157),vf(vd1164,X2))|~rpoint(X2)|~ron(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_40]), c_0_26])])).
cnf(c_0_52,plain,(rinside(X1,X2)|~rcenter(vd1165,X2)|~rless(vf(vd1165,X1),vf(vd1158,vd1156))|~rpoint(X1)|~ron(vd1163,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_25]), c_0_44]), c_0_45])])).
cnf(c_0_53,plain,(rless(X1,X2)|X2=X1|rless(X2,X1)|~rreal(X3)|X2!=X3|~rreal(X4)|X1!=X4), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_54,plain,(rreal(vf(vd1156,vd1158))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_24]), c_0_27]), c_0_26])])).
cnf(c_0_55,plain,(rinside(vd1163,X1)|~rcenter(vd1165,X1)|~rless(vf(vd1158,vd1156),vf(vd1165,X2))|~rpoint(X2)|~ron(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_25]), c_0_44]), c_0_45])])).
fof(c_0_56, plain, (![X522]:![X523]:![X524]:![X525]:(((ron(X523,esk69_2(X522,X523))|((X522=X523|(X523!=X524|~rpoint(X524)))|(X522!=X525|~rpoint(X525))))&(rcenter(X522,esk69_2(X522,X523))|((X522=X523|(X523!=X524|~rpoint(X524)))|(X522!=X525|~rpoint(X525)))))&(rcircle(esk69_2(X522,X523))|((X522=X523|(X523!=X524|~rpoint(X524)))|(X522!=X525|~rpoint(X525)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
cnf(c_0_57,plain,(rinside(X1,X2)|~rcenter(vd1164,X2)|~rless(vf(vd1164,X1),vf(vd1157,vd1156))|~rpoint(X1)|~ron(vd1163,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_25]), c_0_50]), c_0_45])])).
cnf(c_0_58,plain,(rreal(vf(vd1156,vd1157))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_30]), c_0_31]), c_0_26])])).
cnf(c_0_59,plain,(rinside(vd1163,X1)|~rcenter(vd1164,X1)|~rless(vf(vd1157,vd1156),vf(vd1164,X2))|~rpoint(X2)|~ron(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_25]), c_0_50]), c_0_45])])).
fof(c_0_60, plain, (![X442]:![X443]:![X444]:![X445]:(((((rinside(X442,X444)&X444=X445)&rcircle(X445))&X442=X443)&rpoint(X443))=>~ron(X442,X444))), inference(fof_simplification,[status(thm)],[c_0_6])).
cnf(c_0_61,plain,(rinside(vd1163,X1)|~rcenter(vd1165,X1)|~rless(vf(vd1156,vd1158),vf(vd1158,vd1156))|~ron(vd1163,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_33]), c_0_26])])).
cnf(c_0_62,plain,(X1=X2|rless(X2,X1)|rless(X1,X2)|~rreal(X1)|~rreal(X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_53])])).
cnf(c_0_63,plain,(rreal(vf(vd1158,vd1156))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_25]), c_0_45]), c_0_44])])).
cnf(c_0_64,plain,(rinside(vd1163,X1)|~rcenter(vd1165,X1)|~rless(vf(vd1158,vd1156),vf(vd1156,vd1158))|~ron(vd1163,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_33]), c_0_26])])).
cnf(c_0_65,plain,(X2=X4|rcenter(X2,esk69_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_66,plain,(rinside(vd1163,X1)|~rcenter(vd1164,X1)|~rless(vf(vd1156,vd1157),vf(vd1157,vd1156))|~ron(vd1163,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_40]), c_0_26])])).
cnf(c_0_67,plain,(rreal(vf(vd1157,vd1156))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_25]), c_0_45]), c_0_50])])).
cnf(c_0_68,plain,(rinside(vd1163,X1)|~rcenter(vd1164,X1)|~rless(vf(vd1157,vd1156),vf(vd1156,vd1157))|~ron(vd1163,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_40]), c_0_26])])).
fof(c_0_69, plain, (![X56]:![X57]:![X58]:![X59]:(((ron(X57,esk50_2(X56,X57))|((X56=X57|(X57!=X58|~rpoint(X58)))|(X56!=X59|~rpoint(X59))))&(ron(X56,esk50_2(X56,X57))|((X56=X57|(X57!=X58|~rpoint(X58)))|(X56!=X59|~rpoint(X59)))))&(rline(esk50_2(X56,X57))|((X56=X57|(X57!=X58|~rpoint(X58)))|(X56!=X59|~rpoint(X59)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
fof(c_0_70, plain, (![X446]:![X447]:![X448]:![X449]:(((((~rinside(X446,X448)|X448!=X449)|~rcircle(X449))|X446!=X447)|~rpoint(X447))|~ron(X446,X448))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_60])])])])).
cnf(c_0_71,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|rinside(vd1163,X1)|~rcenter(vd1165,X1)|~ron(vd1163,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_54])]), c_0_64])).
cnf(c_0_72,plain,(X1=X2|rcenter(X1,esk69_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_65])])).
cnf(c_0_73,plain,(X2=X4|ron(X4,esk69_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_74,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)|rinside(vd1163,X1)|~rcenter(vd1164,X1)|~ron(vd1163,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_62]), c_0_67]), c_0_58])]), c_0_68])).
cnf(c_0_75,plain,(X2=X4|ron(X2,esk50_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_76,plain,(~ron(X1,X2)|~rpoint(X3)|X1!=X3|~rcircle(X4)|X2!=X4|~rinside(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_77,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vd1165=X1|rinside(vd1163,esk69_2(vd1165,X1))|~rpoint(X1)|~ron(vd1163,esk69_2(vd1165,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_27])])).
cnf(c_0_78,plain,(X1=X2|ron(X2,esk69_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_73])])).
cnf(c_0_79,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)|vd1164=X1|rinside(vd1163,esk69_2(vd1164,X1))|~rpoint(X1)|~ron(vd1163,esk69_2(vd1164,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_72]), c_0_31])])).
cnf(c_0_80,plain,(~rline(X1)|~ron(vd1163,X1)|~ron(vd1164,X1)|~ron(vd1165,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_81,plain,(X1=X2|ron(X1,esk50_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_75])])).
cnf(c_0_82,plain,(X2=X4|rline(esk50_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_83,plain,(~rcircle(X1)|~rinside(X2,X1)|~rpoint(X2)|~ron(X2,X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_76])])).
cnf(c_0_84,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vd1163=vd1165|rinside(vd1163,esk69_2(vd1165,vd1163))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_26]), c_0_27])])).
cnf(c_0_85,plain,(X2=X4|rcircle(esk69_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_86,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)|vd1163=vd1164|rinside(vd1163,esk69_2(vd1164,vd1163))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_78]), c_0_26]), c_0_31])])).
cnf(c_0_87,plain,(vd1163=X1|~rpoint(X1)|~rline(esk50_2(vd1163,X1))|~ron(vd1164,esk50_2(vd1163,X1))|~ron(vd1165,esk50_2(vd1163,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_26])])).
cnf(c_0_88,plain,(X1=X2|rline(esk50_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_82])])).
cnf(c_0_89,plain,(X2=X4|ron(X4,esk50_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_90,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vd1163=vd1165|~rcircle(esk69_2(vd1165,vd1163))|~ron(vd1163,esk69_2(vd1165,vd1163))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_26])])).
cnf(c_0_91,plain,(X1=X2|rcircle(esk69_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_85])])).
fof(c_0_92, plain, (![X51]:![X52]:![X53]:![X54]:(((vf(X51,X52)!=v0|X51=X52)|((X52!=X53|~rpoint(X53))|(X51!=X54|~rpoint(X54))))&((X51!=X52|vf(X51,X52)=v0)|((X52!=X53|~rpoint(X53))|(X51!=X54|~rpoint(X54)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
cnf(c_0_93,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)|vd1163=vd1164|~rcircle(esk69_2(vd1164,vd1163))|~ron(vd1163,esk69_2(vd1164,vd1163))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_86]), c_0_26])])).
cnf(c_0_94,plain,(vd1163=X1|~rpoint(X1)|~ron(vd1164,esk50_2(vd1163,X1))|~ron(vd1165,esk50_2(vd1163,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_26])])).
cnf(c_0_95,plain,(X1=X2|ron(X2,esk50_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_89])])).
cnf(c_0_96,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vd1163=vd1165|~ron(vd1163,esk69_2(vd1165,vd1163))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_26]), c_0_27])])).
cnf(c_0_97,plain,(vf(X2,X4)=v0|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|X2!=X4), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_98,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)|vd1163=vd1164|~ron(vd1163,esk69_2(vd1164,vd1163))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_91]), c_0_26]), c_0_31])])).
cnf(c_0_99,plain,(vd1163=vd1164|~ron(vd1165,esk50_2(vd1163,vd1164))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_31]), c_0_26])])).
cnf(c_0_100,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vd1163=vd1165), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_78]), c_0_26]), c_0_27])])).
cnf(c_0_101,plain,(vf(X1,X1)=v0|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_97])])])).
cnf(c_0_102,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)|vd1163=vd1164), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_78]), c_0_26]), c_0_31])])).
cnf(c_0_103,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vd1164=vd1165|~ron(vd1165,esk50_2(vd1165,vd1164))), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_104,plain,(X2=X4|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|vf(X2,X4)!=v0), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_105,plain,(vf(X1,X1)=vf(X2,X2)|~rpoint(X1)|~rpoint(X2)), inference(spm,[status(thm)],[c_0_101, c_0_101])).
cnf(c_0_106,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)|vf(vd1164,vd1164)=vf(vd1156,vd1157)), inference(spm,[status(thm)],[c_0_30, c_0_102])).
fof(c_0_107, plain, (![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:((((((((((((((((ron(X26,esk8_6(X21,X22,X23,X24,X25,X26))|vangle(X22,X21,X23)!=vangle(X25,X24,X26))|vf(X21,X23)!=vf(X24,X26))|vf(X21,X22)!=vf(X24,X25))|(X26!=X28|~rpoint(X28)))|(X25!=X29|~rpoint(X29)))|(X24!=X30|~rpoint(X30)))|X22=X23)|X21=X23)|X21=X22)|(X23!=X31|~rpoint(X31)))|(X22!=X32|~rpoint(X32)))|(X21!=X33|~rpoint(X33)))|vangle(X24,X25,X26)=vangle(X21,X22,X23))&(((((((((((((ron(X25,esk8_6(X21,X22,X23,X24,X25,X26))|vangle(X22,X21,X23)!=vangle(X25,X24,X26))|vf(X21,X23)!=vf(X24,X26))|vf(X21,X22)!=vf(X24,X25))|(X26!=X28|~rpoint(X28)))|(X25!=X29|~rpoint(X29)))|(X24!=X30|~rpoint(X30)))|X22=X23)|X21=X23)|X21=X22)|(X23!=X31|~rpoint(X31)))|(X22!=X32|~rpoint(X32)))|(X21!=X33|~rpoint(X33)))|vangle(X24,X25,X26)=vangle(X21,X22,X23)))&(((((((((((((ron(X24,esk8_6(X21,X22,X23,X24,X25,X26))|vangle(X22,X21,X23)!=vangle(X25,X24,X26))|vf(X21,X23)!=vf(X24,X26))|vf(X21,X22)!=vf(X24,X25))|(X26!=X28|~rpoint(X28)))|(X25!=X29|~rpoint(X29)))|(X24!=X30|~rpoint(X30)))|X22=X23)|X21=X23)|X21=X22)|(X23!=X31|~rpoint(X31)))|(X22!=X32|~rpoint(X32)))|(X21!=X33|~rpoint(X33)))|vangle(X24,X25,X26)=vangle(X21,X22,X23)))&(((((((((((((rline(esk8_6(X21,X22,X23,X24,X25,X26))|vangle(X22,X21,X23)!=vangle(X25,X24,X26))|vf(X21,X23)!=vf(X24,X26))|vf(X21,X22)!=vf(X24,X25))|(X26!=X28|~rpoint(X28)))|(X25!=X29|~rpoint(X29)))|(X24!=X30|~rpoint(X30)))|X22=X23)|X21=X23)|X21=X22)|(X23!=X31|~rpoint(X31)))|(X22!=X32|~rpoint(X32)))|(X21!=X33|~rpoint(X33)))|vangle(X24,X25,X26)=vangle(X21,X22,X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
cnf(c_0_108,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vd1164=vd1165), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_81]), c_0_31]), c_0_27])])).
cnf(c_0_109,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vf(vd1165,vd1165)=vf(vd1156,vd1158)), inference(spm,[status(thm)],[c_0_24, c_0_100])).
cnf(c_0_110,plain,(X1=X2|vf(X1,X2)!=v0|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_104])])).
cnf(c_0_111,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)|vf(vd1156,vd1157)=vf(X1,X1)|~rpoint(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_31])])).
cnf(c_0_112,plain,(vd1156!=vd1157), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_113,plain,(vangle(X1,X2,X3)=vangle(X4,X5,X6)|X4=X5|X4=X6|X5=X6|ron(X1,esk8_6(X4,X5,X6,X1,X2,X3))|~rpoint(X7)|X4!=X7|~rpoint(X8)|X5!=X8|~rpoint(X9)|X6!=X9|~rpoint(X10)|X1!=X10|~rpoint(X11)|X2!=X11|~rpoint(X12)|X3!=X12|vf(X4,X5)!=vf(X1,X2)|vf(X4,X6)!=vf(X1,X3)|vangle(X5,X4,X6)!=vangle(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_114,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vf(vd1156,vd1157)=vf(vd1156,vd1158)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_108]), c_0_24])).
cnf(c_0_115,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)|vf(vd1156,vd1158)=v0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_109]), c_0_27])])).
cnf(c_0_116,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)|~rpoint(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_50]), c_0_45])]), c_0_112]), c_0_101])).
cnf(c_0_117,plain,(rpoint(vd14)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_118,plain,(vangle(X1,X2,X3)=vangle(X4,X5,X6)|X4=X5|X4=X6|X5=X6|rline(esk8_6(X4,X5,X6,X1,X2,X3))|~rpoint(X7)|X4!=X7|~rpoint(X8)|X5!=X8|~rpoint(X9)|X6!=X9|~rpoint(X10)|X1!=X10|~rpoint(X11)|X2!=X11|~rpoint(X12)|X3!=X12|vf(X4,X5)!=vf(X1,X2)|vf(X4,X6)!=vf(X1,X3)|vangle(X5,X4,X6)!=vangle(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_119,plain,(vangle(X1,X2,X3)=vangle(X4,X5,X6)|X4=X5|X4=X6|X5=X6|ron(X2,esk8_6(X4,X5,X6,X1,X2,X3))|~rpoint(X7)|X4!=X7|~rpoint(X8)|X5!=X8|~rpoint(X9)|X6!=X9|~rpoint(X10)|X1!=X10|~rpoint(X11)|X2!=X11|~rpoint(X12)|X3!=X12|vf(X4,X5)!=vf(X1,X2)|vf(X4,X6)!=vf(X1,X3)|vangle(X5,X4,X6)!=vangle(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_120,plain,(vangle(X1,X2,X3)=vangle(X4,X5,X6)|X4=X5|X4=X6|X5=X6|ron(X3,esk8_6(X4,X5,X6,X1,X2,X3))|~rpoint(X7)|X4!=X7|~rpoint(X8)|X5!=X8|~rpoint(X9)|X6!=X9|~rpoint(X10)|X1!=X10|~rpoint(X11)|X2!=X11|~rpoint(X12)|X3!=X12|vf(X4,X5)!=vf(X1,X2)|vf(X4,X6)!=vf(X1,X3)|vangle(X5,X4,X6)!=vangle(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_121,plain,(vangle(X1,X2,X3)=vangle(X4,X5,X6)|X5=X6|X4=X6|X4=X5|ron(X1,esk8_6(X4,X5,X6,X1,X2,X3))|vangle(X2,X1,X3)!=vangle(X5,X4,X6)|vf(X1,X3)!=vf(X4,X6)|vf(X1,X2)!=vf(X4,X5)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~rpoint(X6)|~rpoint(X5)|~rpoint(X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_113])])])])])])).
cnf(c_0_122,plain,(vangle(vd1157,vd1156,vd1158)=vangle(vd1164,vd1163,vd1165)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_123,plain,(vf(vd1156,vd1158)=vf(vd1158,vd1156)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_114]), c_0_50]), c_0_45])]), c_0_112]), c_0_115])).
cnf(c_0_124,plain,(vf(vd1156,vd1157)=vf(vd1157,vd1156)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
fof(c_0_125, negated_conjecture, (vangle(vd1156,vd1157,vd1158)!=vangle(vd1163,vd1164,vd1165)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_11])])).
cnf(c_0_126,plain,(vangle(X1,X2,X3)=vangle(X4,X5,X6)|X5=X6|X4=X6|X4=X5|rline(esk8_6(X4,X5,X6,X1,X2,X3))|vangle(X2,X1,X3)!=vangle(X5,X4,X6)|vf(X1,X3)!=vf(X4,X6)|vf(X1,X2)!=vf(X4,X5)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~rpoint(X6)|~rpoint(X5)|~rpoint(X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_118])])])])])])).
cnf(c_0_127,plain,(vangle(X1,X2,X3)=vangle(X4,X5,X6)|X5=X6|X4=X6|X4=X5|ron(X2,esk8_6(X4,X5,X6,X1,X2,X3))|vangle(X2,X1,X3)!=vangle(X5,X4,X6)|vf(X1,X3)!=vf(X4,X6)|vf(X1,X2)!=vf(X4,X5)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~rpoint(X6)|~rpoint(X5)|~rpoint(X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_119])])])])])])).
cnf(c_0_128,plain,(vangle(X1,X2,X3)=vangle(X4,X5,X6)|X5=X6|X4=X6|X4=X5|ron(X3,esk8_6(X4,X5,X6,X1,X2,X3))|vangle(X2,X1,X3)!=vangle(X5,X4,X6)|vf(X1,X3)!=vf(X4,X6)|vf(X1,X2)!=vf(X4,X5)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~rpoint(X6)|~rpoint(X5)|~rpoint(X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_120])])])])])])).
cnf(c_0_129,plain,(vangle(vd1163,vd1164,vd1165)=vangle(X1,X2,X3)|X2=X3|X1=X3|X1=X2|ron(vd1163,esk8_6(X1,X2,X3,vd1163,vd1164,vd1165))|vangle(vd1157,vd1156,vd1158)!=vangle(X2,X1,X3)|vf(vd1158,vd1156)!=vf(X1,X3)|vf(vd1157,vd1156)!=vf(X1,X2)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_24]), c_0_30]), c_0_27]), c_0_31]), c_0_26])]), c_0_123]), c_0_124])).
cnf(c_0_130,negated_conjecture,(vangle(vd1156,vd1157,vd1158)!=vangle(vd1163,vd1164,vd1165)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_131,plain,(vd1157!=vd1158), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_132,plain,(vd1156!=vd1158), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_133,plain,(vangle(vd1163,vd1164,vd1165)=vangle(X1,X2,X3)|X2=X3|X1=X3|X1=X2|rline(esk8_6(X1,X2,X3,vd1163,vd1164,vd1165))|vangle(vd1157,vd1156,vd1158)!=vangle(X2,X1,X3)|vf(vd1158,vd1156)!=vf(X1,X3)|vf(vd1157,vd1156)!=vf(X1,X2)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_122]), c_0_24]), c_0_30]), c_0_27]), c_0_31]), c_0_26])]), c_0_123]), c_0_124])).
cnf(c_0_134,plain,(vangle(vd1163,vd1164,vd1165)=vangle(X1,X2,X3)|X2=X3|X1=X3|X1=X2|ron(vd1164,esk8_6(X1,X2,X3,vd1163,vd1164,vd1165))|vangle(vd1157,vd1156,vd1158)!=vangle(X2,X1,X3)|vf(vd1158,vd1156)!=vf(X1,X3)|vf(vd1157,vd1156)!=vf(X1,X2)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_122]), c_0_24]), c_0_30]), c_0_27]), c_0_31]), c_0_26])]), c_0_123]), c_0_124])).
cnf(c_0_135,plain,(vangle(vd1163,vd1164,vd1165)=vangle(X1,X2,X3)|X2=X3|X1=X3|X1=X2|ron(vd1165,esk8_6(X1,X2,X3,vd1163,vd1164,vd1165))|vangle(vd1157,vd1156,vd1158)!=vangle(X2,X1,X3)|vf(vd1158,vd1156)!=vf(X1,X3)|vf(vd1157,vd1156)!=vf(X1,X2)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_122]), c_0_24]), c_0_30]), c_0_27]), c_0_31]), c_0_26])]), c_0_123]), c_0_124])).
cnf(c_0_136,plain,(ron(vd1163,esk8_6(vd1156,vd1157,vd1158,vd1163,vd1164,vd1165))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_129]), c_0_123]), c_0_124]), c_0_44]), c_0_50]), c_0_45])]), c_0_130]), c_0_131]), c_0_132]), c_0_112])).
cnf(c_0_137,plain,(rline(esk8_6(vd1156,vd1157,vd1158,vd1163,vd1164,vd1165))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_133]), c_0_123]), c_0_124]), c_0_44]), c_0_50]), c_0_45])]), c_0_130]), c_0_131]), c_0_132]), c_0_112])).
cnf(c_0_138,plain,(ron(vd1164,esk8_6(vd1156,vd1157,vd1158,vd1163,vd1164,vd1165))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_134]), c_0_123]), c_0_124]), c_0_44]), c_0_50]), c_0_45])]), c_0_130]), c_0_131]), c_0_132]), c_0_112])).
cnf(c_0_139,plain,(ron(vd1165,esk8_6(vd1156,vd1157,vd1158,vd1163,vd1164,vd1165))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_135]), c_0_123]), c_0_124]), c_0_44]), c_0_50]), c_0_45])]), c_0_130]), c_0_131]), c_0_132]), c_0_112])).
cnf(c_0_140,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_136]), c_0_137]), c_0_138]), c_0_139])]), ['proof']).
# SZS output end CNFRefutation.
EOF
