% Problem    : SCT120+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n142.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:16:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X7]:![X4]:![X5]:(c_member(X5,X4,X7)<=>hBOOL(hAPP(X7,X4)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_mem__def)).
fof(c_0_1, axiom, (![X5]:![X7]:![X4]:![X14]:((~(hBOOL(hAPP(X14,X4)))=>hBOOL(hAPP(X7,X4)))=>hBOOL(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X5,tc_HOL_Obool),X7,X14),X4)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_sup1CI)).
fof(c_0_2, axiom, (![X6]:![X5]:c_Lattices_Osemilattice__sup__class_Osup(tc_fun(tc_prod(X5,X5),tc_HOL_Obool),c_Transitive__Closure_Otrancl(X5,X6),c_Relation_OId(X5))=c_Transitive__Closure_Ortrancl(X5,X6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_reflcl__trancl)).
fof(c_0_3, axiom, (![X5]:c_Transitive__Closure_Ortrancl(X5,c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X5,X5),tc_HOL_Obool)))=c_Relation_OId(X5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rtrancl__empty)).
fof(c_0_4, hypothesis, (![X3]:(![X8]:((c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X3,X8))&c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X3,X8)))&c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(X3,X8)))=>((c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,v_F(X3))&c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,v_F(X3)))&c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),v_F(X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_5, hypothesis, (![X3]:((c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(v_P,X3))&c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(v_P,X3)))&c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(v_P,X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_2)).
fof(c_0_6, axiom, (![X6]:![X33]:![X5]:(c_member(tc_prod(X5,X5),X33,X6)=>c_member(tc_prod(X5,X5),X33,c_Transitive__Closure_Otrancl(X5,X6)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_r__into__trancl_H)).
fof(c_0_7, hypothesis, (c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_a,v_b),v_F(v_P))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_3)).
fof(c_0_8, axiom, (![X6]:![X5]:(c_Relation_Otrans(X5,X6)=>c_Transitive__Closure_Otrancl(X5,X6)=X6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_trancl__id)).
fof(c_0_9, axiom, (![X6]:![X5]:(c_Relation_Oirrefl(X5,X6)<=>![X3]:~(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X3,X3),X6)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_irrefl__def)).
fof(c_0_10, axiom, (![X6]:![X5]:(c_Wellfounded_Oacyclic(X5,X6)<=>![X3]:~(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X3,X3),c_Transitive__Closure_Otrancl(X5,X6))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_acyclic__def)).
fof(c_0_11, axiom, (![X19]:![X17]:![X4]:![X5]:(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X4,X17),c_Transitive__Closure_Otrancl(X5,X19))=>?[X10]:(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X4,X10),c_Transitive__Closure_Ortrancl(X5,X19))&c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X10,X17),X19)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_tranclD2)).
fof(c_0_12, axiom, (![X6]:![X5]:(c_Wellfounded_Oacyclic(X5,X6)=>c_Relation_Oantisym(X5,c_Transitive__Closure_Ortrancl(X5,X6)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_acyclic__impl__antisym__rtrancl)).
fof(c_0_13, axiom, (![X19]:![X5]:c_Transitive__Closure_Ortrancl(X5,c_Transitive__Closure_Otrancl(X5,X19))=c_Transitive__Closure_Ortrancl(X5,X19)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_trancl__rtrancl__absorb)).
fof(c_0_14, axiom, (![X6]:![X5]:(c_Relation_Otrans(X5,X6)<=>![X3]:![X9]:(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X3,X9),X6)=>![X10]:(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X9,X10),X6)=>c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X3,X10),X6))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_trans__def)).
fof(c_0_15, axiom, (![X4]:![X7]:![X17]:![X5]:(hBOOL(hAPP(c_Set_Oinsert(X5,X17,X7),X4))<=>(X17=X4|hBOOL(hAPP(X7,X4))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_insert__code)).
fof(c_0_16, axiom, (![X12]:![X13]:![X6]:![X5]:(c_Relation_Oantisym(X5,X6)=>(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X13,X12),X6)=>(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X12,X13),X6)=>X13=X12)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_antisymD)).
fof(c_0_17, axiom, (![X11]:![X19]:![X12]:![X13]:![X5]:(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X13,X12),X19)=>(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X12,X11),X19)=>c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X13,X11),c_Transitive__Closure_Otrancl(X5,X19))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_r__r__into__trancl)).
fof(c_0_18, axiom, (![X19]:![X5]:c_Transitive__Closure_Otrancl(X5,c_Transitive__Closure_Ortrancl(X5,X19))=c_Transitive__Closure_Ortrancl(X5,X19)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rtrancl__trancl__absorb)).
fof(c_0_19, hypothesis, (![X78]:(![X3]:((c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X78,X3))&c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X78,X3)))&c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(X78,X3)))=>![X38]:![X39]:(X38!=X39=>(c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X38,X39),hAPP(X78,v_i))=>c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X38,X39),v_F(X78)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_1)).
fof(c_0_20, axiom, (![X6]:![X7]:![X5]:(c_Relation_Ototal__on(X5,X7,X6)<=>![X3]:(c_member(X5,X3,X7)=>![X8]:(c_member(X5,X8,X7)=>(X3!=X8=>(c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X3,X8),X6)|c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X8,X3),X6))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_total__on__def)).
fof(c_0_21, axiom, (![X4]:![X5]:c_member(X5,X4,c_Orderings_Otop__class_Otop(tc_fun(X5,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__I)).
fof(c_0_22, axiom, (![X7]:![X13]:![X5]:(c_member(X5,X13,X7)=>c_Set_Oinsert(X5,X13,X7)=X7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_insert__absorb)).
fof(c_0_23, conjecture, (c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_a,v_b),hAPP(v_P,v_i))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_4)).
fof(c_0_24, axiom, (![X6]:![X13]:![X5]:c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X13,X13),c_Transitive__Closure_Ortrancl(X5,X6))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rtrancl_Ortrancl__refl)).
fof(c_0_25, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:((~c_member(X10,X9,X8)|hBOOL(hAPP(X8,X9)))&(~hBOOL(hAPP(X11,X12))|c_member(X13,X12,X11)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])).
fof(c_0_26, plain, (![X5]:![X7]:![X4]:![X14]:((~hBOOL(hAPP(X14,X4))=>hBOOL(hAPP(X7,X4)))=>hBOOL(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X5,tc_HOL_Obool),X7,X14),X4)))), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_27, plain, (![X7]:![X8]:c_Lattices_Osemilattice__sup__class_Osup(tc_fun(tc_prod(X8,X8),tc_HOL_Obool),c_Transitive__Closure_Otrancl(X8,X7),c_Relation_OId(X8))=c_Transitive__Closure_Ortrancl(X8,X7)), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_28, plain, (![X6]:c_Transitive__Closure_Ortrancl(X6,c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X6,X6),tc_HOL_Obool)))=c_Relation_OId(X6)), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_29, hypothesis, (![X9]:(((c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,v_F(X9))|((~c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X9,esk87_1(X9)))|~c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X9,esk88_1(X9))))|~c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(X9,esk89_1(X9)))))&(c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,v_F(X9))|((~c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X9,esk87_1(X9)))|~c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X9,esk88_1(X9))))|~c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(X9,esk89_1(X9))))))&(c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),v_F(X9))|((~c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X9,esk87_1(X9)))|~c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X9,esk88_1(X9))))|~c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(X9,esk89_1(X9))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
fof(c_0_30, hypothesis, (![X4]:![X5]:![X6]:((c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(v_P,X4))&c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(v_P,X5)))&c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(v_P,X6)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[c_0_5])])])).
fof(c_0_31, plain, (![X34]:![X35]:![X36]:(~c_member(tc_prod(X36,X36),X35,X34)|c_member(tc_prod(X36,X36),X35,c_Transitive__Closure_Otrancl(X36,X34)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
cnf(c_0_32,plain,(hBOOL(hAPP(X1,X2))|~c_member(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33,hypothesis,(c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_a,v_b),v_F(v_P))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_34, plain, (![X15]:![X16]:![X17]:![X18]:((~hBOOL(hAPP(X18,X17))|hBOOL(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X15,tc_HOL_Obool),X16,X18),X17)))&(~hBOOL(hAPP(X16,X17))|hBOOL(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X15,tc_HOL_Obool),X16,X18),X17))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
cnf(c_0_35,plain,(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(tc_prod(X1,X1),tc_HOL_Obool),c_Transitive__Closure_Otrancl(X1,X2),c_Relation_OId(X1))=c_Transitive__Closure_Ortrancl(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36,plain,(c_Transitive__Closure_Ortrancl(X1,c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X1,X1),tc_HOL_Obool)))=c_Relation_OId(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_37, plain, (![X7]:![X8]:(~c_Relation_Otrans(X8,X7)|c_Transitive__Closure_Otrancl(X8,X7)=X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
cnf(c_0_38,hypothesis,(c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,v_F(X1))|~c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(X1,esk89_1(X1)))|~c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X1,esk88_1(X1)))|~c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X1,esk87_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39,hypothesis,(c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(v_P,X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40,hypothesis,(c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(v_P,X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41,hypothesis,(c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(v_P,X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42,plain,(c_member(tc_prod(X1,X1),X2,c_Transitive__Closure_Otrancl(X1,X3))|~c_member(tc_prod(X1,X1),X2,X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43,plain,(c_member(X1,X2,X3)|~hBOOL(hAPP(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_44,hypothesis,(hBOOL(hAPP(v_F(v_P),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_a,v_b)))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_45, plain, (![X6]:![X5]:(c_Relation_Oirrefl(X5,X6)<=>![X3]:~c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X3,X3),X6))), inference(fof_simplification,[status(thm)],[c_0_9])).
fof(c_0_46, plain, (![X6]:![X5]:(c_Wellfounded_Oacyclic(X5,X6)<=>![X3]:~c_member(tc_prod(X5,X5),c_Product__Type_OPair(X5,X5,X3,X3),c_Transitive__Closure_Otrancl(X5,X6)))), inference(fof_simplification,[status(thm)],[c_0_10])).
cnf(c_0_47,plain,(hBOOL(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool),X2,X3),X4))|~hBOOL(hAPP(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48,plain,(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(tc_prod(X1,X1),tc_HOL_Obool),c_Transitive__Closure_Otrancl(X1,X2),c_Transitive__Closure_Ortrancl(X1,c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X1,X1),tc_HOL_Obool))))=c_Transitive__Closure_Ortrancl(X1,X2)), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_49,plain,(c_Transitive__Closure_Otrancl(X1,X2)=X2|~c_Relation_Otrans(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_50,hypothesis,(c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_41])])).
fof(c_0_51, plain, (![X20]:![X21]:![X22]:![X23]:((c_member(tc_prod(X23,X23),c_Product__Type_OPair(X23,X23,X22,esk66_4(X20,X21,X22,X23)),c_Transitive__Closure_Ortrancl(X23,X20))|~c_member(tc_prod(X23,X23),c_Product__Type_OPair(X23,X23,X22,X21),c_Transitive__Closure_Otrancl(X23,X20)))&(c_member(tc_prod(X23,X23),c_Product__Type_OPair(X23,X23,esk66_4(X20,X21,X22,X23),X21),X20)|~c_member(tc_prod(X23,X23),c_Product__Type_OPair(X23,X23,X22,X21),c_Transitive__Closure_Otrancl(X23,X20))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
cnf(c_0_52,plain,(hBOOL(hAPP(c_Transitive__Closure_Otrancl(X1,X2),X3))|~c_member(tc_prod(X1,X1),X3,X2)), inference(spm,[status(thm)],[c_0_32, c_0_42])).
cnf(c_0_53,hypothesis,(c_member(X1,c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_a,v_b),v_F(v_P))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
fof(c_0_54, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:((~c_Relation_Oirrefl(X8,X7)|~c_member(tc_prod(X8,X8),c_Product__Type_OPair(X8,X8,X9,X9),X7))&(c_member(tc_prod(X11,X11),c_Product__Type_OPair(X11,X11,esk2_2(X10,X11),esk2_2(X10,X11)),X10)|c_Relation_Oirrefl(X11,X10)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_45])])])])])).
fof(c_0_55, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:((~c_Wellfounded_Oacyclic(X8,X7)|~c_member(tc_prod(X8,X8),c_Product__Type_OPair(X8,X8,X9,X9),c_Transitive__Closure_Otrancl(X8,X7)))&(c_member(tc_prod(X11,X11),c_Product__Type_OPair(X11,X11,esk67_2(X10,X11),esk67_2(X10,X11)),c_Transitive__Closure_Otrancl(X11,X10))|c_Wellfounded_Oacyclic(X11,X10)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_46])])])])])).
cnf(c_0_56,plain,(hBOOL(hAPP(c_Transitive__Closure_Ortrancl(X1,X2),X3))|~hBOOL(hAPP(c_Transitive__Closure_Otrancl(X1,X2),X3))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_57,hypothesis,(c_Transitive__Closure_Otrancl(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P))=v_F(v_P)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_58,plain,(c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,esk66_4(X4,X3,X2,X1),X3),X4)|~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X3),c_Transitive__Closure_Otrancl(X1,X4))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59,hypothesis,(hBOOL(hAPP(c_Transitive__Closure_Otrancl(X1,v_F(v_P)),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_a,v_b)))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
fof(c_0_60, plain, (![X7]:![X8]:(~c_Wellfounded_Oacyclic(X8,X7)|c_Relation_Oantisym(X8,c_Transitive__Closure_Ortrancl(X8,X7)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
fof(c_0_61, plain, (![X20]:![X21]:c_Transitive__Closure_Ortrancl(X21,c_Transitive__Closure_Otrancl(X21,X20))=c_Transitive__Closure_Ortrancl(X21,X20)), inference(variable_rename,[status(thm)],[c_0_13])).
cnf(c_0_62,plain,(~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X2),X3)|~c_Relation_Oirrefl(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_63,plain,(c_Wellfounded_Oacyclic(X1,X2)|c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,esk67_2(X2,X1),esk67_2(X2,X1)),c_Transitive__Closure_Otrancl(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_64,hypothesis,(c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,v_F(X1))|~c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(X1,esk89_1(X1)))|~c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X1,esk88_1(X1)))|~c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X1,esk87_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_65,hypothesis,(hBOOL(hAPP(c_Transitive__Closure_Ortrancl(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P)),X1))|~hBOOL(hAPP(v_F(v_P),X1))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_66,plain,(hBOOL(hAPP(X1,c_Product__Type_OPair(X2,X2,esk66_4(X1,X3,X4,X2),X3)))|~c_member(tc_prod(X2,X2),c_Product__Type_OPair(X2,X2,X4,X3),c_Transitive__Closure_Otrancl(X2,X1))), inference(spm,[status(thm)],[c_0_32, c_0_58])).
cnf(c_0_67,hypothesis,(c_member(X1,c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_a,v_b),c_Transitive__Closure_Otrancl(X2,v_F(v_P)))), inference(spm,[status(thm)],[c_0_43, c_0_59])).
cnf(c_0_68,plain,(c_Relation_Oantisym(X1,c_Transitive__Closure_Ortrancl(X1,X2))|~c_Wellfounded_Oacyclic(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_69,plain,(c_Transitive__Closure_Ortrancl(X1,c_Transitive__Closure_Otrancl(X1,X2))=c_Transitive__Closure_Ortrancl(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70,plain,(c_Wellfounded_Oacyclic(X1,X2)|~c_Relation_Oirrefl(X1,c_Transitive__Closure_Otrancl(X1,X2))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_71,hypothesis,(c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_39]), c_0_40]), c_0_41])])).
fof(c_0_72, plain, (![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:((~c_Relation_Otrans(X12,X11)|(~c_member(tc_prod(X12,X12),c_Product__Type_OPair(X12,X12,X13,X14),X11)|(~c_member(tc_prod(X12,X12),c_Product__Type_OPair(X12,X12,X14,X15),X11)|c_member(tc_prod(X12,X12),c_Product__Type_OPair(X12,X12,X13,X15),X11))))&((c_member(tc_prod(X17,X17),c_Product__Type_OPair(X17,X17,esk5_2(X16,X17),esk6_2(X16,X17)),X16)|c_Relation_Otrans(X17,X16))&((c_member(tc_prod(X17,X17),c_Product__Type_OPair(X17,X17,esk6_2(X16,X17),esk7_2(X16,X17)),X16)|c_Relation_Otrans(X17,X16))&(~c_member(tc_prod(X17,X17),c_Product__Type_OPair(X17,X17,esk5_2(X16,X17),esk7_2(X16,X17)),X16)|c_Relation_Otrans(X17,X16)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])])])).
fof(c_0_73, plain, (![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:((~hBOOL(hAPP(c_Set_Oinsert(X21,X20,X19),X18))|(X20=X18|hBOOL(hAPP(X19,X18))))&((X24!=X22|hBOOL(hAPP(c_Set_Oinsert(X25,X24,X23),X22)))&(~hBOOL(hAPP(X23,X22))|hBOOL(hAPP(c_Set_Oinsert(X25,X24,X23),X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).
fof(c_0_74, plain, (![X14]:![X15]:![X16]:![X17]:(~c_Relation_Oantisym(X17,X16)|(~c_member(tc_prod(X17,X17),c_Product__Type_OPair(X17,X17,X15,X14),X16)|(~c_member(tc_prod(X17,X17),c_Product__Type_OPair(X17,X17,X14,X15),X16)|X15=X14)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).
cnf(c_0_75,hypothesis,(c_member(X1,X2,c_Transitive__Closure_Ortrancl(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P)))|~hBOOL(hAPP(v_F(v_P),X2))), inference(spm,[status(thm)],[c_0_43, c_0_65])).
cnf(c_0_76,hypothesis,(hBOOL(hAPP(v_F(v_P),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,esk66_4(v_F(v_P),v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt),v_b)))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_77,plain,(c_Relation_Oantisym(X1,c_Transitive__Closure_Ortrancl(X1,X2))|~c_Wellfounded_Oacyclic(X1,c_Transitive__Closure_Otrancl(X1,X2))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_78,hypothesis,(c_Wellfounded_Oacyclic(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_57]), c_0_71])])).
fof(c_0_79, plain, (![X20]:![X21]:![X22]:![X23]:![X24]:(~c_member(tc_prod(X24,X24),c_Product__Type_OPair(X24,X24,X23,X22),X21)|(~c_member(tc_prod(X24,X24),c_Product__Type_OPair(X24,X24,X22,X20),X21)|c_member(tc_prod(X24,X24),c_Product__Type_OPair(X24,X24,X23,X20),c_Transitive__Closure_Otrancl(X24,X21))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
fof(c_0_80, plain, (![X20]:![X21]:c_Transitive__Closure_Otrancl(X21,c_Transitive__Closure_Ortrancl(X21,X20))=c_Transitive__Closure_Ortrancl(X21,X20)), inference(variable_rename,[status(thm)],[c_0_18])).
cnf(c_0_81,plain,(c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X3),X4)|~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X5,X3),X4)|~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X5),X4)|~c_Relation_Otrans(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_72])).
fof(c_0_82, hypothesis, (![X79]:![X83]:![X84]:(((~c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X79,esk90_1(X79)))|~c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X79,esk91_1(X79))))|~c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(X79,esk92_1(X79))))|(X83=X84|(~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X83,X84),hAPP(X79,v_i))|c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X83,X84),v_F(X79)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])])).
cnf(c_0_83,plain,(hBOOL(hAPP(c_Set_Oinsert(X1,X2,X3),X4))|X2!=X4), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_84,plain,(X1=X2|~c_member(tc_prod(X3,X3),c_Product__Type_OPair(X3,X3,X2,X1),X4)|~c_member(tc_prod(X3,X3),c_Product__Type_OPair(X3,X3,X1,X2),X4)|~c_Relation_Oantisym(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_85,hypothesis,(c_member(X1,c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,esk66_4(v_F(v_P),v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt),v_b),c_Transitive__Closure_Ortrancl(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P)))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_86,hypothesis,(c_Relation_Oantisym(tc_Arrow__Order__Mirabelle_Oalt,c_Transitive__Closure_Ortrancl(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_57]), c_0_78])])).
cnf(c_0_87,plain,(c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X3),c_Transitive__Closure_Otrancl(X1,X4))|~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X5,X3),X4)|~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X5),X4)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_88,plain,(c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,esk66_4(X4,X3,X2,X1)),c_Transitive__Closure_Ortrancl(X1,X4))|~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X3),c_Transitive__Closure_Otrancl(X1,X4))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_89,plain,(c_Transitive__Closure_Otrancl(X1,c_Transitive__Closure_Ortrancl(X1,X2))=c_Transitive__Closure_Ortrancl(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_90,plain,(~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X2),c_Transitive__Closure_Otrancl(X1,X3))|~c_Wellfounded_Oacyclic(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_91,hypothesis,(c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X1,v_b),v_F(v_P))|~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X1,v_a),v_F(v_P))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_53]), c_0_50])])).
cnf(c_0_92,hypothesis,(c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X1,X2),v_F(X3))|X1=X2|~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X1,X2),hAPP(X3,v_i))|~c_Relation_Ototal__on(tc_Arrow__Order__Mirabelle_Oalt,c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_HOL_Obool)),hAPP(X3,esk92_1(X3)))|~c_Relation_Oirrefl(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X3,esk91_1(X3)))|~c_Relation_Otrans(tc_Arrow__Order__Mirabelle_Oalt,hAPP(X3,esk90_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_82])).
fof(c_0_93, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:((~c_Relation_Ototal__on(X11,X10,X9)|(~c_member(X11,X12,X10)|(~c_member(X11,X13,X10)|(X12=X13|(c_member(tc_prod(X11,X11),c_Product__Type_OPair(X11,X11,X12,X13),X9)|c_member(tc_prod(X11,X11),c_Product__Type_OPair(X11,X11,X13,X12),X9))))))&((c_member(X16,esk3_3(X14,X15,X16),X15)|c_Relation_Ototal__on(X16,X15,X14))&((c_member(X16,esk4_3(X14,X15,X16),X15)|c_Relation_Ototal__on(X16,X15,X14))&((esk3_3(X14,X15,X16)!=esk4_3(X14,X15,X16)|c_Relation_Ototal__on(X16,X15,X14))&((~c_member(tc_prod(X16,X16),c_Product__Type_OPair(X16,X16,esk3_3(X14,X15,X16),esk4_3(X14,X15,X16)),X14)|c_Relation_Ototal__on(X16,X15,X14))&(~c_member(tc_prod(X16,X16),c_Product__Type_OPair(X16,X16,esk4_3(X14,X15,X16),esk3_3(X14,X15,X16)),X14)|c_Relation_Ototal__on(X16,X15,X14)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])])])).
fof(c_0_94, plain, (![X6]:![X7]:c_member(X7,X6,c_Orderings_Otop__class_Otop(tc_fun(X7,tc_HOL_Obool)))), inference(variable_rename,[status(thm)],[c_0_21])).
cnf(c_0_95,plain,(hBOOL(hAPP(c_Set_Oinsert(X1,X2,X3),X2))), inference(er,[status(thm)],[c_0_83])).
fof(c_0_96, plain, (![X14]:![X15]:![X16]:(~c_member(X16,X15,X14)|c_Set_Oinsert(X16,X15,X14)=X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])).
cnf(c_0_97,hypothesis,(esk66_4(v_F(v_P),v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt)=v_b|~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_b,esk66_4(v_F(v_P),v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt)),c_Transitive__Closure_Ortrancl(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])])).
cnf(c_0_98,plain,(c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,esk66_4(X3,X4,X5,X1)),c_Transitive__Closure_Ortrancl(X1,X3))|~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X5),c_Transitive__Closure_Ortrancl(X1,X3))|~c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X5,X4),c_Transitive__Closure_Otrancl(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89])).
cnf(c_0_99,hypothesis,(~c_Wellfounded_Oacyclic(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P))|~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X1,X1),v_F(v_P))), inference(spm,[status(thm)],[c_0_90, c_0_57])).
cnf(c_0_100,hypothesis,(~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_b,v_a),v_F(v_P))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_91]), c_0_71])])).
cnf(c_0_101,hypothesis,(X1=X2|c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X1,X2),v_F(v_P))|~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X1,X2),hAPP(v_P,v_i))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_39]), c_0_40]), c_0_41])])).
cnf(c_0_102,plain,(c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X3),X4)|c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X3,X2),X4)|X3=X2|~c_member(X1,X2,X5)|~c_member(X1,X3,X5)|~c_Relation_Ototal__on(X1,X5,X4)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_103,plain,(c_member(X1,X2,c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))), inference(split_conjunct,[status(thm)],[c_0_94])).
fof(c_0_104, negated_conjecture, (~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_a,v_b),hAPP(v_P,v_i))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_23])])).
cnf(c_0_105,plain,(c_member(X1,X2,c_Set_Oinsert(X3,X2,X4))), inference(spm,[status(thm)],[c_0_43, c_0_95])).
cnf(c_0_106,plain,(c_Set_Oinsert(X1,X2,X3)=X3|~c_member(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_96])).
fof(c_0_107, plain, (![X14]:![X15]:![X16]:c_member(tc_prod(X16,X16),c_Product__Type_OPair(X16,X16,X15,X15),c_Transitive__Closure_Ortrancl(X16,X14))), inference(variable_rename,[status(thm)],[c_0_24])).
cnf(c_0_108,hypothesis,(esk66_4(v_F(v_P),v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt)=v_b|~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_b,v_a),c_Transitive__Closure_Ortrancl(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_57]), c_0_53])])).
cnf(c_0_109,hypothesis,(~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X1,X1),v_F(v_P))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_78])])).
cnf(c_0_110,hypothesis,(v_b=v_a|~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_b,v_a),hAPP(v_P,v_i))), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_111,hypothesis,(X1=X2|c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X2,X1),hAPP(v_P,X3))|c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,X1,X2),hAPP(v_P,X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_39]), c_0_103]), c_0_103])])).
cnf(c_0_112,negated_conjecture,(~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_a,v_b),hAPP(v_P,v_i))), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_113,plain,(c_member(X1,X2,X3)|~c_member(X4,X2,X3)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_114,plain,(c_member(tc_prod(X1,X1),c_Product__Type_OPair(X1,X1,X2,X2),c_Transitive__Closure_Ortrancl(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_115,hypothesis,(~c_member(tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),c_Product__Type_OPair(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt,v_b,v_a),c_Transitive__Closure_Ortrancl(tc_Arrow__Order__Mirabelle_Oalt,v_F(v_P)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_108]), c_0_57]), c_0_53])]), c_0_109])).
cnf(c_0_116,hypothesis,(v_b=v_a), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112])).
cnf(c_0_117,plain,(c_member(X1,c_Product__Type_OPair(X2,X2,X3,X3),c_Transitive__Closure_Ortrancl(X2,X4))), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_118,hypothesis,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_116]), c_0_117])]), ['proof']).
# SZS output end CNFRefutation.
EOF
