% Problem    : SWB056+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n187.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:16:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C45_F1_AE_R8_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X11]:![X12]:![X21]:![X14]:![X22]:((((iext(uri_rdf_first,X12,X21)&iext(uri_rdf_rest,X12,X14))&iext(uri_rdf_first,X14,X22))&iext(uri_rdf_rest,X14,uri_rdf_nil))=>(iext(uri_owl_oneOf,X11,X12)<=>(ic(X11)&![X4]:(icext(X11,X4)<=>(X4=X21|X4=X22)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_enum_class_002)).
fof(c_0_1, axiom, (?[X42]:?[X37]:((((((iext(uri_rdf_type,uri_ex_z,uri_ex_e)&iext(uri_owl_differentFrom,uri_ex_z,uri_ex_x))&iext(uri_owl_oneOf,uri_ex_e,X42))&iext(uri_rdf_first,X42,uri_ex_x))&iext(uri_rdf_rest,X42,X37))&iext(uri_rdf_first,X37,uri_ex_y))&iext(uri_rdf_rest,X37,uri_rdf_nil))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', premise_rdfbased_sem_enum_inst_closed)).
fof(c_0_2, axiom, (![X2]:(icext(uri_owl_SymmetricProperty,X2)<=>(ip(X2)&![X4]:![X7]:(iext(X2,X4,X7)=>iext(X2,X7,X4))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_char_symmetric)).
fof(c_0_3, axiom, (![X1]:![X2]:![X3]:(iext(X2,X1,X3)=>ip(X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', simple_iext_property)).
fof(c_0_4, axiom, (![X4]:![X7]:(iext(uri_owl_differentFrom,X4,X7)<=>X4!=X7)), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_eqdis_differentfrom)).
fof(c_0_5, axiom, (![X4]:![X5]:(iext(uri_rdf_type,X4,X5)<=>icext(X5,X4))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', rdfs_cext_def)).
fof(c_0_6, axiom, (ip(uri_owl_differentFrom)), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_prop_differentfrom_type)).
fof(c_0_7, axiom, (![X4]:![X7]:(iext(uri_owl_sameAs,X4,X7)<=>X4=X7)), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_eqdis_sameas)).
fof(c_0_8, conjecture, (iext(uri_owl_sameAs,uri_ex_z,uri_ex_y)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conclusion_rdfbased_sem_enum_inst_closed)).
fof(c_0_9, plain, (![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:((((ic(X23)|~iext(uri_owl_oneOf,X23,X24))|(((~iext(uri_rdf_first,X24,X25)|~iext(uri_rdf_rest,X24,X26))|~iext(uri_rdf_first,X26,X27))|~iext(uri_rdf_rest,X26,uri_rdf_nil)))&((((~icext(X23,X28)|(X28=X25|X28=X27))|~iext(uri_owl_oneOf,X23,X24))|(((~iext(uri_rdf_first,X24,X25)|~iext(uri_rdf_rest,X24,X26))|~iext(uri_rdf_first,X26,X27))|~iext(uri_rdf_rest,X26,uri_rdf_nil)))&((((X29!=X25|icext(X23,X29))|~iext(uri_owl_oneOf,X23,X24))|(((~iext(uri_rdf_first,X24,X25)|~iext(uri_rdf_rest,X24,X26))|~iext(uri_rdf_first,X26,X27))|~iext(uri_rdf_rest,X26,uri_rdf_nil)))&(((X29!=X27|icext(X23,X29))|~iext(uri_owl_oneOf,X23,X24))|(((~iext(uri_rdf_first,X24,X25)|~iext(uri_rdf_rest,X24,X26))|~iext(uri_rdf_first,X26,X27))|~iext(uri_rdf_rest,X26,uri_rdf_nil))))))&((((((esk12_5(X23,X24,X25,X26,X27)!=X25|~icext(X23,esk12_5(X23,X24,X25,X26,X27)))|~ic(X23))|iext(uri_owl_oneOf,X23,X24))|(((~iext(uri_rdf_first,X24,X25)|~iext(uri_rdf_rest,X24,X26))|~iext(uri_rdf_first,X26,X27))|~iext(uri_rdf_rest,X26,uri_rdf_nil)))&((((esk12_5(X23,X24,X25,X26,X27)!=X27|~icext(X23,esk12_5(X23,X24,X25,X26,X27)))|~ic(X23))|iext(uri_owl_oneOf,X23,X24))|(((~iext(uri_rdf_first,X24,X25)|~iext(uri_rdf_rest,X24,X26))|~iext(uri_rdf_first,X26,X27))|~iext(uri_rdf_rest,X26,uri_rdf_nil))))&((((icext(X23,esk12_5(X23,X24,X25,X26,X27))|(esk12_5(X23,X24,X25,X26,X27)=X25|esk12_5(X23,X24,X25,X26,X27)=X27))|~ic(X23))|iext(uri_owl_oneOf,X23,X24))|(((~iext(uri_rdf_first,X24,X25)|~iext(uri_rdf_rest,X24,X26))|~iext(uri_rdf_first,X26,X27))|~iext(uri_rdf_rest,X26,uri_rdf_nil)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])])).
fof(c_0_10, plain, (((((((iext(uri_rdf_type,uri_ex_z,uri_ex_e)&iext(uri_owl_differentFrom,uri_ex_z,uri_ex_x))&iext(uri_owl_oneOf,uri_ex_e,esk192_0))&iext(uri_rdf_first,esk192_0,uri_ex_x))&iext(uri_rdf_rest,esk192_0,esk193_0))&iext(uri_rdf_first,esk193_0,uri_ex_y))&iext(uri_rdf_rest,esk193_0,uri_rdf_nil))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_1])])).
cnf(c_0_11,plain,(X6=X2|X6=X4|~iext(uri_rdf_rest,X1,uri_rdf_nil)|~iext(uri_rdf_first,X1,X2)|~iext(uri_rdf_rest,X3,X1)|~iext(uri_rdf_first,X3,X4)|~iext(uri_owl_oneOf,X5,X3)|~icext(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12,plain,(iext(uri_rdf_rest,esk193_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, (![X8]:![X9]:![X10]:![X11]:(((ip(X8)|~icext(uri_owl_SymmetricProperty,X8))&((~iext(X8,X9,X10)|iext(X8,X10,X9))|~icext(uri_owl_SymmetricProperty,X8)))&(((iext(X11,esk171_1(X11),esk172_1(X11))|~ip(X11))|icext(uri_owl_SymmetricProperty,X11))&((~iext(X11,esk172_1(X11),esk171_1(X11))|~ip(X11))|icext(uri_owl_SymmetricProperty,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
fof(c_0_14, plain, (![X4]:![X5]:![X6]:(~iext(X5,X4,X6)|ip(X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_15,plain,(X1=X2|X3=X2|~icext(X4,X2)|~iext(uri_rdf_first,esk193_0,X3)|~iext(uri_rdf_rest,X5,esk193_0)|~iext(uri_rdf_first,X5,X1)|~iext(uri_owl_oneOf,X4,X5)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16,plain,(iext(uri_rdf_first,esk193_0,uri_ex_y)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,plain,(icext(uri_owl_SymmetricProperty,X1)|~ip(X1)|~iext(X1,esk172_1(X1),esk171_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(ip(X1)|~iext(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, (![X8]:![X9]:![X10]:![X11]:((~iext(uri_owl_differentFrom,X8,X9)|X8!=X9)&(X10=X11|iext(uri_owl_differentFrom,X10,X11)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_20,plain,(uri_ex_y=X1|X2=X1|~icext(X3,X1)|~iext(uri_rdf_rest,X4,esk193_0)|~iext(uri_rdf_first,X4,X2)|~iext(uri_owl_oneOf,X3,X4)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21,plain,(iext(uri_rdf_rest,esk192_0,esk193_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22,plain,(icext(uri_owl_SymmetricProperty,X1)|~iext(X1,esk172_1(X1),esk171_1(X1))), inference(csr,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23,plain,(iext(uri_owl_differentFrom,X1,X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24,plain,(X1!=X2|~iext(uri_owl_differentFrom,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25,plain,(uri_ex_y=X1|X2=X1|~icext(X3,X1)|~iext(uri_rdf_first,esk192_0,X2)|~iext(uri_owl_oneOf,X3,esk192_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26,plain,(iext(uri_rdf_first,esk192_0,uri_ex_x)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_27, plain, (![X6]:![X7]:![X8]:![X9]:((~iext(uri_rdf_type,X6,X7)|icext(X7,X6))&(~icext(X9,X8)|iext(uri_rdf_type,X8,X9)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
cnf(c_0_28,plain,(icext(uri_owl_SymmetricProperty,X1)|iext(X1,esk171_1(X1),esk172_1(X1))|~ip(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_29,plain,(esk172_1(uri_owl_differentFrom)=esk171_1(uri_owl_differentFrom)|icext(uri_owl_SymmetricProperty,uri_owl_differentFrom)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30,plain,(ip(uri_owl_differentFrom)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_31,plain,(~iext(uri_owl_differentFrom,X1,X1)), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(uri_ex_x=X1|uri_ex_y=X1|~icext(X2,X1)|~iext(uri_owl_oneOf,X2,esk192_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33,plain,(iext(uri_owl_oneOf,uri_ex_e,esk192_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_34,plain,(icext(X1,X2)|~iext(uri_rdf_type,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35,plain,(iext(uri_rdf_type,uri_ex_z,uri_ex_e)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_36, plain, (![X8]:![X9]:![X10]:![X11]:((~iext(uri_owl_sameAs,X8,X9)|X8=X9)&(X10!=X11|iext(uri_owl_sameAs,X10,X11)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
cnf(c_0_37,plain,(iext(X1,X2,X3)|~icext(uri_owl_SymmetricProperty,X1)|~iext(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_38,plain,(icext(uri_owl_SymmetricProperty,uri_owl_differentFrom)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])]), c_0_31])).
fof(c_0_39, negated_conjecture, (~iext(uri_owl_sameAs,uri_ex_z,uri_ex_y)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_8])])).
cnf(c_0_40,plain,(uri_ex_y=X1|uri_ex_x=X1|~icext(uri_ex_e,X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41,plain,(icext(uri_ex_e,uri_ex_z)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42,plain,(iext(uri_owl_sameAs,X1,X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43,plain,(iext(uri_owl_differentFrom,X1,X2)|~iext(uri_owl_differentFrom,X2,X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44,plain,(iext(uri_owl_differentFrom,uri_ex_z,uri_ex_x)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_45,negated_conjecture,(~iext(uri_owl_sameAs,uri_ex_z,uri_ex_y)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46,plain,(uri_ex_z=uri_ex_x|uri_ex_y=uri_ex_z), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47,plain,(iext(uri_owl_sameAs,X1,X1)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_48,plain,(iext(uri_owl_differentFrom,uri_ex_x,uri_ex_z)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49,negated_conjecture,(uri_ex_z=uri_ex_x), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
cnf(c_0_50,plain,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49]), c_0_31]), ['proof']).
# SZS output end CNFRefutation.
EOF
