% Problem    : CSR012+1 : TPTP v6.1.0. Bugfixed v3.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n092.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:15:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___212_C18_F1_AE_Q12_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X12]:![X13]:(less_or_equal(X12,X13)<=>(less(X12,X13)|X12=X13))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', less_or_equal)).
fof(c_0_1, axiom, (![X4]:![X5]:![X2]:![X6]:![X7]:(((((happens(X4,X5)&initiates(X4,X2,X5))&less(n0,X7))&trajectory(X2,X5,X6,X7))&~(stoppedIn(X5,X2,plus(X5,X7))))=>holdsAt(X6,plus(X5,X7)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+0.ax', change_holding)).
fof(c_0_2, axiom, (![X12]:(less(X12,n1)<=>less_or_equal(X12,n0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', less1)).
fof(c_0_3, axiom, (![X10]:![X5]:![X11]:![X7]:((holdsAt(waterLevel(X10),X5)&X11=plus(X10,X7))=>trajectory(filling,X5,waterLevel(X11),X7))), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+1.ax', change_of_waterLevel)).
fof(c_0_4, axiom, (![X12]:![X13]:plus(X12,X13)=plus(X13,X12)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', symmetry_of_plus)).
fof(c_0_5, axiom, (![X4]:![X5]:(happens(X4,X5)<=>((X4=tapOn&X5=n0)|((holdsAt(waterLevel(n3),X5)&holdsAt(filling,X5))&X4=overflow)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+1.ax', happens_all_defn)).
fof(c_0_6, axiom, (![X4]:![X2]:![X5]:(initiates(X4,X2,X5)<=>((((X4=tapOn&X2=filling)|(X4=overflow&X2=spilling))|?[X9]:((holdsAt(waterLevel(X9),X5)&X4=tapOff)&X2=waterLevel(X9)))|?[X9]:((holdsAt(waterLevel(X9),X5)&X4=overflow)&X2=waterLevel(X9))))), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+1.ax', initiates_all_defn)).
fof(c_0_7, hypothesis, (holdsAt(waterLevel(n0),n0)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', waterLevel_0)).
fof(c_0_8, axiom, (plus(n0,n1)=n1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', plus0_1)).
fof(c_0_9, conjecture, (holdsAt(waterLevel(n1),n1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', waterLevel_1)).
fof(c_0_10, axiom, (![X1]:![X2]:![X3]:(stoppedIn(X1,X2,X3)<=>?[X4]:?[X5]:(((happens(X4,X5)&less(X1,X5))&less(X5,X3))&terminates(X4,X2,X5)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+0.ax', stoppedin_defn)).
fof(c_0_11, axiom, (~(?[X12]:less(X12,n0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', less0)).
fof(c_0_12, plain, (![X14]:![X15]:![X16]:![X17]:((~less_or_equal(X14,X15)|(less(X14,X15)|X14=X15))&((~less(X16,X17)|less_or_equal(X16,X17))&(X16!=X17|less_or_equal(X16,X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])).
fof(c_0_13, plain, (![X5]:![X4]:![X2]:(epred1_3(X2,X4,X5)<=>((((X4=tapOn&X2=filling)|(X4=overflow&X2=spilling))|?[X9]:((holdsAt(waterLevel(X9),X5)&X4=tapOff)&X2=waterLevel(X9)))|?[X9]:((holdsAt(waterLevel(X9),X5)&X4=overflow)&X2=waterLevel(X9))))), introduced(definition)).
fof(c_0_14, plain, (![X4]:![X5]:![X2]:![X6]:![X7]:(((((happens(X4,X5)&initiates(X4,X2,X5))&less(n0,X7))&trajectory(X2,X5,X6,X7))&~stoppedIn(X5,X2,plus(X5,X7)))=>holdsAt(X6,plus(X5,X7)))), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_15, plain, (![X13]:![X14]:((~less(X13,n1)|less_or_equal(X13,n0))&(~less_or_equal(X14,n0)|less(X14,n1)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_16,plain,(less_or_equal(X1,X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, (![X12]:![X13]:![X14]:![X15]:((~holdsAt(waterLevel(X12),X13)|X14!=plus(X12,X15))|trajectory(filling,X13,waterLevel(X14),X15))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
fof(c_0_18, plain, (![X10]:![X11]:![X12]:![X15]:![X16]:![X17]:![X18]:![X19]:(((((((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X11=overflow|X11=tapOn)))|~epred1_3(X12,X11,X10))&((X11=overflow|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X11=overflow|X11=tapOn)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X11=overflow|X11=tapOn)))|~epred1_3(X12,X11,X10)))&((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(X11=tapOff|(X11=overflow|X11=tapOn)))|~epred1_3(X12,X11,X10))&((X11=overflow|(X11=tapOff|(X11=overflow|X11=tapOn)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(X11=tapOff|(X11=overflow|X11=tapOn)))|~epred1_3(X12,X11,X10))))&((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(X12=waterLevel(esk10_3(X10,X11,X12))|(X11=overflow|X11=tapOn)))|~epred1_3(X12,X11,X10))&((X11=overflow|(X12=waterLevel(esk10_3(X10,X11,X12))|(X11=overflow|X11=tapOn)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(X12=waterLevel(esk10_3(X10,X11,X12))|(X11=overflow|X11=tapOn)))|~epred1_3(X12,X11,X10))))&((((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X12=spilling|X11=tapOn)))|~epred1_3(X12,X11,X10))&((X11=overflow|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X12=spilling|X11=tapOn)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X12=spilling|X11=tapOn)))|~epred1_3(X12,X11,X10)))&((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(X11=tapOff|(X12=spilling|X11=tapOn)))|~epred1_3(X12,X11,X10))&((X11=overflow|(X11=tapOff|(X12=spilling|X11=tapOn)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(X11=tapOff|(X12=spilling|X11=tapOn)))|~epred1_3(X12,X11,X10))))&((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(X12=waterLevel(esk10_3(X10,X11,X12))|(X12=spilling|X11=tapOn)))|~epred1_3(X12,X11,X10))&((X11=overflow|(X12=waterLevel(esk10_3(X10,X11,X12))|(X12=spilling|X11=tapOn)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(X12=waterLevel(esk10_3(X10,X11,X12))|(X12=spilling|X11=tapOn)))|~epred1_3(X12,X11,X10)))))&(((((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X11=overflow|X12=filling)))|~epred1_3(X12,X11,X10))&((X11=overflow|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X11=overflow|X12=filling)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X11=overflow|X12=filling)))|~epred1_3(X12,X11,X10)))&((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(X11=tapOff|(X11=overflow|X12=filling)))|~epred1_3(X12,X11,X10))&((X11=overflow|(X11=tapOff|(X11=overflow|X12=filling)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(X11=tapOff|(X11=overflow|X12=filling)))|~epred1_3(X12,X11,X10))))&((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(X12=waterLevel(esk10_3(X10,X11,X12))|(X11=overflow|X12=filling)))|~epred1_3(X12,X11,X10))&((X11=overflow|(X12=waterLevel(esk10_3(X10,X11,X12))|(X11=overflow|X12=filling)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(X12=waterLevel(esk10_3(X10,X11,X12))|(X11=overflow|X12=filling)))|~epred1_3(X12,X11,X10))))&((((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X12=spilling|X12=filling)))|~epred1_3(X12,X11,X10))&((X11=overflow|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X12=spilling|X12=filling)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(holdsAt(waterLevel(esk10_3(X10,X11,X12)),X10)|(X12=spilling|X12=filling)))|~epred1_3(X12,X11,X10)))&((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(X11=tapOff|(X12=spilling|X12=filling)))|~epred1_3(X12,X11,X10))&((X11=overflow|(X11=tapOff|(X12=spilling|X12=filling)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(X11=tapOff|(X12=spilling|X12=filling)))|~epred1_3(X12,X11,X10))))&((((holdsAt(waterLevel(esk11_3(X10,X11,X12)),X10)|(X12=waterLevel(esk10_3(X10,X11,X12))|(X12=spilling|X12=filling)))|~epred1_3(X12,X11,X10))&((X11=overflow|(X12=waterLevel(esk10_3(X10,X11,X12))|(X12=spilling|X12=filling)))|~epred1_3(X12,X11,X10)))&((X12=waterLevel(esk11_3(X10,X11,X12))|(X12=waterLevel(esk10_3(X10,X11,X12))|(X12=spilling|X12=filling)))|~epred1_3(X12,X11,X10))))))&(((((X16!=tapOn|X17!=filling)|epred1_3(X17,X16,X15))&((X16!=overflow|X17!=spilling)|epred1_3(X17,X16,X15)))&(((~holdsAt(waterLevel(X18),X15)|X16!=tapOff)|X17!=waterLevel(X18))|epred1_3(X17,X16,X15)))&(((~holdsAt(waterLevel(X19),X15)|X16!=overflow)|X17!=waterLevel(X19))|epred1_3(X17,X16,X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])])).
fof(c_0_19, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:(((((~happens(X8,X9)|~initiates(X8,X10,X9))|~less(n0,X12))|~trajectory(X10,X9,X11,X12))|stoppedIn(X9,X10,plus(X9,X12)))|holdsAt(X11,plus(X9,X12)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
cnf(c_0_20,plain,(less(X1,n1)|~less_or_equal(X1,n0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21,plain,(less_or_equal(X1,X1)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(trajectory(filling,X1,waterLevel(X2),X3)|X2!=plus(X4,X3)|~holdsAt(waterLevel(X4),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, (![X14]:![X15]:plus(X14,X15)=plus(X15,X14)), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_24, plain, (![X6]:![X7]:![X8]:![X9]:((((((holdsAt(waterLevel(n3),X7)|X6=tapOn)|~happens(X6,X7))&((holdsAt(filling,X7)|X6=tapOn)|~happens(X6,X7)))&((X6=overflow|X6=tapOn)|~happens(X6,X7)))&((((holdsAt(waterLevel(n3),X7)|X7=n0)|~happens(X6,X7))&((holdsAt(filling,X7)|X7=n0)|~happens(X6,X7)))&((X6=overflow|X7=n0)|~happens(X6,X7))))&(((X8!=tapOn|X9!=n0)|happens(X8,X9))&(((~holdsAt(waterLevel(n3),X9)|~holdsAt(filling,X9))|X8!=overflow)|happens(X8,X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
fof(c_0_25, axiom, (![X4]:![X2]:![X5]:(initiates(X4,X2,X5)<=>epred1_3(X2,X4,X5))), inference(apply_def,[status(thm)],[c_0_6, c_0_13])).
cnf(c_0_26,plain,(epred1_3(X1,X2,X3)|X1!=filling|X2!=tapOn), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(holdsAt(X1,plus(X2,X3))|stoppedIn(X2,X4,plus(X2,X3))|~trajectory(X4,X2,X1,X3)|~less(n0,X3)|~initiates(X5,X4,X2)|~happens(X5,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,plain,(less(n0,n1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_29,plain,(trajectory(filling,X1,waterLevel(plus(X2,X3)),X3)|~holdsAt(waterLevel(X2),X1)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_30,hypothesis,(holdsAt(waterLevel(n0),n0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_31,plain,(plus(n0,n1)=n1), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_32,plain,(plus(X1,X2)=plus(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_33, negated_conjecture, (~holdsAt(waterLevel(n1),n1)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_9])])).
cnf(c_0_34,plain,(happens(X1,X2)|X2!=n0|X1!=tapOn), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_35, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:((~initiates(X6,X7,X8)|epred1_3(X7,X6,X8))&(~epred1_3(X10,X9,X11)|initiates(X9,X10,X11)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])).
cnf(c_0_36,plain,(epred1_3(X1,tapOn,X2)|filling!=X1), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_37,plain,(holdsAt(X1,plus(X2,n1))|stoppedIn(X2,X3,plus(X2,n1))|~trajectory(X3,X2,X1,n1)|~initiates(X4,X3,X2)|~happens(X4,X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_38,hypothesis,(trajectory(filling,n0,waterLevel(plus(n0,X1)),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_39,plain,(plus(n1,n0)=n1), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40,negated_conjecture,(~holdsAt(waterLevel(n1),n1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41,plain,(happens(X1,n0)|tapOn!=X1), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_42,plain,(initiates(X1,X2,X3)|~epred1_3(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43,plain,(epred1_3(filling,tapOn,X1)), inference(er,[status(thm)],[c_0_36])).
fof(c_0_44, plain, (![X6]:![X7]:![X8]:![X11]:![X12]:![X13]:![X14]:![X15]:(((((happens(esk1_3(X6,X7,X8),esk2_3(X6,X7,X8))|~stoppedIn(X6,X7,X8))&(less(X6,esk2_3(X6,X7,X8))|~stoppedIn(X6,X7,X8)))&(less(esk2_3(X6,X7,X8),X8)|~stoppedIn(X6,X7,X8)))&(terminates(esk1_3(X6,X7,X8),X7,esk2_3(X6,X7,X8))|~stoppedIn(X6,X7,X8)))&((((~happens(X14,X15)|~less(X11,X15))|~less(X15,X13))|~terminates(X14,X12,X15))|stoppedIn(X11,X12,X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])])).
cnf(c_0_45,hypothesis,(stoppedIn(n0,filling,n1)|~initiates(X1,filling,n0)|~happens(X1,n0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_32]), c_0_39]), c_0_32]), c_0_39]), c_0_32]), c_0_39]), c_0_40])).
cnf(c_0_46,plain,(happens(tapOn,n0)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_47,plain,(initiates(tapOn,filling,X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48,plain,(less(esk2_3(X1,X2,X3),X3)|~stoppedIn(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_49,hypothesis,(stoppedIn(n0,filling,n1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
cnf(c_0_50,plain,(less_or_equal(X1,n0)|~less(X1,n1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_51,hypothesis,(less(esk2_3(n0,filling,n1),n1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
fof(c_0_52, plain, (![X13]:~less(X13,n0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).
cnf(c_0_53,plain,(less(X1,esk2_3(X1,X2,X3))|~stoppedIn(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54,plain,(X1=X2|less(X1,X2)|~less_or_equal(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_55,hypothesis,(less_or_equal(esk2_3(n0,filling,n1),n0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_56,plain,(~less(X1,n0)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_57,hypothesis,(less(n0,esk2_3(n0,filling,n1))), inference(spm,[status(thm)],[c_0_53, c_0_49])).
cnf(c_0_58,hypothesis,(esk2_3(n0,filling,n1)=n0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56])).
cnf(c_0_59,hypothesis,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_58]), c_0_56]), ['proof']).
# SZS output end CNFRefutation.
EOF
