% Problem    : LAT348+4 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n126.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:08:33 CDT 2014
% CPUTime    : 
# SinE strategy is gf120_h_gu_R02_F100_L20000
# Trying AutoSched0 for 150 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2i
# and selection function SelectGrCQArEqFirst.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X1]:(((((((~(v3_struct_0(X1))&v2_orders_2(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v2_yellow_0(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>((((~(v3_struct_0(k2_yellow_1(k9_waybel_0(X1))))&v4_yellow_0(k2_yellow_1(k9_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&v7_yellow_0(k2_yellow_1(k9_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&v4_waybel_0(k2_yellow_1(k9_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&m1_yellow_0(k2_yellow_1(k9_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', t6_waybel22)).
fof(c_0_1, axiom, (![X1]:((((~(v3_struct_0(X1))&v2_orders_2(X1))&v3_orders_2(X1))&l1_orders_2(X1))=>k9_waybel_0(X1)=k8_waybel_0(k7_lattice3(X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+537.ax', t7_waybel16)).
fof(c_0_2, axiom, (![X1]:((v2_lattice3(X1)&l1_orders_2(X1))=>((~(v3_struct_0(k7_lattice3(X1)))&v1_orders_2(k7_lattice3(X1)))&v1_lattice3(k7_lattice3(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+495.ax', fc5_yellow_7)).
fof(c_0_3, axiom, (![X1]:((((((v2_orders_2(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v1_lattice3(X1))&v1_yellow_0(X1))&l1_orders_2(X1))=>((((~(v3_struct_0(k2_yellow_1(k8_waybel_0(X1))))&v4_yellow_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&v7_yellow_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&v4_waybel_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&m1_yellow_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+509.ax', t8_waybel13)).
fof(c_0_4, axiom, (![X1]:(l1_orders_2(X1)=>(v1_orders_2(k7_lattice3(X1))&l1_orders_2(k7_lattice3(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+335.ax', dt_k7_lattice3)).
fof(c_0_5, axiom, (![X1]:(l1_orders_2(X1)=>u1_struct_0(X1)=u1_struct_0(k7_lattice3(X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+493.ax', t12_yellow_6)).
fof(c_0_6, axiom, (![X1]:((v2_yellow_0(X1)&l1_orders_2(X1))=>(v1_orders_2(k7_lattice3(X1))&v1_yellow_0(k7_lattice3(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+495.ax', fc10_yellow_7)).
fof(c_0_7, axiom, (![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>(v1_orders_2(k7_lattice3(X1))&v4_orders_2(k7_lattice3(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+495.ax', fc3_yellow_7)).
fof(c_0_8, axiom, (![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>(v1_orders_2(k7_lattice3(X1))&v3_orders_2(k7_lattice3(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+495.ax', fc2_yellow_7)).
fof(c_0_9, axiom, (![X1]:((v2_orders_2(X1)&l1_orders_2(X1))=>(v1_orders_2(k7_lattice3(X1))&v2_orders_2(k7_lattice3(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+495.ax', fc1_yellow_7)).
fof(c_0_10, axiom, (![X1]:((((((v2_orders_2(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v2_lattice3(X1))&v2_yellow_0(X1))&l1_orders_2(X1))=>(((((((((~(v3_struct_0(k2_yellow_1(k9_waybel_0(X1))))&v1_orders_2(k2_yellow_1(k9_waybel_0(X1))))&v2_orders_2(k2_yellow_1(k9_waybel_0(X1))))&v3_orders_2(k2_yellow_1(k9_waybel_0(X1))))&v4_orders_2(k2_yellow_1(k9_waybel_0(X1))))&v2_lattice3(k2_yellow_1(k9_waybel_0(X1))))&v3_lattice3(k2_yellow_1(k9_waybel_0(X1))))&v1_yellow_0(k2_yellow_1(k9_waybel_0(X1))))&v24_waybel_0(k2_yellow_1(k9_waybel_0(X1))))&v25_waybel_0(k2_yellow_1(k9_waybel_0(X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+537.ax', fc2_waybel16)).
fof(c_0_11, negated_conjecture, (~(![X1]:(((((((~v3_struct_0(X1)&v2_orders_2(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v2_yellow_0(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>((((~v3_struct_0(k2_yellow_1(k9_waybel_0(X1)))&v4_yellow_0(k2_yellow_1(k9_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&v7_yellow_0(k2_yellow_1(k9_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&v4_waybel_0(k2_yellow_1(k9_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&m1_yellow_0(k2_yellow_1(k9_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
fof(c_0_12, plain, (![X1]:((((~v3_struct_0(X1)&v2_orders_2(X1))&v3_orders_2(X1))&l1_orders_2(X1))=>k9_waybel_0(X1)=k8_waybel_0(k7_lattice3(X1)))), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_13, plain, (![X1]:((v2_lattice3(X1)&l1_orders_2(X1))=>((~v3_struct_0(k7_lattice3(X1))&v1_orders_2(k7_lattice3(X1)))&v1_lattice3(k7_lattice3(X1))))), inference(fof_simplification,[status(thm)],[c_0_2])).
fof(c_0_14, plain, (![X1]:((((((v2_orders_2(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v1_lattice3(X1))&v1_yellow_0(X1))&l1_orders_2(X1))=>((((~v3_struct_0(k2_yellow_1(k8_waybel_0(X1)))&v4_yellow_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&v7_yellow_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&v4_waybel_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1))))&m1_yellow_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1)))))), inference(fof_simplification,[status(thm)],[c_0_3])).
fof(c_0_15, plain, (![X2]:((v1_orders_2(k7_lattice3(X2))|~l1_orders_2(X2))&(l1_orders_2(k7_lattice3(X2))|~l1_orders_2(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
fof(c_0_16, negated_conjecture, ((((((((~v3_struct_0(esk1_0)&v2_orders_2(esk1_0))&v3_orders_2(esk1_0))&v4_orders_2(esk1_0))&v2_yellow_0(esk1_0))&v2_lattice3(esk1_0))&l1_orders_2(esk1_0))&((((v3_struct_0(k2_yellow_1(k9_waybel_0(esk1_0)))|~v4_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0))))|~v7_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0))))|~v4_waybel_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0))))|~m1_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_17, plain, (![X2]:((((v3_struct_0(X2)|~v2_orders_2(X2))|~v3_orders_2(X2))|~l1_orders_2(X2))|k9_waybel_0(X2)=k8_waybel_0(k7_lattice3(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
fof(c_0_18, plain, (![X2]:(~l1_orders_2(X2)|u1_struct_0(X2)=u1_struct_0(k7_lattice3(X2)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
fof(c_0_19, plain, (![X2]:((v1_orders_2(k7_lattice3(X2))|(~v2_yellow_0(X2)|~l1_orders_2(X2)))&(v1_yellow_0(k7_lattice3(X2))|(~v2_yellow_0(X2)|~l1_orders_2(X2))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
fof(c_0_20, plain, (![X2]:(((~v3_struct_0(k7_lattice3(X2))|(~v2_lattice3(X2)|~l1_orders_2(X2)))&(v1_orders_2(k7_lattice3(X2))|(~v2_lattice3(X2)|~l1_orders_2(X2))))&(v1_lattice3(k7_lattice3(X2))|(~v2_lattice3(X2)|~l1_orders_2(X2))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
fof(c_0_21, plain, (![X2]:((v1_orders_2(k7_lattice3(X2))|(~v4_orders_2(X2)|~l1_orders_2(X2)))&(v4_orders_2(k7_lattice3(X2))|(~v4_orders_2(X2)|~l1_orders_2(X2))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
fof(c_0_22, plain, (![X2]:((v1_orders_2(k7_lattice3(X2))|(~v3_orders_2(X2)|~l1_orders_2(X2)))&(v3_orders_2(k7_lattice3(X2))|(~v3_orders_2(X2)|~l1_orders_2(X2))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
fof(c_0_23, plain, (![X2]:((v1_orders_2(k7_lattice3(X2))|(~v2_orders_2(X2)|~l1_orders_2(X2)))&(v2_orders_2(k7_lattice3(X2))|(~v2_orders_2(X2)|~l1_orders_2(X2))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
fof(c_0_24, plain, (![X1]:((((((v2_orders_2(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v2_lattice3(X1))&v2_yellow_0(X1))&l1_orders_2(X1))=>(((((((((~v3_struct_0(k2_yellow_1(k9_waybel_0(X1)))&v1_orders_2(k2_yellow_1(k9_waybel_0(X1))))&v2_orders_2(k2_yellow_1(k9_waybel_0(X1))))&v3_orders_2(k2_yellow_1(k9_waybel_0(X1))))&v4_orders_2(k2_yellow_1(k9_waybel_0(X1))))&v2_lattice3(k2_yellow_1(k9_waybel_0(X1))))&v3_lattice3(k2_yellow_1(k9_waybel_0(X1))))&v1_yellow_0(k2_yellow_1(k9_waybel_0(X1))))&v24_waybel_0(k2_yellow_1(k9_waybel_0(X1))))&v25_waybel_0(k2_yellow_1(k9_waybel_0(X1)))))), inference(fof_simplification,[status(thm)],[c_0_10])).
fof(c_0_25, plain, (![X2]:(((((~v3_struct_0(k2_yellow_1(k8_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v1_lattice3(X2))|~v1_yellow_0(X2))|~l1_orders_2(X2)))&(v4_yellow_0(k2_yellow_1(k8_waybel_0(X2)),k3_yellow_1(u1_struct_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v1_lattice3(X2))|~v1_yellow_0(X2))|~l1_orders_2(X2))))&(v7_yellow_0(k2_yellow_1(k8_waybel_0(X2)),k3_yellow_1(u1_struct_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v1_lattice3(X2))|~v1_yellow_0(X2))|~l1_orders_2(X2))))&(v4_waybel_0(k2_yellow_1(k8_waybel_0(X2)),k3_yellow_1(u1_struct_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v1_lattice3(X2))|~v1_yellow_0(X2))|~l1_orders_2(X2))))&(m1_yellow_0(k2_yellow_1(k8_waybel_0(X2)),k3_yellow_1(u1_struct_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v1_lattice3(X2))|~v1_yellow_0(X2))|~l1_orders_2(X2))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
cnf(c_0_26,plain,(l1_orders_2(k7_lattice3(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27,negated_conjecture,(l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28,plain,(k9_waybel_0(X1)=k8_waybel_0(k7_lattice3(X1))|v3_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v2_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29,negated_conjecture,(v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_30,negated_conjecture,(v2_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_31,negated_conjecture,(~v3_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_32,plain,(u1_struct_0(X1)=u1_struct_0(k7_lattice3(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_33,plain,(v1_yellow_0(k7_lattice3(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34,negated_conjecture,(v2_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_35,plain,(v1_lattice3(k7_lattice3(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36,negated_conjecture,(v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_37,plain,(v4_orders_2(k7_lattice3(X1))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_38,negated_conjecture,(v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_39,plain,(v3_orders_2(k7_lattice3(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_40,plain,(v2_orders_2(k7_lattice3(X1))|~l1_orders_2(X1)|~v2_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_41, plain, (![X2]:((((((((((~v3_struct_0(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2)))&(v1_orders_2(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2))))&(v2_orders_2(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2))))&(v3_orders_2(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2))))&(v4_orders_2(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2))))&(v2_lattice3(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2))))&(v3_lattice3(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2))))&(v1_yellow_0(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2))))&(v24_waybel_0(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2))))&(v25_waybel_0(k2_yellow_1(k9_waybel_0(X2)))|(((((~v2_orders_2(X2)|~v3_orders_2(X2))|~v4_orders_2(X2))|~v2_lattice3(X2))|~v2_yellow_0(X2))|~l1_orders_2(X2))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
cnf(c_0_42,plain,(m1_yellow_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v2_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_43,negated_conjecture,(l1_orders_2(k7_lattice3(esk1_0))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_44,negated_conjecture,(k8_waybel_0(k7_lattice3(esk1_0))=k9_waybel_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_27]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_45,negated_conjecture,(u1_struct_0(k7_lattice3(esk1_0))=u1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_27])).
cnf(c_0_46,negated_conjecture,(v1_yellow_0(k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_27]), c_0_34])])).
cnf(c_0_47,negated_conjecture,(v1_lattice3(k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_27]), c_0_36])])).
cnf(c_0_48,negated_conjecture,(v4_orders_2(k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_27]), c_0_38])])).
cnf(c_0_49,negated_conjecture,(v3_orders_2(k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_27]), c_0_29])])).
cnf(c_0_50,negated_conjecture,(v2_orders_2(k7_lattice3(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_27]), c_0_30])])).
cnf(c_0_51,plain,(~l1_orders_2(X1)|~v2_yellow_0(X1)|~v2_lattice3(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v2_orders_2(X1)|~v3_struct_0(k2_yellow_1(k9_waybel_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52,negated_conjecture,(v3_struct_0(k2_yellow_1(k9_waybel_0(esk1_0)))|~m1_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))|~v4_waybel_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))|~v7_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))|~v4_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_53,negated_conjecture,(m1_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_54,negated_conjecture,(~v3_struct_0(k2_yellow_1(k9_waybel_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_27]), c_0_34]), c_0_36]), c_0_38]), c_0_29]), c_0_30])])).
cnf(c_0_55,plain,(v4_yellow_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v2_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_56,negated_conjecture,(~v4_waybel_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))|~v7_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))|~v4_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53])]), c_0_54])).
cnf(c_0_57,negated_conjecture,(v4_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_58,plain,(v7_yellow_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v2_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_59,negated_conjecture,(~v4_waybel_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))|~v7_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57])])).
cnf(c_0_60,negated_conjecture,(v7_yellow_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_61,plain,(v4_waybel_0(k2_yellow_1(k8_waybel_0(X1)),k3_yellow_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v2_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_62,negated_conjecture,(~v4_waybel_0(k2_yellow_1(k9_waybel_0(esk1_0)),k3_yellow_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_60])])).
cnf(c_0_63,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_48]), c_0_49]), c_0_50])]), c_0_62]), ['proof']).
# SZS output end CNFRefutation.
EOF
