% Problem    : SEU449+3 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n190.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:09:22 CDT 2014
% CPUTime    : 
# SinE strategy is gf120_h_gu_RUU_F100_L00500
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:![X2]:r1_tarski(X1,X1)), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+1.ax', reflexivity_r1_tarski)).
fof(c_0_1, conjecture, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X2))=>![X5]:(m2_relset_1(X5,X1,X2)=>(r1_tarski(X3,k8_relset_2(X2,X1,k8_relset_2(X1,X2,X3,X5),k6_relset_1(X1,X2,X5)))&r1_tarski(X4,k8_relset_2(X1,X2,k8_relset_2(X2,X1,X4,k6_relset_1(X1,X2,X5)),X5))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', t63_relset_2)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X2))=>![X5]:(m2_relset_1(X5,X1,X2)=>(r1_tarski(X3,k8_relset_2(X2,X1,X4,k6_relset_1(X1,X2,X5)))<=>r1_tarski(X4,k8_relset_2(X1,X2,X3,X5))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', t61_relset_2)).
fof(c_0_3, axiom, (![X1]:![X2]:![X3]:![X4]:((m1_subset_1(X3,k1_zfmisc_1(X1))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>m1_subset_1(k8_relset_2(X1,X2,X3,X4),k1_zfmisc_1(X2)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', dt_k8_relset_2)).
fof(c_0_4, axiom, (![X1]:![X2]:![X3]:(m2_relset_1(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+16.ax', dt_m2_relset_1)).
fof(c_0_5, axiom, (![X1]:![X2]:![X3]:(m1_relset_1(X3,X1,X2)=>m2_relset_1(k6_relset_1(X1,X2,X3),X2,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+16.ax', dt_k6_relset_1)).
fof(c_0_6, axiom, (![X1]:![X2]:![X3]:(m2_relset_1(X3,X1,X2)<=>m1_relset_1(X3,X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SET007/SET007+16.ax', redefinition_m2_relset_1)).
fof(c_0_7, plain, (![X1]:r1_tarski(X1,X1)), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_8, negated_conjecture, (~(![X1]:![X2]:![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X2))=>![X5]:(m2_relset_1(X5,X1,X2)=>(r1_tarski(X3,k8_relset_2(X2,X1,k8_relset_2(X1,X2,X3,X5),k6_relset_1(X1,X2,X5)))&r1_tarski(X4,k8_relset_2(X1,X2,k8_relset_2(X2,X1,X4,k6_relset_1(X1,X2,X5)),X5)))))))), inference(assume_negation,[status(cth)],[c_0_1])).
fof(c_0_9, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((((~r1_tarski(X8,k8_relset_2(X7,X6,X9,k6_relset_1(X6,X7,X10)))|r1_tarski(X9,k8_relset_2(X6,X7,X8,X10)))|~m2_relset_1(X10,X6,X7))|~m1_subset_1(X9,k1_zfmisc_1(X7)))|~m1_subset_1(X8,k1_zfmisc_1(X6)))&((((~r1_tarski(X9,k8_relset_2(X6,X7,X8,X10))|r1_tarski(X8,k8_relset_2(X7,X6,X9,k6_relset_1(X6,X7,X10))))|~m2_relset_1(X10,X6,X7))|~m1_subset_1(X9,k1_zfmisc_1(X7)))|~m1_subset_1(X8,k1_zfmisc_1(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
fof(c_0_10, plain, (![X2]:r1_tarski(X2,X2)), inference(variable_rename,[status(thm)],[c_0_7])).
fof(c_0_11, negated_conjecture, ((m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk4_0,k1_zfmisc_1(esk2_0))&(m2_relset_1(esk5_0,esk1_0,esk2_0)&(~r1_tarski(esk3_0,k8_relset_2(esk2_0,esk1_0,k8_relset_2(esk1_0,esk2_0,esk3_0,esk5_0),k6_relset_1(esk1_0,esk2_0,esk5_0)))|~r1_tarski(esk4_0,k8_relset_2(esk1_0,esk2_0,k8_relset_2(esk2_0,esk1_0,esk4_0,k6_relset_1(esk1_0,esk2_0,esk5_0)),esk5_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
cnf(c_0_12,plain,(r1_tarski(X3,k8_relset_2(X2,X4,X1,X5))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X4))|~m2_relset_1(X5,X2,X4)|~r1_tarski(X1,k8_relset_2(X4,X2,X3,k6_relset_1(X2,X4,X5)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13,plain,(r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,negated_conjecture,(~r1_tarski(esk4_0,k8_relset_2(esk1_0,esk2_0,k8_relset_2(esk2_0,esk1_0,esk4_0,k6_relset_1(esk1_0,esk2_0,esk5_0)),esk5_0))|~r1_tarski(esk3_0,k8_relset_2(esk2_0,esk1_0,k8_relset_2(esk1_0,esk2_0,esk3_0,esk5_0),k6_relset_1(esk1_0,esk2_0,esk5_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,plain,(r1_tarski(X1,k8_relset_2(X2,X3,k8_relset_2(X3,X2,X1,k6_relset_1(X2,X3,X4)),X4))|~m2_relset_1(X4,X2,X3)|~m1_subset_1(k8_relset_2(X3,X2,X1,k6_relset_1(X2,X3,X4)),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16,negated_conjecture,(m2_relset_1(esk5_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18,negated_conjecture,(~m1_subset_1(k8_relset_2(esk2_0,esk1_0,esk4_0,k6_relset_1(esk1_0,esk2_0,esk5_0)),k1_zfmisc_1(esk1_0))|~r1_tarski(esk3_0,k8_relset_2(esk2_0,esk1_0,k8_relset_2(esk1_0,esk2_0,esk3_0,esk5_0),k6_relset_1(esk1_0,esk2_0,esk5_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])])).
cnf(c_0_19,plain,(r1_tarski(X1,k8_relset_2(X4,X2,X3,k6_relset_1(X2,X4,X5)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X4))|~m2_relset_1(X5,X2,X4)|~r1_tarski(X3,k8_relset_2(X2,X4,X1,X5))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_21, plain, (![X5]:![X6]:![X7]:![X8]:((~m1_subset_1(X7,k1_zfmisc_1(X5))|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))|m1_subset_1(k8_relset_2(X5,X6,X7,X8),k1_zfmisc_1(X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
cnf(c_0_22,negated_conjecture,(~m1_subset_1(k8_relset_2(esk2_0,esk1_0,esk4_0,k6_relset_1(esk1_0,esk2_0,esk5_0)),k1_zfmisc_1(esk1_0))|~m1_subset_1(k8_relset_2(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_16]), c_0_20]), c_0_13])])).
cnf(c_0_23,plain,(m1_subset_1(k8_relset_2(X1,X2,X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_24, plain, (![X4]:![X5]:![X6]:(~m2_relset_1(X6,X4,X5)|m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
cnf(c_0_25,negated_conjecture,(~m1_subset_1(k6_relset_1(esk1_0,esk2_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))|~m1_subset_1(k8_relset_2(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_17])])).
cnf(c_0_26,plain,(m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m2_relset_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_27, plain, (![X4]:![X5]:![X6]:(~m1_relset_1(X6,X4,X5)|m2_relset_1(k6_relset_1(X4,X5,X6),X5,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
cnf(c_0_28,negated_conjecture,(~m2_relset_1(k6_relset_1(esk1_0,esk2_0,esk5_0),esk2_0,esk1_0)|~m1_subset_1(k8_relset_2(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29,plain,(m2_relset_1(k6_relset_1(X1,X2,X3),X2,X1)|~m1_relset_1(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_30,negated_conjecture,(~m1_relset_1(esk5_0,esk1_0,esk2_0)|~m1_subset_1(k8_relset_2(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_31, plain, (![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:((~m2_relset_1(X6,X4,X5)|m1_relset_1(X6,X4,X5))&(~m1_relset_1(X9,X7,X8)|m2_relset_1(X9,X7,X8)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_32,negated_conjecture,(~m1_relset_1(esk5_0,esk1_0,esk2_0)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_23]), c_0_20])])).
cnf(c_0_33,plain,(m1_relset_1(X1,X2,X3)|~m2_relset_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_34,negated_conjecture,(~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_16])])).
cnf(c_0_35,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_26]), c_0_16])]), ['proof']).
# SZS output end CNFRefutation.
EOF
