% Problem    : NUM924+2 : TPTP v6.1.0. Released v5.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n143.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:26:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___042_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X13]:bit0(X13)=plus_plus_int(X13,X13)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_180_Bit0__def)).
fof(c_0_1, axiom, (![X13]:bit1(X13)=plus_plus_int(plus_plus_int(one_one_int,X13),X13)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_234_Bit1__def)).
fof(c_0_2, conjecture, (ord_less_int(plus_plus_int(power_power_int(s,number_number_of_nat(bit0(bit1(pls)))),one_one_int),zero_zero_int)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_3, axiom, (plus_plus_int(power_power_int(s,number_number_of_nat(bit0(bit1(pls)))),one_one_int)=times_times_int(plus_plus_int(times_times_int(number_number_of_int(bit0(bit0(bit1(pls)))),m),one_one_int),t)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_3_t)).
fof(c_0_4, axiom, (![X14]:![X12]:plus_plus_int(X14,X12)=plus_plus_int(X12,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_132_zadd__commute)).
fof(c_0_5, axiom, (![X18]:![X19]:![X20]:plus_plus_int(plus_plus_int(X18,X19),X20)=plus_plus_int(X18,plus_plus_int(X19,X20))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_130_zadd__assoc)).
fof(c_0_6, axiom, (![X3]:![X4]:![X14]:plus_plus_int(X3,plus_plus_int(X4,X14))=plus_plus_int(X4,plus_plus_int(X3,X14))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_131_zadd__left__commute)).
fof(c_0_7, axiom, (![X14]:![X12]:times_times_int(X14,X12)=times_times_int(X12,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_35_zmult__commute)).
fof(c_0_8, axiom, (![X156]:times_times_int(X156,zero_zero_int)=zero_zero_int), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_362_comm__semiring__1__class_Onormalizing__semiring__rules_I10_J)).
fof(c_0_9, axiom, (ord_less_int(times_times_int(plus_plus_int(times_times_int(number_number_of_int(bit0(bit0(bit1(pls)))),m),one_one_int),t),times_times_int(plus_plus_int(times_times_int(number_number_of_int(bit0(bit0(bit1(pls)))),m),one_one_int),zero_zero_int))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_2__096_I4_A_K_Am_A_L_A1_J_A_K_At_A_060_A_I4_A_K_Am_A_L_A1_J_A_K_A0_096)).
fof(c_0_10, axiom, (pls=zero_zero_int), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_170_Pls__def)).
fof(c_0_11, plain, (![X14]:bit0(X14)=plus_plus_int(X14,X14)), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_12, plain, (![X14]:bit1(X14)=plus_plus_int(plus_plus_int(one_one_int,X14),X14)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_13, negated_conjecture, (~ord_less_int(plus_plus_int(power_power_int(s,number_number_of_nat(bit0(bit1(pls)))),one_one_int),zero_zero_int)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_2])])).
cnf(c_0_14,plain,(plus_plus_int(power_power_int(s,number_number_of_nat(bit0(bit1(pls)))),one_one_int)=times_times_int(plus_plus_int(times_times_int(number_number_of_int(bit0(bit0(bit1(pls)))),m),one_one_int),t)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_15,plain,(bit0(X1)=plus_plus_int(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,plain,(bit1(X1)=plus_plus_int(plus_plus_int(one_one_int,X1),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, (![X15]:![X16]:plus_plus_int(X15,X16)=plus_plus_int(X16,X15)), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_18, plain, (![X21]:![X22]:![X23]:plus_plus_int(plus_plus_int(X21,X22),X23)=plus_plus_int(X21,plus_plus_int(X22,X23))), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_19, plain, (![X15]:![X16]:![X17]:plus_plus_int(X15,plus_plus_int(X16,X17))=plus_plus_int(X16,plus_plus_int(X15,X17))), inference(variable_rename,[status(thm)],[c_0_6])).
fof(c_0_20, plain, (![X15]:![X16]:times_times_int(X15,X16)=times_times_int(X16,X15)), inference(variable_rename,[status(thm)],[c_0_7])).
fof(c_0_21, plain, (![X157]:times_times_int(X157,zero_zero_int)=zero_zero_int), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_22,negated_conjecture,(~ord_less_int(plus_plus_int(power_power_int(s,number_number_of_nat(bit0(bit1(pls)))),one_one_int),zero_zero_int)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23,plain,(ord_less_int(times_times_int(plus_plus_int(times_times_int(number_number_of_int(bit0(bit0(bit1(pls)))),m),one_one_int),t),times_times_int(plus_plus_int(times_times_int(number_number_of_int(bit0(bit0(bit1(pls)))),m),one_one_int),zero_zero_int))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24,plain,(times_times_int(plus_plus_int(times_times_int(number_number_of_int(plus_plus_int(plus_plus_int(plus_plus_int(plus_plus_int(one_one_int,pls),pls),plus_plus_int(plus_plus_int(one_one_int,pls),pls)),plus_plus_int(plus_plus_int(plus_plus_int(one_one_int,pls),pls),plus_plus_int(plus_plus_int(one_one_int,pls),pls)))),m),one_one_int),t)=plus_plus_int(power_power_int(s,number_number_of_nat(plus_plus_int(plus_plus_int(plus_plus_int(one_one_int,pls),pls),plus_plus_int(plus_plus_int(one_one_int,pls),pls)))),one_one_int)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_15]), c_0_15]), c_0_16]), c_0_16]), c_0_16]), c_0_16]), c_0_16]), c_0_16])).
cnf(c_0_25,plain,(plus_plus_int(X1,X2)=plus_plus_int(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26,plain,(plus_plus_int(plus_plus_int(X1,X2),X3)=plus_plus_int(X1,plus_plus_int(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(plus_plus_int(X1,plus_plus_int(X2,X3))=plus_plus_int(X2,plus_plus_int(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,plain,(times_times_int(X1,X2)=times_times_int(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29,plain,(times_times_int(X1,zero_zero_int)=zero_zero_int), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30,plain,(pls=zero_zero_int), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_31,negated_conjecture,(~ord_less_int(plus_plus_int(power_power_int(s,number_number_of_nat(plus_plus_int(plus_plus_int(plus_plus_int(one_one_int,pls),pls),plus_plus_int(plus_plus_int(one_one_int,pls),pls)))),one_one_int),zero_zero_int)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_15]), c_0_16]), c_0_16])).
cnf(c_0_32,plain,(ord_less_int(times_times_int(plus_plus_int(times_times_int(number_number_of_int(plus_plus_int(plus_plus_int(plus_plus_int(plus_plus_int(one_one_int,pls),pls),plus_plus_int(plus_plus_int(one_one_int,pls),pls)),plus_plus_int(plus_plus_int(plus_plus_int(one_one_int,pls),pls),plus_plus_int(plus_plus_int(one_one_int,pls),pls)))),m),one_one_int),t),times_times_int(plus_plus_int(times_times_int(number_number_of_int(plus_plus_int(plus_plus_int(plus_plus_int(plus_plus_int(one_one_int,pls),pls),plus_plus_int(plus_plus_int(one_one_int,pls),pls)),plus_plus_int(plus_plus_int(plus_plus_int(one_one_int,pls),pls),plus_plus_int(plus_plus_int(one_one_int,pls),pls)))),m),one_one_int),zero_zero_int))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_15]), c_0_15]), c_0_15]), c_0_15]), c_0_16]), c_0_16]), c_0_16]), c_0_16]), c_0_16]), c_0_16]), c_0_16]), c_0_16])).
cnf(c_0_33,plain,(times_times_int(t,plus_plus_int(one_one_int,times_times_int(m,number_number_of_int(plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(one_one_int,plus_plus_int(one_one_int,plus_plus_int(one_one_int,one_one_int)))))))))))))))=plus_plus_int(one_one_int,power_power_int(s,number_number_of_nat(plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(one_one_int,one_one_int))))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_26]), c_0_26]), c_0_26]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_28]), c_0_25]), c_0_28]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_27]), c_0_27])).
cnf(c_0_34,plain,(times_times_int(X1,pls)=pls), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30]), c_0_30])).
cnf(c_0_35,negated_conjecture,(~ord_less_int(plus_plus_int(power_power_int(s,number_number_of_nat(plus_plus_int(plus_plus_int(plus_plus_int(one_one_int,pls),pls),plus_plus_int(plus_plus_int(one_one_int,pls),pls)))),one_one_int),pls)), inference(rw,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_36,plain,(ord_less_int(plus_plus_int(one_one_int,power_power_int(s,number_number_of_nat(plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(pls,plus_plus_int(one_one_int,one_one_int)))))))),pls)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_26]), c_0_26]), c_0_26]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_28]), c_0_25]), c_0_28]), c_0_33]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_26]), c_0_26]), c_0_26]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_28]), c_0_25]), c_0_30]), c_0_34])).
cnf(c_0_37,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_27]), c_0_27]), c_0_36])]), ['proof']).
# SZS output end CNFRefutation.
EOF
