% Problem    : SWW313+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n111.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:44:56 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X4]:![X5]:![X6]:(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(tc_Hoare__Mirabelle_Otriple(X6),tc_HOL_Obool)),X5),X4))=>c_Hoare__Mirabelle_Ohoare__derivs(X6,X4,X5))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_asm)).
fof(c_0_1, axiom, (![X4]:![X5]:![X8]:![X6]:(c_Hoare__Mirabelle_Ohoare__derivs(X6,X8,X5)=>(c_Hoare__Mirabelle_Ohoare__derivs(X6,X4,X8)=>c_Hoare__Mirabelle_Ohoare__derivs(X6,X4,X5)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_cut)).
fof(c_0_2, hypothesis, (hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(tc_Hoare__Mirabelle_Otriple(t_a),tc_HOL_Obool)),v_ts),v_G))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_3, hypothesis, (hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(tc_Hoare__Mirabelle_Otriple(t_a),tc_HOL_Obool)),v_G),v_Ga))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_1)).
fof(c_0_4, conjecture, (c_Hoare__Mirabelle_Ohoare__derivs(t_a,v_Ga,v_ts)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_2)).
fof(c_0_5, plain, (![X7]:![X8]:![X9]:(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(tc_Hoare__Mirabelle_Otriple(X9),tc_HOL_Obool)),X8),X7))|c_Hoare__Mirabelle_Ohoare__derivs(X9,X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])).
fof(c_0_6, plain, (![X9]:![X10]:![X11]:![X12]:(~c_Hoare__Mirabelle_Ohoare__derivs(X12,X11,X10)|(~c_Hoare__Mirabelle_Ohoare__derivs(X12,X9,X11)|c_Hoare__Mirabelle_Ohoare__derivs(X12,X9,X10)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
cnf(c_0_7,plain,(c_Hoare__Mirabelle_Ohoare__derivs(X1,X2,X3)|~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(tc_Hoare__Mirabelle_Otriple(X1),tc_HOL_Obool)),X3),X2))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,hypothesis,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(tc_Hoare__Mirabelle_Otriple(t_a),tc_HOL_Obool)),v_ts),v_G))), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_9,plain,(c_Hoare__Mirabelle_Ohoare__derivs(X1,X2,X3)|~c_Hoare__Mirabelle_Ohoare__derivs(X1,X2,X4)|~c_Hoare__Mirabelle_Ohoare__derivs(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,hypothesis,(c_Hoare__Mirabelle_Ohoare__derivs(t_a,v_G,v_ts)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_11,hypothesis,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_fun(tc_Hoare__Mirabelle_Otriple(t_a),tc_HOL_Obool)),v_G),v_Ga))), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_12, negated_conjecture, (~c_Hoare__Mirabelle_Ohoare__derivs(t_a,v_Ga,v_ts)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_4])])).
cnf(c_0_13,hypothesis,(c_Hoare__Mirabelle_Ohoare__derivs(t_a,X1,v_ts)|~c_Hoare__Mirabelle_Ohoare__derivs(t_a,X1,v_G)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_14,hypothesis,(c_Hoare__Mirabelle_Ohoare__derivs(t_a,v_Ga,v_G)), inference(spm,[status(thm)],[c_0_7, c_0_11])).
cnf(c_0_15,negated_conjecture,(~c_Hoare__Mirabelle_Ohoare__derivs(t_a,v_Ga,v_ts)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,hypothesis,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), ['proof']).
# SZS output end CNFRefutation.
EOF
