% Problem    : SWW257+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n089.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:04:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_C45_F1_AE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___200_B07_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# No success with AutoSched1
# Trying AutoSched2 for 23 seconds
# AutoSched2-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X21]:![X22]:![X24]:![X5]:(class_Rings_Ocomm__semiring__1(X5)=>hAPP(hAPP(c_Groups_Otimes__class_Otimes(X5),X24),hAPP(hAPP(c_Groups_Otimes__class_Otimes(X5),X22),X21))=hAPP(hAPP(c_Groups_Otimes__class_Otimes(X5),X22),hAPP(hAPP(c_Groups_Otimes__class_Otimes(X5),X24),X21)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_comm__semiring__1__class_Onormalizing__semiring__rules_I19_J)).
fof(c_0_1, hypothesis, ((?[X70]:c_Groups_Oplus__class_Oplus(tc_Complex_Ocomplex,c_Groups_Oone__class_Oone(tc_Complex_Ocomplex),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),X70),v_k____)),v_a____))=c_Groups_Ozero__class_Ozero(tc_Complex_Ocomplex)=>v_thesis____)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_2, conjecture, (v_thesis____), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_1)).
fof(c_0_3, axiom, (![X12]:![X7]:![X5]:(class_Rings_Ocomm__semiring__1(X5)=>hAPP(hAPP(c_Groups_Otimes__class_Otimes(X5),X7),X12)=hAPP(hAPP(c_Groups_Otimes__class_Otimes(X5),X12),X7))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_comm__semiring__1__class_Onormalizing__semiring__rules_I7_J)).
fof(c_0_4, axiom, (?[X14]:c_Groups_Oplus__class_Oplus(tc_Complex_Ocomplex,c_Groups_Oone__class_Oone(tc_Complex_Ocomplex),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),X14),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),v_a____),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),X14),c_Groups_Ominus__class_Ominus(tc_Nat_Onat,v_k____,c_Groups_Oone__class_Oone(tc_Nat_Onat))))))=c_Groups_Ozero__class_Ozero(tc_Complex_Ocomplex)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact__096EX_Az_O_A1_A_L_Az_A_K_A_Ia_A_K_Az_A_094_A_Ik_A_N_A1_J_J_A_061_A0_096)).
fof(c_0_5, axiom, (class_Rings_Ocomm__semiring__1(tc_Complex_Ocomplex)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Complex__Ocomplex__Rings_Ocomm__semiring__1)).
fof(c_0_6, axiom, (![X6]:![X4]:![X5]:(class_Power_Opower(X5)=>((X4=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=>hAPP(hAPP(c_Power_Opower__class_Opower(X5),X6),X4)=c_Groups_Oone__class_Oone(X5))&(X4!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=>hAPP(hAPP(c_Power_Opower__class_Opower(X5),X6),X4)=hAPP(hAPP(c_Groups_Otimes__class_Otimes(X5),X6),hAPP(hAPP(c_Power_Opower__class_Opower(X5),X6),c_Groups_Ominus__class_Ominus(tc_Nat_Onat,X4,c_Groups_Oone__class_Oone(tc_Nat_Onat)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_realpow__num__eq__if)).
fof(c_0_7, axiom, (class_Power_Opower(tc_Complex_Ocomplex)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Complex__Ocomplex__Power_Opower)).
fof(c_0_8, axiom, (v_k____!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_kas_I2_J)).
fof(c_0_9, plain, (![X25]:![X26]:![X27]:![X28]:(~class_Rings_Ocomm__semiring__1(X28)|hAPP(hAPP(c_Groups_Otimes__class_Otimes(X28),X27),hAPP(hAPP(c_Groups_Otimes__class_Otimes(X28),X26),X25))=hAPP(hAPP(c_Groups_Otimes__class_Otimes(X28),X26),hAPP(hAPP(c_Groups_Otimes__class_Otimes(X28),X27),X25)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])).
fof(c_0_10, hypothesis, (![X71]:(c_Groups_Oplus__class_Oplus(tc_Complex_Ocomplex,c_Groups_Oone__class_Oone(tc_Complex_Ocomplex),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),X71),v_k____)),v_a____))!=c_Groups_Ozero__class_Ozero(tc_Complex_Ocomplex)|v_thesis____)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
fof(c_0_11, negated_conjecture, (~v_thesis____), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_2])])).
fof(c_0_12, plain, (![X13]:![X14]:![X15]:(~class_Rings_Ocomm__semiring__1(X15)|hAPP(hAPP(c_Groups_Otimes__class_Otimes(X15),X14),X13)=hAPP(hAPP(c_Groups_Otimes__class_Otimes(X15),X13),X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
fof(c_0_13, plain, (c_Groups_Oplus__class_Oplus(tc_Complex_Ocomplex,c_Groups_Oone__class_Oone(tc_Complex_Ocomplex),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),esk2_0),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),v_a____),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),esk2_0),c_Groups_Ominus__class_Ominus(tc_Nat_Onat,v_k____,c_Groups_Oone__class_Oone(tc_Nat_Onat))))))=c_Groups_Ozero__class_Ozero(tc_Complex_Ocomplex)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_4])])).
cnf(c_0_14,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(X1),X2),hAPP(hAPP(c_Groups_Otimes__class_Otimes(X1),X3),X4))=hAPP(hAPP(c_Groups_Otimes__class_Otimes(X1),X3),hAPP(hAPP(c_Groups_Otimes__class_Otimes(X1),X2),X4))|~class_Rings_Ocomm__semiring__1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(class_Rings_Ocomm__semiring__1(tc_Complex_Ocomplex)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_16, plain, (![X7]:![X8]:![X9]:(((X8!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|hAPP(hAPP(c_Power_Opower__class_Opower(X9),X7),X8)=c_Groups_Oone__class_Oone(X9))|~class_Power_Opower(X9))&((X8=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|hAPP(hAPP(c_Power_Opower__class_Opower(X9),X7),X8)=hAPP(hAPP(c_Groups_Otimes__class_Otimes(X9),X7),hAPP(hAPP(c_Power_Opower__class_Opower(X9),X7),c_Groups_Ominus__class_Ominus(tc_Nat_Onat,X8,c_Groups_Oone__class_Oone(tc_Nat_Onat)))))|~class_Power_Opower(X9)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_17,hypothesis,(v_thesis____|c_Groups_Oplus__class_Oplus(tc_Complex_Ocomplex,c_Groups_Oone__class_Oone(tc_Complex_Ocomplex),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),X1),v_k____)),v_a____))!=c_Groups_Ozero__class_Ozero(tc_Complex_Ocomplex)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,negated_conjecture,(~v_thesis____), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(X1),X2),X3)=hAPP(hAPP(c_Groups_Otimes__class_Otimes(X1),X3),X2)|~class_Rings_Ocomm__semiring__1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,(c_Groups_Oplus__class_Oplus(tc_Complex_Ocomplex,c_Groups_Oone__class_Oone(tc_Complex_Ocomplex),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),esk2_0),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),v_a____),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),esk2_0),c_Groups_Ominus__class_Ominus(tc_Nat_Onat,v_k____,c_Groups_Oone__class_Oone(tc_Nat_Onat))))))=c_Groups_Ozero__class_Ozero(tc_Complex_Ocomplex)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),X1),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),X2),X3))=hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),X2),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),X1),X3))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22,plain,(hAPP(hAPP(c_Power_Opower__class_Opower(X1),X2),X3)=hAPP(hAPP(c_Groups_Otimes__class_Otimes(X1),X2),hAPP(hAPP(c_Power_Opower__class_Opower(X1),X2),c_Groups_Ominus__class_Ominus(tc_Nat_Onat,X3,c_Groups_Oone__class_Oone(tc_Nat_Onat))))|X3=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)|~class_Power_Opower(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(class_Power_Opower(tc_Complex_Ocomplex)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24,hypothesis,(c_Groups_Oplus__class_Oplus(tc_Complex_Ocomplex,c_Groups_Oone__class_Oone(tc_Complex_Ocomplex),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),X1),v_k____)),v_a____))!=c_Groups_Ozero__class_Ozero(tc_Complex_Ocomplex)), inference(sr,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),X1),X2)=hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),X2),X1)), inference(spm,[status(thm)],[c_0_19, c_0_15])).
cnf(c_0_26,plain,(c_Groups_Oplus__class_Oplus(tc_Complex_Ocomplex,c_Groups_Oone__class_Oone(tc_Complex_Ocomplex),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),v_a____),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),esk2_0),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),esk2_0),c_Groups_Ominus__class_Ominus(tc_Nat_Onat,v_k____,c_Groups_Oone__class_Oone(tc_Nat_Onat))))))=c_Groups_Ozero__class_Ozero(tc_Complex_Ocomplex)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),X1),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),X1),c_Groups_Ominus__class_Ominus(tc_Nat_Onat,X2,c_Groups_Oone__class_Oone(tc_Nat_Onat))))=hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),X1),X2)|c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=X2), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28,hypothesis,(c_Groups_Oplus__class_Oplus(tc_Complex_Ocomplex,c_Groups_Oone__class_Oone(tc_Complex_Ocomplex),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Complex_Ocomplex),v_a____),hAPP(hAPP(c_Power_Opower__class_Opower(tc_Complex_Ocomplex),X1),v_k____)))!=c_Groups_Ozero__class_Ozero(tc_Complex_Ocomplex)), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29,plain,(v_k____!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_30,plain,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29]), ['proof']).
# SZS output end CNFRefutation.
EOF
