% Problem    : SWW373+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n141.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:12:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X17]:![X1]:![X2]:![X82]:![X5]:![X20]:hAPP(hAPP(c_Fun_Ocomp(X20,X5,X82,X2),X1),X17)=hAPP(X2,hAPP(X1,X17))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_o__apply)).
fof(c_0_1, axiom, (![X19]:![X5]:(hAPP(c_Set_OCollect(X5),X19)=c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))<=>![X3]:~(hBOOL(hAPP(X19,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__empty__eq)).
fof(c_0_2, axiom, (![X5]:c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),hAPP(c_COMBK(tc_HOL_Obool,X5),c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__def)).
fof(c_0_3, axiom, (![X19]:![X5]:hAPP(c_Set_OCollect(X5),X19)=X19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__def)).
fof(c_0_4, axiom, (c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_transfer__nat__int__numerals_I1_J)).
fof(c_0_5, axiom, (c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Pls__def)).
fof(c_0_6, axiom, (c_Code__Numeral_Oint__of=hAPP(c_Fun_Ocomp(tc_Nat_Onat,tc_Int_Oint,tc_Code__Numeral_Ocode__numeral,c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint)),c_Code__Numeral_Onat__of)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_int__of__def)).
fof(c_0_7, axiom, (![X108]:hAPP(c_Code__Numeral_Onat__of,hAPP(c_Code__Numeral_Oof__nat,X108))=X108), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_nat__of__of__nat)).
fof(c_0_8, axiom, (![X5]:(class_Rings_Olinordered__semidom(X5)=>~(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(X5),c_Groups_Oone__class_Oone(X5)),c_Groups_Ozero__class_Ozero(X5)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_not__one__le__zero)).
fof(c_0_9, axiom, (![X19]:![X5]:hAPP(c_Set_OCollect(X5),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X5),c_fNot),X19))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),hAPP(c_Set_OCollect(X5),X19))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__neg__eq)).
fof(c_0_10, axiom, (![X328]:![X27]:![X20]:![X5]:hAPP(hAPP(c_COMBK(X5,X20),X27),X328)=X27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBK__1)).
fof(c_0_11, axiom, (![X5]:c_Orderings_Otop__class_Otop(tc_fun(X5,tc_HOL_Obool))=hAPP(c_Set_OCollect(X5),hAPP(c_COMBK(tc_HOL_Obool,X5),c_fTrue))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__def)).
fof(c_0_12, axiom, (![X6]:![X5]:(class_Groups_Ogroup__add(X5)=>(hAPP(c_Groups_Ouminus__class_Ouminus(X5),X6)=c_Groups_Ozero__class_Ozero(X5)<=>X6=c_Groups_Ozero__class_Ozero(X5)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_neg__equal__0__iff__equal)).
fof(c_0_13, axiom, (![X110]:hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),X110)=hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),X110)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_mult__Min)).
fof(c_0_14, axiom, (![X87]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),X87))=X87), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_zminus__zminus)).
fof(c_0_15, axiom, (![X108]:hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_Int_Oint),hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),X108))),c_Groups_Ozero__class_Ozero(tc_Int_Oint)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_negative__zle__0)).
fof(c_0_16, axiom, (![X90]:![X42]:(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X42),X90)=c_Groups_Oone__class_Oone(tc_Int_Oint)<=>((X42=c_Groups_Oone__class_Oone(tc_Int_Oint)&X90=c_Groups_Oone__class_Oone(tc_Int_Oint))|(X42=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin)&X90=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_zmult__eq__1__iff)).
fof(c_0_17, axiom, (![X110]:hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),X110)=X110), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_number__of__is__id)).
fof(c_0_18, axiom, (c_Groups_Oone__class_Oone(tc_Int_Oint)=hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),c_Groups_Oone__class_Oone(tc_Nat_Onat))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_transfer__int__nat__numerals_I2_J)).
fof(c_0_19, axiom, (c_Groups_Oone__class_Oone(tc_Nat_Onat)=hAPP(c_Nat_OSuc,c_Groups_Ozero__class_Ozero(tc_Nat_Onat))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_One__nat__def)).
fof(c_0_20, axiom, (![X10]:![X5]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),X10))=X10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__complement)).
fof(c_0_21, axiom, (![X19]:(~(hBOOL(hAPP(c_fNot,X19)))|~(hBOOL(X19)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__1)).
fof(c_0_22, axiom, (![X5]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X5,tc_HOL_Obool)),c_Orderings_Otop__class_Otop(tc_fun(X5,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__UNIV__eq)).
fof(c_0_23, axiom, (![X108]:hBOOL(hAPP(c_Nat__Numeral_Oneg,hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),hAPP(c_Nat_OSuc,X108)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_neg__zminus__int)).
fof(c_0_24, axiom, (class_Groups_Ogroup__add(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Int__Oint__Groups_Ogroup__add)).
fof(c_0_25, axiom, (~(hBOOL(hAPP(c_Nat__Numeral_Oneg,c_Groups_Ozero__class_Ozero(tc_Int_Oint))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_not__neg__0)).
fof(c_0_26, axiom, (class_Rings_Olinordered__semidom(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Int__Oint__Rings_Olinordered__semidom)).
fof(c_0_27, axiom, (![X49]:![X18]:![X19]:![X82]:![X5]:![X20]:hAPP(hAPP(hAPP(c_COMBB(X20,X5,X82),X19),X18),X49)=hAPP(X19,hAPP(X18,X49))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBB__1)).
fof(c_0_28, axiom, (![X109]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X109),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X109), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Nat_Oadd__0__right)).
fof(c_0_29, axiom, (![X186]:![X5]:![X20]:hAPP(c_Relation_Oconverse(X20,X5),hAPP(c_Relation_Oconverse(X5,X20),X186))=X186), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_converse__converse)).
fof(c_0_30, axiom, (![X204]:![X1]:![X2]:![X20]:![X5]:c_Product__Type_Oprod_Oprod__size(X5,X20,X2,X1,X204)=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(X2,hAPP(c_Product__Type_Ofst(X5,X20),X204))),hAPP(X1,hAPP(c_Product__Type_Osnd(X5,X20),X204)))),hAPP(c_Nat_OSuc,c_Groups_Ozero__class_Ozero(tc_Nat_Onat)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_prod__size__simp)).
fof(c_0_31, axiom, (![X110]:![X108]:![X109]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X109),X108)),X110)=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X109),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X108),X110))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_nat__add__assoc)).
fof(c_0_32, axiom, (![X108]:![X109]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X109),hAPP(c_Nat_OSuc,X108))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X109),X108))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_add__Suc__right)).
fof(c_0_33, axiom, (![X110]:hAPP(c_Code__Numeral_Oof__nat,hAPP(c_Code__Numeral_Onat__of,X110))=X110), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_of__nat__nat__of)).
fof(c_0_34, axiom, (c_Groups_Ozero__class_Ozero(tc_Int_Oint)=hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_transfer__int__nat__numerals_I1_J)).
fof(c_0_35, axiom, (![X110]:c_Int_OBit1(X110)=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),c_Groups_Oone__class_Oone(tc_Int_Oint)),X110)),X110)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Bit1__def)).
fof(c_0_36, axiom, (![X28]:c_Divides_Odiv__class_Omod(tc_Int_Oint,c_Groups_Ozero__class_Ozero(tc_Int_Oint),X28)=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_zmod__zero)).
fof(c_0_37, axiom, (![X6]:![X5]:(class_Groups_Olinordered__ab__group__add(X5)=>(hAPP(c_Groups_Ouminus__class_Ouminus(X5),X6)=X6<=>X6=c_Groups_Ozero__class_Ozero(X5)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_neg__equal__zero)).
fof(c_0_38, axiom, (hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),c_Int_OMin)=c_Int_OBit1(c_Int_OPls)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_minus__Min)).
fof(c_0_39, axiom, (![X110]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),X110),c_Int_OPls)=X110), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_add__Pls__right)).
fof(c_0_40, axiom, (class_Groups_Olinordered__ab__group__add(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Int__Oint__Groups_Olinordered__ab__group__add)).
fof(c_0_41, axiom, (c_Int_OMin!=c_Int_OPls), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rel__simps_I40_J)).
fof(c_0_42, plain, (![X83]:![X84]:![X85]:![X86]:![X87]:![X88]:hAPP(hAPP(c_Fun_Ocomp(X88,X87,X86,X85),X84),X83)=hAPP(X85,hAPP(X84,X83))), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_43, plain, (![X19]:![X5]:(hAPP(c_Set_OCollect(X5),X19)=c_Orderings_Obot__class_Obot(tc_fun(X5,tc_HOL_Obool))<=>![X3]:~hBOOL(hAPP(X19,X3)))), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_44, plain, (![X6]:c_Orderings_Obot__class_Obot(tc_fun(X6,tc_HOL_Obool))=hAPP(c_Set_OCollect(X6),hAPP(c_COMBK(tc_HOL_Obool,X6),c_fFalse))), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_45, plain, (![X20]:![X21]:hAPP(c_Set_OCollect(X21),X20)=X20), inference(variable_rename,[status(thm)],[c_0_3])).
cnf(c_0_46,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_47,plain,(c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_48,plain,(hAPP(hAPP(c_Fun_Ocomp(X1,X2,X3,X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49,plain,(c_Code__Numeral_Oint__of=hAPP(c_Fun_Ocomp(tc_Nat_Onat,tc_Int_Oint,tc_Code__Numeral_Ocode__numeral,c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint)),c_Code__Numeral_Onat__of)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_50, plain, (![X109]:hAPP(c_Code__Numeral_Onat__of,hAPP(c_Code__Numeral_Oof__nat,X109))=X109), inference(variable_rename,[status(thm)],[c_0_7])).
fof(c_0_51, plain, (![X5]:(class_Rings_Olinordered__semidom(X5)=>~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(X5),c_Groups_Oone__class_Oone(X5)),c_Groups_Ozero__class_Ozero(X5))))), inference(fof_simplification,[status(thm)],[c_0_8])).
fof(c_0_52, plain, (![X20]:![X21]:hAPP(c_Set_OCollect(X21),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X21),c_fNot),X20))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X21,tc_HOL_Obool)),hAPP(c_Set_OCollect(X21),X20))), inference(variable_rename,[status(thm)],[c_0_9])).
fof(c_0_53, plain, (![X20]:![X21]:![X22]:![X23]:![X25]:((hAPP(c_Set_OCollect(X21),X20)!=c_Orderings_Obot__class_Obot(tc_fun(X21,tc_HOL_Obool))|~hBOOL(hAPP(X20,X22)))&(hBOOL(hAPP(X23,esk304_1(X23)))|hAPP(c_Set_OCollect(X25),X23)=c_Orderings_Obot__class_Obot(tc_fun(X25,tc_HOL_Obool))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_43])])])])])).
fof(c_0_54, plain, (![X329]:![X330]:![X331]:![X332]:hAPP(hAPP(c_COMBK(X332,X331),X330),X329)=X330), inference(variable_rename,[status(thm)],[c_0_10])).
cnf(c_0_55,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56,plain,(hAPP(c_Set_OCollect(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_57, plain, (![X6]:c_Orderings_Otop__class_Otop(tc_fun(X6,tc_HOL_Obool))=hAPP(c_Set_OCollect(X6),hAPP(c_COMBK(tc_HOL_Obool,X6),c_fTrue))), inference(variable_rename,[status(thm)],[c_0_11])).
fof(c_0_58, plain, (![X7]:![X8]:(((hAPP(c_Groups_Ouminus__class_Ouminus(X8),X7)!=c_Groups_Ozero__class_Ozero(X8)|X7=c_Groups_Ozero__class_Ozero(X8))|~class_Groups_Ogroup__add(X8))&((X7!=c_Groups_Ozero__class_Ozero(X8)|hAPP(c_Groups_Ouminus__class_Ouminus(X8),X7)=c_Groups_Ozero__class_Ozero(X8))|~class_Groups_Ogroup__add(X8)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
fof(c_0_59, plain, (![X111]:hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),X111)=hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),X111)), inference(variable_rename,[status(thm)],[c_0_13])).
fof(c_0_60, plain, (![X88]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),X88))=X88), inference(variable_rename,[status(thm)],[c_0_14])).
fof(c_0_61, plain, (![X109]:hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_Int_Oint),hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),X109))),c_Groups_Ozero__class_Ozero(tc_Int_Oint)))), inference(variable_rename,[status(thm)],[c_0_15])).
fof(c_0_62, plain, (![X91]:![X92]:![X93]:![X94]:(((((X92=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin)|X92=c_Groups_Oone__class_Oone(tc_Int_Oint))|hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X92),X91)!=c_Groups_Oone__class_Oone(tc_Int_Oint))&((X91=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin)|X92=c_Groups_Oone__class_Oone(tc_Int_Oint))|hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X92),X91)!=c_Groups_Oone__class_Oone(tc_Int_Oint)))&(((X92=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin)|X91=c_Groups_Oone__class_Oone(tc_Int_Oint))|hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X92),X91)!=c_Groups_Oone__class_Oone(tc_Int_Oint))&((X91=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin)|X91=c_Groups_Oone__class_Oone(tc_Int_Oint))|hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X92),X91)!=c_Groups_Oone__class_Oone(tc_Int_Oint))))&(((X94!=c_Groups_Oone__class_Oone(tc_Int_Oint)|X93!=c_Groups_Oone__class_Oone(tc_Int_Oint))|hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X94),X93)=c_Groups_Oone__class_Oone(tc_Int_Oint))&((X94!=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin)|X93!=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin))|hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X94),X93)=c_Groups_Oone__class_Oone(tc_Int_Oint))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])])).
fof(c_0_63, plain, (![X111]:hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),X111)=X111), inference(variable_rename,[status(thm)],[c_0_17])).
cnf(c_0_64,plain,(c_Groups_Oone__class_Oone(tc_Int_Oint)=hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),c_Groups_Oone__class_Oone(tc_Nat_Onat))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_65,plain,(c_Groups_Oone__class_Oone(tc_Nat_Onat)=hAPP(c_Nat_OSuc,c_Groups_Ozero__class_Ozero(tc_Nat_Onat))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_66,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_67,plain,(hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),hAPP(c_Code__Numeral_Onat__of,X1))=hAPP(c_Code__Numeral_Oint__of,X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_68,plain,(hAPP(c_Code__Numeral_Onat__of,hAPP(c_Code__Numeral_Oof__nat,X1))=X1), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_69, plain, (![X6]:(~class_Rings_Olinordered__semidom(X6)|~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(X6),c_Groups_Oone__class_Oone(X6)),c_Groups_Ozero__class_Ozero(X6))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_51])])).
fof(c_0_70, plain, (![X11]:![X12]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X12,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X12,tc_HOL_Obool)),X11))=X11), inference(variable_rename,[status(thm)],[c_0_20])).
cnf(c_0_71,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Set_OCollect(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_72, plain, (![X19]:(~hBOOL(hAPP(c_fNot,X19))|~hBOOL(X19))), inference(fof_simplification,[status(thm)],[c_0_21])).
cnf(c_0_73,plain,(hAPP(c_Set_OCollect(X1),X2)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))|hBOOL(hAPP(X2,esk304_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_74,plain,(hAPP(hAPP(c_COMBK(X1,X2),X3),X4)=X3), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_75,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_76, plain, (![X6]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X6,tc_HOL_Obool)),c_Orderings_Otop__class_Otop(tc_fun(X6,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X6,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_22])).
cnf(c_0_77,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue))), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_78, plain, (![X109]:hBOOL(hAPP(c_Nat__Numeral_Oneg,hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),hAPP(c_Nat_OSuc,X109)))))), inference(variable_rename,[status(thm)],[c_0_23])).
cnf(c_0_79,plain,(X2=c_Groups_Ozero__class_Ozero(X1)|~class_Groups_Ogroup__add(X1)|hAPP(c_Groups_Ouminus__class_Ouminus(X1),X2)!=c_Groups_Ozero__class_Ozero(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_80,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),X1)=hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_81,plain,(class_Groups_Ogroup__add(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_82,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),X1))=X1), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_83, plain, (~hBOOL(hAPP(c_Nat__Numeral_Oneg,c_Groups_Ozero__class_Ozero(tc_Int_Oint)))), inference(fof_simplification,[status(thm)],[c_0_25])).
cnf(c_0_84,plain,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_Int_Oint),hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),X1))),c_Groups_Ozero__class_Ozero(tc_Int_Oint)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_85,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X1),X2)=c_Groups_Oone__class_Oone(tc_Int_Oint)|X2!=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin)|X1!=hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),c_Int_OMin)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_86,plain,(hAPP(c_Int_Onumber__class_Onumber__of(tc_Int_Oint),X1)=X1), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_87,plain,(hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),hAPP(c_Nat_OSuc,hAPP(c_Int_Onat,c_Int_OPls)))=c_Groups_Oone__class_Oone(tc_Int_Oint)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_65]), c_0_66])).
cnf(c_0_88,plain,(hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),X1)=hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,X1))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_89,plain,(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(X1),c_Groups_Oone__class_Oone(X1)),c_Groups_Ozero__class_Ozero(X1)))|~class_Rings_Olinordered__semidom(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_90,plain,(class_Rings_Olinordered__semidom(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_91, plain, (![X83]:![X84]:![X85]:![X86]:![X87]:![X88]:hAPP(hAPP(hAPP(c_COMBB(X88,X87,X86),X85),X84),X83)=hAPP(X85,hAPP(X84,X83))), inference(variable_rename,[status(thm)],[c_0_27])).
cnf(c_0_92,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_93,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2)=hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_56]), c_0_56])).
fof(c_0_94, plain, (![X20]:(~hBOOL(hAPP(c_fNot,X20))|~hBOOL(X20))), inference(variable_rename,[status(thm)],[c_0_72])).
cnf(c_0_95,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hBOOL(hAPP(X1,esk304_1(X1)))), inference(rw,[status(thm)],[c_0_73, c_0_56])).
cnf(c_0_96,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)),X2)=c_fFalse), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_97,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_98,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue)=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_77, c_0_56])).
cnf(c_0_99,plain,(hBOOL(hAPP(c_Nat__Numeral_Oneg,hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),hAPP(c_Nat_OSuc,X1)))))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_100,plain,(c_Int_OPls=X1|hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),X1)!=c_Int_OPls), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_47]), c_0_47]), c_0_81])])).
cnf(c_0_101,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_80]), c_0_80])).
cnf(c_0_102,plain,(~hBOOL(hAPP(c_Nat__Numeral_Oneg,c_Groups_Ozero__class_Ozero(tc_Int_Oint)))), inference(split_conjunct,[status(thm)],[c_0_83])).
fof(c_0_103, plain, (![X110]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X110),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X110), inference(variable_rename,[status(thm)],[c_0_28])).
cnf(c_0_104,plain,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_Int_Oint),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),X1))),c_Int_OPls))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_80]), c_0_47])).
cnf(c_0_105,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X1),X2)=c_Groups_Oone__class_Oone(tc_Int_Oint)|c_Int_OMin!=X2|c_Int_OMin!=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86]), c_0_86])).
cnf(c_0_106,plain,(c_Groups_Oone__class_Oone(tc_Int_Oint)=hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,hAPP(c_Nat_OSuc,hAPP(c_Int_Onat,c_Int_OPls))))), inference(rw,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_107,plain,(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_Int_Oint),c_Groups_Oone__class_Oone(tc_Int_Oint)),c_Int_OPls))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_47]), c_0_90])])).
cnf(c_0_108,plain,(hAPP(hAPP(hAPP(c_COMBB(X1,X2,X3),X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_109,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93]), c_0_93])).
fof(c_0_110, plain, (![X187]:![X188]:![X189]:hAPP(c_Relation_Oconverse(X189,X188),hAPP(c_Relation_Oconverse(X188,X189),X187))=X187), inference(variable_rename,[status(thm)],[c_0_29])).
cnf(c_0_111,plain,(~hBOOL(X1)|~hBOOL(hAPP(c_fNot,X1))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_112,plain,(c_fFalse=X1|hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_95]), c_0_96]), c_0_74])).
cnf(c_0_113,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_97, c_0_93])).
cnf(c_0_114,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)),X2)=c_fTrue), inference(spm,[status(thm)],[c_0_74, c_0_98])).
cnf(c_0_115,plain,(hBOOL(hAPP(c_Nat__Numeral_Oneg,hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),hAPP(c_Nat_OSuc,X1)))))), inference(rw,[status(thm)],[c_0_99, c_0_80])).
cnf(c_0_116,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),X1)=c_Int_OPls|X1!=c_Int_OPls), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_117,plain,(~hBOOL(hAPP(c_Nat__Numeral_Oneg,c_Int_OPls))), inference(rw,[status(thm)],[c_0_102, c_0_47])).
fof(c_0_118, plain, (![X205]:![X206]:![X207]:![X208]:![X209]:c_Product__Type_Oprod_Oprod__size(X209,X208,X207,X206,X205)=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(X207,hAPP(c_Product__Type_Ofst(X209,X208),X205))),hAPP(X206,hAPP(c_Product__Type_Osnd(X209,X208),X205)))),hAPP(c_Nat_OSuc,c_Groups_Ozero__class_Ozero(tc_Nat_Onat)))), inference(variable_rename,[status(thm)],[c_0_30])).
fof(c_0_119, plain, (![X111]:![X112]:![X113]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X113),X112)),X111)=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X113),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X112),X111))), inference(variable_rename,[status(thm)],[c_0_31])).
fof(c_0_120, plain, (![X110]:![X111]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X111),hAPP(c_Nat_OSuc,X110))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X111),X110))), inference(variable_rename,[status(thm)],[c_0_32])).
cnf(c_0_121,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X1), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_122,plain,(hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_Int_Oint),hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,X1)))),c_Int_OPls))), inference(rw,[status(thm)],[c_0_104, c_0_88])).
cnf(c_0_123,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),X1),X2)=hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,hAPP(c_Nat_OSuc,hAPP(c_Int_Onat,c_Int_OPls))))|c_Int_OMin!=X2|c_Int_OMin!=X1), inference(rw,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_124,plain,(~hBOOL(hAPP(hAPP(c_Orderings_Oord__class_Oless__eq(tc_Int_Oint),hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,hAPP(c_Nat_OSuc,hAPP(c_Int_Onat,c_Int_OPls))))),c_Int_OPls))), inference(rw,[status(thm)],[c_0_107, c_0_106])).
fof(c_0_125, plain, (![X111]:hAPP(c_Code__Numeral_Oof__nat,hAPP(c_Code__Numeral_Onat__of,X111))=X111), inference(variable_rename,[status(thm)],[c_0_33])).
cnf(c_0_126,plain,(hAPP(c_fNot,hAPP(c_fNot,hAPP(X1,X2)))=hAPP(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_108])).
cnf(c_0_127,plain,(hAPP(c_Relation_Oconverse(X1,X2),hAPP(c_Relation_Oconverse(X2,X1),X3))=X3), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_128,plain,(hAPP(c_fNot,X1)=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_129,plain,(hAPP(c_fNot,c_fTrue)=c_fFalse), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_113]), c_0_96]), c_0_114])).
cnf(c_0_130,plain,(hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),hAPP(c_Nat_OSuc,X1))!=c_Int_OPls), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_117])).
cnf(c_0_131,plain,(c_Product__Type_Oprod_Oprod__size(X1,X2,X3,X4,X5)=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(X3,hAPP(c_Product__Type_Ofst(X1,X2),X5))),hAPP(X4,hAPP(c_Product__Type_Osnd(X1,X2),X5)))),hAPP(c_Nat_OSuc,c_Groups_Ozero__class_Ozero(tc_Nat_Onat)))), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_132,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2)),X3)=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X2),X3))), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_133,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),hAPP(c_Nat_OSuc,X2))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_134,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),hAPP(c_Int_Onat,c_Int_OPls))=X1), inference(rw,[status(thm)],[c_0_121, c_0_66])).
cnf(c_0_135,plain,(hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,X1))!=c_Int_OMin), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_124])).
cnf(c_0_136,plain,(hAPP(c_Code__Numeral_Oof__nat,hAPP(c_Code__Numeral_Onat__of,X1))=X1), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_137,plain,(hAPP(c_fNot,hAPP(c_fNot,X1))=X1), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_138,plain,(hAPP(c_fNot,X1)=c_fFalse|c_fFalse=X1), inference(spm,[status(thm)],[c_0_128, c_0_112])).
cnf(c_0_139,plain,(hAPP(c_fNot,c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_114]), c_0_129])).
cnf(c_0_140,plain,(c_Groups_Ozero__class_Ozero(tc_Int_Oint)=hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_141,plain,(hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,hAPP(c_Nat_OSuc,X1)))!=c_Int_OPls), inference(rw,[status(thm)],[c_0_130, c_0_88])).
cnf(c_0_142,plain,(hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),hAPP(X1,hAPP(c_Product__Type_Ofst(X2,X3),X4))),hAPP(X5,hAPP(c_Product__Type_Osnd(X2,X3),X4))))=c_Product__Type_Oprod_Oprod__size(X2,X3,X1,X5,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_66]), c_0_132]), c_0_133]), c_0_134]), c_0_133])).
cnf(c_0_143,plain,(hAPP(c_Code__Numeral_Oint__of,X1)!=c_Int_OMin), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_144,plain,(c_fFalse=X1|c_fTrue=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_139])).
cnf(c_0_145,plain,(hAPP(c_Nat_Osemiring__1__class_Oof__nat(tc_Int_Oint),hAPP(c_Int_Onat,c_Int_OPls))=c_Int_OPls), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_140, c_0_66]), c_0_47])).
cnf(c_0_146,plain,(hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,c_Product__Type_Oprod_Oprod__size(X1,X2,X3,X4,X5)))!=c_Int_OPls), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_147,plain,(hAPP(c_Code__Numeral_Oint__of,X1)=c_fFalse|c_fTrue!=c_Int_OMin), inference(spm,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_148,plain,(hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,hAPP(c_Int_Onat,c_Int_OPls)))=c_Int_OPls), inference(rw,[status(thm)],[c_0_145, c_0_88])).
cnf(c_0_149,plain,(c_fFalse!=c_Int_OPls|c_fTrue!=c_Int_OMin), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_150,plain,(c_fTrue!=c_Int_OMin), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_147]), c_0_149])).
fof(c_0_151, plain, (![X111]:c_Int_OBit1(X111)=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),c_Groups_Oone__class_Oone(tc_Int_Oint)),X111)),X111)), inference(variable_rename,[status(thm)],[c_0_35])).
fof(c_0_152, plain, (![X29]:c_Divides_Odiv__class_Omod(tc_Int_Oint,c_Groups_Ozero__class_Ozero(tc_Int_Oint),X29)=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), inference(variable_rename,[status(thm)],[c_0_36])).
cnf(c_0_153,plain,(c_fFalse=X1|X1!=c_Int_OMin), inference(spm,[status(thm)],[c_0_150, c_0_144])).
fof(c_0_154, plain, (![X7]:![X8]:(((hAPP(c_Groups_Ouminus__class_Ouminus(X8),X7)!=X7|X7=c_Groups_Ozero__class_Ozero(X8))|~class_Groups_Olinordered__ab__group__add(X8))&((X7!=c_Groups_Ozero__class_Ozero(X8)|hAPP(c_Groups_Ouminus__class_Ouminus(X8),X7)=X7)|~class_Groups_Olinordered__ab__group__add(X8)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).
cnf(c_0_155,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),c_Int_OMin)=c_Int_OBit1(c_Int_OPls)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_156,plain,(c_Int_OBit1(X1)=hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),c_Groups_Oone__class_Oone(tc_Int_Oint)),X1)),X1)), inference(split_conjunct,[status(thm)],[c_0_151])).
fof(c_0_157, plain, (![X111]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),X111),c_Int_OPls)=X111), inference(variable_rename,[status(thm)],[c_0_39])).
cnf(c_0_158,plain,(c_Divides_Odiv__class_Omod(tc_Int_Oint,c_Groups_Ozero__class_Ozero(tc_Int_Oint),X1)=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_152])).
cnf(c_0_159,plain,(c_fFalse=c_Int_OMin), inference(er,[status(thm)],[c_0_153])).
cnf(c_0_160,plain,(X2=c_Groups_Ozero__class_Ozero(X1)|~class_Groups_Olinordered__ab__group__add(X1)|hAPP(c_Groups_Ouminus__class_Ouminus(X1),X2)!=X2), inference(split_conjunct,[status(thm)],[c_0_154])).
cnf(c_0_161,plain,(class_Groups_Olinordered__ab__group__add(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_162,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),c_Groups_Oone__class_Oone(tc_Int_Oint)),c_Int_OPls)),c_Int_OPls)=hAPP(c_Groups_Ouminus__class_Ouminus(tc_Int_Oint),c_Int_OMin)), inference(rw,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_163,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Int_Oint),X1),c_Int_OPls)=X1), inference(split_conjunct,[status(thm)],[c_0_157])).
cnf(c_0_164,plain,(c_Divides_Odiv__class_Omod(tc_Int_Oint,c_Int_OPls,X1)=c_Int_OPls), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_158, c_0_47]), c_0_47])).
cnf(c_0_165,plain,(c_fTrue=X1|c_Int_OMin=X1), inference(rw,[status(thm)],[c_0_144, c_0_159])).
cnf(c_0_166,plain,(c_Int_OMin!=c_Int_OPls), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_167,plain,(c_Int_OPls=X1|hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),X1)!=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_80]), c_0_47]), c_0_161])])).
cnf(c_0_168,plain,(hAPP(hAPP(c_Groups_Otimes__class_Otimes(tc_Int_Oint),c_Int_OMin),c_Int_OMin)=c_Groups_Oone__class_Oone(tc_Int_Oint)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_162, c_0_163]), c_0_163]), c_0_80])).
cnf(c_0_169,plain,(c_fTrue=c_Int_OPls), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_164]), c_0_166])).
cnf(c_0_170,plain,(c_Groups_Oone__class_Oone(tc_Int_Oint)!=c_Int_OMin), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_166])).
cnf(c_0_171,plain,(c_Int_OMin=X1|c_Int_OPls=X1), inference(rw,[status(thm)],[c_0_165, c_0_169])).
cnf(c_0_172,plain,(hAPP(c_Code__Numeral_Oint__of,hAPP(c_Code__Numeral_Oof__nat,hAPP(c_Nat_OSuc,hAPP(c_Int_Onat,c_Int_OPls))))!=c_Int_OMin), inference(rw,[status(thm)],[c_0_170, c_0_106])).
cnf(c_0_173,plain,($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_171]), c_0_106]), c_0_172]), c_0_141]), ['proof']).
# SZS output end CNFRefutation.
EOF
