% Problem    : GEO342+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n118.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:54:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X1]:((ron(vd1399,X1)&ron(vd1401,X1))&rline(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qe(thee(the(408), 1))')).
fof(c_0_1, axiom, (![X123]:![X124]:(((X123!=X124&?[X125]:(X124=X125&rpoint(X125)))&?[X126]:(X123=X126&rpoint(X126)))=>?[X127]:((ron(X124,X127)&ron(X123,X127))&rline(X127)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(51), 0), imp(cond(axiom(51), 0)))')).
fof(c_0_2, axiom, ((((((((((vd1406=vd1407&ron(vd1400,vd1406))&ron(vd1399,vd1406))&rline(vd1407))&vd1400!=vd1401)&vd1399!=vd1401)&vd1399!=vd1400)&?[X2]:(vd1401=X2&rpoint(X2)))&?[X3]:(vd1400=X3&rpoint(X3)))&?[X4]:(vd1399=X4&rpoint(X4)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(pred(conjunct2(405), 5), and(pred(s2(plural(conjunct2(405))), 0), and(pred(s1(plural(conjunct2(405))), 0), and(pred(conjunct2(405), 1), and(pred(conjunct1(405), 9), and(pred(conjunct1(405), 8), and(pred(conjunct1(405), 7), and(qe(s3(plural(405))), and(qe(s2(plural(405))), qe(s1(plural(405))))))))))))')).
fof(c_0_3, negated_conjecture, (~(?[X1]:((ron(vd1399,X1)&ron(vd1401,X1))&rline(X1)))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_4, plain, (![X128]:![X129]:![X130]:![X131]:(((ron(X129,esk8_2(X128,X129))|((X128=X129|(X129!=X130|~rpoint(X130)))|(X128!=X131|~rpoint(X131))))&(ron(X128,esk8_2(X128,X129))|((X128=X129|(X129!=X130|~rpoint(X130)))|(X128!=X131|~rpoint(X131)))))&(rline(esk8_2(X128,X129))|((X128=X129|(X129!=X130|~rpoint(X130)))|(X128!=X131|~rpoint(X131)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
fof(c_0_5, negated_conjecture, (![X2]:((~ron(vd1399,X2)|~ron(vd1401,X2))|~rline(X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
cnf(c_0_6,plain,(X2=X4|rline(esk8_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, plain, ((((((((((vd1406=vd1407&ron(vd1400,vd1406))&ron(vd1399,vd1406))&rline(vd1407))&vd1400!=vd1401)&vd1399!=vd1401)&vd1399!=vd1400)&(vd1401=esk1_0&rpoint(esk1_0)))&(vd1400=esk2_0&rpoint(esk2_0)))&(vd1399=esk3_0&rpoint(esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_2])])).
cnf(c_0_8,negated_conjecture,(~rline(X1)|~ron(vd1401,X1)|~ron(vd1399,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(X1=X2|rline(esk8_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_6])])).
cnf(c_0_10,plain,(X2=X4|ron(X4,esk8_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11,plain,(rpoint(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(vd1399=esk3_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,negated_conjecture,(X1=X2|~rpoint(X2)|~rpoint(X1)|~ron(vd1399,esk8_2(X1,X2))|~ron(vd1401,esk8_2(X1,X2))), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14,plain,(X1=X2|ron(X2,esk8_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_10])])).
cnf(c_0_15,plain,(rpoint(vd1399)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16,plain,(X2=X4|ron(X2,esk8_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,plain,(rpoint(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_18,plain,(vd1401=esk1_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19,negated_conjecture,(X1=vd1399|~rpoint(X1)|~ron(vd1401,esk8_2(X1,vd1399))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
cnf(c_0_20,plain,(X1=X2|ron(X1,esk8_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_16])])).
cnf(c_0_21,plain,(rpoint(vd1401)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22,plain,(vd1399!=vd1401), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_15])]), c_0_22]), ['proof']).
# SZS output end CNFRefutation.
EOF
