% Problem    : SWW474+3 : TPTP v6.1.0. Released v5.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n112.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:41:11 CDT 2014
% CPUTime    : 
# SinE strategy is gf120_h_gu_R02_F100_L20000
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X29]:![X30]:(hBOOL(wT_bodies)=>(hAPP_p799580910on_com(body,X29)=hAPP_com_option_com(some_com,X30)=>hBOOL(hAPP_com_bool(wt,X30))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_297_WT__bodiesD)).
fof(c_0_1, axiom, (![X15]:(hBOOL(hoare_1795711768gleton)=>(hBOOL(wT_bodies)=>(hBOOL(hAPP_com_bool(wt,X15))=>hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(bot_bo784226126e_bool),hAPP_f806699093e_bool(hAPP_H1902130436e_bool(insert1744391420_state,hAPP_c1455475371_state(hoare_Mirabelle_MGT,X15)),bot_bo784226126e_bool))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_66_MGF)).
fof(c_0_2, hypothesis, (hBOOL(wT_bodies)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_1)).
fof(c_0_3, axiom, (![X3]:![X6]:![X4]:(hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X6),X4))=>(hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X3),X6))=>hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X3),X4))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_4_cut)).
fof(c_0_4, hypothesis, (hBOOL(hoare_1795711768gleton)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_5, axiom, (![X3]:hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X3),bot_bo784226126e_bool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_0_empty)).
fof(c_0_6, hypothesis, (hAPP_p799580910on_com(body,pn)=hAPP_com_option_com(some_com,y)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_5)).
fof(c_0_7, conjecture, (hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(hAPP_f360545851e_bool(image_185131637_state(hAPP_f1377420673_state(cOMBB_271860050_pname(hoare_Mirabelle_MGT),body_1)),dom_pname_com(body))),hAPP_f806699093e_bool(hAPP_H1902130436e_bool(insert1744391420_state,hAPP_c1455475371_state(hoare_Mirabelle_MGT,y)),bot_bo784226126e_bool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_7)).
fof(c_0_8, plain, (![X31]:![X32]:(~hBOOL(wT_bodies)|(hAPP_p799580910on_com(body,X31)!=hAPP_com_option_com(some_com,X32)|hBOOL(hAPP_com_bool(wt,X32))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])).
fof(c_0_9, plain, (![X16]:(~hBOOL(hoare_1795711768gleton)|(~hBOOL(wT_bodies)|(~hBOOL(hAPP_com_bool(wt,X16))|hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(bot_bo784226126e_bool),hAPP_f806699093e_bool(hAPP_H1902130436e_bool(insert1744391420_state,hAPP_c1455475371_state(hoare_Mirabelle_MGT,X16)),bot_bo784226126e_bool))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
cnf(c_0_10,plain,(hBOOL(hAPP_com_bool(wt,X1))|hAPP_p799580910on_com(body,X2)!=hAPP_com_option_com(some_com,X1)|~hBOOL(wT_bodies)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11,hypothesis,(hBOOL(wT_bodies)), inference(split_conjunct,[status(thm)],[c_0_2])).
fof(c_0_12, plain, (![X7]:![X8]:![X9]:(~hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X8),X9))|(~hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X7),X8))|hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X7),X9))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
cnf(c_0_13,plain,(hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(bot_bo784226126e_bool),hAPP_f806699093e_bool(hAPP_H1902130436e_bool(insert1744391420_state,hAPP_c1455475371_state(hoare_Mirabelle_MGT,X1)),bot_bo784226126e_bool)))|~hBOOL(hAPP_com_bool(wt,X1))|~hBOOL(wT_bodies)|~hBOOL(hoare_1795711768gleton)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,hypothesis,(hBOOL(hoare_1795711768gleton)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_15, plain, (![X4]:hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X4),bot_bo784226126e_bool))), inference(variable_rename,[status(thm)],[c_0_5])).
cnf(c_0_16,plain,(hBOOL(hAPP_com_bool(wt,X1))|hAPP_com_option_com(some_com,X1)!=hAPP_p799580910on_com(body,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11])])).
cnf(c_0_17,hypothesis,(hAPP_p799580910on_com(body,pn)=hAPP_com_option_com(some_com,y)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_18, negated_conjecture, (~hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(hAPP_f360545851e_bool(image_185131637_state(hAPP_f1377420673_state(cOMBB_271860050_pname(hoare_Mirabelle_MGT),body_1)),dom_pname_com(body))),hAPP_f806699093e_bool(hAPP_H1902130436e_bool(insert1744391420_state,hAPP_c1455475371_state(hoare_Mirabelle_MGT,y)),bot_bo784226126e_bool)))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_7])])).
cnf(c_0_19,plain,(hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X1),X2))|~hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X1),X3))|~hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X3),X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20,plain,(hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(bot_bo784226126e_bool),hAPP_f806699093e_bool(hAPP_H1902130436e_bool(insert1744391420_state,hAPP_c1455475371_state(hoare_Mirabelle_MGT,X1)),bot_bo784226126e_bool)))|~hBOOL(hAPP_com_bool(wt,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_11]), c_0_14])])).
cnf(c_0_21,plain,(hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X1),bot_bo784226126e_bool))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,hypothesis,(hBOOL(hAPP_com_bool(wt,y))|hAPP_p799580910on_com(body,pn)!=hAPP_p799580910on_com(body,X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23,negated_conjecture,(~hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(hAPP_f360545851e_bool(image_185131637_state(hAPP_f1377420673_state(cOMBB_271860050_pname(hoare_Mirabelle_MGT),body_1)),dom_pname_com(body))),hAPP_f806699093e_bool(hAPP_H1902130436e_bool(insert1744391420_state,hAPP_c1455475371_state(hoare_Mirabelle_MGT,y)),bot_bo784226126e_bool)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,plain,(hBOOL(hAPP_f1378282496l_bool(hoare_512830354_state(X1),hAPP_f806699093e_bool(hAPP_H1902130436e_bool(insert1744391420_state,hAPP_c1455475371_state(hoare_Mirabelle_MGT,X2)),bot_bo784226126e_bool)))|~hBOOL(hAPP_com_bool(wt,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
cnf(c_0_25,hypothesis,(hBOOL(hAPP_com_bool(wt,y))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_26,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])]), ['proof']).
# SZS output end CNFRefutation.
EOF
