% Problem    : LCL893+1 : TPTP v6.1.0. Released v5.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n146.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:22:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____102_C18_F1_PI_AE_Q4_CS_SP_S1S
# and selection function SelectComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X9]:![X10]:![X11]:('>='('+'(X9,X10),X11)<=>'>='(X10,'==>'(X9,X11)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_08)).
fof(c_0_1, axiom, (![X1]:'+'(X1,'0')=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_03)).
fof(c_0_2, conjecture, (![X21]:(h(X21)=X21=>X21='0')), file('/export/starexec/sandbox/benchmark/theBenchmark.p', goals_15)).
fof(c_0_3, axiom, (![X1]:'>='(X1,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_05)).
fof(c_0_4, axiom, (![X1]:h(X1)='==>'(h(X1),X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_14)).
fof(c_0_5, axiom, (![X7]:![X8]:(('>='(X7,X8)&'>='(X8,X7))=>X7=X8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_07)).
fof(c_0_6, axiom, (![X1]:'>='(X1,'0')), file('/export/starexec/sandbox/benchmark/theBenchmark.p', sos_09)).
fof(c_0_7, plain, (![X12]:![X13]:![X14]:![X15]:![X16]:![X17]:((~'>='('+'(X12,X13),X14)|'>='(X13,'==>'(X12,X14)))&(~'>='(X16,'==>'(X15,X17))|'>='('+'(X15,X16),X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])).
fof(c_0_8, plain, (![X2]:'+'(X2,'0')=X2), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_9, negated_conjecture, (~(![X21]:(h(X21)=X21=>X21='0'))), inference(assume_negation,[status(cth)],[c_0_2])).
cnf(c_0_10,plain,('>='(X1,'==>'(X2,X3))|~'>='('+'(X2,X1),X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,plain,('+'(X1,'0')=X1), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, plain, (![X2]:'>='(X2,X2)), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_13, plain, (![X2]:h(X2)='==>'(h(X2),X2)), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_14, negated_conjecture, ((h(esk1_0)=esk1_0&esk1_0!='0')), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
cnf(c_0_15,plain,('>='('0','==>'(X1,X2))|~'>='(X1,X2)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_16,plain,('>='(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(h(X1)='==>'(h(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,negated_conjecture,(h(esk1_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, (![X9]:![X10]:((~'>='(X9,X10)|~'>='(X10,X9))|X9=X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
cnf(c_0_20,plain,('>='('0','==>'(X1,X1))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21,negated_conjecture,('==>'(esk1_0,esk1_0)=esk1_0), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_22, plain, (![X2]:'>='(X2,'0')), inference(variable_rename,[status(thm)],[c_0_6])).
cnf(c_0_23,plain,(X1=X2|~'>='(X2,X1)|~'>='(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24,negated_conjecture,('>='('0',esk1_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25,plain,('>='(X1,'0')), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26,negated_conjecture,(esk1_0!='0'), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_27,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])]), c_0_26]), ['proof']).
# SZS output end CNFRefutation.
EOF
