% Problem    : SWV492+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n134.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:25:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:![X2]:(int_less(X1,X2)<=>?[X3]:(plus(X1,X3)=X2&int_less(int_zero,X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', plus_and_inverse)).
fof(c_0_1, axiom, (![X1]:![X2]:(int_less(X1,X2)=>X1!=X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', int_less_irreflexive)).
fof(c_0_2, axiom, (![X1]:![X2]:![X3]:((int_less(X1,X2)&int_less(X2,X3))=>int_less(X1,X3))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', int_less_transitive)).
fof(c_0_3, axiom, (![X4]:![X5]:![X6]:![X7]:((int_less(X4,X5)&int_leq(X6,X7))=>int_leq(plus(X4,X6),plus(X5,X7)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', plus_and_order1)).
fof(c_0_4, axiom, (![X1]:![X2]:plus(X1,X2)=plus(X2,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', plus_commutative)).
fof(c_0_5, axiom, (![X1]:plus(X1,int_zero)=X1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', plus_zero)).
fof(c_0_6, axiom, (![X1]:![X2]:(int_leq(X1,X2)<=>(int_less(X1,X2)|X1=X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', int_leq)).
fof(c_0_7, conjecture, ((![X1]:![X2]:(((int_leq(int_one,X1)&int_less(X1,X2))&int_leq(X2,n))=>a(X1,X2)=real_zero)&![X1]:![X2]:(((int_leq(int_one,X1)&int_leq(X2,n))&X1=X2)=>a(X1,X2)!=real_zero))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', lti)).
fof(c_0_8, hypothesis, (![X1]:![X2]:((((int_leq(int_one,X1)&int_leq(X1,n))&int_leq(int_one,X2))&int_leq(X2,n))=>((![X8]:((int_less(int_zero,X8)&X1=plus(X2,X8))=>![X3]:((int_leq(int_one,X3)&int_leq(X3,X2))=>a(plus(X3,X8),X3)=qr(plus(X3,X8),X3)))&![X3]:((int_leq(int_one,X3)&int_leq(X3,X2))=>a(X3,X3)=real_one))&![X8]:((int_less(int_zero,X8)&X2=plus(X1,X8))=>![X3]:((int_leq(int_one,X3)&int_leq(X3,X1))=>a(X3,plus(X3,X8))=real_zero))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', qih)).
fof(c_0_9, axiom, (![X1]:![X2]:(int_less(X1,X2)|int_leq(X2,X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', int_less_total)).
fof(c_0_10, axiom, (![X1]:(int_less(int_zero,X1)<=>int_leq(int_one,X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', one_successor_of_zero)).
fof(c_0_11, axiom, (real_zero!=real_one), file('/export/starexec/sandbox/benchmark/theBenchmark.p', real_constants)).
fof(c_0_12, axiom, (int_less(int_zero,int_one)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', int_zero_one)).
fof(c_0_13, plain, (![X4]:![X5]:![X7]:![X8]:![X9]:(((plus(X4,esk1_2(X4,X5))=X5|~int_less(X4,X5))&(int_less(int_zero,esk1_2(X4,X5))|~int_less(X4,X5)))&((plus(X7,X9)!=X8|~int_less(int_zero,X9))|int_less(X7,X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])])).
fof(c_0_14, plain, (![X2]:![X1]:(epred1_2(X1,X2)<=>((![X8]:((int_less(int_zero,X8)&X1=plus(X2,X8))=>![X3]:((int_leq(int_one,X3)&int_leq(X3,X2))=>a(plus(X3,X8),X3)=qr(plus(X3,X8),X3)))&![X3]:((int_leq(int_one,X3)&int_leq(X3,X2))=>a(X3,X3)=real_one))&![X8]:((int_less(int_zero,X8)&X2=plus(X1,X8))=>![X3]:((int_leq(int_one,X3)&int_leq(X3,X1))=>a(X3,plus(X3,X8))=real_zero))))), introduced(definition)).
fof(c_0_15, plain, (![X3]:![X4]:(~int_less(X3,X4)|X3!=X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
fof(c_0_16, plain, (![X4]:![X5]:![X6]:((~int_less(X4,X5)|~int_less(X5,X6))|int_less(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
cnf(c_0_17,plain,(int_less(X1,X2)|~int_less(int_zero,X3)|plus(X1,X3)!=X2), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_18, plain, (![X2]:![X1]:(epred1_2(X1,X2)=>((![X8]:((int_less(int_zero,X8)&X1=plus(X2,X8))=>![X3]:((int_leq(int_one,X3)&int_leq(X3,X2))=>a(plus(X3,X8),X3)=qr(plus(X3,X8),X3)))&![X3]:((int_leq(int_one,X3)&int_leq(X3,X2))=>a(X3,X3)=real_one))&![X8]:((int_less(int_zero,X8)&X2=plus(X1,X8))=>![X3]:((int_leq(int_one,X3)&int_leq(X3,X1))=>a(X3,plus(X3,X8))=real_zero))))), inference(split_equiv,[status(thm)],[c_0_14])).
cnf(c_0_19,plain,(X1!=X2|~int_less(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20,plain,(int_less(X1,X2)|~int_less(X3,X2)|~int_less(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21,plain,(int_less(X1,plus(X1,X2))|~int_less(int_zero,X2)), inference(er,[status(thm)],[c_0_17])).
fof(c_0_22, plain, (![X8]:![X9]:![X10]:![X11]:((~int_less(X8,X9)|~int_leq(X10,X11))|int_leq(plus(X8,X10),plus(X9,X11)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
fof(c_0_23, plain, (![X3]:![X4]:plus(X3,X4)=plus(X4,X3)), inference(variable_rename,[status(thm)],[c_0_4])).
fof(c_0_24, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:(((((~int_less(int_zero,X11)|X10!=plus(X9,X11))|((~int_leq(int_one,X12)|~int_leq(X12,X9))|a(plus(X12,X11),X12)=qr(plus(X12,X11),X12)))|~epred1_2(X10,X9))&(((~int_leq(int_one,X13)|~int_leq(X13,X9))|a(X13,X13)=real_one)|~epred1_2(X10,X9)))&(((~int_less(int_zero,X14)|X9!=plus(X10,X14))|((~int_leq(int_one,X15)|~int_leq(X15,X10))|a(X15,plus(X15,X14))=real_zero))|~epred1_2(X10,X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).
cnf(c_0_25,plain,(~int_less(X1,X1)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_26,plain,(int_less(X1,plus(X2,X3))|~int_less(int_zero,X3)|~int_less(X1,X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27,plain,(int_leq(plus(X1,X2),plus(X3,X4))|~int_leq(X2,X4)|~int_less(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28,plain,(plus(X1,X2)=plus(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_29, plain, (![X2]:plus(X2,int_zero)=X2), inference(variable_rename,[status(thm)],[c_0_5])).
cnf(c_0_30,plain,(a(X3,plus(X3,X4))=real_zero|~epred1_2(X1,X2)|~int_leq(X3,X1)|~int_leq(int_one,X3)|X2!=plus(X1,X4)|~int_less(int_zero,X4)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_31, plain, (![X3]:![X4]:![X5]:![X6]:((~int_leq(X3,X4)|(int_less(X3,X4)|X3=X4))&((~int_less(X5,X6)|int_leq(X5,X6))&(X5!=X6|int_leq(X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
cnf(c_0_32,plain,(~int_less(plus(X1,X2),X1)|~int_less(int_zero,X2)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33,plain,(plus(X1,esk1_2(X1,X2))=X2|~int_less(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_34,plain,(int_less(int_zero,esk1_2(X1,X2))|~int_less(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35,plain,(int_leq(plus(X1,X2),plus(X3,X4))|~int_less(X1,X4)|~int_leq(X2,X3)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36,plain,(plus(X1,int_zero)=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_37, negated_conjecture, (~((![X1]:![X2]:(((int_leq(int_one,X1)&int_less(X1,X2))&int_leq(X2,n))=>a(X1,X2)=real_zero)&![X1]:![X2]:(((int_leq(int_one,X1)&int_leq(X2,n))&X1=X2)=>a(X1,X2)!=real_zero)))), inference(assume_negation,[status(cth)],[c_0_7])).
cnf(c_0_38,plain,(a(X1,plus(X1,X2))=real_zero|~epred1_2(X3,plus(X3,X2))|~int_less(int_zero,X2)|~int_leq(int_one,X1)|~int_leq(X1,X3)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_39,plain,(int_leq(X1,X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_40, hypothesis, (![X1]:![X2]:((((int_leq(int_one,X1)&int_leq(X1,n))&int_leq(int_one,X2))&int_leq(X2,n))=>epred1_2(X1,X2))), inference(apply_def,[status(thm)],[c_0_8, c_0_14])).
cnf(c_0_41,plain,(~int_less(X1,X2)|~int_less(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_42,plain,(X1=X2|int_less(X1,X2)|~int_leq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43,plain,(int_leq(plus(X1,X2),X3)|~int_less(X1,esk1_2(X4,X3))|~int_less(X4,X3)|~int_leq(X2,X4)), inference(spm,[status(thm)],[c_0_35, c_0_33])).
cnf(c_0_44,plain,(plus(int_zero,X1)=X1), inference(spm,[status(thm)],[c_0_36, c_0_28])).
fof(c_0_45, negated_conjecture, ((((((((int_leq(int_one,esk4_0)|int_leq(int_one,esk2_0))&(int_leq(esk5_0,n)|int_leq(int_one,esk2_0)))&(esk4_0=esk5_0|int_leq(int_one,esk2_0)))&(a(esk4_0,esk5_0)=real_zero|int_leq(int_one,esk2_0)))&((((int_leq(int_one,esk4_0)|int_less(esk2_0,esk3_0))&(int_leq(esk5_0,n)|int_less(esk2_0,esk3_0)))&(esk4_0=esk5_0|int_less(esk2_0,esk3_0)))&(a(esk4_0,esk5_0)=real_zero|int_less(esk2_0,esk3_0))))&((((int_leq(int_one,esk4_0)|int_leq(esk3_0,n))&(int_leq(esk5_0,n)|int_leq(esk3_0,n)))&(esk4_0=esk5_0|int_leq(esk3_0,n)))&(a(esk4_0,esk5_0)=real_zero|int_leq(esk3_0,n))))&((((int_leq(int_one,esk4_0)|a(esk2_0,esk3_0)!=real_zero)&(int_leq(esk5_0,n)|a(esk2_0,esk3_0)!=real_zero))&(esk4_0=esk5_0|a(esk2_0,esk3_0)!=real_zero))&(a(esk4_0,esk5_0)=real_zero|a(esk2_0,esk3_0)!=real_zero)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])])).
fof(c_0_46, plain, (![X3]:![X4]:(int_less(X3,X4)|int_leq(X4,X3))), inference(variable_rename,[status(thm)],[c_0_9])).
cnf(c_0_47,plain,(a(X1,plus(X1,esk1_2(X2,X3)))=real_zero|~epred1_2(X2,X3)|~int_less(X2,X3)|~int_leq(int_one,X1)|~int_leq(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_33]), c_0_34])).
cnf(c_0_48,plain,(int_leq(X1,X1)), inference(er,[status(thm)],[c_0_39])).
fof(c_0_49, hypothesis, (![X3]:![X4]:((((~int_leq(int_one,X3)|~int_leq(X3,n))|~int_leq(int_one,X4))|~int_leq(X4,n))|epred1_2(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])).
cnf(c_0_50,plain,(X1=X2|~int_less(X2,X1)|~int_leq(X1,X2)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51,plain,(int_leq(X1,X2)|~int_less(X3,X2)|~int_leq(X1,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_34]), c_0_44])).
cnf(c_0_52,negated_conjecture,(int_less(esk2_0,esk3_0)|esk4_0=esk5_0), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_53,plain,(int_leq(X1,X2)|int_less(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_54,plain,(a(X1,X2)=real_zero|~epred1_2(X1,X2)|~int_less(X1,X2)|~int_leq(int_one,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_33]), c_0_48])])).
cnf(c_0_55,hypothesis,(epred1_2(X1,X2)|~int_leq(X2,n)|~int_leq(int_one,X2)|~int_leq(X1,n)|~int_leq(int_one,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56,plain,(X1=X2|~int_leq(X1,X2)|~int_leq(X2,X1)), inference(spm,[status(thm)],[c_0_50, c_0_42])).
cnf(c_0_57,negated_conjecture,(esk4_0=esk5_0|int_leq(X1,esk3_0)|~int_leq(X1,esk2_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_58,plain,(int_leq(X1,X2)|int_leq(X3,X1)|~int_leq(X3,X2)), inference(spm,[status(thm)],[c_0_51, c_0_53])).
cnf(c_0_59,negated_conjecture,(int_leq(esk3_0,n)|esk4_0=esk5_0), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_60,negated_conjecture,(esk4_0=esk5_0|a(esk2_0,esk3_0)!=real_zero), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_61,hypothesis,(a(X1,X2)=real_zero|~int_less(X1,X2)|~int_leq(int_one,X1)|~int_leq(X2,n)|~int_leq(X1,n)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_51])).
cnf(c_0_62,negated_conjecture,(int_leq(int_one,esk2_0)|esk4_0=esk5_0), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_63,negated_conjecture,(esk4_0=esk5_0|X1=esk3_0|~int_leq(esk3_0,X1)|~int_leq(X1,esk2_0)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_64,negated_conjecture,(esk4_0=esk5_0|int_leq(esk3_0,X1)|int_leq(X1,n)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65,plain,(a(X3,X3)=real_one|~epred1_2(X1,X2)|~int_leq(X3,X2)|~int_leq(int_one,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_66,negated_conjecture,(esk4_0=esk5_0|~int_leq(esk2_0,n)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_59]), c_0_62]), c_0_52])).
cnf(c_0_67,negated_conjecture,(esk4_0=esk5_0|X1=esk3_0|int_leq(X1,n)|~int_leq(X1,esk2_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68,plain,(X1=X2|int_less(X3,X2)|~int_less(X3,X1)|~int_leq(X1,X2)), inference(spm,[status(thm)],[c_0_20, c_0_42])).
cnf(c_0_69,negated_conjecture,(int_leq(esk3_0,n)|int_leq(int_one,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_70,hypothesis,(a(X1,X1)=real_one|~int_leq(int_one,X1)|~int_leq(X2,n)|~int_leq(X3,n)|~int_leq(int_one,X2)|~int_leq(int_one,X3)|~int_leq(X1,X2)), inference(spm,[status(thm)],[c_0_65, c_0_55])).
cnf(c_0_71,negated_conjecture,(esk3_0=esk2_0|esk4_0=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_48])])).
cnf(c_0_72,negated_conjecture,(n=esk3_0|int_less(X1,n)|int_leq(int_one,esk4_0)|~int_less(X1,esk3_0)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_73,negated_conjecture,(int_less(esk2_0,esk3_0)|int_leq(int_one,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_74,plain,(int_leq(X1,plus(X2,X3))|~int_less(int_zero,X3)|~int_leq(X1,X2)), inference(spm,[status(thm)],[c_0_35, c_0_44])).
fof(c_0_75, plain, (![X2]:![X3]:((~int_less(int_zero,X2)|int_leq(int_one,X2))&(~int_leq(int_one,X3)|int_less(int_zero,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
cnf(c_0_76,negated_conjecture,(int_less(esk2_0,esk3_0)|int_leq(esk5_0,n)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_77,hypothesis,(a(X1,X1)=real_one|~int_leq(int_one,n)|~int_leq(int_one,X1)|~int_leq(X2,n)|~int_leq(int_one,X2)|~int_leq(X1,n)), inference(spm,[status(thm)],[c_0_70, c_0_48])).
cnf(c_0_78,negated_conjecture,(int_less(esk2_0,esk3_0)|a(esk4_0,esk5_0)=real_zero), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_79,negated_conjecture,(esk4_0=esk5_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_71]), c_0_25])).
cnf(c_0_80,negated_conjecture,(n=esk3_0|int_leq(int_one,esk4_0)|~int_less(n,esk3_0)), inference(spm,[status(thm)],[c_0_25, c_0_72])).
cnf(c_0_81,negated_conjecture,(int_less(X1,esk3_0)|int_leq(int_one,esk4_0)|~int_less(X1,esk2_0)), inference(spm,[status(thm)],[c_0_20, c_0_73])).
cnf(c_0_82,plain,(X1=plus(X2,X3)|~int_less(int_zero,X3)|~int_leq(plus(X2,X3),X1)|~int_leq(X1,X2)), inference(spm,[status(thm)],[c_0_56, c_0_74])).
cnf(c_0_83,negated_conjecture,(int_leq(esk3_0,n)|int_leq(esk5_0,n)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_84,plain,(int_leq(int_one,X1)|~int_less(int_zero,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_85,negated_conjecture,(int_less(X1,esk3_0)|int_leq(esk5_0,n)|~int_less(X1,esk2_0)), inference(spm,[status(thm)],[c_0_20, c_0_76])).
cnf(c_0_86,plain,(~int_less(plus(X1,X2),X2)|~int_less(int_zero,X1)), inference(spm,[status(thm)],[c_0_32, c_0_28])).
cnf(c_0_87,hypothesis,(a(X1,X1)=real_one|~int_leq(int_one,n)|~int_leq(int_one,X1)|~int_leq(X1,n)), inference(spm,[status(thm)],[c_0_77, c_0_48])).
cnf(c_0_88,negated_conjecture,(a(esk5_0,esk5_0)=real_zero|int_less(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_89,plain,(real_zero!=real_one), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_90,negated_conjecture,(int_less(esk2_0,esk3_0)|int_leq(int_one,esk5_0)), inference(rw,[status(thm)],[c_0_73, c_0_79])).
cnf(c_0_91,negated_conjecture,(int_leq(esk3_0,n)|a(esk4_0,esk5_0)=real_zero), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_92,negated_conjecture,(int_leq(int_one,esk4_0)|a(esk2_0,esk3_0)!=real_zero), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_93,negated_conjecture,(int_leq(int_one,esk2_0)|int_leq(int_one,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_94,negated_conjecture,(n=esk3_0|int_leq(int_one,esk4_0)|~int_less(n,esk2_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_95,plain,(X1=X2|~int_less(X3,X2)|~int_leq(X2,X1)|~int_leq(X1,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_33]), c_0_34])).
cnf(c_0_96,negated_conjecture,(n=esk3_0|int_less(X1,n)|int_leq(esk5_0,n)|~int_less(X1,esk3_0)), inference(spm,[status(thm)],[c_0_68, c_0_83])).
cnf(c_0_97,negated_conjecture,(int_leq(esk5_0,n)|int_leq(int_one,esk3_0)|~int_less(int_zero,esk2_0)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_98,plain,(int_less(int_zero,X1)|~int_leq(int_one,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_99,negated_conjecture,(int_leq(int_one,esk2_0)|int_leq(esk5_0,n)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_100,plain,(plus(X1,X2)=X2|~int_less(int_zero,X1)|~int_leq(plus(X1,X2),X2)), inference(spm,[status(thm)],[c_0_86, c_0_42])).
cnf(c_0_101,plain,(int_leq(plus(X1,X2),X3)|~int_less(X1,X3)|~int_leq(X2,int_zero)), inference(spm,[status(thm)],[c_0_27, c_0_36])).
cnf(c_0_102,negated_conjecture,(a(esk4_0,esk5_0)=real_zero|a(esk2_0,esk3_0)!=real_zero), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_103,negated_conjecture,(int_leq(int_one,esk2_0)|a(esk4_0,esk5_0)=real_zero), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_104,hypothesis,(int_less(esk2_0,esk3_0)|~int_leq(int_one,n)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89]), c_0_76]), c_0_90])).
cnf(c_0_105,negated_conjecture,(a(esk5_0,esk5_0)=real_zero|int_leq(esk3_0,n)), inference(rw,[status(thm)],[c_0_91, c_0_79])).
cnf(c_0_106,negated_conjecture,(int_leq(esk3_0,n)|int_leq(int_one,esk5_0)), inference(rw,[status(thm)],[c_0_69, c_0_79])).
cnf(c_0_107,negated_conjecture,(int_leq(int_one,esk5_0)|a(esk2_0,esk3_0)!=real_zero), inference(rw,[status(thm)],[c_0_92, c_0_79])).
cnf(c_0_108,negated_conjecture,(int_leq(int_one,esk5_0)|int_leq(int_one,esk2_0)), inference(rw,[status(thm)],[c_0_93, c_0_79])).
cnf(c_0_109,negated_conjecture,(n=esk3_0|int_leq(esk2_0,n)|int_leq(int_one,esk4_0)), inference(spm,[status(thm)],[c_0_94, c_0_53])).
cnf(c_0_110,plain,(X1=X2|int_leq(X2,X3)|~int_leq(X2,X1)|~int_leq(X1,X3)), inference(spm,[status(thm)],[c_0_95, c_0_53])).
cnf(c_0_111,plain,(int_leq(X1,int_zero)|int_leq(int_one,X1)), inference(spm,[status(thm)],[c_0_84, c_0_53])).
cnf(c_0_112,negated_conjecture,(n=esk3_0|int_leq(esk5_0,n)|int_leq(int_one,n)|~int_less(int_zero,esk3_0)), inference(spm,[status(thm)],[c_0_84, c_0_96])).
cnf(c_0_113,negated_conjecture,(int_leq(int_one,esk3_0)|int_leq(esk5_0,n)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99])).
cnf(c_0_114,plain,(int_less(X1,plus(X2,X1))|~int_less(int_zero,X2)), inference(spm,[status(thm)],[c_0_21, c_0_28])).
cnf(c_0_115,plain,(plus(X1,X2)=X2|~int_less(int_zero,X1)|~int_less(X1,X2)|~int_leq(X2,int_zero)), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_116,negated_conjecture,(a(esk5_0,esk5_0)=real_zero|a(esk2_0,esk3_0)!=real_zero), inference(rw,[status(thm)],[c_0_102, c_0_79])).
cnf(c_0_117,negated_conjecture,(a(esk5_0,esk5_0)=real_zero|int_leq(int_one,esk2_0)), inference(rw,[status(thm)],[c_0_103, c_0_79])).
cnf(c_0_118,negated_conjecture,(int_leq(esk5_0,n)|a(esk2_0,esk3_0)!=real_zero), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_119,hypothesis,(X1=esk3_0|~int_leq(int_one,n)|~int_leq(esk3_0,X1)|~int_leq(X1,esk2_0)), inference(spm,[status(thm)],[c_0_95, c_0_104])).
cnf(c_0_120,hypothesis,(int_leq(esk3_0,n)|~int_leq(int_one,n)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_105]), c_0_89]), c_0_83]), c_0_106])).
cnf(c_0_121,plain,(int_leq(X1,X2)|~int_less(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_122,hypothesis,(int_leq(int_one,esk5_0)|~int_leq(esk2_0,n)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_61]), c_0_106]), c_0_108]), c_0_90])).
cnf(c_0_123,negated_conjecture,(n=esk3_0|int_leq(int_one,esk5_0)|int_leq(esk2_0,n)), inference(rw,[status(thm)],[c_0_109, c_0_79])).
cnf(c_0_124,plain,(int_leq(X1,X2)|~int_less(int_zero,X2)|~int_leq(X1,int_zero)), inference(spm,[status(thm)],[c_0_101, c_0_44])).
cnf(c_0_125,plain,(int_less(int_zero,int_one)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_126,plain,(X1=int_one|int_leq(X1,int_zero)|int_leq(int_one,X2)|~int_leq(X1,X2)), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_127,negated_conjecture,(n=esk3_0|int_leq(int_one,n)|int_leq(esk5_0,n)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_98]), c_0_113])).
cnf(c_0_128,plain,(~int_less(int_zero,X1)|~int_less(X1,X2)|~int_leq(X2,int_zero)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_25])).
cnf(c_0_129,plain,(plus(X1,X2)=X3|~int_less(X1,X3)|~int_leq(X3,plus(X1,X2))|~int_leq(X2,int_zero)), inference(spm,[status(thm)],[c_0_56, c_0_101])).
cnf(c_0_130,hypothesis,(a(esk5_0,esk5_0)=real_zero|~int_leq(esk2_0,n)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_61]), c_0_105]), c_0_117]), c_0_88])).
cnf(c_0_131,negated_conjecture,(int_leq(esk5_0,n)|~int_leq(esk2_0,n)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_61]), c_0_83]), c_0_99]), c_0_76])).
cnf(c_0_132,hypothesis,(n=esk3_0|~int_leq(int_one,n)|~int_leq(n,esk2_0)), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_133,plain,(int_leq(X1,X2)|int_leq(X2,X1)), inference(spm,[status(thm)],[c_0_121, c_0_53])).
cnf(c_0_134,negated_conjecture,(n=esk3_0|int_leq(int_one,esk5_0)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_135,plain,(int_leq(X1,int_one)|~int_leq(X1,int_zero)), inference(spm,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_136,negated_conjecture,(n=esk3_0|esk5_0=int_one|int_leq(int_one,n)|int_leq(esk5_0,int_zero)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_137,plain,(~int_less(X1,X2)|~int_leq(X2,int_zero)|~int_leq(int_one,X1)), inference(spm,[status(thm)],[c_0_128, c_0_98])).
cnf(c_0_138,hypothesis,(int_leq(int_one,esk2_0)|~int_leq(int_one,n)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_117]), c_0_89]), c_0_99]), c_0_108])).
cnf(c_0_139,plain,(X1=X2|~int_less(int_zero,X2)|~int_leq(X1,int_zero)|~int_leq(X2,X1)), inference(spm,[status(thm)],[c_0_129, c_0_44])).
cnf(c_0_140,negated_conjecture,(int_leq(int_one,esk4_0)|int_leq(int_one,esk3_0)|~int_less(int_zero,esk2_0)), inference(spm,[status(thm)],[c_0_84, c_0_81])).
cnf(c_0_141,hypothesis,(~int_leq(int_one,n)|~int_leq(esk2_0,n)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_130]), c_0_89]), c_0_131]), c_0_122])).
cnf(c_0_142,hypothesis,(n=esk3_0|int_leq(esk2_0,n)|~int_leq(int_one,n)), inference(spm,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_143,negated_conjecture,(n=esk3_0|esk5_0=int_one|~int_leq(esk5_0,int_one)), inference(spm,[status(thm)],[c_0_56, c_0_134])).
cnf(c_0_144,negated_conjecture,(esk5_0=int_one|n=esk3_0|int_leq(int_one,n)|int_leq(esk5_0,int_one)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_145,hypothesis,(~int_leq(esk3_0,int_zero)|~int_leq(int_one,n)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_104]), c_0_138])).
cnf(c_0_146,plain,(X1=int_one|~int_leq(X1,int_zero)|~int_leq(int_one,X1)), inference(spm,[status(thm)],[c_0_139, c_0_125])).
cnf(c_0_147,negated_conjecture,(esk5_0=int_one|int_leq(int_one,esk3_0)|int_leq(int_one,n)|int_leq(esk5_0,int_zero)), inference(spm,[status(thm)],[c_0_126, c_0_113])).
cnf(c_0_148,negated_conjecture,(int_leq(int_one,esk3_0)|int_leq(int_one,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_98]), c_0_93])).
cnf(c_0_149,hypothesis,(n=esk3_0|~int_leq(int_one,n)), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_150,negated_conjecture,(esk5_0=int_one|n=esk3_0|int_leq(int_one,n)), inference(spm,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_151,hypothesis,(int_leq(int_one,esk3_0)|~int_leq(int_one,n)), inference(spm,[status(thm)],[c_0_145, c_0_111])).
cnf(c_0_152,negated_conjecture,(esk5_0=int_one|int_leq(int_one,n)|int_leq(int_one,esk3_0)|~int_leq(int_one,esk5_0)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_153,negated_conjecture,(int_leq(int_one,esk5_0)|int_leq(int_one,esk3_0)), inference(rw,[status(thm)],[c_0_148, c_0_79])).
cnf(c_0_154,negated_conjecture,(esk5_0=int_one|n=esk3_0), inference(spm,[status(thm)],[c_0_149, c_0_150])).
cnf(c_0_155,negated_conjecture,(esk5_0=int_one|int_leq(int_one,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_153])).
cnf(c_0_156,hypothesis,(esk5_0=int_one|~int_leq(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_154]), c_0_155])).
cnf(c_0_157,hypothesis,(int_leq(esk2_0,esk3_0)|~int_leq(int_one,n)), inference(spm,[status(thm)],[c_0_121, c_0_104])).
cnf(c_0_158,negated_conjecture,(int_leq(esk5_0,n)|~int_leq(esk3_0,int_zero)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_76]), c_0_99])).
cnf(c_0_159,hypothesis,(esk5_0=int_one|~int_leq(int_one,n)), inference(spm,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_160,negated_conjecture,(esk3_0=int_one|int_leq(esk5_0,n)|int_leq(int_one,n)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_83]), c_0_158])).
cnf(c_0_161,negated_conjecture,(esk5_0=int_one), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_154]), c_0_155])).
cnf(c_0_162,negated_conjecture,(n=esk3_0|int_leq(esk5_0,n)|~int_less(n,esk3_0)), inference(spm,[status(thm)],[c_0_25, c_0_96])).
cnf(c_0_163,negated_conjecture,(esk3_0=int_one|int_leq(int_one,n)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_160, c_0_161])])).
cnf(c_0_164,negated_conjecture,(n=esk3_0|int_leq(esk5_0,n)|~int_less(n,esk2_0)), inference(spm,[status(thm)],[c_0_162, c_0_85])).
cnf(c_0_165,negated_conjecture,(~int_leq(esk2_0,n)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_161]), c_0_141])).
cnf(c_0_166,hypothesis,(esk3_0=int_one|n=esk3_0), inference(spm,[status(thm)],[c_0_149, c_0_163])).
cnf(c_0_167,negated_conjecture,(n=esk3_0|int_leq(esk2_0,n)|int_leq(esk5_0,n)), inference(spm,[status(thm)],[c_0_164, c_0_53])).
cnf(c_0_168,negated_conjecture,(esk3_0=int_one|~int_leq(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_169,negated_conjecture,(n=esk3_0|int_leq(esk5_0,n)), inference(spm,[status(thm)],[c_0_131, c_0_167])).
cnf(c_0_170,hypothesis,(esk3_0=int_one), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_157]), c_0_163])).
cnf(c_0_171,hypothesis,(~int_less(esk3_0,esk2_0)|~int_leq(int_one,n)), inference(spm,[status(thm)],[c_0_41, c_0_104])).
cnf(c_0_172,negated_conjecture,(n=int_one|int_leq(int_one,n)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_169, c_0_161]), c_0_170])).
cnf(c_0_173,hypothesis,(~int_less(int_one,esk2_0)|~int_leq(int_one,n)), inference(rw,[status(thm)],[c_0_171, c_0_170])).
cnf(c_0_174,hypothesis,(n=int_one), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_149, c_0_170]), c_0_172])).
cnf(c_0_175,hypothesis,(~int_less(int_one,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_173, c_0_174]), c_0_48])])).
cnf(c_0_176,hypothesis,(int_leq(int_one,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_138, c_0_174]), c_0_48])])).
cnf(c_0_177,negated_conjecture,(~int_leq(esk2_0,int_one)), inference(rw,[status(thm)],[c_0_165, c_0_174])).
cnf(c_0_178,hypothesis,(esk2_0=int_one), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_42]), c_0_176])])).
cnf(c_0_179,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_177, c_0_178]), c_0_48])]), ['proof']).
# SZS output end CNFRefutation.
EOF
