% Problem    : SWV468+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n143.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:35:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2f
# and selection function SelectCQIAr.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X12]:![X13]:![X5]:![X6]:(((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~(leq(host(X7),host(X14))))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~(setIn(X7,alive))|~(setIn(X14,alive)))))&![X7]:![X14]:((setIn(X14,alive)&index(status,host(X14))=elec_2)=>~(elem(m_Ack(X7,X14),queue(host(X7))))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X14]:((((setIn(X7,alive)&setIn(X14,alive))&setIn(host(X14),index(down,host(X7))))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X14]:(((setIn(X14,alive)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&setIn(host(X14),index(down,host(X7)))))))&![X7]:![X15]:![X14]:(((((~(leq(host(X14),host(X7)))&setIn(X7,alive))&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>~(elem(m_Down(X15),queue(host(X7))))))&![X7]:![X14]:(((((~(leq(host(X7),host(X14)))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:(((setIn(X14,alive)&elem(m_Ack(X14,X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~(elem(m_Ack(X7,X15),queue(host(X7))))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&elem(m_Down(X15),queue(host(X7))))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&host(X15)=host(X14))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&elem(m_Down(X15),queue(host(X7)))))))&![X7]:![X14]:(((~(leq(index(pendack,host(X14)),host(X7)))&setIn(X14,alive))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X14]:(((((~(leq(host(X7),host(X14)))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~(leq(index(pendack,host(X7)),index(pendack,host(X14))))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:((((![X16]:((~(leq(host(X14),X16))&leq(s(zero),X16))=>(setIn(X16,index(down,host(X14)))|X16=host(X15)))&~(leq(host(X14),host(X7))))&elem(m_Down(X15),queue(host(X14))))&index(status,host(X14))=elec_1)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12))=>(setIn(X5,alive)=>(~(leq(host(X5),host(X6)))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>((![X7]:((~(leq(host(X5),X7))&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1)=>(leq(nbr_proc,host(X5))=>![X7]:(~(setIn(host(X7),setEmpty))=>![X16]:(host(X5)=host(X16)=>![X17]:(host(X5)!=host(X17)=>(((~(leq(index(pendack,host(X17)),host(X16)))&setIn(X17,alive))&index(status,host(X17))=elec_2)=>~((setIn(X16,alive)&host(X5)=host(X16)))))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj)).
fof(c_0_1, axiom, (![X5]:![X6]:((leq(X5,X6)&leq(X6,X5))<=>X5=X6)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_61)).
fof(c_0_2, axiom, (![X3]:leq(host(X3),nbr_proc)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_04)).
fof(c_0_3, axiom, (![X5]:![X4]:tail(cons(X5,X4))=X4), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_36)).
fof(c_0_4, axiom, (![X5]:![X4]:q_nil!=cons(X5,X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_41)).
fof(c_0_5, axiom, (![X3]:leq(s(zero),host(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_02)).
fof(c_0_6, axiom, (![X5]:![X6]:![X4]:(elem(X5,cons(X6,X4))<=>(X5=X6|elem(X5,X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_46)).
fof(c_0_7, axiom, (![X4]:(X4=q_nil|X4=cons(head(X4),tail(X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_39)).
fof(c_0_8, axiom, (![X5]:![X4]:head(cons(X5,X4))=X5), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_35)).
fof(c_0_9, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)<=>((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:![X14]:((setIn(X14,alive)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X14),queue(host(X7)))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X14]:((((setIn(X7,alive)&setIn(X14,alive))&setIn(host(X14),index(down,host(X7))))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X14]:(((setIn(X14,alive)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&setIn(host(X14),index(down,host(X7)))))))&![X7]:![X15]:![X14]:(((((~leq(host(X14),host(X7))&setIn(X7,alive))&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>~elem(m_Down(X15),queue(host(X7)))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:(((setIn(X14,alive)&elem(m_Ack(X14,X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X15),queue(host(X7)))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&elem(m_Down(X15),queue(host(X7))))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&host(X15)=host(X14))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&elem(m_Down(X15),queue(host(X7)))))))&![X7]:![X14]:(((~leq(index(pendack,host(X14)),host(X7))&setIn(X14,alive))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~leq(index(pendack,host(X7)),index(pendack,host(X14)))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:((((![X16]:((~leq(host(X14),X16)&leq(s(zero),X16))=>(setIn(X16,index(down,host(X14)))|X16=host(X15)))&~leq(host(X14),host(X7)))&elem(m_Down(X15),queue(host(X14))))&index(status,host(X14))=elec_1)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12)))), introduced(definition)).
fof(c_0_10, negated_conjecture, (~(![X12]:![X5]:![X6]:(epred1_3(X5,X6,X12)=>(setIn(X5,alive)=>(~leq(host(X5),host(X6))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>((![X7]:((~leq(host(X5),X7)&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1)=>(leq(nbr_proc,host(X5))=>![X7]:(~setIn(host(X7),setEmpty)=>![X16]:(host(X5)=host(X16)=>![X17]:(host(X5)!=host(X17)=>(((~leq(index(pendack,host(X17)),host(X16))&setIn(X17,alive))&index(status,host(X17))=elec_2)=>~((setIn(X16,alive)&host(X5)=host(X16))))))))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_9])).
fof(c_0_11, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)=>((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:![X14]:((setIn(X14,alive)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X14),queue(host(X7)))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X14]:((((setIn(X7,alive)&setIn(X14,alive))&setIn(host(X14),index(down,host(X7))))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X14]:(((setIn(X14,alive)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&setIn(host(X14),index(down,host(X7)))))))&![X7]:![X15]:![X14]:(((((~leq(host(X14),host(X7))&setIn(X7,alive))&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>~elem(m_Down(X15),queue(host(X7)))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:(((setIn(X14,alive)&elem(m_Ack(X14,X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X15),queue(host(X7)))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&elem(m_Down(X15),queue(host(X7))))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&host(X15)=host(X14))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&elem(m_Down(X15),queue(host(X7)))))))&![X7]:![X14]:(((~leq(index(pendack,host(X14)),host(X7))&setIn(X14,alive))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~leq(index(pendack,host(X7)),index(pendack,host(X14)))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:((((![X16]:((~leq(host(X14),X16)&leq(s(zero),X16))=>(setIn(X16,index(down,host(X14)))|X16=host(X15)))&~leq(host(X14),host(X7)))&elem(m_Down(X15),queue(host(X14))))&index(status,host(X14))=elec_1)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12)))), inference(split_equiv,[status(thm)],[c_0_9])).
fof(c_0_12, plain, (![X7]:![X8]:![X9]:![X10]:(((~leq(X7,X8)|~leq(X8,X7))|X7=X8)&((leq(X9,X10)|X9!=X10)&(leq(X10,X9)|X9!=X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
fof(c_0_13, negated_conjecture, (![X21]:(epred1_3(esk6_0,esk7_0,esk5_0)&(setIn(esk6_0,alive)&(~leq(host(esk6_0),host(esk7_0))&(((index(ldr,host(esk6_0))!=host(esk7_0)|index(status,host(esk6_0))!=norm)&(index(status,host(esk6_0))!=wait|host(esk7_0)!=host(index(elid,host(esk6_0)))))&((((leq(host(esk6_0),X21)|~leq(s(zero),X21))|(setIn(X21,index(down,host(esk6_0)))|X21=host(esk7_0)))&index(status,host(esk6_0))=elec_1)&(leq(nbr_proc,host(esk6_0))&(~setIn(host(esk8_0),setEmpty)&(host(esk6_0)=host(esk9_0)&(host(esk6_0)!=host(esk10_0)&(((~leq(index(pendack,host(esk10_0)),host(esk9_0))&setIn(esk10_0,alive))&index(status,host(esk10_0))=elec_2)&(setIn(esk9_0,alive)&host(esk6_0)=host(esk9_0))))))))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
fof(c_0_14, plain, (![X4]:leq(host(X4),nbr_proc)), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_15, plain, (![X17]:![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:![X34]:![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:![X43]:![X44]:![X45]:![X46]:![X47]:![X48]:![X49]:![X50]:![X51]:![X52]:![X53]:![X54]:![X55]:![X56]:![X57]:((((((((((((((((((~elem(m_Ldr(X21),queue(host(X20)))|~leq(host(X20),host(X21)))|~epred1_3(X19,X18,X17))&(((X23=X22|host(X23)!=host(X22))|(~setIn(X22,alive)|~setIn(X23,alive)))|~epred1_3(X19,X18,X17)))&(((~setIn(X25,alive)|index(status,host(X25))!=elec_2)|~elem(m_Ack(X24,X25),queue(host(X24))))|~epred1_3(X19,X18,X17)))&((((index(status,host(X26))!=elec_1|~setIn(X26,alive))|index(elid,host(X26))=X26)|~epred1_3(X19,X18,X17))&(((index(status,host(X26))!=elec_2|~setIn(X26,alive))|index(elid,host(X26))=X26)|~epred1_3(X19,X18,X17))))&(((((~setIn(X27,alive)|~setIn(X28,alive))|~setIn(host(X28),index(down,host(X27))))|index(status,host(X28))!=elec_2)|leq(index(pendack,host(X28)),host(X27)))|~epred1_3(X19,X18,X17)))&((((~setIn(X30,alive)|index(status,host(X30))!=norm)|index(ldr,host(X30))!=host(X30))|(~setIn(X29,alive)|~setIn(host(X30),index(down,host(X29)))))|~epred1_3(X19,X18,X17)))&((((((leq(host(X33),host(X31))|~setIn(X31,alive))|~setIn(X33,alive))|host(X33)!=host(X32))|index(status,host(X33))!=elec_2)|~elem(m_Down(X32),queue(host(X31))))|~epred1_3(X19,X18,X17)))&((((((leq(host(X34),host(X35))|~setIn(X34,alive))|~setIn(X35,alive))|index(status,host(X34))!=elec_2)|index(status,host(X35))!=elec_2)|leq(index(pendack,host(X35)),host(X34)))|~epred1_3(X19,X18,X17)))&((((~setIn(X38,alive)|~elem(m_Ack(X38,X37),queue(host(X38))))|host(X37)!=host(X36))|((~setIn(X36,alive)|index(ldr,host(X36))!=host(X36))|index(status,host(X36))!=norm))|~epred1_3(X19,X18,X17)))&((((((~setIn(X39,alive)|~setIn(X41,alive))|host(X41)!=host(X40))|index(status,host(X39))!=elec_2)|index(status,host(X41))!=elec_2)|~elem(m_Ack(X39,X40),queue(host(X39))))|~epred1_3(X19,X18,X17)))&((((((~setIn(X42,alive)|~setIn(X44,alive))|~elem(m_Down(X43),queue(host(X42))))|host(X44)!=host(X43))|index(status,host(X44))!=elec_2)|leq(index(pendack,host(X44)),host(X42)))|~epred1_3(X19,X18,X17)))&(((((~setIn(X47,alive)|host(X46)!=host(X47))|index(status,host(X47))!=norm)|index(ldr,host(X47))!=host(X47))|(~setIn(X45,alive)|~elem(m_Down(X46),queue(host(X45)))))|~epred1_3(X19,X18,X17)))&((((leq(index(pendack,host(X49)),host(X48))|~setIn(X49,alive))|index(status,host(X49))!=elec_2)|((~setIn(X48,alive)|index(ldr,host(X48))!=host(X48))|index(status,host(X48))!=norm))|~epred1_3(X19,X18,X17)))&((((((leq(host(X50),host(X51))|~setIn(X50,alive))|~setIn(X51,alive))|index(status,host(X50))!=elec_2)|index(status,host(X51))!=elec_2)|~leq(index(pendack,host(X50)),index(pendack,host(X51))))|~epred1_3(X19,X18,X17)))&(((((~setIn(X54,alive)|~elem(m_Down(X53),queue(host(X54))))|host(X53)!=host(X52))|index(status,host(X54))!=elec_2)|((~setIn(X52,alive)|index(ldr,host(X52))!=host(X52))|index(status,host(X52))!=norm))|~epred1_3(X19,X18,X17)))&(((((((~leq(host(X57),esk11_6(X17,X18,X19,X55,X56,X57))|leq(host(X57),host(X55)))|~elem(m_Down(X56),queue(host(X57))))|index(status,host(X57))!=elec_1)|((~setIn(X55,alive)|index(ldr,host(X55))!=host(X55))|index(status,host(X55))!=norm))|~epred1_3(X19,X18,X17))&(((((leq(s(zero),esk11_6(X17,X18,X19,X55,X56,X57))|leq(host(X57),host(X55)))|~elem(m_Down(X56),queue(host(X57))))|index(status,host(X57))!=elec_1)|((~setIn(X55,alive)|index(ldr,host(X55))!=host(X55))|index(status,host(X55))!=norm))|~epred1_3(X19,X18,X17)))&((((((~setIn(esk11_6(X17,X18,X19,X55,X56,X57),index(down,host(X57)))|leq(host(X57),host(X55)))|~elem(m_Down(X56),queue(host(X57))))|index(status,host(X57))!=elec_1)|((~setIn(X55,alive)|index(ldr,host(X55))!=host(X55))|index(status,host(X55))!=norm))|~epred1_3(X19,X18,X17))&(((((esk11_6(X17,X18,X19,X55,X56,X57)!=host(X56)|leq(host(X57),host(X55)))|~elem(m_Down(X56),queue(host(X57))))|index(status,host(X57))!=elec_1)|((~setIn(X55,alive)|index(ldr,host(X55))!=host(X55))|index(status,host(X55))!=norm))|~epred1_3(X19,X18,X17)))))&(queue(host(X19))=cons(m_Down(X18),X17)|~epred1_3(X19,X18,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])])).
cnf(c_0_16,plain,(X1=X2|~leq(X2,X1)|~leq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17,negated_conjecture,(leq(nbr_proc,host(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18,plain,(leq(host(X1),nbr_proc)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19,negated_conjecture,(~leq(index(pendack,host(esk10_0)),host(esk9_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,negated_conjecture,(host(esk6_0)=host(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_21, plain, (![X6]:![X7]:tail(cons(X6,X7))=X7), inference(variable_rename,[status(thm)],[c_0_3])).
cnf(c_0_22,plain,(queue(host(X1))=cons(m_Down(X2),X3)|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23,negated_conjecture,(epred1_3(esk6_0,esk7_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24,negated_conjecture,(host(esk6_0)=nbr_proc), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])])).
fof(c_0_25, plain, (![X6]:![X7]:q_nil!=cons(X6,X7)), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_26,plain,(leq(index(pendack,host(X4)),host(X5))|~epred1_3(X1,X2,X3)|index(status,host(X4))!=elec_2|~setIn(host(X4),index(down,host(X5)))|~setIn(X4,alive)|~setIn(X5,alive)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27,negated_conjecture,(index(status,host(esk10_0))=elec_2), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_28,negated_conjecture,(setIn(esk10_0,alive)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_29,negated_conjecture,(~leq(index(pendack,host(esk10_0)),host(esk6_0))), inference(rw,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_30,negated_conjecture,(X1=host(esk7_0)|setIn(X1,index(down,host(esk6_0)))|leq(host(esk6_0),X1)|~leq(s(zero),X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_31, plain, (![X4]:leq(s(zero),host(X4))), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_32, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((~elem(X7,cons(X8,X9))|(X7=X8|elem(X7,X9)))&((X10!=X11|elem(X10,cons(X11,X12)))&(~elem(X10,X12)|elem(X10,cons(X11,X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
fof(c_0_33, plain, (![X5]:(X5=q_nil|X5=cons(head(X5),tail(X5)))), inference(variable_rename,[status(thm)],[c_0_7])).
cnf(c_0_34,plain,(tail(cons(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_35,negated_conjecture,(cons(m_Down(esk7_0),esk5_0)=queue(nbr_proc)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
cnf(c_0_36,plain,(q_nil!=cons(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37,negated_conjecture,(leq(index(pendack,host(esk10_0)),host(X1))|~epred1_3(X2,X3,X4)|~setIn(host(esk10_0),index(down,host(X1)))|~setIn(X1,alive)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_38,negated_conjecture,(setIn(esk6_0,alive)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_39,negated_conjecture,(~leq(index(pendack,host(esk10_0)),nbr_proc)), inference(rw,[status(thm)],[c_0_29, c_0_24])).
cnf(c_0_40,negated_conjecture,(host(esk7_0)=X1|leq(nbr_proc,X1)|setIn(X1,index(down,nbr_proc))|~leq(s(zero),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_24]), c_0_24])).
cnf(c_0_41,plain,(leq(s(zero),host(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42,plain,(elem(X1,cons(X2,X3))|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43,plain,(X1=cons(head(X1),tail(X1))|X1=q_nil), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44,negated_conjecture,(tail(queue(nbr_proc))=esk5_0), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_45,negated_conjecture,(queue(nbr_proc)!=q_nil), inference(spm,[status(thm)],[c_0_36, c_0_35])).
fof(c_0_46, plain, (![X6]:![X7]:head(cons(X6,X7))=X6), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_47,negated_conjecture,(~epred1_3(X1,X2,X3)|~setIn(host(esk10_0),index(down,nbr_proc))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_24]), c_0_24]), c_0_39])).
cnf(c_0_48,negated_conjecture,(host(esk7_0)=host(X1)|leq(nbr_proc,host(X1))|setIn(host(X1),index(down,nbr_proc))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_49,plain,(elem(X1,cons(X1,X2))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_50,negated_conjecture,(cons(head(queue(nbr_proc)),esk5_0)=queue(nbr_proc)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
cnf(c_0_51,plain,(head(cons(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52,negated_conjecture,(host(esk10_0)=host(esk7_0)|leq(nbr_proc,host(esk10_0))|~epred1_3(X1,X2,X3)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53,negated_conjecture,(host(esk6_0)!=host(esk10_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54,plain,(leq(index(pendack,host(X4)),host(X5))|~epred1_3(X1,X2,X3)|index(status,host(X4))!=elec_2|host(X4)!=host(X6)|~elem(m_Down(X6),queue(host(X5)))|~setIn(X4,alive)|~setIn(X5,alive)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_55,negated_conjecture,(elem(head(queue(nbr_proc)),queue(nbr_proc))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_56,negated_conjecture,(head(queue(nbr_proc))=m_Down(esk7_0)), inference(spm,[status(thm)],[c_0_51, c_0_35])).
cnf(c_0_57,negated_conjecture,(host(esk10_0)=host(esk7_0)|leq(nbr_proc,host(esk10_0))), inference(spm,[status(thm)],[c_0_52, c_0_23])).
cnf(c_0_58,negated_conjecture,(host(esk10_0)!=nbr_proc), inference(rw,[status(thm)],[c_0_53, c_0_24])).
cnf(c_0_59,negated_conjecture,(leq(index(pendack,host(X1)),nbr_proc)|index(status,host(X1))!=elec_2|host(X1)!=host(X2)|~epred1_3(X3,X4,X5)|~setIn(X1,alive)|~elem(m_Down(X2),queue(nbr_proc))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_24]), c_0_38])])).
cnf(c_0_60,negated_conjecture,(elem(m_Down(esk7_0),queue(nbr_proc))), inference(rw,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_61,negated_conjecture,(host(esk10_0)=host(esk7_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_57]), c_0_18])]), c_0_58])).
cnf(c_0_62,negated_conjecture,(leq(index(pendack,host(X1)),nbr_proc)|index(status,host(X1))!=elec_2|host(X1)!=host(esk7_0)|~epred1_3(X2,X3,X4)|~setIn(X1,alive)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_63,negated_conjecture,(index(status,host(esk7_0))=elec_2), inference(rw,[status(thm)],[c_0_27, c_0_61])).
cnf(c_0_64,negated_conjecture,(~leq(index(pendack,host(esk7_0)),nbr_proc)), inference(rw,[status(thm)],[c_0_39, c_0_61])).
cnf(c_0_65,negated_conjecture,(~epred1_3(X1,X2,X3)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_61]), c_0_63]), c_0_28])]), c_0_64])).
cnf(c_0_66,negated_conjecture,($false), inference(sr,[status(thm)],[c_0_23, c_0_65]), ['proof']).
# SZS output end CNFRefutation.
EOF
