% Problem    : SWV467+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n116.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:26:56 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2f
# and selection function SelectCQIAr.
#
# Presaturation interreduction done
# No success with AutoSched1
# Trying AutoSched2 for 23 seconds
# AutoSched2-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2g
# and selection function SelectCQArNpEqFirst.
#
# Presaturation interreduction done
# No success with AutoSched2
# Trying AutoSched3 for 17 seconds
# AutoSched3-Mode selected heuristic G_E___200_C45_F1_SE_CS_SP_PI_CO_S0V
# and selection function PSelectComplexExceptRRHorn.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X12]:![X13]:![X5]:![X6]:((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~(leq(host(X7),host(X14))))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:(elem(m_Halt(X14),queue(host(X7)))=>~(leq(host(X7),host(X14)))))&![X7]:![X15]:![X14]:(elem(m_Ack(X14,X7),queue(host(X15)))=>~(leq(host(X7),host(X14)))))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~(setIn(X7,alive))|~(setIn(X14,alive)))))&![X7]:![X14]:((setIn(X14,alive)&elem(m_Ack(X14,X7),queue(host(X14))))=>leq(host(X7),index(pendack,host(X14)))))&![X7]:![X14]:((setIn(X14,alive)&index(status,host(X14))=elec_1)=>~(elem(m_Ack(X14,X7),queue(host(X14))))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X15]:![X14]:(((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X14]:(((((~(leq(host(X7),host(X14)))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~(elem(m_Ack(X7,X15),queue(host(X7))))))&![X7]:![X14]:(((((~(leq(host(X7),host(X14)))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~(leq(index(pendack,host(X7)),index(pendack,host(X14))))))&![X7]:![X15]:![X14]:((((~(leq(index(pendack,host(X14)),host(X7)))&setIn(X14,alive))&elem(m_Halt(X14),queue(host(X15))))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X16]:![X15]:![X14]:(((((![X17]:((~(leq(host(X14),X17))&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&elem(m_Down(X15),queue(host(X14))))&host(X14)=nbr_proc)&host(X14)=host(X16))&index(status,host(X14))=elec_1)=>~((setIn(X7,alive)&elem(m_Down(X16),queue(host(X7)))))))&![X7]:![X16]:![X15]:![X14]:(((((((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&elem(m_Ack(X14,X16),queue(host(X14))))&leq(nbr_proc,s(index(pendack,host(X14)))))&index(status,host(X14))=elec_2)&host(X16)=index(pendack,host(X14)))&host(X15)=s(index(pendack,host(X14))))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12))=>(setIn(X5,alive)=>(~(leq(host(X5),host(X6)))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>((![X7]:((~(leq(host(X5),X7))&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1)=>(leq(nbr_proc,host(X5))=>![X7]:(~(setIn(host(X7),setEmpty))=>![X17]:(host(X5)=host(X17)=>![X18]:![X19]:![X20]:(host(X5)!=host(X20)=>(((((((setIn(X20,alive)&elem(m_Down(X19),queue(host(X20))))&elem(m_Ack(X20,X18),queue(host(X20))))&leq(nbr_proc,s(index(pendack,host(X20)))))&index(status,host(X20))=elec_2)&host(X18)=index(pendack,host(X20)))&host(X19)=s(index(pendack,host(X20))))=>~((setIn(X17,alive)&host(X5)=host(X17)))))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj)).
fof(c_0_1, axiom, (![X5]:![X6]:((leq(X5,X6)&leq(X6,X5))<=>X5=X6)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_61)).
fof(c_0_2, axiom, (![X3]:leq(host(X3),nbr_proc)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_04)).
fof(c_0_3, axiom, (![X5]:![X6]:![X4]:(elem(X5,cons(X6,X4))<=>(X5=X6|elem(X5,X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_46)).
fof(c_0_4, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)<=>(((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:(elem(m_Halt(X14),queue(host(X7)))=>~leq(host(X7),host(X14))))&![X7]:![X15]:![X14]:(elem(m_Ack(X14,X7),queue(host(X15)))=>~leq(host(X7),host(X14))))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:![X14]:((setIn(X14,alive)&elem(m_Ack(X14,X7),queue(host(X14))))=>leq(host(X7),index(pendack,host(X14)))))&![X7]:![X14]:((setIn(X14,alive)&index(status,host(X14))=elec_1)=>~elem(m_Ack(X14,X7),queue(host(X14)))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X15]:![X14]:(((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X15),queue(host(X7)))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~leq(index(pendack,host(X7)),index(pendack,host(X14)))))&![X7]:![X15]:![X14]:((((~leq(index(pendack,host(X14)),host(X7))&setIn(X14,alive))&elem(m_Halt(X14),queue(host(X15))))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X16]:![X15]:![X14]:(((((![X17]:((~leq(host(X14),X17)&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&elem(m_Down(X15),queue(host(X14))))&host(X14)=nbr_proc)&host(X14)=host(X16))&index(status,host(X14))=elec_1)=>~((setIn(X7,alive)&elem(m_Down(X16),queue(host(X7)))))))&![X7]:![X16]:![X15]:![X14]:(((((((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&elem(m_Ack(X14,X16),queue(host(X14))))&leq(nbr_proc,s(index(pendack,host(X14)))))&index(status,host(X14))=elec_2)&host(X16)=index(pendack,host(X14)))&host(X15)=s(index(pendack,host(X14))))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12)))), introduced(definition)).
fof(c_0_5, negated_conjecture, (~(![X12]:![X5]:![X6]:(epred1_3(X5,X6,X12)=>(setIn(X5,alive)=>(~leq(host(X5),host(X6))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>((![X7]:((~leq(host(X5),X7)&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1)=>(leq(nbr_proc,host(X5))=>![X7]:(~setIn(host(X7),setEmpty)=>![X17]:(host(X5)=host(X17)=>![X18]:![X19]:![X20]:(host(X5)!=host(X20)=>(((((((setIn(X20,alive)&elem(m_Down(X19),queue(host(X20))))&elem(m_Ack(X20,X18),queue(host(X20))))&leq(nbr_proc,s(index(pendack,host(X20)))))&index(status,host(X20))=elec_2)&host(X18)=index(pendack,host(X20)))&host(X19)=s(index(pendack,host(X20))))=>~((setIn(X17,alive)&host(X5)=host(X17))))))))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_4])).
fof(c_0_6, negated_conjecture, (![X24]:(epred1_3(esk6_0,esk7_0,esk5_0)&(setIn(esk6_0,alive)&(~leq(host(esk6_0),host(esk7_0))&(((index(ldr,host(esk6_0))!=host(esk7_0)|index(status,host(esk6_0))!=norm)&(index(status,host(esk6_0))!=wait|host(esk7_0)!=host(index(elid,host(esk6_0)))))&((((leq(host(esk6_0),X24)|~leq(s(zero),X24))|(setIn(X24,index(down,host(esk6_0)))|X24=host(esk7_0)))&index(status,host(esk6_0))=elec_1)&(leq(nbr_proc,host(esk6_0))&(~setIn(host(esk8_0),setEmpty)&(host(esk6_0)=host(esk9_0)&(host(esk6_0)!=host(esk12_0)&(((((((setIn(esk12_0,alive)&elem(m_Down(esk11_0),queue(host(esk12_0))))&elem(m_Ack(esk12_0,esk10_0),queue(host(esk12_0))))&leq(nbr_proc,s(index(pendack,host(esk12_0)))))&index(status,host(esk12_0))=elec_2)&host(esk10_0)=index(pendack,host(esk12_0)))&host(esk11_0)=s(index(pendack,host(esk12_0))))&(setIn(esk9_0,alive)&host(esk6_0)=host(esk9_0))))))))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
fof(c_0_7, plain, (![X7]:![X8]:![X9]:![X10]:(((~leq(X7,X8)|~leq(X8,X7))|X7=X8)&((leq(X9,X10)|X9!=X10)&(leq(X10,X9)|X9!=X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
fof(c_0_8, plain, (![X4]:leq(host(X4),nbr_proc)), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_9,negated_conjecture,(host(esk11_0)=s(index(pendack,host(esk12_0)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,negated_conjecture,(host(esk10_0)=index(pendack,host(esk12_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,plain,(X1=X2|~leq(X2,X1)|~leq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12,plain,(leq(host(X1),nbr_proc)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_13, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)=>(((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:(elem(m_Halt(X14),queue(host(X7)))=>~leq(host(X7),host(X14))))&![X7]:![X15]:![X14]:(elem(m_Ack(X14,X7),queue(host(X15)))=>~leq(host(X7),host(X14))))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:![X14]:((setIn(X14,alive)&elem(m_Ack(X14,X7),queue(host(X14))))=>leq(host(X7),index(pendack,host(X14)))))&![X7]:![X14]:((setIn(X14,alive)&index(status,host(X14))=elec_1)=>~elem(m_Ack(X14,X7),queue(host(X14)))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X15]:![X14]:(((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X15),queue(host(X7)))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~leq(index(pendack,host(X7)),index(pendack,host(X14)))))&![X7]:![X15]:![X14]:((((~leq(index(pendack,host(X14)),host(X7))&setIn(X14,alive))&elem(m_Halt(X14),queue(host(X15))))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X16]:![X15]:![X14]:(((((![X17]:((~leq(host(X14),X17)&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&elem(m_Down(X15),queue(host(X14))))&host(X14)=nbr_proc)&host(X14)=host(X16))&index(status,host(X14))=elec_1)=>~((setIn(X7,alive)&elem(m_Down(X16),queue(host(X7)))))))&![X7]:![X16]:![X15]:![X14]:(((((((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&elem(m_Ack(X14,X16),queue(host(X14))))&leq(nbr_proc,s(index(pendack,host(X14)))))&index(status,host(X14))=elec_2)&host(X16)=index(pendack,host(X14)))&host(X15)=s(index(pendack,host(X14))))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12)))), inference(split_equiv,[status(thm)],[c_0_4])).
cnf(c_0_14,negated_conjecture,(leq(nbr_proc,s(index(pendack,host(esk12_0))))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15,negated_conjecture,(s(host(esk10_0))=host(esk11_0)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_16,plain,(host(X1)=nbr_proc|~leq(nbr_proc,host(X1))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17,negated_conjecture,(host(esk6_0)=host(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_18,negated_conjecture,(leq(nbr_proc,host(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_19, plain, (![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:![X34]:![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:![X43]:![X44]:![X45]:![X46]:![X47]:![X48]:![X49]:![X50]:![X51]:![X52]:![X53]:![X55]:![X56]:![X57]:![X58]:(((((((((((((((((~elem(m_Ldr(X22),queue(host(X21)))|~leq(host(X21),host(X22)))|~epred1_3(X20,X19,X18))&((~elem(m_Down(X24),queue(host(X23)))|host(X24)!=host(X23))|~epred1_3(X20,X19,X18)))&((~elem(m_Halt(X26),queue(host(X25)))|~leq(host(X25),host(X26)))|~epred1_3(X20,X19,X18)))&((~elem(m_Ack(X29,X27),queue(host(X28)))|~leq(host(X27),host(X29)))|~epred1_3(X20,X19,X18)))&(((X31=X30|host(X31)!=host(X30))|(~setIn(X30,alive)|~setIn(X31,alive)))|~epred1_3(X20,X19,X18)))&(((~setIn(X33,alive)|~elem(m_Ack(X33,X32),queue(host(X33))))|leq(host(X32),index(pendack,host(X33))))|~epred1_3(X20,X19,X18)))&(((~setIn(X35,alive)|index(status,host(X35))!=elec_1)|~elem(m_Ack(X35,X34),queue(host(X35))))|~epred1_3(X20,X19,X18)))&((((index(status,host(X36))!=elec_1|~setIn(X36,alive))|index(elid,host(X36))=X36)|~epred1_3(X20,X19,X18))&(((index(status,host(X36))!=elec_2|~setIn(X36,alive))|index(elid,host(X36))=X36)|~epred1_3(X20,X19,X18))))&((((~setIn(X39,alive)|~elem(m_Down(X38),queue(host(X39))))|host(X38)!=host(X37))|((~setIn(X37,alive)|index(ldr,host(X37))!=host(X37))|index(status,host(X37))!=norm))|~epred1_3(X20,X19,X18)))&((((((leq(host(X40),host(X41))|~setIn(X40,alive))|~setIn(X41,alive))|index(status,host(X40))!=elec_2)|index(status,host(X41))!=elec_2)|leq(index(pendack,host(X41)),host(X40)))|~epred1_3(X20,X19,X18)))&((((((~setIn(X42,alive)|~setIn(X44,alive))|host(X44)!=host(X43))|index(status,host(X42))!=elec_2)|index(status,host(X44))!=elec_2)|~elem(m_Ack(X42,X43),queue(host(X42))))|~epred1_3(X20,X19,X18)))&((((((leq(host(X45),host(X46))|~setIn(X45,alive))|~setIn(X46,alive))|index(status,host(X45))!=elec_2)|index(status,host(X46))!=elec_2)|~leq(index(pendack,host(X45)),index(pendack,host(X46))))|~epred1_3(X20,X19,X18)))&(((((leq(index(pendack,host(X49)),host(X47))|~setIn(X49,alive))|~elem(m_Halt(X49),queue(host(X48))))|index(status,host(X49))!=elec_2)|((~setIn(X47,alive)|index(ldr,host(X47))!=host(X47))|index(status,host(X47))!=norm))|~epred1_3(X20,X19,X18)))&((((((((~leq(host(X53),esk13_7(X18,X19,X20,X50,X51,X52,X53))|~elem(m_Down(X52),queue(host(X53))))|host(X53)!=nbr_proc)|host(X53)!=host(X51))|index(status,host(X53))!=elec_1)|(~setIn(X50,alive)|~elem(m_Down(X51),queue(host(X50)))))|~epred1_3(X20,X19,X18))&((((((leq(s(zero),esk13_7(X18,X19,X20,X50,X51,X52,X53))|~elem(m_Down(X52),queue(host(X53))))|host(X53)!=nbr_proc)|host(X53)!=host(X51))|index(status,host(X53))!=elec_1)|(~setIn(X50,alive)|~elem(m_Down(X51),queue(host(X50)))))|~epred1_3(X20,X19,X18)))&(((((((~setIn(esk13_7(X18,X19,X20,X50,X51,X52,X53),index(down,host(X53)))|~elem(m_Down(X52),queue(host(X53))))|host(X53)!=nbr_proc)|host(X53)!=host(X51))|index(status,host(X53))!=elec_1)|(~setIn(X50,alive)|~elem(m_Down(X51),queue(host(X50)))))|~epred1_3(X20,X19,X18))&((((((esk13_7(X18,X19,X20,X50,X51,X52,X53)!=host(X52)|~elem(m_Down(X52),queue(host(X53))))|host(X53)!=nbr_proc)|host(X53)!=host(X51))|index(status,host(X53))!=elec_1)|(~setIn(X50,alive)|~elem(m_Down(X51),queue(host(X50)))))|~epred1_3(X20,X19,X18)))))&((((((((~setIn(X57,alive)|~elem(m_Down(X56),queue(host(X57))))|~elem(m_Ack(X57,X55),queue(host(X57))))|~leq(nbr_proc,s(index(pendack,host(X57)))))|index(status,host(X57))!=elec_2)|host(X55)!=index(pendack,host(X57)))|host(X56)!=s(index(pendack,host(X57))))|((~setIn(X58,alive)|index(ldr,host(X58))!=host(X58))|index(status,host(X58))!=norm))|~epred1_3(X20,X19,X18)))&(queue(host(X20))=cons(m_Down(X19),X18)|~epred1_3(X20,X19,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])])).
cnf(c_0_20,negated_conjecture,(leq(nbr_proc,host(esk11_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_10]), c_0_15])).
cnf(c_0_21,negated_conjecture,(index(status,host(esk6_0))=elec_1), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22,negated_conjecture,(host(esk6_0)=nbr_proc), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])])).
cnf(c_0_23,negated_conjecture,(X1=host(esk7_0)|setIn(X1,index(down,host(esk6_0)))|leq(host(esk6_0),X1)|~leq(s(zero),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_24, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((~elem(X7,cons(X8,X9))|(X7=X8|elem(X7,X9)))&((X10!=X11|elem(X10,cons(X11,X12)))&(~elem(X10,X12)|elem(X10,cons(X11,X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_25,plain,(~epred1_3(X1,X2,X3)|~elem(m_Down(X4),queue(host(X5)))|~setIn(X5,alive)|index(status,host(X6))!=elec_1|host(X6)!=host(X4)|host(X6)!=nbr_proc|~elem(m_Down(X7),queue(host(X6)))|esk13_7(X3,X2,X1,X5,X4,X7,X6)!=host(X7)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26,negated_conjecture,(host(esk11_0)=nbr_proc), inference(spm,[status(thm)],[c_0_16, c_0_20])).
cnf(c_0_27,negated_conjecture,(index(status,nbr_proc)=elec_1), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28,negated_conjecture,(host(esk7_0)=X1|leq(nbr_proc,X1)|setIn(X1,index(down,nbr_proc))|~leq(s(zero),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_22]), c_0_22])).
cnf(c_0_29,plain,(leq(s(zero),esk13_7(X3,X2,X1,X5,X4,X7,X6))|~epred1_3(X1,X2,X3)|~elem(m_Down(X4),queue(host(X5)))|~setIn(X5,alive)|index(status,host(X6))!=elec_1|host(X6)!=host(X4)|host(X6)!=nbr_proc|~elem(m_Down(X7),queue(host(X6)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30,plain,(~epred1_3(X1,X2,X3)|~elem(m_Down(X4),queue(host(X5)))|~setIn(X5,alive)|index(status,host(X6))!=elec_1|host(X6)!=host(X4)|host(X6)!=nbr_proc|~elem(m_Down(X7),queue(host(X6)))|~setIn(esk13_7(X3,X2,X1,X5,X4,X7,X6),index(down,host(X6)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31,plain,(~epred1_3(X1,X2,X3)|~elem(m_Down(X4),queue(host(X5)))|~setIn(X5,alive)|index(status,host(X6))!=elec_1|host(X6)!=host(X4)|host(X6)!=nbr_proc|~elem(m_Down(X7),queue(host(X6)))|~leq(host(X6),esk13_7(X3,X2,X1,X5,X4,X7,X6))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32,plain,(elem(X1,cons(X2,X3))|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33,plain,(esk13_7(X1,X2,X3,X4,X5,X6,esk11_0)!=host(X6)|host(X5)!=nbr_proc|~epred1_3(X3,X2,X1)|~setIn(X4,alive)|~elem(m_Down(X5),queue(host(X4)))|~elem(m_Down(X6),queue(nbr_proc))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_34,negated_conjecture,(elem(m_Down(esk11_0),queue(host(esk12_0)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_35,negated_conjecture,(setIn(esk12_0,alive)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_36,plain,(esk13_7(X1,X2,X3,X4,X5,X6,X7)=host(esk7_0)|leq(nbr_proc,esk13_7(X1,X2,X3,X4,X5,X6,X7))|setIn(esk13_7(X1,X2,X3,X4,X5,X6,X7),index(down,nbr_proc))|index(status,host(X7))!=elec_1|host(X5)!=host(X7)|host(X7)!=nbr_proc|~epred1_3(X3,X2,X1)|~setIn(X4,alive)|~elem(m_Down(X6),queue(host(X7)))|~elem(m_Down(X5),queue(host(X4)))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_37,plain,(host(X1)!=nbr_proc|~epred1_3(X2,X3,X4)|~setIn(esk13_7(X4,X3,X2,X5,X1,X6,esk11_0),index(down,nbr_proc))|~setIn(X5,alive)|~elem(m_Down(X1),queue(host(X5)))|~elem(m_Down(X6),queue(nbr_proc))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_26]), c_0_27])])).
cnf(c_0_38,plain,(host(X1)!=nbr_proc|~epred1_3(X2,X3,X4)|~leq(nbr_proc,esk13_7(X4,X3,X2,X5,X1,X6,esk11_0))|~setIn(X5,alive)|~elem(m_Down(X1),queue(host(X5)))|~elem(m_Down(X6),queue(nbr_proc))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_26]), c_0_27])])).
cnf(c_0_39,plain,(elem(X1,cons(X1,X2))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_40,plain,(queue(host(X1))=cons(m_Down(X2),X3)|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_41,negated_conjecture,(esk13_7(X1,X2,X3,esk12_0,esk11_0,X4,esk11_0)!=host(X4)|~epred1_3(X3,X2,X1)|~elem(m_Down(X4),queue(nbr_proc))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_26]), c_0_35])])).
cnf(c_0_42,negated_conjecture,(esk13_7(X1,X2,X3,X4,X5,X6,esk11_0)=host(esk7_0)|host(X5)!=nbr_proc|~epred1_3(X3,X2,X1)|~setIn(X4,alive)|~elem(m_Down(X5),queue(host(X4)))|~elem(m_Down(X6),queue(nbr_proc))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_26]), c_0_27])]), c_0_37]), c_0_38])).
cnf(c_0_43,plain,(elem(m_Down(X1),queue(host(X2)))|~epred1_3(X2,X1,X3)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44,negated_conjecture,(epred1_3(esk6_0,esk7_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_45,negated_conjecture,(host(esk7_0)!=host(X1)|~epred1_3(X2,X3,X4)|~elem(m_Down(X1),queue(nbr_proc))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_26]), c_0_35]), c_0_34])])).
cnf(c_0_46,negated_conjecture,(elem(m_Down(esk7_0),queue(nbr_proc))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_22])).
cnf(c_0_47,negated_conjecture,(~epred1_3(X1,X2,X3)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_48,negated_conjecture,($false), inference(sr,[status(thm)],[c_0_44, c_0_47]), ['proof']).
# SZS output end CNFRefutation.
EOF
