% Problem    : CSR038+5 : TPTP v6.1.0. Bugfixed v3.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n117.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:52:46 CDT 2014
% CPUTime    : 
# SinE strategy is gf500_h_gu_R04_F100_L20000
# Trying AutoSched0 for 146 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0YP
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (?[X23]:(mtvisible(c_knowledgefragmentd3mt)=>(tptp_8_968(X23,c_tptpnsubcollectionofwithrelationtotypefnissuingaprescriptionproductscorrectivelensprescription_2804)&tptpcol_16_7738(X23)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', query238)).
fof(c_0_1, axiom, (![X23]:(isa(X23,c_tptpcol_16_7738)=>tptpcol_16_7738(X23))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+4.ax', ax4_470078)).
fof(c_0_2, axiom, (![X6]:(isa(X6,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription))=>tptp_8_968(f_relationexistsallfn(X6,c_tptp_8_968,c_tptpcol_16_7738,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription)),X6))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+4.ax', ax4_59513)).
fof(c_0_3, axiom, (isa(c_tptpnsubcollectionofwithrelationtotypefnissuingaprescriptionproductscorrectivelensprescription_2804,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+4.ax', ax4_19965)).
fof(c_0_4, axiom, (![X6]:![X9]:![X10]:![X11]:((isa(X6,X9)&relationexistsall(X10,X11,X9))=>isa(f_relationexistsallfn(X6,X10,X11,X9),X11))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+4.ax', ax4_46691)).
fof(c_0_5, axiom, (relationexistsall(c_tptp_8_968,c_tptpcol_16_7738,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription))), file('/export/starexec/sandbox/benchmark/Axioms/CSR002+4.ax', ax4_59514)).
fof(c_0_6, negated_conjecture, (~(?[X23]:(mtvisible(c_knowledgefragmentd3mt)=>(tptp_8_968(X23,c_tptpnsubcollectionofwithrelationtotypefnissuingaprescriptionproductscorrectivelensprescription_2804)&tptpcol_16_7738(X23))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_7, negated_conjecture, (![X24]:(mtvisible(c_knowledgefragmentd3mt)&(~tptp_8_968(X24,c_tptpnsubcollectionofwithrelationtotypefnissuingaprescriptionproductscorrectivelensprescription_2804)|~tptpcol_16_7738(X24)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
fof(c_0_8, plain, (![X24]:(~isa(X24,c_tptpcol_16_7738)|tptpcol_16_7738(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
cnf(c_0_9,negated_conjecture,(~tptpcol_16_7738(X1)|~tptp_8_968(X1,c_tptpnsubcollectionofwithrelationtotypefnissuingaprescriptionproductscorrectivelensprescription_2804)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,plain,(tptpcol_16_7738(X1)|~isa(X1,c_tptpcol_16_7738)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, plain, (![X7]:(~isa(X7,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription))|tptp_8_968(f_relationexistsallfn(X7,c_tptp_8_968,c_tptpcol_16_7738,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription)),X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
cnf(c_0_12,negated_conjecture,(~tptp_8_968(X1,c_tptpnsubcollectionofwithrelationtotypefnissuingaprescriptionproductscorrectivelensprescription_2804)|~isa(X1,c_tptpcol_16_7738)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13,plain,(tptp_8_968(f_relationexistsallfn(X1,c_tptp_8_968,c_tptpcol_16_7738,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription)),X1)|~isa(X1,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14,plain,(isa(c_tptpnsubcollectionofwithrelationtotypefnissuingaprescriptionproductscorrectivelensprescription_2804,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription))), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_15, plain, (![X12]:![X13]:![X14]:![X15]:((~isa(X12,X13)|~relationexistsall(X14,X15,X13))|isa(f_relationexistsallfn(X12,X14,X15,X13),X15))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
cnf(c_0_16,negated_conjecture,(~isa(f_relationexistsallfn(c_tptpnsubcollectionofwithrelationtotypefnissuingaprescriptionproductscorrectivelensprescription_2804,c_tptp_8_968,c_tptpcol_16_7738,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription)),c_tptpcol_16_7738)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
cnf(c_0_17,plain,(isa(f_relationexistsallfn(X1,X2,X3,X4),X3)|~relationexistsall(X2,X3,X4)|~isa(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18,plain,(relationexistsall(c_tptp_8_968,c_tptpcol_16_7738,f_subcollectionofwithrelationtotypefn(c_issuingaprescription,c_products,c_correctivelensprescription))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_14])]), ['proof']).
# SZS output end CNFRefutation.
EOF
