% Problem    : GEO319+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n089.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:17:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X185]:![X186]:((X185=X186&rreal(X186))=>vplus(X185,v0)=X185)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(141), 0), imp(cond(axiom(141), 0)))')).
fof(c_0_1, axiom, (![X181]:![X182]:![X183]:![X184]:((((X183=X184&rreal(X184))&X181=X182)&rreal(X182))=>vplus(X181,X183)=vplus(X183,X181))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(143), 0), imp(cond(axiom(143), 0)))')).
fof(c_0_2, axiom, (![X171]:![X172]:![X173]:![X174]:![X175]:![X176]:(((((((vplus(X171,X175)=vplus(X173,X175)&X175=X176)&rreal(X176))&X173=X174)&rreal(X174))&X171=X172)&rreal(X172))=>X171=X173)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(156, 0), 0)')).
fof(c_0_3, axiom, (rreal(vd684)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(axiom(137), 1)')).
fof(c_0_4, axiom, (v0=vd684), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(axiom(137), 2)')).
fof(c_0_5, axiom, (![X187]:![X188]:![X189]:![X190]:![X191]:(((((((((((X190=X191&ron(X189,X190))&ron(X188,X190))&ron(X187,X190))&rline(X191))&X188!=X189)&X187!=X189)&X187!=X188)&?[X192]:(X189=X192&rpoint(X192)))&?[X193]:(X188=X193&rpoint(X193)))&?[X194]:(X187=X194&rpoint(X194)))=>((rR(X188,X187,X189)|rR(X187,X188,X189))|rR(X189,X187,X188)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(91), 0), imp(cond(axiom(91), 0)))')).
fof(c_0_6, axiom, (((((((vf(vd1287,vd1288)=vf(vd1287,vd1289)&vd1288!=vd1289)&vd1287!=vd1289)&vd1287!=vd1288)&?[X7]:(vd1289=X7&rpoint(X7)))&?[X8]:(vd1288=X8&rpoint(X8)))&?[X9]:(vd1287=X9&rpoint(X9)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(holds(conjunct2(345), 1294, 0), and(pred(conjunct1(345), 9), and(pred(conjunct1(345), 8), and(pred(conjunct1(345), 7), and(qe(s3(plural(345))), and(qe(s2(plural(345))), qe(s1(plural(345)))))))))')).
fof(c_0_7, axiom, (![X10]:![X11]:![X12]:![X13]:![X14]:((((((((X13=X14&ron(X12,X13))&ron(X10,X13))&rR(X11,X10,X12))&rline(X14))&?[X15]:(X12=X15&rpoint(X15)))&?[X16]:(X11=X16&rpoint(X16)))&?[X17]:(X10=X17&rpoint(X17)))=>ron(X11,X13))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(85), 0), imp(cond(axiom(85), 0)))')).
fof(c_0_8, conjecture, (vd1329!=vd1302), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct1(conjunct2(conjunct2(368))), 1333, 0)')).
fof(c_0_9, axiom, (![X64]:![X65]:((?[X66]:(X65=X66&rpoint(X66))&?[X67]:(X64=X67&rpoint(X67)))=>?[X68]:(vf(X64,X65)=X68&rreal(X68)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(160), 0), imp(cond(axiom(160), 0)))')).
fof(c_0_10, axiom, (rpoint(vd1329)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(367, 0)')).
fof(c_0_11, axiom, (![X420]:![X421]:![X422]:![X423]:((((((rR(X423,X420,X421)&rR(X421,X420,X422))&?[X424]:(X423=X424&rpoint(X424)))&?[X425]:(X422=X425&rpoint(X425)))&?[X426]:(X421=X426&rpoint(X426)))&?[X427]:(X420=X427&rpoint(X427)))=>rR(X423,X420,X422))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(87), 0), imp(cond(axiom(87), 0)))')).
fof(c_0_12, axiom, (![X199]:![X200]:![X201]:((((rR(X200,X199,X201)&?[X202]:(X201=X202&rpoint(X202)))&?[X203]:(X200=X203&rpoint(X203)))&?[X204]:(X199=X204&rpoint(X204)))=>vplus(vf(X199,X200),vf(X200,X201))=vf(X199,X201))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(180), 0), imp(cond(axiom(180), 0)))')).
fof(c_0_13, axiom, (ron(vd1288,vd1295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(s2(plural(346)), 0)')).
fof(c_0_14, axiom, (rline(vd1295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(346, 0)')).
fof(c_0_15, axiom, ((((vd1302=vd1303&rR(vd1289,vd1287,vd1302))&rpoint(vd1303))&rR(vd1288,vd1287,vd1299))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(pred(conjunct2(348), 4), and(holds(conjunct2(348), 1304, 0), and(pred(conjunct2(348), 1), holds(conjunct1(348), 1301, 0))))')).
fof(c_0_16, axiom, (ron(vd1329,vd1295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(conjunct1(368), 0)')).
fof(c_0_17, axiom, (rR(vd1329,vd1287,vd1299)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conjunct1(367), 1330, 0)')).
fof(c_0_18, axiom, (rpoint(vd1300)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(348, 1)')).
fof(c_0_19, axiom, (vd1299=vd1300), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(348, 4)')).
fof(c_0_20, axiom, (ron(vd1287,vd1295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(s1(plural(346)), 0)')).
fof(c_0_21, axiom, (![X1]:![X2]:![X3]:((((rR(X2,X1,X3)&?[X4]:(X3=X4&rpoint(X4)))&?[X5]:(X2=X5&rpoint(X5)))&?[X6]:(X1=X6&rpoint(X6)))=>(((~(rR(X1,X2,X3))&X1!=X3)&X1!=X2)&rR(X2,X3,X1)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(81), 0), imp(cond(axiom(81), 0)))')).
fof(c_0_22, axiom, (![X77]:![X78]:![X79]:![X80]:((((((rR(X79,X78,X80)&rR(X78,X77,X79))&?[X81]:(X80=X81&rpoint(X81)))&?[X82]:(X79=X82&rpoint(X82)))&?[X83]:(X78=X83&rpoint(X83)))&?[X84]:(X77=X84&rpoint(X84)))=>rR(X78,X77,X80))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(89), 0), imp(cond(axiom(89), 0)))')).
fof(c_0_23, axiom, (![X60]:![X61]:((?[X62]:(X61=X62&rpoint(X62))&?[X63]:(X60=X63&rpoint(X63)))=>vf(X60,X61)=vf(X61,X60))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(166), 0), imp(cond(axiom(166), 0)))')).
fof(c_0_24, axiom, (![X109]:![X110]:((?[X111]:(X110=X111&rpoint(X111))&?[X112]:(X109=X112&rpoint(X112)))=>(vf(X109,X110)=v0<=>X109=X110))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(162), 0), imp(cond(axiom(162), 0)))')).
fof(c_0_25, plain, (![X187]:![X188]:((X187!=X188|~rreal(X188))|vplus(X187,v0)=X187)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])).
fof(c_0_26, plain, (![X185]:![X186]:![X187]:![X188]:((((X187!=X188|~rreal(X188))|X185!=X186)|~rreal(X186))|vplus(X185,X187)=vplus(X187,X185))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
fof(c_0_27, plain, (![X177]:![X178]:![X179]:![X180]:![X181]:![X182]:(((((((vplus(X177,X181)!=vplus(X179,X181)|X181!=X182)|~rreal(X182))|X179!=X180)|~rreal(X180))|X177!=X178)|~rreal(X178))|X177=X179)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_28,plain,(vplus(X1,v0)=X1|~rreal(X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29,plain,(vplus(X1,X2)=vplus(X2,X1)|~rreal(X3)|X1!=X3|~rreal(X4)|X2!=X4), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_30,plain,(rreal(vd684)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_31,plain,(v0=vd684), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_32, plain, (![X195]:![X196]:![X197]:![X198]:![X199]:![X200]:![X201]:![X202]:(((((((((((X198!=X199|~ron(X197,X198))|~ron(X196,X198))|~ron(X195,X198))|~rline(X199))|X196=X197)|X195=X197)|X195=X196)|(X197!=X200|~rpoint(X200)))|(X196!=X201|~rpoint(X201)))|(X195!=X202|~rpoint(X202)))|((rR(X196,X195,X197)|rR(X195,X196,X197))|rR(X197,X195,X196)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
fof(c_0_33, plain, (((((((vf(vd1287,vd1288)=vf(vd1287,vd1289)&vd1288!=vd1289)&vd1287!=vd1289)&vd1287!=vd1288)&(vd1289=esk1_0&rpoint(esk1_0)))&(vd1288=esk2_0&rpoint(esk2_0)))&(vd1287=esk3_0&rpoint(esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_6])])).
fof(c_0_34, plain, (![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:((((((((X21!=X22|~ron(X20,X21))|~ron(X18,X21))|~rR(X19,X18,X20))|~rline(X22))|(X20!=X23|~rpoint(X23)))|(X19!=X24|~rpoint(X24)))|(X18!=X25|~rpoint(X25)))|ron(X19,X21))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
fof(c_0_35, negated_conjecture, (vd1329=vd1302), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_8])])).
cnf(c_0_36,plain,(X1=X2|~rreal(X3)|X1!=X3|~rreal(X4)|X2!=X4|~rreal(X5)|X6!=X5|vplus(X1,X6)!=vplus(X2,X6)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37,plain,(vplus(X1,v0)=X1|~rreal(X1)), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_38,plain,(vplus(X1,X2)=vplus(X2,X1)|~rreal(X2)|~rreal(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_29])])).
cnf(c_0_39,plain,(rreal(v0)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_40, plain, (![X69]:![X70]:![X71]:![X72]:((vf(X69,X70)=esk8_2(X69,X70)|((X70!=X71|~rpoint(X71))|(X69!=X72|~rpoint(X72))))&(rreal(esk8_2(X69,X70))|((X70!=X71|~rpoint(X71))|(X69!=X72|~rpoint(X72)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
cnf(c_0_41,plain,(rR(X1,X2,X3)|rR(X2,X3,X1)|rR(X3,X2,X1)|X2=X3|X2=X1|X3=X1|~rpoint(X4)|X2!=X4|~rpoint(X5)|X3!=X5|~rpoint(X6)|X1!=X6|~rline(X7)|~ron(X2,X8)|~ron(X3,X8)|~ron(X1,X8)|X8!=X7), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42,plain,(rpoint(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43,plain,(vd1288=esk2_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44,plain,(ron(X1,X2)|~rpoint(X3)|X4!=X3|~rpoint(X5)|X1!=X5|~rpoint(X6)|X7!=X6|~rline(X8)|~rR(X1,X4,X7)|~ron(X4,X2)|~ron(X7,X2)|X2!=X8), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45,plain,(rpoint(vd1329)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_46,negated_conjecture,(vd1329=vd1302), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47,plain,(rpoint(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_48,plain,(vd1289=esk1_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_49,plain,(rpoint(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_50,plain,(vd1287=esk3_0), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_51, plain, (![X428]:![X429]:![X430]:![X431]:![X432]:![X433]:![X434]:![X435]:((((((~rR(X431,X428,X429)|~rR(X429,X428,X430))|(X431!=X432|~rpoint(X432)))|(X430!=X433|~rpoint(X433)))|(X429!=X434|~rpoint(X434)))|(X428!=X435|~rpoint(X435)))|rR(X431,X428,X430))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_52,plain,(X1=X2|vplus(X1,X3)!=vplus(X2,X3)|~rreal(X3)|~rreal(X2)|~rreal(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_36])])])).
cnf(c_0_53,plain,(vplus(v0,X1)=X1|~rreal(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
fof(c_0_54, plain, (![X205]:![X206]:![X207]:![X208]:![X209]:![X210]:((((~rR(X206,X205,X207)|(X207!=X208|~rpoint(X208)))|(X206!=X209|~rpoint(X209)))|(X205!=X210|~rpoint(X210)))|vplus(vf(X205,X206),vf(X206,X207))=vf(X205,X207))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_55,plain,(rreal(esk8_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_56,plain,(vf(X2,X4)=esk8_2(X2,X4)|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_57,plain,(X1=X2|X3=X2|X3=X1|rR(X2,X1,X3)|rR(X1,X2,X3)|rR(X3,X1,X2)|~rline(X4)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~ron(X2,X4)|~ron(X1,X4)|~ron(X3,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_41])])])])).
cnf(c_0_58,plain,(ron(vd1288,vd1295)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_59,plain,(rline(vd1295)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_60,plain,(rpoint(vd1288)), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_61,plain,(ron(X1,X2)|~rline(X2)|~rpoint(X3)|~rpoint(X1)|~rpoint(X4)|~rR(X1,X4,X3)|~ron(X3,X2)|~ron(X4,X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_44])])])])).
cnf(c_0_62,plain,(rR(vd1289,vd1287,vd1302)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_63,plain,(rpoint(vd1302)), inference(rw,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_64,plain,(rpoint(vd1289)), inference(rw,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_65,plain,(rpoint(vd1287)), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_66,plain,(ron(vd1329,vd1295)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_67,plain,(rR(X1,X2,X3)|~rpoint(X4)|X2!=X4|~rpoint(X5)|X6!=X5|~rpoint(X7)|X3!=X7|~rpoint(X8)|X1!=X8|~rR(X6,X2,X3)|~rR(X1,X2,X6)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_68,plain,(rR(vd1329,vd1287,vd1299)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_69,plain,(rpoint(vd1300)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_70,plain,(vd1299=vd1300), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_71,plain,(v0=X1|vplus(X1,X2)!=X2|~rreal(X2)|~rreal(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_39])])).
cnf(c_0_72,plain,(vplus(vf(X1,X2),vf(X2,X3))=vf(X1,X3)|~rpoint(X4)|X1!=X4|~rpoint(X5)|X2!=X5|~rpoint(X6)|X3!=X6|~rR(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_73,plain,(rreal(esk8_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_55])])).
cnf(c_0_74,plain,(esk8_2(X1,X2)=vf(X1,X2)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_56])])).
cnf(c_0_75,plain,(X1=vd1288|X2=vd1288|X2=X1|rR(vd1288,X1,X2)|rR(X1,vd1288,X2)|rR(X2,X1,vd1288)|~rpoint(X2)|~rpoint(X1)|~ron(X1,vd1295)|~ron(X2,vd1295)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_60])])).
cnf(c_0_76,plain,(ron(vd1287,vd1295)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_77,plain,(vd1287!=vd1288), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_78,plain,(ron(vd1289,X1)|~rline(X1)|~ron(vd1302,X1)|~ron(vd1287,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_64]), c_0_65])])).
cnf(c_0_79,plain,(ron(vd1302,vd1295)), inference(rw,[status(thm)],[c_0_66, c_0_46])).
fof(c_0_80, plain, (![X1]:![X2]:![X3]:((((rR(X2,X1,X3)&?[X4]:(X3=X4&rpoint(X4)))&?[X5]:(X2=X5&rpoint(X5)))&?[X6]:(X1=X6&rpoint(X6)))=>(((~rR(X1,X2,X3)&X1!=X3)&X1!=X2)&rR(X2,X3,X1)))), inference(fof_simplification,[status(thm)],[c_0_21])).
fof(c_0_81, plain, (![X85]:![X86]:![X87]:![X88]:![X89]:![X90]:![X91]:![X92]:((((((~rR(X87,X86,X88)|~rR(X86,X85,X87))|(X88!=X89|~rpoint(X89)))|(X87!=X90|~rpoint(X90)))|(X86!=X91|~rpoint(X91)))|(X85!=X92|~rpoint(X92)))|rR(X86,X85,X88))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
cnf(c_0_82,plain,(rR(X1,X2,X3)|~rpoint(X1)|~rpoint(X3)|~rpoint(X4)|~rpoint(X2)|~rR(X4,X2,X3)|~rR(X1,X2,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_67])])])])).
cnf(c_0_83,plain,(rR(vd1302,vd1287,vd1299)), inference(rw,[status(thm)],[c_0_68, c_0_46])).
cnf(c_0_84,plain,(rpoint(vd1299)), inference(rw,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_85,plain,(v0=X1|vplus(X2,X1)!=X2|~rreal(X2)|~rreal(X1)), inference(spm,[status(thm)],[c_0_71, c_0_38])).
cnf(c_0_86,plain,(vplus(vf(X1,X2),vf(X2,X3))=vf(X1,X3)|~rpoint(X3)|~rpoint(X2)|~rpoint(X1)|~rR(X2,X1,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_72])])])).
cnf(c_0_87,plain,(rreal(vf(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_88,plain,(X1=vd1287|X1=vd1288|rR(X1,vd1287,vd1288)|rR(vd1287,vd1288,X1)|rR(vd1288,vd1287,X1)|~rpoint(X1)|~ron(X1,vd1295)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_65])]), c_0_77])).
cnf(c_0_89,plain,(ron(vd1289,vd1295)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_59]), c_0_79]), c_0_76])])).
cnf(c_0_90,plain,(vd1287!=vd1289), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_91,plain,(vd1288!=vd1289), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_92, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((((~rR(X7,X8,X9)|(((~rR(X8,X7,X9)|(X9!=X10|~rpoint(X10)))|(X8!=X11|~rpoint(X11)))|(X7!=X12|~rpoint(X12))))&(X7!=X9|(((~rR(X8,X7,X9)|(X9!=X10|~rpoint(X10)))|(X8!=X11|~rpoint(X11)))|(X7!=X12|~rpoint(X12)))))&(X7!=X8|(((~rR(X8,X7,X9)|(X9!=X10|~rpoint(X10)))|(X8!=X11|~rpoint(X11)))|(X7!=X12|~rpoint(X12)))))&(rR(X8,X9,X7)|(((~rR(X8,X7,X9)|(X9!=X10|~rpoint(X10)))|(X8!=X11|~rpoint(X11)))|(X7!=X12|~rpoint(X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_80])])])])).
cnf(c_0_93,plain,(rR(X1,X2,X3)|~rpoint(X4)|X2!=X4|~rpoint(X5)|X1!=X5|~rpoint(X6)|X7!=X6|~rpoint(X8)|X3!=X8|~rR(X1,X2,X7)|~rR(X7,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_94,plain,(rR(X1,vd1287,vd1299)|~rpoint(X1)|~rR(X1,vd1287,vd1302)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84]), c_0_63]), c_0_65])])).
cnf(c_0_95,plain,(vf(X1,X2)=v0|vf(X3,X2)!=vf(X3,X1)|~rpoint(X2)|~rpoint(X1)|~rpoint(X3)|~rR(X1,X3,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87]), c_0_87])).
cnf(c_0_96,plain,(rR(vd1288,vd1287,vd1289)|rR(vd1287,vd1288,vd1289)|rR(vd1289,vd1287,vd1288)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_64])]), c_0_90]), c_0_91])).
cnf(c_0_97,plain,(vf(vd1287,vd1288)=vf(vd1287,vd1289)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_98,plain,(~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rR(X4,X2,X6)|~rR(X2,X4,X6)), inference(split_conjunct,[status(thm)],[c_0_92])).
fof(c_0_99, plain, (![X64]:![X65]:![X66]:![X67]:(((X65!=X66|~rpoint(X66))|(X64!=X67|~rpoint(X67)))|vf(X64,X65)=vf(X65,X64))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
cnf(c_0_100,plain,(rR(X1,X2,X3)|~rpoint(X3)|~rpoint(X4)|~rpoint(X1)|~rpoint(X2)|~rR(X4,X1,X3)|~rR(X1,X2,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_93])])])])).
cnf(c_0_101,plain,(rR(vd1289,vd1287,vd1299)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_62]), c_0_64])])).
cnf(c_0_102,plain,(vf(vd1288,vd1289)=v0|rR(vd1289,vd1287,vd1288)|rR(vd1287,vd1288,vd1289)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_97]), c_0_64]), c_0_60]), c_0_65])])).
cnf(c_0_103,plain,(~rpoint(X1)|~rpoint(X2)|~rpoint(X3)|~rR(X2,X3,X1)|~rR(X3,X2,X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_98])])])).
cnf(c_0_104,plain,(rR(vd1288,vd1287,vd1299)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_105, plain, (![X113]:![X114]:![X115]:![X116]:(((vf(X113,X114)!=v0|X113=X114)|((X114!=X115|~rpoint(X115))|(X113!=X116|~rpoint(X116))))&((X113!=X114|vf(X113,X114)=v0)|((X114!=X115|~rpoint(X115))|(X113!=X116|~rpoint(X116)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])])).
cnf(c_0_106,plain,(vf(X1,X2)=vf(X2,X1)|~rpoint(X3)|X1!=X3|~rpoint(X4)|X2!=X4), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_107,plain,(rR(vd1287,X1,vd1299)|~rpoint(X1)|~rR(vd1287,X1,vd1289)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_84]), c_0_64]), c_0_65])])).
cnf(c_0_108,plain,(vf(vd1288,vd1289)=v0|vf(vd1289,vd1288)=v0|rR(vd1287,vd1288,vd1289)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_102]), c_0_97]), c_0_60]), c_0_64]), c_0_65])])).
cnf(c_0_109,plain,(~rR(vd1287,vd1288,vd1299)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_84]), c_0_60]), c_0_65])])).
cnf(c_0_110,plain,(X2=X4|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3|vf(X2,X4)!=v0), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_111,plain,(vf(X1,X2)=vf(X2,X1)|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_106])])).
cnf(c_0_112,plain,(vf(vd1289,vd1288)=v0|vf(vd1288,vd1289)=v0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_60])]), c_0_109])).
cnf(c_0_113,plain,(X1=X2|vf(X1,X2)!=v0|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_110])])).
cnf(c_0_114,plain,(vf(vd1289,vd1288)=v0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_64]), c_0_60])])).
cnf(c_0_115,plain,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_60]), c_0_64])]), c_0_91]), ['proof']).
# SZS output end CNFRefutation.
EOF
