% Problem    : SWW329+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n140.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:15:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X16]:![X10]:(hAPP(c_Set_OCollect(X10),X16)=c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))<=>![X3]:~(hBOOL(hAPP(X16,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__empty__eq)).
fof(c_0_1, axiom, (![X10]:c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),hAPP(c_COMBK(tc_HOL_Obool,X10),c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__def)).
fof(c_0_2, axiom, (![X16]:![X10]:hAPP(c_Set_OCollect(X10),X16)=X16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__def)).
fof(c_0_3, axiom, (![X16]:![X10]:hAPP(c_Set_OCollect(X10),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X10),c_fNot),X16))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),hAPP(c_Set_OCollect(X10),X16))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__neg__eq)).
fof(c_0_4, axiom, (![X312]:![X66]:![X10]:![X29]:hAPP(hAPP(c_COMBK(X29,X10),X66),X312)=X66), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBK__1)).
fof(c_0_5, axiom, (![X19]:![X10]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),X19))=X19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__complement)).
fof(c_0_6, axiom, (![X16]:(~(hBOOL(hAPP(c_fNot,X16)))|~(hBOOL(X16)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__1)).
fof(c_0_7, axiom, (![X10]:c_Orderings_Otop__class_Otop(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),hAPP(c_COMBK(tc_HOL_Obool,X10),c_fTrue))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__def)).
fof(c_0_8, axiom, (![X58]:hAPP(c_Nat_OSuc,X58)!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Suc__neq__Zero)).
fof(c_0_9, axiom, (hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_nat__0)).
fof(c_0_10, axiom, (c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Pls__def)).
fof(c_0_11, axiom, (![X54]:![X15]:![X16]:![X24]:![X10]:![X20]:hAPP(hAPP(hAPP(c_COMBB(X20,X10,X24),X16),X15),X54)=hAPP(X16,hAPP(X15,X54))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBB__1)).
fof(c_0_12, axiom, (![X10]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X10,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__empty__eq)).
fof(c_0_13, axiom, (![X49]:![X50]:![X10]:(class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(X10)=>(X50=hAPP(hAPP(c_Groups_Oplus__class_Oplus(X10),X50),X49)<=>X49=c_Groups_Ozero__class_Ozero(X10)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_add__0__iff)).
fof(c_0_14, axiom, (![X28]:![X58]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X58),hAPP(c_Nat_OSuc,X28))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X58),X28))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_add__Suc__right)).
fof(c_0_15, axiom, (![X249]:![X29]:hAPP(c_List_Orev(X29),hAPP(c_List_Orev(X29),X249))=X249), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rev__rev__ident)).
fof(c_0_16, axiom, (class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Nat__Onat__Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct)).
fof(c_0_17, axiom, (![X58]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X58),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X58), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Nat_Oadd__0__right)).
fof(c_0_18, axiom, (![X28]:X28!=hAPP(c_Nat_OSuc,X28)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_n__not__Suc__n)).
fof(c_0_19, plain, (![X16]:![X10]:(hAPP(c_Set_OCollect(X10),X16)=c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))<=>![X3]:~hBOOL(hAPP(X16,X3)))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_20, plain, (![X11]:c_Orderings_Obot__class_Obot(tc_fun(X11,tc_HOL_Obool))=hAPP(c_Set_OCollect(X11),hAPP(c_COMBK(tc_HOL_Obool,X11),c_fFalse))), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_21, plain, (![X17]:![X18]:hAPP(c_Set_OCollect(X18),X17)=X17), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_22, plain, (![X17]:![X18]:hAPP(c_Set_OCollect(X18),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X18),c_fNot),X17))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X18,tc_HOL_Obool)),hAPP(c_Set_OCollect(X18),X17))), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_23, plain, (![X17]:![X18]:![X19]:![X20]:![X22]:((hAPP(c_Set_OCollect(X18),X17)!=c_Orderings_Obot__class_Obot(tc_fun(X18,tc_HOL_Obool))|~hBOOL(hAPP(X17,X19)))&(hBOOL(hAPP(X20,esk161_1(X20)))|hAPP(c_Set_OCollect(X22),X20)=c_Orderings_Obot__class_Obot(tc_fun(X22,tc_HOL_Obool))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])])).
fof(c_0_24, plain, (![X313]:![X314]:![X315]:![X316]:hAPP(hAPP(c_COMBK(X316,X315),X314),X313)=X314), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_25,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26,plain,(hAPP(c_Set_OCollect(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_27, plain, (![X20]:![X21]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X21,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X21,tc_HOL_Obool)),X20))=X20), inference(variable_rename,[status(thm)],[c_0_5])).
cnf(c_0_28,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Set_OCollect(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, plain, (![X16]:(~hBOOL(hAPP(c_fNot,X16))|~hBOOL(X16))), inference(fof_simplification,[status(thm)],[c_0_6])).
cnf(c_0_30,plain,(hAPP(c_Set_OCollect(X1),X2)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))|hBOOL(hAPP(X2,esk161_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31,plain,(hAPP(hAPP(c_COMBK(X1,X2),X3),X4)=X3), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_33, plain, (![X11]:c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool))=hAPP(c_Set_OCollect(X11),hAPP(c_COMBK(tc_HOL_Obool,X11),c_fTrue))), inference(variable_rename,[status(thm)],[c_0_7])).
fof(c_0_34, plain, (![X59]:hAPP(c_Nat_OSuc,X59)!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_35,plain,(hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_36,plain,(c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_37, plain, (![X55]:![X56]:![X57]:![X58]:![X59]:![X60]:hAPP(hAPP(hAPP(c_COMBB(X60,X59,X58),X57),X56),X55)=hAPP(X57,hAPP(X56,X55))), inference(variable_rename,[status(thm)],[c_0_11])).
cnf(c_0_38,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_39,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2)=hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_26]), c_0_26])).
fof(c_0_40, plain, (![X17]:(~hBOOL(hAPP(c_fNot,X17))|~hBOOL(X17))), inference(variable_rename,[status(thm)],[c_0_29])).
cnf(c_0_41,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hBOOL(hAPP(X1,esk161_1(X1)))), inference(rw,[status(thm)],[c_0_30, c_0_26])).
cnf(c_0_42,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)),X2)=c_fFalse), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_43, plain, (![X11]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X11,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X11,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_12])).
cnf(c_0_44,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue))), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_45, plain, (![X51]:![X52]:![X53]:(((X52!=hAPP(hAPP(c_Groups_Oplus__class_Oplus(X53),X52),X51)|X51=c_Groups_Ozero__class_Ozero(X53))|~class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(X53))&((X51!=c_Groups_Ozero__class_Ozero(X53)|X52=hAPP(hAPP(c_Groups_Oplus__class_Oplus(X53),X52),X51))|~class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(X53)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
fof(c_0_46, plain, (![X59]:![X60]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X60),hAPP(c_Nat_OSuc,X59))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X60),X59))), inference(variable_rename,[status(thm)],[c_0_14])).
cnf(c_0_47,plain,(hAPP(c_Nat_OSuc,X1)!=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_49,plain,(hAPP(hAPP(hAPP(c_COMBB(X1,X2,X3),X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_50,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39]), c_0_39])).
fof(c_0_51, plain, (![X250]:![X251]:hAPP(c_List_Orev(X251),hAPP(c_List_Orev(X251),X250))=X250), inference(variable_rename,[status(thm)],[c_0_15])).
cnf(c_0_52,plain,(~hBOOL(X1)|~hBOOL(hAPP(c_fNot,X1))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_53,plain,(c_fFalse=X1|hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_41]), c_0_42]), c_0_31])).
cnf(c_0_54,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue)=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_44, c_0_26])).
cnf(c_0_56,plain,(X2=c_Groups_Ozero__class_Ozero(X1)|~class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(X1)|X3!=hAPP(hAPP(c_Groups_Oplus__class_Oplus(X1),X3),X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),hAPP(c_Nat_OSuc,X2))=hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58,plain,(class_Semiring__Normalization_Ocomm__semiring__1__cancel__crossproduct(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_59,plain,(hAPP(c_Nat_OSuc,X1)!=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_47, c_0_48])).
fof(c_0_60, plain, (![X59]:hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X59),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X59), inference(variable_rename,[status(thm)],[c_0_17])).
cnf(c_0_61,plain,(hAPP(c_fNot,hAPP(c_fNot,hAPP(X1,X2)))=hAPP(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_49])).
cnf(c_0_62,plain,(hAPP(c_List_Orev(X1),hAPP(c_List_Orev(X1),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_63,plain,(hAPP(c_fNot,X1)=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_64,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_54, c_0_39])).
cnf(c_0_65,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)),X2)=c_fTrue), inference(spm,[status(thm)],[c_0_31, c_0_55])).
cnf(c_0_66,plain,(hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2))!=X1), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_48]), c_0_58])]), c_0_59])).
cnf(c_0_67,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=X1), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_68, plain, (![X29]:X29!=hAPP(c_Nat_OSuc,X29)), inference(variable_rename,[status(thm)],[c_0_18])).
cnf(c_0_69,plain,(hAPP(c_fNot,hAPP(c_fNot,X1))=X1), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_70,plain,(hAPP(c_fNot,X1)=c_fFalse|c_fFalse=X1), inference(spm,[status(thm)],[c_0_63, c_0_53])).
cnf(c_0_71,plain,(hAPP(c_fNot,c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_64]), c_0_65]), c_0_42])).
cnf(c_0_72,plain,(hAPP(c_Nat_OSuc,hAPP(c_Nat_OSuc,hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),X2)))!=X1), inference(spm,[status(thm)],[c_0_66, c_0_57])).
cnf(c_0_73,plain,(hAPP(hAPP(c_Groups_Oplus__class_Oplus(tc_Nat_Onat),X1),hAPP(c_Int_Onat,c_Int_OPls))=X1), inference(rw,[status(thm)],[c_0_67, c_0_48])).
cnf(c_0_74,plain,(X1!=hAPP(c_Nat_OSuc,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_75,plain,(c_fFalse=X1|c_fTrue=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])).
cnf(c_0_76,plain,(hAPP(c_Nat_OSuc,hAPP(c_Nat_OSuc,X1))!=X1), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_77,plain,(hAPP(c_Nat_OSuc,X1)=c_fFalse|c_fTrue!=X1), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_78,plain,(hAPP(c_Nat_OSuc,c_fFalse)!=X1|c_fTrue!=X1), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_79,plain,(c_fTrue!=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_75]), c_0_74])).
cnf(c_0_80,plain,($false), inference(sr,[status(thm)],[c_0_65, c_0_79]), ['proof']).
# SZS output end CNFRefutation.
EOF
