% Problem    : SET083+1 : TPTP v6.1.0. Bugfixed v5.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n091.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:44:37 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_B42_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X1]:![X2]:((singleton(X1)=singleton(X2)&member(X1,universal_class))=>X1=X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', singleton_identified_by_element1)).
fof(c_0_1, axiom, (![X1]:singleton(X1)=unordered_pair(X1,X1)), file('/export/starexec/sandbox/benchmark/Axioms/SET005+0.ax', singleton_set_defn)).
fof(c_0_2, axiom, (![X3]:![X1]:![X2]:(member(X3,unordered_pair(X1,X2))<=>(member(X3,universal_class)&(X3=X1|X3=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SET005+0.ax', unordered_pair_defn)).
fof(c_0_3, negated_conjecture, (~(![X1]:![X2]:((singleton(X1)=singleton(X2)&member(X1,universal_class))=>X1=X2))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_4, negated_conjecture, (((singleton(esk8_0)=singleton(esk9_0)&member(esk8_0,universal_class))&esk8_0!=esk9_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
fof(c_0_5, plain, (![X2]:singleton(X2)=unordered_pair(X2,X2)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_6, plain, (![X4]:![X5]:![X6]:![X7]:![X8]:![X9]:(((member(X4,universal_class)|~member(X4,unordered_pair(X5,X6)))&((X4=X5|X4=X6)|~member(X4,unordered_pair(X5,X6))))&(((X7!=X8|~member(X7,universal_class))|member(X7,unordered_pair(X8,X9)))&((X7!=X9|~member(X7,universal_class))|member(X7,unordered_pair(X8,X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
cnf(c_0_7,negated_conjecture,(singleton(esk8_0)=singleton(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8,plain,(singleton(X1)=unordered_pair(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9,plain,(member(X1,unordered_pair(X2,X3))|~member(X1,universal_class)|X1!=X3), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10,plain,(X1=X3|X1=X2|~member(X1,unordered_pair(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,negated_conjecture,(unordered_pair(esk9_0,esk9_0)=unordered_pair(esk8_0,esk8_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8]), c_0_8])).
cnf(c_0_12,plain,(member(X1,unordered_pair(X2,X1))|~member(X1,universal_class)), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_13,negated_conjecture,(member(esk8_0,universal_class)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14,negated_conjecture,(X1=esk9_0|~member(X1,unordered_pair(esk8_0,esk8_0))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15,negated_conjecture,(member(esk8_0,unordered_pair(X1,esk8_0))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16,negated_conjecture,(esk8_0!=esk9_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17,negated_conjecture,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), ['proof']).
# SZS output end CNFRefutation.
EOF
