% Problem    : GEO278+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n130.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:01:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X105]:![X106]:![X107]:![X108]:![X109]:(((((((vf(X107,X105)=vf(X107,X106)&ron(X106,X109))&ron(X105,X108))&rcenter(X107,X109))&rcenter(X107,X108))&?[X110]:(X106=X110&rpoint(X110)))&?[X111]:(X105=X111&rpoint(X111)))=>X108=X109)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(182), 0), imp(cond(axiom(182), 0)))')).
fof(c_0_1, conjecture, (![X1]:(((ron(vd1096,X1)&rcenter(vd1089,X1))&rcircle(X1))=>X1=vd1102)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(theu(the(231), 1), imp(the(231)))')).
fof(c_0_2, axiom, (((ron(vd1096,vd1102)&rcenter(vd1089,vd1102))&rcircle(vd1102))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(pred(comma_conjunct2(the(231)), 0), and(pred(comma_conjunct1(the(231)), 0), pred(the(231), 0)))')).
fof(c_0_3, axiom, (rpoint(vd1096)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'pred(229, 0)')).
fof(c_0_4, plain, (![X112]:![X113]:![X114]:![X115]:![X116]:![X117]:![X118]:(((((((vf(X114,X112)!=vf(X114,X113)|~ron(X113,X116))|~ron(X112,X115))|~rcenter(X114,X116))|~rcenter(X114,X115))|(X113!=X117|~rpoint(X117)))|(X112!=X118|~rpoint(X118)))|X115=X116)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])).
fof(c_0_5, negated_conjecture, (~(![X1]:(((ron(vd1096,X1)&rcenter(vd1089,X1))&rcircle(X1))=>X1=vd1102))), inference(assume_negation,[status(cth)],[c_0_1])).
cnf(c_0_6,plain,(X1=X2|~rpoint(X3)|X4!=X3|~rpoint(X5)|X6!=X5|~rcenter(X7,X1)|~rcenter(X7,X2)|~ron(X4,X1)|~ron(X6,X2)|vf(X7,X4)!=vf(X7,X6)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, negated_conjecture, ((((ron(vd1096,esk1_0)&rcenter(vd1089,esk1_0))&rcircle(esk1_0))&esk1_0!=vd1102)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_8,plain,(X1=X2|vf(X3,X4)!=vf(X3,X5)|~rpoint(X5)|~rpoint(X4)|~rcenter(X3,X2)|~rcenter(X3,X1)|~ron(X5,X2)|~ron(X4,X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_6])])).
cnf(c_0_9,negated_conjecture,(rcenter(vd1089,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,negated_conjecture,(X1=esk1_0|vf(vd1089,X2)!=vf(vd1089,X3)|~rpoint(X3)|~rpoint(X2)|~rcenter(vd1089,X1)|~ron(X3,esk1_0)|~ron(X2,X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_11,plain,(rcenter(vd1089,vd1102)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_12,negated_conjecture,(esk1_0!=vd1102), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13,negated_conjecture,(vf(vd1089,X1)!=vf(vd1089,X2)|~rpoint(X2)|~rpoint(X1)|~ron(X2,esk1_0)|~ron(X1,vd1102)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])).
cnf(c_0_14,negated_conjecture,(~rpoint(X1)|~ron(X1,esk1_0)|~ron(X1,vd1102)), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_15,negated_conjecture,(ron(vd1096,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16,plain,(rpoint(vd1096)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_17,plain,(ron(vd1096,vd1102)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_18,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])]), ['proof']).
# SZS output end CNFRefutation.
EOF
