% Problem    : SWB049+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n103.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:27:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_C45_F1_AE_R8_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X4]:![X7]:(iext(uri_owl_topObjectProperty,X4,X7)<=>(ir(X4)&ir(X7)))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_prop_topobjectproperty_ext)).
fof(c_0_1, axiom, (![X4]:ir(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', simple_ir)).
fof(c_0_2, axiom, (![X2]:(icext(uri_owl_ReflexiveProperty,X2)<=>(ip(X2)&![X4]:iext(X2,X4,X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_char_reflexive)).
fof(c_0_3, axiom, (![X1]:![X2]:![X3]:(iext(X2,X1,X3)=>ip(X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', simple_iext_property)).
fof(c_0_4, axiom, (![X33]:![X34]:(iext(uri_owl_equivalentProperty,X33,X34)<=>((ip(X33)&ip(X34))&![X4]:![X7]:(iext(X33,X4,X7)<=>iext(X34,X4,X7))))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', owl_eqdis_equivalentproperty)).
fof(c_0_5, axiom, (iext(uri_owl_equivalentProperty,uri_ex_p,uri_owl_topObjectProperty)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', premise_rdfbased_sem_char_reflexive_ext)).
fof(c_0_6, axiom, (![X4]:![X5]:(iext(uri_rdf_type,X4,X5)<=>icext(X5,X4))), file('/export/starexec/sandbox/benchmark/Axioms/SWB001+0.ax', rdfs_cext_def)).
fof(c_0_7, conjecture, (iext(uri_rdf_type,uri_ex_p,uri_owl_ReflexiveProperty)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conclusion_rdfbased_sem_char_reflexive_ext)).
fof(c_0_8, plain, (![X8]:![X9]:![X10]:![X11]:(((ir(X8)|~iext(uri_owl_topObjectProperty,X8,X9))&(ir(X9)|~iext(uri_owl_topObjectProperty,X8,X9)))&((~ir(X10)|~ir(X11))|iext(uri_owl_topObjectProperty,X10,X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])).
fof(c_0_9, plain, (![X5]:ir(X5)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_10, plain, (![X5]:![X6]:![X7]:(((ip(X5)|~icext(uri_owl_ReflexiveProperty,X5))&(iext(X5,X6,X6)|~icext(uri_owl_ReflexiveProperty,X5)))&((~ip(X7)|~iext(X7,esk170_1(X7),esk170_1(X7)))|icext(uri_owl_ReflexiveProperty,X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
fof(c_0_11, plain, (![X4]:![X5]:![X6]:(~iext(X5,X4,X6)|ip(X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
fof(c_0_12, plain, (![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:((((ip(X35)|~iext(uri_owl_equivalentProperty,X35,X36))&(ip(X36)|~iext(uri_owl_equivalentProperty,X35,X36)))&(((~iext(X35,X37,X38)|iext(X36,X37,X38))|~iext(uri_owl_equivalentProperty,X35,X36))&((~iext(X36,X39,X40)|iext(X35,X39,X40))|~iext(uri_owl_equivalentProperty,X35,X36))))&((((~iext(X41,esk117_2(X41,X42),esk118_2(X41,X42))|~iext(X42,esk117_2(X41,X42),esk118_2(X41,X42)))|(~ip(X41)|~ip(X42)))|iext(uri_owl_equivalentProperty,X41,X42))&(((iext(X41,esk117_2(X41,X42),esk118_2(X41,X42))|iext(X42,esk117_2(X41,X42),esk118_2(X41,X42)))|(~ip(X41)|~ip(X42)))|iext(uri_owl_equivalentProperty,X41,X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
cnf(c_0_13,plain,(iext(uri_owl_topObjectProperty,X1,X2)|~ir(X2)|~ir(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(ir(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15,plain,(icext(uri_owl_ReflexiveProperty,X1)|~iext(X1,esk170_1(X1),esk170_1(X1))|~ip(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(ip(X1)|~iext(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17,plain,(iext(X1,X3,X4)|~iext(uri_owl_equivalentProperty,X1,X2)|~iext(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18,plain,(iext(uri_owl_equivalentProperty,uri_ex_p,uri_owl_topObjectProperty)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19,plain,(iext(uri_owl_topObjectProperty,X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_14])])).
fof(c_0_20, plain, (![X6]:![X7]:![X8]:![X9]:((~iext(uri_rdf_type,X6,X7)|icext(X7,X6))&(~icext(X9,X8)|iext(uri_rdf_type,X8,X9)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_21,plain,(icext(uri_owl_ReflexiveProperty,X1)|~iext(X1,esk170_1(X1),esk170_1(X1))), inference(csr,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22,plain,(iext(uri_ex_p,X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
fof(c_0_23, negated_conjecture, (~iext(uri_rdf_type,uri_ex_p,uri_owl_ReflexiveProperty)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_7])])).
cnf(c_0_24,plain,(iext(uri_rdf_type,X1,X2)|~icext(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25,plain,(icext(uri_owl_ReflexiveProperty,uri_ex_p)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26,negated_conjecture,(~iext(uri_rdf_type,uri_ex_p,uri_owl_ReflexiveProperty)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), ['proof']).
# SZS output end CNFRefutation.
EOF
