% Problem    : GEO284+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n134.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:56:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___100_C18_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X39]:![X40]:(((X39!=X40&?[X41]:(X40=X41&rpoint(X41)))&?[X42]:(X39=X42&rpoint(X42)))=>?[X43]:((ron(X40,X43)&rcenter(X39,X43))&rcircle(X43)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(53), 0), imp(cond(axiom(53), 0)))')).
fof(c_0_1, axiom, (((((((((((rless(vf(vd1127,vd1129),vf(vd1123,vd1125))&vd1127!=vd1129)&vd1123!=vd1125)&vd1129=vd1130)&rpoint(vd1130))&vd1127=vd1128)&rpoint(vd1128))&vd1125=vd1126)&rpoint(vd1126))&vd1123=vd1124)&rpoint(vd1124))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'and(and(and(and(holds(conjunct2(conjunct2(conjunct2(conjunct2(conjunct2(conjunct2(242)))))), 1133, 0), and(holds(conjunct1(conjunct2(conjunct2(conjunct2(conjunct2(conjunct2(242)))))), 1132, 0), and(holds(conjunct1(conjunct2(conjunct2(conjunct2(conjunct2(242))))), 1131, 0), and(pred(conjunct1(conjunct2(conjunct2(conjunct2(242)))), 2), pred(conjunct1(conjunct2(conjunct2(conjunct2(242)))), 1))))), and(pred(conjunct1(conjunct2(conjunct2(242))), 2), pred(conjunct1(conjunct2(conjunct2(242))), 1))), and(pred(conjunct1(conjunct2(242)), 2), pred(conjunct1(conjunct2(242)), 1))), and(pred(conjunct1(242), 2), pred(conjunct1(242), 1)))')).
fof(c_0_2, axiom, (vd1123=vd1129), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(258, 1150, 0)')).
fof(c_0_3, conjecture, (?[X1]:?[X2]:(((((ron(vd1127,X2)&rcenter(vd1123,X2))&rcircle(X2))&ron(vd1125,X1))&ron(vd1123,X1))&rline(X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qe(259)')).
fof(c_0_4, axiom, (![X53]:![X54]:(((X53!=X54&?[X55]:(X54=X55&rpoint(X55)))&?[X56]:(X53=X56&rpoint(X56)))=>?[X57]:((ron(X54,X57)&ron(X53,X57))&rline(X57)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'qu(cond(axiom(51), 0), imp(cond(axiom(51), 0)))')).
fof(c_0_5, plain, (![X44]:![X45]:![X46]:![X47]:(((ron(X45,esk3_2(X44,X45))|((X44=X45|(X45!=X46|~rpoint(X46)))|(X44!=X47|~rpoint(X47))))&(rcenter(X44,esk3_2(X44,X45))|((X44=X45|(X45!=X46|~rpoint(X46)))|(X44!=X47|~rpoint(X47)))))&(rcircle(esk3_2(X44,X45))|((X44=X45|(X45!=X46|~rpoint(X46)))|(X44!=X47|~rpoint(X47)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])).
cnf(c_0_6,plain,(vd1129=vd1130), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_7,plain,(vd1123=vd1129), inference(split_conjunct,[status(thm)],[c_0_2])).
fof(c_0_8, plain, ((~epred2_0<=>![X2]:((~ron(vd1127,X2)|~rcenter(vd1123,X2))|~rcircle(X2)))), introduced(definition)).
cnf(c_0_9,plain,(X2=X4|rcenter(X2,esk3_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,plain,(rpoint(vd1130)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_11,plain,(vd1130=vd1123), inference(rw,[status(thm)],[c_0_6, c_0_7])).
fof(c_0_12, negated_conjecture, (~(?[X1]:?[X2]:(((((ron(vd1127,X2)&rcenter(vd1123,X2))&rcircle(X2))&ron(vd1125,X1))&ron(vd1123,X1))&rline(X1)))), inference(assume_negation,[status(cth)],[c_0_3])).
cnf(c_0_13,negated_conjecture,(epred2_0|~rcircle(X1)|~rcenter(vd1123,X1)|~ron(vd1127,X1)), inference(split_equiv,[status(thm)],[c_0_8])).
cnf(c_0_14,plain,(X1=X2|rcenter(X1,esk3_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_9])])).
cnf(c_0_15,plain,(rpoint(vd1123)), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_16,plain,(X2=X4|rcircle(esk3_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_17, negated_conjecture, (![X3]:![X4]:(((((~ron(vd1127,X4)|~rcenter(vd1123,X4))|~rcircle(X4))|~ron(vd1125,X3))|~ron(vd1123,X3))|~rline(X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
cnf(c_0_18,negated_conjecture,(vd1123=X1|epred2_0|~rpoint(X1)|~rcircle(esk3_2(vd1123,X1))|~ron(vd1127,esk3_2(vd1123,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
cnf(c_0_19,plain,(X1=X2|rcircle(esk3_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_16])])).
cnf(c_0_20,plain,(X2=X4|ron(X4,esk3_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_21,plain,(rpoint(vd1128)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_22,plain,(vd1127=vd1128), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_23,plain,(vd1127!=vd1129), inference(split_conjunct,[status(thm)],[c_0_1])).
fof(c_0_24, plain, (![X58]:![X59]:![X60]:![X61]:(((ron(X59,esk5_2(X58,X59))|((X58=X59|(X59!=X60|~rpoint(X60)))|(X58!=X61|~rpoint(X61))))&(ron(X58,esk5_2(X58,X59))|((X58=X59|(X59!=X60|~rpoint(X60)))|(X58!=X61|~rpoint(X61)))))&(rline(esk5_2(X58,X59))|((X58=X59|(X59!=X60|~rpoint(X60)))|(X58!=X61|~rpoint(X61)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
fof(c_0_25, plain, ((~epred1_0<=>![X1]:((~ron(vd1125,X1)|~ron(vd1123,X1))|~rline(X1)))), introduced(definition)).
cnf(c_0_26,negated_conjecture,(~rline(X1)|~ron(vd1123,X1)|~ron(vd1125,X1)|~rcircle(X2)|~rcenter(vd1123,X2)|~ron(vd1127,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27,negated_conjecture,(vd1123=X1|epred2_0|~rpoint(X1)|~ron(vd1127,esk3_2(vd1123,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_15])])).
cnf(c_0_28,plain,(X1=X2|ron(X2,esk3_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_20])])).
cnf(c_0_29,plain,(rpoint(vd1127)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_30,plain,(vd1127!=vd1123), inference(rw,[status(thm)],[c_0_23, c_0_7])).
cnf(c_0_31,plain,(X2=X4|rline(esk5_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32,negated_conjecture,(~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_26, c_0_25]), c_0_8])).
cnf(c_0_33,negated_conjecture,(epred2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_15])]), c_0_30])).
cnf(c_0_34,negated_conjecture,(epred1_0|~rline(X1)|~ron(vd1125,X1)|~ron(vd1123,X1)), inference(split_equiv,[status(thm)],[c_0_25])).
cnf(c_0_35,plain,(X1=X2|rline(esk5_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])])).
cnf(c_0_36,negated_conjecture,(~epred1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33])])).
cnf(c_0_37,plain,(X2=X4|ron(X4,esk5_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38,plain,(rpoint(vd1126)), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_39,plain,(vd1125=vd1126), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_40,negated_conjecture,(X1=X2|~rpoint(X2)|~rpoint(X1)|~ron(vd1125,esk5_2(X1,X2))|~ron(vd1123,esk5_2(X1,X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_41,plain,(X1=X2|ron(X2,esk5_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_37])])).
cnf(c_0_42,plain,(rpoint(vd1125)), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43,plain,(X2=X4|ron(X2,esk5_2(X2,X4))|~rpoint(X1)|X2!=X1|~rpoint(X3)|X4!=X3), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_44,negated_conjecture,(X1=vd1125|~rpoint(X1)|~ron(vd1123,esk5_2(X1,vd1125))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])])).
cnf(c_0_45,plain,(X1=X2|ron(X1,esk5_2(X1,X2))|~rpoint(X2)|~rpoint(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_43])])).
cnf(c_0_46,plain,(vd1123!=vd1125), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_47,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_15]), c_0_42])]), c_0_46]), ['proof']).
# SZS output end CNFRefutation.
EOF
