% Problem    : NUM840+1 : TPTP v6.1.0. Released v4.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n109.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:37:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0U
# and selection function SelectComplexExceptRRHorn.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X27]:![X28]:(~(greater(X27,X28))|~(less(X27,X28)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(goal(130), 0), 2)')).
fof(c_0_1, axiom, (![X23]:![X24]:(less(X23,X24)=>greater(X24,X23))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(147, 0), 0)')).
fof(c_0_2, axiom, (![X44]:![X45]:vplus(X45,X44)=vplus(X44,X45)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(61, 0), 0)')).
fof(c_0_3, axiom, (less(vplus(vd328,vd330),vplus(vd329,vd330))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(antec(conjunct2(conjunct2(204))), 335, 0)')).
fof(c_0_4, axiom, (![X1]:![X2]:![X3]:(less(X1,X2)=>less(vplus(X1,X3),vplus(X2,X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(goal(193), 0), 0)')).
fof(c_0_5, axiom, (![X25]:![X26]:(greater(X25,X26)=>less(X26,X25))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(140, 0), 0)')).
fof(c_0_6, axiom, (![X27]:![X28]:((X27=X28|greater(X27,X28))|less(X27,X28))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(goal(130), 0), 0)')).
fof(c_0_7, axiom, (![X27]:![X28]:(X27!=X28|~(less(X27,X28)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'ass(cond(goal(130), 0), 1)')).
fof(c_0_8, conjecture, (less(vd328,vd329)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', 'holds(conseq(conjunct2(conjunct2(204))), 336, 0)')).
fof(c_0_9, plain, (![X27]:![X28]:(~greater(X27,X28)|~less(X27,X28))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_10, plain, (![X29]:![X30]:(~greater(X29,X30)|~less(X29,X30))), inference(variable_rename,[status(thm)],[c_0_9])).
fof(c_0_11, plain, (![X25]:![X26]:(~less(X25,X26)|greater(X26,X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
fof(c_0_12, plain, (![X46]:![X47]:vplus(X47,X46)=vplus(X46,X47)), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_13,plain,(~less(X1,X2)|~greater(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,plain,(greater(X1,X2)|~less(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,plain,(less(vplus(vd328,vd330),vplus(vd329,vd330))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_16,plain,(vplus(X1,X2)=vplus(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, (![X4]:![X5]:![X6]:(~less(X4,X5)|less(vplus(X4,X6),vplus(X5,X6)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_18,plain,(~less(X1,X2)|~less(X2,X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19,plain,(less(vplus(vd328,vd330),vplus(vd330,vd329))), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20,plain,(less(vplus(X1,X2),vplus(X3,X2))|~less(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, plain, (![X27]:![X28]:(~greater(X27,X28)|less(X28,X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
fof(c_0_22, plain, (![X29]:![X30]:((X29=X30|greater(X29,X30))|less(X29,X30))), inference(variable_rename,[status(thm)],[c_0_6])).
fof(c_0_23, plain, (![X27]:![X28]:(X27!=X28|~less(X27,X28))), inference(fof_simplification,[status(thm)],[c_0_7])).
cnf(c_0_24,plain,(~less(vplus(vd330,vd329),vplus(vd328,vd330))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25,plain,(less(vplus(X1,X2),vplus(X3,X1))|~less(X2,X3)), inference(spm,[status(thm)],[c_0_20, c_0_16])).
cnf(c_0_26,plain,(less(X1,X2)|~greater(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27,plain,(less(X1,X2)|greater(X1,X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, negated_conjecture, (~less(vd328,vd329)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_8])])).
fof(c_0_29, plain, (![X29]:![X30]:(X29!=X30|~less(X29,X30))), inference(variable_rename,[status(thm)],[c_0_23])).
cnf(c_0_30,plain,(~less(vd329,vd328)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31,plain,(X1=X2|less(X1,X2)|less(X2,X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32,negated_conjecture,(~less(vd328,vd329)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33,plain,(~less(X1,X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34,plain,(vd329=vd328), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_35,plain,(~less(X1,X1)), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_36,plain,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_34]), c_0_16]), c_0_35]), ['proof']).
# SZS output end CNFRefutation.
EOF
