% Problem    : CSR017+1 : TPTP v6.1.0. Bugfixed v3.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n185.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:23:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___212_C18_F1_AE_Q12_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X2]:![X5]:(((~(holdsAt(X2,X5))&~(releasedAt(X2,plus(X5,n1))))&~(?[X4]:(happens(X4,X5)&initiates(X4,X2,X5))))=>~(holdsAt(X2,plus(X5,n1))))), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+0.ax', keep_not_holding)).
fof(c_0_1, hypothesis, (![X2]:![X5]:~(releasedAt(X2,X5))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', not_releasedAt)).
fof(c_0_2, axiom, (![X9]:![X10]:plus(X9,X10)=plus(X10,X9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', symmetry_of_plus)).
fof(c_0_3, axiom, (plus(n0,n1)=n1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', plus0_1)).
fof(c_0_4, axiom, (![X9]:![X10]:(less_or_equal(X9,X10)<=>(less(X9,X10)|X9=X10))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', less_or_equal)).
fof(c_0_5, axiom, (![X4]:![X2]:![X5]:(initiates(X4,X2,X5)<=>((((X4=push&X2=forwards)&~(happens(pull,X5)))|((X4=pull&X2=backwards)&~(happens(push,X5))))|((X4=pull&X2=spinning)&happens(push,X5))))), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+2.ax', initiates_all_defn)).
fof(c_0_6, conjecture, (~(holdsAt(spinning,n1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', not_spinning_1)).
fof(c_0_7, hypothesis, (~(holdsAt(spinning,n0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', not_splinning_0)).
fof(c_0_8, axiom, (![X9]:(less(X9,n1)<=>less_or_equal(X9,n0))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', less1)).
fof(c_0_9, axiom, (![X4]:![X5]:(happens(X4,X5)<=>((((X4=push&X5=n0)|(X4=pull&X5=n1))|(X4=pull&X5=n2))|(X4=push&X5=n2)))), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+2.ax', happens_all_defn)).
fof(c_0_10, axiom, (![X9]:![X10]:(less(X9,X10)<=>(~(less(X10,X9))&X10!=X9))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', less_property)).
fof(c_0_11, axiom, (![X9]:(less(X9,n2)<=>less_or_equal(X9,n1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', less2)).
fof(c_0_12, axiom, (forwards!=spinning), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+2.ax', forwards_not_spinning)).
fof(c_0_13, axiom, (push!=pull), file('/export/starexec/sandbox/benchmark/Axioms/CSR001+2.ax', push_not_pull)).
fof(c_0_14, plain, (![X2]:![X5]:(((~holdsAt(X2,X5)&~releasedAt(X2,plus(X5,n1)))&~(?[X4]:(happens(X4,X5)&initiates(X4,X2,X5))))=>~holdsAt(X2,plus(X5,n1)))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_15, hypothesis, (![X2]:![X5]:~releasedAt(X2,X5)), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_16, plain, (![X6]:![X7]:(((happens(esk6_2(X6,X7),X7)|(holdsAt(X6,X7)|releasedAt(X6,plus(X7,n1))))|~holdsAt(X6,plus(X7,n1)))&((initiates(esk6_2(X6,X7),X6,X7)|(holdsAt(X6,X7)|releasedAt(X6,plus(X7,n1))))|~holdsAt(X6,plus(X7,n1))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).
fof(c_0_17, hypothesis, (![X6]:![X7]:~releasedAt(X6,X7)), inference(variable_rename,[status(thm)],[c_0_15])).
cnf(c_0_18,plain,(releasedAt(X1,plus(X2,n1))|holdsAt(X1,X2)|initiates(esk6_2(X1,X2),X1,X2)|~holdsAt(X1,plus(X2,n1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_19,hypothesis,(~releasedAt(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_20, plain, (![X11]:![X12]:plus(X11,X12)=plus(X12,X11)), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_21,plain,(releasedAt(X1,plus(X2,n1))|holdsAt(X1,X2)|happens(esk6_2(X1,X2),X2)|~holdsAt(X1,plus(X2,n1))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, (![X2]:![X4]:![X5]:(epred1_3(X5,X4,X2)<=>((((X4=push&X2=forwards)&~happens(pull,X5))|((X4=pull&X2=backwards)&~happens(push,X5)))|((X4=pull&X2=spinning)&happens(push,X5))))), introduced(definition)).
cnf(c_0_23,plain,(holdsAt(X1,X2)|initiates(esk6_2(X1,X2),X1,X2)|~holdsAt(X1,plus(X2,n1))), inference(sr,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24,plain,(plus(X1,X2)=plus(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25,plain,(plus(n0,n1)=n1), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_26, plain, (![X11]:![X12]:![X13]:![X14]:((~less_or_equal(X11,X12)|(less(X11,X12)|X11=X12))&((~less(X13,X14)|less_or_equal(X13,X14))&(X13!=X14|less_or_equal(X13,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
cnf(c_0_27,plain,(holdsAt(X1,X2)|happens(esk6_2(X1,X2),X2)|~holdsAt(X1,plus(X2,n1))), inference(sr,[status(thm)],[c_0_21, c_0_19])).
fof(c_0_28, plain, (![X4]:![X2]:![X5]:(initiates(X4,X2,X5)<=>epred1_3(X5,X4,X2))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5]), c_0_22])).
cnf(c_0_29,plain,(holdsAt(X1,X2)|initiates(esk6_2(X1,X2),X1,X2)|~holdsAt(X1,plus(n1,X2))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30,plain,(plus(n1,n0)=n1), inference(rw,[status(thm)],[c_0_25, c_0_24])).
fof(c_0_31, negated_conjecture, (holdsAt(spinning,n1)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_6])])).
fof(c_0_32, hypothesis, (~holdsAt(spinning,n0)), inference(fof_simplification,[status(thm)],[c_0_7])).
fof(c_0_33, plain, (![X10]:![X11]:((~less(X10,n1)|less_or_equal(X10,n0))&(~less_or_equal(X11,n0)|less(X11,n1)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
cnf(c_0_34,plain,(less_or_equal(X1,X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_35, plain, (![X6]:![X7]:![X8]:![X9]:(((((((X6=push|(X6=pull|(X6=pull|X6=push)))|~happens(X6,X7))&((X7=n2|(X6=pull|(X6=pull|X6=push)))|~happens(X6,X7)))&(((X6=push|(X7=n2|(X6=pull|X6=push)))|~happens(X6,X7))&((X7=n2|(X7=n2|(X6=pull|X6=push)))|~happens(X6,X7))))&((((X6=push|(X6=pull|(X7=n1|X6=push)))|~happens(X6,X7))&((X7=n2|(X6=pull|(X7=n1|X6=push)))|~happens(X6,X7)))&(((X6=push|(X7=n2|(X7=n1|X6=push)))|~happens(X6,X7))&((X7=n2|(X7=n2|(X7=n1|X6=push)))|~happens(X6,X7)))))&(((((X6=push|(X6=pull|(X6=pull|X7=n0)))|~happens(X6,X7))&((X7=n2|(X6=pull|(X6=pull|X7=n0)))|~happens(X6,X7)))&(((X6=push|(X7=n2|(X6=pull|X7=n0)))|~happens(X6,X7))&((X7=n2|(X7=n2|(X6=pull|X7=n0)))|~happens(X6,X7))))&((((X6=push|(X6=pull|(X7=n1|X7=n0)))|~happens(X6,X7))&((X7=n2|(X6=pull|(X7=n1|X7=n0)))|~happens(X6,X7)))&(((X6=push|(X7=n2|(X7=n1|X7=n0)))|~happens(X6,X7))&((X7=n2|(X7=n2|(X7=n1|X7=n0)))|~happens(X6,X7))))))&(((((X8!=push|X9!=n0)|happens(X8,X9))&((X8!=pull|X9!=n1)|happens(X8,X9)))&((X8!=pull|X9!=n2)|happens(X8,X9)))&((X8!=push|X9!=n2)|happens(X8,X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
cnf(c_0_36,plain,(holdsAt(X1,X2)|happens(esk6_2(X1,X2),X2)|~holdsAt(X1,plus(n1,X2))), inference(spm,[status(thm)],[c_0_27, c_0_24])).
fof(c_0_37, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:(((((((((X7=pull|(X7=pull|X7=push))|~epred1_3(X8,X7,X6))&((X6=spinning|(X7=pull|X7=push))|~epred1_3(X8,X7,X6)))&((happens(push,X8)|(X7=pull|X7=push))|~epred1_3(X8,X7,X6)))&((((X7=pull|(X6=backwards|X7=push))|~epred1_3(X8,X7,X6))&((X6=spinning|(X6=backwards|X7=push))|~epred1_3(X8,X7,X6)))&((happens(push,X8)|(X6=backwards|X7=push))|~epred1_3(X8,X7,X6))))&((((X7=pull|(~happens(push,X8)|X7=push))|~epred1_3(X8,X7,X6))&((X6=spinning|(~happens(push,X8)|X7=push))|~epred1_3(X8,X7,X6)))&((happens(push,X8)|(~happens(push,X8)|X7=push))|~epred1_3(X8,X7,X6))))&((((((X7=pull|(X7=pull|X6=forwards))|~epred1_3(X8,X7,X6))&((X6=spinning|(X7=pull|X6=forwards))|~epred1_3(X8,X7,X6)))&((happens(push,X8)|(X7=pull|X6=forwards))|~epred1_3(X8,X7,X6)))&((((X7=pull|(X6=backwards|X6=forwards))|~epred1_3(X8,X7,X6))&((X6=spinning|(X6=backwards|X6=forwards))|~epred1_3(X8,X7,X6)))&((happens(push,X8)|(X6=backwards|X6=forwards))|~epred1_3(X8,X7,X6))))&((((X7=pull|(~happens(push,X8)|X6=forwards))|~epred1_3(X8,X7,X6))&((X6=spinning|(~happens(push,X8)|X6=forwards))|~epred1_3(X8,X7,X6)))&((happens(push,X8)|(~happens(push,X8)|X6=forwards))|~epred1_3(X8,X7,X6)))))&((((((X7=pull|(X7=pull|~happens(pull,X8)))|~epred1_3(X8,X7,X6))&((X6=spinning|(X7=pull|~happens(pull,X8)))|~epred1_3(X8,X7,X6)))&((happens(push,X8)|(X7=pull|~happens(pull,X8)))|~epred1_3(X8,X7,X6)))&((((X7=pull|(X6=backwards|~happens(pull,X8)))|~epred1_3(X8,X7,X6))&((X6=spinning|(X6=backwards|~happens(pull,X8)))|~epred1_3(X8,X7,X6)))&((happens(push,X8)|(X6=backwards|~happens(pull,X8)))|~epred1_3(X8,X7,X6))))&((((X7=pull|(~happens(push,X8)|~happens(pull,X8)))|~epred1_3(X8,X7,X6))&((X6=spinning|(~happens(push,X8)|~happens(pull,X8)))|~epred1_3(X8,X7,X6)))&((happens(push,X8)|(~happens(push,X8)|~happens(pull,X8)))|~epred1_3(X8,X7,X6)))))&(((((X10!=push|X9!=forwards)|happens(pull,X11))|epred1_3(X11,X10,X9))&(((X10!=pull|X9!=backwards)|happens(push,X11))|epred1_3(X11,X10,X9)))&(((X10!=pull|X9!=spinning)|~happens(push,X11))|epred1_3(X11,X10,X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])).
fof(c_0_38, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:![X11]:((~initiates(X6,X7,X8)|epred1_3(X8,X6,X7))&(~epred1_3(X11,X9,X10)|initiates(X9,X10,X11)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])])).
cnf(c_0_39,plain,(holdsAt(X1,n0)|initiates(esk6_2(X1,n0),X1,n0)|~holdsAt(X1,n1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_40,negated_conjecture,(holdsAt(spinning,n1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41,hypothesis,(~holdsAt(spinning,n0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42,plain,(less(X1,n1)|~less_or_equal(X1,n0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43,plain,(less_or_equal(X1,X1)), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_44,plain,(X1=push|X2=n1|X2=n2|X2=n2|~happens(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45,plain,(holdsAt(X1,n0)|happens(esk6_2(X1,n0),n0)|~holdsAt(X1,n1)), inference(spm,[status(thm)],[c_0_36, c_0_30])).
cnf(c_0_46,plain,(X3=forwards|X2=pull|X2=pull|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47,plain,(epred1_3(X1,X2,X3)|~initiates(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48,negated_conjecture,(initiates(esk6_2(spinning,n0),spinning,n0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
fof(c_0_49, plain, (![X9]:![X10]:(less(X9,X10)<=>(~less(X10,X9)&X10!=X9))), inference(fof_simplification,[status(thm)],[c_0_10])).
fof(c_0_50, plain, (![X10]:![X11]:((~less(X10,n2)|less_or_equal(X10,n1))&(~less_or_equal(X11,n1)|less(X11,n2)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
cnf(c_0_51,plain,(less_or_equal(X1,X2)|~less(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_52,plain,(less(n0,n1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_53,plain,(n1=X2|push=X1|n2=X2|~happens(X1,X2)), inference(cn,[status(thm)],[c_0_44])).
cnf(c_0_54,negated_conjecture,(happens(esk6_2(spinning,n0),n0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_40]), c_0_41])).
cnf(c_0_55,plain,(forwards=X3|pull=X2|~epred1_3(X1,X2,X3)), inference(cn,[status(thm)],[c_0_46])).
cnf(c_0_56,negated_conjecture,(epred1_3(n0,esk6_2(spinning,n0),spinning)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_57,plain,(forwards!=spinning), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_58, plain, (![X11]:![X12]:![X13]:![X14]:(((~less(X12,X11)|~less(X11,X12))&(X12!=X11|~less(X11,X12)))&((less(X14,X13)|X14=X13)|less(X13,X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_49])])])])])).
cnf(c_0_59,plain,(less(X1,n2)|~less_or_equal(X1,n1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_60,plain,(less_or_equal(n0,n1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_61,negated_conjecture,(esk6_2(spinning,n0)=push|n0=n1|n0=n2), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62,plain,(esk6_2(spinning,n0)=pull), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
cnf(c_0_63,plain,(push!=pull), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_64,plain,(~less(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65,plain,(less(n0,n2)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66,negated_conjecture,(n0=n2|n0=n1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).
cnf(c_0_67,plain,(~less(X1,X1)), inference(er,[status(thm)],[c_0_64])).
cnf(c_0_68,negated_conjecture,(n0=n1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_69,plain,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_68]), c_0_67]), ['proof']).
# SZS output end CNFRefutation.
EOF
