% Problem    : LCL461+1 : TPTP v6.1.0. Released v3.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n104.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:57:31 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____011_C18_F1_PI_AE_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, ((modus_ponens<=>![X1]:![X2]:((is_a_theorem(X1)&is_a_theorem(implies(X1,X2)))=>is_a_theorem(X2)))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', modus_ponens)).
fof(c_0_1, axiom, ((and_3<=>![X1]:![X2]:is_a_theorem(implies(X1,implies(X2,and(X1,X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', and_3)).
fof(c_0_2, axiom, (modus_ponens), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_modus_ponens)).
fof(c_0_3, axiom, (and_3), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_and_3)).
fof(c_0_4, axiom, ((implies_2<=>![X1]:![X2]:is_a_theorem(implies(implies(X1,implies(X1,X2)),implies(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', implies_2)).
fof(c_0_5, axiom, ((op_equiv=>![X1]:![X2]:equiv(X1,X2)=and(implies(X1,X2),implies(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+1.ax', op_equiv)).
fof(c_0_6, axiom, (implies_2), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_implies_2)).
fof(c_0_7, axiom, (op_equiv), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_op_equiv)).
fof(c_0_8, axiom, ((and_1<=>![X1]:![X2]:is_a_theorem(implies(and(X1,X2),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', and_1)).
fof(c_0_9, axiom, ((substitution_of_equivalents<=>![X1]:![X2]:(is_a_theorem(equiv(X1,X2))=>X1=X2))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', substitution_of_equivalents)).
fof(c_0_10, axiom, (and_1), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_and_1)).
fof(c_0_11, axiom, ((op_implies_and=>![X1]:![X2]:implies(X1,X2)=not(and(X1,not(X2))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+1.ax', op_implies_and)).
fof(c_0_12, axiom, (substitution_of_equivalents), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', substitution_of_equivalents)).
fof(c_0_13, axiom, ((op_or=>![X1]:![X2]:or(X1,X2)=not(and(not(X1),not(X2))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+1.ax', op_or)).
fof(c_0_14, axiom, (op_implies_and), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_op_implies_and)).
fof(c_0_15, axiom, ((or_1<=>![X1]:![X2]:is_a_theorem(implies(X1,or(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', or_1)).
fof(c_0_16, axiom, (op_or), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_op_or)).
fof(c_0_17, axiom, ((modus_tollens<=>![X1]:![X2]:is_a_theorem(implies(implies(not(X2),not(X1)),implies(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', modus_tollens)).
fof(c_0_18, axiom, (or_1), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_or_1)).
fof(c_0_19, axiom, (modus_tollens), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_modus_tollens)).
fof(c_0_20, axiom, ((kn3<=>![X4]:![X5]:![X6]:is_a_theorem(implies(implies(X4,X5),implies(not(and(X5,X6)),not(and(X6,X4))))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', kn3)).
fof(c_0_21, conjecture, (kn3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rosser_kn3)).
fof(c_0_22, axiom, ((implies_3<=>![X1]:![X2]:![X3]:is_a_theorem(implies(implies(X1,X2),implies(implies(X2,X3),implies(X1,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+0.ax', implies_3)).
fof(c_0_23, axiom, (implies_3), file('/export/starexec/sandbox/benchmark/Axioms/LCL006+2.ax', hilbert_implies_3)).
fof(c_0_24, plain, (![X3]:![X4]:((~modus_ponens|((~is_a_theorem(X3)|~is_a_theorem(implies(X3,X4)))|is_a_theorem(X4)))&(((is_a_theorem(esk1_0)|modus_ponens)&(is_a_theorem(implies(esk1_0,esk2_0))|modus_ponens))&(~is_a_theorem(esk2_0)|modus_ponens)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])).
fof(c_0_25, plain, (![X3]:![X4]:((~and_3|is_a_theorem(implies(X3,implies(X4,and(X3,X4)))))&(~is_a_theorem(implies(esk18_0,implies(esk19_0,and(esk18_0,esk19_0))))|and_3))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
cnf(c_0_26,plain,(is_a_theorem(X1)|~is_a_theorem(implies(X2,X1))|~is_a_theorem(X2)|~modus_ponens), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_27,plain,(modus_ponens), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_28,plain,(is_a_theorem(implies(X1,implies(X2,and(X1,X2))))|~and_3), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29,plain,(and_3), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_30, plain, (![X3]:![X4]:((~implies_2|is_a_theorem(implies(implies(X3,implies(X3,X4)),implies(X3,X4))))&(~is_a_theorem(implies(implies(esk9_0,implies(esk9_0,esk10_0)),implies(esk9_0,esk10_0)))|implies_2))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_31,plain,(is_a_theorem(X1)|~is_a_theorem(implies(X2,X1))|~is_a_theorem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27])])).
cnf(c_0_32,plain,(is_a_theorem(implies(X1,implies(X2,and(X1,X2))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29])])).
fof(c_0_33, plain, (![X3]:![X4]:(~op_equiv|equiv(X3,X4)=and(implies(X3,X4),implies(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_34,plain,(is_a_theorem(implies(implies(X1,implies(X1,X2)),implies(X1,X2)))|~implies_2), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_35,plain,(implies_2), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_36,plain,(is_a_theorem(implies(X1,and(X2,X1)))|~is_a_theorem(X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37,plain,(equiv(X1,X2)=and(implies(X1,X2),implies(X2,X1))|~op_equiv), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_38,plain,(op_equiv), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_39,plain,(is_a_theorem(implies(implies(X1,implies(X1,X2)),implies(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35])])).
fof(c_0_40, plain, (![X3]:![X4]:((~and_1|is_a_theorem(implies(and(X3,X4),X3)))&(~is_a_theorem(implies(and(esk14_0,esk15_0),esk14_0))|and_1))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
fof(c_0_41, plain, (![X3]:![X4]:((~substitution_of_equivalents|(~is_a_theorem(equiv(X3,X4))|X3=X4))&((is_a_theorem(equiv(esk3_0,esk4_0))|substitution_of_equivalents)&(esk3_0!=esk4_0|substitution_of_equivalents)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
cnf(c_0_42,plain,(is_a_theorem(and(X1,X2))|~is_a_theorem(X2)|~is_a_theorem(X1)), inference(spm,[status(thm)],[c_0_31, c_0_36])).
cnf(c_0_43,plain,(and(implies(X1,X2),implies(X2,X1))=equiv(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38])])).
cnf(c_0_44,plain,(is_a_theorem(implies(X1,X2))|~is_a_theorem(implies(X1,implies(X1,X2)))), inference(spm,[status(thm)],[c_0_31, c_0_39])).
cnf(c_0_45,plain,(is_a_theorem(implies(and(X1,X2),X1))|~and_1), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46,plain,(and_1), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_47, plain, (![X3]:![X4]:(~op_implies_and|implies(X3,X4)=not(and(X3,not(X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_48,plain,(X1=X2|~is_a_theorem(equiv(X1,X2))|~substitution_of_equivalents), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49,plain,(substitution_of_equivalents), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_50,plain,(is_a_theorem(equiv(X1,X2))|~is_a_theorem(implies(X2,X1))|~is_a_theorem(implies(X1,X2))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_51,plain,(is_a_theorem(implies(X1,and(X1,X1)))), inference(spm,[status(thm)],[c_0_44, c_0_32])).
cnf(c_0_52,plain,(is_a_theorem(implies(and(X1,X2),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])])).
fof(c_0_53, plain, (![X3]:![X4]:(~op_or|or(X3,X4)=not(and(not(X3),not(X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
cnf(c_0_54,plain,(implies(X1,X2)=not(and(X1,not(X2)))|~op_implies_and), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_55,plain,(op_implies_and), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_56, plain, (![X3]:![X4]:((~or_1|is_a_theorem(implies(X3,or(X3,X4))))&(~is_a_theorem(implies(esk20_0,or(esk20_0,esk21_0)))|or_1))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])).
cnf(c_0_57,plain,(X1=X2|~is_a_theorem(equiv(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49])])).
cnf(c_0_58,plain,(is_a_theorem(equiv(and(X1,X1),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])])).
cnf(c_0_59,plain,(or(X1,X2)=not(and(not(X1),not(X2)))|~op_or), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60,plain,(not(and(X1,not(X2)))=implies(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55])])).
cnf(c_0_61,plain,(op_or), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_62, plain, (![X3]:![X4]:((~modus_tollens|is_a_theorem(implies(implies(not(X4),not(X3)),implies(X3,X4))))&(~is_a_theorem(implies(implies(not(esk6_0),not(esk5_0)),implies(esk5_0,esk6_0)))|modus_tollens))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])).
cnf(c_0_63,plain,(is_a_theorem(implies(X1,or(X1,X2)))|~or_1), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64,plain,(or_1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_65,plain,(and(X1,X1)=X1), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_66,plain,(implies(not(X1),X2)=or(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_60]), c_0_61])])).
cnf(c_0_67,plain,(is_a_theorem(implies(implies(not(X1),not(X2)),implies(X2,X1)))|~modus_tollens), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_68,plain,(modus_tollens), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_69,plain,(is_a_theorem(implies(X1,or(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_70,plain,(not(not(X1))=or(X1,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_65]), c_0_66])).
cnf(c_0_71,plain,(is_a_theorem(implies(or(X1,not(X2)),implies(X2,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_66]), c_0_68])])).
cnf(c_0_72,plain,(is_a_theorem(implies(X1,not(not(X1))))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_73,plain,(is_a_theorem(implies(X1,X2))|~is_a_theorem(or(X2,not(X1)))), inference(spm,[status(thm)],[c_0_31, c_0_71])).
cnf(c_0_74,plain,(is_a_theorem(or(X1,not(or(X1,X1))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_70]), c_0_66])).
cnf(c_0_75,plain,(is_a_theorem(implies(or(X1,X1),X1))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_76,plain,(is_a_theorem(equiv(X1,or(X1,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_75]), c_0_69])])).
cnf(c_0_77,plain,(or(X1,X1)=X1), inference(spm,[status(thm)],[c_0_57, c_0_76])).
cnf(c_0_78,plain,(is_a_theorem(implies(or(X1,not(not(X2))),or(X2,X1)))), inference(spm,[status(thm)],[c_0_71, c_0_66])).
cnf(c_0_79,plain,(not(not(X1))=X1), inference(rw,[status(thm)],[c_0_70, c_0_77])).
fof(c_0_80, plain, (![X7]:![X8]:![X9]:((~kn3|is_a_theorem(implies(implies(X7,X8),implies(not(and(X8,X9)),not(and(X9,X7))))))&(~is_a_theorem(implies(implies(esk36_0,esk37_0),implies(not(and(esk37_0,esk38_0)),not(and(esk38_0,esk36_0)))))|kn3))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).
fof(c_0_81, negated_conjecture, (~kn3), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_21])])).
cnf(c_0_82,plain,(is_a_theorem(implies(or(X1,X2),or(X2,X1)))), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_83,plain,(kn3|~is_a_theorem(implies(implies(esk36_0,esk37_0),implies(not(and(esk37_0,esk38_0)),not(and(esk38_0,esk36_0)))))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_84,negated_conjecture,(~kn3), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_85,plain,(is_a_theorem(equiv(or(X1,X2),or(X2,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_82]), c_0_82])])).
cnf(c_0_86,plain,(~is_a_theorem(implies(implies(esk36_0,esk37_0),or(and(esk37_0,esk38_0),not(and(esk38_0,esk36_0)))))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_66]), c_0_84])).
cnf(c_0_87,plain,(not(and(X1,X2))=implies(X1,not(X2))), inference(spm,[status(thm)],[c_0_60, c_0_79])).
cnf(c_0_88,plain,(or(not(X1),X2)=implies(X1,X2)), inference(spm,[status(thm)],[c_0_66, c_0_79])).
cnf(c_0_89,plain,(or(X1,X2)=or(X2,X1)), inference(spm,[status(thm)],[c_0_57, c_0_85])).
cnf(c_0_90,plain,(~is_a_theorem(implies(implies(esk36_0,esk37_0),or(and(esk37_0,esk38_0),implies(esk38_0,not(esk36_0)))))), inference(rw,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_91,plain,(or(X1,not(X2))=implies(X2,X1)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_92,plain,(~is_a_theorem(implies(implies(esk36_0,esk37_0),or(implies(esk38_0,not(esk36_0)),and(esk37_0,esk38_0))))), inference(rw,[status(thm)],[c_0_90, c_0_89])).
cnf(c_0_93,plain,(implies(X1,not(X2))=implies(X2,not(X1))), inference(spm,[status(thm)],[c_0_88, c_0_91])).
cnf(c_0_94,plain,(not(implies(X1,not(X2)))=and(X1,X2)), inference(spm,[status(thm)],[c_0_79, c_0_87])).
fof(c_0_95, plain, (![X4]:![X5]:![X6]:((~implies_3|is_a_theorem(implies(implies(X4,X5),implies(implies(X5,X6),implies(X4,X6)))))&(~is_a_theorem(implies(implies(esk11_0,esk12_0),implies(implies(esk12_0,esk13_0),implies(esk11_0,esk13_0))))|implies_3))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).
cnf(c_0_96,plain,(~is_a_theorem(implies(implies(esk36_0,esk37_0),or(implies(esk36_0,not(esk38_0)),and(esk37_0,esk38_0))))), inference(rw,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_97,plain,(or(implies(X1,not(X2)),X3)=implies(and(X1,X2),X3)), inference(spm,[status(thm)],[c_0_88, c_0_87])).
cnf(c_0_98,plain,(or(X1,and(X2,X3))=implies(implies(X2,not(X3)),X1)), inference(spm,[status(thm)],[c_0_91, c_0_94])).
cnf(c_0_99,plain,(is_a_theorem(implies(implies(X1,X2),implies(implies(X2,X3),implies(X1,X3))))|~implies_3), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_100,plain,(implies_3), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_101,plain,(~is_a_theorem(implies(implies(esk36_0,esk37_0),implies(and(esk36_0,esk38_0),and(esk37_0,esk38_0))))), inference(rw,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102,plain,(implies(implies(X1,not(X2)),implies(X3,not(X4)))=implies(and(X3,X4),and(X1,X2))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_103,plain,(is_a_theorem(implies(implies(X1,X2),implies(implies(X2,X3),implies(X1,X3))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_104,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102]), c_0_103])]), ['proof']).
# SZS output end CNFRefutation.
EOF
