% Problem    : SWV174+1 : TPTP v6.1.0. Bugfixed v3.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n144.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:37:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (((((leq(n0,pv10)&leq(pv10,n135299))&![X14]:((leq(n0,X14)&leq(X14,pred(pv10)))=>![X18]:((leq(n0,X18)&leq(X18,n4))=>a_select3(q_init,X14,X18)=init)))&![X4]:((leq(n0,X4)&leq(X4,n4))=>a_select3(center_init,X4,n0)=init))=>![X20]:![X21]:((((leq(n0,X20)&leq(n0,X21))&leq(X20,n135299))&leq(X21,n4))=>(gt(pv10,X20)=>a_select3(q_init,X20,X21)=init)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', cl5_nebula_init_0046)).
fof(c_0_1, axiom, (![X1]:~(gt(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', irreflexivity_gt)).
fof(c_0_2, axiom, (![X1]:![X2]:(leq(succ(X1),X2)=>gt(X2,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', leq_succ_gt)).
fof(c_0_3, axiom, (![X1]:plus(X1,n1)=succ(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', succ_plus_1_r)).
fof(c_0_4, axiom, (![X1]:![X2]:(leq(X1,pred(X2))<=>gt(X2,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', leq_gt_pred)).
fof(c_0_5, axiom, (![X1]:minus(X1,n1)=pred(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', pred_minus_1)).
fof(c_0_6, axiom, (![X1]:plus(n1,X1)=succ(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', succ_plus_1_l)).
fof(c_0_7, axiom, (![X1]:![X2]:![X3]:((leq(X1,X2)&leq(X2,X3))=>leq(X1,X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', transitivity_leq)).
fof(c_0_8, axiom, (![X1]:succ(pred(X1))=X1), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', succ_pred)).
fof(c_0_9, axiom, (![X1]:![X2]:((leq(X1,X2)&X1!=X2)=>gt(X2,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', leq_gt2)).
fof(c_0_10, axiom, (![X1]:![X2]:(gt(X2,X1)=>leq(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', leq_gt1)).
fof(c_0_11, axiom, (![X1]:leq(X1,X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWV003+0.ax', reflexivity_leq)).
fof(c_0_12, negated_conjecture, (~(((((leq(n0,pv10)&leq(pv10,n135299))&![X14]:((leq(n0,X14)&leq(X14,pred(pv10)))=>![X18]:((leq(n0,X18)&leq(X18,n4))=>a_select3(q_init,X14,X18)=init)))&![X4]:((leq(n0,X4)&leq(X4,n4))=>a_select3(center_init,X4,n0)=init))=>![X20]:![X21]:((((leq(n0,X20)&leq(n0,X21))&leq(X20,n135299))&leq(X21,n4))=>(gt(pv10,X20)=>a_select3(q_init,X20,X21)=init))))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_13, plain, (![X1]:~gt(X1,X1)), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_14, plain, (![X3]:![X4]:(~leq(succ(X3),X4)|gt(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
fof(c_0_15, plain, (![X2]:plus(X2,n1)=succ(X2)), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_16, plain, (![X3]:![X4]:![X5]:![X6]:((~leq(X3,pred(X4))|gt(X4,X3))&(~gt(X6,X5)|leq(X5,pred(X6))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
fof(c_0_17, plain, (![X2]:minus(X2,n1)=pred(X2)), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_18, negated_conjecture, (![X22]:![X23]:![X24]:((((leq(n0,pv10)&leq(pv10,n135299))&((~leq(n0,X22)|~leq(X22,pred(pv10)))|((~leq(n0,X23)|~leq(X23,n4))|a_select3(q_init,X22,X23)=init)))&((~leq(n0,X24)|~leq(X24,n4))|a_select3(center_init,X24,n0)=init))&((((leq(n0,esk24_0)&leq(n0,esk25_0))&leq(esk24_0,n135299))&leq(esk25_0,n4))&(gt(pv10,esk24_0)&a_select3(q_init,esk24_0,esk25_0)!=init)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
fof(c_0_19, plain, (![X2]:~gt(X2,X2)), inference(variable_rename,[status(thm)],[c_0_13])).
cnf(c_0_20,plain,(gt(X1,X2)|~leq(succ(X2),X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21,plain,(plus(X1,n1)=succ(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_22, plain, (![X2]:plus(n1,X2)=succ(X2)), inference(variable_rename,[status(thm)],[c_0_6])).
fof(c_0_23, plain, (![X4]:![X5]:![X6]:((~leq(X4,X5)|~leq(X5,X6))|leq(X4,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
cnf(c_0_24,plain,(leq(X1,pred(X2))|~gt(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25,plain,(minus(X1,n1)=pred(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_26, plain, (![X2]:succ(pred(X2))=X2), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_27,negated_conjecture,(a_select3(q_init,X1,X2)=init|~leq(X2,n4)|~leq(n0,X2)|~leq(X1,pred(pv10))|~leq(n0,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28,plain,(~gt(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29,plain,(gt(X1,X2)|~leq(plus(X2,n1),X1)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_30,plain,(plus(n1,X1)=succ(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31,plain,(leq(X1,X2)|~leq(X3,X2)|~leq(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32,plain,(leq(X1,minus(X2,n1))|~gt(X2,X1)), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33,plain,(succ(pred(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34,negated_conjecture,(a_select3(q_init,X1,X2)=init|~leq(X2,n4)|~leq(n0,X2)|~leq(n0,X1)|~leq(X1,minus(pv10,n1))), inference(rw,[status(thm)],[c_0_27, c_0_25])).
fof(c_0_35, plain, (![X3]:![X4]:((~leq(X3,X4)|X3=X4)|gt(X4,X3))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
fof(c_0_36, plain, (![X3]:![X4]:(~gt(X4,X3)|leq(X3,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
cnf(c_0_37,plain,(~leq(plus(X1,n1),X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_38,plain,(plus(n1,X1)=plus(X1,n1)), inference(rw,[status(thm)],[c_0_30, c_0_21])).
cnf(c_0_39,plain,(leq(X1,minus(X2,n1))|~leq(X1,X3)|~gt(X2,X3)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40,negated_conjecture,(gt(pv10,esk24_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_41,plain,(plus(minus(X1,n1),n1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_21]), c_0_25])).
cnf(c_0_42,negated_conjecture,(a_select3(q_init,X1,X2)=init|~leq(X2,n4)|~leq(n0,X2)|~leq(n0,X1)|~gt(pv10,X1)), inference(spm,[status(thm)],[c_0_34, c_0_32])).
cnf(c_0_43,plain,(gt(X1,X2)|X2=X1|~leq(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44,plain,(leq(X1,X2)|~gt(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45,plain,(~leq(plus(n1,X1),X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46,negated_conjecture,(leq(X1,minus(pv10,n1))|~leq(X1,esk24_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_47,plain,(plus(n1,minus(X1,n1))=X1), inference(rw,[status(thm)],[c_0_41, c_0_38])).
cnf(c_0_48,negated_conjecture,(a_select3(q_init,esk24_0,esk25_0)!=init), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_49,negated_conjecture,(a_select3(q_init,X1,X2)=init|pv10=X1|~leq(X2,n4)|~leq(n0,X2)|~leq(n0,X1)|~leq(X1,pv10)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50,negated_conjecture,(leq(esk25_0,n4)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_51,negated_conjecture,(leq(n0,esk25_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_52,negated_conjecture,(leq(n0,esk24_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_53,negated_conjecture,(leq(esk24_0,pv10)), inference(spm,[status(thm)],[c_0_44, c_0_40])).
fof(c_0_54, plain, (![X2]:leq(X2,X2)), inference(variable_rename,[status(thm)],[c_0_11])).
cnf(c_0_55,negated_conjecture,(~leq(pv10,esk24_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])).
cnf(c_0_56,negated_conjecture,(esk24_0=pv10), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_53])])).
cnf(c_0_57,plain,(leq(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_58,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56]), c_0_57])]), ['proof']).
# SZS output end CNFRefutation.
EOF
