% Problem    : SWB010+2 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n140.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:28:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C48_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X7]:![X6]:![X8]:((((ir(X6)&ip(X7))&ir(X8))&~(iext(X7,X6,X8)))=>?[X4]:((iext(uri_owl_sourceIndividual,X4,X6)&iext(uri_owl_assertionProperty,X4,X7))&iext(uri_owl_targetIndividual,X4,X8)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', owl_npa_object_fi)).
fof(c_0_1, conjecture, (?[X9]:(((iext(uri_rdf_type,X9,uri_owl_NegativePropertyAssertion)&iext(uri_owl_sourceIndividual,X9,uri_ex_s))&iext(uri_owl_assertionProperty,X9,uri_ex_p))&iext(uri_owl_targetIndividual,X9,uri_ex_o))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_conclusion_fullish_010_Negative_Property_Assertions)).
fof(c_0_2, axiom, (![X1]:ir(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', simple_ir)).
fof(c_0_3, axiom, (![X1]:![X3]:(iext(uri_owl_onProperty,X1,X3)=>(icext(uri_owl_Restriction,X1)&ip(X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', owl_prop_onproperty_ext)).
fof(c_0_4, axiom, (?[X10]:?[X11]:?[X12]:?[X13]:(((((((iext(uri_rdf_type,uri_ex_p,uri_owl_ObjectProperty)&iext(uri_rdf_type,uri_ex_s,X10))&iext(uri_owl_onProperty,X10,uri_ex_p))&iext(uri_owl_allValuesFrom,X10,X11))&iext(uri_owl_complementOf,X11,X12))&iext(uri_owl_oneOf,X12,X13))&iext(uri_rdf_first,X13,uri_ex_o))&iext(uri_rdf_rest,X13,uri_rdf_nil))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', testcase_premise_fullish_010_Negative_Property_Assertions)).
fof(c_0_5, axiom, (![X1]:![X3]:(iext(uri_owl_sourceIndividual,X1,X3)=>(icext(uri_owl_NegativePropertyAssertion,X1)&ir(X3)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', owl_prop_sourceindividual_ext)).
fof(c_0_6, axiom, (![X1]:![X2]:(iext(uri_rdf_type,X1,X2)<=>icext(X2,X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', rdfs_cext_def)).
fof(c_0_7, axiom, (![X4]:![X2]:(iext(uri_owl_complementOf,X4,X2)=>((ic(X4)&ic(X2))&![X1]:(icext(X4,X1)<=>~(icext(X2,X1)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', owl_bool_complementof_class)).
fof(c_0_8, axiom, (![X4]:![X5]:![X6]:((iext(uri_rdf_first,X5,X6)&iext(uri_rdf_rest,X5,uri_rdf_nil))=>(iext(uri_owl_oneOf,X4,X5)<=>(ic(X4)&![X1]:(icext(X4,X1)<=>X1=X6))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', owl_enum_class_001)).
fof(c_0_9, axiom, (![X4]:![X7]:![X2]:((iext(uri_owl_allValuesFrom,X4,X2)&iext(uri_owl_onProperty,X4,X7))=>![X1]:(icext(X4,X1)<=>![X3]:(iext(X7,X1,X3)=>icext(X2,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', owl_restrict_allvaluesfrom)).
fof(c_0_10, plain, (![X7]:![X6]:![X8]:((((ir(X6)&ip(X7))&ir(X8))&~iext(X7,X6,X8))=>?[X4]:((iext(uri_owl_sourceIndividual,X4,X6)&iext(uri_owl_assertionProperty,X4,X7))&iext(uri_owl_targetIndividual,X4,X8)))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_11, negated_conjecture, (~(?[X9]:(((iext(uri_rdf_type,X9,uri_owl_NegativePropertyAssertion)&iext(uri_owl_sourceIndividual,X9,uri_ex_s))&iext(uri_owl_assertionProperty,X9,uri_ex_p))&iext(uri_owl_targetIndividual,X9,uri_ex_o)))), inference(assume_negation,[status(cth)],[c_0_1])).
fof(c_0_12, plain, (![X9]:![X10]:![X11]:(((iext(uri_owl_sourceIndividual,esk3_3(X9,X10,X11),X10)|(((~ir(X10)|~ip(X9))|~ir(X11))|iext(X9,X10,X11)))&(iext(uri_owl_assertionProperty,esk3_3(X9,X10,X11),X9)|(((~ir(X10)|~ip(X9))|~ir(X11))|iext(X9,X10,X11))))&(iext(uri_owl_targetIndividual,esk3_3(X9,X10,X11),X11)|(((~ir(X10)|~ip(X9))|~ir(X11))|iext(X9,X10,X11))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
fof(c_0_13, plain, (![X2]:ir(X2)), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_14, negated_conjecture, (![X10]:(((~iext(uri_rdf_type,X10,uri_owl_NegativePropertyAssertion)|~iext(uri_owl_sourceIndividual,X10,uri_ex_s))|~iext(uri_owl_assertionProperty,X10,uri_ex_p))|~iext(uri_owl_targetIndividual,X10,uri_ex_o))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).
cnf(c_0_15,plain,(iext(X1,X2,X3)|iext(uri_owl_targetIndividual,esk3_3(X1,X2,X3),X3)|~ir(X3)|~ip(X1)|~ir(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16,plain,(ir(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, (![X4]:![X5]:((icext(uri_owl_Restriction,X4)|~iext(uri_owl_onProperty,X4,X5))&(ip(X5)|~iext(uri_owl_onProperty,X4,X5)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
fof(c_0_18, plain, ((((((((iext(uri_rdf_type,uri_ex_p,uri_owl_ObjectProperty)&iext(uri_rdf_type,uri_ex_s,esk4_0))&iext(uri_owl_onProperty,esk4_0,uri_ex_p))&iext(uri_owl_allValuesFrom,esk4_0,esk5_0))&iext(uri_owl_complementOf,esk5_0,esk6_0))&iext(uri_owl_oneOf,esk6_0,esk7_0))&iext(uri_rdf_first,esk7_0,uri_ex_o))&iext(uri_rdf_rest,esk7_0,uri_rdf_nil))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_4])])).
cnf(c_0_19,negated_conjecture,(~iext(uri_owl_targetIndividual,X1,uri_ex_o)|~iext(uri_owl_assertionProperty,X1,uri_ex_p)|~iext(uri_owl_sourceIndividual,X1,uri_ex_s)|~iext(uri_rdf_type,X1,uri_owl_NegativePropertyAssertion)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20,plain,(iext(X1,X2,X3)|iext(uri_owl_targetIndividual,esk3_3(X1,X2,X3),X3)|~ip(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16]), c_0_16])])).
cnf(c_0_21,plain,(iext(X1,X2,X3)|iext(uri_owl_assertionProperty,esk3_3(X1,X2,X3),X1)|~ir(X3)|~ip(X1)|~ir(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22,plain,(ip(X2)|~iext(uri_owl_onProperty,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,plain,(iext(uri_owl_onProperty,esk4_0,uri_ex_p)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, (![X4]:![X5]:((icext(uri_owl_NegativePropertyAssertion,X4)|~iext(uri_owl_sourceIndividual,X4,X5))&(ir(X5)|~iext(uri_owl_sourceIndividual,X4,X5)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_25,plain,(iext(X1,X2,X3)|iext(uri_owl_sourceIndividual,esk3_3(X1,X2,X3),X2)|~ir(X3)|~ip(X1)|~ir(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_26,negated_conjecture,(iext(X1,X2,uri_ex_o)|~ip(X1)|~iext(uri_rdf_type,esk3_3(X1,X2,uri_ex_o),uri_owl_NegativePropertyAssertion)|~iext(uri_owl_sourceIndividual,esk3_3(X1,X2,uri_ex_o),uri_ex_s)|~iext(uri_owl_assertionProperty,esk3_3(X1,X2,uri_ex_o),uri_ex_p)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_27,plain,(iext(X1,X2,X3)|iext(uri_owl_assertionProperty,esk3_3(X1,X2,X3),X1)|~ip(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_16]), c_0_16])])).
cnf(c_0_28,plain,(ip(uri_ex_p)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_29, plain, (![X3]:![X4]:![X5]:![X6]:((~iext(uri_rdf_type,X3,X4)|icext(X4,X3))&(~icext(X6,X5)|iext(uri_rdf_type,X5,X6)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_30,plain,(icext(uri_owl_NegativePropertyAssertion,X1)|~iext(uri_owl_sourceIndividual,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31,plain,(iext(X1,X2,X3)|iext(uri_owl_sourceIndividual,esk3_3(X1,X2,X3),X2)|~ip(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_16]), c_0_16])])).
fof(c_0_32, plain, (![X4]:![X2]:(iext(uri_owl_complementOf,X4,X2)=>((ic(X4)&ic(X2))&![X1]:(icext(X4,X1)<=>~icext(X2,X1))))), inference(fof_simplification,[status(thm)],[c_0_7])).
fof(c_0_33, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:((((ic(X7)|~iext(uri_owl_oneOf,X7,X8))|(~iext(uri_rdf_first,X8,X9)|~iext(uri_rdf_rest,X8,uri_rdf_nil)))&((((~icext(X7,X10)|X10=X9)|~iext(uri_owl_oneOf,X7,X8))|(~iext(uri_rdf_first,X8,X9)|~iext(uri_rdf_rest,X8,uri_rdf_nil)))&(((X11!=X9|icext(X7,X11))|~iext(uri_owl_oneOf,X7,X8))|(~iext(uri_rdf_first,X8,X9)|~iext(uri_rdf_rest,X8,uri_rdf_nil)))))&(((((~icext(X7,esk1_3(X7,X8,X9))|esk1_3(X7,X8,X9)!=X9)|~ic(X7))|iext(uri_owl_oneOf,X7,X8))|(~iext(uri_rdf_first,X8,X9)|~iext(uri_rdf_rest,X8,uri_rdf_nil)))&((((icext(X7,esk1_3(X7,X8,X9))|esk1_3(X7,X8,X9)=X9)|~ic(X7))|iext(uri_owl_oneOf,X7,X8))|(~iext(uri_rdf_first,X8,X9)|~iext(uri_rdf_rest,X8,uri_rdf_nil)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])])).
fof(c_0_34, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:(((~icext(X8,X11)|(~iext(X9,X11,X12)|icext(X10,X12)))|(~iext(uri_owl_allValuesFrom,X8,X10)|~iext(uri_owl_onProperty,X8,X9)))&(((iext(X9,X13,esk2_4(X8,X9,X10,X13))|icext(X8,X13))|(~iext(uri_owl_allValuesFrom,X8,X10)|~iext(uri_owl_onProperty,X8,X9)))&((~icext(X10,esk2_4(X8,X9,X10,X13))|icext(X8,X13))|(~iext(uri_owl_allValuesFrom,X8,X10)|~iext(uri_owl_onProperty,X8,X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])])).
cnf(c_0_35,negated_conjecture,(iext(uri_ex_p,X1,uri_ex_o)|~iext(uri_rdf_type,esk3_3(uri_ex_p,X1,uri_ex_o),uri_owl_NegativePropertyAssertion)|~iext(uri_owl_sourceIndividual,esk3_3(uri_ex_p,X1,uri_ex_o),uri_ex_s)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_36,plain,(iext(uri_rdf_type,X1,X2)|~icext(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37,plain,(icext(uri_owl_NegativePropertyAssertion,esk3_3(X1,X2,X3))|iext(X1,X2,X3)|~ip(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_38, plain, (![X5]:![X6]:![X7]:![X8]:(((ic(X5)|~iext(uri_owl_complementOf,X5,X6))&(ic(X6)|~iext(uri_owl_complementOf,X5,X6)))&(((~icext(X5,X7)|~icext(X6,X7))|~iext(uri_owl_complementOf,X5,X6))&((icext(X6,X8)|icext(X5,X8))|~iext(uri_owl_complementOf,X5,X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])])])).
cnf(c_0_39,plain,(icext(X3,X4)|~iext(uri_rdf_rest,X1,uri_rdf_nil)|~iext(uri_rdf_first,X1,X2)|~iext(uri_owl_oneOf,X3,X1)|X4!=X2), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40,plain,(icext(X3,X4)|~iext(uri_owl_onProperty,X1,X2)|~iext(uri_owl_allValuesFrom,X1,X3)|~iext(X2,X5,X4)|~icext(X1,X5)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41,negated_conjecture,(iext(uri_ex_p,uri_ex_s,uri_ex_o)|~iext(uri_rdf_type,esk3_3(uri_ex_p,uri_ex_s,uri_ex_o),uri_owl_NegativePropertyAssertion)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_31]), c_0_28])])).
cnf(c_0_42,plain,(iext(uri_rdf_type,esk3_3(X1,X2,X3),uri_owl_NegativePropertyAssertion)|iext(X1,X2,X3)|~ip(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43,plain,(~iext(uri_owl_complementOf,X1,X2)|~icext(X2,X3)|~icext(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44,plain,(iext(uri_owl_complementOf,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_45,plain,(icext(X1,X2)|~iext(uri_rdf_rest,X3,uri_rdf_nil)|~iext(uri_rdf_first,X3,X2)|~iext(uri_owl_oneOf,X1,X3)), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_46,plain,(iext(uri_rdf_first,esk7_0,uri_ex_o)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_47,plain,(iext(uri_rdf_rest,esk7_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_48,plain,(icext(X1,X3)|icext(X2,X3)|~iext(uri_owl_complementOf,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49,plain,(icext(X1,X2)|~icext(esk4_0,X3)|~iext(uri_owl_allValuesFrom,esk4_0,X1)|~iext(uri_ex_p,X3,X2)), inference(spm,[status(thm)],[c_0_40, c_0_23])).
cnf(c_0_50,negated_conjecture,(iext(uri_ex_p,uri_ex_s,uri_ex_o)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_28])])).
cnf(c_0_51,plain,(~icext(esk6_0,X1)|~icext(esk5_0,X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_52,plain,(icext(X1,uri_ex_o)|~iext(uri_owl_oneOf,X1,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
cnf(c_0_53,plain,(iext(uri_owl_oneOf,esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_54,plain,(icext(esk5_0,X1)|icext(esk6_0,X1)), inference(spm,[status(thm)],[c_0_48, c_0_44])).
cnf(c_0_55,plain,(icext(X1,X2)|~iext(uri_rdf_type,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_56,negated_conjecture,(icext(X1,uri_ex_o)|~icext(esk4_0,uri_ex_s)|~iext(uri_owl_allValuesFrom,esk4_0,X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_57,plain,(iext(uri_rdf_type,uri_ex_s,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_58,plain,(~icext(esk5_0,uri_ex_o)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])])).
cnf(c_0_59,plain,(icext(esk5_0,X1)|iext(uri_rdf_type,X1,esk6_0)), inference(spm,[status(thm)],[c_0_36, c_0_54])).
cnf(c_0_60,plain,(~icext(esk5_0,X1)|~iext(uri_rdf_type,X1,esk6_0)), inference(spm,[status(thm)],[c_0_51, c_0_55])).
cnf(c_0_61,negated_conjecture,(icext(X1,uri_ex_o)|~iext(uri_owl_allValuesFrom,esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_55]), c_0_57])])).
cnf(c_0_62,plain,(iext(uri_rdf_type,uri_ex_o,esk6_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63,plain,(iext(uri_owl_allValuesFrom,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_64,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63])]), ['proof']).
# SZS output end CNFRefutation.
EOF
