% Problem    : SWW378+1 : TPTP v6.1.0. Released v5.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n102.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 07:05:16 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X10]:c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),hAPP(c_COMBK(tc_HOL_Obool,X10),c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__def)).
fof(c_0_1, axiom, (![X16]:![X10]:hAPP(c_Set_OCollect(X10),X16)=X16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__def)).
fof(c_0_2, axiom, (![X325]:![X19]:![X28]:![X10]:hAPP(hAPP(c_COMBK(X10,X28),X19),X325)=X19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBK__1)).
fof(c_0_3, axiom, (![X10]:c_Orderings_Otop__class_Otop(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),hAPP(c_COMBK(tc_HOL_Obool,X10),c_fTrue))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__def)).
fof(c_0_4, axiom, (![X10]:(class_Lattices_Oboolean__algebra(X10)=>hAPP(c_Groups_Ouminus__class_Ouminus(X10),c_Orderings_Obot__class_Obot(X10))=c_Orderings_Otop__class_Otop(X10))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_compl__bot__eq)).
fof(c_0_5, axiom, (![X25]:![X28]:![X10]:(class_Orderings_Obot(X10)=>hAPP(c_Orderings_Obot__class_Obot(tc_fun(X28,X10)),X25)=c_Orderings_Obot__class_Obot(X10))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_bot__fun__def)).
fof(c_0_6, axiom, (class_Lattices_Oboolean__algebra(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Lattices_Oboolean__algebra)).
fof(c_0_7, axiom, (class_Orderings_Obot(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Obot)).
fof(c_0_8, axiom, (![X48]:![X10]:hAPP(c_Set_OCollect(X10),hAPP(c_fequal,X48))=hAPP(hAPP(c_Set_Oinsert(X10),X48),c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_singleton__conv2)).
fof(c_0_9, axiom, (![X25]:![X28]:![X10]:(class_Orderings_Otop(X10)=>hAPP(c_Orderings_Otop__class_Otop(tc_fun(X28,X10)),X25)=c_Orderings_Otop__class_Otop(X10))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_top__apply)).
fof(c_0_10, axiom, (![X27]:![X10]:hAPP(c_Set_OCollect(X10),hAPP(hAPP(c_COMBC(X10,tc_fun(X10,tc_HOL_Obool),tc_HOL_Obool),c_member(X10)),X27))=X27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__mem__eq)).
fof(c_0_11, axiom, (![X48]:![X10]:(class_Complete__Lattice_Ocomplete__lattice(X10)=>c_Complete__Lattice_OSup__class_OSup(X10,hAPP(hAPP(c_Set_Oinsert(X10),X48),c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))))=X48)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Sup__singleton)).
fof(c_0_12, axiom, (![X33]:![X10]:(class_Lattices_Obounded__lattice__top(X10)=>hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(X10),c_Orderings_Otop__class_Otop(X10)),X33)=c_Orderings_Otop__class_Otop(X10))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_sup__top__left)).
fof(c_0_13, axiom, (class_Orderings_Otop(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Orderings_Otop)).
fof(c_0_14, axiom, (![X54]:![X15]:![X16]:![X10]:![X40]:![X28]:hAPP(hAPP(hAPP(c_COMBC(X28,X40,X10),X16),X15),X54)=hAPP(hAPP(X16,X54),X15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBC__1)).
fof(c_0_15, axiom, (![X323]:![X324]:(class_Complete__Lattice_Ocomplete__lattice(X324)=>class_Complete__Lattice_Ocomplete__lattice(tc_fun(X323,X324)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_fun__Complete__Lattice_Ocomplete__lattice)).
fof(c_0_16, axiom, (![X16]:![X10]:(c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),X16)<=>![X3]:~(hBOOL(hAPP(X16,X3))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__Collect__eq)).
fof(c_0_17, axiom, (![X16]:![X10]:hAPP(c_Set_OCollect(X10),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X10),c_fNot),X16))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),hAPP(c_Set_OCollect(X10),X16))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Collect__neg__eq)).
fof(c_0_18, axiom, (![X1]:![X2]:(![X3]:hAPP(X2,X3)=hAPP(X1,X3)=>X2=X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_ext)).
fof(c_0_19, axiom, (class_Lattices_Obounded__lattice__top(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Lattices_Obounded__lattice__top)).
fof(c_0_20, axiom, (![X26]:![X27]:![X10]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X10,tc_HOL_Obool)),X27),X26)=c_Complete__Lattice_OSup__class_OSup(tc_fun(X10,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X10,tc_HOL_Obool)),X27),hAPP(hAPP(c_Set_Oinsert(tc_fun(X10,tc_HOL_Obool)),X26),c_Orderings_Obot__class_Obot(tc_fun(tc_fun(X10,tc_HOL_Obool),tc_HOL_Obool)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Un__eq__Union)).
fof(c_0_21, axiom, (![X27]:![X48]:![X10]:(hBOOL(hAPP(hAPP(c_member(X10),X48),X27))=>hAPP(hAPP(c_Set_Oinsert(X10),X48),X27)=X27)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_insert__absorb)).
fof(c_0_22, axiom, (![X10]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),c_Orderings_Otop__class_Otop(tc_fun(X10,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__UNIV__eq)).
fof(c_0_23, axiom, (![X27]:![X48]:![X10]:c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))!=hAPP(hAPP(c_Set_Oinsert(X10),X48),X27)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__not__insert)).
fof(c_0_24, axiom, (![X26]:![X48]:![X27]:![X10]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X10,tc_HOL_Obool)),X27),hAPP(hAPP(c_Set_Oinsert(X10),X48),X26))=hAPP(hAPP(c_Set_Oinsert(X10),X48),hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X10,tc_HOL_Obool)),X27),X26))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Un__insert__right)).
fof(c_0_25, axiom, (class_Complete__Lattice_Ocomplete__lattice(tc_HOL_Obool)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_HOL__Obool__Complete__Lattice_Ocomplete__lattice)).
fof(c_0_26, axiom, (![X27]:![X10]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),X27))=X27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_double__complement)).
fof(c_0_27, axiom, (![X16]:(~(hBOOL(hAPP(c_fNot,X16)))|~(hBOOL(X16)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fNot__1)).
fof(c_0_28, axiom, (![X26]:![X27]:![X10]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X10,tc_HOL_Obool)),X27),X26)=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X10,tc_HOL_Obool)),X26),X27)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Un__commute)).
fof(c_0_29, axiom, (![X54]:![X15]:![X16]:![X40]:![X10]:![X28]:hAPP(hAPP(hAPP(c_COMBB(X28,X10,X40),X16),X15),X54)=hAPP(X16,hAPP(X15,X54))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__COMBB__1)).
fof(c_0_30, axiom, (![X54]:![X10]:(c_Equiv__Relations_Oequivp(X10,X54)<=>![X3]:![X109]:(hBOOL(hAPP(hAPP(X54,X3),X109))<=>hAPP(X54,X3)=hAPP(X54,X109)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_equivp__def)).
fof(c_0_31, axiom, (![X10]:c_Equiv__Relations_Oequivp(X10,c_fequal)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_identity__equivp)).
fof(c_0_32, axiom, (![X250]:![X10]:hAPP(c_List_Orev(X10),hAPP(c_List_Orev(X10),X250))=X250), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_rev__rev__ident)).
fof(c_0_33, axiom, (![X10]:(class_HOL_Oequal(X10)=>c_HOL_Oequal__class_Oequal(X10)=c_fequal)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_equal)).
fof(c_0_34, conjecture, (c_Hoare__Mirabelle_Ohoare__derivs(tc_Com_Ostate,hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate),tc_HOL_Obool)),v_G),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate)),hAPP(hAPP(c_COMBB(tc_Com_Ocom,tc_fun(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate)),tc_Com_Opname),hAPP(c_Hoare__Mirabelle_Otriple_Otriple(tc_Com_Ostate),c_fequal)),c_Com_Ocom_OBODY)),hAPP(hAPP(c_COMBB(tc_Com_Ocom,tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Com_Opname),c_Natural_Oevalc),c_Com_Ocom_OBODY))),v_Procs)),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate)),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_Com_Ocom,tc_fun(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate))),hAPP(hAPP(c_COMBB(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_fun(tc_Com_Ocom,tc_fun(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate))),tc_Com_Opname),c_Hoare__Mirabelle_Otriple_Otriple(tc_Com_Ostate)),v_Sledgehammer_058P_H__2)),hAPP(hAPP(c_COMBB(tc_Option_Ooption(tc_Com_Ocom),tc_Com_Ocom,tc_Com_Opname),c_Option_Othe(tc_Com_Ocom)),c_Com_Obody))),v_Sledgehammer_058Q_H__2)),v_Procs))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_1)).
fof(c_0_35, axiom, (![X27]:![X2]:![X28]:![X10]:(c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_Oimage(X28,X10,X2),X27)<=>X27=c_Orderings_Obot__class_Obot(tc_fun(X28,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty__is__image)).
fof(c_0_36, axiom, (![X11]:![X10]:c_Hoare__Mirabelle_Ohoare__derivs(X10,X11,c_Orderings_Obot__class_Obot(tc_fun(tc_Hoare__Mirabelle_Otriple(X10),tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_empty)).
fof(c_0_37, axiom, (~(hBOOL(hAPP(hAPP(c_HOL_Oequal__class_Oequal(tc_Int_Oint),c_Int_OPls),c_Int_OMin)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_eq__int__code_I2_J)).
fof(c_0_38, axiom, (class_HOL_Oequal(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', arity_Int__Oint__HOL_Oequal)).
fof(c_0_39, axiom, (![X46]:![X25]:![X10]:(hBOOL(hAPP(hAPP(c_member(X10),X25),hAPP(c_fequal,X46)))<=>X25=X46)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_eq__mem)).
fof(c_0_40, axiom, (![X27]:![X10]:(?[X3]:hBOOL(hAPP(hAPP(c_member(X10),X3),X27))<=>X27!=c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_ex__in__conv)).
fof(c_0_41, axiom, (![X27]:![X10]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X10,tc_HOL_Obool)),X27),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X10,tc_HOL_Obool)),X27))=c_Orderings_Otop__class_Otop(tc_fun(X10,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Compl__partition)).
fof(c_0_42, axiom, (![X10]:c_Orderings_Otop__class_Otop(tc_fun(X10,tc_HOL_Obool))!=c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_UNIV__not__empty)).
fof(c_0_43, axiom, (![X48]:![X10]:hAPP(c_Set_OCollect(X10),hAPP(hAPP(c_COMBC(X10,X10,tc_HOL_Obool),c_fequal),X48))=hAPP(hAPP(c_Set_Oinsert(X10),X48),c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_singleton__conv)).
fof(c_0_44, axiom, (~(hBOOL(c_fFalse))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', help_c__fFalse__1)).
fof(c_0_45, axiom, (![X47]:![X10]:hAPP(c_List_Olexn(X10,X47),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=c_Orderings_Obot__class_Obot(tc_fun(tc_prod(tc_List_Olist(X10),tc_List_Olist(X10)),tc_HOL_Obool))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_lexn_Osimps_I1_J)).
fof(c_0_46, axiom, (hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_nat__0)).
fof(c_0_47, axiom, (c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_Pls__def)).
fof(c_0_48, plain, (![X11]:c_Orderings_Obot__class_Obot(tc_fun(X11,tc_HOL_Obool))=hAPP(c_Set_OCollect(X11),hAPP(c_COMBK(tc_HOL_Obool,X11),c_fFalse))), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_49, plain, (![X17]:![X18]:hAPP(c_Set_OCollect(X18),X17)=X17), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_50, plain, (![X326]:![X327]:![X328]:![X329]:hAPP(hAPP(c_COMBK(X329,X328),X327),X326)=X327), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_51,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_52,plain,(hAPP(c_Set_OCollect(X1),X2)=X2), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_53, plain, (![X11]:c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool))=hAPP(c_Set_OCollect(X11),hAPP(c_COMBK(tc_HOL_Obool,X11),c_fTrue))), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_54, plain, (![X11]:(~class_Lattices_Oboolean__algebra(X11)|hAPP(c_Groups_Ouminus__class_Ouminus(X11),c_Orderings_Obot__class_Obot(X11))=c_Orderings_Otop__class_Otop(X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
fof(c_0_55, plain, (![X29]:![X30]:![X31]:(~class_Orderings_Obot(X31)|hAPP(c_Orderings_Obot__class_Obot(tc_fun(X30,X31)),X29)=c_Orderings_Obot__class_Obot(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
cnf(c_0_56,plain,(hAPP(hAPP(c_COMBK(X1,X2),X3),X4)=X3), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fFalse)=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_58,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_59,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(X1),c_Orderings_Obot__class_Obot(X1))=c_Orderings_Otop__class_Otop(X1)|~class_Lattices_Oboolean__algebra(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_60,plain,(class_Lattices_Oboolean__algebra(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_61,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,X2)),X3)=c_Orderings_Obot__class_Obot(X2)|~class_Orderings_Obot(X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62,plain,(class_Orderings_Obot(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_63,plain,(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)),X2)=c_fFalse), inference(spm,[status(thm)],[c_0_56, c_0_57])).
fof(c_0_64, plain, (![X49]:![X50]:hAPP(c_Set_OCollect(X50),hAPP(c_fequal,X49))=hAPP(hAPP(c_Set_Oinsert(X50),X49),c_Orderings_Obot__class_Obot(tc_fun(X50,tc_HOL_Obool)))), inference(variable_rename,[status(thm)],[c_0_8])).
fof(c_0_65, plain, (![X29]:![X30]:![X31]:(~class_Orderings_Otop(X31)|hAPP(c_Orderings_Otop__class_Otop(tc_fun(X30,X31)),X29)=c_Orderings_Otop__class_Otop(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
cnf(c_0_66,plain,(hAPP(c_COMBK(tc_HOL_Obool,X1),c_fTrue)=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_58, c_0_52])).
cnf(c_0_67,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_Orderings_Obot__class_Obot(tc_HOL_Obool))=c_Orderings_Otop__class_Otop(tc_HOL_Obool)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_68,plain,(c_Orderings_Obot__class_Obot(tc_HOL_Obool)=c_fFalse), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).
fof(c_0_69, plain, (![X28]:![X29]:hAPP(c_Set_OCollect(X29),hAPP(hAPP(c_COMBC(X29,tc_fun(X29,tc_HOL_Obool),tc_HOL_Obool),c_member(X29)),X28))=X28), inference(variable_rename,[status(thm)],[c_0_10])).
fof(c_0_70, plain, (![X49]:![X50]:(~class_Complete__Lattice_Ocomplete__lattice(X50)|c_Complete__Lattice_OSup__class_OSup(X50,hAPP(hAPP(c_Set_Oinsert(X50),X49),c_Orderings_Obot__class_Obot(tc_fun(X50,tc_HOL_Obool))))=X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).
cnf(c_0_71,plain,(hAPP(c_Set_OCollect(X1),hAPP(c_fequal,X2))=hAPP(hAPP(c_Set_Oinsert(X1),X2),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_72, plain, (![X34]:![X35]:(~class_Lattices_Obounded__lattice__top(X35)|hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(X35),c_Orderings_Otop__class_Otop(X35)),X34)=c_Orderings_Otop__class_Otop(X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
cnf(c_0_73,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,X2)),X3)=c_Orderings_Otop__class_Otop(X2)|~class_Orderings_Otop(X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_74,plain,(class_Orderings_Otop(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_75,plain,(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)),X2)=c_fTrue), inference(spm,[status(thm)],[c_0_56, c_0_66])).
cnf(c_0_76,plain,(c_Orderings_Otop__class_Otop(tc_HOL_Obool)=hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fFalse)), inference(rw,[status(thm)],[c_0_67, c_0_68])).
fof(c_0_77, plain, (![X55]:![X56]:![X57]:![X58]:![X59]:![X60]:hAPP(hAPP(hAPP(c_COMBC(X60,X59,X58),X57),X56),X55)=hAPP(hAPP(X57,X55),X56)), inference(variable_rename,[status(thm)],[c_0_14])).
cnf(c_0_78,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBC(X1,tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool),c_member(X1)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_79,plain,(c_Complete__Lattice_OSup__class_OSup(X1,hAPP(hAPP(c_Set_Oinsert(X1),X2),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))))=X2|~class_Complete__Lattice_Ocomplete__lattice(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_80,plain,(hAPP(hAPP(c_Set_Oinsert(X1),X2),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=hAPP(c_fequal,X2)), inference(rw,[status(thm)],[c_0_71, c_0_52])).
fof(c_0_81, plain, (![X325]:![X326]:(~class_Complete__Lattice_Ocomplete__lattice(X326)|class_Complete__Lattice_Ocomplete__lattice(tc_fun(X325,X326)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
fof(c_0_82, plain, (![X16]:![X10]:(c_Orderings_Obot__class_Obot(tc_fun(X10,tc_HOL_Obool))=hAPP(c_Set_OCollect(X10),X16)<=>![X3]:~hBOOL(hAPP(X16,X3)))), inference(fof_simplification,[status(thm)],[c_0_16])).
fof(c_0_83, plain, (![X17]:![X18]:hAPP(c_Set_OCollect(X18),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X18),c_fNot),X17))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X18,tc_HOL_Obool)),hAPP(c_Set_OCollect(X18),X17))), inference(variable_rename,[status(thm)],[c_0_17])).
fof(c_0_84, plain, (![X4]:![X5]:(hAPP(X5,esk1_2(X4,X5))!=hAPP(X4,esk1_2(X4,X5))|X5=X4)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
cnf(c_0_85,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(X1),c_Orderings_Otop__class_Otop(X1)),X2)=c_Orderings_Otop__class_Otop(X1)|~class_Lattices_Obounded__lattice__top(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_86,plain,(class_Lattices_Obounded__lattice__top(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_87,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_HOL_Obool),c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75]), c_0_76])).
fof(c_0_88, plain, (![X28]:![X29]:![X30]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X30,tc_HOL_Obool)),X29),X28)=c_Complete__Lattice_OSup__class_OSup(tc_fun(X30,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X30,tc_HOL_Obool)),X29),hAPP(hAPP(c_Set_Oinsert(tc_fun(X30,tc_HOL_Obool)),X28),c_Orderings_Obot__class_Obot(tc_fun(tc_fun(X30,tc_HOL_Obool),tc_HOL_Obool)))))), inference(variable_rename,[status(thm)],[c_0_20])).
fof(c_0_89, plain, (![X49]:![X50]:![X51]:(~hBOOL(hAPP(hAPP(c_member(X51),X50),X49))|hAPP(hAPP(c_Set_Oinsert(X51),X50),X49)=X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).
cnf(c_0_90,plain,(hAPP(hAPP(hAPP(c_COMBC(X1,X2,X3),X4),X5),X6)=hAPP(hAPP(X4,X6),X5)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_91,plain,(hAPP(hAPP(c_COMBC(X1,tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool),c_member(X1)),X2)=X2), inference(rw,[status(thm)],[c_0_78, c_0_52])).
cnf(c_0_92,plain,(c_Complete__Lattice_OSup__class_OSup(X1,hAPP(c_fequal,X2))=X2|~class_Complete__Lattice_Ocomplete__lattice(X1)), inference(rw,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_93,plain,(class_Complete__Lattice_Ocomplete__lattice(tc_fun(X1,X2))|~class_Complete__Lattice_Ocomplete__lattice(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_94, plain, (![X17]:![X18]:![X19]:![X20]:![X22]:((c_Orderings_Obot__class_Obot(tc_fun(X18,tc_HOL_Obool))!=hAPP(c_Set_OCollect(X18),X17)|~hBOOL(hAPP(X17,X19)))&(hBOOL(hAPP(X20,esk194_1(X20)))|c_Orderings_Obot__class_Obot(tc_fun(X22,tc_HOL_Obool))=hAPP(c_Set_OCollect(X22),X20)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_82])])])])])).
fof(c_0_95, plain, (![X11]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X11,tc_HOL_Obool)),c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X11,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_22])).
cnf(c_0_96,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Set_OCollect(X1),X2))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_97,plain,(X1=X2|hAPP(X1,esk1_2(X2,X1))!=hAPP(X2,esk1_2(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_98,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue),X1)=c_fTrue), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_76]), c_0_87]), c_0_76]), c_0_87])).
fof(c_0_99, plain, (![X49]:![X50]:![X51]:c_Orderings_Obot__class_Obot(tc_fun(X51,tc_HOL_Obool))!=hAPP(hAPP(c_Set_Oinsert(X51),X50),X49)), inference(variable_rename,[status(thm)],[c_0_23])).
fof(c_0_100, plain, (![X49]:![X50]:![X51]:![X52]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X52,tc_HOL_Obool)),X51),hAPP(hAPP(c_Set_Oinsert(X52),X50),X49))=hAPP(hAPP(c_Set_Oinsert(X52),X50),hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X52,tc_HOL_Obool)),X51),X49))), inference(variable_rename,[status(thm)],[c_0_24])).
cnf(c_0_101,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X3)=c_Complete__Lattice_OSup__class_OSup(tc_fun(X1,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(hAPP(c_Set_Oinsert(tc_fun(X1,tc_HOL_Obool)),X3),c_Orderings_Obot__class_Obot(tc_fun(tc_fun(X1,tc_HOL_Obool),tc_HOL_Obool)))))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_102,plain,(hAPP(hAPP(c_Set_Oinsert(X1),X2),X3)=X3|~hBOOL(hAPP(hAPP(c_member(X1),X2),X3))), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_103,plain,(hAPP(hAPP(c_member(X1),X2),X3)=hAPP(X3,X2)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_104,plain,(c_Complete__Lattice_OSup__class_OSup(tc_fun(X1,X2),hAPP(c_fequal,X3))=X3|~class_Complete__Lattice_Ocomplete__lattice(X2)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_105,plain,(class_Complete__Lattice_Ocomplete__lattice(tc_HOL_Obool)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_106, plain, (![X28]:![X29]:hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X29,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X29,tc_HOL_Obool)),X28))=X28), inference(variable_rename,[status(thm)],[c_0_26])).
fof(c_0_107, plain, (![X16]:(~hBOOL(hAPP(c_fNot,X16))|~hBOOL(X16))), inference(fof_simplification,[status(thm)],[c_0_27])).
cnf(c_0_108,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_OCollect(X1),X2)|hBOOL(hAPP(X2,esk194_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_109,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_110,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2)=hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_52]), c_0_52])).
cnf(c_0_111,plain,(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue)=X1|hAPP(X1,esk1_2(X1,hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue)))!=c_fTrue), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_112,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))!=hAPP(hAPP(c_Set_Oinsert(X1),X2),X3)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_113,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(hAPP(c_Set_Oinsert(X1),X3),X4))=hAPP(hAPP(c_Set_Oinsert(X1),X3),hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X4))), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_114,plain,(c_Complete__Lattice_OSup__class_OSup(tc_fun(X1,tc_HOL_Obool),hAPP(hAPP(c_Set_Oinsert(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(c_fequal,X3)))=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X3)), inference(rw,[status(thm)],[c_0_101, c_0_80])).
cnf(c_0_115,plain,(hAPP(hAPP(c_Set_Oinsert(X1),X2),X3)=X3|~hBOOL(hAPP(X3,X2))), inference(rw,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_116,plain,(c_Complete__Lattice_OSup__class_OSup(tc_fun(X1,tc_HOL_Obool),hAPP(c_fequal,X2))=X2), inference(spm,[status(thm)],[c_0_104, c_0_105])).
fof(c_0_117, plain, (![X28]:![X29]:![X30]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X30,tc_HOL_Obool)),X29),X28)=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X30,tc_HOL_Obool)),X28),X29)), inference(variable_rename,[status(thm)],[c_0_28])).
fof(c_0_118, plain, (![X55]:![X56]:![X57]:![X58]:![X59]:![X60]:hAPP(hAPP(hAPP(c_COMBB(X60,X59,X58),X57),X56),X55)=hAPP(X57,hAPP(X56,X55))), inference(variable_rename,[status(thm)],[c_0_29])).
cnf(c_0_119,plain,(hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_106])).
fof(c_0_120, plain, (![X17]:(~hBOOL(hAPP(c_fNot,X17))|~hBOOL(X17))), inference(variable_rename,[status(thm)],[c_0_107])).
cnf(c_0_121,plain,(X1=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|hBOOL(hAPP(X1,esk194_1(X1)))), inference(rw,[status(thm)],[c_0_108, c_0_52])).
cnf(c_0_122,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool)))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_123,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue)), inference(spm,[status(thm)],[c_0_111, c_0_75])).
cnf(c_0_124,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(hAPP(c_Set_Oinsert(X1),X3),X4))!=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_125,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X3)=X3|~hBOOL(hAPP(hAPP(c_fequal,X3),X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_116])).
cnf(c_0_126,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X3)=hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X3),X2)), inference(split_conjunct,[status(thm)],[c_0_117])).
fof(c_0_127, plain, (![X110]:![X111]:![X112]:![X113]:![X114]:![X115]:![X116]:![X119]:((((~hBOOL(hAPP(hAPP(X110,X112),X113))|hAPP(X110,X112)=hAPP(X110,X113))|~c_Equiv__Relations_Oequivp(X111,X110))&((hAPP(X110,X114)!=hAPP(X110,X115)|hBOOL(hAPP(hAPP(X110,X114),X115)))|~c_Equiv__Relations_Oequivp(X111,X110)))&(((~hBOOL(hAPP(hAPP(X116,esk279_1(X116)),esk280_1(X116)))|hAPP(X116,esk279_1(X116))!=hAPP(X116,esk280_1(X116)))|c_Equiv__Relations_Oequivp(X119,X116))&((hBOOL(hAPP(hAPP(X116,esk279_1(X116)),esk280_1(X116)))|hAPP(X116,esk279_1(X116))=hAPP(X116,esk280_1(X116)))|c_Equiv__Relations_Oequivp(X119,X116))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])])])).
fof(c_0_128, plain, (![X11]:c_Equiv__Relations_Oequivp(X11,c_fequal)), inference(variable_rename,[status(thm)],[c_0_31])).
cnf(c_0_129,plain,(hAPP(hAPP(hAPP(c_COMBB(X1,X2,X3),X4),X5),X6)=hAPP(X4,hAPP(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_130,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_110]), c_0_110])).
fof(c_0_131, plain, (![X251]:![X252]:hAPP(c_List_Orev(X252),hAPP(c_List_Orev(X252),X251))=X251), inference(variable_rename,[status(thm)],[c_0_32])).
cnf(c_0_132,plain,(~hBOOL(X1)|~hBOOL(hAPP(c_fNot,X1))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_133,plain,(c_fFalse=X1|hBOOL(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_121]), c_0_63]), c_0_56])).
cnf(c_0_134,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue))=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_122, c_0_123])).
fof(c_0_135, plain, (![X11]:(~class_HOL_Oequal(X11)|c_HOL_Oequal__class_Oequal(X11)=c_fequal)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])).
cnf(c_0_136,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(c_fequal,X3))!=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_124, c_0_80])).
cnf(c_0_137,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),X3)=X2|~hBOOL(hAPP(hAPP(c_fequal,X2),X3))), inference(spm,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_138,plain,(hBOOL(hAPP(hAPP(X2,X3),X4))|~c_Equiv__Relations_Oequivp(X1,X2)|hAPP(X2,X3)!=hAPP(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_139,plain,(c_Equiv__Relations_Oequivp(X1,c_fequal)), inference(split_conjunct,[status(thm)],[c_0_128])).
fof(c_0_140, negated_conjecture, (~c_Hoare__Mirabelle_Ohoare__derivs(tc_Com_Ostate,hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate),tc_HOL_Obool)),v_G),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate)),hAPP(hAPP(c_COMBB(tc_Com_Ocom,tc_fun(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate)),tc_Com_Opname),hAPP(c_Hoare__Mirabelle_Otriple_Otriple(tc_Com_Ostate),c_fequal)),c_Com_Ocom_OBODY)),hAPP(hAPP(c_COMBB(tc_Com_Ocom,tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Com_Opname),c_Natural_Oevalc),c_Com_Ocom_OBODY))),v_Procs)),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate)),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_Com_Ocom,tc_fun(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate))),hAPP(hAPP(c_COMBB(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_fun(tc_Com_Ocom,tc_fun(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate))),tc_Com_Opname),c_Hoare__Mirabelle_Otriple_Otriple(tc_Com_Ostate)),v_Sledgehammer_058P_H__2)),hAPP(hAPP(c_COMBB(tc_Option_Ooption(tc_Com_Ocom),tc_Com_Ocom,tc_Com_Opname),c_Option_Othe(tc_Com_Ocom)),c_Com_Obody))),v_Sledgehammer_058Q_H__2)),v_Procs))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_34])])).
fof(c_0_141, plain, (![X29]:![X30]:![X31]:![X32]:![X33]:![X34]:![X35]:![X36]:((c_Orderings_Obot__class_Obot(tc_fun(X32,tc_HOL_Obool))!=hAPP(c_Set_Oimage(X31,X32,X30),X29)|X29=c_Orderings_Obot__class_Obot(tc_fun(X31,tc_HOL_Obool)))&(X33!=c_Orderings_Obot__class_Obot(tc_fun(X35,tc_HOL_Obool))|c_Orderings_Obot__class_Obot(tc_fun(X36,tc_HOL_Obool))=hAPP(c_Set_Oimage(X35,X36,X34),X33)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])])).
fof(c_0_142, plain, (![X12]:![X13]:c_Hoare__Mirabelle_Ohoare__derivs(X13,X12,c_Orderings_Obot__class_Obot(tc_fun(tc_Hoare__Mirabelle_Otriple(X13),tc_HOL_Obool)))), inference(variable_rename,[status(thm)],[c_0_36])).
cnf(c_0_143,plain,(hAPP(c_fNot,hAPP(c_fNot,hAPP(X1,X2)))=hAPP(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_129])).
cnf(c_0_144,plain,(hAPP(c_List_Orev(X1),hAPP(c_List_Orev(X1),X2))=X2), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_145,plain,(hAPP(c_fNot,X1)=c_fFalse|~hBOOL(X1)), inference(spm,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_146,plain,(hAPP(c_fNot,c_fTrue)=c_fFalse), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_134]), c_0_63]), c_0_98])).
fof(c_0_147, plain, (~hBOOL(hAPP(hAPP(c_HOL_Oequal__class_Oequal(tc_Int_Oint),c_Int_OPls),c_Int_OMin))), inference(fof_simplification,[status(thm)],[c_0_37])).
cnf(c_0_148,plain,(c_HOL_Oequal__class_Oequal(X1)=c_fequal|~class_HOL_Oequal(X1)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_149,plain,(class_HOL_Oequal(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_150,plain,(X1!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|~hBOOL(hAPP(hAPP(c_fequal,X1),hAPP(c_fequal,X3)))), inference(spm,[status(thm)],[c_0_136, c_0_137])).
cnf(c_0_151,plain,(hBOOL(hAPP(hAPP(c_fequal,X1),X2))|hAPP(c_fequal,X1)!=hAPP(c_fequal,X2)), inference(spm,[status(thm)],[c_0_138, c_0_139])).
cnf(c_0_152,negated_conjecture,(~c_Hoare__Mirabelle_Ohoare__derivs(tc_Com_Ostate,hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate),tc_HOL_Obool)),v_G),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate)),hAPP(hAPP(c_COMBB(tc_Com_Ocom,tc_fun(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate)),tc_Com_Opname),hAPP(c_Hoare__Mirabelle_Otriple_Otriple(tc_Com_Ostate),c_fequal)),c_Com_Ocom_OBODY)),hAPP(hAPP(c_COMBB(tc_Com_Ocom,tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Com_Opname),c_Natural_Oevalc),c_Com_Ocom_OBODY))),v_Procs)),hAPP(c_Set_Oimage(tc_Com_Opname,tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate)),hAPP(hAPP(c_COMBS(tc_Com_Opname,tc_Com_Ocom,tc_fun(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate))),hAPP(hAPP(c_COMBB(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_fun(tc_Com_Ocom,tc_fun(tc_fun(tc_Com_Ostate,tc_fun(tc_Com_Ostate,tc_HOL_Obool)),tc_Hoare__Mirabelle_Otriple(tc_Com_Ostate))),tc_Com_Opname),c_Hoare__Mirabelle_Otriple_Otriple(tc_Com_Ostate)),v_Sledgehammer_058P_H__2)),hAPP(hAPP(c_COMBB(tc_Option_Ooption(tc_Com_Ocom),tc_Com_Ocom,tc_Com_Opname),c_Option_Othe(tc_Com_Ocom)),c_Com_Obody))),v_Sledgehammer_058Q_H__2)),v_Procs))), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_153,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))=hAPP(c_Set_Oimage(X2,X1,X3),X4)|X4!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_154,plain,(c_Hoare__Mirabelle_Ohoare__derivs(X1,X2,c_Orderings_Obot__class_Obot(tc_fun(tc_Hoare__Mirabelle_Otriple(X1),tc_HOL_Obool)))), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_155,plain,(hAPP(c_fNot,hAPP(c_fNot,X1))=X1), inference(spm,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_156,plain,(hAPP(c_fNot,X1)=c_fFalse|c_fFalse=X1), inference(spm,[status(thm)],[c_0_145, c_0_133])).
cnf(c_0_157,plain,(hAPP(c_fNot,c_fFalse)=c_fTrue), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_87]), c_0_146])).
cnf(c_0_158,plain,(~hBOOL(hAPP(hAPP(c_HOL_Oequal__class_Oequal(tc_Int_Oint),c_Int_OPls),c_Int_OMin))), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_159,plain,(c_HOL_Oequal__class_Oequal(tc_Int_Oint)=c_fequal), inference(spm,[status(thm)],[c_0_148, c_0_149])).
fof(c_0_160, plain, (![X47]:![X48]:![X49]:![X50]:![X51]:![X52]:((~hBOOL(hAPP(hAPP(c_member(X49),X48),hAPP(c_fequal,X47)))|X48=X47)&(X51!=X50|hBOOL(hAPP(hAPP(c_member(X52),X51),hAPP(c_fequal,X50)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])])])).
cnf(c_0_161,plain,(hAPP(c_fequal,X1)!=hAPP(c_fequal,hAPP(c_fequal,X2))|X1!=c_Orderings_Obot__class_Obot(tc_fun(X3,tc_HOL_Obool))), inference(spm,[status(thm)],[c_0_150, c_0_151])).
cnf(c_0_162,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(tc_Com_Opname,tc_HOL_Obool))!=v_Procs), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_154])])).
cnf(c_0_163,plain,(c_fFalse=X1|c_fTrue=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_157])).
cnf(c_0_164,plain,(~hBOOL(hAPP(hAPP(c_fequal,c_Int_OPls),c_Int_OMin))), inference(rw,[status(thm)],[c_0_158, c_0_159])).
fof(c_0_165, plain, (![X28]:![X29]:![X30]:![X31]:![X32]:((~hBOOL(hAPP(hAPP(c_member(X29),X30),X28))|X28!=c_Orderings_Obot__class_Obot(tc_fun(X29,tc_HOL_Obool)))&(X31=c_Orderings_Obot__class_Obot(tc_fun(X32,tc_HOL_Obool))|hBOOL(hAPP(hAPP(c_member(X32),esk9_2(X31,X32)),X31))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])])])).
cnf(c_0_166,plain,(hBOOL(hAPP(hAPP(c_member(X1),X2),hAPP(c_fequal,X3)))|X2!=X3), inference(split_conjunct,[status(thm)],[c_0_160])).
fof(c_0_167, plain, (![X28]:![X29]:hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X29,tc_HOL_Obool)),X28),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X29,tc_HOL_Obool)),X28))=c_Orderings_Otop__class_Otop(tc_fun(X29,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_41])).
cnf(c_0_168,plain,(hAPP(c_fequal,c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))!=hAPP(c_fequal,hAPP(c_fequal,X2))), inference(er,[status(thm)],[c_0_161])).
cnf(c_0_169,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(tc_Com_Opname,tc_HOL_Obool))=c_fFalse|v_Procs!=c_fTrue), inference(spm,[status(thm)],[c_0_162, c_0_163])).
cnf(c_0_170,plain,(hAPP(c_fequal,c_Int_OMin)!=hAPP(c_fequal,c_Int_OPls)), inference(spm,[status(thm)],[c_0_164, c_0_151])).
cnf(c_0_171,plain,(X1!=c_Orderings_Obot__class_Obot(tc_fun(X2,tc_HOL_Obool))|~hBOOL(hAPP(hAPP(c_member(X2),X3),X1))), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_172,plain,(hBOOL(hAPP(hAPP(c_member(X1),X2),hAPP(c_fequal,X2)))), inference(er,[status(thm)],[c_0_166])).
cnf(c_0_173,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(c_Groups_Ouminus__class_Ouminus(tc_fun(X1,tc_HOL_Obool)),X2))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_167])).
cnf(c_0_174,negated_conjecture,(hAPP(c_fequal,hAPP(c_fequal,X1))!=hAPP(c_fequal,c_fFalse)|v_Procs!=c_fTrue), inference(spm,[status(thm)],[c_0_168, c_0_169])).
cnf(c_0_175,plain,(hAPP(c_fequal,c_Int_OMin)=c_fFalse|hAPP(c_fequal,c_Int_OPls)!=c_fTrue), inference(spm,[status(thm)],[c_0_170, c_0_163])).
cnf(c_0_176,plain,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))!=hAPP(c_fequal,X2)), inference(spm,[status(thm)],[c_0_171, c_0_172])).
cnf(c_0_177,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_173, c_0_110])).
fof(c_0_178, plain, (![X11]:c_Orderings_Otop__class_Otop(tc_fun(X11,tc_HOL_Obool))!=c_Orderings_Obot__class_Obot(tc_fun(X11,tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_42])).
fof(c_0_179, plain, (![X49]:![X50]:hAPP(c_Set_OCollect(X50),hAPP(hAPP(c_COMBC(X50,X50,tc_HOL_Obool),c_fequal),X49))=hAPP(hAPP(c_Set_Oinsert(X50),X49),c_Orderings_Obot__class_Obot(tc_fun(X50,tc_HOL_Obool)))), inference(variable_rename,[status(thm)],[c_0_43])).
fof(c_0_180, plain, (~hBOOL(c_fFalse)), inference(fof_simplification,[status(thm)],[c_0_44])).
cnf(c_0_181,negated_conjecture,(hAPP(c_fequal,c_Int_OPls)!=c_fTrue|v_Procs!=c_fTrue), inference(spm,[status(thm)],[c_0_174, c_0_175])).
cnf(c_0_182,negated_conjecture,(hAPP(c_fequal,X1)!=c_fFalse|v_Procs!=c_fTrue), inference(spm,[status(thm)],[c_0_176, c_0_169])).
cnf(c_0_183,plain,(hAPP(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_fun(X1,tc_HOL_Obool)),X2),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),X2))=hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue)), inference(rw,[status(thm)],[c_0_177, c_0_123])).
cnf(c_0_184,plain,(c_Orderings_Otop__class_Otop(tc_fun(X1,tc_HOL_Obool))!=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_178])).
fof(c_0_185, plain, (![X48]:![X49]:hAPP(c_List_Olexn(X49,X48),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=c_Orderings_Obot__class_Obot(tc_fun(tc_prod(tc_List_Olist(X49),tc_List_Olist(X49)),tc_HOL_Obool))), inference(variable_rename,[status(thm)],[c_0_45])).
cnf(c_0_186,plain,(hAPP(c_Int_Onat,c_Groups_Ozero__class_Ozero(tc_Int_Oint))=c_Groups_Ozero__class_Ozero(tc_Nat_Onat)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_187,plain,(c_Int_OPls=c_Groups_Ozero__class_Ozero(tc_Int_Oint)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_188,plain,(hAPP(c_Set_OCollect(X1),hAPP(hAPP(c_COMBC(X1,X1,tc_HOL_Obool),c_fequal),X2))=hAPP(hAPP(c_Set_Oinsert(X1),X2),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))), inference(split_conjunct,[status(thm)],[c_0_179])).
cnf(c_0_189,plain,(~hBOOL(c_fFalse)), inference(split_conjunct,[status(thm)],[c_0_180])).
cnf(c_0_190,negated_conjecture,(v_Procs!=c_fTrue), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_163]), c_0_182])).
cnf(c_0_191,plain,(X1=hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue)|~hBOOL(hAPP(hAPP(c_fequal,X1),hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X2),c_fNot),X1)))), inference(spm,[status(thm)],[c_0_183, c_0_137])).
cnf(c_0_192,plain,(hAPP(hAPP(c_COMBB(tc_HOL_Obool,tc_HOL_Obool,X1),c_fNot),c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool)))=hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue)), inference(spm,[status(thm)],[c_0_130, c_0_134])).
cnf(c_0_193,plain,(hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue)!=c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))), inference(rw,[status(thm)],[c_0_184, c_0_123])).
cnf(c_0_194,plain,(hAPP(c_List_Olexn(X1,X2),c_Groups_Ozero__class_Ozero(tc_Nat_Onat))=c_Orderings_Obot__class_Obot(tc_fun(tc_prod(tc_List_Olist(X1),tc_List_Olist(X1)),tc_HOL_Obool))), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_195,plain,(c_Groups_Ozero__class_Ozero(tc_Nat_Onat)=hAPP(c_Int_Onat,c_Int_OPls)), inference(rw,[status(thm)],[c_0_186, c_0_187])).
cnf(c_0_196,plain,(hAPP(hAPP(c_COMBC(X1,X1,tc_HOL_Obool),c_fequal),X2)=hAPP(c_fequal,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_188, c_0_52]), c_0_80])).
cnf(c_0_197,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))!=v_Procs), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_121]), c_0_63]), c_0_189])).
cnf(c_0_198,negated_conjecture,(v_Procs=c_fFalse), inference(spm,[status(thm)],[c_0_190, c_0_163])).
cnf(c_0_199,plain,(~hBOOL(hAPP(hAPP(c_fequal,c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))),hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192]), c_0_193])).
cnf(c_0_200,plain,(c_Orderings_Obot__class_Obot(tc_fun(tc_prod(tc_List_Olist(X1),tc_List_Olist(X1)),tc_HOL_Obool))=hAPP(c_List_Olexn(X1,X2),hAPP(c_Int_Onat,c_Int_OPls))), inference(rw,[status(thm)],[c_0_194, c_0_195])).
cnf(c_0_201,plain,(hAPP(hAPP(c_fequal,X1),X2)=hAPP(hAPP(c_fequal,X2),X1)), inference(spm,[status(thm)],[c_0_90, c_0_196])).
cnf(c_0_202,negated_conjecture,(c_Orderings_Obot__class_Obot(tc_fun(X1,tc_HOL_Obool))!=c_fFalse), inference(rw,[status(thm)],[c_0_197, c_0_198])).
cnf(c_0_203,plain,(~hBOOL(hAPP(hAPP(c_fequal,hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue)),hAPP(c_List_Olexn(X1,X2),hAPP(c_Int_Onat,c_Int_OPls))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200]), c_0_201])).
cnf(c_0_204,plain,(hAPP(c_fequal,X1)!=c_fTrue), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_163]), c_0_202])).
cnf(c_0_205,plain,(hAPP(c_fequal,hAPP(c_List_Olexn(X1,X2),hAPP(c_Int_Onat,c_Int_OPls)))!=hAPP(c_fequal,hAPP(c_Lattices_Osemilattice__sup__class_Osup(tc_HOL_Obool),c_fTrue))), inference(spm,[status(thm)],[c_0_203, c_0_151])).
cnf(c_0_206,plain,(hAPP(c_fequal,X1)=c_fFalse), inference(spm,[status(thm)],[c_0_204, c_0_163])).
cnf(c_0_207,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_205, c_0_206]), c_0_206])]), ['proof']).
# SZS output end CNFRefutation.
EOF
