% Problem    : SWV474+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n114.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:59:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2f
# and selection function SelectCQIAr.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X12]:![X13]:![X5]:![X6]:(((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~(leq(host(X7),host(X14))))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:(host(X14)=nbr_proc=>~(elem(m_NotNorm(X7),queue(host(X14))))))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~(setIn(X7,alive))|~(setIn(X14,alive)))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X15]:![X14]:(((((host(X14)!=host(X15)&setIn(X14,alive))&host(X14)=host(X7))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X15))=>~(elem(m_Down(X7),queue(host(X15))))))&![X7]:![X15]:![X14]:(((((host(X14)!=host(X15)&setIn(X14,alive))&host(X14)=host(X7))&index(status,host(X14))=wait)&host(index(elid,host(X14)))=host(X15))=>~(elem(m_Down(X7),queue(host(X15))))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&elem(m_Down(X16),queue(host(X15)))))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&setIn(host(X16),index(down,host(X15)))))))&![X7]:![X16]:![X15]:![X14]:(((((![X17]:((~(leq(host(X14),X17))&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&elem(m_Down(X15),queue(host(X14))))&host(X14)=nbr_proc)&host(X14)=host(X16))&index(status,host(X14))=elec_1)=>~((setIn(X7,alive)&elem(m_Down(X16),queue(host(X7)))))))&queue(host(X5))=cons(m_Down(X6),X12))=>(setIn(X5,alive)=>(~(leq(host(X5),host(X6)))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>(~((![X7]:((~(leq(host(X5),X7))&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1))=>![X7]:(host(X5)!=host(X7)=>![X18]:![X19]:![X20]:(host(X5)=host(X20)=>(((((![X17]:((~(leq(host(X20),X17))&leq(s(zero),X17))=>((setIn(X17,index(down,host(X5)))|X17=host(X6))|X17=host(X19)))&elem(m_Down(X19),X12))&host(X20)=host(X18))&host(X20)=nbr_proc)&index(status,host(X20))=elec_1)=>~((setIn(X7,alive)&elem(m_Down(X18),queue(host(X7)))))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj)).
fof(c_0_1, axiom, (![X5]:![X6]:![X4]:(elem(X5,cons(X6,X4))<=>(X5=X6|elem(X5,X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_46)).
fof(c_0_2, axiom, (![X3]:leq(s(zero),host(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_02)).
fof(c_0_3, axiom, (![X5]:![X6]:((leq(X5,X6)&leq(X6,X5))<=>X5=X6)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_61)).
fof(c_0_4, axiom, (![X3]:leq(host(X3),nbr_proc)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_04)).
fof(c_0_5, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)<=>((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:(host(X14)=nbr_proc=>~elem(m_NotNorm(X7),queue(host(X14)))))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X15]:![X14]:(((((host(X14)!=host(X15)&setIn(X14,alive))&host(X14)=host(X7))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X15))=>~elem(m_Down(X7),queue(host(X15)))))&![X7]:![X15]:![X14]:(((((host(X14)!=host(X15)&setIn(X14,alive))&host(X14)=host(X7))&index(status,host(X14))=wait)&host(index(elid,host(X14)))=host(X15))=>~elem(m_Down(X7),queue(host(X15)))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&elem(m_Down(X16),queue(host(X15)))))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&setIn(host(X16),index(down,host(X15)))))))&![X7]:![X16]:![X15]:![X14]:(((((![X17]:((~leq(host(X14),X17)&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&elem(m_Down(X15),queue(host(X14))))&host(X14)=nbr_proc)&host(X14)=host(X16))&index(status,host(X14))=elec_1)=>~((setIn(X7,alive)&elem(m_Down(X16),queue(host(X7)))))))&queue(host(X5))=cons(m_Down(X6),X12)))), introduced(definition)).
fof(c_0_6, negated_conjecture, (~(![X12]:![X5]:![X6]:(epred1_3(X5,X6,X12)=>(setIn(X5,alive)=>(~leq(host(X5),host(X6))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>(~((![X7]:((~leq(host(X5),X7)&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1))=>![X7]:(host(X5)!=host(X7)=>![X18]:![X19]:![X20]:(host(X5)=host(X20)=>(((((![X17]:((~leq(host(X20),X17)&leq(s(zero),X17))=>((setIn(X17,index(down,host(X5)))|X17=host(X6))|X17=host(X19)))&elem(m_Down(X19),X12))&host(X20)=host(X18))&host(X20)=nbr_proc)&index(status,host(X20))=elec_1)=>~((setIn(X7,alive)&elem(m_Down(X18),queue(host(X7))))))))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_5])).
fof(c_0_7, negated_conjecture, (![X29]:(epred1_3(esk6_0,esk7_0,esk5_0)&(setIn(esk6_0,alive)&(~leq(host(esk6_0),host(esk7_0))&(((index(ldr,host(esk6_0))!=host(esk7_0)|index(status,host(esk6_0))!=norm)&(index(status,host(esk6_0))!=wait|host(esk7_0)!=host(index(elid,host(esk6_0)))))&((((~leq(host(esk6_0),esk8_0)|index(status,host(esk6_0))!=elec_1)&(leq(s(zero),esk8_0)|index(status,host(esk6_0))!=elec_1))&((~setIn(esk8_0,index(down,host(esk6_0)))|index(status,host(esk6_0))!=elec_1)&(esk8_0!=host(esk7_0)|index(status,host(esk6_0))!=elec_1)))&(host(esk6_0)!=host(esk9_0)&(host(esk6_0)=host(esk12_0)&(((((((leq(host(esk12_0),X29)|~leq(s(zero),X29))|((setIn(X29,index(down,host(esk6_0)))|X29=host(esk7_0))|X29=host(esk11_0)))&elem(m_Down(esk11_0),esk5_0))&host(esk12_0)=host(esk10_0))&host(esk12_0)=nbr_proc)&index(status,host(esk12_0))=elec_1)&(setIn(esk9_0,alive)&elem(m_Down(esk10_0),queue(host(esk9_0))))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
fof(c_0_8, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)=>((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:(host(X14)=nbr_proc=>~elem(m_NotNorm(X7),queue(host(X14)))))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X15]:![X14]:(((((host(X14)!=host(X15)&setIn(X14,alive))&host(X14)=host(X7))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X15))=>~elem(m_Down(X7),queue(host(X15)))))&![X7]:![X15]:![X14]:(((((host(X14)!=host(X15)&setIn(X14,alive))&host(X14)=host(X7))&index(status,host(X14))=wait)&host(index(elid,host(X14)))=host(X15))=>~elem(m_Down(X7),queue(host(X15)))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&elem(m_Down(X16),queue(host(X15)))))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&setIn(host(X16),index(down,host(X15)))))))&![X7]:![X16]:![X15]:![X14]:(((((![X17]:((~leq(host(X14),X17)&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&elem(m_Down(X15),queue(host(X14))))&host(X14)=nbr_proc)&host(X14)=host(X16))&index(status,host(X14))=elec_1)=>~((setIn(X7,alive)&elem(m_Down(X16),queue(host(X7)))))))&queue(host(X5))=cons(m_Down(X6),X12)))), inference(split_equiv,[status(thm)],[c_0_5])).
cnf(c_0_9,negated_conjecture,(host(esk6_0)=host(esk12_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10,negated_conjecture,(host(esk12_0)=nbr_proc), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11,negated_conjecture,(index(status,host(esk12_0))=elec_1), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_12, plain, (![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:![X34]:![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:![X43]:![X44]:![X45]:![X46]:![X47]:((((((((((((~elem(m_Ldr(X22),queue(host(X21)))|~leq(host(X21),host(X22)))|~epred1_3(X20,X19,X18))&((~elem(m_Down(X24),queue(host(X23)))|host(X24)!=host(X23))|~epred1_3(X20,X19,X18)))&((host(X26)!=nbr_proc|~elem(m_NotNorm(X25),queue(host(X26))))|~epred1_3(X20,X19,X18)))&(((X28=X27|host(X28)!=host(X27))|(~setIn(X27,alive)|~setIn(X28,alive)))|~epred1_3(X20,X19,X18)))&((((index(status,host(X29))!=elec_1|~setIn(X29,alive))|index(elid,host(X29))=X29)|~epred1_3(X20,X19,X18))&(((index(status,host(X29))!=elec_2|~setIn(X29,alive))|index(elid,host(X29))=X29)|~epred1_3(X20,X19,X18))))&((((((host(X32)=host(X31)|~setIn(X32,alive))|host(X32)!=host(X30))|index(status,host(X32))!=norm)|index(ldr,host(X32))!=host(X31))|~elem(m_Down(X30),queue(host(X31))))|~epred1_3(X20,X19,X18)))&((((((host(X35)=host(X34)|~setIn(X35,alive))|host(X35)!=host(X33))|index(status,host(X35))!=wait)|host(index(elid,host(X35)))!=host(X34))|~elem(m_Down(X33),queue(host(X34))))|~epred1_3(X20,X19,X18)))&((((((host(X38)=host(X36)|~setIn(X36,alive))|~setIn(X38,alive))|host(X37)!=host(X36))|host(X39)!=host(X38))|(~elem(m_Down(X39),queue(host(X36)))|~elem(m_Down(X37),queue(host(X38)))))|~epred1_3(X20,X19,X18)))&((((((host(X42)=host(X40)|~setIn(X40,alive))|~setIn(X42,alive))|host(X41)!=host(X40))|host(X43)!=host(X42))|(~elem(m_Down(X43),queue(host(X40)))|~setIn(host(X41),index(down,host(X42)))))|~epred1_3(X20,X19,X18)))&((((((((~leq(host(X47),esk13_7(X18,X19,X20,X44,X45,X46,X47))|~elem(m_Down(X46),queue(host(X47))))|host(X47)!=nbr_proc)|host(X47)!=host(X45))|index(status,host(X47))!=elec_1)|(~setIn(X44,alive)|~elem(m_Down(X45),queue(host(X44)))))|~epred1_3(X20,X19,X18))&((((((leq(s(zero),esk13_7(X18,X19,X20,X44,X45,X46,X47))|~elem(m_Down(X46),queue(host(X47))))|host(X47)!=nbr_proc)|host(X47)!=host(X45))|index(status,host(X47))!=elec_1)|(~setIn(X44,alive)|~elem(m_Down(X45),queue(host(X44)))))|~epred1_3(X20,X19,X18)))&(((((((~setIn(esk13_7(X18,X19,X20,X44,X45,X46,X47),index(down,host(X47)))|~elem(m_Down(X46),queue(host(X47))))|host(X47)!=nbr_proc)|host(X47)!=host(X45))|index(status,host(X47))!=elec_1)|(~setIn(X44,alive)|~elem(m_Down(X45),queue(host(X44)))))|~epred1_3(X20,X19,X18))&((((((esk13_7(X18,X19,X20,X44,X45,X46,X47)!=host(X46)|~elem(m_Down(X46),queue(host(X47))))|host(X47)!=nbr_proc)|host(X47)!=host(X45))|index(status,host(X47))!=elec_1)|(~setIn(X44,alive)|~elem(m_Down(X45),queue(host(X44)))))|~epred1_3(X20,X19,X18)))))&(queue(host(X20))=cons(m_Down(X19),X18)|~epred1_3(X20,X19,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])])).
cnf(c_0_13,negated_conjecture,(host(esk12_0)=host(esk10_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14,negated_conjecture,(X1=host(esk11_0)|X1=host(esk7_0)|setIn(X1,index(down,host(esk6_0)))|leq(host(esk12_0),X1)|~leq(s(zero),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15,negated_conjecture,(host(esk6_0)=nbr_proc), inference(rw,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_16,negated_conjecture,(leq(s(zero),esk8_0)|index(status,host(esk6_0))!=elec_1), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17,negated_conjecture,(index(status,nbr_proc)=elec_1), inference(rw,[status(thm)],[c_0_11, c_0_10])).
cnf(c_0_18,negated_conjecture,(index(status,host(esk6_0))!=elec_1|esk8_0!=host(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19,negated_conjecture,(index(status,host(esk6_0))!=elec_1|~leq(host(esk6_0),esk8_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20,negated_conjecture,(index(status,host(esk6_0))!=elec_1|~setIn(esk8_0,index(down,host(esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_21, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((~elem(X7,cons(X8,X9))|(X7=X8|elem(X7,X9)))&((X10!=X11|elem(X10,cons(X11,X12)))&(~elem(X10,X12)|elem(X10,cons(X11,X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
cnf(c_0_22,plain,(host(X5)=host(X7)|~epred1_3(X1,X2,X3)|~setIn(host(X4),index(down,host(X5)))|~elem(m_Down(X6),queue(host(X7)))|host(X6)!=host(X5)|host(X4)!=host(X7)|~setIn(X5,alive)|~setIn(X7,alive)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,negated_conjecture,(elem(m_Down(esk10_0),queue(host(esk9_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24,negated_conjecture,(host(esk10_0)=nbr_proc), inference(rw,[status(thm)],[c_0_13, c_0_10])).
cnf(c_0_25,negated_conjecture,(setIn(esk9_0,alive)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_26,negated_conjecture,(host(esk6_0)!=host(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27,negated_conjecture,(host(esk11_0)=X1|host(esk7_0)=X1|leq(nbr_proc,X1)|setIn(X1,index(down,nbr_proc))|~leq(s(zero),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_10]), c_0_15])).
cnf(c_0_28,negated_conjecture,(leq(s(zero),esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_15]), c_0_17])])).
cnf(c_0_29,negated_conjecture,(host(esk7_0)!=esk8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_15]), c_0_17])])).
cnf(c_0_30,negated_conjecture,(~leq(nbr_proc,esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_15]), c_0_17]), c_0_15])])).
cnf(c_0_31,negated_conjecture,(~setIn(esk8_0,index(down,nbr_proc))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_15]), c_0_17]), c_0_15])])).
cnf(c_0_32,plain,(host(X5)=host(X7)|~epred1_3(X1,X2,X3)|~elem(m_Down(X4),queue(host(X5)))|~elem(m_Down(X6),queue(host(X7)))|host(X6)!=host(X5)|host(X4)!=host(X7)|~setIn(X5,alive)|~setIn(X7,alive)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_33,plain,(elem(X1,cons(X2,X3))|~elem(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34,negated_conjecture,(elem(m_Down(esk11_0),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_35,plain,(queue(host(X1))=cons(m_Down(X2),X3)|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_36,negated_conjecture,(epred1_3(esk6_0,esk7_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_37,plain,(elem(X1,cons(X2,X3))|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_38,negated_conjecture,(host(X1)=host(esk9_0)|host(X2)!=host(esk9_0)|host(X1)!=nbr_proc|~epred1_3(X3,X4,X5)|~setIn(host(X2),index(down,host(X1)))|~setIn(X1,alive)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_39,negated_conjecture,(setIn(esk6_0,alive)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_40,negated_conjecture,(host(esk9_0)!=nbr_proc), inference(rw,[status(thm)],[c_0_26, c_0_15])).
cnf(c_0_41,negated_conjecture,(host(esk11_0)=esk8_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30]), c_0_31])).
fof(c_0_42, plain, (![X4]:leq(s(zero),host(X4))), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_43,negated_conjecture,(host(X1)=host(esk9_0)|host(X2)!=host(esk9_0)|host(X1)!=nbr_proc|~epred1_3(X3,X4,X5)|~setIn(X1,alive)|~elem(m_Down(X2),queue(host(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_44,negated_conjecture,(elem(m_Down(esk11_0),cons(X1,esk5_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_45,negated_conjecture,(cons(m_Down(esk7_0),esk5_0)=queue(nbr_proc)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_15])).
cnf(c_0_46,plain,(elem(X1,cons(X1,X2))), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_47,negated_conjecture,(host(X1)!=host(esk9_0)|~epred1_3(X2,X3,X4)|~setIn(host(X1),index(down,nbr_proc))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_15]), c_0_39])]), c_0_40])).
cnf(c_0_48,negated_conjecture,(host(esk7_0)=X1|esk8_0=X1|leq(nbr_proc,X1)|setIn(X1,index(down,nbr_proc))|~leq(s(zero),X1)), inference(rw,[status(thm)],[c_0_27, c_0_41])).
cnf(c_0_49,plain,(leq(s(zero),host(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50,negated_conjecture,(host(X1)!=host(esk9_0)|~epred1_3(X2,X3,X4)|~elem(m_Down(X1),queue(nbr_proc))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_15]), c_0_39])]), c_0_40])).
cnf(c_0_51,negated_conjecture,(elem(m_Down(esk11_0),queue(nbr_proc))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_52,negated_conjecture,(elem(m_Down(esk7_0),queue(nbr_proc))), inference(spm,[status(thm)],[c_0_46, c_0_45])).
cnf(c_0_53,negated_conjecture,(~epred1_3(X1,X2,X3)|~setIn(host(esk9_0),index(down,nbr_proc))), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_54,negated_conjecture,(host(esk7_0)=host(X1)|host(X1)=esk8_0|leq(nbr_proc,host(X1))|setIn(host(X1),index(down,nbr_proc))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_55,negated_conjecture,(host(esk9_0)!=esk8_0|~epred1_3(X1,X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_41])).
cnf(c_0_56,negated_conjecture,(host(esk9_0)!=host(esk7_0)|~epred1_3(X1,X2,X3)), inference(spm,[status(thm)],[c_0_50, c_0_52])).
fof(c_0_57, plain, (![X7]:![X8]:![X9]:![X10]:(((~leq(X7,X8)|~leq(X8,X7))|X7=X8)&((leq(X9,X10)|X9!=X10)&(leq(X10,X9)|X9!=X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_58,negated_conjecture,(leq(nbr_proc,host(esk9_0))|~epred1_3(X1,X2,X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56])).
fof(c_0_59, plain, (![X4]:leq(host(X4),nbr_proc)), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_60,plain,(X1=X2|~leq(X2,X1)|~leq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_61,negated_conjecture,(leq(nbr_proc,host(esk9_0))), inference(spm,[status(thm)],[c_0_58, c_0_36])).
cnf(c_0_62,plain,(leq(host(X1),nbr_proc)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_63,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])]), c_0_40]), ['proof']).
# SZS output end CNFRefutation.
EOF
