% Problem    : SWV455+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n106.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:03:46 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X12]:![X13]:![X5]:![X6]:(((((((((![X7]:![X14]:(setIn(X14,alive)=>~(elem(m_Down(X14),queue(host(X7)))))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>~(setIn(X14,alive))))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:(elem(m_Halt(X14),queue(host(X7)))=>~(leq(host(X7),host(X14)))))&![X7]:![X15]:![X14]:(elem(m_Ack(X14,X7),queue(host(X15)))=>~(leq(host(X7),host(X14)))))&![X7]:![X14]:(((~(setIn(X7,alive))&leq(X14,X7))&host(X14)=host(X7))=>~(setIn(X14,alive))))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~(setIn(X7,alive))|~(setIn(X14,alive)))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&elem(m_Down(X16),queue(host(X15)))))))&queue(host(X5))=cons(m_Down(X6),X12))=>(setIn(X5,alive)=>(~(leq(host(X5),host(X6)))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>(~((![X7]:((~(leq(host(X5),X7))&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1))=>![X7]:(host(X5)!=host(X7)=>![X17]:![X18]:(host(X5)=host(X18)=>![X19]:(((((host(X18)!=host(X7)&setIn(X7,alive))&setIn(X18,alive))&host(X17)=host(X7))&host(X19)=host(X18))=>~((elem(m_Down(X17),X12)&elem(m_Down(X19),queue(host(X7)))))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj)).
fof(c_0_1, axiom, (![X5]:![X6]:![X4]:(elem(X5,cons(X6,X4))<=>(X5=X6|elem(X5,X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_46)).
fof(c_0_2, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)<=>((((((((![X7]:![X14]:(setIn(X14,alive)=>~elem(m_Down(X14),queue(host(X7))))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>~setIn(X14,alive)))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:(elem(m_Halt(X14),queue(host(X7)))=>~leq(host(X7),host(X14))))&![X7]:![X15]:![X14]:(elem(m_Ack(X14,X7),queue(host(X15)))=>~leq(host(X7),host(X14))))&![X7]:![X14]:(((~setIn(X7,alive)&leq(X14,X7))&host(X14)=host(X7))=>~setIn(X14,alive)))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&elem(m_Down(X16),queue(host(X15)))))))&queue(host(X5))=cons(m_Down(X6),X12)))), introduced(definition)).
fof(c_0_3, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)=>((((((((![X7]:![X14]:(setIn(X14,alive)=>~elem(m_Down(X14),queue(host(X7))))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>~setIn(X14,alive)))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:(elem(m_Halt(X14),queue(host(X7)))=>~leq(host(X7),host(X14))))&![X7]:![X15]:![X14]:(elem(m_Ack(X14,X7),queue(host(X15)))=>~leq(host(X7),host(X14))))&![X7]:![X14]:(((~setIn(X7,alive)&leq(X14,X7))&host(X14)=host(X7))=>~setIn(X14,alive)))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&elem(m_Down(X16),queue(host(X15)))))))&queue(host(X5))=cons(m_Down(X6),X12)))), inference(split_equiv,[status(thm)],[c_0_2])).
fof(c_0_4, negated_conjecture, (~(![X12]:![X5]:![X6]:(epred1_3(X5,X6,X12)=>(setIn(X5,alive)=>(~leq(host(X5),host(X6))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>(~((![X7]:((~leq(host(X5),X7)&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1))=>![X7]:(host(X5)!=host(X7)=>![X17]:![X18]:(host(X5)=host(X18)=>![X19]:(((((host(X18)!=host(X7)&setIn(X7,alive))&setIn(X18,alive))&host(X17)=host(X7))&host(X19)=host(X18))=>~((elem(m_Down(X17),X12)&elem(m_Down(X19),queue(host(X7))))))))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_2])).
fof(c_0_5, plain, (![X17]:![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:![X34]:![X35]:![X36]:![X37]:![X38]:((((((((((~setIn(X21,alive)|~elem(m_Down(X21),queue(host(X20))))|~epred1_3(X19,X18,X17))&((~elem(m_Down(X23),queue(host(X22)))|~setIn(X23,alive))|~epred1_3(X19,X18,X17)))&((~elem(m_Down(X25),queue(host(X24)))|host(X25)!=host(X24))|~epred1_3(X19,X18,X17)))&((~elem(m_Halt(X27),queue(host(X26)))|~leq(host(X26),host(X27)))|~epred1_3(X19,X18,X17)))&((~elem(m_Ack(X30,X28),queue(host(X29)))|~leq(host(X28),host(X30)))|~epred1_3(X19,X18,X17)))&((((setIn(X31,alive)|~leq(X32,X31))|host(X32)!=host(X31))|~setIn(X32,alive))|~epred1_3(X19,X18,X17)))&(((X34=X33|host(X34)!=host(X33))|(~setIn(X33,alive)|~setIn(X34,alive)))|~epred1_3(X19,X18,X17)))&((((((host(X37)=host(X35)|~setIn(X35,alive))|~setIn(X37,alive))|host(X36)!=host(X35))|host(X38)!=host(X37))|(~elem(m_Down(X38),queue(host(X35)))|~elem(m_Down(X36),queue(host(X37)))))|~epred1_3(X19,X18,X17)))&(queue(host(X19))=cons(m_Down(X18),X17)|~epred1_3(X19,X18,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
fof(c_0_6, negated_conjecture, ((epred1_3(esk6_0,esk7_0,esk5_0)&(setIn(esk6_0,alive)&(~leq(host(esk6_0),host(esk7_0))&(((index(ldr,host(esk6_0))!=host(esk7_0)|index(status,host(esk6_0))!=norm)&(index(status,host(esk6_0))!=wait|host(esk7_0)!=host(index(elid,host(esk6_0)))))&((((~leq(host(esk6_0),esk8_0)|index(status,host(esk6_0))!=elec_1)&(leq(s(zero),esk8_0)|index(status,host(esk6_0))!=elec_1))&((~setIn(esk8_0,index(down,host(esk6_0)))|index(status,host(esk6_0))!=elec_1)&(esk8_0!=host(esk7_0)|index(status,host(esk6_0))!=elec_1)))&(host(esk6_0)!=host(esk9_0)&(host(esk6_0)=host(esk11_0)&(((((host(esk11_0)!=host(esk9_0)&setIn(esk9_0,alive))&setIn(esk11_0,alive))&host(esk10_0)=host(esk9_0))&host(esk12_0)=host(esk11_0))&(elem(m_Down(esk10_0),esk5_0)&elem(m_Down(esk12_0),queue(host(esk9_0))))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_7,plain,(X4=X5|~epred1_3(X1,X2,X3)|~setIn(X4,alive)|~setIn(X5,alive)|host(X4)!=host(X5)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8,negated_conjecture,(epred1_3(esk6_0,esk7_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,plain,(host(X5)=host(X7)|~epred1_3(X1,X2,X3)|~elem(m_Down(X4),queue(host(X5)))|~elem(m_Down(X6),queue(host(X7)))|host(X6)!=host(X5)|host(X4)!=host(X7)|~setIn(X5,alive)|~setIn(X7,alive)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10,negated_conjecture,(host(esk10_0)=host(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11,negated_conjecture,(setIn(esk9_0,alive)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12,negated_conjecture,(elem(m_Down(esk12_0),queue(host(esk9_0)))), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_13, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((~elem(X7,cons(X8,X9))|(X7=X8|elem(X7,X9)))&((X10!=X11|elem(X10,cons(X11,X12)))&(~elem(X10,X12)|elem(X10,cons(X11,X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
cnf(c_0_14,plain,(queue(host(X1))=cons(m_Down(X2),X3)|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15,negated_conjecture,(host(esk6_0)=host(esk11_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16,negated_conjecture,(X1=X2|host(X1)!=host(X2)|~setIn(X2,alive)|~setIn(X1,alive)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_17,negated_conjecture,(setIn(esk11_0,alive)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_18,negated_conjecture,(host(X1)=host(esk10_0)|host(X2)!=host(esk10_0)|host(X1)!=host(X3)|~epred1_3(X4,X5,X6)|~setIn(X1,alive)|~elem(m_Down(X3),queue(host(esk10_0)))|~elem(m_Down(X2),queue(host(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])])).
cnf(c_0_19,negated_conjecture,(elem(m_Down(esk12_0),queue(host(esk10_0)))), inference(rw,[status(thm)],[c_0_12, c_0_10])).
cnf(c_0_20,negated_conjecture,(host(esk12_0)=host(esk11_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21,plain,(elem(X1,cons(X2,X3))|~elem(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22,negated_conjecture,(cons(m_Down(esk7_0),esk5_0)=queue(host(esk11_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_8]), c_0_15])).
cnf(c_0_23,negated_conjecture,(host(esk11_0)!=host(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24,negated_conjecture,(X1=esk11_0|host(X1)!=host(esk11_0)|~setIn(X1,alive)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_25,negated_conjecture,(setIn(esk6_0,alive)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26,negated_conjecture,(host(X1)=host(esk10_0)|host(X2)!=host(esk10_0)|host(X1)!=host(esk11_0)|~epred1_3(X3,X4,X5)|~setIn(X1,alive)|~elem(m_Down(X2),queue(host(X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_27,negated_conjecture,(elem(X1,queue(host(esk11_0)))|~elem(X1,esk5_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28,negated_conjecture,(host(esk11_0)!=host(esk10_0)), inference(rw,[status(thm)],[c_0_23, c_0_10])).
cnf(c_0_29,negated_conjecture,(esk6_0=esk11_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_15])])).
cnf(c_0_30,negated_conjecture,(host(X1)!=host(esk10_0)|~epred1_3(X2,X3,X4)|~elem(m_Down(X1),esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_17])]), c_0_28])).
cnf(c_0_31,negated_conjecture,(elem(m_Down(esk10_0),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_32,negated_conjecture,(epred1_3(esk11_0,esk7_0,esk5_0)), inference(rw,[status(thm)],[c_0_8, c_0_29])).
cnf(c_0_33,negated_conjecture,(~epred1_3(X1,X2,X3)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34,negated_conjecture,($false), inference(sr,[status(thm)],[c_0_32, c_0_33]), ['proof']).
# SZS output end CNFRefutation.
EOF
