% Problem    : SET603+4 : TPTP v6.1.0. Released v2.2.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n123.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:31:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_B02_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X4]:equal_set(difference(X4,empty_set),X4)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', thI30)).
fof(c_0_1, axiom, (![X1]:![X2]:(equal_set(X1,X2)<=>(subset(X1,X2)&subset(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SET006+0.ax', equal_set)).
fof(c_0_2, axiom, (![X1]:![X2]:(subset(X1,X2)<=>![X3]:(member(X3,X1)=>member(X3,X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SET006+0.ax', subset)).
fof(c_0_3, axiom, (![X2]:![X1]:![X4]:(member(X2,difference(X4,X1))<=>(member(X2,X4)&~(member(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SET006+0.ax', difference)).
fof(c_0_4, axiom, (![X3]:~(member(X3,empty_set))), file('/export/starexec/sandbox/benchmark/Axioms/SET006+0.ax', empty_set)).
fof(c_0_5, negated_conjecture, (~(![X4]:equal_set(difference(X4,empty_set),X4))), inference(assume_negation,[status(cth)],[c_0_0])).
fof(c_0_6, negated_conjecture, (~equal_set(difference(esk4_0,empty_set),esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
fof(c_0_7, plain, (![X3]:![X4]:![X5]:![X6]:(((subset(X3,X4)|~equal_set(X3,X4))&(subset(X4,X3)|~equal_set(X3,X4)))&((~subset(X5,X6)|~subset(X6,X5))|equal_set(X5,X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
cnf(c_0_8,negated_conjecture,(~equal_set(difference(esk4_0,empty_set),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9,plain,(equal_set(X1,X2)|~subset(X2,X1)|~subset(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:![X7]:![X8]:((~subset(X4,X5)|(~member(X6,X4)|member(X6,X5)))&((member(esk1_2(X7,X8),X7)|subset(X7,X8))&(~member(esk1_2(X7,X8),X8)|subset(X7,X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
fof(c_0_11, plain, (![X2]:![X1]:![X4]:(member(X2,difference(X4,X1))<=>(member(X2,X4)&~member(X2,X1)))), inference(fof_simplification,[status(thm)],[c_0_3])).
cnf(c_0_12,negated_conjecture,(~subset(esk4_0,difference(esk4_0,empty_set))|~subset(difference(esk4_0,empty_set),esk4_0)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13,plain,(subset(X1,X2)|member(esk1_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14,plain,(subset(X1,X2)|~member(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, (![X5]:![X6]:![X7]:![X8]:![X9]:![X10]:(((member(X5,X7)|~member(X5,difference(X7,X6)))&(~member(X5,X6)|~member(X5,difference(X7,X6))))&((~member(X8,X10)|member(X8,X9))|member(X8,difference(X10,X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
cnf(c_0_16,negated_conjecture,(member(esk1_2(difference(esk4_0,empty_set),esk4_0),difference(esk4_0,empty_set))|~subset(esk4_0,difference(esk4_0,empty_set))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17,negated_conjecture,(~member(esk1_2(difference(esk4_0,empty_set),esk4_0),esk4_0)|~subset(esk4_0,difference(esk4_0,empty_set))), inference(spm,[status(thm)],[c_0_12, c_0_14])).
fof(c_0_18, plain, (![X3]:~member(X3,empty_set)), inference(fof_simplification,[status(thm)],[c_0_4])).
cnf(c_0_19,plain,(member(X1,X2)|~member(X1,difference(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20,negated_conjecture,(member(esk1_2(difference(esk4_0,empty_set),esk4_0),difference(esk4_0,empty_set))|~member(esk1_2(esk4_0,difference(esk4_0,empty_set)),difference(esk4_0,empty_set))), inference(spm,[status(thm)],[c_0_16, c_0_14])).
cnf(c_0_21,negated_conjecture,(~member(esk1_2(esk4_0,difference(esk4_0,empty_set)),difference(esk4_0,empty_set))|~member(esk1_2(difference(esk4_0,empty_set),esk4_0),esk4_0)), inference(spm,[status(thm)],[c_0_17, c_0_14])).
fof(c_0_22, plain, (![X4]:~member(X4,empty_set)), inference(variable_rename,[status(thm)],[c_0_18])).
cnf(c_0_23,negated_conjecture,(~member(esk1_2(esk4_0,difference(esk4_0,empty_set)),difference(esk4_0,empty_set))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_24,plain,(member(X1,difference(X2,X3))|member(X1,X3)|~member(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25,plain,(~member(X1,empty_set)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26,negated_conjecture,(~member(esk1_2(esk4_0,difference(esk4_0,empty_set)),esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_27,negated_conjecture,(member(esk1_2(difference(esk4_0,empty_set),esk4_0),difference(esk4_0,empty_set))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_13]), c_0_26])).
cnf(c_0_28,negated_conjecture,(member(esk1_2(esk4_0,difference(esk4_0,empty_set)),esk4_0)|~member(esk1_2(difference(esk4_0,empty_set),esk4_0),esk4_0)), inference(spm,[status(thm)],[c_0_17, c_0_13])).
cnf(c_0_29,negated_conjecture,(member(esk1_2(difference(esk4_0,empty_set),esk4_0),esk4_0)), inference(spm,[status(thm)],[c_0_19, c_0_27])).
cnf(c_0_30,negated_conjecture,($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29])]), c_0_26]), ['proof']).
# SZS output end CNFRefutation.
EOF
