% Problem    : NUM924+6 : TPTP v6.1.0. Released v5.3.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n119.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:21:51 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X14]:plus_plus(int,X14,pls)=ti(int,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_127_add__Pls__right)).
fof(c_0_1, axiom, (![X14]:number_number_of(int,X14)=ti(int,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_22_number__of__is__id)).
fof(c_0_2, hypothesis, (![X1]:bit0(ti(int,X1))=bit0(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tsy_c_Int_OBit0_arg1)).
fof(c_0_3, axiom, (![X14]:bit0(X14)=plus_plus(int,X14,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_130_Bit0__def)).
fof(c_0_4, axiom, (![X12]:![X11]:plus_plus(int,number_number_of(int,X12),number_number_of(int,X11))=number_number_of(int,plus_plus(int,X12,X11))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_67_plus__numeral__code_I9_J)).
fof(c_0_5, axiom, (![X15]:![X11]:plus_plus(int,X15,X11)=plus_plus(int,X11,X15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_96_zadd__commute)).
fof(c_0_6, axiom, (![X14]:bit1(X14)=plus_plus(int,plus_plus(int,one_one(int),X14),X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_162_Bit1__def)).
fof(c_0_7, axiom, (![X14]:plus_plus(int,pls,X14)=ti(int,X14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_128_add__Pls)).
fof(c_0_8, conjecture, (ord_less(int,plus_plus(int,power_power(int,s,number_number_of(nat,bit0(bit1(pls)))),one_one(int)),zero_zero(int))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj_0)).
fof(c_0_9, axiom, (plus_plus(int,power_power(int,s,number_number_of(nat,bit0(bit1(pls)))),one_one(int))=times_times(int,plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)),t)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_3_t)).
fof(c_0_10, axiom, (![X18]:![X19]:![X20]:plus_plus(int,plus_plus(int,X18,X19),X20)=plus_plus(int,X18,plus_plus(int,X19,X20))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_94_zadd__assoc)).
fof(c_0_11, axiom, (![X5]:![X6]:![X15]:plus_plus(int,X5,plus_plus(int,X6,X15))=plus_plus(int,X6,plus_plus(int,X5,X15))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_95_zadd__left__commute)).
fof(c_0_12, axiom, (![X15]:![X11]:times_times(int,X15,X11)=times_times(int,X11,X15)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_23_zmult__commute)).
fof(c_0_13, axiom, (ord_less(int,times_times(int,plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)),t),times_times(int,plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)),zero_zero(int)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_2__096_I4_A_K_Am_A_L_A1_J_A_K_At_A_060_A_I4_A_K_Am_A_L_A1_J_A_K_A0_096)).
fof(c_0_14, axiom, (![X11]:times_times(int,pls,X11)=pls), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_62_mult__Pls)).
fof(c_0_15, axiom, (pls=zero_zero(int)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', fact_120_Pls__def)).
fof(c_0_16, plain, (![X15]:plus_plus(int,X15,pls)=ti(int,X15)), inference(variable_rename,[status(thm)],[c_0_0])).
fof(c_0_17, plain, (![X15]:number_number_of(int,X15)=ti(int,X15)), inference(variable_rename,[status(thm)],[c_0_1])).
fof(c_0_18, hypothesis, (![X2]:bit0(ti(int,X2))=bit0(X2)), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_19, plain, (![X15]:bit0(X15)=plus_plus(int,X15,X15)), inference(variable_rename,[status(thm)],[c_0_3])).
fof(c_0_20, plain, (![X13]:![X14]:plus_plus(int,number_number_of(int,X13),number_number_of(int,X14))=number_number_of(int,plus_plus(int,X13,X14))), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_21,plain,(plus_plus(int,X1,pls)=ti(int,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(number_number_of(int,X1)=ti(int,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, (![X16]:![X17]:plus_plus(int,X16,X17)=plus_plus(int,X17,X16)), inference(variable_rename,[status(thm)],[c_0_5])).
fof(c_0_24, plain, (![X15]:bit1(X15)=plus_plus(int,plus_plus(int,one_one(int),X15),X15)), inference(variable_rename,[status(thm)],[c_0_6])).
fof(c_0_25, plain, (![X15]:plus_plus(int,pls,X15)=ti(int,X15)), inference(variable_rename,[status(thm)],[c_0_7])).
cnf(c_0_26,hypothesis,(bit0(ti(int,X1))=bit0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27,plain,(bit0(X1)=plus_plus(int,X1,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28,plain,(plus_plus(int,number_number_of(int,X1),number_number_of(int,X2))=number_number_of(int,plus_plus(int,X1,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29,plain,(number_number_of(int,X1)=plus_plus(int,X1,pls)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_30,plain,(plus_plus(int,X1,X2)=plus_plus(int,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_31, negated_conjecture, (~ord_less(int,plus_plus(int,power_power(int,s,number_number_of(nat,bit0(bit1(pls)))),one_one(int)),zero_zero(int))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_8])])).
cnf(c_0_32,plain,(plus_plus(int,power_power(int,s,number_number_of(nat,bit0(bit1(pls)))),one_one(int))=times_times(int,plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)),t)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_33,plain,(bit1(X1)=plus_plus(int,plus_plus(int,one_one(int),X1),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34,plain,(plus_plus(int,pls,X1)=ti(int,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_35, plain, (![X21]:![X22]:![X23]:plus_plus(int,plus_plus(int,X21,X22),X23)=plus_plus(int,X21,plus_plus(int,X22,X23))), inference(variable_rename,[status(thm)],[c_0_10])).
fof(c_0_36, plain, (![X16]:![X17]:![X18]:plus_plus(int,X16,plus_plus(int,X17,X18))=plus_plus(int,X17,plus_plus(int,X16,X18))), inference(variable_rename,[status(thm)],[c_0_11])).
cnf(c_0_37,hypothesis,(plus_plus(int,ti(int,X1),ti(int,X1))=plus_plus(int,X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_27])).
cnf(c_0_38,plain,(plus_plus(int,plus_plus(int,X1,pls),plus_plus(int,X2,pls))=plus_plus(int,pls,plus_plus(int,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29]), c_0_29]), c_0_29]), c_0_30])).
fof(c_0_39, plain, (![X16]:![X17]:times_times(int,X16,X17)=times_times(int,X17,X16)), inference(variable_rename,[status(thm)],[c_0_12])).
cnf(c_0_40,negated_conjecture,(~ord_less(int,plus_plus(int,power_power(int,s,number_number_of(nat,bit0(bit1(pls)))),one_one(int)),zero_zero(int))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41,plain,(ord_less(int,times_times(int,plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)),t),times_times(int,plus_plus(int,times_times(int,number_number_of(int,bit0(bit0(bit1(pls)))),m),one_one(int)),zero_zero(int)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_42,plain,(times_times(int,plus_plus(int,times_times(int,number_number_of(int,plus_plus(int,plus_plus(int,plus_plus(int,plus_plus(int,one_one(int),pls),pls),plus_plus(int,plus_plus(int,one_one(int),pls),pls)),plus_plus(int,plus_plus(int,plus_plus(int,one_one(int),pls),pls),plus_plus(int,plus_plus(int,one_one(int),pls),pls)))),m),one_one(int)),t)=plus_plus(int,power_power(int,s,number_number_of(nat,plus_plus(int,plus_plus(int,plus_plus(int,one_one(int),pls),pls),plus_plus(int,plus_plus(int,one_one(int),pls),pls)))),one_one(int))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_27]), c_0_27]), c_0_27]), c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_43,plain,(plus_plus(int,X1,pls)=plus_plus(int,pls,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_22]), c_0_29])).
cnf(c_0_44,plain,(plus_plus(int,plus_plus(int,X1,X2),X3)=plus_plus(int,X1,plus_plus(int,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45,plain,(plus_plus(int,X1,plus_plus(int,X2,X3))=plus_plus(int,X2,plus_plus(int,X1,X3))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46,hypothesis,(plus_plus(int,pls,plus_plus(int,X1,X1))=plus_plus(int,X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_22]), c_0_29]), c_0_22]), c_0_29]), c_0_38])).
cnf(c_0_47,plain,(times_times(int,X1,X2)=times_times(int,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_48, plain, (![X12]:times_times(int,pls,X12)=pls), inference(variable_rename,[status(thm)],[c_0_14])).
cnf(c_0_49,negated_conjecture,(~ord_less(int,plus_plus(int,power_power(int,s,number_number_of(nat,plus_plus(int,plus_plus(int,plus_plus(int,one_one(int),pls),pls),plus_plus(int,plus_plus(int,one_one(int),pls),pls)))),one_one(int)),zero_zero(int))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_27]), c_0_33]), c_0_33])).
cnf(c_0_50,plain,(pls=zero_zero(int)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_51,plain,(ord_less(int,times_times(int,plus_plus(int,times_times(int,number_number_of(int,plus_plus(int,plus_plus(int,plus_plus(int,plus_plus(int,one_one(int),pls),pls),plus_plus(int,plus_plus(int,one_one(int),pls),pls)),plus_plus(int,plus_plus(int,plus_plus(int,one_one(int),pls),pls),plus_plus(int,plus_plus(int,one_one(int),pls),pls)))),m),one_one(int)),t),times_times(int,plus_plus(int,times_times(int,number_number_of(int,plus_plus(int,plus_plus(int,plus_plus(int,plus_plus(int,one_one(int),pls),pls),plus_plus(int,plus_plus(int,one_one(int),pls),pls)),plus_plus(int,plus_plus(int,plus_plus(int,one_one(int),pls),pls),plus_plus(int,plus_plus(int,one_one(int),pls),pls)))),m),one_one(int)),zero_zero(int)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_27]), c_0_27]), c_0_27]), c_0_27]), c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_52,plain,(times_times(int,t,plus_plus(int,one_one(int),times_times(int,m,plus_plus(int,one_one(int),plus_plus(int,one_one(int),plus_plus(int,one_one(int),one_one(int)))))))=plus_plus(int,one_one(int),power_power(int,s,number_number_of(nat,plus_plus(int,one_one(int),one_one(int)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_44]), c_0_29]), c_0_44]), c_0_44]), c_0_44]), c_0_43]), c_0_45]), c_0_46]), c_0_47]), c_0_30]), c_0_47]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_30])).
cnf(c_0_53,plain,(times_times(int,pls,X1)=pls), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54,negated_conjecture,(~ord_less(int,plus_plus(int,power_power(int,s,number_number_of(nat,plus_plus(int,plus_plus(int,plus_plus(int,one_one(int),pls),pls),plus_plus(int,plus_plus(int,one_one(int),pls),pls)))),one_one(int)),pls)), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55,plain,(ord_less(int,plus_plus(int,one_one(int),power_power(int,s,number_number_of(nat,plus_plus(int,one_one(int),one_one(int))))),pls)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_44]), c_0_29]), c_0_44]), c_0_44]), c_0_44]), c_0_43]), c_0_45]), c_0_46]), c_0_47]), c_0_30]), c_0_47]), c_0_52]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_44]), c_0_29]), c_0_44]), c_0_44]), c_0_44]), c_0_43]), c_0_45]), c_0_46]), c_0_47]), c_0_30]), c_0_50]), c_0_47]), c_0_53])).
cnf(c_0_56,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_30]), c_0_55])]), ['proof']).
# SZS output end CNFRefutation.
EOF
