% Problem    : SWV469+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n116.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:16:36 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X12]:![X13]:![X5]:![X6]:(((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~(leq(host(X7),host(X14))))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~(setIn(X7,alive))|~(setIn(X14,alive)))))&![X7]:![X14]:((setIn(X14,alive)&index(status,host(X14))=elec_2)=>~(elem(m_Ack(X7,X14),queue(host(X7))))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X14]:((((setIn(X7,alive)&setIn(X14,alive))&setIn(host(X14),index(down,host(X7))))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X14]:(((setIn(X14,alive)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&setIn(host(X14),index(down,host(X7)))))))&![X7]:![X15]:![X14]:(((((~(leq(host(X14),host(X7)))&setIn(X7,alive))&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>~(elem(m_Down(X15),queue(host(X7))))))&![X7]:![X14]:(((((~(leq(host(X7),host(X14)))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:(((setIn(X14,alive)&elem(m_Ack(X14,X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~(elem(m_Ack(X7,X15),queue(host(X7))))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&elem(m_Down(X15),queue(host(X7))))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&host(X15)=host(X14))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&elem(m_Down(X15),queue(host(X7)))))))&![X7]:![X14]:(((~(leq(index(pendack,host(X14)),host(X7)))&setIn(X14,alive))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X14]:(((((~(leq(host(X7),host(X14)))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~(leq(index(pendack,host(X7)),index(pendack,host(X14))))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:((((![X16]:((~(leq(host(X14),X16))&leq(s(zero),X16))=>(setIn(X16,index(down,host(X14)))|X16=host(X15)))&~(leq(host(X14),host(X7))))&elem(m_Down(X15),queue(host(X14))))&index(status,host(X14))=elec_1)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12))=>(setIn(X5,alive)=>(~(leq(host(X5),host(X6)))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>((![X7]:((~(leq(host(X5),X7))&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1)=>(~(leq(nbr_proc,host(X5)))=>![X7]:(host(X5)!=host(X7)=>![X17]:(host(X5)=host(X17)=>((~(leq(s(host(X5)),host(X7)))&setIn(X17,alive))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm)))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj)).
fof(c_0_1, axiom, (![X5]:![X6]:![X4]:(elem(X5,cons(X6,X4))<=>(X5=X6|elem(X5,X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_46)).
fof(c_0_2, axiom, (![X5]:![X6]:(leq(X5,X6)|leq(X6,X5))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_60)).
fof(c_0_3, axiom, (![X5]:![X6]:(leq(X5,s(X6))<=>(X5=s(X6)|leq(X5,X6)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_64)).
fof(c_0_4, axiom, (![X3]:leq(s(zero),host(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_02)).
fof(c_0_5, axiom, (![X5]:leq(X5,X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_59)).
fof(c_0_6, axiom, (![X5]:![X6]:((leq(X5,X6)&leq(X6,X5))<=>X5=X6)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_61)).
fof(c_0_7, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)<=>((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:![X14]:((setIn(X14,alive)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X14),queue(host(X7)))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X14]:((((setIn(X7,alive)&setIn(X14,alive))&setIn(host(X14),index(down,host(X7))))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X14]:(((setIn(X14,alive)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&setIn(host(X14),index(down,host(X7)))))))&![X7]:![X15]:![X14]:(((((~leq(host(X14),host(X7))&setIn(X7,alive))&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>~elem(m_Down(X15),queue(host(X7)))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:(((setIn(X14,alive)&elem(m_Ack(X14,X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X15),queue(host(X7)))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&elem(m_Down(X15),queue(host(X7))))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&host(X15)=host(X14))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&elem(m_Down(X15),queue(host(X7)))))))&![X7]:![X14]:(((~leq(index(pendack,host(X14)),host(X7))&setIn(X14,alive))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~leq(index(pendack,host(X7)),index(pendack,host(X14)))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:((((![X16]:((~leq(host(X14),X16)&leq(s(zero),X16))=>(setIn(X16,index(down,host(X14)))|X16=host(X15)))&~leq(host(X14),host(X7)))&elem(m_Down(X15),queue(host(X14))))&index(status,host(X14))=elec_1)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12)))), introduced(definition)).
fof(c_0_8, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)=>((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:![X14]:((setIn(X14,alive)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X14),queue(host(X7)))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X14]:((((setIn(X7,alive)&setIn(X14,alive))&setIn(host(X14),index(down,host(X7))))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X14]:(((setIn(X14,alive)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&setIn(host(X14),index(down,host(X7)))))))&![X7]:![X15]:![X14]:(((((~leq(host(X14),host(X7))&setIn(X7,alive))&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>~elem(m_Down(X15),queue(host(X7)))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:(((setIn(X14,alive)&elem(m_Ack(X14,X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&host(X14)=host(X15))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~elem(m_Ack(X7,X15),queue(host(X7)))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&elem(m_Down(X15),queue(host(X7))))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&host(X15)=host(X14))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&elem(m_Down(X15),queue(host(X7)))))))&![X7]:![X14]:(((~leq(index(pendack,host(X14)),host(X7))&setIn(X14,alive))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X14]:(((((~leq(host(X7),host(X14))&setIn(X7,alive))&setIn(X14,alive))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_2)=>~leq(index(pendack,host(X7)),index(pendack,host(X14)))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=elec_2)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X15]:![X14]:((((![X16]:((~leq(host(X14),X16)&leq(s(zero),X16))=>(setIn(X16,index(down,host(X14)))|X16=host(X15)))&~leq(host(X14),host(X7)))&elem(m_Down(X15),queue(host(X14))))&index(status,host(X14))=elec_1)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12)))), inference(split_equiv,[status(thm)],[c_0_7])).
fof(c_0_9, negated_conjecture, (~(![X12]:![X5]:![X6]:(epred1_3(X5,X6,X12)=>(setIn(X5,alive)=>(~leq(host(X5),host(X6))=>(~(((index(ldr,host(X5))=host(X6)&index(status,host(X5))=norm)|(index(status,host(X5))=wait&host(X6)=host(index(elid,host(X5))))))=>((![X7]:((~leq(host(X5),X7)&leq(s(zero),X7))=>(setIn(X7,index(down,host(X5)))|X7=host(X6)))&index(status,host(X5))=elec_1)=>(~leq(nbr_proc,host(X5))=>![X7]:(host(X5)!=host(X7)=>![X17]:(host(X5)=host(X17)=>((~leq(s(host(X5)),host(X7))&setIn(X17,alive))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_7])).
fof(c_0_10, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((~elem(X7,cons(X8,X9))|(X7=X8|elem(X7,X9)))&((X10!=X11|elem(X10,cons(X11,X12)))&(~elem(X10,X12)|elem(X10,cons(X11,X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
fof(c_0_11, plain, (![X17]:![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:![X34]:![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:![X43]:![X44]:![X45]:![X46]:![X47]:![X48]:![X49]:![X50]:![X51]:![X52]:![X53]:![X54]:![X55]:![X56]:![X57]:((((((((((((((((((~elem(m_Ldr(X21),queue(host(X20)))|~leq(host(X20),host(X21)))|~epred1_3(X19,X18,X17))&(((X23=X22|host(X23)!=host(X22))|(~setIn(X22,alive)|~setIn(X23,alive)))|~epred1_3(X19,X18,X17)))&(((~setIn(X25,alive)|index(status,host(X25))!=elec_2)|~elem(m_Ack(X24,X25),queue(host(X24))))|~epred1_3(X19,X18,X17)))&((((index(status,host(X26))!=elec_1|~setIn(X26,alive))|index(elid,host(X26))=X26)|~epred1_3(X19,X18,X17))&(((index(status,host(X26))!=elec_2|~setIn(X26,alive))|index(elid,host(X26))=X26)|~epred1_3(X19,X18,X17))))&(((((~setIn(X27,alive)|~setIn(X28,alive))|~setIn(host(X28),index(down,host(X27))))|index(status,host(X28))!=elec_2)|leq(index(pendack,host(X28)),host(X27)))|~epred1_3(X19,X18,X17)))&((((~setIn(X30,alive)|index(status,host(X30))!=norm)|index(ldr,host(X30))!=host(X30))|(~setIn(X29,alive)|~setIn(host(X30),index(down,host(X29)))))|~epred1_3(X19,X18,X17)))&((((((leq(host(X33),host(X31))|~setIn(X31,alive))|~setIn(X33,alive))|host(X33)!=host(X32))|index(status,host(X33))!=elec_2)|~elem(m_Down(X32),queue(host(X31))))|~epred1_3(X19,X18,X17)))&((((((leq(host(X34),host(X35))|~setIn(X34,alive))|~setIn(X35,alive))|index(status,host(X34))!=elec_2)|index(status,host(X35))!=elec_2)|leq(index(pendack,host(X35)),host(X34)))|~epred1_3(X19,X18,X17)))&((((~setIn(X38,alive)|~elem(m_Ack(X38,X37),queue(host(X38))))|host(X37)!=host(X36))|((~setIn(X36,alive)|index(ldr,host(X36))!=host(X36))|index(status,host(X36))!=norm))|~epred1_3(X19,X18,X17)))&((((((~setIn(X39,alive)|~setIn(X41,alive))|host(X41)!=host(X40))|index(status,host(X39))!=elec_2)|index(status,host(X41))!=elec_2)|~elem(m_Ack(X39,X40),queue(host(X39))))|~epred1_3(X19,X18,X17)))&((((((~setIn(X42,alive)|~setIn(X44,alive))|~elem(m_Down(X43),queue(host(X42))))|host(X44)!=host(X43))|index(status,host(X44))!=elec_2)|leq(index(pendack,host(X44)),host(X42)))|~epred1_3(X19,X18,X17)))&(((((~setIn(X47,alive)|host(X46)!=host(X47))|index(status,host(X47))!=norm)|index(ldr,host(X47))!=host(X47))|(~setIn(X45,alive)|~elem(m_Down(X46),queue(host(X45)))))|~epred1_3(X19,X18,X17)))&((((leq(index(pendack,host(X49)),host(X48))|~setIn(X49,alive))|index(status,host(X49))!=elec_2)|((~setIn(X48,alive)|index(ldr,host(X48))!=host(X48))|index(status,host(X48))!=norm))|~epred1_3(X19,X18,X17)))&((((((leq(host(X50),host(X51))|~setIn(X50,alive))|~setIn(X51,alive))|index(status,host(X50))!=elec_2)|index(status,host(X51))!=elec_2)|~leq(index(pendack,host(X50)),index(pendack,host(X51))))|~epred1_3(X19,X18,X17)))&(((((~setIn(X54,alive)|~elem(m_Down(X53),queue(host(X54))))|host(X53)!=host(X52))|index(status,host(X54))!=elec_2)|((~setIn(X52,alive)|index(ldr,host(X52))!=host(X52))|index(status,host(X52))!=norm))|~epred1_3(X19,X18,X17)))&(((((((~leq(host(X57),esk10_6(X17,X18,X19,X55,X56,X57))|leq(host(X57),host(X55)))|~elem(m_Down(X56),queue(host(X57))))|index(status,host(X57))!=elec_1)|((~setIn(X55,alive)|index(ldr,host(X55))!=host(X55))|index(status,host(X55))!=norm))|~epred1_3(X19,X18,X17))&(((((leq(s(zero),esk10_6(X17,X18,X19,X55,X56,X57))|leq(host(X57),host(X55)))|~elem(m_Down(X56),queue(host(X57))))|index(status,host(X57))!=elec_1)|((~setIn(X55,alive)|index(ldr,host(X55))!=host(X55))|index(status,host(X55))!=norm))|~epred1_3(X19,X18,X17)))&((((((~setIn(esk10_6(X17,X18,X19,X55,X56,X57),index(down,host(X57)))|leq(host(X57),host(X55)))|~elem(m_Down(X56),queue(host(X57))))|index(status,host(X57))!=elec_1)|((~setIn(X55,alive)|index(ldr,host(X55))!=host(X55))|index(status,host(X55))!=norm))|~epred1_3(X19,X18,X17))&(((((esk10_6(X17,X18,X19,X55,X56,X57)!=host(X56)|leq(host(X57),host(X55)))|~elem(m_Down(X56),queue(host(X57))))|index(status,host(X57))!=elec_1)|((~setIn(X55,alive)|index(ldr,host(X55))!=host(X55))|index(status,host(X55))!=norm))|~epred1_3(X19,X18,X17)))))&(queue(host(X19))=cons(m_Down(X18),X17)|~epred1_3(X19,X18,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])])).
fof(c_0_12, negated_conjecture, (![X21]:(epred1_3(esk6_0,esk7_0,esk5_0)&(setIn(esk6_0,alive)&(~leq(host(esk6_0),host(esk7_0))&(((index(ldr,host(esk6_0))!=host(esk7_0)|index(status,host(esk6_0))!=norm)&(index(status,host(esk6_0))!=wait|host(esk7_0)!=host(index(elid,host(esk6_0)))))&((((leq(host(esk6_0),X21)|~leq(s(zero),X21))|(setIn(X21,index(down,host(esk6_0)))|X21=host(esk7_0)))&index(status,host(esk6_0))=elec_1)&(~leq(nbr_proc,host(esk6_0))&(host(esk6_0)!=host(esk8_0)&(host(esk6_0)=host(esk9_0)&((~leq(s(host(esk6_0)),host(esk8_0))&setIn(esk9_0,alive))&((setIn(esk8_0,alive)&index(ldr,host(esk8_0))=host(esk8_0))&index(status,host(esk8_0))=norm))))))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
fof(c_0_13, plain, (![X7]:![X8]:(leq(X7,X8)|leq(X8,X7))), inference(variable_rename,[status(thm)],[c_0_2])).
cnf(c_0_14,plain,(elem(X1,cons(X2,X3))|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(queue(host(X1))=cons(m_Down(X2),X3)|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16,negated_conjecture,(epred1_3(esk6_0,esk7_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, (![X7]:![X8]:![X9]:![X10]:((~leq(X7,s(X8))|(X7=s(X8)|leq(X7,X8)))&((X9!=s(X10)|leq(X9,s(X10)))&(~leq(X9,X10)|leq(X9,s(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_18,negated_conjecture,(~leq(s(host(esk6_0)),host(esk8_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19,plain,(leq(X1,X2)|leq(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_20, plain, (![X4]:leq(s(zero),host(X4))), inference(variable_rename,[status(thm)],[c_0_4])).
cnf(c_0_21,plain,(elem(X1,cons(X1,X2))), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_22,negated_conjecture,(cons(m_Down(esk7_0),esk5_0)=queue(host(esk6_0))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23,plain,(leq(X1,X2)|X1=s(X2)|~leq(X1,s(X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24,negated_conjecture,(leq(host(esk8_0),s(host(esk6_0)))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_25, plain, (![X6]:leq(X6,X6)), inference(variable_rename,[status(thm)],[c_0_5])).
cnf(c_0_26,negated_conjecture,(X1=host(esk7_0)|setIn(X1,index(down,host(esk6_0)))|leq(host(esk6_0),X1)|~leq(s(zero),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27,plain,(leq(s(zero),host(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28,plain,(~epred1_3(X1,X2,X3)|~elem(m_Down(X4),queue(host(X5)))|~setIn(X5,alive)|index(ldr,host(X6))!=host(X6)|index(status,host(X6))!=norm|host(X4)!=host(X6)|~setIn(X6,alive)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_29,negated_conjecture,(elem(m_Down(esk7_0),queue(host(esk6_0)))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_30,negated_conjecture,(setIn(esk6_0,alive)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_31, plain, (![X7]:![X8]:![X9]:![X10]:(((~leq(X7,X8)|~leq(X8,X7))|X7=X8)&((leq(X9,X10)|X9!=X10)&(leq(X10,X9)|X9!=X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
cnf(c_0_32,negated_conjecture,(s(host(esk6_0))=host(esk8_0)|leq(host(esk8_0),host(esk6_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_33,plain,(leq(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34,plain,(~epred1_3(X1,X2,X3)|~setIn(host(X4),index(down,host(X5)))|~setIn(X5,alive)|index(ldr,host(X4))!=host(X4)|index(status,host(X4))!=norm|~setIn(X4,alive)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_35,negated_conjecture,(host(esk7_0)=host(X1)|leq(host(esk6_0),host(X1))|setIn(host(X1),index(down,host(esk6_0)))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36,plain,(index(ldr,host(X1))!=host(X1)|index(status,host(X1))!=norm|host(esk7_0)!=host(X1)|~epred1_3(X2,X3,X4)|~setIn(X1,alive)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_37,plain,(X1=X2|~leq(X2,X1)|~leq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38,negated_conjecture,(leq(host(esk8_0),host(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_32]), c_0_33])])).
cnf(c_0_39,negated_conjecture,(host(esk6_0)!=host(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_40,plain,(host(esk7_0)=host(X1)|leq(host(esk6_0),host(X1))|index(ldr,host(X1))!=host(X1)|index(status,host(X1))!=norm|~epred1_3(X2,X3,X4)|~setIn(X1,alive)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_30])])).
cnf(c_0_41,negated_conjecture,(index(ldr,host(X1))!=host(X1)|index(status,host(X1))!=norm|host(esk7_0)!=host(X1)|~setIn(X1,alive)), inference(spm,[status(thm)],[c_0_36, c_0_16])).
cnf(c_0_42,negated_conjecture,(~leq(host(esk6_0),host(esk8_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_43,negated_conjecture,(leq(host(esk6_0),host(X1))|index(ldr,host(X1))!=host(X1)|index(status,host(X1))!=norm|~setIn(X1,alive)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_16]), c_0_41])).
cnf(c_0_44,negated_conjecture,(index(ldr,host(esk8_0))=host(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_45,negated_conjecture,(index(status,host(esk8_0))=norm), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_46,negated_conjecture,(setIn(esk8_0,alive)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_47,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_45]), c_0_46])]), ['proof']).
# SZS output end CNFRefutation.
EOF
