% Problem    : ALG167+1 : TPTP v6.1.0. Released v2.7.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n148.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:17:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4c
# and selection function SelectCQPrecWNTNp.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, ((((e20=op2(op2(e23,op2(e23,e23)),op2(e23,op2(e23,e23)))&e21=op2(e23,op2(e23,op2(e23,e23))))&e22=op2(e23,e23))&e24=op2(e23,op2(e23,e23)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax15)).
fof(c_0_1, axiom, (((((((((((((((((((((((((op2(op2(e20,e20),op2(e20,e20))=e20&op2(op2(e21,e20),op2(e20,e21))=e20)&op2(op2(e22,e20),op2(e20,e22))=e20)&op2(op2(e23,e20),op2(e20,e23))=e20)&op2(op2(e24,e20),op2(e20,e24))=e20)&op2(op2(e20,e21),op2(e21,e20))=e21)&op2(op2(e21,e21),op2(e21,e21))=e21)&op2(op2(e22,e21),op2(e21,e22))=e21)&op2(op2(e23,e21),op2(e21,e23))=e21)&op2(op2(e24,e21),op2(e21,e24))=e21)&op2(op2(e20,e22),op2(e22,e20))=e22)&op2(op2(e21,e22),op2(e22,e21))=e22)&op2(op2(e22,e22),op2(e22,e22))=e22)&op2(op2(e23,e22),op2(e22,e23))=e22)&op2(op2(e24,e22),op2(e22,e24))=e22)&op2(op2(e20,e23),op2(e23,e20))=e23)&op2(op2(e21,e23),op2(e23,e21))=e23)&op2(op2(e22,e23),op2(e23,e22))=e23)&op2(op2(e23,e23),op2(e23,e23))=e23)&op2(op2(e24,e23),op2(e23,e24))=e23)&op2(op2(e20,e24),op2(e24,e20))=e24)&op2(op2(e21,e24),op2(e24,e21))=e24)&op2(op2(e22,e24),op2(e24,e22))=e24)&op2(op2(e23,e24),op2(e24,e23))=e24)&op2(op2(e24,e24),op2(e24,e24))=e24)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax6)).
fof(c_0_2, axiom, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op2(e20,e20)!=op2(e21,e20)&op2(e20,e20)!=op2(e22,e20))&op2(e20,e20)!=op2(e23,e20))&op2(e20,e20)!=op2(e24,e20))&op2(e21,e20)!=op2(e22,e20))&op2(e21,e20)!=op2(e23,e20))&op2(e21,e20)!=op2(e24,e20))&op2(e22,e20)!=op2(e23,e20))&op2(e22,e20)!=op2(e24,e20))&op2(e23,e20)!=op2(e24,e20))&op2(e20,e21)!=op2(e21,e21))&op2(e20,e21)!=op2(e22,e21))&op2(e20,e21)!=op2(e23,e21))&op2(e20,e21)!=op2(e24,e21))&op2(e21,e21)!=op2(e22,e21))&op2(e21,e21)!=op2(e23,e21))&op2(e21,e21)!=op2(e24,e21))&op2(e22,e21)!=op2(e23,e21))&op2(e22,e21)!=op2(e24,e21))&op2(e23,e21)!=op2(e24,e21))&op2(e20,e22)!=op2(e21,e22))&op2(e20,e22)!=op2(e22,e22))&op2(e20,e22)!=op2(e23,e22))&op2(e20,e22)!=op2(e24,e22))&op2(e21,e22)!=op2(e22,e22))&op2(e21,e22)!=op2(e23,e22))&op2(e21,e22)!=op2(e24,e22))&op2(e22,e22)!=op2(e23,e22))&op2(e22,e22)!=op2(e24,e22))&op2(e23,e22)!=op2(e24,e22))&op2(e20,e23)!=op2(e21,e23))&op2(e20,e23)!=op2(e22,e23))&op2(e20,e23)!=op2(e23,e23))&op2(e20,e23)!=op2(e24,e23))&op2(e21,e23)!=op2(e22,e23))&op2(e21,e23)!=op2(e23,e23))&op2(e21,e23)!=op2(e24,e23))&op2(e22,e23)!=op2(e23,e23))&op2(e22,e23)!=op2(e24,e23))&op2(e23,e23)!=op2(e24,e23))&op2(e20,e24)!=op2(e21,e24))&op2(e20,e24)!=op2(e22,e24))&op2(e20,e24)!=op2(e23,e24))&op2(e20,e24)!=op2(e24,e24))&op2(e21,e24)!=op2(e22,e24))&op2(e21,e24)!=op2(e23,e24))&op2(e21,e24)!=op2(e24,e24))&op2(e22,e24)!=op2(e23,e24))&op2(e22,e24)!=op2(e24,e24))&op2(e23,e24)!=op2(e24,e24))&op2(e20,e20)!=op2(e20,e21))&op2(e20,e20)!=op2(e20,e22))&op2(e20,e20)!=op2(e20,e23))&op2(e20,e20)!=op2(e20,e24))&op2(e20,e21)!=op2(e20,e22))&op2(e20,e21)!=op2(e20,e23))&op2(e20,e21)!=op2(e20,e24))&op2(e20,e22)!=op2(e20,e23))&op2(e20,e22)!=op2(e20,e24))&op2(e20,e23)!=op2(e20,e24))&op2(e21,e20)!=op2(e21,e21))&op2(e21,e20)!=op2(e21,e22))&op2(e21,e20)!=op2(e21,e23))&op2(e21,e20)!=op2(e21,e24))&op2(e21,e21)!=op2(e21,e22))&op2(e21,e21)!=op2(e21,e23))&op2(e21,e21)!=op2(e21,e24))&op2(e21,e22)!=op2(e21,e23))&op2(e21,e22)!=op2(e21,e24))&op2(e21,e23)!=op2(e21,e24))&op2(e22,e20)!=op2(e22,e21))&op2(e22,e20)!=op2(e22,e22))&op2(e22,e20)!=op2(e22,e23))&op2(e22,e20)!=op2(e22,e24))&op2(e22,e21)!=op2(e22,e22))&op2(e22,e21)!=op2(e22,e23))&op2(e22,e21)!=op2(e22,e24))&op2(e22,e22)!=op2(e22,e23))&op2(e22,e22)!=op2(e22,e24))&op2(e22,e23)!=op2(e22,e24))&op2(e23,e20)!=op2(e23,e21))&op2(e23,e20)!=op2(e23,e22))&op2(e23,e20)!=op2(e23,e23))&op2(e23,e20)!=op2(e23,e24))&op2(e23,e21)!=op2(e23,e22))&op2(e23,e21)!=op2(e23,e23))&op2(e23,e21)!=op2(e23,e24))&op2(e23,e22)!=op2(e23,e23))&op2(e23,e22)!=op2(e23,e24))&op2(e23,e23)!=op2(e23,e24))&op2(e24,e20)!=op2(e24,e21))&op2(e24,e20)!=op2(e24,e22))&op2(e24,e20)!=op2(e24,e23))&op2(e24,e20)!=op2(e24,e24))&op2(e24,e21)!=op2(e24,e22))&op2(e24,e21)!=op2(e24,e23))&op2(e24,e21)!=op2(e24,e24))&op2(e24,e22)!=op2(e24,e23))&op2(e24,e22)!=op2(e24,e24))&op2(e24,e23)!=op2(e24,e24))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax8)).
fof(c_0_3, axiom, ((((e10=op1(op1(e13,op1(e13,e13)),op1(e13,op1(e13,e13)))&e11=op1(e13,op1(e13,op1(e13,e13))))&e12=op1(e13,e13))&e14=op1(e13,op1(e13,e13)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax14)).
fof(c_0_4, axiom, ((((((((((((((((((((((((((((((((((((((((((((((((((((((op2(e20,e20)=e20|op2(e20,e21)=e20)|op2(e20,e22)=e20)|op2(e20,e23)=e20)|op2(e20,e24)=e20)&((((op2(e20,e20)=e20|op2(e21,e20)=e20)|op2(e22,e20)=e20)|op2(e23,e20)=e20)|op2(e24,e20)=e20))&((((op2(e20,e20)=e21|op2(e20,e21)=e21)|op2(e20,e22)=e21)|op2(e20,e23)=e21)|op2(e20,e24)=e21))&((((op2(e20,e20)=e21|op2(e21,e20)=e21)|op2(e22,e20)=e21)|op2(e23,e20)=e21)|op2(e24,e20)=e21))&((((op2(e20,e20)=e22|op2(e20,e21)=e22)|op2(e20,e22)=e22)|op2(e20,e23)=e22)|op2(e20,e24)=e22))&((((op2(e20,e20)=e22|op2(e21,e20)=e22)|op2(e22,e20)=e22)|op2(e23,e20)=e22)|op2(e24,e20)=e22))&((((op2(e20,e20)=e23|op2(e20,e21)=e23)|op2(e20,e22)=e23)|op2(e20,e23)=e23)|op2(e20,e24)=e23))&((((op2(e20,e20)=e23|op2(e21,e20)=e23)|op2(e22,e20)=e23)|op2(e23,e20)=e23)|op2(e24,e20)=e23))&((((op2(e20,e20)=e24|op2(e20,e21)=e24)|op2(e20,e22)=e24)|op2(e20,e23)=e24)|op2(e20,e24)=e24))&((((op2(e20,e20)=e24|op2(e21,e20)=e24)|op2(e22,e20)=e24)|op2(e23,e20)=e24)|op2(e24,e20)=e24))&((((op2(e21,e20)=e20|op2(e21,e21)=e20)|op2(e21,e22)=e20)|op2(e21,e23)=e20)|op2(e21,e24)=e20))&((((op2(e20,e21)=e20|op2(e21,e21)=e20)|op2(e22,e21)=e20)|op2(e23,e21)=e20)|op2(e24,e21)=e20))&((((op2(e21,e20)=e21|op2(e21,e21)=e21)|op2(e21,e22)=e21)|op2(e21,e23)=e21)|op2(e21,e24)=e21))&((((op2(e20,e21)=e21|op2(e21,e21)=e21)|op2(e22,e21)=e21)|op2(e23,e21)=e21)|op2(e24,e21)=e21))&((((op2(e21,e20)=e22|op2(e21,e21)=e22)|op2(e21,e22)=e22)|op2(e21,e23)=e22)|op2(e21,e24)=e22))&((((op2(e20,e21)=e22|op2(e21,e21)=e22)|op2(e22,e21)=e22)|op2(e23,e21)=e22)|op2(e24,e21)=e22))&((((op2(e21,e20)=e23|op2(e21,e21)=e23)|op2(e21,e22)=e23)|op2(e21,e23)=e23)|op2(e21,e24)=e23))&((((op2(e20,e21)=e23|op2(e21,e21)=e23)|op2(e22,e21)=e23)|op2(e23,e21)=e23)|op2(e24,e21)=e23))&((((op2(e21,e20)=e24|op2(e21,e21)=e24)|op2(e21,e22)=e24)|op2(e21,e23)=e24)|op2(e21,e24)=e24))&((((op2(e20,e21)=e24|op2(e21,e21)=e24)|op2(e22,e21)=e24)|op2(e23,e21)=e24)|op2(e24,e21)=e24))&((((op2(e22,e20)=e20|op2(e22,e21)=e20)|op2(e22,e22)=e20)|op2(e22,e23)=e20)|op2(e22,e24)=e20))&((((op2(e20,e22)=e20|op2(e21,e22)=e20)|op2(e22,e22)=e20)|op2(e23,e22)=e20)|op2(e24,e22)=e20))&((((op2(e22,e20)=e21|op2(e22,e21)=e21)|op2(e22,e22)=e21)|op2(e22,e23)=e21)|op2(e22,e24)=e21))&((((op2(e20,e22)=e21|op2(e21,e22)=e21)|op2(e22,e22)=e21)|op2(e23,e22)=e21)|op2(e24,e22)=e21))&((((op2(e22,e20)=e22|op2(e22,e21)=e22)|op2(e22,e22)=e22)|op2(e22,e23)=e22)|op2(e22,e24)=e22))&((((op2(e20,e22)=e22|op2(e21,e22)=e22)|op2(e22,e22)=e22)|op2(e23,e22)=e22)|op2(e24,e22)=e22))&((((op2(e22,e20)=e23|op2(e22,e21)=e23)|op2(e22,e22)=e23)|op2(e22,e23)=e23)|op2(e22,e24)=e23))&((((op2(e20,e22)=e23|op2(e21,e22)=e23)|op2(e22,e22)=e23)|op2(e23,e22)=e23)|op2(e24,e22)=e23))&((((op2(e22,e20)=e24|op2(e22,e21)=e24)|op2(e22,e22)=e24)|op2(e22,e23)=e24)|op2(e22,e24)=e24))&((((op2(e20,e22)=e24|op2(e21,e22)=e24)|op2(e22,e22)=e24)|op2(e23,e22)=e24)|op2(e24,e22)=e24))&((((op2(e23,e20)=e20|op2(e23,e21)=e20)|op2(e23,e22)=e20)|op2(e23,e23)=e20)|op2(e23,e24)=e20))&((((op2(e20,e23)=e20|op2(e21,e23)=e20)|op2(e22,e23)=e20)|op2(e23,e23)=e20)|op2(e24,e23)=e20))&((((op2(e23,e20)=e21|op2(e23,e21)=e21)|op2(e23,e22)=e21)|op2(e23,e23)=e21)|op2(e23,e24)=e21))&((((op2(e20,e23)=e21|op2(e21,e23)=e21)|op2(e22,e23)=e21)|op2(e23,e23)=e21)|op2(e24,e23)=e21))&((((op2(e23,e20)=e22|op2(e23,e21)=e22)|op2(e23,e22)=e22)|op2(e23,e23)=e22)|op2(e23,e24)=e22))&((((op2(e20,e23)=e22|op2(e21,e23)=e22)|op2(e22,e23)=e22)|op2(e23,e23)=e22)|op2(e24,e23)=e22))&((((op2(e23,e20)=e23|op2(e23,e21)=e23)|op2(e23,e22)=e23)|op2(e23,e23)=e23)|op2(e23,e24)=e23))&((((op2(e20,e23)=e23|op2(e21,e23)=e23)|op2(e22,e23)=e23)|op2(e23,e23)=e23)|op2(e24,e23)=e23))&((((op2(e23,e20)=e24|op2(e23,e21)=e24)|op2(e23,e22)=e24)|op2(e23,e23)=e24)|op2(e23,e24)=e24))&((((op2(e20,e23)=e24|op2(e21,e23)=e24)|op2(e22,e23)=e24)|op2(e23,e23)=e24)|op2(e24,e23)=e24))&((((op2(e24,e20)=e20|op2(e24,e21)=e20)|op2(e24,e22)=e20)|op2(e24,e23)=e20)|op2(e24,e24)=e20))&((((op2(e20,e24)=e20|op2(e21,e24)=e20)|op2(e22,e24)=e20)|op2(e23,e24)=e20)|op2(e24,e24)=e20))&((((op2(e24,e20)=e21|op2(e24,e21)=e21)|op2(e24,e22)=e21)|op2(e24,e23)=e21)|op2(e24,e24)=e21))&((((op2(e20,e24)=e21|op2(e21,e24)=e21)|op2(e22,e24)=e21)|op2(e23,e24)=e21)|op2(e24,e24)=e21))&((((op2(e24,e20)=e22|op2(e24,e21)=e22)|op2(e24,e22)=e22)|op2(e24,e23)=e22)|op2(e24,e24)=e22))&((((op2(e20,e24)=e22|op2(e21,e24)=e22)|op2(e22,e24)=e22)|op2(e23,e24)=e22)|op2(e24,e24)=e22))&((((op2(e24,e20)=e23|op2(e24,e21)=e23)|op2(e24,e22)=e23)|op2(e24,e23)=e23)|op2(e24,e24)=e23))&((((op2(e20,e24)=e23|op2(e21,e24)=e23)|op2(e22,e24)=e23)|op2(e23,e24)=e23)|op2(e24,e24)=e23))&((((op2(e24,e20)=e24|op2(e24,e21)=e24)|op2(e24,e22)=e24)|op2(e24,e23)=e24)|op2(e24,e24)=e24))&((((op2(e20,e24)=e24|op2(e21,e24)=e24)|op2(e22,e24)=e24)|op2(e23,e24)=e24)|op2(e24,e24)=e24))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax5)).
fof(c_0_5, axiom, ((((((((((e20!=e21&e20!=e22)&e20!=e23)&e20!=e24)&e21!=e22)&e21!=e23)&e21!=e24)&e22!=e23)&e22!=e24)&e23!=e24)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax10)).
fof(c_0_6, axiom, (((((((((op2(e20,e20)!=e20&op2(e20,e21)!=e21)&op2(e20,e22)!=e22)&op2(e20,e23)!=e23)&op2(e20,e24)!=e24)|((((op2(e21,e20)!=e20&op2(e21,e21)!=e21)&op2(e21,e22)!=e22)&op2(e21,e23)!=e23)&op2(e21,e24)!=e24))|((((op2(e22,e20)!=e20&op2(e22,e21)!=e21)&op2(e22,e22)!=e22)&op2(e22,e23)!=e23)&op2(e22,e24)!=e24))|((((op2(e23,e20)!=e20&op2(e23,e21)!=e21)&op2(e23,e22)!=e22)&op2(e23,e23)!=e23)&op2(e23,e24)!=e24))|((((op2(e24,e20)!=e20&op2(e24,e21)!=e21)&op2(e24,e22)!=e22)&op2(e24,e23)!=e23)&op2(e24,e24)!=e24))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax13)).
fof(c_0_7, axiom, (((((((((((((((((((((((((((((op2(e20,e20)=e20|op2(e20,e20)=e21)|op2(e20,e20)=e22)|op2(e20,e20)=e23)|op2(e20,e20)=e24)&((((op2(e20,e21)=e20|op2(e20,e21)=e21)|op2(e20,e21)=e22)|op2(e20,e21)=e23)|op2(e20,e21)=e24))&((((op2(e20,e22)=e20|op2(e20,e22)=e21)|op2(e20,e22)=e22)|op2(e20,e22)=e23)|op2(e20,e22)=e24))&((((op2(e20,e23)=e20|op2(e20,e23)=e21)|op2(e20,e23)=e22)|op2(e20,e23)=e23)|op2(e20,e23)=e24))&((((op2(e20,e24)=e20|op2(e20,e24)=e21)|op2(e20,e24)=e22)|op2(e20,e24)=e23)|op2(e20,e24)=e24))&((((op2(e21,e20)=e20|op2(e21,e20)=e21)|op2(e21,e20)=e22)|op2(e21,e20)=e23)|op2(e21,e20)=e24))&((((op2(e21,e21)=e20|op2(e21,e21)=e21)|op2(e21,e21)=e22)|op2(e21,e21)=e23)|op2(e21,e21)=e24))&((((op2(e21,e22)=e20|op2(e21,e22)=e21)|op2(e21,e22)=e22)|op2(e21,e22)=e23)|op2(e21,e22)=e24))&((((op2(e21,e23)=e20|op2(e21,e23)=e21)|op2(e21,e23)=e22)|op2(e21,e23)=e23)|op2(e21,e23)=e24))&((((op2(e21,e24)=e20|op2(e21,e24)=e21)|op2(e21,e24)=e22)|op2(e21,e24)=e23)|op2(e21,e24)=e24))&((((op2(e22,e20)=e20|op2(e22,e20)=e21)|op2(e22,e20)=e22)|op2(e22,e20)=e23)|op2(e22,e20)=e24))&((((op2(e22,e21)=e20|op2(e22,e21)=e21)|op2(e22,e21)=e22)|op2(e22,e21)=e23)|op2(e22,e21)=e24))&((((op2(e22,e22)=e20|op2(e22,e22)=e21)|op2(e22,e22)=e22)|op2(e22,e22)=e23)|op2(e22,e22)=e24))&((((op2(e22,e23)=e20|op2(e22,e23)=e21)|op2(e22,e23)=e22)|op2(e22,e23)=e23)|op2(e22,e23)=e24))&((((op2(e22,e24)=e20|op2(e22,e24)=e21)|op2(e22,e24)=e22)|op2(e22,e24)=e23)|op2(e22,e24)=e24))&((((op2(e23,e20)=e20|op2(e23,e20)=e21)|op2(e23,e20)=e22)|op2(e23,e20)=e23)|op2(e23,e20)=e24))&((((op2(e23,e21)=e20|op2(e23,e21)=e21)|op2(e23,e21)=e22)|op2(e23,e21)=e23)|op2(e23,e21)=e24))&((((op2(e23,e22)=e20|op2(e23,e22)=e21)|op2(e23,e22)=e22)|op2(e23,e22)=e23)|op2(e23,e22)=e24))&((((op2(e23,e23)=e20|op2(e23,e23)=e21)|op2(e23,e23)=e22)|op2(e23,e23)=e23)|op2(e23,e23)=e24))&((((op2(e23,e24)=e20|op2(e23,e24)=e21)|op2(e23,e24)=e22)|op2(e23,e24)=e23)|op2(e23,e24)=e24))&((((op2(e24,e20)=e20|op2(e24,e20)=e21)|op2(e24,e20)=e22)|op2(e24,e20)=e23)|op2(e24,e20)=e24))&((((op2(e24,e21)=e20|op2(e24,e21)=e21)|op2(e24,e21)=e22)|op2(e24,e21)=e23)|op2(e24,e21)=e24))&((((op2(e24,e22)=e20|op2(e24,e22)=e21)|op2(e24,e22)=e22)|op2(e24,e22)=e23)|op2(e24,e22)=e24))&((((op2(e24,e23)=e20|op2(e24,e23)=e21)|op2(e24,e23)=e22)|op2(e24,e23)=e23)|op2(e24,e23)=e24))&((((op2(e24,e24)=e20|op2(e24,e24)=e21)|op2(e24,e24)=e22)|op2(e24,e24)=e23)|op2(e24,e24)=e24))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax4)).
fof(c_0_8, axiom, (((((((((((((((((((((((((op1(op1(e10,e10),op1(e10,e10))=e10&op1(op1(e11,e10),op1(e10,e11))=e10)&op1(op1(e12,e10),op1(e10,e12))=e10)&op1(op1(e13,e10),op1(e10,e13))=e10)&op1(op1(e14,e10),op1(e10,e14))=e10)&op1(op1(e10,e11),op1(e11,e10))=e11)&op1(op1(e11,e11),op1(e11,e11))=e11)&op1(op1(e12,e11),op1(e11,e12))=e11)&op1(op1(e13,e11),op1(e11,e13))=e11)&op1(op1(e14,e11),op1(e11,e14))=e11)&op1(op1(e10,e12),op1(e12,e10))=e12)&op1(op1(e11,e12),op1(e12,e11))=e12)&op1(op1(e12,e12),op1(e12,e12))=e12)&op1(op1(e13,e12),op1(e12,e13))=e12)&op1(op1(e14,e12),op1(e12,e14))=e12)&op1(op1(e10,e13),op1(e13,e10))=e13)&op1(op1(e11,e13),op1(e13,e11))=e13)&op1(op1(e12,e13),op1(e13,e12))=e13)&op1(op1(e13,e13),op1(e13,e13))=e13)&op1(op1(e14,e13),op1(e13,e14))=e13)&op1(op1(e10,e14),op1(e14,e10))=e14)&op1(op1(e11,e14),op1(e14,e11))=e14)&op1(op1(e12,e14),op1(e14,e12))=e14)&op1(op1(e13,e14),op1(e14,e13))=e14)&op1(op1(e14,e14),op1(e14,e14))=e14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax3)).
fof(c_0_9, axiom, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op1(e10,e10)!=op1(e11,e10)&op1(e10,e10)!=op1(e12,e10))&op1(e10,e10)!=op1(e13,e10))&op1(e10,e10)!=op1(e14,e10))&op1(e11,e10)!=op1(e12,e10))&op1(e11,e10)!=op1(e13,e10))&op1(e11,e10)!=op1(e14,e10))&op1(e12,e10)!=op1(e13,e10))&op1(e12,e10)!=op1(e14,e10))&op1(e13,e10)!=op1(e14,e10))&op1(e10,e11)!=op1(e11,e11))&op1(e10,e11)!=op1(e12,e11))&op1(e10,e11)!=op1(e13,e11))&op1(e10,e11)!=op1(e14,e11))&op1(e11,e11)!=op1(e12,e11))&op1(e11,e11)!=op1(e13,e11))&op1(e11,e11)!=op1(e14,e11))&op1(e12,e11)!=op1(e13,e11))&op1(e12,e11)!=op1(e14,e11))&op1(e13,e11)!=op1(e14,e11))&op1(e10,e12)!=op1(e11,e12))&op1(e10,e12)!=op1(e12,e12))&op1(e10,e12)!=op1(e13,e12))&op1(e10,e12)!=op1(e14,e12))&op1(e11,e12)!=op1(e12,e12))&op1(e11,e12)!=op1(e13,e12))&op1(e11,e12)!=op1(e14,e12))&op1(e12,e12)!=op1(e13,e12))&op1(e12,e12)!=op1(e14,e12))&op1(e13,e12)!=op1(e14,e12))&op1(e10,e13)!=op1(e11,e13))&op1(e10,e13)!=op1(e12,e13))&op1(e10,e13)!=op1(e13,e13))&op1(e10,e13)!=op1(e14,e13))&op1(e11,e13)!=op1(e12,e13))&op1(e11,e13)!=op1(e13,e13))&op1(e11,e13)!=op1(e14,e13))&op1(e12,e13)!=op1(e13,e13))&op1(e12,e13)!=op1(e14,e13))&op1(e13,e13)!=op1(e14,e13))&op1(e10,e14)!=op1(e11,e14))&op1(e10,e14)!=op1(e12,e14))&op1(e10,e14)!=op1(e13,e14))&op1(e10,e14)!=op1(e14,e14))&op1(e11,e14)!=op1(e12,e14))&op1(e11,e14)!=op1(e13,e14))&op1(e11,e14)!=op1(e14,e14))&op1(e12,e14)!=op1(e13,e14))&op1(e12,e14)!=op1(e14,e14))&op1(e13,e14)!=op1(e14,e14))&op1(e10,e10)!=op1(e10,e11))&op1(e10,e10)!=op1(e10,e12))&op1(e10,e10)!=op1(e10,e13))&op1(e10,e10)!=op1(e10,e14))&op1(e10,e11)!=op1(e10,e12))&op1(e10,e11)!=op1(e10,e13))&op1(e10,e11)!=op1(e10,e14))&op1(e10,e12)!=op1(e10,e13))&op1(e10,e12)!=op1(e10,e14))&op1(e10,e13)!=op1(e10,e14))&op1(e11,e10)!=op1(e11,e11))&op1(e11,e10)!=op1(e11,e12))&op1(e11,e10)!=op1(e11,e13))&op1(e11,e10)!=op1(e11,e14))&op1(e11,e11)!=op1(e11,e12))&op1(e11,e11)!=op1(e11,e13))&op1(e11,e11)!=op1(e11,e14))&op1(e11,e12)!=op1(e11,e13))&op1(e11,e12)!=op1(e11,e14))&op1(e11,e13)!=op1(e11,e14))&op1(e12,e10)!=op1(e12,e11))&op1(e12,e10)!=op1(e12,e12))&op1(e12,e10)!=op1(e12,e13))&op1(e12,e10)!=op1(e12,e14))&op1(e12,e11)!=op1(e12,e12))&op1(e12,e11)!=op1(e12,e13))&op1(e12,e11)!=op1(e12,e14))&op1(e12,e12)!=op1(e12,e13))&op1(e12,e12)!=op1(e12,e14))&op1(e12,e13)!=op1(e12,e14))&op1(e13,e10)!=op1(e13,e11))&op1(e13,e10)!=op1(e13,e12))&op1(e13,e10)!=op1(e13,e13))&op1(e13,e10)!=op1(e13,e14))&op1(e13,e11)!=op1(e13,e12))&op1(e13,e11)!=op1(e13,e13))&op1(e13,e11)!=op1(e13,e14))&op1(e13,e12)!=op1(e13,e13))&op1(e13,e12)!=op1(e13,e14))&op1(e13,e13)!=op1(e13,e14))&op1(e14,e10)!=op1(e14,e11))&op1(e14,e10)!=op1(e14,e12))&op1(e14,e10)!=op1(e14,e13))&op1(e14,e10)!=op1(e14,e14))&op1(e14,e11)!=op1(e14,e12))&op1(e14,e11)!=op1(e14,e13))&op1(e14,e11)!=op1(e14,e14))&op1(e14,e12)!=op1(e14,e13))&op1(e14,e12)!=op1(e14,e14))&op1(e14,e13)!=op1(e14,e14))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax7)).
fof(c_0_10, axiom, ((((((((((((((((((((((((((((((((((((((((((((((((((((((op1(e10,e10)=e10|op1(e10,e11)=e10)|op1(e10,e12)=e10)|op1(e10,e13)=e10)|op1(e10,e14)=e10)&((((op1(e10,e10)=e10|op1(e11,e10)=e10)|op1(e12,e10)=e10)|op1(e13,e10)=e10)|op1(e14,e10)=e10))&((((op1(e10,e10)=e11|op1(e10,e11)=e11)|op1(e10,e12)=e11)|op1(e10,e13)=e11)|op1(e10,e14)=e11))&((((op1(e10,e10)=e11|op1(e11,e10)=e11)|op1(e12,e10)=e11)|op1(e13,e10)=e11)|op1(e14,e10)=e11))&((((op1(e10,e10)=e12|op1(e10,e11)=e12)|op1(e10,e12)=e12)|op1(e10,e13)=e12)|op1(e10,e14)=e12))&((((op1(e10,e10)=e12|op1(e11,e10)=e12)|op1(e12,e10)=e12)|op1(e13,e10)=e12)|op1(e14,e10)=e12))&((((op1(e10,e10)=e13|op1(e10,e11)=e13)|op1(e10,e12)=e13)|op1(e10,e13)=e13)|op1(e10,e14)=e13))&((((op1(e10,e10)=e13|op1(e11,e10)=e13)|op1(e12,e10)=e13)|op1(e13,e10)=e13)|op1(e14,e10)=e13))&((((op1(e10,e10)=e14|op1(e10,e11)=e14)|op1(e10,e12)=e14)|op1(e10,e13)=e14)|op1(e10,e14)=e14))&((((op1(e10,e10)=e14|op1(e11,e10)=e14)|op1(e12,e10)=e14)|op1(e13,e10)=e14)|op1(e14,e10)=e14))&((((op1(e11,e10)=e10|op1(e11,e11)=e10)|op1(e11,e12)=e10)|op1(e11,e13)=e10)|op1(e11,e14)=e10))&((((op1(e10,e11)=e10|op1(e11,e11)=e10)|op1(e12,e11)=e10)|op1(e13,e11)=e10)|op1(e14,e11)=e10))&((((op1(e11,e10)=e11|op1(e11,e11)=e11)|op1(e11,e12)=e11)|op1(e11,e13)=e11)|op1(e11,e14)=e11))&((((op1(e10,e11)=e11|op1(e11,e11)=e11)|op1(e12,e11)=e11)|op1(e13,e11)=e11)|op1(e14,e11)=e11))&((((op1(e11,e10)=e12|op1(e11,e11)=e12)|op1(e11,e12)=e12)|op1(e11,e13)=e12)|op1(e11,e14)=e12))&((((op1(e10,e11)=e12|op1(e11,e11)=e12)|op1(e12,e11)=e12)|op1(e13,e11)=e12)|op1(e14,e11)=e12))&((((op1(e11,e10)=e13|op1(e11,e11)=e13)|op1(e11,e12)=e13)|op1(e11,e13)=e13)|op1(e11,e14)=e13))&((((op1(e10,e11)=e13|op1(e11,e11)=e13)|op1(e12,e11)=e13)|op1(e13,e11)=e13)|op1(e14,e11)=e13))&((((op1(e11,e10)=e14|op1(e11,e11)=e14)|op1(e11,e12)=e14)|op1(e11,e13)=e14)|op1(e11,e14)=e14))&((((op1(e10,e11)=e14|op1(e11,e11)=e14)|op1(e12,e11)=e14)|op1(e13,e11)=e14)|op1(e14,e11)=e14))&((((op1(e12,e10)=e10|op1(e12,e11)=e10)|op1(e12,e12)=e10)|op1(e12,e13)=e10)|op1(e12,e14)=e10))&((((op1(e10,e12)=e10|op1(e11,e12)=e10)|op1(e12,e12)=e10)|op1(e13,e12)=e10)|op1(e14,e12)=e10))&((((op1(e12,e10)=e11|op1(e12,e11)=e11)|op1(e12,e12)=e11)|op1(e12,e13)=e11)|op1(e12,e14)=e11))&((((op1(e10,e12)=e11|op1(e11,e12)=e11)|op1(e12,e12)=e11)|op1(e13,e12)=e11)|op1(e14,e12)=e11))&((((op1(e12,e10)=e12|op1(e12,e11)=e12)|op1(e12,e12)=e12)|op1(e12,e13)=e12)|op1(e12,e14)=e12))&((((op1(e10,e12)=e12|op1(e11,e12)=e12)|op1(e12,e12)=e12)|op1(e13,e12)=e12)|op1(e14,e12)=e12))&((((op1(e12,e10)=e13|op1(e12,e11)=e13)|op1(e12,e12)=e13)|op1(e12,e13)=e13)|op1(e12,e14)=e13))&((((op1(e10,e12)=e13|op1(e11,e12)=e13)|op1(e12,e12)=e13)|op1(e13,e12)=e13)|op1(e14,e12)=e13))&((((op1(e12,e10)=e14|op1(e12,e11)=e14)|op1(e12,e12)=e14)|op1(e12,e13)=e14)|op1(e12,e14)=e14))&((((op1(e10,e12)=e14|op1(e11,e12)=e14)|op1(e12,e12)=e14)|op1(e13,e12)=e14)|op1(e14,e12)=e14))&((((op1(e13,e10)=e10|op1(e13,e11)=e10)|op1(e13,e12)=e10)|op1(e13,e13)=e10)|op1(e13,e14)=e10))&((((op1(e10,e13)=e10|op1(e11,e13)=e10)|op1(e12,e13)=e10)|op1(e13,e13)=e10)|op1(e14,e13)=e10))&((((op1(e13,e10)=e11|op1(e13,e11)=e11)|op1(e13,e12)=e11)|op1(e13,e13)=e11)|op1(e13,e14)=e11))&((((op1(e10,e13)=e11|op1(e11,e13)=e11)|op1(e12,e13)=e11)|op1(e13,e13)=e11)|op1(e14,e13)=e11))&((((op1(e13,e10)=e12|op1(e13,e11)=e12)|op1(e13,e12)=e12)|op1(e13,e13)=e12)|op1(e13,e14)=e12))&((((op1(e10,e13)=e12|op1(e11,e13)=e12)|op1(e12,e13)=e12)|op1(e13,e13)=e12)|op1(e14,e13)=e12))&((((op1(e13,e10)=e13|op1(e13,e11)=e13)|op1(e13,e12)=e13)|op1(e13,e13)=e13)|op1(e13,e14)=e13))&((((op1(e10,e13)=e13|op1(e11,e13)=e13)|op1(e12,e13)=e13)|op1(e13,e13)=e13)|op1(e14,e13)=e13))&((((op1(e13,e10)=e14|op1(e13,e11)=e14)|op1(e13,e12)=e14)|op1(e13,e13)=e14)|op1(e13,e14)=e14))&((((op1(e10,e13)=e14|op1(e11,e13)=e14)|op1(e12,e13)=e14)|op1(e13,e13)=e14)|op1(e14,e13)=e14))&((((op1(e14,e10)=e10|op1(e14,e11)=e10)|op1(e14,e12)=e10)|op1(e14,e13)=e10)|op1(e14,e14)=e10))&((((op1(e10,e14)=e10|op1(e11,e14)=e10)|op1(e12,e14)=e10)|op1(e13,e14)=e10)|op1(e14,e14)=e10))&((((op1(e14,e10)=e11|op1(e14,e11)=e11)|op1(e14,e12)=e11)|op1(e14,e13)=e11)|op1(e14,e14)=e11))&((((op1(e10,e14)=e11|op1(e11,e14)=e11)|op1(e12,e14)=e11)|op1(e13,e14)=e11)|op1(e14,e14)=e11))&((((op1(e14,e10)=e12|op1(e14,e11)=e12)|op1(e14,e12)=e12)|op1(e14,e13)=e12)|op1(e14,e14)=e12))&((((op1(e10,e14)=e12|op1(e11,e14)=e12)|op1(e12,e14)=e12)|op1(e13,e14)=e12)|op1(e14,e14)=e12))&((((op1(e14,e10)=e13|op1(e14,e11)=e13)|op1(e14,e12)=e13)|op1(e14,e13)=e13)|op1(e14,e14)=e13))&((((op1(e10,e14)=e13|op1(e11,e14)=e13)|op1(e12,e14)=e13)|op1(e13,e14)=e13)|op1(e14,e14)=e13))&((((op1(e14,e10)=e14|op1(e14,e11)=e14)|op1(e14,e12)=e14)|op1(e14,e13)=e14)|op1(e14,e14)=e14))&((((op1(e10,e14)=e14|op1(e11,e14)=e14)|op1(e12,e14)=e14)|op1(e13,e14)=e14)|op1(e14,e14)=e14))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax2)).
fof(c_0_11, axiom, ((((((((((e10!=e11&e10!=e12)&e10!=e13)&e10!=e14)&e11!=e12)&e11!=e13)&e11!=e14)&e12!=e13)&e12!=e14)&e13!=e14)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax9)).
fof(c_0_12, axiom, (((((((((op1(e10,e10)!=e10&op1(e10,e11)!=e11)&op1(e10,e12)!=e12)&op1(e10,e13)!=e13)&op1(e10,e14)!=e14)|((((op1(e11,e10)!=e10&op1(e11,e11)!=e11)&op1(e11,e12)!=e12)&op1(e11,e13)!=e13)&op1(e11,e14)!=e14))|((((op1(e12,e10)!=e10&op1(e12,e11)!=e11)&op1(e12,e12)!=e12)&op1(e12,e13)!=e13)&op1(e12,e14)!=e14))|((((op1(e13,e10)!=e10&op1(e13,e11)!=e11)&op1(e13,e12)!=e12)&op1(e13,e13)!=e13)&op1(e13,e14)!=e14))|((((op1(e14,e10)!=e10&op1(e14,e11)!=e11)&op1(e14,e12)!=e12)&op1(e14,e13)!=e13)&op1(e14,e14)!=e14))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax12)).
fof(c_0_13, axiom, (((((((((((((((((((((((((((((op1(e10,e10)=e10|op1(e10,e10)=e11)|op1(e10,e10)=e12)|op1(e10,e10)=e13)|op1(e10,e10)=e14)&((((op1(e10,e11)=e10|op1(e10,e11)=e11)|op1(e10,e11)=e12)|op1(e10,e11)=e13)|op1(e10,e11)=e14))&((((op1(e10,e12)=e10|op1(e10,e12)=e11)|op1(e10,e12)=e12)|op1(e10,e12)=e13)|op1(e10,e12)=e14))&((((op1(e10,e13)=e10|op1(e10,e13)=e11)|op1(e10,e13)=e12)|op1(e10,e13)=e13)|op1(e10,e13)=e14))&((((op1(e10,e14)=e10|op1(e10,e14)=e11)|op1(e10,e14)=e12)|op1(e10,e14)=e13)|op1(e10,e14)=e14))&((((op1(e11,e10)=e10|op1(e11,e10)=e11)|op1(e11,e10)=e12)|op1(e11,e10)=e13)|op1(e11,e10)=e14))&((((op1(e11,e11)=e10|op1(e11,e11)=e11)|op1(e11,e11)=e12)|op1(e11,e11)=e13)|op1(e11,e11)=e14))&((((op1(e11,e12)=e10|op1(e11,e12)=e11)|op1(e11,e12)=e12)|op1(e11,e12)=e13)|op1(e11,e12)=e14))&((((op1(e11,e13)=e10|op1(e11,e13)=e11)|op1(e11,e13)=e12)|op1(e11,e13)=e13)|op1(e11,e13)=e14))&((((op1(e11,e14)=e10|op1(e11,e14)=e11)|op1(e11,e14)=e12)|op1(e11,e14)=e13)|op1(e11,e14)=e14))&((((op1(e12,e10)=e10|op1(e12,e10)=e11)|op1(e12,e10)=e12)|op1(e12,e10)=e13)|op1(e12,e10)=e14))&((((op1(e12,e11)=e10|op1(e12,e11)=e11)|op1(e12,e11)=e12)|op1(e12,e11)=e13)|op1(e12,e11)=e14))&((((op1(e12,e12)=e10|op1(e12,e12)=e11)|op1(e12,e12)=e12)|op1(e12,e12)=e13)|op1(e12,e12)=e14))&((((op1(e12,e13)=e10|op1(e12,e13)=e11)|op1(e12,e13)=e12)|op1(e12,e13)=e13)|op1(e12,e13)=e14))&((((op1(e12,e14)=e10|op1(e12,e14)=e11)|op1(e12,e14)=e12)|op1(e12,e14)=e13)|op1(e12,e14)=e14))&((((op1(e13,e10)=e10|op1(e13,e10)=e11)|op1(e13,e10)=e12)|op1(e13,e10)=e13)|op1(e13,e10)=e14))&((((op1(e13,e11)=e10|op1(e13,e11)=e11)|op1(e13,e11)=e12)|op1(e13,e11)=e13)|op1(e13,e11)=e14))&((((op1(e13,e12)=e10|op1(e13,e12)=e11)|op1(e13,e12)=e12)|op1(e13,e12)=e13)|op1(e13,e12)=e14))&((((op1(e13,e13)=e10|op1(e13,e13)=e11)|op1(e13,e13)=e12)|op1(e13,e13)=e13)|op1(e13,e13)=e14))&((((op1(e13,e14)=e10|op1(e13,e14)=e11)|op1(e13,e14)=e12)|op1(e13,e14)=e13)|op1(e13,e14)=e14))&((((op1(e14,e10)=e10|op1(e14,e10)=e11)|op1(e14,e10)=e12)|op1(e14,e10)=e13)|op1(e14,e10)=e14))&((((op1(e14,e11)=e10|op1(e14,e11)=e11)|op1(e14,e11)=e12)|op1(e14,e11)=e13)|op1(e14,e11)=e14))&((((op1(e14,e12)=e10|op1(e14,e12)=e11)|op1(e14,e12)=e12)|op1(e14,e12)=e13)|op1(e14,e12)=e14))&((((op1(e14,e13)=e10|op1(e14,e13)=e11)|op1(e14,e13)=e12)|op1(e14,e13)=e13)|op1(e14,e13)=e14))&((((op1(e14,e14)=e10|op1(e14,e14)=e11)|op1(e14,e14)=e12)|op1(e14,e14)=e13)|op1(e14,e14)=e14))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax1)).
fof(c_0_14, conjecture, ((((((((((((((((((((((((((((((((((h1(op1(e10,e10))=op2(h1(e10),h1(e10))&h1(op1(e10,e11))=op2(h1(e10),h1(e11)))&h1(op1(e10,e12))=op2(h1(e10),h1(e12)))&h1(op1(e10,e13))=op2(h1(e10),h1(e13)))&h1(op1(e10,e14))=op2(h1(e10),h1(e14)))&h1(op1(e11,e10))=op2(h1(e11),h1(e10)))&h1(op1(e11,e11))=op2(h1(e11),h1(e11)))&h1(op1(e11,e12))=op2(h1(e11),h1(e12)))&h1(op1(e11,e13))=op2(h1(e11),h1(e13)))&h1(op1(e11,e14))=op2(h1(e11),h1(e14)))&h1(op1(e12,e10))=op2(h1(e12),h1(e10)))&h1(op1(e12,e11))=op2(h1(e12),h1(e11)))&h1(op1(e12,e12))=op2(h1(e12),h1(e12)))&h1(op1(e12,e13))=op2(h1(e12),h1(e13)))&h1(op1(e12,e14))=op2(h1(e12),h1(e14)))&h1(op1(e13,e10))=op2(h1(e13),h1(e10)))&h1(op1(e13,e11))=op2(h1(e13),h1(e11)))&h1(op1(e13,e12))=op2(h1(e13),h1(e12)))&h1(op1(e13,e13))=op2(h1(e13),h1(e13)))&h1(op1(e13,e14))=op2(h1(e13),h1(e14)))&h1(op1(e14,e10))=op2(h1(e14),h1(e10)))&h1(op1(e14,e11))=op2(h1(e14),h1(e11)))&h1(op1(e14,e12))=op2(h1(e14),h1(e12)))&h1(op1(e14,e13))=op2(h1(e14),h1(e13)))&h1(op1(e14,e14))=op2(h1(e14),h1(e14)))&((((h1(e10)=e20|h1(e11)=e20)|h1(e12)=e20)|h1(e13)=e20)|h1(e14)=e20))&((((h1(e10)=e21|h1(e11)=e21)|h1(e12)=e21)|h1(e13)=e21)|h1(e14)=e21))&((((h1(e10)=e22|h1(e11)=e22)|h1(e12)=e22)|h1(e13)=e22)|h1(e14)=e22))&((((h1(e10)=e23|h1(e11)=e23)|h1(e12)=e23)|h1(e13)=e23)|h1(e14)=e23))&((((h1(e10)=e24|h1(e11)=e24)|h1(e12)=e24)|h1(e13)=e24)|h1(e14)=e24))|(((((((((((((((((((((((((((((h2(op1(e10,e10))=op2(h2(e10),h2(e10))&h2(op1(e10,e11))=op2(h2(e10),h2(e11)))&h2(op1(e10,e12))=op2(h2(e10),h2(e12)))&h2(op1(e10,e13))=op2(h2(e10),h2(e13)))&h2(op1(e10,e14))=op2(h2(e10),h2(e14)))&h2(op1(e11,e10))=op2(h2(e11),h2(e10)))&h2(op1(e11,e11))=op2(h2(e11),h2(e11)))&h2(op1(e11,e12))=op2(h2(e11),h2(e12)))&h2(op1(e11,e13))=op2(h2(e11),h2(e13)))&h2(op1(e11,e14))=op2(h2(e11),h2(e14)))&h2(op1(e12,e10))=op2(h2(e12),h2(e10)))&h2(op1(e12,e11))=op2(h2(e12),h2(e11)))&h2(op1(e12,e12))=op2(h2(e12),h2(e12)))&h2(op1(e12,e13))=op2(h2(e12),h2(e13)))&h2(op1(e12,e14))=op2(h2(e12),h2(e14)))&h2(op1(e13,e10))=op2(h2(e13),h2(e10)))&h2(op1(e13,e11))=op2(h2(e13),h2(e11)))&h2(op1(e13,e12))=op2(h2(e13),h2(e12)))&h2(op1(e13,e13))=op2(h2(e13),h2(e13)))&h2(op1(e13,e14))=op2(h2(e13),h2(e14)))&h2(op1(e14,e10))=op2(h2(e14),h2(e10)))&h2(op1(e14,e11))=op2(h2(e14),h2(e11)))&h2(op1(e14,e12))=op2(h2(e14),h2(e12)))&h2(op1(e14,e13))=op2(h2(e14),h2(e13)))&h2(op1(e14,e14))=op2(h2(e14),h2(e14)))&((((h2(e10)=e20|h2(e11)=e20)|h2(e12)=e20)|h2(e13)=e20)|h2(e14)=e20))&((((h2(e10)=e21|h2(e11)=e21)|h2(e12)=e21)|h2(e13)=e21)|h2(e14)=e21))&((((h2(e10)=e22|h2(e11)=e22)|h2(e12)=e22)|h2(e13)=e22)|h2(e14)=e22))&((((h2(e10)=e23|h2(e11)=e23)|h2(e12)=e23)|h2(e13)=e23)|h2(e14)=e23))&((((h2(e10)=e24|h2(e11)=e24)|h2(e12)=e24)|h2(e13)=e24)|h2(e14)=e24)))|(((((((((((((((((((((((((((((h3(op1(e10,e10))=op2(h3(e10),h3(e10))&h3(op1(e10,e11))=op2(h3(e10),h3(e11)))&h3(op1(e10,e12))=op2(h3(e10),h3(e12)))&h3(op1(e10,e13))=op2(h3(e10),h3(e13)))&h3(op1(e10,e14))=op2(h3(e10),h3(e14)))&h3(op1(e11,e10))=op2(h3(e11),h3(e10)))&h3(op1(e11,e11))=op2(h3(e11),h3(e11)))&h3(op1(e11,e12))=op2(h3(e11),h3(e12)))&h3(op1(e11,e13))=op2(h3(e11),h3(e13)))&h3(op1(e11,e14))=op2(h3(e11),h3(e14)))&h3(op1(e12,e10))=op2(h3(e12),h3(e10)))&h3(op1(e12,e11))=op2(h3(e12),h3(e11)))&h3(op1(e12,e12))=op2(h3(e12),h3(e12)))&h3(op1(e12,e13))=op2(h3(e12),h3(e13)))&h3(op1(e12,e14))=op2(h3(e12),h3(e14)))&h3(op1(e13,e10))=op2(h3(e13),h3(e10)))&h3(op1(e13,e11))=op2(h3(e13),h3(e11)))&h3(op1(e13,e12))=op2(h3(e13),h3(e12)))&h3(op1(e13,e13))=op2(h3(e13),h3(e13)))&h3(op1(e13,e14))=op2(h3(e13),h3(e14)))&h3(op1(e14,e10))=op2(h3(e14),h3(e10)))&h3(op1(e14,e11))=op2(h3(e14),h3(e11)))&h3(op1(e14,e12))=op2(h3(e14),h3(e12)))&h3(op1(e14,e13))=op2(h3(e14),h3(e13)))&h3(op1(e14,e14))=op2(h3(e14),h3(e14)))&((((h3(e10)=e20|h3(e11)=e20)|h3(e12)=e20)|h3(e13)=e20)|h3(e14)=e20))&((((h3(e10)=e21|h3(e11)=e21)|h3(e12)=e21)|h3(e13)=e21)|h3(e14)=e21))&((((h3(e10)=e22|h3(e11)=e22)|h3(e12)=e22)|h3(e13)=e22)|h3(e14)=e22))&((((h3(e10)=e23|h3(e11)=e23)|h3(e12)=e23)|h3(e13)=e23)|h3(e14)=e23))&((((h3(e10)=e24|h3(e11)=e24)|h3(e12)=e24)|h3(e13)=e24)|h3(e14)=e24)))|(((((((((((((((((((((((((((((h4(op1(e10,e10))=op2(h4(e10),h4(e10))&h4(op1(e10,e11))=op2(h4(e10),h4(e11)))&h4(op1(e10,e12))=op2(h4(e10),h4(e12)))&h4(op1(e10,e13))=op2(h4(e10),h4(e13)))&h4(op1(e10,e14))=op2(h4(e10),h4(e14)))&h4(op1(e11,e10))=op2(h4(e11),h4(e10)))&h4(op1(e11,e11))=op2(h4(e11),h4(e11)))&h4(op1(e11,e12))=op2(h4(e11),h4(e12)))&h4(op1(e11,e13))=op2(h4(e11),h4(e13)))&h4(op1(e11,e14))=op2(h4(e11),h4(e14)))&h4(op1(e12,e10))=op2(h4(e12),h4(e10)))&h4(op1(e12,e11))=op2(h4(e12),h4(e11)))&h4(op1(e12,e12))=op2(h4(e12),h4(e12)))&h4(op1(e12,e13))=op2(h4(e12),h4(e13)))&h4(op1(e12,e14))=op2(h4(e12),h4(e14)))&h4(op1(e13,e10))=op2(h4(e13),h4(e10)))&h4(op1(e13,e11))=op2(h4(e13),h4(e11)))&h4(op1(e13,e12))=op2(h4(e13),h4(e12)))&h4(op1(e13,e13))=op2(h4(e13),h4(e13)))&h4(op1(e13,e14))=op2(h4(e13),h4(e14)))&h4(op1(e14,e10))=op2(h4(e14),h4(e10)))&h4(op1(e14,e11))=op2(h4(e14),h4(e11)))&h4(op1(e14,e12))=op2(h4(e14),h4(e12)))&h4(op1(e14,e13))=op2(h4(e14),h4(e13)))&h4(op1(e14,e14))=op2(h4(e14),h4(e14)))&((((h4(e10)=e20|h4(e11)=e20)|h4(e12)=e20)|h4(e13)=e20)|h4(e14)=e20))&((((h4(e10)=e21|h4(e11)=e21)|h4(e12)=e21)|h4(e13)=e21)|h4(e14)=e21))&((((h4(e10)=e22|h4(e11)=e22)|h4(e12)=e22)|h4(e13)=e22)|h4(e14)=e22))&((((h4(e10)=e23|h4(e11)=e23)|h4(e12)=e23)|h4(e13)=e23)|h4(e14)=e23))&((((h4(e10)=e24|h4(e11)=e24)|h4(e12)=e24)|h4(e13)=e24)|h4(e14)=e24)))|(((((((((((((((((((((((((((((h5(op1(e10,e10))=op2(h5(e10),h5(e10))&h5(op1(e10,e11))=op2(h5(e10),h5(e11)))&h5(op1(e10,e12))=op2(h5(e10),h5(e12)))&h5(op1(e10,e13))=op2(h5(e10),h5(e13)))&h5(op1(e10,e14))=op2(h5(e10),h5(e14)))&h5(op1(e11,e10))=op2(h5(e11),h5(e10)))&h5(op1(e11,e11))=op2(h5(e11),h5(e11)))&h5(op1(e11,e12))=op2(h5(e11),h5(e12)))&h5(op1(e11,e13))=op2(h5(e11),h5(e13)))&h5(op1(e11,e14))=op2(h5(e11),h5(e14)))&h5(op1(e12,e10))=op2(h5(e12),h5(e10)))&h5(op1(e12,e11))=op2(h5(e12),h5(e11)))&h5(op1(e12,e12))=op2(h5(e12),h5(e12)))&h5(op1(e12,e13))=op2(h5(e12),h5(e13)))&h5(op1(e12,e14))=op2(h5(e12),h5(e14)))&h5(op1(e13,e10))=op2(h5(e13),h5(e10)))&h5(op1(e13,e11))=op2(h5(e13),h5(e11)))&h5(op1(e13,e12))=op2(h5(e13),h5(e12)))&h5(op1(e13,e13))=op2(h5(e13),h5(e13)))&h5(op1(e13,e14))=op2(h5(e13),h5(e14)))&h5(op1(e14,e10))=op2(h5(e14),h5(e10)))&h5(op1(e14,e11))=op2(h5(e14),h5(e11)))&h5(op1(e14,e12))=op2(h5(e14),h5(e12)))&h5(op1(e14,e13))=op2(h5(e14),h5(e13)))&h5(op1(e14,e14))=op2(h5(e14),h5(e14)))&((((h5(e10)=e20|h5(e11)=e20)|h5(e12)=e20)|h5(e13)=e20)|h5(e14)=e20))&((((h5(e10)=e21|h5(e11)=e21)|h5(e12)=e21)|h5(e13)=e21)|h5(e14)=e21))&((((h5(e10)=e22|h5(e11)=e22)|h5(e12)=e22)|h5(e13)=e22)|h5(e14)=e22))&((((h5(e10)=e23|h5(e11)=e23)|h5(e12)=e23)|h5(e13)=e23)|h5(e14)=e23))&((((h5(e10)=e24|h5(e11)=e24)|h5(e12)=e24)|h5(e13)=e24)|h5(e14)=e24)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
fof(c_0_15, axiom, (((((h5(e13)=e24&h5(e10)=op2(op2(e24,op2(e24,e24)),op2(e24,op2(e24,e24))))&h5(e11)=op2(e24,op2(e24,op2(e24,e24))))&h5(e12)=op2(e24,e24))&h5(e14)=op2(e24,op2(e24,e24)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', ax20)).
cnf(c_0_16,plain,(e24=op2(e23,op2(e23,e23))), inference(split_conjunct,[status(thm)],[c_0_0])).
cnf(c_0_17,plain,(e22=op2(e23,e23)), inference(split_conjunct,[status(thm)],[c_0_0])).
cnf(c_0_18,plain,(e20=op2(op2(e23,op2(e23,e23)),op2(e23,op2(e23,e23)))), inference(split_conjunct,[status(thm)],[c_0_0])).
cnf(c_0_19,plain,(op2(e23,e22)=e24), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20,plain,(op2(op2(e24,e24),op2(e24,e24))=e24), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_21,plain,(op2(e24,e24)=e20), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_17]), c_0_19]), c_0_17]), c_0_19])).
fof(c_0_22, plain, ((epred4_0<=>(((((((op2(e20,e20)!=e20&op2(e20,e21)!=e21)&op2(e20,e22)!=e22)&op2(e20,e23)!=e23)&op2(e20,e24)!=e24)|((((op2(e21,e20)!=e20&op2(e21,e21)!=e21)&op2(e21,e22)!=e22)&op2(e21,e23)!=e23)&op2(e21,e24)!=e24))|((((op2(e22,e20)!=e20&op2(e22,e21)!=e21)&op2(e22,e22)!=e22)&op2(e22,e23)!=e23)&op2(e22,e24)!=e24))|((((op2(e23,e20)!=e20&op2(e23,e21)!=e21)&op2(e23,e22)!=e22)&op2(e23,e23)!=e23)&op2(e23,e24)!=e24)))), introduced(definition)).
cnf(c_0_23,plain,(op2(e23,e23)!=op2(e24,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_24,plain,(e14=op1(e13,op1(e13,e13))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_25,plain,(e12=op1(e13,e13)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_26,plain,(op2(e24,e23)=e24|op2(e23,e23)=e24|op2(e22,e23)=e24|op2(e21,e23)=e24|op2(e20,e23)=e24), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_27,plain,(e22!=e24), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_28,plain,(op2(e20,e20)!=op2(e20,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_29,plain,(op2(e20,e20)=e24), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_21])).
fof(c_0_30, axiom, ((epred4_0|((((op2(e24,e20)!=e20&op2(e24,e21)!=e21)&op2(e24,e22)!=e22)&op2(e24,e23)!=e23)&op2(e24,e24)!=e24))), inference(apply_def,[status(thm)],[c_0_6, c_0_22])).
cnf(c_0_31,plain,(op2(e24,e23)=e24|op2(e24,e23)=e23|op2(e24,e23)=e22|op2(e24,e23)=e21|op2(e24,e23)=e20), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_32,plain,(op2(e24,e23)!=e22), inference(rw,[status(thm)],[c_0_23, c_0_17])).
cnf(c_0_33,plain,(op2(e24,e23)!=op2(e24,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_34,plain,(e10=op1(op1(e13,op1(e13,e13)),op1(e13,op1(e13,e13)))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_35,plain,(op1(e13,e12)=e14), inference(rw,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_36, plain, ((epred3_0<=>(((((op2(e20,e20)!=e20&op2(e20,e21)!=e21)&op2(e20,e22)!=e22)&op2(e20,e23)!=e23)&op2(e20,e24)!=e24)|((((op2(e21,e20)!=e20&op2(e21,e21)!=e21)&op2(e21,e22)!=e22)&op2(e21,e23)!=e23)&op2(e21,e24)!=e24)))), introduced(definition)).
cnf(c_0_37,plain,(e21=op2(e23,op2(e23,op2(e23,e23)))), inference(split_conjunct,[status(thm)],[c_0_0])).
cnf(c_0_38,plain,(op2(op2(e22,e23),op2(e23,e22))=e23), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_39,plain,(op2(e24,e23)=e24|op2(e22,e23)=e24|op2(e21,e23)=e24|op2(e20,e23)=e24), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_17]), c_0_27])).
cnf(c_0_40,plain,(op2(e20,e23)!=e24), inference(rw,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_41, plain, ((((((op2(e24,e20)!=e20|epred4_0)&(op2(e24,e21)!=e21|epred4_0))&(op2(e24,e22)!=e22|epred4_0))&(op2(e24,e23)!=e23|epred4_0))&(op2(e24,e24)!=e24|epred4_0))), inference(distribute,[status(thm)],[c_0_30])).
cnf(c_0_42,plain,(op2(e24,e23)=e24|op2(e24,e23)=e23|op2(e24,e23)=e21|op2(e24,e23)=e20), inference(sr,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_43,plain,(op2(e24,e23)!=e20), inference(rw,[status(thm)],[c_0_33, c_0_21])).
cnf(c_0_44,plain,(op1(op1(e14,e14),op1(e14,e14))=e14), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_45,plain,(op1(e14,e14)=e10), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_25]), c_0_35]), c_0_25]), c_0_35])).
fof(c_0_46, plain, ((epred2_0<=>(((((((op1(e10,e10)!=e10&op1(e10,e11)!=e11)&op1(e10,e12)!=e12)&op1(e10,e13)!=e13)&op1(e10,e14)!=e14)|((((op1(e11,e10)!=e10&op1(e11,e11)!=e11)&op1(e11,e12)!=e12)&op1(e11,e13)!=e13)&op1(e11,e14)!=e14))|((((op1(e12,e10)!=e10&op1(e12,e11)!=e11)&op1(e12,e12)!=e12)&op1(e12,e13)!=e13)&op1(e12,e14)!=e14))|((((op1(e13,e10)!=e10&op1(e13,e11)!=e11)&op1(e13,e12)!=e12)&op1(e13,e13)!=e13)&op1(e13,e14)!=e14)))), introduced(definition)).
cnf(c_0_47,plain,(op1(e13,e13)!=op1(e14,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_48, plain, ((epred4_0=>((epred3_0|((((op2(e22,e20)!=e20&op2(e22,e21)!=e21)&op2(e22,e22)!=e22)&op2(e22,e23)!=e23)&op2(e22,e24)!=e24))|((((op2(e23,e20)!=e20&op2(e23,e21)!=e21)&op2(e23,e22)!=e22)&op2(e23,e23)!=e23)&op2(e23,e24)!=e24)))), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_22]), c_0_36])).
cnf(c_0_49,plain,(op2(e24,e24)=e24|op2(e23,e24)=e24|op2(e22,e24)=e24|op2(e21,e24)=e24|op2(e20,e24)=e24), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_50,plain,(op2(e23,e24)=e21), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_17]), c_0_19])).
cnf(c_0_51,plain,(e21!=e24), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_52,plain,(e20!=e24), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_53,plain,(op2(e20,e20)!=op2(e20,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_54,plain,(op2(op2(e22,e23),e24)=e23), inference(rw,[status(thm)],[c_0_38, c_0_19])).
cnf(c_0_55,plain,(op2(e21,e23)=e24|op2(e22,e23)=e24|op2(e24,e23)=e24), inference(sr,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_56,plain,(e20!=e23), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_57,plain,(op2(op2(e23,e24),op2(e24,e23))=e24), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_58,plain,(epred4_0|op2(e24,e23)!=e23), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_59,plain,(op2(e24,e23)=e21|op2(e24,e23)=e23|op2(e24,e23)=e24), inference(sr,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_60,plain,(op1(e14,e13)=e14|op1(e13,e13)=e14|op1(e12,e13)=e14|op1(e11,e13)=e14|op1(e10,e13)=e14), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_61,plain,(e12!=e14), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_62,plain,(op1(e10,e10)!=op1(e10,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_63,plain,(op1(e10,e10)=e14), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45]), c_0_45])).
cnf(c_0_64,plain,(e11=op1(e13,op1(e13,op1(e13,e13)))), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_65, axiom, ((epred2_0|((((op1(e14,e10)!=e10&op1(e14,e11)!=e11)&op1(e14,e12)!=e12)&op1(e14,e13)!=e13)&op1(e14,e14)!=e14))), inference(apply_def,[status(thm)],[c_0_12, c_0_46])).
cnf(c_0_66,plain,(op1(e14,e13)=e14|op1(e14,e13)=e13|op1(e14,e13)=e12|op1(e14,e13)=e11|op1(e14,e13)=e10), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_67,plain,(op1(e14,e13)!=e12), inference(rw,[status(thm)],[c_0_47, c_0_25])).
cnf(c_0_68,plain,(op1(e14,e13)!=op1(e14,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_69,plain,(op2(e23,e22)!=op2(e24,e22)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_70,plain,(op2(op2(e23,e23),op2(e23,e23))=e23), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_71,plain,(op2(e23,e21)!=op2(e23,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_72,plain,(op2(e23,e21)!=op2(e23,e22)), inference(split_conjunct,[status(thm)],[c_0_2])).
fof(c_0_73, plain, (((((((((((op2(e23,e20)!=e20|(op2(e22,e20)!=e20|epred3_0))|~epred4_0)&((op2(e23,e21)!=e21|(op2(e22,e20)!=e20|epred3_0))|~epred4_0))&((op2(e23,e22)!=e22|(op2(e22,e20)!=e20|epred3_0))|~epred4_0))&((op2(e23,e23)!=e23|(op2(e22,e20)!=e20|epred3_0))|~epred4_0))&((op2(e23,e24)!=e24|(op2(e22,e20)!=e20|epred3_0))|~epred4_0))&((((((op2(e23,e20)!=e20|(op2(e22,e21)!=e21|epred3_0))|~epred4_0)&((op2(e23,e21)!=e21|(op2(e22,e21)!=e21|epred3_0))|~epred4_0))&((op2(e23,e22)!=e22|(op2(e22,e21)!=e21|epred3_0))|~epred4_0))&((op2(e23,e23)!=e23|(op2(e22,e21)!=e21|epred3_0))|~epred4_0))&((op2(e23,e24)!=e24|(op2(e22,e21)!=e21|epred3_0))|~epred4_0)))&((((((op2(e23,e20)!=e20|(op2(e22,e22)!=e22|epred3_0))|~epred4_0)&((op2(e23,e21)!=e21|(op2(e22,e22)!=e22|epred3_0))|~epred4_0))&((op2(e23,e22)!=e22|(op2(e22,e22)!=e22|epred3_0))|~epred4_0))&((op2(e23,e23)!=e23|(op2(e22,e22)!=e22|epred3_0))|~epred4_0))&((op2(e23,e24)!=e24|(op2(e22,e22)!=e22|epred3_0))|~epred4_0)))&((((((op2(e23,e20)!=e20|(op2(e22,e23)!=e23|epred3_0))|~epred4_0)&((op2(e23,e21)!=e21|(op2(e22,e23)!=e23|epred3_0))|~epred4_0))&((op2(e23,e22)!=e22|(op2(e22,e23)!=e23|epred3_0))|~epred4_0))&((op2(e23,e23)!=e23|(op2(e22,e23)!=e23|epred3_0))|~epred4_0))&((op2(e23,e24)!=e24|(op2(e22,e23)!=e23|epred3_0))|~epred4_0)))&((((((op2(e23,e20)!=e20|(op2(e22,e24)!=e24|epred3_0))|~epred4_0)&((op2(e23,e21)!=e21|(op2(e22,e24)!=e24|epred3_0))|~epred4_0))&((op2(e23,e22)!=e22|(op2(e22,e24)!=e24|epred3_0))|~epred4_0))&((op2(e23,e23)!=e23|(op2(e22,e24)!=e24|epred3_0))|~epred4_0))&((op2(e23,e24)!=e24|(op2(e22,e24)!=e24|epred3_0))|~epred4_0)))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_48])])).
cnf(c_0_74,plain,(op2(e22,e24)=e24|op2(e21,e24)=e24|op2(e20,e24)=e24), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50]), c_0_21]), c_0_51]), c_0_52])).
cnf(c_0_75,plain,(op2(e20,e24)!=e24), inference(rw,[status(thm)],[c_0_53, c_0_29])).
cnf(c_0_76,plain,(op2(op2(e23,e21),op2(e21,e23))=e21), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_77,plain,(op2(e24,e23)=e24|op2(e21,e23)=e24), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_21]), c_0_56])).
cnf(c_0_78,plain,(op2(e23,e24)=e20|op2(e23,e23)=e20|op2(e23,e22)=e20|op2(e23,e21)=e20|op2(e23,e20)=e20), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_79,plain,(e20!=e22), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_80,plain,(e20!=e21), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_81,plain,(op2(e20,e24)!=op2(e23,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_82,plain,(op2(e21,op2(e24,e23))=e24), inference(rw,[status(thm)],[c_0_57, c_0_50])).
cnf(c_0_83,plain,(op2(e24,e23)=e24|op2(e24,e23)=e21|epred4_0), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_84,plain,(op1(e14,e13)=e14|op1(e12,e13)=e14|op1(e11,e13)=e14|op1(e10,e13)=e14), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_25]), c_0_61])).
cnf(c_0_85,plain,(op1(e10,e13)!=e14), inference(rw,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_86,plain,(op1(e14,e14)=e14|op1(e13,e14)=e14|op1(e12,e14)=e14|op1(e11,e14)=e14|op1(e10,e14)=e14), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_87,plain,(op1(e13,e14)=e11), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_25]), c_0_35])).
cnf(c_0_88,plain,(e11!=e14), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_89,plain,(e10!=e14), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_90,plain,(op1(e10,e10)!=op1(e10,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_91, plain, ((((((op1(e14,e10)!=e10|epred2_0)&(op1(e14,e11)!=e11|epred2_0))&(op1(e14,e12)!=e12|epred2_0))&(op1(e14,e13)!=e13|epred2_0))&(op1(e14,e14)!=e14|epred2_0))), inference(distribute,[status(thm)],[c_0_65])).
cnf(c_0_92,plain,(op1(e14,e13)=e14|op1(e14,e13)=e13|op1(e14,e13)=e11|op1(e14,e13)=e10), inference(sr,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_93,plain,(op1(e14,e13)!=e10), inference(rw,[status(thm)],[c_0_68, c_0_45])).
cnf(c_0_94,plain,(op1(e13,e10)!=op1(e13,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_95,plain,(op1(op1(e13,e13),op1(e13,e13))=e13), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_96,plain,(op2(e24,e22)=e24|op2(e24,e22)=e23|op2(e24,e22)=e22|op2(e24,e22)=e21|op2(e24,e22)=e20), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_97,plain,(op2(e24,e22)!=e24), inference(rw,[status(thm)],[c_0_69, c_0_19])).
cnf(c_0_98,plain,(op2(e22,e22)!=op2(e24,e22)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_99,plain,(op2(e22,e22)=e23), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_17]), c_0_17])).
cnf(c_0_100,plain,(op2(e24,e22)!=op2(e24,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_101,plain,(op2(e23,e21)=e24|op2(e23,e21)=e23|op2(e23,e21)=e22|op2(e23,e21)=e21|op2(e23,e21)=e20), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_102,plain,(op2(e23,e21)!=e22), inference(rw,[status(thm)],[c_0_71, c_0_17])).
cnf(c_0_103,plain,(op2(e23,e21)!=e24), inference(rw,[status(thm)],[c_0_72, c_0_19])).
cnf(c_0_104,plain,(op2(e23,e21)!=op2(e23,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_105,plain,(epred3_0|~epred4_0|op2(e22,e24)!=e24|op2(e23,e20)!=e20), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_106,plain,(op2(e21,e24)=e24|op2(e22,e24)=e24), inference(sr,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_107,plain,(op2(op2(e23,e21),e24)=e21|op2(e24,e23)=e24), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_108,plain,(op2(e23,e21)=e20|op2(e23,e20)=e20), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_19]), c_0_17]), c_0_50]), c_0_52]), c_0_79]), c_0_80])).
cnf(c_0_109,plain,(op2(e20,e24)!=e21), inference(rw,[status(thm)],[c_0_81, c_0_50])).
cnf(c_0_110,plain,(op2(op2(e21,e21),op2(e21,e21))=e21), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_111,plain,(op2(e24,e23)=e24|op2(e21,e21)=e24|epred4_0), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_112,plain,(op2(e21,e23)!=op2(e21,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
fof(c_0_113, plain, ((epred1_0<=>(((((op1(e10,e10)!=e10&op1(e10,e11)!=e11)&op1(e10,e12)!=e12)&op1(e10,e13)!=e13)&op1(e10,e14)!=e14)|((((op1(e11,e10)!=e10&op1(e11,e11)!=e11)&op1(e11,e12)!=e12)&op1(e11,e13)!=e13)&op1(e11,e14)!=e14)))), introduced(definition)).
cnf(c_0_114,plain,(op1(e11,e13)!=op1(e11,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_115,plain,(op1(e11,e13)=e14|op1(e12,e13)=e14|op1(e14,e13)=e14), inference(sr,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_116,plain,(op1(e12,e14)=e14|op1(e11,e14)=e14|op1(e10,e14)=e14), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_87]), c_0_45]), c_0_88]), c_0_89])).
cnf(c_0_117,plain,(op1(e10,e14)!=e14), inference(rw,[status(thm)],[c_0_90, c_0_63])).
cnf(c_0_118,plain,(op1(op1(e13,e14),op1(e14,e13))=e14), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_119,plain,(epred2_0|op1(e14,e13)!=e13), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_120,plain,(op1(e14,e13)=e11|op1(e14,e13)=e13|op1(e14,e13)=e14), inference(sr,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_121,plain,(op1(e13,e10)=e14|op1(e13,e10)=e13|op1(e13,e10)=e12|op1(e13,e10)=e11|op1(e13,e10)=e10), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_122,plain,(op1(e13,e10)!=e12), inference(rw,[status(thm)],[c_0_94, c_0_25])).
cnf(c_0_123,plain,(op1(e13,e10)!=op1(e13,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_124,plain,(op1(e13,e10)!=op1(e13,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_125,plain,(op1(e14,e14)=e13|op1(e14,e13)=e13|op1(e14,e12)=e13|op1(e14,e11)=e13|op1(e14,e10)=e13), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_126,plain,(e10!=e13), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_127,plain,(op1(e12,e12)!=op1(e14,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_128,plain,(op1(e12,e12)=e13), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_25]), c_0_25])).
cnf(c_0_129,plain,(op1(e13,e14)=e13|op1(e13,e13)=e13|op1(e13,e12)=e13|op1(e13,e11)=e13|op1(e13,e10)=e13), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_130,plain,(e13!=e14), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_131,plain,(e12!=e13), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_132,plain,(e11!=e13), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_133,plain,(op2(e24,e22)=e23|op2(e24,e22)=e22|op2(e24,e22)=e21|op2(e24,e22)=e20), inference(sr,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_134,plain,(op2(e24,e22)!=e23), inference(rw,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_135,plain,(op2(e24,e22)!=e20), inference(rw,[status(thm)],[c_0_100, c_0_21])).
cnf(c_0_136,plain,(op2(e23,e21)=e23|op2(e23,e21)=e21|op2(e23,e21)=e20), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_101, c_0_102]), c_0_103])).
cnf(c_0_137,plain,(op2(e23,e21)!=e21), inference(rw,[status(thm)],[c_0_104, c_0_50])).
cnf(c_0_138,plain,(op2(op2(e24,e23),op2(e23,e24))=e23), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_139,plain,(op2(e21,e24)=e24|epred3_0|op2(e23,e20)!=e20|~epred4_0), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_140,plain,(op2(e23,e20)=e20|op2(e24,e23)=e24), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109])).
cnf(c_0_141,plain,(op2(e24,e23)=e24|epred4_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_21]), c_0_80])).
cnf(c_0_142,plain,(op2(e24,e23)=e24|op2(e21,e24)!=e24), inference(spm,[status(thm)],[c_0_112, c_0_77])).
fof(c_0_143, plain, ((epred2_0=>((epred1_0|((((op1(e12,e10)!=e10&op1(e12,e11)!=e11)&op1(e12,e12)!=e12)&op1(e12,e13)!=e13)&op1(e12,e14)!=e14))|((((op1(e13,e10)!=e10&op1(e13,e11)!=e11)&op1(e13,e12)!=e12)&op1(e13,e13)!=e13)&op1(e13,e14)!=e14)))), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_46]), c_0_113])).
cnf(c_0_144,plain,(op1(op1(e12,e13),op1(e13,e12))=e13), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_145,plain,(op1(e14,e13)=e14|op1(e12,e13)=e14|op1(e11,e14)!=e14), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_146,plain,(op1(e11,e14)=e14|op1(e12,e14)=e14), inference(sr,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_147,plain,(op1(e11,op1(e14,e13))=e14), inference(rw,[status(thm)],[c_0_118, c_0_87])).
cnf(c_0_148,plain,(op1(e14,e13)=e14|op1(e14,e13)=e11|epred2_0), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_149,plain,(op1(e13,e10)=e14|op1(e13,e10)=e13|op1(e13,e10)=e11|op1(e13,e10)=e10), inference(sr,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_150,plain,(op1(e13,e10)!=e14), inference(rw,[status(thm)],[c_0_123, c_0_35])).
cnf(c_0_151,plain,(op1(e13,e10)!=e11), inference(rw,[status(thm)],[c_0_124, c_0_87])).
cnf(c_0_152,plain,(op1(op1(e14,e13),op1(e13,e14))=e13), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_153,plain,(op1(e14,e13)=e13|op1(e14,e12)=e13|op1(e14,e11)=e13|op1(e14,e10)=e13), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_45]), c_0_126])).
cnf(c_0_154,plain,(op1(e14,e12)!=e13), inference(rw,[status(thm)],[c_0_127, c_0_128])).
cnf(c_0_155,plain,(op1(e13,e10)!=op1(e14,e10)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_156,plain,(op1(e13,e11)=e13|op1(e13,e10)=e13), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_35]), c_0_25]), c_0_87]), c_0_130]), c_0_131]), c_0_132])).
cnf(c_0_157,plain,(op2(e24,e24)=e23|op2(e23,e24)=e23|op2(e22,e24)=e23|op2(e21,e24)=e23|op2(e20,e24)=e23), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_158,plain,(e21!=e23), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_159,plain,(op2(e22,e22)!=op2(e22,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_160,plain,(op2(op2(e22,e24),op2(e24,e22))=e24), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_161,plain,(op2(e24,e22)=e21|op2(e24,e22)=e22), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_133, c_0_134]), c_0_135])).
cnf(c_0_162,plain,(op2(e24,e24)=e22|op2(e23,e24)=e22|op2(e22,e24)=e22|op2(e21,e24)=e22|op2(e20,e24)=e22), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_163,plain,(e21!=e22), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_164,plain,(op2(e23,e21)!=op2(e24,e21)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_165,plain,(op2(e23,e21)=e20|op2(e23,e21)=e23), inference(sr,[status(thm)],[c_0_136, c_0_137])).
cnf(c_0_166,plain,(op2(op2(e24,e23),e21)=e23), inference(rw,[status(thm)],[c_0_138, c_0_50])).
cnf(c_0_167,plain,(op2(e24,e23)=e24|epred3_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_141]), c_0_142])).
cnf(c_0_168,plain,(op2(e24,e23)=e20|op2(e23,e23)=e20|op2(e22,e23)=e20|op2(e21,e23)=e20|op2(e20,e23)=e20), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_169,plain,(op1(e14,e14)=e14|op1(e14,e13)=e14|op1(e14,e12)=e14|op1(e14,e11)=e14|op1(e14,e10)=e14), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_170,plain,(op1(e13,e12)!=op1(e14,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_171,plain,(op1(e10,e10)!=op1(e14,e10)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_172, plain, (((((((((((op1(e13,e10)!=e10|(op1(e12,e10)!=e10|epred1_0))|~epred2_0)&((op1(e13,e11)!=e11|(op1(e12,e10)!=e10|epred1_0))|~epred2_0))&((op1(e13,e12)!=e12|(op1(e12,e10)!=e10|epred1_0))|~epred2_0))&((op1(e13,e13)!=e13|(op1(e12,e10)!=e10|epred1_0))|~epred2_0))&((op1(e13,e14)!=e14|(op1(e12,e10)!=e10|epred1_0))|~epred2_0))&((((((op1(e13,e10)!=e10|(op1(e12,e11)!=e11|epred1_0))|~epred2_0)&((op1(e13,e11)!=e11|(op1(e12,e11)!=e11|epred1_0))|~epred2_0))&((op1(e13,e12)!=e12|(op1(e12,e11)!=e11|epred1_0))|~epred2_0))&((op1(e13,e13)!=e13|(op1(e12,e11)!=e11|epred1_0))|~epred2_0))&((op1(e13,e14)!=e14|(op1(e12,e11)!=e11|epred1_0))|~epred2_0)))&((((((op1(e13,e10)!=e10|(op1(e12,e12)!=e12|epred1_0))|~epred2_0)&((op1(e13,e11)!=e11|(op1(e12,e12)!=e12|epred1_0))|~epred2_0))&((op1(e13,e12)!=e12|(op1(e12,e12)!=e12|epred1_0))|~epred2_0))&((op1(e13,e13)!=e13|(op1(e12,e12)!=e12|epred1_0))|~epred2_0))&((op1(e13,e14)!=e14|(op1(e12,e12)!=e12|epred1_0))|~epred2_0)))&((((((op1(e13,e10)!=e10|(op1(e12,e13)!=e13|epred1_0))|~epred2_0)&((op1(e13,e11)!=e11|(op1(e12,e13)!=e13|epred1_0))|~epred2_0))&((op1(e13,e12)!=e12|(op1(e12,e13)!=e13|epred1_0))|~epred2_0))&((op1(e13,e13)!=e13|(op1(e12,e13)!=e13|epred1_0))|~epred2_0))&((op1(e13,e14)!=e14|(op1(e12,e13)!=e13|epred1_0))|~epred2_0)))&((((((op1(e13,e10)!=e10|(op1(e12,e14)!=e14|epred1_0))|~epred2_0)&((op1(e13,e11)!=e11|(op1(e12,e14)!=e14|epred1_0))|~epred2_0))&((op1(e13,e12)!=e12|(op1(e12,e14)!=e14|epred1_0))|~epred2_0))&((op1(e13,e13)!=e13|(op1(e12,e14)!=e14|epred1_0))|~epred2_0))&((op1(e13,e14)!=e14|(op1(e12,e14)!=e14|epred1_0))|~epred2_0)))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_143])])).
cnf(c_0_173,plain,(op1(op1(e12,e13),e14)=e13), inference(rw,[status(thm)],[c_0_144, c_0_35])).
cnf(c_0_174,plain,(op1(e12,e14)=e14|op1(e12,e13)=e14|op1(e14,e13)=e14), inference(spm,[status(thm)],[c_0_145, c_0_146])).
cnf(c_0_175,plain,(op1(op1(e11,e11),op1(e11,e11))=e11), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_176,plain,(op1(e14,e13)=e14|op1(e11,e11)=e14|epred2_0), inference(spm,[status(thm)],[c_0_147, c_0_148])).
cnf(c_0_177,plain,(e10!=e11), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_178,plain,(op1(e13,e10)!=op1(e13,e11)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_179,plain,(op1(e13,e10)=e10|op1(e13,e10)=e13), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_149, c_0_150]), c_0_151])).
cnf(c_0_180,plain,(op1(op1(e14,e13),e11)=e13), inference(rw,[status(thm)],[c_0_152, c_0_87])).
cnf(c_0_181,plain,(op1(e14,e10)=e13|op1(e14,e11)=e13|op1(e14,e13)=e13), inference(sr,[status(thm)],[c_0_153, c_0_154])).
cnf(c_0_182,plain,(op1(e13,e11)=e13|op1(e14,e10)!=e13), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_183,plain,(op2(e22,e24)=e23|op2(e21,e24)=e23|op2(e20,e24)=e23), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_157, c_0_50]), c_0_21]), c_0_158]), c_0_56])).
cnf(c_0_184,plain,(op2(e22,e24)!=e23), inference(rw,[status(thm)],[c_0_159, c_0_99])).
cnf(c_0_185,plain,(op2(op2(e22,e24),e22)=e24|op2(e24,e22)=e21), inference(spm,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_186,plain,(op2(e22,e24)=e22|op2(e21,e24)=e22|op2(e20,e24)=e22), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_162, c_0_50]), c_0_21]), c_0_163]), c_0_79])).
cnf(c_0_187,plain,(e23!=e24), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_188,plain,(op2(e23,e24)=e23|op2(e23,e23)=e23|op2(e23,e22)=e23|op2(e23,e21)=e23|op2(e23,e20)=e23), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_189,plain,(e22!=e23), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_190,plain,(op2(e23,e21)=e20|op2(e24,e21)!=e23), inference(spm,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_191,plain,(op2(e24,e21)=e23|epred3_0), inference(spm,[status(thm)],[c_0_166, c_0_167])).
cnf(c_0_192,plain,(op2(op2(e23,e22),op2(e22,e23))=e22), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_193,plain,(op2(e24,e23)=e20|op2(e22,e23)=e20|op2(e21,e23)=e20|op2(e20,e23)=e20), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_168, c_0_17]), c_0_79])).
cnf(c_0_194,plain,(op1(e12,e13)!=op1(e13,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_195,plain,(op1(e14,e13)=e14|op1(e14,e12)=e14|op1(e14,e11)=e14|op1(e14,e10)=e14), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_169, c_0_45]), c_0_89])).
cnf(c_0_196,plain,(op1(e14,e12)!=e14), inference(rw,[status(thm)],[c_0_170, c_0_35])).
cnf(c_0_197,plain,(op1(e14,e10)!=e14), inference(rw,[status(thm)],[c_0_171, c_0_63])).
cnf(c_0_198,plain,(epred1_0|~epred2_0|op1(e12,e14)!=e14|op1(e13,e10)!=e10), inference(split_conjunct,[status(thm)],[c_0_172])).
cnf(c_0_199,plain,(op1(e14,e13)=e14|op1(e12,e14)=e14), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_45]), c_0_126])).
cnf(c_0_200,plain,(op1(e14,e13)=e14|epred2_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176]), c_0_45]), c_0_177])).
cnf(c_0_201,plain,(op1(e13,e10)=e10|op1(e13,e11)!=e13), inference(spm,[status(thm)],[c_0_178, c_0_179])).
cnf(c_0_202,plain,(op1(e14,e11)=e13|op1(e13,e11)=e13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181]), c_0_182])).
cnf(c_0_203,plain,(op2(e24,e24)=e24|op2(e24,e23)=e24|op2(e24,e22)=e24|op2(e24,e21)=e24|op2(e24,e20)=e24), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_204,plain,(op2(e20,e20)!=op2(e24,e20)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_205,plain,(op2(e20,e24)=e23|op2(e21,e24)=e23), inference(sr,[status(thm)],[c_0_183, c_0_184])).
cnf(c_0_206,plain,(op2(e21,e24)=e24|op2(e24,e22)=e21), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_106]), c_0_97])).
cnf(c_0_207,plain,(op2(e20,e24)=e22|op2(e21,e24)=e22|op2(e24,e22)=e21), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_186]), c_0_99]), c_0_187])).
cnf(c_0_208,plain,(op2(e23,e21)=e23|op2(e23,e20)=e23), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_188, c_0_19]), c_0_17]), c_0_50]), c_0_187]), c_0_189]), c_0_158])).
cnf(c_0_209,plain,(op2(op2(e21,e23),op2(e23,e21))=e23), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_210,plain,(op2(e23,e21)=e20|epred3_0), inference(spm,[status(thm)],[c_0_190, c_0_191])).
cnf(c_0_211,plain,(op2(e24,op2(e22,e23))=e22), inference(rw,[status(thm)],[c_0_192, c_0_19])).
cnf(c_0_212,plain,(op2(e20,e23)=e20|op2(e21,e23)=e20|op2(e22,e23)=e20), inference(sr,[status(thm)],[c_0_193, c_0_43])).
cnf(c_0_213,plain,(op1(e12,e13)=e14|op1(e12,e13)=e13|op1(e12,e13)=e12|op1(e12,e13)=e11|op1(e12,e13)=e10), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_214,plain,(op1(e12,e13)!=e12), inference(rw,[status(thm)],[c_0_194, c_0_25])).
cnf(c_0_215,plain,(op1(e12,e12)!=op1(e12,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_216,plain,(op1(e14,e11)=e14|op1(e14,e13)=e14), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_195, c_0_196]), c_0_197])).
cnf(c_0_217,plain,(op1(e14,e13)=e14|epred1_0|op1(e13,e10)!=e10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199]), c_0_200])).
cnf(c_0_218,plain,(op1(e14,e11)=e13|op1(e13,e10)=e10), inference(spm,[status(thm)],[c_0_201, c_0_202])).
cnf(c_0_219,plain,(op2(e22,e23)!=op2(e23,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_220,plain,(op2(e24,e23)=e24|op2(e24,e21)=e24|op2(e24,e20)=e24), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_203, c_0_21]), c_0_97]), c_0_52])).
cnf(c_0_221,plain,(op2(e24,e20)!=e24), inference(rw,[status(thm)],[c_0_204, c_0_29])).
cnf(c_0_222,plain,(op2(e24,e24)=e22|op2(e24,e23)=e22|op2(e24,e22)=e22|op2(e24,e21)=e22|op2(e24,e20)=e22), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_223,plain,(op2(e24,e22)=e21|op2(e20,e24)=e23), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206]), c_0_187])).
cnf(c_0_224,plain,(op2(e20,e24)=e22|op2(e24,e22)=e21), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207]), c_0_27])).
cnf(c_0_225,plain,(op2(op2(e24,e22),op2(e22,e24))=e22), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_226,plain,(op2(e21,e24)=e24|epred3_0), inference(spm,[status(thm)],[c_0_82, c_0_167])).
cnf(c_0_227,plain,(op2(op2(e23,e20),op2(e20,e23))=e20), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_228,plain,(op2(e23,e20)=e23|op2(e23,e20)=e20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_208]), c_0_56])).
cnf(c_0_229,plain,(op2(op2(e21,e23),e20)=e23|epred3_0), inference(spm,[status(thm)],[c_0_209, c_0_210])).
cnf(c_0_230,plain,(op2(e21,e23)=e20|op2(e20,e23)=e20|op2(e24,e20)=e22), inference(spm,[status(thm)],[c_0_211, c_0_212])).
cnf(c_0_231,plain,(op2(e23,e20)!=op2(e23,e21)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_232,plain,(op1(e10,e13)!=op1(e13,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_233,plain,(op1(op1(e13,e12),op1(e12,e13))=e12), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_234,plain,(op1(e12,e13)=e14|op1(e12,e13)=e13|op1(e12,e13)=e11|op1(e12,e13)=e10), inference(sr,[status(thm)],[c_0_213, c_0_214])).
cnf(c_0_235,plain,(op1(e12,e13)!=e13), inference(rw,[status(thm)],[c_0_215, c_0_128])).
cnf(c_0_236,plain,(op1(op1(e11,e14),op1(e14,e11))=e14), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_237,plain,(op1(e14,e11)=e14|op1(e11,e14)=e14), inference(spm,[status(thm)],[c_0_147, c_0_216])).
cnf(c_0_238,plain,(op1(e14,e11)=e13|op1(e14,e13)=e14|epred1_0), inference(spm,[status(thm)],[c_0_217, c_0_218])).
cnf(c_0_239,plain,(op2(e22,e23)=e24|op2(e22,e23)=e23|op2(e22,e23)=e22|op2(e22,e23)=e21|op2(e22,e23)=e20), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_240,plain,(op2(e22,e23)!=e22), inference(rw,[status(thm)],[c_0_219, c_0_17])).
cnf(c_0_241,plain,(op2(e22,e22)!=op2(e22,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_242,plain,(op2(e24,e21)=e24|op2(e24,e23)=e24), inference(sr,[status(thm)],[c_0_220, c_0_221])).
cnf(c_0_243,plain,(op2(e24,e22)=e22|op2(e24,e21)=e22|op2(e24,e20)=e22), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_222, c_0_21]), c_0_32]), c_0_79])).
cnf(c_0_244,plain,(op2(e24,e22)=e21), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224]), c_0_189])).
cnf(c_0_245,plain,(op2(op2(e24,e22),e24)=e22|op2(e21,e24)=e24), inference(spm,[status(thm)],[c_0_225, c_0_106])).
cnf(c_0_246,plain,(op2(e22,e24)=e24|op2(e22,e23)=e24|op2(e22,e22)=e24|op2(e22,e21)=e24|op2(e22,e20)=e24), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_247,plain,(op2(e20,e20)!=op2(e22,e20)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_248,plain,(op2(e22,e24)=e21|op2(e22,e23)=e21|op2(e22,e22)=e21|op2(e22,e21)=e21|op2(e22,e20)=e21), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_249,plain,(op2(e22,e24)!=op2(e23,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_250,plain,(op2(op2(e24,e20),op2(e20,e24))=e20), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_251,plain,(op2(e20,e24)=e23|epred3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_226]), c_0_187])).
cnf(c_0_252,plain,(op2(e23,op2(e20,e23))=e20|op2(e23,e20)=e20), inference(spm,[status(thm)],[c_0_227, c_0_228])).
cnf(c_0_253,plain,(op2(e24,e20)=e22|op2(e20,e23)=e20|epred3_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230]), c_0_29]), c_0_187])).
cnf(c_0_254,plain,(epred3_0|op2(e23,e20)!=e20), inference(spm,[status(thm)],[c_0_231, c_0_210])).
cnf(c_0_255,plain,(op1(e10,e13)=e14|op1(e10,e13)=e13|op1(e10,e13)=e12|op1(e10,e13)=e11|op1(e10,e13)=e10), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_256,plain,(op1(e10,e13)!=e12), inference(rw,[status(thm)],[c_0_232, c_0_25])).
cnf(c_0_257,plain,(op1(e13,e14)=e10|op1(e13,e13)=e10|op1(e13,e12)=e10|op1(e13,e11)=e10|op1(e13,e10)=e10), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_258,plain,(e10!=e12), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_259,plain,(op1(e14,e14)=e13|op1(e13,e14)=e13|op1(e12,e14)=e13|op1(e11,e14)=e13|op1(e10,e14)=e13), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_260,plain,(op1(e12,e12)!=op1(e12,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_261,plain,(op1(e14,e12)=e10|op1(e13,e12)=e10|op1(e12,e12)=e10|op1(e11,e12)=e10|op1(e10,e12)=e10), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_262,plain,(op1(e14,e12)!=op1(e14,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_263,plain,(op1(e14,op1(e12,e13))=e12), inference(rw,[status(thm)],[c_0_233, c_0_35])).
cnf(c_0_264,plain,(op1(e12,e13)=e10|op1(e12,e13)=e11|op1(e12,e13)=e14), inference(sr,[status(thm)],[c_0_234, c_0_235])).
cnf(c_0_265,plain,(op1(e14,op1(e14,e11))=e14|op1(e14,e11)=e14), inference(spm,[status(thm)],[c_0_236, c_0_237])).
cnf(c_0_266,plain,(op1(e14,e11)=e13|epred1_0), inference(spm,[status(thm)],[c_0_180, c_0_238])).
cnf(c_0_267,plain,(op2(e22,e23)=e24|op2(e22,e23)=e23|op2(e22,e23)=e21|op2(e22,e23)=e20), inference(sr,[status(thm)],[c_0_239, c_0_240])).
cnf(c_0_268,plain,(op2(e22,e23)!=e23), inference(rw,[status(thm)],[c_0_241, c_0_99])).
cnf(c_0_269,plain,(op2(e24,e21)=e24|op2(e24,e21)=e23), inference(spm,[status(thm)],[c_0_166, c_0_242])).
cnf(c_0_270,plain,(op2(e24,e20)=e22|op2(e24,e21)=e22), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244]), c_0_163])).
cnf(c_0_271,plain,(op2(e21,e24)=e24|op2(e21,e24)=e22), inference(rw,[status(thm)],[c_0_245, c_0_244])).
cnf(c_0_272,plain,(op2(e20,e23)!=op2(e23,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_273,plain,(op2(e22,e24)=e24|op2(e22,e23)=e24|op2(e22,e21)=e24|op2(e22,e20)=e24), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_246, c_0_99]), c_0_187])).
cnf(c_0_274,plain,(op2(e22,e20)!=e24), inference(rw,[status(thm)],[c_0_247, c_0_29])).
cnf(c_0_275,plain,(op2(e22,e24)=e21|op2(e22,e23)=e21|op2(e22,e21)=e21|op2(e22,e20)=e21), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_248, c_0_99]), c_0_158])).
cnf(c_0_276,plain,(op2(e22,e24)!=e21), inference(rw,[status(thm)],[c_0_249, c_0_50])).
cnf(c_0_277,plain,(op2(op2(e24,e20),e23)=e20|epred3_0), inference(spm,[status(thm)],[c_0_250, c_0_251])).
cnf(c_0_278,plain,(op2(e24,e20)=e22|epred3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253]), c_0_254])).
cnf(c_0_279,plain,(op1(op1(e10,e13),op1(e13,e10))=e13), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_280,plain,(op1(e10,e13)=e14|op1(e10,e13)=e13|op1(e10,e13)=e11|op1(e10,e13)=e10), inference(sr,[status(thm)],[c_0_255, c_0_256])).
cnf(c_0_281,plain,(op1(e13,e11)=e10|op1(e13,e10)=e10), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_257, c_0_35]), c_0_25]), c_0_87]), c_0_89]), c_0_258]), c_0_177])).
cnf(c_0_282,plain,(op1(e14,e13)=e10|op1(e13,e13)=e10|op1(e12,e13)=e10|op1(e11,e13)=e10|op1(e10,e13)=e10), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_283,plain,(op1(e14,e14)=e12|op1(e13,e14)=e12|op1(e12,e14)=e12|op1(e11,e14)=e12|op1(e10,e14)=e12), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_284,plain,(e11!=e12), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_285,plain,(op1(e12,e14)=e13|op1(e11,e14)=e13|op1(e10,e14)=e13), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_259, c_0_87]), c_0_45]), c_0_132]), c_0_126])).
cnf(c_0_286,plain,(op1(e12,e14)!=e13), inference(rw,[status(thm)],[c_0_260, c_0_128])).
cnf(c_0_287,plain,(op1(e14,e12)=e11|op1(e13,e12)=e11|op1(e12,e12)=e11|op1(e11,e12)=e11|op1(e10,e12)=e11), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_288,plain,(op1(e14,e12)=e12|op1(e13,e12)=e12|op1(e12,e12)=e12|op1(e11,e12)=e12|op1(e10,e12)=e12), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_289,plain,(op1(e14,e12)=e10|op1(e11,e12)=e10|op1(e10,e12)=e10), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_261, c_0_128]), c_0_35]), c_0_126]), c_0_89])).
cnf(c_0_290,plain,(op1(e14,e12)!=e10), inference(rw,[status(thm)],[c_0_262, c_0_45])).
cnf(c_0_291,plain,(op1(e12,e13)=e14|op1(e12,e13)=e11|op1(e14,e10)=e12), inference(spm,[status(thm)],[c_0_263, c_0_264])).
cnf(c_0_292,plain,(op1(e14,e13)=e14|epred1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_266]), c_0_130])).
cnf(c_0_293,plain,(op2(e22,e23)=e20|op2(e22,e23)=e21|op2(e22,e23)=e24), inference(sr,[status(thm)],[c_0_267, c_0_268])).
cnf(c_0_294,plain,(op2(e24,e20)=e22), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_270]), c_0_27]), c_0_189])).
cnf(c_0_295,plain,(op2(e21,e24)=e24|op2(e20,e24)=e23), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_271]), c_0_189])).
cnf(c_0_296,plain,(op2(e20,e23)=e24|op2(e20,e23)=e23|op2(e20,e23)=e22|op2(e20,e23)=e21|op2(e20,e23)=e20), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_297,plain,(op2(e20,e23)!=e22), inference(rw,[status(thm)],[c_0_272, c_0_17])).
cnf(c_0_298,plain,(op2(e24,e20)!=op2(e24,e22)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_299,plain,(op2(e22,e24)=e22|op2(e22,e23)=e22|op2(e22,e22)=e22|op2(e22,e21)=e22|op2(e22,e20)=e22), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_300,plain,(op2(e22,e21)=e24|op2(e22,e23)=e24|op2(e22,e24)=e24), inference(sr,[status(thm)],[c_0_273, c_0_274])).
cnf(c_0_301,plain,(op2(e22,e20)=e21|op2(e22,e21)=e21|op2(e22,e23)=e21), inference(sr,[status(thm)],[c_0_275, c_0_276])).
cnf(c_0_302,plain,(op2(e22,e23)=e20|epred3_0), inference(spm,[status(thm)],[c_0_277, c_0_278])).
cnf(c_0_303,plain,(op1(op1(e13,e10),op1(e10,e13))=e10), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_304,plain,(op1(op1(e10,e13),e13)=e13|op1(e13,e11)=e13), inference(spm,[status(thm)],[c_0_279, c_0_156])).
cnf(c_0_305,plain,(op1(e10,e13)=e10|op1(e10,e13)=e11|op1(e10,e13)=e13), inference(sr,[status(thm)],[c_0_280, c_0_85])).
cnf(c_0_306,plain,(op1(e13,e11)!=op1(e14,e11)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_307,plain,(op1(e13,e11)=e13|op1(e13,e11)=e10), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_156]), c_0_126])).
cnf(c_0_308,plain,(op1(e14,e13)=e10|op1(e12,e13)=e10|op1(e11,e13)=e10|op1(e10,e13)=e10), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_282, c_0_25]), c_0_258])).
cnf(c_0_309,plain,(op1(e12,e14)=e12|op1(e11,e14)=e12|op1(e10,e14)=e12), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_283, c_0_87]), c_0_45]), c_0_284]), c_0_258])).
cnf(c_0_310,plain,(op1(e10,e14)=e13|op1(e11,e14)=e13), inference(sr,[status(thm)],[c_0_285, c_0_286])).
cnf(c_0_311,plain,(op1(e14,e12)=e11|op1(e11,e12)=e11|op1(e10,e12)=e11), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_287, c_0_128]), c_0_35]), c_0_132]), c_0_88])).
cnf(c_0_312,plain,(op1(e14,e12)=e12|op1(e11,e12)=e12|op1(e10,e12)=e12), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_288, c_0_128]), c_0_35]), c_0_131]), c_0_61])).
cnf(c_0_313,plain,(op1(e10,e12)=e10|op1(e11,e12)=e10), inference(sr,[status(thm)],[c_0_289, c_0_290])).
cnf(c_0_314,plain,(op1(e14,e10)=e12|op1(e12,e13)=e14|op1(e11,e14)=e13), inference(spm,[status(thm)],[c_0_173, c_0_291])).
cnf(c_0_315,plain,(op1(e11,e14)=e14|epred1_0), inference(spm,[status(thm)],[c_0_147, c_0_292])).
cnf(c_0_316,plain,(op2(e22,e20)!=op2(e22,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_317,plain,(op2(e22,e23)=e24|op2(e22,e23)=e20|op2(e21,e24)=e23), inference(spm,[status(thm)],[c_0_54, c_0_293])).
cnf(c_0_318,plain,(op2(e22,op2(e20,e24))=e20), inference(rw,[status(thm)],[c_0_250, c_0_294])).
cnf(c_0_319,plain,(op2(e20,e24)=e23), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_295]), c_0_187])).
cnf(c_0_320,plain,(op2(e20,e24)!=op2(e21,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_321,plain,(op2(e20,e23)=e24|op2(e20,e23)=e23|op2(e20,e23)=e21|op2(e20,e23)=e20), inference(sr,[status(thm)],[c_0_296, c_0_297])).
cnf(c_0_322,plain,(op2(e20,e23)!=op2(e22,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_323,plain,(op2(e24,e22)=e21|op2(e24,e20)!=e22), inference(spm,[status(thm)],[c_0_298, c_0_161])).
cnf(c_0_324,plain,(op2(e22,e24)=e22|op2(e22,e21)=e22|op2(e22,e20)=e22), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_299, c_0_99]), c_0_189]), c_0_240])).
cnf(c_0_325,plain,(op2(e22,e24)=e24|op2(e22,e21)=e24), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_300]), c_0_21]), c_0_56])).
cnf(c_0_326,plain,(op2(e22,e21)=e21|op2(e22,e20)=e21|epred3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_302]), c_0_80])).
cnf(c_0_327,plain,(op1(e10,op1(e10,e13))=e10|op1(e13,e11)=e10), inference(spm,[status(thm)],[c_0_303, c_0_281])).
cnf(c_0_328,plain,(op1(e10,e13)=e11|op1(e10,e13)=e10|op1(e13,e11)=e13), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_305]), c_0_25]), c_0_131])).
cnf(c_0_329,plain,(op1(e13,e11)=e10|op1(e14,e11)!=e13), inference(spm,[status(thm)],[c_0_306, c_0_307])).
cnf(c_0_330,plain,(op1(e10,e12)!=op1(e10,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_331,plain,(op1(e10,e13)=e10|op1(e11,e13)=e10|op1(e12,e13)=e10), inference(sr,[status(thm)],[c_0_308, c_0_93])).
cnf(c_0_332,plain,(op1(e11,e14)=e13|op1(e11,e14)=e12|op1(e12,e14)=e12), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_310]), c_0_131])).
cnf(c_0_333,plain,(op1(e11,e12)=e12|op1(e14,e12)=e12|op1(e11,e12)=e11|op1(e14,e12)=e11), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_312]), c_0_284])).
cnf(c_0_334,plain,(op1(e11,e12)=e10|op1(e11,e12)=e11|op1(e14,e12)=e11), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_313]), c_0_177])).
cnf(c_0_335,plain,(op1(e12,e13)=e14|op1(e14,e10)=e12|epred1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_315]), c_0_130])).
cnf(c_0_336,plain,(op2(e22,e20)!=op2(e22,e22)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_337,plain,(op2(e22,e20)!=op2(e24,e20)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_338,plain,(op2(e21,e24)=e23|op2(e22,e23)=e24|op2(e22,e20)!=e20), inference(spm,[status(thm)],[c_0_316, c_0_317])).
cnf(c_0_339,plain,(op2(e22,e23)=e20), inference(rw,[status(thm)],[c_0_318, c_0_319])).
cnf(c_0_340,plain,(op2(e21,e24)!=e23), inference(rw,[status(thm)],[c_0_320, c_0_319])).
cnf(c_0_341,plain,(op2(e20,e23)!=op2(e20,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_342,plain,(op2(e20,e23)=e20|op2(e20,e23)=e21|op2(e20,e23)=e23), inference(sr,[status(thm)],[c_0_321, c_0_40])).
cnf(c_0_343,plain,(op2(e21,e24)=e23|op2(e22,e23)=e24|op2(e20,e23)!=e20), inference(spm,[status(thm)],[c_0_322, c_0_317])).
cnf(c_0_344,plain,(op2(e24,e22)=e21|epred3_0), inference(spm,[status(thm)],[c_0_323, c_0_278])).
cnf(c_0_345,plain,(op2(e22,e24)=e24|op2(e22,e20)=e22|op2(e22,e24)=e22), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_325]), c_0_27])).
cnf(c_0_346,plain,(op2(e22,e20)=e21|op2(e22,e24)=e24|epred3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_51])).
cnf(c_0_347,plain,(op1(e10,op1(e10,e13))=e10|op1(e14,e11)=e13), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_202]), c_0_126])).
cnf(c_0_348,plain,(op1(e10,e13)=e10|op1(e10,e13)=e11|op1(e14,e11)!=e13), inference(spm,[status(thm)],[c_0_306, c_0_328])).
cnf(c_0_349,plain,(op1(e13,op1(e10,e13))=e10|op1(e13,e10)=e10), inference(spm,[status(thm)],[c_0_303, c_0_179])).
cnf(c_0_350,plain,(op1(op1(e10,e13),e10)=e13|op1(e13,e11)=e10), inference(spm,[status(thm)],[c_0_279, c_0_281])).
cnf(c_0_351,plain,(op1(op1(e11,e13),op1(e13,e11))=e13), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_352,plain,(op1(e13,e11)=e10|epred1_0), inference(spm,[status(thm)],[c_0_329, c_0_266])).
cnf(c_0_353,plain,(op1(e12,e13)=e10|op1(e11,e13)=e10|op1(e10,e12)!=e10), inference(spm,[status(thm)],[c_0_330, c_0_331])).
cnf(c_0_354,plain,(op1(op1(e14,e12),op1(e12,e14))=e12), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_355,plain,(op1(e12,e14)=e12|epred1_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_315]), c_0_130]), c_0_61])).
cnf(c_0_356,plain,(op1(e11,e12)=e10|op1(e11,e12)=e12|op1(e14,e12)=e12), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_313]), c_0_258])).
cnf(c_0_357,plain,(op1(e11,e12)=e11|op1(e14,e12)=e11|op1(e14,e12)=e12), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_334]), c_0_258]), c_0_177])).
cnf(c_0_358,plain,(op1(e14,e10)!=op1(e14,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_359,plain,(op1(e14,e10)=e12|epred1_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_335]), c_0_45]), c_0_126])).
cnf(c_0_360,plain,(op2(e24,e22)=e22|op2(e23,e22)=e22|op2(e22,e22)=e22|op2(e21,e22)=e22|op2(e20,e22)=e22), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_361,plain,(op2(e22,e20)=e24|op2(e22,e20)=e23|op2(e22,e20)=e22|op2(e22,e20)=e21|op2(e22,e20)=e20), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_362,plain,(op2(e22,e20)!=e23), inference(rw,[status(thm)],[c_0_336, c_0_99])).
cnf(c_0_363,plain,(op2(e22,e20)!=e22), inference(rw,[status(thm)],[c_0_337, c_0_294])).
cnf(c_0_364,plain,(op2(e22,e20)!=e20), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_338, c_0_339]), c_0_340]), c_0_52])).
cnf(c_0_365,plain,(op2(e20,e23)=e21|op2(e20,e23)=e20|op2(e20,e24)!=e23), inference(spm,[status(thm)],[c_0_341, c_0_342])).
cnf(c_0_366,plain,(op2(e20,e23)!=e20), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_343, c_0_339]), c_0_340]), c_0_52])).
cnf(c_0_367,plain,(op2(e21,op2(e22,e24))=e22|epred3_0), inference(spm,[status(thm)],[c_0_225, c_0_344])).
cnf(c_0_368,plain,(op2(e22,e24)=e22|op2(e22,e24)=e24|epred3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_346]), c_0_163])).
cnf(c_0_369,plain,(op2(e24,e21)!=op2(e24,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_370,plain,(op2(e24,e22)=e20|op2(e23,e22)=e20|op2(e22,e22)=e20|op2(e21,e22)=e20|op2(e20,e22)=e20), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_371,plain,(op2(e24,e24)=e23|op2(e24,e23)=e23|op2(e24,e22)=e23|op2(e24,e21)=e23|op2(e24,e20)=e23), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_372, plain, ((epred6_0<=>((((((((((((((((((((((((((((h1(op1(e10,e10))=op2(h1(e10),h1(e10))&h1(op1(e10,e11))=op2(h1(e10),h1(e11)))&h1(op1(e10,e12))=op2(h1(e10),h1(e12)))&h1(op1(e10,e13))=op2(h1(e10),h1(e13)))&h1(op1(e10,e14))=op2(h1(e10),h1(e14)))&h1(op1(e11,e10))=op2(h1(e11),h1(e10)))&h1(op1(e11,e11))=op2(h1(e11),h1(e11)))&h1(op1(e11,e12))=op2(h1(e11),h1(e12)))&h1(op1(e11,e13))=op2(h1(e11),h1(e13)))&h1(op1(e11,e14))=op2(h1(e11),h1(e14)))&h1(op1(e12,e10))=op2(h1(e12),h1(e10)))&h1(op1(e12,e11))=op2(h1(e12),h1(e11)))&h1(op1(e12,e12))=op2(h1(e12),h1(e12)))&h1(op1(e12,e13))=op2(h1(e12),h1(e13)))&h1(op1(e12,e14))=op2(h1(e12),h1(e14)))&h1(op1(e13,e10))=op2(h1(e13),h1(e10)))&h1(op1(e13,e11))=op2(h1(e13),h1(e11)))&h1(op1(e13,e12))=op2(h1(e13),h1(e12)))&h1(op1(e13,e13))=op2(h1(e13),h1(e13)))&h1(op1(e13,e14))=op2(h1(e13),h1(e14)))&h1(op1(e14,e10))=op2(h1(e14),h1(e10)))&h1(op1(e14,e11))=op2(h1(e14),h1(e11)))&h1(op1(e14,e12))=op2(h1(e14),h1(e12)))&h1(op1(e14,e13))=op2(h1(e14),h1(e13)))&h1(op1(e14,e14))=op2(h1(e14),h1(e14)))&((((h1(e10)=e20|h1(e11)=e20)|h1(e12)=e20)|h1(e13)=e20)|h1(e14)=e20))&((((h1(e10)=e21|h1(e11)=e21)|h1(e12)=e21)|h1(e13)=e21)|h1(e14)=e21))&((((h1(e10)=e22|h1(e11)=e22)|h1(e12)=e22)|h1(e13)=e22)|h1(e14)=e22))&((((h1(e10)=e23|h1(e11)=e23)|h1(e12)=e23)|h1(e13)=e23)|h1(e14)=e23)))), introduced(definition)).
fof(c_0_373, plain, ((epred8_0<=>((((((((((((((((((((((((((((h2(op1(e10,e10))=op2(h2(e10),h2(e10))&h2(op1(e10,e11))=op2(h2(e10),h2(e11)))&h2(op1(e10,e12))=op2(h2(e10),h2(e12)))&h2(op1(e10,e13))=op2(h2(e10),h2(e13)))&h2(op1(e10,e14))=op2(h2(e10),h2(e14)))&h2(op1(e11,e10))=op2(h2(e11),h2(e10)))&h2(op1(e11,e11))=op2(h2(e11),h2(e11)))&h2(op1(e11,e12))=op2(h2(e11),h2(e12)))&h2(op1(e11,e13))=op2(h2(e11),h2(e13)))&h2(op1(e11,e14))=op2(h2(e11),h2(e14)))&h2(op1(e12,e10))=op2(h2(e12),h2(e10)))&h2(op1(e12,e11))=op2(h2(e12),h2(e11)))&h2(op1(e12,e12))=op2(h2(e12),h2(e12)))&h2(op1(e12,e13))=op2(h2(e12),h2(e13)))&h2(op1(e12,e14))=op2(h2(e12),h2(e14)))&h2(op1(e13,e10))=op2(h2(e13),h2(e10)))&h2(op1(e13,e11))=op2(h2(e13),h2(e11)))&h2(op1(e13,e12))=op2(h2(e13),h2(e12)))&h2(op1(e13,e13))=op2(h2(e13),h2(e13)))&h2(op1(e13,e14))=op2(h2(e13),h2(e14)))&h2(op1(e14,e10))=op2(h2(e14),h2(e10)))&h2(op1(e14,e11))=op2(h2(e14),h2(e11)))&h2(op1(e14,e12))=op2(h2(e14),h2(e12)))&h2(op1(e14,e13))=op2(h2(e14),h2(e13)))&h2(op1(e14,e14))=op2(h2(e14),h2(e14)))&((((h2(e10)=e20|h2(e11)=e20)|h2(e12)=e20)|h2(e13)=e20)|h2(e14)=e20))&((((h2(e10)=e21|h2(e11)=e21)|h2(e12)=e21)|h2(e13)=e21)|h2(e14)=e21))&((((h2(e10)=e22|h2(e11)=e22)|h2(e12)=e22)|h2(e13)=e22)|h2(e14)=e22))&((((h2(e10)=e23|h2(e11)=e23)|h2(e12)=e23)|h2(e13)=e23)|h2(e14)=e23)))), introduced(definition)).
fof(c_0_374, plain, ((epred10_0<=>((((((((((((((((((((((((((((h3(op1(e10,e10))=op2(h3(e10),h3(e10))&h3(op1(e10,e11))=op2(h3(e10),h3(e11)))&h3(op1(e10,e12))=op2(h3(e10),h3(e12)))&h3(op1(e10,e13))=op2(h3(e10),h3(e13)))&h3(op1(e10,e14))=op2(h3(e10),h3(e14)))&h3(op1(e11,e10))=op2(h3(e11),h3(e10)))&h3(op1(e11,e11))=op2(h3(e11),h3(e11)))&h3(op1(e11,e12))=op2(h3(e11),h3(e12)))&h3(op1(e11,e13))=op2(h3(e11),h3(e13)))&h3(op1(e11,e14))=op2(h3(e11),h3(e14)))&h3(op1(e12,e10))=op2(h3(e12),h3(e10)))&h3(op1(e12,e11))=op2(h3(e12),h3(e11)))&h3(op1(e12,e12))=op2(h3(e12),h3(e12)))&h3(op1(e12,e13))=op2(h3(e12),h3(e13)))&h3(op1(e12,e14))=op2(h3(e12),h3(e14)))&h3(op1(e13,e10))=op2(h3(e13),h3(e10)))&h3(op1(e13,e11))=op2(h3(e13),h3(e11)))&h3(op1(e13,e12))=op2(h3(e13),h3(e12)))&h3(op1(e13,e13))=op2(h3(e13),h3(e13)))&h3(op1(e13,e14))=op2(h3(e13),h3(e14)))&h3(op1(e14,e10))=op2(h3(e14),h3(e10)))&h3(op1(e14,e11))=op2(h3(e14),h3(e11)))&h3(op1(e14,e12))=op2(h3(e14),h3(e12)))&h3(op1(e14,e13))=op2(h3(e14),h3(e13)))&h3(op1(e14,e14))=op2(h3(e14),h3(e14)))&((((h3(e10)=e20|h3(e11)=e20)|h3(e12)=e20)|h3(e13)=e20)|h3(e14)=e20))&((((h3(e10)=e21|h3(e11)=e21)|h3(e12)=e21)|h3(e13)=e21)|h3(e14)=e21))&((((h3(e10)=e22|h3(e11)=e22)|h3(e12)=e22)|h3(e13)=e22)|h3(e14)=e22))&((((h3(e10)=e23|h3(e11)=e23)|h3(e12)=e23)|h3(e13)=e23)|h3(e14)=e23)))), introduced(definition)).
fof(c_0_375, plain, ((epred12_0<=>((((((((((((((((((((((((((((h4(op1(e10,e10))=op2(h4(e10),h4(e10))&h4(op1(e10,e11))=op2(h4(e10),h4(e11)))&h4(op1(e10,e12))=op2(h4(e10),h4(e12)))&h4(op1(e10,e13))=op2(h4(e10),h4(e13)))&h4(op1(e10,e14))=op2(h4(e10),h4(e14)))&h4(op1(e11,e10))=op2(h4(e11),h4(e10)))&h4(op1(e11,e11))=op2(h4(e11),h4(e11)))&h4(op1(e11,e12))=op2(h4(e11),h4(e12)))&h4(op1(e11,e13))=op2(h4(e11),h4(e13)))&h4(op1(e11,e14))=op2(h4(e11),h4(e14)))&h4(op1(e12,e10))=op2(h4(e12),h4(e10)))&h4(op1(e12,e11))=op2(h4(e12),h4(e11)))&h4(op1(e12,e12))=op2(h4(e12),h4(e12)))&h4(op1(e12,e13))=op2(h4(e12),h4(e13)))&h4(op1(e12,e14))=op2(h4(e12),h4(e14)))&h4(op1(e13,e10))=op2(h4(e13),h4(e10)))&h4(op1(e13,e11))=op2(h4(e13),h4(e11)))&h4(op1(e13,e12))=op2(h4(e13),h4(e12)))&h4(op1(e13,e13))=op2(h4(e13),h4(e13)))&h4(op1(e13,e14))=op2(h4(e13),h4(e14)))&h4(op1(e14,e10))=op2(h4(e14),h4(e10)))&h4(op1(e14,e11))=op2(h4(e14),h4(e11)))&h4(op1(e14,e12))=op2(h4(e14),h4(e12)))&h4(op1(e14,e13))=op2(h4(e14),h4(e13)))&h4(op1(e14,e14))=op2(h4(e14),h4(e14)))&((((h4(e10)=e20|h4(e11)=e20)|h4(e12)=e20)|h4(e13)=e20)|h4(e14)=e20))&((((h4(e10)=e21|h4(e11)=e21)|h4(e12)=e21)|h4(e13)=e21)|h4(e14)=e21))&((((h4(e10)=e22|h4(e11)=e22)|h4(e12)=e22)|h4(e13)=e22)|h4(e14)=e22))&((((h4(e10)=e23|h4(e11)=e23)|h4(e12)=e23)|h4(e13)=e23)|h4(e14)=e23)))), introduced(definition)).
fof(c_0_376, plain, ((epred14_0<=>((((((((((((((((((((((((((((h5(op1(e10,e10))=op2(h5(e10),h5(e10))&h5(op1(e10,e11))=op2(h5(e10),h5(e11)))&h5(op1(e10,e12))=op2(h5(e10),h5(e12)))&h5(op1(e10,e13))=op2(h5(e10),h5(e13)))&h5(op1(e10,e14))=op2(h5(e10),h5(e14)))&h5(op1(e11,e10))=op2(h5(e11),h5(e10)))&h5(op1(e11,e11))=op2(h5(e11),h5(e11)))&h5(op1(e11,e12))=op2(h5(e11),h5(e12)))&h5(op1(e11,e13))=op2(h5(e11),h5(e13)))&h5(op1(e11,e14))=op2(h5(e11),h5(e14)))&h5(op1(e12,e10))=op2(h5(e12),h5(e10)))&h5(op1(e12,e11))=op2(h5(e12),h5(e11)))&h5(op1(e12,e12))=op2(h5(e12),h5(e12)))&h5(op1(e12,e13))=op2(h5(e12),h5(e13)))&h5(op1(e12,e14))=op2(h5(e12),h5(e14)))&h5(op1(e13,e10))=op2(h5(e13),h5(e10)))&h5(op1(e13,e11))=op2(h5(e13),h5(e11)))&h5(op1(e13,e12))=op2(h5(e13),h5(e12)))&h5(op1(e13,e13))=op2(h5(e13),h5(e13)))&h5(op1(e13,e14))=op2(h5(e13),h5(e14)))&h5(op1(e14,e10))=op2(h5(e14),h5(e10)))&h5(op1(e14,e11))=op2(h5(e14),h5(e11)))&h5(op1(e14,e12))=op2(h5(e14),h5(e12)))&h5(op1(e14,e13))=op2(h5(e14),h5(e13)))&h5(op1(e14,e14))=op2(h5(e14),h5(e14)))&((((h5(e10)=e20|h5(e11)=e20)|h5(e12)=e20)|h5(e13)=e20)|h5(e14)=e20))&((((h5(e10)=e21|h5(e11)=e21)|h5(e12)=e21)|h5(e13)=e21)|h5(e14)=e21))&((((h5(e10)=e22|h5(e11)=e22)|h5(e12)=e22)|h5(e13)=e22)|h5(e14)=e22))&((((h5(e10)=e23|h5(e11)=e23)|h5(e12)=e23)|h5(e13)=e23)|h5(e14)=e23)))), introduced(definition)).
cnf(c_0_377,plain,(op1(e10,e13)!=op1(e11,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_378,plain,(op1(e10,e13)=e11|op1(e10,e13)=e10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_305]), c_0_348])).
cnf(c_0_379,plain,(op1(e12,e13)=e10|op1(e11,e13)=e10|op1(e13,e10)=e10), inference(spm,[status(thm)],[c_0_349, c_0_331])).
cnf(c_0_380,plain,(op1(e12,e13)=e10|op1(e11,e13)=e10|op1(e13,e11)=e10), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_331]), c_0_63]), c_0_130])).
cnf(c_0_381,plain,(op1(e10,e13)!=op1(e12,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_382,plain,(op1(op1(e11,e13),e10)=e13|epred1_0), inference(spm,[status(thm)],[c_0_351, c_0_352])).
cnf(c_0_383,plain,(op1(e11,e12)=e10|op1(e11,e13)=e10|op1(e12,e13)=e10), inference(spm,[status(thm)],[c_0_353, c_0_313])).
cnf(c_0_384,plain,(op1(op1(e14,e12),e12)=e12|epred1_0), inference(spm,[status(thm)],[c_0_354, c_0_355])).
cnf(c_0_385,plain,(op1(e14,e12)=e11|op1(e14,e12)=e12), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_357]), c_0_177]), c_0_284])).
cnf(c_0_386,plain,(epred1_0|op1(e14,e12)!=e12), inference(spm,[status(thm)],[c_0_358, c_0_359])).
cnf(c_0_387,plain,(op2(e21,e22)!=op2(e21,e24)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_388,plain,(op2(e24,e22)=e22|op2(e21,e22)=e22|op2(e20,e22)=e22), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_360, c_0_99]), c_0_19]), c_0_189]), c_0_27])).
fof(c_0_389, plain, ((epred3_0=>(((((op2(e20,e20)!=e20&op2(e20,e21)!=e21)&op2(e20,e22)!=e22)&op2(e20,e23)!=e23)&op2(e20,e24)!=e24)|((((op2(e21,e20)!=e20&op2(e21,e21)!=e21)&op2(e21,e22)!=e22)&op2(e21,e23)!=e23)&op2(e21,e24)!=e24)))), inference(split_equiv,[status(thm)],[c_0_36])).
cnf(c_0_390,plain,(op2(e24,e21)!=op2(e24,e22)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_391,plain,(op2(e22,e20)!=op2(e22,e21)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_392,plain,(op2(e22,e20)=e21), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_361, c_0_362]), c_0_274]), c_0_363]), c_0_364])).
cnf(c_0_393,plain,(op2(e20,e21)!=op2(e20,e23)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_394,plain,(op2(e20,e23)=e21), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_365, c_0_319])]), c_0_366])).
cnf(c_0_395,plain,(op2(op2(e22,e24),e21)=e24|epred3_0), inference(spm,[status(thm)],[c_0_160, c_0_344])).
cnf(c_0_396,plain,(op2(e22,e24)=e24|op2(e21,e22)=e22|epred3_0), inference(spm,[status(thm)],[c_0_367, c_0_368])).
cnf(c_0_397,plain,(epred3_0|op2(e24,e21)!=e24), inference(spm,[status(thm)],[c_0_369, c_0_167])).
cnf(c_0_398,plain,(op2(e24,e22)=e20|op2(e21,e22)=e20|op2(e20,e22)=e20), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_370, c_0_99]), c_0_19]), c_0_56]), c_0_52])).
cnf(c_0_399,plain,(op2(e24,e23)=e23|op2(e24,e22)=e23|op2(e24,e21)=e23|op2(e24,e20)=e23), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_371, c_0_21]), c_0_56])).
fof(c_0_400, plain, ((epred13_0<=>((((((((((((((((((((((((((h5(op1(e10,e10))=op2(h5(e10),h5(e10))&h5(op1(e10,e11))=op2(h5(e10),h5(e11)))&h5(op1(e10,e12))=op2(h5(e10),h5(e12)))&h5(op1(e10,e13))=op2(h5(e10),h5(e13)))&h5(op1(e10,e14))=op2(h5(e10),h5(e14)))&h5(op1(e11,e10))=op2(h5(e11),h5(e10)))&h5(op1(e11,e11))=op2(h5(e11),h5(e11)))&h5(op1(e11,e12))=op2(h5(e11),h5(e12)))&h5(op1(e11,e13))=op2(h5(e11),h5(e13)))&h5(op1(e11,e14))=op2(h5(e11),h5(e14)))&h5(op1(e12,e10))=op2(h5(e12),h5(e10)))&h5(op1(e12,e11))=op2(h5(e12),h5(e11)))&h5(op1(e12,e12))=op2(h5(e12),h5(e12)))&h5(op1(e12,e13))=op2(h5(e12),h5(e13)))&h5(op1(e12,e14))=op2(h5(e12),h5(e14)))&h5(op1(e13,e10))=op2(h5(e13),h5(e10)))&h5(op1(e13,e11))=op2(h5(e13),h5(e11)))&h5(op1(e13,e12))=op2(h5(e13),h5(e12)))&h5(op1(e13,e13))=op2(h5(e13),h5(e13)))&h5(op1(e13,e14))=op2(h5(e13),h5(e14)))&h5(op1(e14,e10))=op2(h5(e14),h5(e10)))&h5(op1(e14,e11))=op2(h5(e14),h5(e11)))&h5(op1(e14,e12))=op2(h5(e14),h5(e12)))&h5(op1(e14,e13))=op2(h5(e14),h5(e13)))&h5(op1(e14,e14))=op2(h5(e14),h5(e14)))&((((h5(e10)=e20|h5(e11)=e20)|h5(e12)=e20)|h5(e13)=e20)|h5(e14)=e20))&((((h5(e10)=e21|h5(e11)=e21)|h5(e12)=e21)|h5(e13)=e21)|h5(e14)=e21)))), introduced(definition)).
fof(c_0_401, negated_conjecture, (~((((((epred6_0&((((h1(e10)=e24|h1(e11)=e24)|h1(e12)=e24)|h1(e13)=e24)|h1(e14)=e24))|(epred8_0&((((h2(e10)=e24|h2(e11)=e24)|h2(e12)=e24)|h2(e13)=e24)|h2(e14)=e24)))|(epred10_0&((((h3(e10)=e24|h3(e11)=e24)|h3(e12)=e24)|h3(e13)=e24)|h3(e14)=e24)))|(epred12_0&((((h4(e10)=e24|h4(e11)=e24)|h4(e12)=e24)|h4(e13)=e24)|h4(e14)=e24)))|(epred14_0&((((h5(e10)=e24|h5(e11)=e24)|h5(e12)=e24)|h5(e13)=e24)|h5(e14)=e24))))), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_14]), c_0_372]), c_0_373]), c_0_374]), c_0_375]), c_0_376])).
cnf(c_0_402,plain,(op1(e10,e13)=e11|op1(e11,e13)!=e10), inference(spm,[status(thm)],[c_0_377, c_0_378])).
cnf(c_0_403,plain,(op1(e11,e13)=e10|op1(e12,e13)=e10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_379]), c_0_380])).
cnf(c_0_404,plain,(op1(e10,e13)=e11|op1(e12,e13)!=e10), inference(spm,[status(thm)],[c_0_381, c_0_378])).
cnf(c_0_405,plain,(op1(op1(e13,e11),op1(e11,e13))=e11), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_406,plain,(op1(e10,e13)=e10|op1(e10,e13)=e11|epred1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_352]), c_0_126])).
cnf(c_0_407,plain,(op1(e12,e13)=e10|op1(e11,e12)=e10|epred1_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_383]), c_0_63]), c_0_130])).
cnf(c_0_408,plain,(op1(e11,e12)=e12|epred1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_385]), c_0_386])).
cnf(c_0_409,plain,(op2(e20,e22)=e22|op2(e24,e22)=e22|op2(e21,e24)!=e22), inference(spm,[status(thm)],[c_0_387, c_0_388])).
fof(c_0_410, plain, (((((((((((op2(e21,e20)!=e20|op2(e20,e20)!=e20)|~epred3_0)&((op2(e21,e21)!=e21|op2(e20,e20)!=e20)|~epred3_0))&((op2(e21,e22)!=e22|op2(e20,e20)!=e20)|~epred3_0))&((op2(e21,e23)!=e23|op2(e20,e20)!=e20)|~epred3_0))&((op2(e21,e24)!=e24|op2(e20,e20)!=e20)|~epred3_0))&((((((op2(e21,e20)!=e20|op2(e20,e21)!=e21)|~epred3_0)&((op2(e21,e21)!=e21|op2(e20,e21)!=e21)|~epred3_0))&((op2(e21,e22)!=e22|op2(e20,e21)!=e21)|~epred3_0))&((op2(e21,e23)!=e23|op2(e20,e21)!=e21)|~epred3_0))&((op2(e21,e24)!=e24|op2(e20,e21)!=e21)|~epred3_0)))&((((((op2(e21,e20)!=e20|op2(e20,e22)!=e22)|~epred3_0)&((op2(e21,e21)!=e21|op2(e20,e22)!=e22)|~epred3_0))&((op2(e21,e22)!=e22|op2(e20,e22)!=e22)|~epred3_0))&((op2(e21,e23)!=e23|op2(e20,e22)!=e22)|~epred3_0))&((op2(e21,e24)!=e24|op2(e20,e22)!=e22)|~epred3_0)))&((((((op2(e21,e20)!=e20|op2(e20,e23)!=e23)|~epred3_0)&((op2(e21,e21)!=e21|op2(e20,e23)!=e23)|~epred3_0))&((op2(e21,e22)!=e22|op2(e20,e23)!=e23)|~epred3_0))&((op2(e21,e23)!=e23|op2(e20,e23)!=e23)|~epred3_0))&((op2(e21,e24)!=e24|op2(e20,e23)!=e23)|~epred3_0)))&((((((op2(e21,e20)!=e20|op2(e20,e24)!=e24)|~epred3_0)&((op2(e21,e21)!=e21|op2(e20,e24)!=e24)|~epred3_0))&((op2(e21,e22)!=e22|op2(e20,e24)!=e24)|~epred3_0))&((op2(e21,e23)!=e23|op2(e20,e24)!=e24)|~epred3_0))&((op2(e21,e24)!=e24|op2(e20,e24)!=e24)|~epred3_0)))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_389])])).
cnf(c_0_411,plain,(op2(e24,e21)=e21|op2(e23,e21)=e21|op2(e22,e21)=e21|op2(e21,e21)=e21|op2(e20,e21)=e21), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_412,plain,(op2(e24,e21)!=e21), inference(rw,[status(thm)],[c_0_390, c_0_244])).
cnf(c_0_413,plain,(op2(e22,e21)!=e21), inference(rw,[status(thm)],[c_0_391, c_0_392])).
cnf(c_0_414,plain,(op2(e20,e21)!=e21), inference(rw,[status(thm)],[c_0_393, c_0_394])).
cnf(c_0_415,plain,(op2(e20,e22)!=op2(e21,e22)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_416,plain,(op2(e21,e22)=e22|epred3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_396]), c_0_397])).
cnf(c_0_417,plain,(op2(e20,e22)=e20|op2(e21,e22)=e20), inference(sr,[status(thm)],[c_0_398, c_0_135])).
cnf(c_0_418,plain,(op2(e24,e20)=e23|op2(e24,e21)=e23|op2(e24,e23)=e23), inference(sr,[status(thm)],[c_0_399, c_0_134])).
fof(c_0_419, plain, ((((epred13_0&((((h5(e10)=e22|h5(e11)=e22)|h5(e12)=e22)|h5(e13)=e22)|h5(e14)=e22))&((((h5(e10)=e23|h5(e11)=e23)|h5(e12)=e23)|h5(e13)=e23)|h5(e14)=e23))=>epred14_0)), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_376]), c_0_400])).
fof(c_0_420, negated_conjecture, ((((((((((h1(e10)!=e24|~epred6_0)&(h1(e11)!=e24|~epred6_0))&(h1(e12)!=e24|~epred6_0))&(h1(e13)!=e24|~epred6_0))&(h1(e14)!=e24|~epred6_0))&(((((h2(e10)!=e24|~epred8_0)&(h2(e11)!=e24|~epred8_0))&(h2(e12)!=e24|~epred8_0))&(h2(e13)!=e24|~epred8_0))&(h2(e14)!=e24|~epred8_0)))&(((((h3(e10)!=e24|~epred10_0)&(h3(e11)!=e24|~epred10_0))&(h3(e12)!=e24|~epred10_0))&(h3(e13)!=e24|~epred10_0))&(h3(e14)!=e24|~epred10_0)))&(((((h4(e10)!=e24|~epred12_0)&(h4(e11)!=e24|~epred12_0))&(h4(e12)!=e24|~epred12_0))&(h4(e13)!=e24|~epred12_0))&(h4(e14)!=e24|~epred12_0)))&(((((h5(e10)!=e24|~epred14_0)&(h5(e11)!=e24|~epred14_0))&(h5(e12)!=e24|~epred14_0))&(h5(e13)!=e24|~epred14_0))&(h5(e14)!=e24|~epred14_0)))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_401])])).
cnf(c_0_421,plain,(op1(e10,e13)=e11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_403]), c_0_404])).
cnf(c_0_422,plain,(op1(e13,op1(e11,e13))=e11|op1(e13,e11)=e10), inference(spm,[status(thm)],[c_0_405, c_0_307])).
cnf(c_0_423,plain,(op1(e10,e13)=e11|epred1_0|op1(e12,e13)!=e10), inference(spm,[status(thm)],[c_0_381, c_0_406])).
cnf(c_0_424,plain,(op1(e12,e13)=e10|epred1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_407, c_0_408]), c_0_258])).
cnf(c_0_425,plain,(op2(op2(e22,e24),e21)=e24), inference(rw,[status(thm)],[c_0_160, c_0_244])).
cnf(c_0_426,plain,(op2(e21,e24)=e22|op2(e22,e24)=e22), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_186, c_0_319]), c_0_189])).
cnf(c_0_427,plain,(op2(e20,e22)=e22|op2(e21,e24)!=e22), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_409, c_0_244]), c_0_163])).
cnf(c_0_428,plain,(~epred3_0|op2(e20,e22)!=e22|op2(e21,e21)!=e21), inference(split_conjunct,[status(thm)],[c_0_410])).
cnf(c_0_429,plain,(op2(e21,e21)=e21), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_411, c_0_137]), c_0_412]), c_0_413]), c_0_414])).
cnf(c_0_430,plain,(epred3_0|op2(e20,e22)!=e22), inference(spm,[status(thm)],[c_0_415, c_0_416])).
cnf(c_0_431,plain,(op2(op2(e20,e23),op2(e23,e20))=e23), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_432,plain,(op2(e20,e22)=e20|op2(e20,e22)=e22|op2(e24,e22)=e22), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_417]), c_0_79])).
cnf(c_0_433,plain,(op2(e20,e24)=e22|op2(e20,e23)=e22|op2(e20,e22)=e22|op2(e20,e21)=e22|op2(e20,e20)=e22), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_434,plain,(op2(e24,e21)=e23|op2(e24,e20)=e23|op2(e21,e23)=e24), inference(spm,[status(thm)],[c_0_82, c_0_418])).
fof(c_0_435, plain, (((((((((((h5(e10)!=e23|(h5(e10)!=e22|~epred13_0))|epred14_0)&((h5(e11)!=e23|(h5(e10)!=e22|~epred13_0))|epred14_0))&((h5(e12)!=e23|(h5(e10)!=e22|~epred13_0))|epred14_0))&((h5(e13)!=e23|(h5(e10)!=e22|~epred13_0))|epred14_0))&((h5(e14)!=e23|(h5(e10)!=e22|~epred13_0))|epred14_0))&((((((h5(e10)!=e23|(h5(e11)!=e22|~epred13_0))|epred14_0)&((h5(e11)!=e23|(h5(e11)!=e22|~epred13_0))|epred14_0))&((h5(e12)!=e23|(h5(e11)!=e22|~epred13_0))|epred14_0))&((h5(e13)!=e23|(h5(e11)!=e22|~epred13_0))|epred14_0))&((h5(e14)!=e23|(h5(e11)!=e22|~epred13_0))|epred14_0)))&((((((h5(e10)!=e23|(h5(e12)!=e22|~epred13_0))|epred14_0)&((h5(e11)!=e23|(h5(e12)!=e22|~epred13_0))|epred14_0))&((h5(e12)!=e23|(h5(e12)!=e22|~epred13_0))|epred14_0))&((h5(e13)!=e23|(h5(e12)!=e22|~epred13_0))|epred14_0))&((h5(e14)!=e23|(h5(e12)!=e22|~epred13_0))|epred14_0)))&((((((h5(e10)!=e23|(h5(e13)!=e22|~epred13_0))|epred14_0)&((h5(e11)!=e23|(h5(e13)!=e22|~epred13_0))|epred14_0))&((h5(e12)!=e23|(h5(e13)!=e22|~epred13_0))|epred14_0))&((h5(e13)!=e23|(h5(e13)!=e22|~epred13_0))|epred14_0))&((h5(e14)!=e23|(h5(e13)!=e22|~epred13_0))|epred14_0)))&((((((h5(e10)!=e23|(h5(e14)!=e22|~epred13_0))|epred14_0)&((h5(e11)!=e23|(h5(e14)!=e22|~epred13_0))|epred14_0))&((h5(e12)!=e23|(h5(e14)!=e22|~epred13_0))|epred14_0))&((h5(e13)!=e23|(h5(e14)!=e22|~epred13_0))|epred14_0))&((h5(e14)!=e23|(h5(e14)!=e22|~epred13_0))|epred14_0)))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_419])])).
cnf(c_0_436,negated_conjecture,(~epred14_0|h5(e13)!=e24), inference(split_conjunct,[status(thm)],[c_0_420])).
cnf(c_0_437,plain,(h5(e13)=e24), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_438,plain,(op1(e12,e13)!=e11), inference(rw,[status(thm)],[c_0_381, c_0_421])).
cnf(c_0_439,plain,(op1(e12,e13)=e10|op1(e13,e11)=e10), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_403]), c_0_151])).
cnf(c_0_440,plain,(op1(e13,e11)=e13|op1(e11,e13)=e13), inference(rw,[status(thm)],[c_0_304, c_0_421])).
cnf(c_0_441,plain,(op1(op1(e10,e13),e13)=e13|op1(e13,e10)=e10), inference(spm,[status(thm)],[c_0_279, c_0_179])).
cnf(c_0_442,plain,(op1(e10,e13)=e11|epred1_0), inference(spm,[status(thm)],[c_0_423, c_0_424])).
fof(c_0_443, plain, ((((((((((((((((((((((((((((h5(op1(e10,e10))=op2(h5(e10),h5(e10))&h5(op1(e10,e11))=op2(h5(e10),h5(e11)))&h5(op1(e10,e12))=op2(h5(e10),h5(e12)))&h5(op1(e10,e13))=op2(h5(e10),h5(e13)))&h5(op1(e10,e14))=op2(h5(e10),h5(e14)))&h5(op1(e11,e10))=op2(h5(e11),h5(e10)))&h5(op1(e11,e11))=op2(h5(e11),h5(e11)))&h5(op1(e11,e12))=op2(h5(e11),h5(e12)))&h5(op1(e11,e13))=op2(h5(e11),h5(e13)))&h5(op1(e11,e14))=op2(h5(e11),h5(e14)))&h5(op1(e12,e10))=op2(h5(e12),h5(e10)))&h5(op1(e12,e11))=op2(h5(e12),h5(e11)))&h5(op1(e12,e12))=op2(h5(e12),h5(e12)))&h5(op1(e12,e13))=op2(h5(e12),h5(e13)))&h5(op1(e12,e14))=op2(h5(e12),h5(e14)))&h5(op1(e13,e10))=op2(h5(e13),h5(e10)))&h5(op1(e13,e11))=op2(h5(e13),h5(e11)))&h5(op1(e13,e12))=op2(h5(e13),h5(e12)))&h5(op1(e13,e13))=op2(h5(e13),h5(e13)))&h5(op1(e13,e14))=op2(h5(e13),h5(e14)))&h5(op1(e14,e10))=op2(h5(e14),h5(e10)))&h5(op1(e14,e11))=op2(h5(e14),h5(e11)))&h5(op1(e14,e12))=op2(h5(e14),h5(e12)))&h5(op1(e14,e13))=op2(h5(e14),h5(e13)))&h5(op1(e14,e14))=op2(h5(e14),h5(e14)))&((((h5(e10)=e20|h5(e11)=e20)|h5(e12)=e20)|h5(e13)=e20)|h5(e14)=e20))&((((h5(e10)=e21|h5(e11)=e21)|h5(e12)=e21)|h5(e13)=e21)|h5(e14)=e21))=>epred13_0)), inference(split_equiv,[status(thm)],[c_0_400])).
cnf(c_0_444,plain,(op2(e23,e20)=e23|op2(e24,e21)!=e23), inference(spm,[status(thm)],[c_0_164, c_0_208])).
cnf(c_0_445,plain,(op2(e22,e21)!=op2(e24,e21)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_446,plain,(op2(e21,e24)=e22|op2(e22,e21)=e24), inference(spm,[status(thm)],[c_0_425, c_0_426])).
cnf(c_0_447,plain,(op2(e21,e24)=e24|op2(e20,e22)=e22), inference(spm,[status(thm)],[c_0_427, c_0_271])).
cnf(c_0_448,plain,(op2(e20,e22)!=e22), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_428, c_0_429])]), c_0_430])).
cnf(c_0_449,plain,(op2(op2(e20,e23),e23)=e23|op2(e23,e20)=e20), inference(spm,[status(thm)],[c_0_431, c_0_228])).
cnf(c_0_450,plain,(op2(e20,e22)=e22|op2(e20,e22)=e20), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_432, c_0_244]), c_0_163])).
cnf(c_0_451,plain,(op2(e20,e24)=e22|op2(e20,e22)=e22|op2(e20,e21)=e22|op2(e20,e20)=e22), inference(sr,[status(thm)],[c_0_433, c_0_297])).
cnf(c_0_452,plain,(op2(e24,e20)=e23|op2(e24,e21)=e23|op2(e21,e24)!=e24), inference(spm,[status(thm)],[c_0_112, c_0_434])).
cnf(c_0_453,plain,(epred14_0|~epred13_0|h5(e14)!=e22|h5(e10)!=e23), inference(split_conjunct,[status(thm)],[c_0_435])).
cnf(c_0_454,negated_conjecture,(~epred14_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_436, c_0_437])])).
cnf(c_0_455,plain,(op1(e12,e10)!=op1(e12,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_456,plain,(op1(e10,e10)!=op1(e12,e10)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_457,plain,(op1(e14,e10)=e12|op1(e12,e13)=e14), inference(sr,[status(thm)],[c_0_291, c_0_438])).
cnf(c_0_458,plain,(op1(e12,e10)!=op1(e12,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_459,plain,(op1(e11,e13)=e13|op1(e12,e13)=e10), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_439, c_0_440]), c_0_126])).
cnf(c_0_460,plain,(op1(e10,e14)!=op1(e14,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_461,plain,(op1(e14,e10)!=op1(e14,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_462,plain,(op1(e13,e10)=e10|op1(e11,e13)=e13|epred1_0), inference(spm,[status(thm)],[c_0_441, c_0_442])).
fof(c_0_463, plain, (((((((((((h5(e10)!=e21|(h5(e10)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0)&((h5(e11)!=e21|(h5(e10)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e12)!=e21|(h5(e10)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e13)!=e21|(h5(e10)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e14)!=e21|(h5(e10)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((((((h5(e10)!=e21|(h5(e11)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0)&((h5(e11)!=e21|(h5(e11)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e12)!=e21|(h5(e11)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e13)!=e21|(h5(e11)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e14)!=e21|(h5(e11)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0)))&((((((h5(e10)!=e21|(h5(e12)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0)&((h5(e11)!=e21|(h5(e12)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e12)!=e21|(h5(e12)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e13)!=e21|(h5(e12)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e14)!=e21|(h5(e12)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0)))&((((((h5(e10)!=e21|(h5(e13)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0)&((h5(e11)!=e21|(h5(e13)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e12)!=e21|(h5(e13)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e13)!=e21|(h5(e13)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e14)!=e21|(h5(e13)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0)))&((((((h5(e10)!=e21|(h5(e14)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0)&((h5(e11)!=e21|(h5(e14)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e12)!=e21|(h5(e14)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e13)!=e21|(h5(e14)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0))&((h5(e14)!=e21|(h5(e14)!=e20|((((((((((((((((((((((((h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11)))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12)))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13)))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14)))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10)))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11)))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12)))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13)))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14)))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10)))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11)))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12)))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13)))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14)))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10)))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11)))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12)))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13)))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14)))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10)))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11)))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12)))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13)))|h5(op1(e14,e14))!=op2(h5(e14),h5(e14)))))|epred13_0)))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_443])])).
cnf(c_0_464,plain,(op2(e23,e21)=e20|op2(e23,e20)!=e23), inference(spm,[status(thm)],[c_0_231, c_0_165])).
cnf(c_0_465,plain,(op2(e24,e21)=e24|op2(e23,e20)=e23), inference(spm,[status(thm)],[c_0_444, c_0_269])).
cnf(c_0_466,plain,(op2(e21,e24)=e22|op2(e24,e21)!=e24), inference(spm,[status(thm)],[c_0_445, c_0_446])).
cnf(c_0_467,plain,(op2(e21,e24)=e24), inference(sr,[status(thm)],[c_0_447, c_0_448])).
cnf(c_0_468,plain,(op2(op2(e23,e20),e21)=e20), inference(rw,[status(thm)],[c_0_227, c_0_394])).
cnf(c_0_469,plain,(op2(e23,e20)=e20|op2(e21,e23)=e23), inference(rw,[status(thm)],[c_0_449, c_0_394])).
cnf(c_0_470,plain,(op2(e20,e21)!=op2(e20,e22)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_471,plain,(op2(e20,e22)=e20), inference(sr,[status(thm)],[c_0_450, c_0_448])).
cnf(c_0_472,plain,(op2(op2(e22,e20),op2(e20,e22))=e20), inference(split_conjunct,[status(thm)],[c_0_1])).
cnf(c_0_473,plain,(op2(e20,e21)=e22|op2(e20,e22)=e22|op2(e20,e24)=e22), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_451, c_0_29]), c_0_27])).
cnf(c_0_474,plain,(op2(e24,e21)=e23|op2(e21,e24)!=e24), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_452, c_0_294]), c_0_189])).
cnf(c_0_475,plain,(h5(e10)!=e23|h5(e14)!=e22|~epred13_0), inference(sr,[status(thm)],[c_0_453, c_0_454])).
cnf(c_0_476,plain,(h5(e14)=op2(e24,op2(e24,e24))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_477,plain,(h5(e10)=op2(op2(e24,op2(e24,e24)),op2(e24,op2(e24,e24)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_478,plain,(op1(e12,e10)=e14|op1(e12,e10)=e13|op1(e12,e10)=e12|op1(e12,e10)=e11|op1(e12,e10)=e10), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_479,plain,(op1(e12,e10)!=e13), inference(rw,[status(thm)],[c_0_455, c_0_128])).
cnf(c_0_480,plain,(op1(e12,e10)!=e14), inference(rw,[status(thm)],[c_0_456, c_0_63])).
cnf(c_0_481,plain,(op1(e12,e10)!=op1(e14,e10)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_482,plain,(op1(e14,e10)=e12), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_457]), c_0_45]), c_0_126])).
cnf(c_0_483,plain,(op1(e12,e13)=e14|op1(e12,e13)=e11|op1(e12,e10)!=e10), inference(spm,[status(thm)],[c_0_458, c_0_264])).
cnf(c_0_484,plain,(op1(e12,e13)=e10), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_403, c_0_459]), c_0_126])).
cnf(c_0_485,plain,(op1(e10,e12)!=op1(e13,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_486,plain,(op1(e10,e12)!=op1(e12,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_487,plain,(op1(e10,e14)=e10|op1(e10,e13)=e10|op1(e10,e12)=e10|op1(e10,e11)=e10|op1(e10,e10)=e10), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_488,plain,(op1(e10,e14)!=e10), inference(rw,[status(thm)],[c_0_460, c_0_45])).
cnf(c_0_489,plain,(op1(e10,e11)!=op1(e13,e11)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_490,plain,(op1(e14,e10)=e10|op1(e13,e10)=e10|op1(e12,e10)=e10|op1(e11,e10)=e10|op1(e10,e10)=e10), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_491,plain,(op1(e14,e10)!=e10), inference(rw,[status(thm)],[c_0_461, c_0_45])).
cnf(c_0_492,plain,(op1(e11,e13)=e13|epred1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_462]), c_0_352])).
cnf(c_0_493,plain,(epred13_0|h5(op1(e14,e14))!=op2(h5(e14),h5(e14))|h5(op1(e14,e13))!=op2(h5(e14),h5(e13))|h5(op1(e14,e12))!=op2(h5(e14),h5(e12))|h5(op1(e14,e11))!=op2(h5(e14),h5(e11))|h5(op1(e14,e10))!=op2(h5(e14),h5(e10))|h5(op1(e13,e14))!=op2(h5(e13),h5(e14))|h5(op1(e13,e13))!=op2(h5(e13),h5(e13))|h5(op1(e13,e12))!=op2(h5(e13),h5(e12))|h5(op1(e13,e11))!=op2(h5(e13),h5(e11))|h5(op1(e13,e10))!=op2(h5(e13),h5(e10))|h5(op1(e12,e14))!=op2(h5(e12),h5(e14))|h5(op1(e12,e13))!=op2(h5(e12),h5(e13))|h5(op1(e12,e12))!=op2(h5(e12),h5(e12))|h5(op1(e12,e11))!=op2(h5(e12),h5(e11))|h5(op1(e12,e10))!=op2(h5(e12),h5(e10))|h5(op1(e11,e14))!=op2(h5(e11),h5(e14))|h5(op1(e11,e13))!=op2(h5(e11),h5(e13))|h5(op1(e11,e12))!=op2(h5(e11),h5(e12))|h5(op1(e11,e11))!=op2(h5(e11),h5(e11))|h5(op1(e11,e10))!=op2(h5(e11),h5(e10))|h5(op1(e10,e14))!=op2(h5(e10),h5(e14))|h5(op1(e10,e13))!=op2(h5(e10),h5(e13))|h5(op1(e10,e12))!=op2(h5(e10),h5(e12))|h5(op1(e10,e11))!=op2(h5(e10),h5(e11))|h5(op1(e10,e10))!=op2(h5(e10),h5(e10))|h5(e12)!=e20|h5(e11)!=e21), inference(split_conjunct,[status(thm)],[c_0_463])).
cnf(c_0_494,plain,(h5(e11)=op2(e24,op2(e24,op2(e24,e24)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_495,plain,(h5(e12)=op2(e24,e24)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_496,plain,(op2(e24,e21)=e24|op2(e23,e21)=e20), inference(spm,[status(thm)],[c_0_464, c_0_465])).
cnf(c_0_497,plain,(op2(e24,e21)!=e24), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_466, c_0_467]), c_0_27])).
cnf(c_0_498,plain,(op2(e21,e23)=e23|op2(e20,e21)=e20), inference(spm,[status(thm)],[c_0_468, c_0_469])).
cnf(c_0_499,plain,(op2(e20,e21)!=e20), inference(rw,[status(thm)],[c_0_470, c_0_471])).
cnf(c_0_500,plain,(op2(e20,e22)=e22|op2(e21,e22)=e22), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_388, c_0_244]), c_0_163])).
cnf(c_0_501,plain,(op2(e20,e22)=e22|op2(e24,e23)=e24), inference(spm,[status(thm)],[c_0_142, c_0_447])).
cnf(c_0_502,plain,(op2(e21,op2(e20,e22))=e20), inference(rw,[status(thm)],[c_0_472, c_0_392])).
cnf(c_0_503,plain,(op2(e20,e22)=e22|op2(e20,e21)=e22), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_473, c_0_319]), c_0_189])).
cnf(c_0_504,plain,(op2(e20,e22)=e22|op2(e24,e21)=e23), inference(spm,[status(thm)],[c_0_474, c_0_447])).
cnf(c_0_505,plain,(op2(op2(e24,e20),op2(e24,e20))!=e23|op2(e24,e20)!=e22|~epred13_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_475, c_0_476]), c_0_477]), c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_506,plain,(op1(e12,e10)=e10|op1(e12,e10)=e11|op1(e12,e10)=e12), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_478, c_0_479]), c_0_480])).
cnf(c_0_507,plain,(op1(e12,e10)!=e12), inference(rw,[status(thm)],[c_0_481, c_0_482])).
cnf(c_0_508,plain,(op1(e12,e10)!=e10), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_483, c_0_484]), c_0_484]), c_0_89]), c_0_177])).
cnf(c_0_509,plain,(op1(e10,e12)=e14|op1(e10,e12)=e13|op1(e10,e12)=e12|op1(e10,e12)=e11|op1(e10,e12)=e10), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_510,plain,(op1(e10,e12)!=e14), inference(rw,[status(thm)],[c_0_485, c_0_35])).
cnf(c_0_511,plain,(op1(e10,e12)!=e13), inference(rw,[status(thm)],[c_0_486, c_0_128])).
cnf(c_0_512,plain,(op1(e10,e11)=e10|op1(e10,e12)=e10|op1(e10,e13)=e10), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_487, c_0_63]), c_0_89]), c_0_488])).
cnf(c_0_513,plain,(epred1_0|op1(e10,e11)!=e10), inference(spm,[status(thm)],[c_0_489, c_0_352])).
cnf(c_0_514,plain,(op1(e14,e10)=e11|op1(e13,e10)=e11|op1(e12,e10)=e11|op1(e11,e10)=e11|op1(e10,e10)=e11), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_515,plain,(op1(e11,e10)=e10|op1(e12,e10)=e10|op1(e13,e10)=e10), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[c_0_490, c_0_491]), c_0_63]), c_0_89])).
cnf(c_0_516,plain,(op1(e13,e10)=e13|epred1_0), inference(spm,[status(thm)],[c_0_382, c_0_492])).
cnf(c_0_517,plain,(epred1_0|op1(e12,e10)!=e10), inference(spm,[status(thm)],[c_0_458, c_0_424])).
cnf(c_0_518,plain,(op1(e10,e14)=e12|op1(e10,e13)=e12|op1(e10,e12)=e12|op1(e10,e11)=e12|op1(e10,e10)=e12), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_519,plain,(epred13_0|h5(op1(e10,e10))!=op2(op2(op2(e24,e20),op2(e24,e20)),op2(op2(e24,e20),op2(e24,e20)))|h5(op1(e10,e11))!=op2(op2(op2(e24,e20),op2(e24,e20)),op2(e24,op2(e24,e20)))|h5(op1(e11,e10))!=op2(op2(e24,op2(e24,e20)),op2(op2(e24,e20),op2(e24,e20)))|h5(op1(e10,e14))!=op2(op2(op2(e24,e20),op2(e24,e20)),op2(e24,e20))|h5(op1(e11,e11))!=op2(op2(e24,op2(e24,e20)),op2(e24,op2(e24,e20)))|h5(op1(e14,e10))!=op2(op2(e24,e20),op2(op2(e24,e20),op2(e24,e20)))|h5(op1(e10,e12))!=op2(op2(op2(e24,e20),op2(e24,e20)),e20)|h5(op1(e10,e13))!=op2(op2(op2(e24,e20),op2(e24,e20)),e24)|h5(op1(e11,e14))!=op2(op2(e24,op2(e24,e20)),op2(e24,e20))|h5(op1(e12,e10))!=op2(e20,op2(op2(e24,e20),op2(e24,e20)))|h5(op1(e13,e10))!=op2(e24,op2(op2(e24,e20),op2(e24,e20)))|h5(op1(e14,e11))!=op2(op2(e24,e20),op2(e24,op2(e24,e20)))|h5(op1(e11,e12))!=op2(op2(e24,op2(e24,e20)),e20)|h5(op1(e11,e13))!=op2(op2(e24,op2(e24,e20)),e24)|h5(op1(e12,e11))!=op2(e20,op2(e24,op2(e24,e20)))|h5(op1(e13,e11))!=op2(e24,op2(e24,op2(e24,e20)))|h5(op1(e12,e14))!=op2(e20,op2(e24,e20))|h5(op1(e14,e12))!=op2(op2(e24,e20),e20)|h5(op1(e14,e13))!=op2(op2(e24,e20),e24)|h5(op1(e12,e13))!=op2(e20,e24)|op2(e24,op2(e24,e20))!=e21|op2(e20,e20)!=e24), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_493, c_0_494]), c_0_21]), c_0_495]), c_0_21]), c_0_477]), c_0_21]), c_0_21]), c_0_477]), c_0_21]), c_0_21]), c_0_477]), c_0_21]), c_0_21]), c_0_494]), c_0_21]), c_0_477]), c_0_21]), c_0_21]), c_0_495]), c_0_21]), c_0_477]), c_0_21]), c_0_21]), c_0_437]), c_0_477]), c_0_21]), c_0_21]), c_0_476]), c_0_21]), c_0_494]), c_0_21]), c_0_477]), c_0_21]), c_0_21]), c_0_494]), c_0_21]), c_0_494]), c_0_21]), c_0_494]), c_0_21]), c_0_495]), c_0_21]), c_0_494]), c_0_21]), c_0_437]), c_0_494]), c_0_21]), c_0_476]), c_0_21]), c_0_495]), c_0_21]), c_0_477]), c_0_21]), c_0_21]), c_0_495]), c_0_21]), c_0_494]), c_0_21]), c_0_128]), c_0_437]), c_0_495]), c_0_21]), c_0_495]), c_0_21]), c_0_495]), c_0_21]), c_0_437]), c_0_495]), c_0_21]), c_0_476]), c_0_21]), c_0_437]), c_0_477]), c_0_21]), c_0_21]), c_0_437]), c_0_494]), c_0_21]), c_0_35]), c_0_476]), c_0_21]), c_0_437]), c_0_495]), c_0_21]), c_0_25]), c_0_495]), c_0_21]), c_0_437]), c_0_437]), c_0_21]), c_0_87]), c_0_494]), c_0_21]), c_0_437]), c_0_476]), c_0_21]), c_0_476]), c_0_21]), c_0_477]), c_0_21]), c_0_21]), c_0_476]), c_0_21]), c_0_494]), c_0_21]), c_0_476]), c_0_21]), c_0_495]), c_0_21]), c_0_476]), c_0_21]), c_0_437]), c_0_45]), c_0_477]), c_0_21]), c_0_21]), c_0_476]), c_0_21]), c_0_476]), c_0_21])])).
cnf(c_0_520,plain,(op2(e23,e21)=e20), inference(sr,[status(thm)],[c_0_496, c_0_497])).
cnf(c_0_521,plain,(op2(e21,e23)=e23), inference(sr,[status(thm)],[c_0_498, c_0_499])).
cnf(c_0_522,plain,(op2(e23,e20)=e23), inference(sr,[status(thm)],[c_0_465, c_0_497])).
cnf(c_0_523,plain,(op2(e21,e22)=e22), inference(sr,[status(thm)],[c_0_500, c_0_448])).
cnf(c_0_524,plain,(op2(e24,e23)=e24), inference(sr,[status(thm)],[c_0_501, c_0_448])).
cnf(c_0_525,plain,(op2(e22,e21)=e24), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_446, c_0_467]), c_0_27])).
cnf(c_0_526,plain,(op2(e21,e20)=e20), inference(rw,[status(thm)],[c_0_502, c_0_471])).
cnf(c_0_527,plain,(op2(e20,e21)=e22), inference(sr,[status(thm)],[c_0_503, c_0_448])).
cnf(c_0_528,plain,(op2(e24,e21)=e23), inference(sr,[status(thm)],[c_0_504, c_0_448])).
cnf(c_0_529,plain,(op2(e22,e24)=e22), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_426, c_0_467]), c_0_27])).
cnf(c_0_530,plain,(~epred13_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_505, c_0_294]), c_0_294]), c_0_99]), c_0_294])])).
fof(c_0_531, plain, ((epred1_0=>(((((op1(e10,e10)!=e10&op1(e10,e11)!=e11)&op1(e10,e12)!=e12)&op1(e10,e13)!=e13)&op1(e10,e14)!=e14)|((((op1(e11,e10)!=e10&op1(e11,e11)!=e11)&op1(e11,e12)!=e12)&op1(e11,e13)!=e13)&op1(e11,e14)!=e14)))), inference(split_equiv,[status(thm)],[c_0_113])).
cnf(c_0_532,plain,(op1(op1(e12,e10),op1(e10,e12))=e10), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_533,plain,(op1(e12,e10)=e11), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_506, c_0_507]), c_0_508])).
cnf(c_0_534,plain,(op1(e10,e12)=e10|op1(e10,e12)=e11|op1(e10,e12)=e12), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_509, c_0_510]), c_0_511])).
cnf(c_0_535,plain,(op1(e10,e12)!=e11), inference(rw,[status(thm)],[c_0_330, c_0_421])).
cnf(c_0_536,plain,(op1(e11,e14)!=op1(e13,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_537,plain,(op1(e14,e12)!=e12), inference(rw,[status(thm)],[c_0_358, c_0_482])).
cnf(c_0_538,plain,(op1(op1(e10,e12),op1(e12,e10))=e12), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_539,plain,(op1(e10,e12)=e10|epred1_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_512]), c_0_177]), c_0_513])).
cnf(c_0_540,plain,(op1(e11,e10)=e11|op1(e12,e10)=e11|op1(e14,e10)=e11), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[c_0_514, c_0_151]), c_0_63]), c_0_88])).
cnf(c_0_541,plain,(op1(e11,e10)=e10|epred1_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_515, c_0_516]), c_0_126]), c_0_517])).
cnf(c_0_542,plain,(op1(e12,e13)=e14|op1(e12,e13)=e11|op1(e10,e14)=e13), inference(spm,[status(thm)],[c_0_173, c_0_264])).
cnf(c_0_543,plain,(op1(e10,e14)=e12|op1(e10,e12)=e12|op1(e10,e11)=e12|op1(e10,e10)=e12), inference(sr,[status(thm)],[c_0_518, c_0_256])).
cnf(c_0_544,plain,(op1(e12,e14)=e12|op1(e12,e13)=e12|op1(e12,e12)=e12|op1(e12,e11)=e12|op1(e12,e10)=e12), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_545,plain,(op1(e10,e13)!=op1(e14,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_546,plain,(op1(e14,e13)=e13|op1(e13,e13)=e13|op1(e12,e13)=e13|op1(e11,e13)=e13|op1(e10,e13)=e13), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_547,plain,(h5(op1(e10,e11))!=e20|h5(op1(e11,e10))!=e23|h5(op1(e10,e14))!=e24|h5(op1(e11,e11))!=e21|h5(op1(e14,e10))!=e20|h5(op1(e10,e12))!=e23|h5(op1(e10,e13))!=e21|h5(op1(e11,e14))!=e22|h5(op1(e12,e10))!=e21|h5(op1(e13,e10))!=e24|h5(op1(e14,e11))!=e24|h5(op1(e11,e12))!=e20|h5(op1(e11,e13))!=e24|h5(op1(e12,e11))!=e22|h5(op1(e13,e11))!=e23|h5(op1(e12,e14))!=e20|h5(op1(e14,e12))!=e21|h5(op1(e14,e13))!=e22|h5(op1(e12,e13))!=e23), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_519, c_0_63]), c_0_476]), c_0_21]), c_0_294]), c_0_294]), c_0_99]), c_0_294]), c_0_294]), c_0_99]), c_0_17]), c_0_294]), c_0_294]), c_0_294]), c_0_99]), c_0_294]), c_0_244]), c_0_520]), c_0_294]), c_0_244]), c_0_294]), c_0_294]), c_0_99]), c_0_521]), c_0_294]), c_0_294]), c_0_99]), c_0_294]), c_0_19]), c_0_294]), c_0_244]), c_0_294]), c_0_244]), c_0_429]), c_0_294]), c_0_294]), c_0_294]), c_0_99]), c_0_339]), c_0_294]), c_0_294]), c_0_99]), c_0_522]), c_0_294]), c_0_294]), c_0_99]), c_0_50]), c_0_294]), c_0_244]), c_0_294]), c_0_523]), c_0_294]), c_0_294]), c_0_99]), c_0_394]), c_0_294]), c_0_294]), c_0_99]), c_0_524]), c_0_294]), c_0_294]), c_0_244]), c_0_525]), c_0_294]), c_0_244]), c_0_526]), c_0_294]), c_0_244]), c_0_467]), c_0_294]), c_0_244]), c_0_527]), c_0_294]), c_0_244]), c_0_528]), c_0_294]), c_0_471]), c_0_294]), c_0_392]), c_0_294]), c_0_529]), c_0_319]), c_0_294]), c_0_244]), c_0_29])]), c_0_530])).
fof(c_0_548, plain, (((((((((((op1(e11,e10)!=e10|op1(e10,e10)!=e10)|~epred1_0)&((op1(e11,e11)!=e11|op1(e10,e10)!=e10)|~epred1_0))&((op1(e11,e12)!=e12|op1(e10,e10)!=e10)|~epred1_0))&((op1(e11,e13)!=e13|op1(e10,e10)!=e10)|~epred1_0))&((op1(e11,e14)!=e14|op1(e10,e10)!=e10)|~epred1_0))&((((((op1(e11,e10)!=e10|op1(e10,e11)!=e11)|~epred1_0)&((op1(e11,e11)!=e11|op1(e10,e11)!=e11)|~epred1_0))&((op1(e11,e12)!=e12|op1(e10,e11)!=e11)|~epred1_0))&((op1(e11,e13)!=e13|op1(e10,e11)!=e11)|~epred1_0))&((op1(e11,e14)!=e14|op1(e10,e11)!=e11)|~epred1_0)))&((((((op1(e11,e10)!=e10|op1(e10,e12)!=e12)|~epred1_0)&((op1(e11,e11)!=e11|op1(e10,e12)!=e12)|~epred1_0))&((op1(e11,e12)!=e12|op1(e10,e12)!=e12)|~epred1_0))&((op1(e11,e13)!=e13|op1(e10,e12)!=e12)|~epred1_0))&((op1(e11,e14)!=e14|op1(e10,e12)!=e12)|~epred1_0)))&((((((op1(e11,e10)!=e10|op1(e10,e13)!=e13)|~epred1_0)&((op1(e11,e11)!=e11|op1(e10,e13)!=e13)|~epred1_0))&((op1(e11,e12)!=e12|op1(e10,e13)!=e13)|~epred1_0))&((op1(e11,e13)!=e13|op1(e10,e13)!=e13)|~epred1_0))&((op1(e11,e14)!=e14|op1(e10,e13)!=e13)|~epred1_0)))&((((((op1(e11,e10)!=e10|op1(e10,e14)!=e14)|~epred1_0)&((op1(e11,e11)!=e11|op1(e10,e14)!=e14)|~epred1_0))&((op1(e11,e12)!=e12|op1(e10,e14)!=e14)|~epred1_0))&((op1(e11,e13)!=e13|op1(e10,e14)!=e14)|~epred1_0))&((op1(e11,e14)!=e14|op1(e10,e14)!=e14)|~epred1_0)))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_531])])).
cnf(c_0_549,plain,(op1(e11,op1(e10,e12))=e10), inference(rw,[status(thm)],[c_0_532, c_0_533])).
cnf(c_0_550,plain,(op1(e10,e12)=e12|op1(e10,e12)=e10), inference(sr,[status(thm)],[c_0_534, c_0_535])).
cnf(c_0_551,plain,(op1(e11,e14)=e11|op1(e11,e13)=e11|op1(e11,e12)=e11|op1(e11,e11)=e11|op1(e11,e10)=e11), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_552,plain,(op1(e11,e14)!=e11), inference(rw,[status(thm)],[c_0_536, c_0_87])).
cnf(c_0_553,plain,(op1(e11,e10)!=op1(e12,e10)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_554,plain,(op1(e11,e12)!=op1(e14,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_555,plain,(op1(e14,e12)=e11), inference(sr,[status(thm)],[c_0_385, c_0_537])).
cnf(c_0_556,plain,(op1(e10,op1(e12,e10))=e12|epred1_0), inference(spm,[status(thm)],[c_0_538, c_0_539])).
cnf(c_0_557,plain,(op1(e14,e10)=e11|op1(e12,e10)=e11|epred1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_540, c_0_541]), c_0_177])).
cnf(c_0_558,plain,(op1(e10,e14)=e13|op1(e12,e13)=e14), inference(sr,[status(thm)],[c_0_542, c_0_438])).
cnf(c_0_559,plain,(op1(e10,e11)!=op1(e12,e11)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_560,plain,(op1(e10,e11)=e12|op1(e10,e12)=e12|op1(e10,e14)=e12), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_543, c_0_63]), c_0_61])).
cnf(c_0_561,plain,(op1(e12,e14)=e12|op1(e12,e11)=e12|op1(e12,e10)=e12), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_544, c_0_128]), c_0_131]), c_0_214])).
cnf(c_0_562,plain,(op1(e14,e13)=e14|op1(e14,e13)=e11|op1(e13,e11)=e13), inference(spm,[status(thm)],[c_0_180, c_0_120])).
cnf(c_0_563,plain,(op1(e14,e13)!=e11), inference(rw,[status(thm)],[c_0_545, c_0_421])).
cnf(c_0_564,plain,(op1(e14,e13)=e13|op1(e12,e13)=e13|op1(e11,e13)=e13|op1(e10,e13)=e13), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_546, c_0_25]), c_0_131])).
cnf(c_0_565,plain,(h5(op1(e10,e11))!=e20|h5(op1(e11,e10))!=e23|h5(op1(e10,e14))!=e24|h5(op1(e11,e11))!=e21|h5(op1(e14,e10))!=e20|h5(op1(e10,e12))!=e23|h5(op1(e11,e14))!=e22|h5(op1(e12,e10))!=e21|h5(op1(e13,e10))!=e24|h5(op1(e14,e11))!=e24|h5(op1(e11,e12))!=e20|h5(op1(e11,e13))!=e24|h5(op1(e12,e11))!=e22|h5(op1(e13,e11))!=e23|h5(op1(e12,e14))!=e20|h5(op1(e14,e12))!=e21|h5(op1(e14,e13))!=e22|h5(op1(e12,e13))!=e23), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_547, c_0_421]), c_0_494]), c_0_21]), c_0_294]), c_0_244])])).
cnf(c_0_566,plain,(~epred1_0|op1(e10,e12)!=e12|op1(e11,e11)!=e11), inference(split_conjunct,[status(thm)],[c_0_548])).
cnf(c_0_567,plain,(op1(e10,e12)=e12|op1(e11,e10)=e10), inference(spm,[status(thm)],[c_0_549, c_0_550])).
cnf(c_0_568,plain,(op1(e11,e10)=e11|op1(e11,e11)=e11|op1(e11,e12)=e11|op1(e11,e13)=e11), inference(sr,[status(thm)],[c_0_551, c_0_552])).
cnf(c_0_569,plain,(op1(e11,e13)!=e11), inference(rw,[status(thm)],[c_0_377, c_0_421])).
cnf(c_0_570,plain,(op1(e11,e10)!=e11), inference(rw,[status(thm)],[c_0_553, c_0_533])).
cnf(c_0_571,plain,(op1(e11,e12)!=e11), inference(rw,[status(thm)],[c_0_554, c_0_555])).
cnf(c_0_572,plain,(op1(e10,e11)!=op1(e10,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_573,plain,(op1(e14,e10)=e11|op1(e10,e11)=e12|epred1_0), inference(spm,[status(thm)],[c_0_556, c_0_557])).
cnf(c_0_574,plain,(op1(e14,e11)=e13|op1(e10,e11)=e10), inference(rw,[status(thm)],[c_0_347, c_0_421])).
cnf(c_0_575,plain,(op1(e13,e11)=e10|op1(e10,e11)=e10), inference(rw,[status(thm)],[c_0_327, c_0_421])).
cnf(c_0_576,plain,(op1(e10,e14)=e13), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_558, c_0_484]), c_0_89])).
cnf(c_0_577,plain,(op1(e10,e14)=e12|op1(e10,e12)=e12|op1(e12,e11)!=e12), inference(spm,[status(thm)],[c_0_559, c_0_560])).
cnf(c_0_578,plain,(op1(e12,e11)=e12|op1(e12,e14)=e12|op1(e14,e10)!=e12), inference(spm,[status(thm)],[c_0_481, c_0_561])).
cnf(c_0_579,plain,(op1(e13,e11)=e13|op1(e14,e13)=e14), inference(sr,[status(thm)],[c_0_562, c_0_563])).
cnf(c_0_580,plain,(op1(e10,e13)=e13|op1(e11,e13)=e13|op1(e14,e13)=e13), inference(sr,[status(thm)],[c_0_564, c_0_235])).
cnf(c_0_581,plain,(h5(op1(e10,e11))!=e20|h5(op1(e11,e10))!=e23|h5(op1(e10,e14))!=e24|h5(op1(e11,e11))!=e21|h5(op1(e10,e12))!=e23|h5(op1(e11,e14))!=e22|h5(op1(e12,e10))!=e21|h5(op1(e13,e10))!=e24|h5(op1(e14,e11))!=e24|h5(op1(e11,e12))!=e20|h5(op1(e11,e13))!=e24|h5(op1(e12,e11))!=e22|h5(op1(e13,e11))!=e23|h5(op1(e12,e14))!=e20|h5(op1(e14,e12))!=e21|h5(op1(e14,e13))!=e22|h5(op1(e12,e13))!=e23), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_565, c_0_482]), c_0_495]), c_0_21])])).
cnf(c_0_582,plain,(op1(e10,e10)!=op1(e10,e11)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_583,plain,(op1(e13,e11)=e10|op1(e11,e10)=e13), inference(rw,[status(thm)],[c_0_350, c_0_421])).
cnf(c_0_584,plain,(op1(e11,e10)=e10|op1(e11,e11)!=e11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_567]), c_0_541])).
cnf(c_0_585,plain,(op1(e11,e11)=e11), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_568, c_0_569]), c_0_570]), c_0_571])).
cnf(c_0_586,plain,(op1(e14,e10)=e11|epred1_0|op1(e10,e12)!=e12), inference(spm,[status(thm)],[c_0_572, c_0_573])).
cnf(c_0_587,plain,(op1(e14,e11)=e13|op1(e10,e12)!=e10), inference(spm,[status(thm)],[c_0_572, c_0_574])).
cnf(c_0_588,plain,(op1(e10,e11)=e10|op1(e11,e13)=e13), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_575]), c_0_126])).
cnf(c_0_589,plain,(op1(e10,e12)=e12|op1(e10,e11)=e12), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_560, c_0_576]), c_0_131])).
cnf(c_0_590,plain,(op1(e10,e12)=e12|op1(e12,e11)!=e12), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_577, c_0_576]), c_0_131])).
cnf(c_0_591,plain,(op1(e12,e14)=e12|op1(e12,e11)=e12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_578, c_0_482])])).
cnf(c_0_592,plain,(op1(e14,e13)=e14|op1(e10,e11)=e10), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_575, c_0_579]), c_0_126])).
cnf(c_0_593,plain,(op1(e14,e13)=e13|op1(e11,e13)=e13), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_580, c_0_421]), c_0_132])).
cnf(c_0_594,plain,(h5(op1(e10,e11))!=e20|h5(op1(e11,e10))!=e23|h5(op1(e10,e14))!=e24|h5(op1(e11,e11))!=e21|h5(op1(e10,e12))!=e23|h5(op1(e11,e14))!=e22|h5(op1(e12,e10))!=e21|h5(op1(e13,e10))!=e24|h5(op1(e14,e11))!=e24|h5(op1(e11,e12))!=e20|h5(op1(e11,e13))!=e24|h5(op1(e12,e11))!=e22|h5(op1(e13,e11))!=e23|h5(op1(e12,e14))!=e20|h5(op1(e14,e13))!=e22|h5(op1(e12,e13))!=e23), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_581, c_0_555]), c_0_494]), c_0_21]), c_0_294]), c_0_244])])).
cnf(c_0_595,plain,(op1(e10,e11)=e14|op1(e10,e11)=e13|op1(e10,e11)=e12|op1(e10,e11)=e11|op1(e10,e11)=e10), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_596,plain,(op1(e10,e11)!=e14), inference(rw,[status(thm)],[c_0_582, c_0_63])).
cnf(c_0_597,plain,(op1(e10,e11)!=op1(e10,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_598,plain,(op1(e11,e10)=e13|op1(e10,e11)!=e10), inference(spm,[status(thm)],[c_0_489, c_0_583])).
cnf(c_0_599,plain,(op1(e11,e10)=e10), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_584, c_0_585])])).
cnf(c_0_600,plain,(op1(e10,e11)!=op1(e10,e14)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_601,plain,(epred1_0|op1(e10,e12)!=e12), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_586, c_0_482]), c_0_284])).
cnf(c_0_602,plain,(op1(e10,e12)=e12|op1(e14,e11)=e13), inference(spm,[status(thm)],[c_0_587, c_0_550])).
cnf(c_0_603,plain,(op1(e11,e12)=e12|op1(e14,e12)=e12|op1(e11,e11)!=e11|~epred1_0), inference(spm,[status(thm)],[c_0_566, c_0_312])).
cnf(c_0_604,plain,(op1(e10,e12)=e12|op1(e11,e13)=e13), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_588, c_0_589]), c_0_258])).
cnf(c_0_605,plain,(op1(e12,e14)=e12|op1(e10,e12)=e12), inference(spm,[status(thm)],[c_0_590, c_0_591])).
cnf(c_0_606,plain,(op1(e10,e12)=e12|op1(e14,e13)=e14), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_592, c_0_589]), c_0_258])).
cnf(c_0_607,plain,(op1(e14,e10)=e13|op1(e13,e10)=e13|op1(e12,e10)=e13|op1(e11,e10)=e13|op1(e10,e10)=e13), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_608,plain,(op1(e11,e12)!=op1(e13,e12)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_609,plain,(op1(e10,e10)!=op1(e11,e10)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_610,plain,(op1(e14,e13)=e13|op1(e13,e11)=e10), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_593]), c_0_25]), c_0_284])).
cnf(c_0_611,plain,(h5(op1(e10,e11))!=e20|h5(op1(e11,e10))!=e23|h5(op1(e11,e11))!=e21|h5(op1(e10,e12))!=e23|h5(op1(e11,e14))!=e22|h5(op1(e12,e10))!=e21|h5(op1(e13,e10))!=e24|h5(op1(e14,e11))!=e24|h5(op1(e11,e12))!=e20|h5(op1(e11,e13))!=e24|h5(op1(e12,e11))!=e22|h5(op1(e13,e11))!=e23|h5(op1(e12,e14))!=e20|h5(op1(e14,e13))!=e22), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_594, c_0_484]), c_0_477]), c_0_21]), c_0_294]), c_0_21]), c_0_294]), c_0_99])]), c_0_576]), c_0_437])])).
cnf(c_0_612,plain,(op1(e10,e11)=e10|op1(e10,e11)=e11|op1(e10,e11)=e12|op1(e10,e11)=e13), inference(sr,[status(thm)],[c_0_595, c_0_596])).
cnf(c_0_613,plain,(op1(e10,e11)!=e11), inference(rw,[status(thm)],[c_0_597, c_0_421])).
cnf(c_0_614,plain,(op1(e10,e11)!=e10), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_598, c_0_599]), c_0_126])).
cnf(c_0_615,plain,(op1(e10,e11)!=e13), inference(rw,[status(thm)],[c_0_600, c_0_576])).
cnf(c_0_616,plain,(op1(e10,e12)!=e12), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_566, c_0_585])]), c_0_601])).
cnf(c_0_617,plain,(op1(e14,e11)=e13|op1(e11,e11)!=e11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_602]), c_0_266])).
cnf(c_0_618,plain,(op1(e11,e12)=e12|op1(e11,e11)!=e11), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_603, c_0_555]), c_0_284]), c_0_408])).
cnf(c_0_619,plain,(op1(e11,e13)=e13|op1(e11,e11)!=e11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_604]), c_0_492])).
cnf(c_0_620,plain,(op1(e12,e14)=e12|op1(e11,e11)!=e11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_605]), c_0_355])).
cnf(c_0_621,plain,(op1(e14,e13)=e14|op1(e11,e11)!=e11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_606]), c_0_292])).
cnf(c_0_622,plain,(op1(e11,e10)=e13|op1(e13,e10)=e13|op1(e14,e10)=e13), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[c_0_607, c_0_479]), c_0_63]), c_0_130])).
cnf(c_0_623,plain,(op1(e11,e14)=e14|op1(e11,e13)=e14|op1(e11,e12)=e14|op1(e11,e11)=e14|op1(e11,e10)=e14), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_624,plain,(op1(e11,e12)!=e14), inference(rw,[status(thm)],[c_0_608, c_0_35])).
cnf(c_0_625,plain,(op1(e11,e10)!=e14), inference(rw,[status(thm)],[c_0_609, c_0_63])).
cnf(c_0_626,plain,(op1(op1(e12,e14),op1(e14,e12))=e14), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_627,plain,(op1(e14,e13)=e13|op1(e10,e11)!=e10), inference(spm,[status(thm)],[c_0_489, c_0_610])).
cnf(c_0_628,plain,(op1(e10,e12)=e10|op1(e10,e11)=e10), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_512, c_0_421]), c_0_177])).
cnf(c_0_629,plain,(h5(op1(e10,e11))!=e20|h5(op1(e11,e10))!=e23|h5(op1(e11,e11))!=e21|h5(op1(e10,e12))!=e23|h5(op1(e11,e14))!=e22|h5(op1(e13,e10))!=e24|h5(op1(e14,e11))!=e24|h5(op1(e11,e12))!=e20|h5(op1(e11,e13))!=e24|h5(op1(e12,e11))!=e22|h5(op1(e13,e11))!=e23|h5(op1(e12,e14))!=e20|h5(op1(e14,e13))!=e22), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_611, c_0_533]), c_0_494]), c_0_21]), c_0_294]), c_0_244])])).
cnf(c_0_630,plain,(op1(e10,e11)=e12), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_612, c_0_613]), c_0_614]), c_0_615])).
cnf(c_0_631,plain,(op1(e10,e12)=e10), inference(sr,[status(thm)],[c_0_550, c_0_616])).
cnf(c_0_632,plain,(op1(e14,e11)=e13), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_617, c_0_585])])).
cnf(c_0_633,plain,(op1(e11,e12)=e12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_618, c_0_585])])).
cnf(c_0_634,plain,(op1(e11,e13)=e13), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_619, c_0_585])])).
cnf(c_0_635,plain,(op1(e12,e14)=e12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_620, c_0_585])])).
cnf(c_0_636,plain,(op1(e14,e13)=e14), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_621, c_0_585])])).
cnf(c_0_637,plain,(op1(e13,e10)=e13|op1(e11,e10)=e13), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_622, c_0_482]), c_0_131])).
cnf(c_0_638,plain,(op1(e11,e13)!=op1(e14,e13)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_639,plain,(op1(e11,e11)=e14|op1(e11,e13)=e14|op1(e11,e14)=e14), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_623, c_0_624]), c_0_625])).
cnf(c_0_640,plain,(op1(op1(e12,e14),e11)=e14|op1(e14,e12)=e12), inference(spm,[status(thm)],[c_0_626, c_0_385])).
cnf(c_0_641,plain,(op1(e10,e12)=e10|op1(e14,e13)=e13), inference(spm,[status(thm)],[c_0_627, c_0_628])).
cnf(c_0_642,plain,(op1(e11,e10)=e13|op1(e11,e13)=e13), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_583]), c_0_126])).
cnf(c_0_643,plain,(h5(op1(e11,e14))!=e22|h5(op1(e13,e10))!=e24|h5(op1(e12,e11))!=e22|h5(op1(e13,e11))!=e23), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_629, c_0_585]), c_0_494]), c_0_21]), c_0_294]), c_0_244])]), c_0_630]), c_0_495]), c_0_21]), c_0_599]), c_0_477]), c_0_21]), c_0_294]), c_0_21]), c_0_294]), c_0_99]), c_0_631]), c_0_477]), c_0_21]), c_0_294]), c_0_21]), c_0_294]), c_0_99]), c_0_632]), c_0_437]), c_0_633]), c_0_495]), c_0_21]), c_0_634]), c_0_437]), c_0_635]), c_0_495]), c_0_21]), c_0_636]), c_0_476]), c_0_21]), c_0_294])])).
cnf(c_0_644,plain,(op1(e13,e10)=e13), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_637, c_0_599]), c_0_126])).
cnf(c_0_645,plain,(op1(e13,e11)=e10), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_583, c_0_599]), c_0_126])).
cnf(c_0_646,plain,(op1(e11,e14)=e14|op1(e11,e11)=e14|op1(e14,e13)!=e14), inference(spm,[status(thm)],[c_0_638, c_0_639])).
cnf(c_0_647,plain,(op1(op1(e12,e14),e11)=e14), inference(sr,[status(thm)],[c_0_640, c_0_537])).
cnf(c_0_648,plain,(op1(e14,e13)=e13|op1(e12,e14)=e12), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_605, c_0_641]), c_0_258])).
cnf(c_0_649,plain,(op1(e11,e10)=e13|op1(e14,e13)!=e13), inference(spm,[status(thm)],[c_0_638, c_0_642])).
cnf(c_0_650,plain,(h5(op1(e11,e14))!=e22|h5(op1(e12,e11))!=e22), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_643, c_0_644]), c_0_437])]), c_0_645]), c_0_477]), c_0_21]), c_0_294]), c_0_21]), c_0_294]), c_0_99])])).
cnf(c_0_651,plain,(op1(e11,e14)=e14), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_646, c_0_585]), c_0_88]), c_0_636])])).
cnf(c_0_652,plain,(op1(e14,e13)=e13|op1(e12,e11)=e14), inference(spm,[status(thm)],[c_0_647, c_0_648])).
cnf(c_0_653,plain,(op1(e14,e13)!=e13), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_649, c_0_599]), c_0_126])).
cnf(c_0_654,plain,(h5(op1(e12,e11))!=e22), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_650, c_0_651]), c_0_476]), c_0_21]), c_0_294])])).
cnf(c_0_655,plain,(op1(e12,e11)=e14), inference(sr,[status(thm)],[c_0_652, c_0_653])).
cnf(c_0_656,plain,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_654, c_0_655]), c_0_476]), c_0_21]), c_0_294])]), ['proof']).
# SZS output end CNFRefutation.
EOF
