% Problem    : ALG220+1 : TPTP v6.1.0. Released v3.4.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n127.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:21:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___200_B02_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X1]:((((((((((~(v3_struct_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&l3_vectsp_1(X1))=>![X2]:((((((((~(v3_struct_0(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v3_vectsp_2(X2,X1))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))=>![X3]:((((((((~(v3_struct_0(X3))&v3_rlvect_1(X3))&v4_rlvect_1(X3))&v5_rlvect_1(X3))&v6_rlvect_1(X3))&v3_vectsp_2(X3,X1))&v5_vectsp_2(X3,X1))&l1_vectsp_2(X3,X1))=>((m1_rmod_2(X3,X1,X2)&m1_rmod_2(X2,X1,X3))=>X3=X2))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', t33_rmod_2)).
fof(c_0_1, axiom, (![X1]:![X2]:(((((((((((((((((~(v3_struct_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&l3_vectsp_1(X1))&~(v3_struct_0(X2)))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))=>![X3]:(m1_rmod_2(X3,X1,X2)=>((((((~(v3_struct_0(X3))&v3_rlvect_1(X3))&v4_rlvect_1(X3))&v5_rlvect_1(X3))&v6_rlvect_1(X3))&v5_vectsp_2(X3,X1))&l1_vectsp_2(X3,X1))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', dt_m1_rmod_2)).
fof(c_0_2, axiom, (![X1]:(((((((((((((~(v3_struct_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v7_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&~(v10_vectsp_1(X1)))&v2_vectsp_2(X1))&l3_vectsp_1(X1))=>![X2]:(((((((~(v3_struct_0(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(r1_tarski(X3,X4)=>m1_rmod_2(k1_rmod_5(X1,X2,X3),X1,k1_rmod_5(X1,X2,X4)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', t15_rmod_5)).
fof(c_0_3, conjecture, (![X1]:(((((((((((((~(v3_struct_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v7_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&~(v10_vectsp_1(X1)))&v2_vectsp_2(X1))&l3_vectsp_1(X1))=>![X2]:((((((((~(v3_struct_0(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v3_vectsp_2(X2,X1))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>((k1_rmod_5(X1,X2,X3)=X2&r1_tarski(X3,X4))=>k1_rmod_5(X1,X2,X4)=X2)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', t16_rmod_5)).
fof(c_0_4, axiom, (![X1]:![X2]:![X3]:(((((((((((((((((((((~(v3_struct_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v7_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&~(v10_vectsp_1(X1)))&v2_vectsp_2(X1))&l3_vectsp_1(X1))&~(v3_struct_0(X2)))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(v3_vectsp_2(k1_rmod_5(X1,X2,X3),X1)&m1_rmod_2(k1_rmod_5(X1,X2,X3),X1,X2)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', dt_k1_rmod_5)).
fof(c_0_5, plain, (![X1]:((((((((((~v3_struct_0(X1)&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&l3_vectsp_1(X1))=>![X2]:((((((((~v3_struct_0(X2)&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v3_vectsp_2(X2,X1))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))=>![X3]:((((((((~v3_struct_0(X3)&v3_rlvect_1(X3))&v4_rlvect_1(X3))&v5_rlvect_1(X3))&v6_rlvect_1(X3))&v3_vectsp_2(X3,X1))&v5_vectsp_2(X3,X1))&l1_vectsp_2(X3,X1))=>((m1_rmod_2(X3,X1,X2)&m1_rmod_2(X2,X1,X3))=>X3=X2))))), inference(fof_simplification,[status(thm)],[c_0_0])).
fof(c_0_6, plain, (![X1]:![X2]:(((((((((((((((((~v3_struct_0(X1)&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&l3_vectsp_1(X1))&~v3_struct_0(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))=>![X3]:(m1_rmod_2(X3,X1,X2)=>((((((~v3_struct_0(X3)&v3_rlvect_1(X3))&v4_rlvect_1(X3))&v5_rlvect_1(X3))&v6_rlvect_1(X3))&v5_vectsp_2(X3,X1))&l1_vectsp_2(X3,X1))))), inference(fof_simplification,[status(thm)],[c_0_1])).
fof(c_0_7, plain, (![X1]:(((((((((((((~v3_struct_0(X1)&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v7_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&~v10_vectsp_1(X1))&v2_vectsp_2(X1))&l3_vectsp_1(X1))=>![X2]:(((((((~v3_struct_0(X2)&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(r1_tarski(X3,X4)=>m1_rmod_2(k1_rmod_5(X1,X2,X3),X1,k1_rmod_5(X1,X2,X4)))))))), inference(fof_simplification,[status(thm)],[c_0_2])).
fof(c_0_8, negated_conjecture, (~(![X1]:(((((((((((((~v3_struct_0(X1)&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v7_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&~v10_vectsp_1(X1))&v2_vectsp_2(X1))&l3_vectsp_1(X1))=>![X2]:((((((((~v3_struct_0(X2)&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v3_vectsp_2(X2,X1))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>((k1_rmod_5(X1,X2,X3)=X2&r1_tarski(X3,X4))=>k1_rmod_5(X1,X2,X4)=X2))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_3])])).
fof(c_0_9, plain, (![X4]:![X5]:![X6]:((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|~l3_vectsp_1(X4))|((((((((v3_struct_0(X5)|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v3_vectsp_2(X5,X4))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4))|((((((((v3_struct_0(X6)|~v3_rlvect_1(X6))|~v4_rlvect_1(X6))|~v5_rlvect_1(X6))|~v6_rlvect_1(X6))|~v3_vectsp_2(X6,X4))|~v5_vectsp_2(X6,X4))|~l1_vectsp_2(X6,X4))|((~m1_rmod_2(X6,X4,X5)|~m1_rmod_2(X5,X4,X6))|X6=X5))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
fof(c_0_10, plain, (![X4]:![X5]:![X6]:((((((((~v3_struct_0(X6)|~m1_rmod_2(X6,X4,X5))|((((((((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|~l3_vectsp_1(X4))|v3_struct_0(X5))|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4)))&((v3_rlvect_1(X6)|~m1_rmod_2(X6,X4,X5))|((((((((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|~l3_vectsp_1(X4))|v3_struct_0(X5))|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4))))&((v4_rlvect_1(X6)|~m1_rmod_2(X6,X4,X5))|((((((((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|~l3_vectsp_1(X4))|v3_struct_0(X5))|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4))))&((v5_rlvect_1(X6)|~m1_rmod_2(X6,X4,X5))|((((((((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|~l3_vectsp_1(X4))|v3_struct_0(X5))|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4))))&((v6_rlvect_1(X6)|~m1_rmod_2(X6,X4,X5))|((((((((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|~l3_vectsp_1(X4))|v3_struct_0(X5))|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4))))&((v5_vectsp_2(X6,X4)|~m1_rmod_2(X6,X4,X5))|((((((((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|~l3_vectsp_1(X4))|v3_struct_0(X5))|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4))))&((l1_vectsp_2(X6,X4)|~m1_rmod_2(X6,X4,X5))|((((((((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|~l3_vectsp_1(X4))|v3_struct_0(X5))|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
fof(c_0_11, plain, (![X5]:![X6]:![X7]:![X8]:(((((((((((((v3_struct_0(X5)|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v4_group_1(X5))|~v7_group_1(X5))|~v6_vectsp_1(X5))|~v7_vectsp_1(X5))|~v8_vectsp_1(X5))|v10_vectsp_1(X5))|~v2_vectsp_2(X5))|~l3_vectsp_1(X5))|(((((((v3_struct_0(X6)|~v3_rlvect_1(X6))|~v4_rlvect_1(X6))|~v5_rlvect_1(X6))|~v6_rlvect_1(X6))|~v5_vectsp_2(X6,X5))|~l1_vectsp_2(X6,X5))|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|(~r1_tarski(X7,X8)|m1_rmod_2(k1_rmod_5(X5,X6,X7),X5,k1_rmod_5(X5,X6,X8)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
fof(c_0_12, negated_conjecture, ((((((((((((((~v3_struct_0(esk1_0)&v3_rlvect_1(esk1_0))&v4_rlvect_1(esk1_0))&v5_rlvect_1(esk1_0))&v6_rlvect_1(esk1_0))&v4_group_1(esk1_0))&v7_group_1(esk1_0))&v6_vectsp_1(esk1_0))&v7_vectsp_1(esk1_0))&v8_vectsp_1(esk1_0))&~v10_vectsp_1(esk1_0))&v2_vectsp_2(esk1_0))&l3_vectsp_1(esk1_0))&((((((((~v3_struct_0(esk2_0)&v3_rlvect_1(esk2_0))&v4_rlvect_1(esk2_0))&v5_rlvect_1(esk2_0))&v6_rlvect_1(esk2_0))&v3_vectsp_2(esk2_0,esk1_0))&v5_vectsp_2(esk2_0,esk1_0))&l1_vectsp_2(esk2_0,esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&((k1_rmod_5(esk1_0,esk2_0,esk3_0)=esk2_0&r1_tarski(esk3_0,esk4_0))&k1_rmod_5(esk1_0,esk2_0,esk4_0)!=esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
cnf(c_0_13,plain,(X1=X2|v3_struct_0(X1)|v3_struct_0(X2)|v3_struct_0(X3)|~m1_rmod_2(X2,X3,X1)|~m1_rmod_2(X1,X3,X2)|~l1_vectsp_2(X1,X3)|~v5_vectsp_2(X1,X3)|~v3_vectsp_2(X1,X3)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l1_vectsp_2(X2,X3)|~v5_vectsp_2(X2,X3)|~v3_vectsp_2(X2,X3)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~l3_vectsp_1(X3)|~v8_vectsp_1(X3)|~v7_vectsp_1(X3)|~v6_vectsp_1(X3)|~v4_group_1(X3)|~v6_rlvect_1(X3)|~v5_rlvect_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14,plain,(v3_struct_0(X1)|v3_struct_0(X2)|l1_vectsp_2(X3,X2)|~l1_vectsp_2(X1,X2)|~v5_vectsp_2(X1,X2)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l3_vectsp_1(X2)|~v8_vectsp_1(X2)|~v7_vectsp_1(X2)|~v6_vectsp_1(X2)|~v4_group_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~m1_rmod_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15,plain,(v3_struct_0(X1)|v3_struct_0(X2)|v5_vectsp_2(X3,X2)|~l1_vectsp_2(X1,X2)|~v5_vectsp_2(X1,X2)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l3_vectsp_1(X2)|~v8_vectsp_1(X2)|~v7_vectsp_1(X2)|~v6_vectsp_1(X2)|~v4_group_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~m1_rmod_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16,plain,(v3_struct_0(X1)|v3_struct_0(X2)|v6_rlvect_1(X3)|~l1_vectsp_2(X1,X2)|~v5_vectsp_2(X1,X2)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l3_vectsp_1(X2)|~v8_vectsp_1(X2)|~v7_vectsp_1(X2)|~v6_vectsp_1(X2)|~v4_group_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~m1_rmod_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17,plain,(v3_struct_0(X1)|v3_struct_0(X2)|v5_rlvect_1(X3)|~l1_vectsp_2(X1,X2)|~v5_vectsp_2(X1,X2)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l3_vectsp_1(X2)|~v8_vectsp_1(X2)|~v7_vectsp_1(X2)|~v6_vectsp_1(X2)|~v4_group_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~m1_rmod_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18,plain,(v3_struct_0(X1)|v3_struct_0(X2)|v4_rlvect_1(X3)|~l1_vectsp_2(X1,X2)|~v5_vectsp_2(X1,X2)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l3_vectsp_1(X2)|~v8_vectsp_1(X2)|~v7_vectsp_1(X2)|~v6_vectsp_1(X2)|~v4_group_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~m1_rmod_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19,plain,(v3_struct_0(X1)|v3_struct_0(X2)|v3_rlvect_1(X3)|~l1_vectsp_2(X1,X2)|~v5_vectsp_2(X1,X2)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l3_vectsp_1(X2)|~v8_vectsp_1(X2)|~v7_vectsp_1(X2)|~v6_vectsp_1(X2)|~v4_group_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~m1_rmod_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20,plain,(v3_struct_0(X1)|v3_struct_0(X2)|~l1_vectsp_2(X1,X2)|~v5_vectsp_2(X1,X2)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l3_vectsp_1(X2)|~v8_vectsp_1(X2)|~v7_vectsp_1(X2)|~v6_vectsp_1(X2)|~v4_group_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~m1_rmod_2(X3,X2,X1)|~v3_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21,plain,(m1_rmod_2(k1_rmod_5(X1,X2,X3),X1,k1_rmod_5(X1,X2,X4))|v3_struct_0(X2)|v10_vectsp_1(X1)|v3_struct_0(X1)|~r1_tarski(X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_vectsp_2(X2,X1)|~v5_vectsp_2(X2,X1)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~l3_vectsp_1(X1)|~v2_vectsp_2(X1)|~v8_vectsp_1(X1)|~v7_vectsp_1(X1)|~v6_vectsp_1(X1)|~v7_group_1(X1)|~v4_group_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22,negated_conjecture,(k1_rmod_5(esk1_0,esk2_0,esk3_0)=esk2_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23,negated_conjecture,(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24,negated_conjecture,(l1_vectsp_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25,negated_conjecture,(v5_vectsp_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_26,negated_conjecture,(l3_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27,negated_conjecture,(v2_vectsp_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_28,negated_conjecture,(v8_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_29,negated_conjecture,(v7_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30,negated_conjecture,(v6_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_31,negated_conjecture,(v7_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_32,negated_conjecture,(v4_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_33,negated_conjecture,(v6_rlvect_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_34,negated_conjecture,(v6_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_35,negated_conjecture,(v5_rlvect_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_36,negated_conjecture,(v5_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37,negated_conjecture,(v4_rlvect_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_38,negated_conjecture,(v4_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_39,negated_conjecture,(v3_rlvect_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_40,negated_conjecture,(v3_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_41,negated_conjecture,(~v10_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42,negated_conjecture,(~v3_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_43,negated_conjecture,(~v3_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_44, plain, (![X1]:![X2]:![X3]:(((((((((((((((((((((~v3_struct_0(X1)&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v4_group_1(X1))&v7_group_1(X1))&v6_vectsp_1(X1))&v7_vectsp_1(X1))&v8_vectsp_1(X1))&~v10_vectsp_1(X1))&v2_vectsp_2(X1))&l3_vectsp_1(X1))&~v3_struct_0(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v5_rlvect_1(X2))&v6_rlvect_1(X2))&v5_vectsp_2(X2,X1))&l1_vectsp_2(X2,X1))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(v3_vectsp_2(k1_rmod_5(X1,X2,X3),X1)&m1_rmod_2(k1_rmod_5(X1,X2,X3),X1,X2)))), inference(fof_simplification,[status(thm)],[c_0_4])).
cnf(c_0_45,plain,(X1=X2|v3_struct_0(X2)|v3_struct_0(X3)|~m1_rmod_2(X2,X3,X1)|~m1_rmod_2(X1,X3,X2)|~l1_vectsp_2(X2,X3)|~v5_vectsp_2(X2,X3)|~v3_vectsp_2(X2,X3)|~v3_vectsp_2(X1,X3)|~l3_vectsp_1(X3)|~v8_vectsp_1(X3)|~v7_vectsp_1(X3)|~v6_vectsp_1(X3)|~v4_group_1(X3)|~v6_rlvect_1(X3)|~v6_rlvect_1(X2)|~v5_rlvect_1(X3)|~v5_rlvect_1(X2)|~v4_rlvect_1(X3)|~v4_rlvect_1(X2)|~v3_rlvect_1(X3)|~v3_rlvect_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_46,negated_conjecture,(m1_rmod_2(esk2_0,esk1_0,k1_rmod_5(esk1_0,esk2_0,X1))|~r1_tarski(esk3_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])]), c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_47,negated_conjecture,(v3_vectsp_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_48, plain, (![X4]:![X5]:![X6]:((v3_vectsp_2(k1_rmod_5(X4,X5,X6),X4)|((((((((((((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v7_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|v10_vectsp_1(X4))|~v2_vectsp_2(X4))|~l3_vectsp_1(X4))|v3_struct_0(X5))|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))))&(m1_rmod_2(k1_rmod_5(X4,X5,X6),X4,X5)|((((((((((((((((((((v3_struct_0(X4)|~v3_rlvect_1(X4))|~v4_rlvect_1(X4))|~v5_rlvect_1(X4))|~v6_rlvect_1(X4))|~v4_group_1(X4))|~v7_group_1(X4))|~v6_vectsp_1(X4))|~v7_vectsp_1(X4))|~v8_vectsp_1(X4))|v10_vectsp_1(X4))|~v2_vectsp_2(X4))|~l3_vectsp_1(X4))|v3_struct_0(X5))|~v3_rlvect_1(X5))|~v4_rlvect_1(X5))|~v5_rlvect_1(X5))|~v6_rlvect_1(X5))|~v5_vectsp_2(X5,X4))|~l1_vectsp_2(X5,X4))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_44])])])).
cnf(c_0_49,negated_conjecture,(k1_rmod_5(esk1_0,esk2_0,X1)=esk2_0|~m1_rmod_2(k1_rmod_5(esk1_0,esk2_0,X1),esk1_0,esk2_0)|~r1_tarski(esk3_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v3_vectsp_2(k1_rmod_5(esk1_0,esk2_0,X1),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_24]), c_0_25]), c_0_47]), c_0_26]), c_0_28]), c_0_29]), c_0_30]), c_0_32]), c_0_34]), c_0_33]), c_0_36]), c_0_35]), c_0_38]), c_0_37]), c_0_40]), c_0_39])]), c_0_43]), c_0_42])).
cnf(c_0_50,plain,(v3_struct_0(X2)|v10_vectsp_1(X3)|v3_struct_0(X3)|m1_rmod_2(k1_rmod_5(X3,X2,X1),X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_vectsp_2(X2,X3)|~v5_vectsp_2(X2,X3)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~l3_vectsp_1(X3)|~v2_vectsp_2(X3)|~v8_vectsp_1(X3)|~v7_vectsp_1(X3)|~v6_vectsp_1(X3)|~v7_group_1(X3)|~v4_group_1(X3)|~v6_rlvect_1(X3)|~v5_rlvect_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_51,negated_conjecture,(k1_rmod_5(esk1_0,esk2_0,X1)=esk2_0|~r1_tarski(esk3_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v3_vectsp_2(k1_rmod_5(esk1_0,esk2_0,X1),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_34]), c_0_33]), c_0_36]), c_0_35]), c_0_38]), c_0_37]), c_0_40]), c_0_39])]), c_0_41]), c_0_43]), c_0_42])).
cnf(c_0_52,plain,(v3_struct_0(X2)|v10_vectsp_1(X3)|v3_struct_0(X3)|v3_vectsp_2(k1_rmod_5(X3,X2,X1),X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_vectsp_2(X2,X3)|~v5_vectsp_2(X2,X3)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~l3_vectsp_1(X3)|~v2_vectsp_2(X3)|~v8_vectsp_1(X3)|~v7_vectsp_1(X3)|~v6_vectsp_1(X3)|~v7_group_1(X3)|~v4_group_1(X3)|~v6_rlvect_1(X3)|~v5_rlvect_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_53,negated_conjecture,(k1_rmod_5(esk1_0,esk2_0,esk4_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_54,negated_conjecture,(k1_rmod_5(esk1_0,esk2_0,X1)=esk2_0|~r1_tarski(esk3_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_34]), c_0_33]), c_0_36]), c_0_35]), c_0_38]), c_0_37]), c_0_40]), c_0_39])]), c_0_41]), c_0_43]), c_0_42])).
cnf(c_0_55,negated_conjecture,(r1_tarski(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_56,negated_conjecture,(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_57,negated_conjecture,($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56])]), ['proof']).
# SZS output end CNFRefutation.
EOF
