% Problem    : PUZ073+1 : TPTP v6.1.0. Released v3.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n113.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 05:23:21 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___207_C18_F1_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X2]:![X1]:((patient(X2)&oxygen(X1))=>~(adj(X2,X1)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49932)).
fof(c_0_1, axiom, (![X1]:![X2]:((patient(X2)&patient(X1))=>X2=X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof44978)).
fof(c_0_2, axiom, (![X1]:![X2]:((oxygen(X2)&oxygen(X1))=>X2=X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof44977)).
fof(c_0_3, axiom, (![X3]:![X4]:(edge(X4,X3)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60))|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65))|(X4=cell56&X3=cell66))|(X4=cell57&X3=cell67))|(X4=cell58&X3=cell68))|(X4=cell59&X3=cell69))|(X4=cell60&X3=cell70))|(X4=cell61&X3=cell71))|(X4=cell62&X3=cell72))|(X4=cell63&X3=cell73))|(X4=cell64&X3=cell74))|(X4=cell65&X3=cell75))|(X4=cell66&X3=cell76))|(X4=cell67&X3=cell77))|(X4=cell68&X3=cell78))|(X4=cell69&X3=cell79))|(X4=cell70&X3=cell80))|(X4=cell71&X3=cell81))|(X4=cell72&X3=cell82))|(X4=cell73&X3=cell83))|(X4=cell74&X3=cell84))|(X4=cell75&X3=cell85))|(X4=cell76&X3=cell86))|(X4=cell77&X3=cell87))|(X4=cell78&X3=cell88))|(X4=cell79&X3=cell89))|(X4=cell80&X3=cell90))|(X4=cell81&X3=cell91))|(X4=cell82&X3=cell92))|(X4=cell83&X3=cell93))|(X4=cell84&X3=cell94))|(X4=cell85&X3=cell95))|(X4=cell86&X3=cell96))|(X4=cell87&X3=cell97))|(X4=cell88&X3=cell98))|(X4=cell89&X3=cell99))|(X4=cell90&X3=cell100)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof44979)).
fof(c_0_4, axiom, (![X1]:![X2]:(adj(X2,X1)<=>(edge(X2,X1)|edge(X1,X2)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof44980)).
fof(c_0_5, axiom, (((((patient(cell12)|patient(cell11))|patient(cell2))|patient(cell13))|patient(cell22))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof44975)).
fof(c_0_6, axiom, (((((oxygen(cell89)|oxygen(cell88))|oxygen(cell79))|oxygen(cell90))|oxygen(cell99))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof44976)).
fof(c_0_7, axiom, (cell99!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49918)).
fof(c_0_8, axiom, (cell96!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45368)).
fof(c_0_9, axiom, (cell97!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45272)).
fof(c_0_10, axiom, (cell98!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45175)).
fof(c_0_11, axiom, (cell100!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45077)).
fof(c_0_12, axiom, (cell89!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49892)).
fof(c_0_13, axiom, (cell88!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49900)).
fof(c_0_14, axiom, (cell90!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49913)).
fof(c_0_15, axiom, (cell86!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46012)).
fof(c_0_16, axiom, (cell87!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45923)).
fof(c_0_17, axiom, (cell89!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49889)).
fof(c_0_18, axiom, (cell88!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49897)).
fof(c_0_19, axiom, (cell90!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49910)).
fof(c_0_20, axiom, (cell86!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46009)).
fof(c_0_21, axiom, (cell87!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45920)).
fof(c_0_22, axiom, (cell81!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46358)).
fof(c_0_23, axiom, (cell82!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46442)).
fof(c_0_24, axiom, (cell83!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46273)).
fof(c_0_25, axiom, (cell84!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46187)).
fof(c_0_26, axiom, (cell85!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46100)).
fof(c_0_27, axiom, (cell96!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45365)).
fof(c_0_28, axiom, (cell97!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45269)).
fof(c_0_29, axiom, (cell98!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45172)).
fof(c_0_30, axiom, (cell100!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45074)).
fof(c_0_31, axiom, (cell91!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45742)).
fof(c_0_32, axiom, (cell92!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45833)).
fof(c_0_33, axiom, (cell93!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45650)).
fof(c_0_34, axiom, (cell94!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45557)).
fof(c_0_35, axiom, (cell95!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45463)).
fof(c_0_36, axiom, (cell81!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46355)).
fof(c_0_37, axiom, (cell82!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46439)).
fof(c_0_38, axiom, (cell83!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46270)).
fof(c_0_39, axiom, (cell84!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46184)).
fof(c_0_40, axiom, (cell85!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46097)).
fof(c_0_41, axiom, (cell79!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49907)).
fof(c_0_42, axiom, (cell76!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46768)).
fof(c_0_43, axiom, (cell77!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46688)).
fof(c_0_44, axiom, (cell78!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46607)).
fof(c_0_45, axiom, (cell80!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46525)).
fof(c_0_46, axiom, (cell91!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45739)).
fof(c_0_47, axiom, (cell92!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45830)).
fof(c_0_48, axiom, (cell93!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45647)).
fof(c_0_49, axiom, (cell94!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45554)).
fof(c_0_50, axiom, (cell95!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45460)).
fof(c_0_51, axiom, (cell76!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46765)).
fof(c_0_52, axiom, (cell77!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46685)).
fof(c_0_53, axiom, (cell78!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46604)).
fof(c_0_54, axiom, (cell80!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46522)).
fof(c_0_55, axiom, (cell71!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47078)).
fof(c_0_56, axiom, (cell72!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47153)).
fof(c_0_57, axiom, (cell73!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47002)).
fof(c_0_58, axiom, (cell74!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46925)).
fof(c_0_59, axiom, (cell75!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46847)).
fof(c_0_60, axiom, (cell71!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47075)).
fof(c_0_61, axiom, (cell72!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47150)).
fof(c_0_62, axiom, (cell73!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46999)).
fof(c_0_63, axiom, (cell74!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46922)).
fof(c_0_64, axiom, (cell75!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46844)).
fof(c_0_65, axiom, (cell66!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47513)).
fof(c_0_66, axiom, (cell67!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47443)).
fof(c_0_67, axiom, (cell68!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47372)).
fof(c_0_68, axiom, (cell69!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47300)).
fof(c_0_69, axiom, (cell70!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47227)).
fof(c_0_70, axiom, (cell66!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47510)).
fof(c_0_71, axiom, (cell67!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47440)).
fof(c_0_72, axiom, (cell68!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47369)).
fof(c_0_73, axiom, (cell69!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47297)).
fof(c_0_74, axiom, (cell70!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47224)).
fof(c_0_75, axiom, (cell61!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47783)).
fof(c_0_76, axiom, (cell62!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47848)).
fof(c_0_77, axiom, (cell63!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47717)).
fof(c_0_78, axiom, (cell64!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47650)).
fof(c_0_79, axiom, (cell65!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47582)).
fof(c_0_80, axiom, (cell79!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49904)).
fof(c_0_81, axiom, (cell61!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47780)).
fof(c_0_82, axiom, (cell62!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47845)).
fof(c_0_83, axiom, (cell63!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47714)).
fof(c_0_84, axiom, (cell64!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47647)).
fof(c_0_85, axiom, (cell65!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47579)).
fof(c_0_86, axiom, (cell56!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48158)).
fof(c_0_87, axiom, (cell57!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48098)).
fof(c_0_88, axiom, (cell58!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48037)).
fof(c_0_89, axiom, (cell59!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47975)).
fof(c_0_90, axiom, (cell60!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47912)).
fof(c_0_91, axiom, (cell56!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48155)).
fof(c_0_92, axiom, (cell57!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48095)).
fof(c_0_93, axiom, (cell58!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48034)).
fof(c_0_94, axiom, (cell59!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47972)).
fof(c_0_95, axiom, (cell60!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47909)).
fof(c_0_96, axiom, (cell51!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48388)).
fof(c_0_97, axiom, (cell52!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48443)).
fof(c_0_98, axiom, (cell53!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48332)).
fof(c_0_99, axiom, (cell54!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48275)).
fof(c_0_100, axiom, (cell55!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48217)).
fof(c_0_101, axiom, (cell51!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48385)).
fof(c_0_102, axiom, (cell52!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48440)).
fof(c_0_103, axiom, (cell53!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48329)).
fof(c_0_104, axiom, (cell54!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48272)).
fof(c_0_105, axiom, (cell55!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48214)).
fof(c_0_106, axiom, (cell46!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48703)).
fof(c_0_107, axiom, (cell47!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48653)).
fof(c_0_108, axiom, (cell48!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48602)).
fof(c_0_109, axiom, (cell49!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48550)).
fof(c_0_110, axiom, (cell50!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48497)).
fof(c_0_111, axiom, (cell46!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48700)).
fof(c_0_112, axiom, (cell47!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48650)).
fof(c_0_113, axiom, (cell48!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48599)).
fof(c_0_114, axiom, (cell49!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48547)).
fof(c_0_115, axiom, (cell50!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48494)).
fof(c_0_116, axiom, (cell41!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48893)).
fof(c_0_117, axiom, (cell42!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48938)).
fof(c_0_118, axiom, (cell43!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48847)).
fof(c_0_119, axiom, (cell44!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48800)).
fof(c_0_120, axiom, (cell45!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48752)).
fof(c_0_121, axiom, (cell41!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48890)).
fof(c_0_122, axiom, (cell42!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48935)).
fof(c_0_123, axiom, (cell43!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48844)).
fof(c_0_124, axiom, (cell44!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48797)).
fof(c_0_125, axiom, (cell45!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48749)).
fof(c_0_126, axiom, (cell36!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49148)).
fof(c_0_127, axiom, (cell37!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49108)).
fof(c_0_128, axiom, (cell38!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49067)).
fof(c_0_129, axiom, (cell39!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49025)).
fof(c_0_130, axiom, (cell40!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48982)).
fof(c_0_131, axiom, (cell36!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49145)).
fof(c_0_132, axiom, (cell37!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49105)).
fof(c_0_133, axiom, (cell38!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49064)).
fof(c_0_134, axiom, (cell39!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49022)).
fof(c_0_135, axiom, (cell40!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48979)).
fof(c_0_136, axiom, (cell31!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49298)).
fof(c_0_137, axiom, (cell32!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49333)).
fof(c_0_138, axiom, (cell33!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49262)).
fof(c_0_139, axiom, (cell34!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49225)).
fof(c_0_140, axiom, (cell35!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49187)).
fof(c_0_141, axiom, (cell31!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49295)).
fof(c_0_142, axiom, (cell32!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49330)).
fof(c_0_143, axiom, (cell33!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49259)).
fof(c_0_144, axiom, (cell34!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49222)).
fof(c_0_145, axiom, (cell35!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49184)).
fof(c_0_146, axiom, (cell26!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49493)).
fof(c_0_147, axiom, (cell27!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49463)).
fof(c_0_148, axiom, (cell28!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49432)).
fof(c_0_149, axiom, (cell29!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49400)).
fof(c_0_150, axiom, (cell30!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49367)).
fof(c_0_151, axiom, (cell26!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49490)).
fof(c_0_152, axiom, (cell27!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49460)).
fof(c_0_153, axiom, (cell28!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49429)).
fof(c_0_154, axiom, (cell29!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49397)).
fof(c_0_155, axiom, (cell30!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49364)).
fof(c_0_156, axiom, (cell2!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49929)).
fof(c_0_157, axiom, (cell21!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49603)).
fof(c_0_158, axiom, (cell23!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49577)).
fof(c_0_159, axiom, (cell24!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49550)).
fof(c_0_160, axiom, (cell25!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49522)).
fof(c_0_161, axiom, (cell99!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49920)).
fof(c_0_162, axiom, (cell21!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49600)).
fof(c_0_163, axiom, (cell23!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49574)).
fof(c_0_164, axiom, (cell24!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49547)).
fof(c_0_165, axiom, (cell25!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49519)).
fof(c_0_166, axiom, (cell16!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49718)).
fof(c_0_167, axiom, (cell17!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49697)).
fof(c_0_168, axiom, (cell18!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49675)).
fof(c_0_169, axiom, (cell19!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49652)).
fof(c_0_170, axiom, (cell20!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49628)).
fof(c_0_171, axiom, (cell16!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49715)).
fof(c_0_172, axiom, (cell17!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49694)).
fof(c_0_173, axiom, (cell18!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49672)).
fof(c_0_174, axiom, (cell19!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49649)).
fof(c_0_175, axiom, (cell20!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49625)).
fof(c_0_176, axiom, (cell12!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49922)).
fof(c_0_177, axiom, (cell11!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49925)).
fof(c_0_178, axiom, (cell2!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49928)).
fof(c_0_179, axiom, (cell14!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49757)).
fof(c_0_180, axiom, (cell15!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49738)).
fof(c_0_181, axiom, (cell99!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49916)).
fof(c_0_182, axiom, (cell99!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49917)).
fof(c_0_183, axiom, (cell99!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49919)).
fof(c_0_184, axiom, (cell14!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49754)).
fof(c_0_185, axiom, (cell15!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49735)).
fof(c_0_186, axiom, (cell6!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49837)).
fof(c_0_187, axiom, (cell7!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49823)).
fof(c_0_188, axiom, (cell8!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49808)).
fof(c_0_189, axiom, (cell9!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49792)).
fof(c_0_190, axiom, (cell10!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49775)).
fof(c_0_191, axiom, (cell6!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49834)).
fof(c_0_192, axiom, (cell7!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49820)).
fof(c_0_193, axiom, (cell8!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49805)).
fof(c_0_194, axiom, (cell9!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49789)).
fof(c_0_195, axiom, (cell10!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49772)).
fof(c_0_196, axiom, (cell1!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49883)).
fof(c_0_197, axiom, (cell3!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49873)).
fof(c_0_198, axiom, (cell4!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49862)).
fof(c_0_199, axiom, (cell5!=cell2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49850)).
fof(c_0_200, axiom, (cell1!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49880)).
fof(c_0_201, axiom, (cell3!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49870)).
fof(c_0_202, axiom, (cell4!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49859)).
fof(c_0_203, axiom, (cell5!=cell99), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49847)).
fof(c_0_204, axiom, (cell12!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49921)).
fof(c_0_205, axiom, (cell89!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49891)).
fof(c_0_206, axiom, (cell98!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45174)).
fof(c_0_207, axiom, (cell100!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45076)).
fof(c_0_208, axiom, (cell3!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49872)).
fof(c_0_209, axiom, (cell4!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49861)).
fof(c_0_210, axiom, (cell6!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49836)).
fof(c_0_211, axiom, (cell7!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49822)).
fof(c_0_212, axiom, (cell8!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49807)).
fof(c_0_213, axiom, (cell9!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49791)).
fof(c_0_214, axiom, (cell12!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49924)).
fof(c_0_215, axiom, (cell89!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49894)).
fof(c_0_216, axiom, (cell98!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45177)).
fof(c_0_217, axiom, (cell100!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45079)).
fof(c_0_218, axiom, (cell3!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49875)).
fof(c_0_219, axiom, (cell4!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49864)).
fof(c_0_220, axiom, (cell6!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49839)).
fof(c_0_221, axiom, (cell7!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49825)).
fof(c_0_222, axiom, (cell8!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49810)).
fof(c_0_223, axiom, (cell9!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49794)).
fof(c_0_224, axiom, (cell6!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49838)).
fof(c_0_225, axiom, (cell12!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49923)).
fof(c_0_226, axiom, (cell89!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49893)).
fof(c_0_227, axiom, (cell98!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45176)).
fof(c_0_228, axiom, (cell100!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45078)).
fof(c_0_229, axiom, (cell89!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49890)).
fof(c_0_230, axiom, (cell98!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45173)).
fof(c_0_231, axiom, (cell100!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45075)).
fof(c_0_232, axiom, (cell3!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49871)).
fof(c_0_233, axiom, (cell4!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49860)).
fof(c_0_234, axiom, (cell6!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49835)).
fof(c_0_235, axiom, (cell7!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49821)).
fof(c_0_236, axiom, (cell8!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49806)).
fof(c_0_237, axiom, (cell9!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49790)).
fof(c_0_238, axiom, (cell89!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49888)).
fof(c_0_239, axiom, (cell88!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49896)).
fof(c_0_240, axiom, (cell86!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46008)).
fof(c_0_241, axiom, (cell87!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45919)).
fof(c_0_242, axiom, (cell96!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45364)).
fof(c_0_243, axiom, (cell97!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45268)).
fof(c_0_244, axiom, (cell98!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45171)).
fof(c_0_245, axiom, (cell100!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45073)).
fof(c_0_246, axiom, (cell81!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46354)).
fof(c_0_247, axiom, (cell82!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46438)).
fof(c_0_248, axiom, (cell83!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46269)).
fof(c_0_249, axiom, (cell84!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46183)).
fof(c_0_250, axiom, (cell85!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46096)).
fof(c_0_251, axiom, (cell91!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45738)).
fof(c_0_252, axiom, (cell92!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45829)).
fof(c_0_253, axiom, (cell93!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45646)).
fof(c_0_254, axiom, (cell94!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45553)).
fof(c_0_255, axiom, (cell95!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45459)).
fof(c_0_256, axiom, (cell79!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49903)).
fof(c_0_257, axiom, (cell76!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46764)).
fof(c_0_258, axiom, (cell77!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46684)).
fof(c_0_259, axiom, (cell78!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46603)).
fof(c_0_260, axiom, (cell80!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46521)).
fof(c_0_261, axiom, (cell71!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47074)).
fof(c_0_262, axiom, (cell72!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47149)).
fof(c_0_263, axiom, (cell73!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46998)).
fof(c_0_264, axiom, (cell74!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46921)).
fof(c_0_265, axiom, (cell75!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46843)).
fof(c_0_266, axiom, (cell66!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47509)).
fof(c_0_267, axiom, (cell67!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47439)).
fof(c_0_268, axiom, (cell68!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47368)).
fof(c_0_269, axiom, (cell69!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47296)).
fof(c_0_270, axiom, (cell70!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47223)).
fof(c_0_271, axiom, (cell61!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47779)).
fof(c_0_272, axiom, (cell62!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47844)).
fof(c_0_273, axiom, (cell63!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47713)).
fof(c_0_274, axiom, (cell64!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47646)).
fof(c_0_275, axiom, (cell65!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47578)).
fof(c_0_276, axiom, (cell56!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48154)).
fof(c_0_277, axiom, (cell57!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48094)).
fof(c_0_278, axiom, (cell58!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48033)).
fof(c_0_279, axiom, (cell59!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47971)).
fof(c_0_280, axiom, (cell60!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47908)).
fof(c_0_281, axiom, (cell51!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48384)).
fof(c_0_282, axiom, (cell52!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48439)).
fof(c_0_283, axiom, (cell53!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48328)).
fof(c_0_284, axiom, (cell54!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48271)).
fof(c_0_285, axiom, (cell55!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48213)).
fof(c_0_286, axiom, (cell46!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48699)).
fof(c_0_287, axiom, (cell47!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48649)).
fof(c_0_288, axiom, (cell48!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48598)).
fof(c_0_289, axiom, (cell49!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48546)).
fof(c_0_290, axiom, (cell50!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48493)).
fof(c_0_291, axiom, (cell41!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48889)).
fof(c_0_292, axiom, (cell42!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48934)).
fof(c_0_293, axiom, (cell43!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48843)).
fof(c_0_294, axiom, (cell44!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48796)).
fof(c_0_295, axiom, (cell45!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48748)).
fof(c_0_296, axiom, (cell36!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49144)).
fof(c_0_297, axiom, (cell37!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49104)).
fof(c_0_298, axiom, (cell38!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49063)).
fof(c_0_299, axiom, (cell39!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49021)).
fof(c_0_300, axiom, (cell40!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48978)).
fof(c_0_301, axiom, (cell31!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49294)).
fof(c_0_302, axiom, (cell32!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49329)).
fof(c_0_303, axiom, (cell33!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49258)).
fof(c_0_304, axiom, (cell34!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49221)).
fof(c_0_305, axiom, (cell35!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49183)).
fof(c_0_306, axiom, (cell26!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49489)).
fof(c_0_307, axiom, (cell27!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49459)).
fof(c_0_308, axiom, (cell28!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49428)).
fof(c_0_309, axiom, (cell29!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49396)).
fof(c_0_310, axiom, (cell30!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49363)).
fof(c_0_311, axiom, (cell90!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49915)).
fof(c_0_312, axiom, (cell21!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49599)).
fof(c_0_313, axiom, (cell23!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49573)).
fof(c_0_314, axiom, (cell24!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49546)).
fof(c_0_315, axiom, (cell25!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49518)).
fof(c_0_316, axiom, (cell16!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49714)).
fof(c_0_317, axiom, (cell17!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49693)).
fof(c_0_318, axiom, (cell18!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49671)).
fof(c_0_319, axiom, (cell19!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49648)).
fof(c_0_320, axiom, (cell20!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49624)).
fof(c_0_321, axiom, (cell90!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49911)).
fof(c_0_322, axiom, (cell90!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49912)).
fof(c_0_323, axiom, (cell90!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49914)).
fof(c_0_324, axiom, (cell14!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49753)).
fof(c_0_325, axiom, (cell15!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49734)).
fof(c_0_326, axiom, (cell6!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49833)).
fof(c_0_327, axiom, (cell7!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49819)).
fof(c_0_328, axiom, (cell8!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49804)).
fof(c_0_329, axiom, (cell9!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49788)).
fof(c_0_330, axiom, (cell10!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49771)).
fof(c_0_331, axiom, (cell1!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49879)).
fof(c_0_332, axiom, (cell3!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49869)).
fof(c_0_333, axiom, (cell4!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49858)).
fof(c_0_334, axiom, (cell5!=cell90), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49846)).
fof(c_0_335, axiom, (cell80!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46523)).
fof(c_0_336, axiom, (cell80!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46526)).
fof(c_0_337, axiom, (cell80!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46527)).
fof(c_0_338, axiom, (cell80!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46524)).
fof(c_0_339, axiom, (cell89!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49887)).
fof(c_0_340, axiom, (cell88!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49895)).
fof(c_0_341, axiom, (cell86!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46007)).
fof(c_0_342, axiom, (cell87!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45918)).
fof(c_0_343, axiom, (cell96!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45363)).
fof(c_0_344, axiom, (cell97!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45267)).
fof(c_0_345, axiom, (cell98!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45170)).
fof(c_0_346, axiom, (cell100!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45072)).
fof(c_0_347, axiom, (cell81!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46353)).
fof(c_0_348, axiom, (cell82!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46437)).
fof(c_0_349, axiom, (cell83!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46268)).
fof(c_0_350, axiom, (cell84!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46182)).
fof(c_0_351, axiom, (cell85!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46095)).
fof(c_0_352, axiom, (cell91!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45737)).
fof(c_0_353, axiom, (cell92!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45828)).
fof(c_0_354, axiom, (cell93!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45645)).
fof(c_0_355, axiom, (cell94!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45552)).
fof(c_0_356, axiom, (cell95!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45458)).
fof(c_0_357, axiom, (cell76!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46763)).
fof(c_0_358, axiom, (cell77!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46683)).
fof(c_0_359, axiom, (cell78!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46602)).
fof(c_0_360, axiom, (cell80!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46520)).
fof(c_0_361, axiom, (cell71!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47073)).
fof(c_0_362, axiom, (cell72!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47148)).
fof(c_0_363, axiom, (cell73!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46997)).
fof(c_0_364, axiom, (cell74!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46920)).
fof(c_0_365, axiom, (cell75!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46842)).
fof(c_0_366, axiom, (cell66!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47508)).
fof(c_0_367, axiom, (cell67!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47438)).
fof(c_0_368, axiom, (cell68!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47367)).
fof(c_0_369, axiom, (cell69!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47295)).
fof(c_0_370, axiom, (cell70!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47222)).
fof(c_0_371, axiom, (cell61!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47778)).
fof(c_0_372, axiom, (cell62!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47843)).
fof(c_0_373, axiom, (cell63!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47712)).
fof(c_0_374, axiom, (cell64!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47645)).
fof(c_0_375, axiom, (cell65!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47577)).
fof(c_0_376, axiom, (cell56!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48153)).
fof(c_0_377, axiom, (cell57!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48093)).
fof(c_0_378, axiom, (cell58!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48032)).
fof(c_0_379, axiom, (cell59!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47970)).
fof(c_0_380, axiom, (cell60!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47907)).
fof(c_0_381, axiom, (cell51!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48383)).
fof(c_0_382, axiom, (cell52!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48438)).
fof(c_0_383, axiom, (cell53!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48327)).
fof(c_0_384, axiom, (cell54!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48270)).
fof(c_0_385, axiom, (cell55!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48212)).
fof(c_0_386, axiom, (cell46!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48698)).
fof(c_0_387, axiom, (cell47!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48648)).
fof(c_0_388, axiom, (cell48!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48597)).
fof(c_0_389, axiom, (cell49!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48545)).
fof(c_0_390, axiom, (cell50!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48492)).
fof(c_0_391, axiom, (cell41!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48888)).
fof(c_0_392, axiom, (cell42!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48933)).
fof(c_0_393, axiom, (cell43!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48842)).
fof(c_0_394, axiom, (cell44!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48795)).
fof(c_0_395, axiom, (cell45!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48747)).
fof(c_0_396, axiom, (cell36!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49143)).
fof(c_0_397, axiom, (cell37!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49103)).
fof(c_0_398, axiom, (cell38!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49062)).
fof(c_0_399, axiom, (cell39!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49020)).
fof(c_0_400, axiom, (cell40!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48977)).
fof(c_0_401, axiom, (cell31!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49293)).
fof(c_0_402, axiom, (cell32!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49328)).
fof(c_0_403, axiom, (cell33!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49257)).
fof(c_0_404, axiom, (cell34!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49220)).
fof(c_0_405, axiom, (cell35!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49182)).
fof(c_0_406, axiom, (cell26!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49488)).
fof(c_0_407, axiom, (cell27!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49458)).
fof(c_0_408, axiom, (cell28!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49427)).
fof(c_0_409, axiom, (cell29!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49395)).
fof(c_0_410, axiom, (cell30!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49362)).
fof(c_0_411, axiom, (cell79!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49909)).
fof(c_0_412, axiom, (cell21!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49598)).
fof(c_0_413, axiom, (cell23!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49572)).
fof(c_0_414, axiom, (cell24!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49545)).
fof(c_0_415, axiom, (cell25!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49517)).
fof(c_0_416, axiom, (cell16!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49713)).
fof(c_0_417, axiom, (cell17!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49692)).
fof(c_0_418, axiom, (cell18!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49670)).
fof(c_0_419, axiom, (cell19!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49647)).
fof(c_0_420, axiom, (cell20!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49623)).
fof(c_0_421, axiom, (cell79!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49905)).
fof(c_0_422, axiom, (cell79!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49906)).
fof(c_0_423, axiom, (cell79!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49908)).
fof(c_0_424, axiom, (cell14!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49752)).
fof(c_0_425, axiom, (cell15!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49733)).
fof(c_0_426, axiom, (cell6!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49832)).
fof(c_0_427, axiom, (cell7!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49818)).
fof(c_0_428, axiom, (cell8!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49803)).
fof(c_0_429, axiom, (cell9!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49787)).
fof(c_0_430, axiom, (cell10!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49770)).
fof(c_0_431, axiom, (cell1!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49878)).
fof(c_0_432, axiom, (cell3!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49868)).
fof(c_0_433, axiom, (cell4!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49857)).
fof(c_0_434, axiom, (cell5!=cell79), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49845)).
fof(c_0_435, axiom, (cell78!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46605)).
fof(c_0_436, axiom, (cell78!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46608)).
fof(c_0_437, axiom, (cell78!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46609)).
fof(c_0_438, axiom, (cell78!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46606)).
fof(c_0_439, axiom, (cell69!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47298)).
fof(c_0_440, axiom, (cell69!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47301)).
fof(c_0_441, axiom, (cell69!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47302)).
fof(c_0_442, axiom, (cell69!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47299)).
fof(c_0_443, axiom, (cell89!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49886)).
fof(c_0_444, axiom, (cell86!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46006)).
fof(c_0_445, axiom, (cell87!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45917)).
fof(c_0_446, axiom, (cell96!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45362)).
fof(c_0_447, axiom, (cell97!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45266)).
fof(c_0_448, axiom, (cell98!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45169)).
fof(c_0_449, axiom, (cell100!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45071)).
fof(c_0_450, axiom, (cell81!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46352)).
fof(c_0_451, axiom, (cell82!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46436)).
fof(c_0_452, axiom, (cell83!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46267)).
fof(c_0_453, axiom, (cell84!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46181)).
fof(c_0_454, axiom, (cell85!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46094)).
fof(c_0_455, axiom, (cell91!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45736)).
fof(c_0_456, axiom, (cell92!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45827)).
fof(c_0_457, axiom, (cell93!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45644)).
fof(c_0_458, axiom, (cell94!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45551)).
fof(c_0_459, axiom, (cell95!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45457)).
fof(c_0_460, axiom, (cell76!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46762)).
fof(c_0_461, axiom, (cell77!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46682)).
fof(c_0_462, axiom, (cell78!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46601)).
fof(c_0_463, axiom, (cell80!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46519)).
fof(c_0_464, axiom, (cell71!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47072)).
fof(c_0_465, axiom, (cell72!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47147)).
fof(c_0_466, axiom, (cell73!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46996)).
fof(c_0_467, axiom, (cell74!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46919)).
fof(c_0_468, axiom, (cell75!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46841)).
fof(c_0_469, axiom, (cell66!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47507)).
fof(c_0_470, axiom, (cell67!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47437)).
fof(c_0_471, axiom, (cell68!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47366)).
fof(c_0_472, axiom, (cell69!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47294)).
fof(c_0_473, axiom, (cell70!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47221)).
fof(c_0_474, axiom, (cell61!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47777)).
fof(c_0_475, axiom, (cell62!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47842)).
fof(c_0_476, axiom, (cell63!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47711)).
fof(c_0_477, axiom, (cell64!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47644)).
fof(c_0_478, axiom, (cell65!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47576)).
fof(c_0_479, axiom, (cell56!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48152)).
fof(c_0_480, axiom, (cell57!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48092)).
fof(c_0_481, axiom, (cell58!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48031)).
fof(c_0_482, axiom, (cell59!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47969)).
fof(c_0_483, axiom, (cell60!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47906)).
fof(c_0_484, axiom, (cell51!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48382)).
fof(c_0_485, axiom, (cell52!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48437)).
fof(c_0_486, axiom, (cell53!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48326)).
fof(c_0_487, axiom, (cell54!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48269)).
fof(c_0_488, axiom, (cell55!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48211)).
fof(c_0_489, axiom, (cell46!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48697)).
fof(c_0_490, axiom, (cell47!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48647)).
fof(c_0_491, axiom, (cell48!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48596)).
fof(c_0_492, axiom, (cell49!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48544)).
fof(c_0_493, axiom, (cell50!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48491)).
fof(c_0_494, axiom, (cell41!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48887)).
fof(c_0_495, axiom, (cell42!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48932)).
fof(c_0_496, axiom, (cell43!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48841)).
fof(c_0_497, axiom, (cell44!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48794)).
fof(c_0_498, axiom, (cell45!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48746)).
fof(c_0_499, axiom, (cell36!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49142)).
fof(c_0_500, axiom, (cell37!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49102)).
fof(c_0_501, axiom, (cell38!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49061)).
fof(c_0_502, axiom, (cell39!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49019)).
fof(c_0_503, axiom, (cell40!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48976)).
fof(c_0_504, axiom, (cell31!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49292)).
fof(c_0_505, axiom, (cell32!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49327)).
fof(c_0_506, axiom, (cell33!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49256)).
fof(c_0_507, axiom, (cell34!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49219)).
fof(c_0_508, axiom, (cell35!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49181)).
fof(c_0_509, axiom, (cell26!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49487)).
fof(c_0_510, axiom, (cell27!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49457)).
fof(c_0_511, axiom, (cell28!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49426)).
fof(c_0_512, axiom, (cell29!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49394)).
fof(c_0_513, axiom, (cell30!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49361)).
fof(c_0_514, axiom, (cell88!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49902)).
fof(c_0_515, axiom, (cell21!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49597)).
fof(c_0_516, axiom, (cell23!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49571)).
fof(c_0_517, axiom, (cell24!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49544)).
fof(c_0_518, axiom, (cell25!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49516)).
fof(c_0_519, axiom, (cell16!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49712)).
fof(c_0_520, axiom, (cell17!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49691)).
fof(c_0_521, axiom, (cell18!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49669)).
fof(c_0_522, axiom, (cell19!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49646)).
fof(c_0_523, axiom, (cell20!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49622)).
fof(c_0_524, axiom, (cell88!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49898)).
fof(c_0_525, axiom, (cell88!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49899)).
fof(c_0_526, axiom, (cell88!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49901)).
fof(c_0_527, axiom, (cell14!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49751)).
fof(c_0_528, axiom, (cell15!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49732)).
fof(c_0_529, axiom, (cell6!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49831)).
fof(c_0_530, axiom, (cell7!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49817)).
fof(c_0_531, axiom, (cell8!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49802)).
fof(c_0_532, axiom, (cell9!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49786)).
fof(c_0_533, axiom, (cell10!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49769)).
fof(c_0_534, axiom, (cell1!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49877)).
fof(c_0_535, axiom, (cell3!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49867)).
fof(c_0_536, axiom, (cell4!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49856)).
fof(c_0_537, axiom, (cell5!=cell88), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49844)).
fof(c_0_538, axiom, (cell87!=cell12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45921)).
fof(c_0_539, axiom, (cell87!=cell13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45924)).
fof(c_0_540, axiom, (cell87!=cell22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45925)).
fof(c_0_541, axiom, (cell87!=cell11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45922)).
fof(c_0_542, axiom, (cell86!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46005)).
fof(c_0_543, axiom, (cell87!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45916)).
fof(c_0_544, axiom, (cell96!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45361)).
fof(c_0_545, axiom, (cell97!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45265)).
fof(c_0_546, axiom, (cell98!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45168)).
fof(c_0_547, axiom, (cell100!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45070)).
fof(c_0_548, axiom, (cell81!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46351)).
fof(c_0_549, axiom, (cell82!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46435)).
fof(c_0_550, axiom, (cell83!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46266)).
fof(c_0_551, axiom, (cell84!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46180)).
fof(c_0_552, axiom, (cell85!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46093)).
fof(c_0_553, axiom, (cell91!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45735)).
fof(c_0_554, axiom, (cell92!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45826)).
fof(c_0_555, axiom, (cell93!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45643)).
fof(c_0_556, axiom, (cell94!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45550)).
fof(c_0_557, axiom, (cell95!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof45456)).
fof(c_0_558, axiom, (cell76!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46761)).
fof(c_0_559, axiom, (cell77!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46681)).
fof(c_0_560, axiom, (cell78!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46600)).
fof(c_0_561, axiom, (cell80!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46518)).
fof(c_0_562, axiom, (cell71!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47071)).
fof(c_0_563, axiom, (cell72!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47146)).
fof(c_0_564, axiom, (cell73!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46995)).
fof(c_0_565, axiom, (cell74!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46918)).
fof(c_0_566, axiom, (cell75!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof46840)).
fof(c_0_567, axiom, (cell66!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47506)).
fof(c_0_568, axiom, (cell67!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47436)).
fof(c_0_569, axiom, (cell68!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47365)).
fof(c_0_570, axiom, (cell70!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47220)).
fof(c_0_571, axiom, (cell61!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47776)).
fof(c_0_572, axiom, (cell62!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47841)).
fof(c_0_573, axiom, (cell63!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47710)).
fof(c_0_574, axiom, (cell64!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47643)).
fof(c_0_575, axiom, (cell65!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47575)).
fof(c_0_576, axiom, (cell56!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48151)).
fof(c_0_577, axiom, (cell57!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48091)).
fof(c_0_578, axiom, (cell58!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48030)).
fof(c_0_579, axiom, (cell59!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47968)).
fof(c_0_580, axiom, (cell60!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47905)).
fof(c_0_581, axiom, (cell69!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof47293)).
fof(c_0_582, axiom, (cell51!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48381)).
fof(c_0_583, axiom, (cell52!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48436)).
fof(c_0_584, axiom, (cell53!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48325)).
fof(c_0_585, axiom, (cell54!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48268)).
fof(c_0_586, axiom, (cell55!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48210)).
fof(c_0_587, axiom, (cell46!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48696)).
fof(c_0_588, axiom, (cell47!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48646)).
fof(c_0_589, axiom, (cell48!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48595)).
fof(c_0_590, axiom, (cell49!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48543)).
fof(c_0_591, axiom, (cell50!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48490)).
fof(c_0_592, axiom, (cell41!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48886)).
fof(c_0_593, axiom, (cell42!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48931)).
fof(c_0_594, axiom, (cell43!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48840)).
fof(c_0_595, axiom, (cell44!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48793)).
fof(c_0_596, axiom, (cell45!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48745)).
fof(c_0_597, axiom, (cell36!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49141)).
fof(c_0_598, axiom, (cell37!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49101)).
fof(c_0_599, axiom, (cell38!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49060)).
fof(c_0_600, axiom, (cell39!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49018)).
fof(c_0_601, axiom, (cell40!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof48975)).
fof(c_0_602, axiom, (cell31!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49291)).
fof(c_0_603, axiom, (cell32!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49326)).
fof(c_0_604, axiom, (cell33!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49255)).
fof(c_0_605, axiom, (cell34!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49218)).
fof(c_0_606, axiom, (cell35!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49180)).
fof(c_0_607, axiom, (cell26!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49486)).
fof(c_0_608, axiom, (cell27!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49456)).
fof(c_0_609, axiom, (cell28!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49425)).
fof(c_0_610, axiom, (cell29!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49393)).
fof(c_0_611, axiom, (cell30!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49360)).
fof(c_0_612, axiom, (cell21!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49596)).
fof(c_0_613, axiom, (cell23!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49570)).
fof(c_0_614, axiom, (cell24!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49543)).
fof(c_0_615, axiom, (cell25!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49515)).
fof(c_0_616, axiom, (cell16!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49711)).
fof(c_0_617, axiom, (cell17!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49690)).
fof(c_0_618, axiom, (cell18!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49668)).
fof(c_0_619, axiom, (cell19!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49645)).
fof(c_0_620, axiom, (cell20!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49621)).
fof(c_0_621, axiom, (cell14!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49750)).
fof(c_0_622, axiom, (cell15!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49731)).
fof(c_0_623, axiom, (cell6!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49830)).
fof(c_0_624, axiom, (cell7!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49816)).
fof(c_0_625, axiom, (cell8!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49801)).
fof(c_0_626, axiom, (cell9!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49785)).
fof(c_0_627, axiom, (cell10!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49768)).
fof(c_0_628, axiom, (cell1!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49876)).
fof(c_0_629, axiom, (cell3!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49866)).
fof(c_0_630, axiom, (cell4!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49855)).
fof(c_0_631, axiom, (cell5!=cell89), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof49843)).
fof(c_0_632, negated_conjecture, (~(![X2]:![X1]:((patient(X2)&oxygen(X1))=>~adj(X2,X1)))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
fof(c_0_633, plain, (![X3]:![X4]:((~patient(X4)|~patient(X3))|X4=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
fof(c_0_634, negated_conjecture, (((patient(esk1_0)&oxygen(esk2_0))&adj(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_632])])])).
fof(c_0_635, plain, (![X4]:![X3]:(epred36_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60))|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65))|(X4=cell56&X3=cell66))|(X4=cell57&X3=cell67))|(X4=cell58&X3=cell68))|(X4=cell59&X3=cell69))|(X4=cell60&X3=cell70))|(X4=cell61&X3=cell71))|(X4=cell62&X3=cell72))|(X4=cell63&X3=cell73))|(X4=cell64&X3=cell74))|(X4=cell65&X3=cell75))|(X4=cell66&X3=cell76))|(X4=cell67&X3=cell77))|(X4=cell68&X3=cell78))|(X4=cell69&X3=cell79))|(X4=cell70&X3=cell80))|(X4=cell71&X3=cell81))|(X4=cell72&X3=cell82))|(X4=cell73&X3=cell83))|(X4=cell74&X3=cell84))|(X4=cell75&X3=cell85))|(X4=cell76&X3=cell86))|(X4=cell77&X3=cell87))|(X4=cell78&X3=cell88))|(X4=cell79&X3=cell89))|(X4=cell80&X3=cell90))|(X4=cell81&X3=cell91))|(X4=cell82&X3=cell92))|(X4=cell83&X3=cell93))|(X4=cell84&X3=cell94))|(X4=cell85&X3=cell95))|(X4=cell86&X3=cell96))|(X4=cell87&X3=cell97))|(X4=cell88&X3=cell98))|(X4=cell89&X3=cell99))|(X4=cell90&X3=cell100)))), introduced(definition)).
cnf(c_0_636,plain,(X1=X2|~patient(X2)|~patient(X1)), inference(split_conjunct,[status(thm)],[c_0_633])).
cnf(c_0_637,negated_conjecture,(patient(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_634])).
fof(c_0_638, plain, (![X3]:![X4]:((~oxygen(X4)|~oxygen(X3))|X4=X3)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
fof(c_0_639, axiom, (![X3]:![X4]:(edge(X4,X3)<=>epred36_2(X3,X4))), inference(apply_def,[status(thm)],[c_0_3, c_0_635])).
fof(c_0_640, plain, (![X3]:![X4]:![X5]:![X6]:((~adj(X4,X3)|(edge(X4,X3)|edge(X3,X4)))&((~edge(X6,X5)|adj(X6,X5))&(~edge(X5,X6)|adj(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
cnf(c_0_641,negated_conjecture,(X1=esk1_0|~patient(X1)), inference(spm,[status(thm)],[c_0_636, c_0_637])).
cnf(c_0_642,plain,(patient(cell22)|patient(cell13)|patient(cell2)|patient(cell11)|patient(cell12)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_643,plain,(X1=X2|~oxygen(X2)|~oxygen(X1)), inference(split_conjunct,[status(thm)],[c_0_638])).
cnf(c_0_644,negated_conjecture,(oxygen(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_634])).
fof(c_0_645, plain, (![X4]:![X3]:(epred35_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60))|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65))|(X4=cell56&X3=cell66))|(X4=cell57&X3=cell67))|(X4=cell58&X3=cell68))|(X4=cell59&X3=cell69))|(X4=cell60&X3=cell70))|(X4=cell61&X3=cell71))|(X4=cell62&X3=cell72))|(X4=cell63&X3=cell73))|(X4=cell64&X3=cell74))|(X4=cell65&X3=cell75))|(X4=cell66&X3=cell76))|(X4=cell67&X3=cell77))|(X4=cell68&X3=cell78))|(X4=cell69&X3=cell79))|(X4=cell70&X3=cell80))|(X4=cell71&X3=cell81))|(X4=cell72&X3=cell82))|(X4=cell73&X3=cell83))|(X4=cell74&X3=cell84))|(X4=cell75&X3=cell85))|(X4=cell76&X3=cell86))|(X4=cell77&X3=cell87))|(X4=cell78&X3=cell88))|(X4=cell79&X3=cell89))|(X4=cell80&X3=cell90))|(X4=cell81&X3=cell91))|(X4=cell82&X3=cell92))|(X4=cell83&X3=cell93))|(X4=cell84&X3=cell94))|(X4=cell85&X3=cell95)))), introduced(definition)).
fof(c_0_646, plain, (![X5]:![X6]:![X7]:![X8]:((~edge(X6,X5)|epred36_2(X5,X6))&(~epred36_2(X7,X8)|edge(X8,X7)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_639])])])])).
cnf(c_0_647,plain,(edge(X1,X2)|edge(X2,X1)|~adj(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_640])).
cnf(c_0_648,negated_conjecture,(adj(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_634])).
cnf(c_0_649,negated_conjecture,(esk1_0=cell2|patient(cell12)|patient(cell11)|patient(cell13)|patient(cell22)), inference(spm,[status(thm)],[c_0_641, c_0_642])).
cnf(c_0_650,negated_conjecture,(X1=esk2_0|~oxygen(X1)), inference(spm,[status(thm)],[c_0_643, c_0_644])).
cnf(c_0_651,plain,(oxygen(cell99)|oxygen(cell90)|oxygen(cell79)|oxygen(cell88)|oxygen(cell89)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_652, plain, (![X4]:![X3]:(epred36_2(X3,X4)<=>(((((epred35_2(X3,X4)|(X4=cell86&X3=cell96))|(X4=cell87&X3=cell97))|(X4=cell88&X3=cell98))|(X4=cell89&X3=cell99))|(X4=cell90&X3=cell100)))), inference(apply_def,[status(thm)],[c_0_635, c_0_645])).
cnf(c_0_653,plain,(epred36_2(X1,X2)|~edge(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_646])).
cnf(c_0_654,negated_conjecture,(edge(esk1_0,esk2_0)|edge(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_647, c_0_648])).
cnf(c_0_655,negated_conjecture,(esk1_0=cell11|esk1_0=cell2|patient(cell12)|patient(cell13)|patient(cell22)), inference(spm,[status(thm)],[c_0_641, c_0_649])).
cnf(c_0_656,negated_conjecture,(esk2_0=cell99|oxygen(cell89)|oxygen(cell88)|oxygen(cell79)|oxygen(cell90)), inference(spm,[status(thm)],[c_0_650, c_0_651])).
fof(c_0_657, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell90|(X5=cell89|(X5=cell88|(X5=cell87|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell89|(X5=cell88|(X5=cell87|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell90|(X6=cell99|(X5=cell88|(X5=cell87|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell88|(X5=cell87|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))&((((X5=cell90|(X5=cell89|(X6=cell98|(X5=cell87|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell89|(X6=cell98|(X5=cell87|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell90|(X6=cell99|(X6=cell98|(X5=cell87|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X5=cell87|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))))&(((((X5=cell90|(X5=cell89|(X5=cell88|(X6=cell97|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell89|(X5=cell88|(X6=cell97|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell90|(X6=cell99|(X5=cell88|(X6=cell97|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell88|(X6=cell97|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))&((((X5=cell90|(X5=cell89|(X6=cell98|(X6=cell97|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell89|(X6=cell98|(X6=cell97|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell90|(X6=cell99|(X6=cell98|(X6=cell97|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X6=cell97|(X5=cell86|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))))&((((((X5=cell90|(X5=cell89|(X5=cell88|(X5=cell87|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell89|(X5=cell88|(X5=cell87|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell90|(X6=cell99|(X5=cell88|(X5=cell87|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell88|(X5=cell87|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))&((((X5=cell90|(X5=cell89|(X6=cell98|(X5=cell87|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell89|(X6=cell98|(X5=cell87|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell90|(X6=cell99|(X6=cell98|(X5=cell87|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X5=cell87|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))))&(((((X5=cell90|(X5=cell89|(X5=cell88|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell89|(X5=cell88|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell90|(X6=cell99|(X5=cell88|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell88|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))&((((X5=cell90|(X5=cell89|(X6=cell98|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell89|(X6=cell98|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell90|(X6=cell99|(X6=cell98|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))))))&((((((~epred35_2(X8,X7)|epred36_2(X8,X7))&((X7!=cell86|X8!=cell96)|epred36_2(X8,X7)))&((X7!=cell87|X8!=cell97)|epred36_2(X8,X7)))&((X7!=cell88|X8!=cell98)|epred36_2(X8,X7)))&((X7!=cell89|X8!=cell99)|epred36_2(X8,X7)))&((X7!=cell90|X8!=cell100)|epred36_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_652])])])])])).
cnf(c_0_658,negated_conjecture,(epred36_2(esk2_0,esk1_0)|edge(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_653, c_0_654])).
cnf(c_0_659,negated_conjecture,(esk1_0=cell22|esk1_0=cell11|esk1_0=cell2|patient(cell12)|patient(cell13)), inference(spm,[status(thm)],[c_0_641, c_0_655])).
cnf(c_0_660,negated_conjecture,(esk2_0=cell90|esk2_0=cell99|oxygen(cell89)|oxygen(cell88)|oxygen(cell79)), inference(spm,[status(thm)],[c_0_650, c_0_656])).
fof(c_0_661, plain, (![X4]:![X3]:(epred34_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60))|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65))|(X4=cell56&X3=cell66))|(X4=cell57&X3=cell67))|(X4=cell58&X3=cell68))|(X4=cell59&X3=cell69))|(X4=cell60&X3=cell70))|(X4=cell61&X3=cell71))|(X4=cell62&X3=cell72))|(X4=cell63&X3=cell73))|(X4=cell64&X3=cell74))|(X4=cell65&X3=cell75))|(X4=cell66&X3=cell76))|(X4=cell67&X3=cell77))|(X4=cell68&X3=cell78))|(X4=cell69&X3=cell79))|(X4=cell70&X3=cell80))|(X4=cell71&X3=cell81))|(X4=cell72&X3=cell82))|(X4=cell73&X3=cell83))|(X4=cell74&X3=cell84))|(X4=cell75&X3=cell85))|(X4=cell76&X3=cell86))|(X4=cell77&X3=cell87))|(X4=cell78&X3=cell88))|(X4=cell79&X3=cell89))|(X4=cell80&X3=cell90)))), introduced(definition)).
cnf(c_0_662,plain,(epred35_2(X1,X2)|X1=cell96|X1=cell97|X1=cell98|X1=cell99|X1=cell100|~epred36_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_657])).
cnf(c_0_663,negated_conjecture,(epred36_2(esk1_0,esk2_0)|epred36_2(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_653, c_0_658])).
cnf(c_0_664,negated_conjecture,(esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|esk1_0=cell2|patient(cell12)), inference(spm,[status(thm)],[c_0_641, c_0_659])).
cnf(c_0_665,negated_conjecture,(esk2_0=cell79|esk2_0=cell90|esk2_0=cell99|oxygen(cell89)|oxygen(cell88)), inference(spm,[status(thm)],[c_0_650, c_0_660])).
fof(c_0_666, plain, (![X4]:![X3]:(epred35_2(X3,X4)<=>(((((epred34_2(X3,X4)|(X4=cell81&X3=cell91))|(X4=cell82&X3=cell92))|(X4=cell83&X3=cell93))|(X4=cell84&X3=cell94))|(X4=cell85&X3=cell95)))), inference(apply_def,[status(thm)],[c_0_645, c_0_661])).
cnf(c_0_667,plain,(esk1_0=cell99|esk1_0=cell96|esk1_0=cell97|esk1_0=cell98|esk1_0=cell100|epred36_2(esk2_0,esk1_0)|epred35_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_662, c_0_663])).
cnf(c_0_668,negated_conjecture,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|esk1_0=cell2), inference(spm,[status(thm)],[c_0_641, c_0_664])).
cnf(c_0_669,plain,(cell99!=cell2), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_670,plain,(cell96!=cell2), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_671,plain,(cell97!=cell2), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_672,plain,(cell98!=cell2), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_673,plain,(cell100!=cell2), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_674,plain,(epred35_2(X1,X2)|X2=cell86|X2=cell87|X2=cell88|X2=cell89|X2=cell90|~epred36_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_657])).
cnf(c_0_675,negated_conjecture,(esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|esk2_0=cell99|oxygen(cell89)), inference(spm,[status(thm)],[c_0_650, c_0_665])).
fof(c_0_676, plain, (![X4]:![X3]:(epred33_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60))|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65))|(X4=cell56&X3=cell66))|(X4=cell57&X3=cell67))|(X4=cell58&X3=cell68))|(X4=cell59&X3=cell69))|(X4=cell60&X3=cell70))|(X4=cell61&X3=cell71))|(X4=cell62&X3=cell72))|(X4=cell63&X3=cell73))|(X4=cell64&X3=cell74))|(X4=cell65&X3=cell75))|(X4=cell66&X3=cell76))|(X4=cell67&X3=cell77))|(X4=cell68&X3=cell78))|(X4=cell69&X3=cell79))|(X4=cell70&X3=cell80))|(X4=cell71&X3=cell81))|(X4=cell72&X3=cell82))|(X4=cell73&X3=cell83))|(X4=cell74&X3=cell84))|(X4=cell75&X3=cell85)))), introduced(definition)).
fof(c_0_677, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell85|(X5=cell84|(X5=cell83|(X5=cell82|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell84|(X5=cell83|(X5=cell82|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell85|(X6=cell94|(X5=cell83|(X5=cell82|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell83|(X5=cell82|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))&((((X5=cell85|(X5=cell84|(X6=cell93|(X5=cell82|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell84|(X6=cell93|(X5=cell82|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell85|(X6=cell94|(X6=cell93|(X5=cell82|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X5=cell82|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))))&(((((X5=cell85|(X5=cell84|(X5=cell83|(X6=cell92|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell84|(X5=cell83|(X6=cell92|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell85|(X6=cell94|(X5=cell83|(X6=cell92|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell83|(X6=cell92|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))&((((X5=cell85|(X5=cell84|(X6=cell93|(X6=cell92|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell84|(X6=cell93|(X6=cell92|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell85|(X6=cell94|(X6=cell93|(X6=cell92|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X6=cell92|(X5=cell81|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))))&((((((X5=cell85|(X5=cell84|(X5=cell83|(X5=cell82|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell84|(X5=cell83|(X5=cell82|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell85|(X6=cell94|(X5=cell83|(X5=cell82|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell83|(X5=cell82|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))&((((X5=cell85|(X5=cell84|(X6=cell93|(X5=cell82|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell84|(X6=cell93|(X5=cell82|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell85|(X6=cell94|(X6=cell93|(X5=cell82|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X5=cell82|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))))&(((((X5=cell85|(X5=cell84|(X5=cell83|(X6=cell92|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell84|(X5=cell83|(X6=cell92|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell85|(X6=cell94|(X5=cell83|(X6=cell92|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell83|(X6=cell92|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))&((((X5=cell85|(X5=cell84|(X6=cell93|(X6=cell92|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell84|(X6=cell93|(X6=cell92|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell85|(X6=cell94|(X6=cell93|(X6=cell92|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X6=cell92|(X6=cell91|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))))))&((((((~epred34_2(X8,X7)|epred35_2(X8,X7))&((X7!=cell81|X8!=cell91)|epred35_2(X8,X7)))&((X7!=cell82|X8!=cell92)|epred35_2(X8,X7)))&((X7!=cell83|X8!=cell93)|epred35_2(X8,X7)))&((X7!=cell84|X8!=cell94)|epred35_2(X8,X7)))&((X7!=cell85|X8!=cell95)|epred35_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_666])])])])])).
cnf(c_0_678,negated_conjecture,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred36_2(esk2_0,cell2)|epred35_2(cell2,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_667, c_0_668]), c_0_669]), c_0_670]), c_0_671]), c_0_672]), c_0_673])).
cnf(c_0_679,plain,(cell89!=cell2), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_680,plain,(cell88!=cell2), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_681,plain,(cell90!=cell2), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_682,plain,(cell86!=cell2), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_683,plain,(cell87!=cell2), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_684,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell90|esk2_0=cell86|esk2_0=cell87|epred36_2(esk2_0,esk1_0)|epred35_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_674, c_0_663])).
cnf(c_0_685,negated_conjecture,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|esk2_0=cell99), inference(spm,[status(thm)],[c_0_650, c_0_675])).
cnf(c_0_686,plain,(cell89!=cell99), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_687,plain,(cell88!=cell99), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_688,plain,(cell90!=cell99), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_689,plain,(cell86!=cell99), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_690,plain,(cell87!=cell99), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_691, plain, (![X4]:![X3]:(epred34_2(X3,X4)<=>(((((epred33_2(X3,X4)|(X4=cell76&X3=cell86))|(X4=cell77&X3=cell87))|(X4=cell78&X3=cell88))|(X4=cell79&X3=cell89))|(X4=cell80&X3=cell90)))), inference(apply_def,[status(thm)],[c_0_661, c_0_676])).
cnf(c_0_692,plain,(epred34_2(X1,X2)|X2=cell81|X2=cell82|X2=cell83|X2=cell84|X2=cell85|~epred35_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_677])).
cnf(c_0_693,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred35_2(esk2_0,cell2)|epred35_2(cell2,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_674, c_0_678]), c_0_679]), c_0_680]), c_0_681]), c_0_682]), c_0_683])).
cnf(c_0_694,plain,(cell81!=cell2), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_695,plain,(cell82!=cell2), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_696,plain,(cell83!=cell2), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_697,plain,(cell84!=cell2), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_698,plain,(cell85!=cell2), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_699,plain,(epred35_2(X1,X2)|X1=cell96|X1=cell97|X1=cell98|X2=cell89|X1=cell100|~epred36_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_657])).
cnf(c_0_700,negated_conjecture,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred36_2(cell99,esk1_0)|epred35_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_684, c_0_685]), c_0_686]), c_0_687]), c_0_688]), c_0_689]), c_0_690])).
cnf(c_0_701,plain,(cell96!=cell99), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_702,plain,(cell97!=cell99), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_703,plain,(cell98!=cell99), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_704,plain,(cell100!=cell99), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_705, plain, (![X4]:![X3]:(epred32_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60))|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65))|(X4=cell56&X3=cell66))|(X4=cell57&X3=cell67))|(X4=cell58&X3=cell68))|(X4=cell59&X3=cell69))|(X4=cell60&X3=cell70))|(X4=cell61&X3=cell71))|(X4=cell62&X3=cell72))|(X4=cell63&X3=cell73))|(X4=cell64&X3=cell74))|(X4=cell65&X3=cell75))|(X4=cell66&X3=cell76))|(X4=cell67&X3=cell77))|(X4=cell68&X3=cell78))|(X4=cell69&X3=cell79))|(X4=cell70&X3=cell80)))), introduced(definition)).
fof(c_0_706, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell80|(X5=cell79|(X5=cell78|(X5=cell77|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X5=cell79|(X5=cell78|(X5=cell77|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell80|(X6=cell89|(X5=cell78|(X5=cell77|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X6=cell89|(X5=cell78|(X5=cell77|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))&((((X5=cell80|(X5=cell79|(X6=cell88|(X5=cell77|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X5=cell79|(X6=cell88|(X5=cell77|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell80|(X6=cell89|(X6=cell88|(X5=cell77|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X6=cell89|(X6=cell88|(X5=cell77|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))))&(((((X5=cell80|(X5=cell79|(X5=cell78|(X6=cell87|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X5=cell79|(X5=cell78|(X6=cell87|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell80|(X6=cell89|(X5=cell78|(X6=cell87|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X6=cell89|(X5=cell78|(X6=cell87|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))&((((X5=cell80|(X5=cell79|(X6=cell88|(X6=cell87|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X5=cell79|(X6=cell88|(X6=cell87|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell80|(X6=cell89|(X6=cell88|(X6=cell87|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X6=cell89|(X6=cell88|(X6=cell87|(X5=cell76|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))))&((((((X5=cell80|(X5=cell79|(X5=cell78|(X5=cell77|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X5=cell79|(X5=cell78|(X5=cell77|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell80|(X6=cell89|(X5=cell78|(X5=cell77|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X6=cell89|(X5=cell78|(X5=cell77|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))&((((X5=cell80|(X5=cell79|(X6=cell88|(X5=cell77|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X5=cell79|(X6=cell88|(X5=cell77|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell80|(X6=cell89|(X6=cell88|(X5=cell77|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X6=cell89|(X6=cell88|(X5=cell77|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))))&(((((X5=cell80|(X5=cell79|(X5=cell78|(X6=cell87|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X5=cell79|(X5=cell78|(X6=cell87|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell80|(X6=cell89|(X5=cell78|(X6=cell87|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X6=cell89|(X5=cell78|(X6=cell87|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))&((((X5=cell80|(X5=cell79|(X6=cell88|(X6=cell87|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X5=cell79|(X6=cell88|(X6=cell87|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell80|(X6=cell89|(X6=cell88|(X6=cell87|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell90|(X6=cell89|(X6=cell88|(X6=cell87|(X6=cell86|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))))))&((((((~epred33_2(X8,X7)|epred34_2(X8,X7))&((X7!=cell76|X8!=cell86)|epred34_2(X8,X7)))&((X7!=cell77|X8!=cell87)|epred34_2(X8,X7)))&((X7!=cell78|X8!=cell88)|epred34_2(X8,X7)))&((X7!=cell79|X8!=cell89)|epred34_2(X8,X7)))&((X7!=cell80|X8!=cell90)|epred34_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_691])])])])])).
cnf(c_0_707,plain,(epred34_2(X1,X2)|X1=cell91|X1=cell92|X1=cell93|X1=cell94|X1=cell95|~epred35_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_677])).
cnf(c_0_708,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred35_2(cell2,esk2_0)|epred34_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_693]), c_0_694]), c_0_695]), c_0_696]), c_0_697]), c_0_698])).
cnf(c_0_709,plain,(cell91!=cell2), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_710,plain,(cell92!=cell2), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_711,plain,(cell93!=cell2), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_712,plain,(cell94!=cell2), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_713,plain,(cell95!=cell2), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_714,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred35_2(cell99,esk1_0)|epred35_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_699, c_0_700]), c_0_701]), c_0_702]), c_0_703]), c_0_704])).
cnf(c_0_715,plain,(cell81!=cell99), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_716,plain,(cell82!=cell99), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_717,plain,(cell83!=cell99), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_718,plain,(cell84!=cell99), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_719,plain,(cell85!=cell99), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_720, plain, (![X4]:![X3]:(epred33_2(X3,X4)<=>(((((epred32_2(X3,X4)|(X4=cell71&X3=cell81))|(X4=cell72&X3=cell82))|(X4=cell73&X3=cell83))|(X4=cell74&X3=cell84))|(X4=cell75&X3=cell85)))), inference(apply_def,[status(thm)],[c_0_676, c_0_705])).
cnf(c_0_721,plain,(epred33_2(X1,X2)|X2=cell76|X2=cell77|X2=cell78|X2=cell79|X2=cell80|~epred34_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_706])).
cnf(c_0_722,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred34_2(cell2,esk2_0)|epred34_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_708]), c_0_709]), c_0_710]), c_0_711]), c_0_712]), c_0_713])).
cnf(c_0_723,plain,(cell79!=cell2), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_724,plain,(cell76!=cell2), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_725,plain,(cell77!=cell2), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_726,plain,(cell78!=cell2), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_727,plain,(cell80!=cell2), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_728,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred35_2(cell99,esk1_0)|epred34_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_714]), c_0_715]), c_0_716]), c_0_717]), c_0_718]), c_0_719])).
cnf(c_0_729,plain,(cell91!=cell99), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_730,plain,(cell92!=cell99), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_731,plain,(cell93!=cell99), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_732,plain,(cell94!=cell99), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_733,plain,(cell95!=cell99), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_734, plain, (![X4]:![X3]:(epred31_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60))|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65))|(X4=cell56&X3=cell66))|(X4=cell57&X3=cell67))|(X4=cell58&X3=cell68))|(X4=cell59&X3=cell69))|(X4=cell60&X3=cell70))|(X4=cell61&X3=cell71))|(X4=cell62&X3=cell72))|(X4=cell63&X3=cell73))|(X4=cell64&X3=cell74))|(X4=cell65&X3=cell75)))), introduced(definition)).
fof(c_0_735, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell75|(X5=cell74|(X5=cell73|(X5=cell72|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X5=cell74|(X5=cell73|(X5=cell72|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell75|(X6=cell84|(X5=cell73|(X5=cell72|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X6=cell84|(X5=cell73|(X5=cell72|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))&((((X5=cell75|(X5=cell74|(X6=cell83|(X5=cell72|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X5=cell74|(X6=cell83|(X5=cell72|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell75|(X6=cell84|(X6=cell83|(X5=cell72|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X6=cell84|(X6=cell83|(X5=cell72|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))))&(((((X5=cell75|(X5=cell74|(X5=cell73|(X6=cell82|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X5=cell74|(X5=cell73|(X6=cell82|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell75|(X6=cell84|(X5=cell73|(X6=cell82|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X6=cell84|(X5=cell73|(X6=cell82|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))&((((X5=cell75|(X5=cell74|(X6=cell83|(X6=cell82|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X5=cell74|(X6=cell83|(X6=cell82|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell75|(X6=cell84|(X6=cell83|(X6=cell82|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X6=cell84|(X6=cell83|(X6=cell82|(X5=cell71|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))))&((((((X5=cell75|(X5=cell74|(X5=cell73|(X5=cell72|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X5=cell74|(X5=cell73|(X5=cell72|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell75|(X6=cell84|(X5=cell73|(X5=cell72|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X6=cell84|(X5=cell73|(X5=cell72|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))&((((X5=cell75|(X5=cell74|(X6=cell83|(X5=cell72|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X5=cell74|(X6=cell83|(X5=cell72|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell75|(X6=cell84|(X6=cell83|(X5=cell72|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X6=cell84|(X6=cell83|(X5=cell72|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))))&(((((X5=cell75|(X5=cell74|(X5=cell73|(X6=cell82|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X5=cell74|(X5=cell73|(X6=cell82|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell75|(X6=cell84|(X5=cell73|(X6=cell82|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X6=cell84|(X5=cell73|(X6=cell82|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))&((((X5=cell75|(X5=cell74|(X6=cell83|(X6=cell82|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X5=cell74|(X6=cell83|(X6=cell82|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell75|(X6=cell84|(X6=cell83|(X6=cell82|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell85|(X6=cell84|(X6=cell83|(X6=cell82|(X6=cell81|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))))))&((((((~epred32_2(X8,X7)|epred33_2(X8,X7))&((X7!=cell71|X8!=cell81)|epred33_2(X8,X7)))&((X7!=cell72|X8!=cell82)|epred33_2(X8,X7)))&((X7!=cell73|X8!=cell83)|epred33_2(X8,X7)))&((X7!=cell74|X8!=cell84)|epred33_2(X8,X7)))&((X7!=cell75|X8!=cell85)|epred33_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_720])])])])])).
cnf(c_0_736,plain,(epred33_2(X1,X2)|X1=cell86|X1=cell87|X1=cell88|X1=cell89|X1=cell90|~epred34_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_706])).
cnf(c_0_737,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred34_2(cell2,esk2_0)|epred33_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_721, c_0_722]), c_0_723]), c_0_724]), c_0_725]), c_0_726]), c_0_727])).
cnf(c_0_738,plain,(epred33_2(X1,X2)|X2=cell76|X2=cell77|X2=cell78|X1=cell89|X2=cell80|~epred34_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_706])).
cnf(c_0_739,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred34_2(cell99,esk1_0)|epred34_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_728]), c_0_729]), c_0_730]), c_0_731]), c_0_732]), c_0_733])).
cnf(c_0_740,plain,(cell76!=cell99), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_741,plain,(cell77!=cell99), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_742,plain,(cell78!=cell99), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_743,plain,(cell80!=cell99), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_744, plain, (![X4]:![X3]:(epred32_2(X3,X4)<=>(((((epred31_2(X3,X4)|(X4=cell66&X3=cell76))|(X4=cell67&X3=cell77))|(X4=cell68&X3=cell78))|(X4=cell69&X3=cell79))|(X4=cell70&X3=cell80)))), inference(apply_def,[status(thm)],[c_0_705, c_0_734])).
cnf(c_0_745,plain,(epred32_2(X1,X2)|X2=cell71|X2=cell72|X2=cell73|X2=cell74|X2=cell75|~epred33_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_735])).
cnf(c_0_746,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred33_2(cell2,esk2_0)|epred33_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_736, c_0_737]), c_0_679]), c_0_680]), c_0_681]), c_0_682]), c_0_683])).
cnf(c_0_747,plain,(cell71!=cell2), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_748,plain,(cell72!=cell2), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_749,plain,(cell73!=cell2), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_750,plain,(cell74!=cell2), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_751,plain,(cell75!=cell2), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_752,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred34_2(cell99,esk1_0)|epred33_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_738, c_0_739]), c_0_740]), c_0_741]), c_0_742]), c_0_743])).
fof(c_0_753, plain, (![X4]:![X3]:(epred30_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60))|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65))|(X4=cell56&X3=cell66))|(X4=cell57&X3=cell67))|(X4=cell58&X3=cell68))|(X4=cell59&X3=cell69))|(X4=cell60&X3=cell70)))), introduced(definition)).
fof(c_0_754, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell70|(X5=cell69|(X5=cell68|(X5=cell67|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X5=cell69|(X5=cell68|(X5=cell67|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell70|(X6=cell79|(X5=cell68|(X5=cell67|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X6=cell79|(X5=cell68|(X5=cell67|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))&((((X5=cell70|(X5=cell69|(X6=cell78|(X5=cell67|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X5=cell69|(X6=cell78|(X5=cell67|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell70|(X6=cell79|(X6=cell78|(X5=cell67|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X6=cell79|(X6=cell78|(X5=cell67|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))))&(((((X5=cell70|(X5=cell69|(X5=cell68|(X6=cell77|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X5=cell69|(X5=cell68|(X6=cell77|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell70|(X6=cell79|(X5=cell68|(X6=cell77|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X6=cell79|(X5=cell68|(X6=cell77|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))&((((X5=cell70|(X5=cell69|(X6=cell78|(X6=cell77|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X5=cell69|(X6=cell78|(X6=cell77|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell70|(X6=cell79|(X6=cell78|(X6=cell77|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X6=cell79|(X6=cell78|(X6=cell77|(X5=cell66|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))))&((((((X5=cell70|(X5=cell69|(X5=cell68|(X5=cell67|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X5=cell69|(X5=cell68|(X5=cell67|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell70|(X6=cell79|(X5=cell68|(X5=cell67|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X6=cell79|(X5=cell68|(X5=cell67|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))&((((X5=cell70|(X5=cell69|(X6=cell78|(X5=cell67|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X5=cell69|(X6=cell78|(X5=cell67|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell70|(X6=cell79|(X6=cell78|(X5=cell67|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X6=cell79|(X6=cell78|(X5=cell67|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))))&(((((X5=cell70|(X5=cell69|(X5=cell68|(X6=cell77|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X5=cell69|(X5=cell68|(X6=cell77|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell70|(X6=cell79|(X5=cell68|(X6=cell77|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X6=cell79|(X5=cell68|(X6=cell77|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))&((((X5=cell70|(X5=cell69|(X6=cell78|(X6=cell77|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X5=cell69|(X6=cell78|(X6=cell77|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell70|(X6=cell79|(X6=cell78|(X6=cell77|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell80|(X6=cell79|(X6=cell78|(X6=cell77|(X6=cell76|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))))))&((((((~epred31_2(X8,X7)|epred32_2(X8,X7))&((X7!=cell66|X8!=cell76)|epred32_2(X8,X7)))&((X7!=cell67|X8!=cell77)|epred32_2(X8,X7)))&((X7!=cell68|X8!=cell78)|epred32_2(X8,X7)))&((X7!=cell69|X8!=cell79)|epred32_2(X8,X7)))&((X7!=cell70|X8!=cell80)|epred32_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_744])])])])])).
cnf(c_0_755,plain,(epred32_2(X1,X2)|X1=cell81|X1=cell82|X1=cell83|X1=cell84|X1=cell85|~epred33_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_735])).
cnf(c_0_756,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred33_2(cell2,esk2_0)|epred32_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_745, c_0_746]), c_0_747]), c_0_748]), c_0_749]), c_0_750]), c_0_751])).
cnf(c_0_757,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred33_2(cell99,esk1_0)|epred33_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_736, c_0_752]), c_0_686]), c_0_687]), c_0_688]), c_0_689]), c_0_690])).
cnf(c_0_758,plain,(cell71!=cell99), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_759,plain,(cell72!=cell99), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_760,plain,(cell73!=cell99), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_761,plain,(cell74!=cell99), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_762,plain,(cell75!=cell99), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_763, plain, (![X4]:![X3]:(epred31_2(X3,X4)<=>(((((epred30_2(X3,X4)|(X4=cell61&X3=cell71))|(X4=cell62&X3=cell72))|(X4=cell63&X3=cell73))|(X4=cell64&X3=cell74))|(X4=cell65&X3=cell75)))), inference(apply_def,[status(thm)],[c_0_734, c_0_753])).
cnf(c_0_764,plain,(epred31_2(X1,X2)|X2=cell66|X2=cell67|X2=cell68|X2=cell69|X2=cell70|~epred32_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_754])).
cnf(c_0_765,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred32_2(cell2,esk2_0)|epred32_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_755, c_0_756]), c_0_694]), c_0_695]), c_0_696]), c_0_697]), c_0_698])).
cnf(c_0_766,plain,(cell66!=cell2), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_767,plain,(cell67!=cell2), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_768,plain,(cell68!=cell2), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_769,plain,(cell69!=cell2), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_770,plain,(cell70!=cell2), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_771,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred33_2(cell99,esk1_0)|epred32_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_745, c_0_757]), c_0_758]), c_0_759]), c_0_760]), c_0_761]), c_0_762])).
fof(c_0_772, plain, (![X4]:![X3]:(epred29_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60))|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65)))), introduced(definition)).
fof(c_0_773, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell65|(X5=cell64|(X5=cell63|(X5=cell62|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X5=cell64|(X5=cell63|(X5=cell62|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell65|(X6=cell74|(X5=cell63|(X5=cell62|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X6=cell74|(X5=cell63|(X5=cell62|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))&((((X5=cell65|(X5=cell64|(X6=cell73|(X5=cell62|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X5=cell64|(X6=cell73|(X5=cell62|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell65|(X6=cell74|(X6=cell73|(X5=cell62|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X6=cell74|(X6=cell73|(X5=cell62|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))))&(((((X5=cell65|(X5=cell64|(X5=cell63|(X6=cell72|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X5=cell64|(X5=cell63|(X6=cell72|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell65|(X6=cell74|(X5=cell63|(X6=cell72|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X6=cell74|(X5=cell63|(X6=cell72|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))&((((X5=cell65|(X5=cell64|(X6=cell73|(X6=cell72|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X5=cell64|(X6=cell73|(X6=cell72|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell65|(X6=cell74|(X6=cell73|(X6=cell72|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X6=cell74|(X6=cell73|(X6=cell72|(X5=cell61|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))))&((((((X5=cell65|(X5=cell64|(X5=cell63|(X5=cell62|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X5=cell64|(X5=cell63|(X5=cell62|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell65|(X6=cell74|(X5=cell63|(X5=cell62|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X6=cell74|(X5=cell63|(X5=cell62|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))&((((X5=cell65|(X5=cell64|(X6=cell73|(X5=cell62|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X5=cell64|(X6=cell73|(X5=cell62|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell65|(X6=cell74|(X6=cell73|(X5=cell62|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X6=cell74|(X6=cell73|(X5=cell62|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))))&(((((X5=cell65|(X5=cell64|(X5=cell63|(X6=cell72|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X5=cell64|(X5=cell63|(X6=cell72|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell65|(X6=cell74|(X5=cell63|(X6=cell72|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X6=cell74|(X5=cell63|(X6=cell72|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))&((((X5=cell65|(X5=cell64|(X6=cell73|(X6=cell72|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X5=cell64|(X6=cell73|(X6=cell72|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell65|(X6=cell74|(X6=cell73|(X6=cell72|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell75|(X6=cell74|(X6=cell73|(X6=cell72|(X6=cell71|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))))))&((((((~epred30_2(X8,X7)|epred31_2(X8,X7))&((X7!=cell61|X8!=cell71)|epred31_2(X8,X7)))&((X7!=cell62|X8!=cell72)|epred31_2(X8,X7)))&((X7!=cell63|X8!=cell73)|epred31_2(X8,X7)))&((X7!=cell64|X8!=cell74)|epred31_2(X8,X7)))&((X7!=cell65|X8!=cell75)|epred31_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_763])])])])])).
cnf(c_0_774,plain,(epred31_2(X1,X2)|X1=cell76|X1=cell77|X1=cell78|X1=cell79|X1=cell80|~epred32_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_754])).
cnf(c_0_775,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred32_2(cell2,esk2_0)|epred31_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_765]), c_0_766]), c_0_767]), c_0_768]), c_0_769]), c_0_770])).
cnf(c_0_776,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred32_2(cell99,esk1_0)|epred32_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_755, c_0_771]), c_0_715]), c_0_716]), c_0_717]), c_0_718]), c_0_719])).
cnf(c_0_777,plain,(cell66!=cell99), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_778,plain,(cell67!=cell99), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_779,plain,(cell68!=cell99), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_780,plain,(cell69!=cell99), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_781,plain,(cell70!=cell99), inference(split_conjunct,[status(thm)],[c_0_74])).
fof(c_0_782, plain, (![X4]:![X3]:(epred30_2(X3,X4)<=>(((((epred29_2(X3,X4)|(X4=cell56&X3=cell66))|(X4=cell57&X3=cell67))|(X4=cell58&X3=cell68))|(X4=cell59&X3=cell69))|(X4=cell60&X3=cell70)))), inference(apply_def,[status(thm)],[c_0_753, c_0_772])).
cnf(c_0_783,plain,(epred30_2(X1,X2)|X2=cell61|X2=cell62|X2=cell63|X2=cell64|X2=cell65|~epred31_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_773])).
cnf(c_0_784,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred31_2(cell2,esk2_0)|epred31_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_774, c_0_775]), c_0_723]), c_0_724]), c_0_725]), c_0_726]), c_0_727])).
cnf(c_0_785,plain,(cell61!=cell2), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_786,plain,(cell62!=cell2), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_787,plain,(cell63!=cell2), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_788,plain,(cell64!=cell2), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_789,plain,(cell65!=cell2), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_790,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred32_2(cell99,esk1_0)|epred31_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_776]), c_0_777]), c_0_778]), c_0_779]), c_0_780]), c_0_781])).
cnf(c_0_791,plain,(cell79!=cell99), inference(split_conjunct,[status(thm)],[c_0_80])).
fof(c_0_792, plain, (![X4]:![X3]:(epred28_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55))|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60)))), introduced(definition)).
fof(c_0_793, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell60|(X5=cell59|(X5=cell58|(X5=cell57|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X5=cell59|(X5=cell58|(X5=cell57|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell60|(X6=cell69|(X5=cell58|(X5=cell57|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X6=cell69|(X5=cell58|(X5=cell57|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))&((((X5=cell60|(X5=cell59|(X6=cell68|(X5=cell57|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X5=cell59|(X6=cell68|(X5=cell57|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell60|(X6=cell69|(X6=cell68|(X5=cell57|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X6=cell69|(X6=cell68|(X5=cell57|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))))&(((((X5=cell60|(X5=cell59|(X5=cell58|(X6=cell67|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X5=cell59|(X5=cell58|(X6=cell67|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell60|(X6=cell69|(X5=cell58|(X6=cell67|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X6=cell69|(X5=cell58|(X6=cell67|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))&((((X5=cell60|(X5=cell59|(X6=cell68|(X6=cell67|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X5=cell59|(X6=cell68|(X6=cell67|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell60|(X6=cell69|(X6=cell68|(X6=cell67|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X6=cell69|(X6=cell68|(X6=cell67|(X5=cell56|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))))&((((((X5=cell60|(X5=cell59|(X5=cell58|(X5=cell57|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X5=cell59|(X5=cell58|(X5=cell57|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell60|(X6=cell69|(X5=cell58|(X5=cell57|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X6=cell69|(X5=cell58|(X5=cell57|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))&((((X5=cell60|(X5=cell59|(X6=cell68|(X5=cell57|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X5=cell59|(X6=cell68|(X5=cell57|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell60|(X6=cell69|(X6=cell68|(X5=cell57|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X6=cell69|(X6=cell68|(X5=cell57|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))))&(((((X5=cell60|(X5=cell59|(X5=cell58|(X6=cell67|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X5=cell59|(X5=cell58|(X6=cell67|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell60|(X6=cell69|(X5=cell58|(X6=cell67|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X6=cell69|(X5=cell58|(X6=cell67|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))&((((X5=cell60|(X5=cell59|(X6=cell68|(X6=cell67|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X5=cell59|(X6=cell68|(X6=cell67|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell60|(X6=cell69|(X6=cell68|(X6=cell67|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell70|(X6=cell69|(X6=cell68|(X6=cell67|(X6=cell66|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))))))&((((((~epred29_2(X8,X7)|epred30_2(X8,X7))&((X7!=cell56|X8!=cell66)|epred30_2(X8,X7)))&((X7!=cell57|X8!=cell67)|epred30_2(X8,X7)))&((X7!=cell58|X8!=cell68)|epred30_2(X8,X7)))&((X7!=cell59|X8!=cell69)|epred30_2(X8,X7)))&((X7!=cell60|X8!=cell70)|epred30_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_782])])])])])).
cnf(c_0_794,plain,(epred30_2(X1,X2)|X1=cell71|X1=cell72|X1=cell73|X1=cell74|X1=cell75|~epred31_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_773])).
cnf(c_0_795,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred31_2(cell2,esk2_0)|epred30_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_783, c_0_784]), c_0_785]), c_0_786]), c_0_787]), c_0_788]), c_0_789])).
cnf(c_0_796,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred31_2(cell99,esk1_0)|epred31_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_774, c_0_790]), c_0_791]), c_0_740]), c_0_741]), c_0_742]), c_0_743])).
cnf(c_0_797,plain,(cell61!=cell99), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_798,plain,(cell62!=cell99), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_799,plain,(cell63!=cell99), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_800,plain,(cell64!=cell99), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_801,plain,(cell65!=cell99), inference(split_conjunct,[status(thm)],[c_0_85])).
fof(c_0_802, plain, (![X4]:![X3]:(epred29_2(X3,X4)<=>(((((epred28_2(X3,X4)|(X4=cell51&X3=cell61))|(X4=cell52&X3=cell62))|(X4=cell53&X3=cell63))|(X4=cell54&X3=cell64))|(X4=cell55&X3=cell65)))), inference(apply_def,[status(thm)],[c_0_772, c_0_792])).
cnf(c_0_803,plain,(epred29_2(X1,X2)|X2=cell56|X2=cell57|X2=cell58|X2=cell59|X2=cell60|~epred30_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_793])).
cnf(c_0_804,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred30_2(cell2,esk2_0)|epred30_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_795]), c_0_747]), c_0_748]), c_0_749]), c_0_750]), c_0_751])).
cnf(c_0_805,plain,(cell56!=cell2), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_806,plain,(cell57!=cell2), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_807,plain,(cell58!=cell2), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_808,plain,(cell59!=cell2), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_809,plain,(cell60!=cell2), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_810,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred31_2(cell99,esk1_0)|epred30_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_783, c_0_796]), c_0_797]), c_0_798]), c_0_799]), c_0_800]), c_0_801])).
fof(c_0_811, plain, (![X4]:![X3]:(epred27_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50))|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55)))), introduced(definition)).
fof(c_0_812, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell55|(X5=cell54|(X5=cell53|(X5=cell52|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X5=cell54|(X5=cell53|(X5=cell52|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell55|(X6=cell64|(X5=cell53|(X5=cell52|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X6=cell64|(X5=cell53|(X5=cell52|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell63|(X5=cell52|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X5=cell54|(X6=cell63|(X5=cell52|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell55|(X6=cell64|(X6=cell63|(X5=cell52|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X6=cell64|(X6=cell63|(X5=cell52|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))))&(((((X5=cell55|(X5=cell54|(X5=cell53|(X6=cell62|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X5=cell54|(X5=cell53|(X6=cell62|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell55|(X6=cell64|(X5=cell53|(X6=cell62|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X6=cell64|(X5=cell53|(X6=cell62|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell63|(X6=cell62|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X5=cell54|(X6=cell63|(X6=cell62|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell55|(X6=cell64|(X6=cell63|(X6=cell62|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X6=cell64|(X6=cell63|(X6=cell62|(X5=cell51|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))))&((((((X5=cell55|(X5=cell54|(X5=cell53|(X5=cell52|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X5=cell54|(X5=cell53|(X5=cell52|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell55|(X6=cell64|(X5=cell53|(X5=cell52|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X6=cell64|(X5=cell53|(X5=cell52|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell63|(X5=cell52|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X5=cell54|(X6=cell63|(X5=cell52|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell55|(X6=cell64|(X6=cell63|(X5=cell52|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X6=cell64|(X6=cell63|(X5=cell52|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))))&(((((X5=cell55|(X5=cell54|(X5=cell53|(X6=cell62|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X5=cell54|(X5=cell53|(X6=cell62|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell55|(X6=cell64|(X5=cell53|(X6=cell62|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X6=cell64|(X5=cell53|(X6=cell62|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell63|(X6=cell62|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X5=cell54|(X6=cell63|(X6=cell62|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell55|(X6=cell64|(X6=cell63|(X6=cell62|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell65|(X6=cell64|(X6=cell63|(X6=cell62|(X6=cell61|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))))))&((((((~epred28_2(X8,X7)|epred29_2(X8,X7))&((X7!=cell51|X8!=cell61)|epred29_2(X8,X7)))&((X7!=cell52|X8!=cell62)|epred29_2(X8,X7)))&((X7!=cell53|X8!=cell63)|epred29_2(X8,X7)))&((X7!=cell54|X8!=cell64)|epred29_2(X8,X7)))&((X7!=cell55|X8!=cell65)|epred29_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_802])])])])])).
cnf(c_0_813,plain,(epred29_2(X1,X2)|X1=cell66|X1=cell67|X1=cell68|X1=cell69|X1=cell70|~epred30_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_793])).
cnf(c_0_814,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred30_2(cell2,esk2_0)|epred29_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_803, c_0_804]), c_0_805]), c_0_806]), c_0_807]), c_0_808]), c_0_809])).
cnf(c_0_815,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred30_2(cell99,esk1_0)|epred30_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_810]), c_0_758]), c_0_759]), c_0_760]), c_0_761]), c_0_762])).
cnf(c_0_816,plain,(cell56!=cell99), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_817,plain,(cell57!=cell99), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_818,plain,(cell58!=cell99), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_819,plain,(cell59!=cell99), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_820,plain,(cell60!=cell99), inference(split_conjunct,[status(thm)],[c_0_95])).
fof(c_0_821, plain, (![X4]:![X3]:(epred28_2(X3,X4)<=>(((((epred27_2(X3,X4)|(X4=cell46&X3=cell56))|(X4=cell47&X3=cell57))|(X4=cell48&X3=cell58))|(X4=cell49&X3=cell59))|(X4=cell50&X3=cell60)))), inference(apply_def,[status(thm)],[c_0_792, c_0_811])).
cnf(c_0_822,plain,(epred28_2(X1,X2)|X2=cell51|X2=cell52|X2=cell53|X2=cell54|X2=cell55|~epred29_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_812])).
cnf(c_0_823,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred29_2(cell2,esk2_0)|epred29_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_813, c_0_814]), c_0_766]), c_0_767]), c_0_768]), c_0_769]), c_0_770])).
cnf(c_0_824,plain,(cell51!=cell2), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_825,plain,(cell52!=cell2), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_826,plain,(cell53!=cell2), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_827,plain,(cell54!=cell2), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_828,plain,(cell55!=cell2), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_829,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred30_2(cell99,esk1_0)|epred29_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_803, c_0_815]), c_0_816]), c_0_817]), c_0_818]), c_0_819]), c_0_820])).
fof(c_0_830, plain, (![X4]:![X3]:(epred26_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45))|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50)))), introduced(definition)).
fof(c_0_831, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell50|(X5=cell49|(X5=cell48|(X5=cell47|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X5=cell49|(X5=cell48|(X5=cell47|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell50|(X6=cell59|(X5=cell48|(X5=cell47|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X6=cell59|(X5=cell48|(X5=cell47|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))&((((X5=cell50|(X5=cell49|(X6=cell58|(X5=cell47|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X5=cell49|(X6=cell58|(X5=cell47|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell50|(X6=cell59|(X6=cell58|(X5=cell47|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X6=cell59|(X6=cell58|(X5=cell47|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))))&(((((X5=cell50|(X5=cell49|(X5=cell48|(X6=cell57|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X5=cell49|(X5=cell48|(X6=cell57|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell50|(X6=cell59|(X5=cell48|(X6=cell57|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X6=cell59|(X5=cell48|(X6=cell57|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))&((((X5=cell50|(X5=cell49|(X6=cell58|(X6=cell57|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X5=cell49|(X6=cell58|(X6=cell57|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell50|(X6=cell59|(X6=cell58|(X6=cell57|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X6=cell59|(X6=cell58|(X6=cell57|(X5=cell46|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))))&((((((X5=cell50|(X5=cell49|(X5=cell48|(X5=cell47|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X5=cell49|(X5=cell48|(X5=cell47|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell50|(X6=cell59|(X5=cell48|(X5=cell47|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X6=cell59|(X5=cell48|(X5=cell47|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))&((((X5=cell50|(X5=cell49|(X6=cell58|(X5=cell47|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X5=cell49|(X6=cell58|(X5=cell47|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell50|(X6=cell59|(X6=cell58|(X5=cell47|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X6=cell59|(X6=cell58|(X5=cell47|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))))&(((((X5=cell50|(X5=cell49|(X5=cell48|(X6=cell57|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X5=cell49|(X5=cell48|(X6=cell57|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell50|(X6=cell59|(X5=cell48|(X6=cell57|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X6=cell59|(X5=cell48|(X6=cell57|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))&((((X5=cell50|(X5=cell49|(X6=cell58|(X6=cell57|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X5=cell49|(X6=cell58|(X6=cell57|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell50|(X6=cell59|(X6=cell58|(X6=cell57|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell60|(X6=cell59|(X6=cell58|(X6=cell57|(X6=cell56|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))))))&((((((~epred27_2(X8,X7)|epred28_2(X8,X7))&((X7!=cell46|X8!=cell56)|epred28_2(X8,X7)))&((X7!=cell47|X8!=cell57)|epred28_2(X8,X7)))&((X7!=cell48|X8!=cell58)|epred28_2(X8,X7)))&((X7!=cell49|X8!=cell59)|epred28_2(X8,X7)))&((X7!=cell50|X8!=cell60)|epred28_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_821])])])])])).
cnf(c_0_832,plain,(epred28_2(X1,X2)|X1=cell61|X1=cell62|X1=cell63|X1=cell64|X1=cell65|~epred29_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_812])).
cnf(c_0_833,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred29_2(cell2,esk2_0)|epred28_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_822, c_0_823]), c_0_824]), c_0_825]), c_0_826]), c_0_827]), c_0_828])).
cnf(c_0_834,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred29_2(cell99,esk1_0)|epred29_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_813, c_0_829]), c_0_777]), c_0_778]), c_0_779]), c_0_780]), c_0_781])).
cnf(c_0_835,plain,(cell51!=cell99), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_836,plain,(cell52!=cell99), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_837,plain,(cell53!=cell99), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_838,plain,(cell54!=cell99), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_839,plain,(cell55!=cell99), inference(split_conjunct,[status(thm)],[c_0_105])).
fof(c_0_840, plain, (![X4]:![X3]:(epred27_2(X3,X4)<=>(((((epred26_2(X3,X4)|(X4=cell41&X3=cell51))|(X4=cell42&X3=cell52))|(X4=cell43&X3=cell53))|(X4=cell44&X3=cell54))|(X4=cell45&X3=cell55)))), inference(apply_def,[status(thm)],[c_0_811, c_0_830])).
cnf(c_0_841,plain,(epred27_2(X1,X2)|X2=cell46|X2=cell47|X2=cell48|X2=cell49|X2=cell50|~epred28_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_831])).
cnf(c_0_842,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred28_2(cell2,esk2_0)|epred28_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_832, c_0_833]), c_0_785]), c_0_786]), c_0_787]), c_0_788]), c_0_789])).
cnf(c_0_843,plain,(cell46!=cell2), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_844,plain,(cell47!=cell2), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_845,plain,(cell48!=cell2), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_846,plain,(cell49!=cell2), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_847,plain,(cell50!=cell2), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_848,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred29_2(cell99,esk1_0)|epred28_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_822, c_0_834]), c_0_835]), c_0_836]), c_0_837]), c_0_838]), c_0_839])).
fof(c_0_849, plain, (![X4]:![X3]:(epred25_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40))|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45)))), introduced(definition)).
fof(c_0_850, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell45|(X5=cell44|(X5=cell43|(X5=cell42|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X5=cell44|(X5=cell43|(X5=cell42|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell45|(X6=cell54|(X5=cell43|(X5=cell42|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X6=cell54|(X5=cell43|(X5=cell42|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))&((((X5=cell45|(X5=cell44|(X6=cell53|(X5=cell42|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X5=cell44|(X6=cell53|(X5=cell42|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell45|(X6=cell54|(X6=cell53|(X5=cell42|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X6=cell54|(X6=cell53|(X5=cell42|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))))&(((((X5=cell45|(X5=cell44|(X5=cell43|(X6=cell52|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X5=cell44|(X5=cell43|(X6=cell52|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell45|(X6=cell54|(X5=cell43|(X6=cell52|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X6=cell54|(X5=cell43|(X6=cell52|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))&((((X5=cell45|(X5=cell44|(X6=cell53|(X6=cell52|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X5=cell44|(X6=cell53|(X6=cell52|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell45|(X6=cell54|(X6=cell53|(X6=cell52|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X6=cell54|(X6=cell53|(X6=cell52|(X5=cell41|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))))&((((((X5=cell45|(X5=cell44|(X5=cell43|(X5=cell42|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X5=cell44|(X5=cell43|(X5=cell42|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell45|(X6=cell54|(X5=cell43|(X5=cell42|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X6=cell54|(X5=cell43|(X5=cell42|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))&((((X5=cell45|(X5=cell44|(X6=cell53|(X5=cell42|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X5=cell44|(X6=cell53|(X5=cell42|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell45|(X6=cell54|(X6=cell53|(X5=cell42|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X6=cell54|(X6=cell53|(X5=cell42|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))))&(((((X5=cell45|(X5=cell44|(X5=cell43|(X6=cell52|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X5=cell44|(X5=cell43|(X6=cell52|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell45|(X6=cell54|(X5=cell43|(X6=cell52|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X6=cell54|(X5=cell43|(X6=cell52|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))&((((X5=cell45|(X5=cell44|(X6=cell53|(X6=cell52|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X5=cell44|(X6=cell53|(X6=cell52|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell45|(X6=cell54|(X6=cell53|(X6=cell52|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell55|(X6=cell54|(X6=cell53|(X6=cell52|(X6=cell51|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))))))&((((((~epred26_2(X8,X7)|epred27_2(X8,X7))&((X7!=cell41|X8!=cell51)|epred27_2(X8,X7)))&((X7!=cell42|X8!=cell52)|epred27_2(X8,X7)))&((X7!=cell43|X8!=cell53)|epred27_2(X8,X7)))&((X7!=cell44|X8!=cell54)|epred27_2(X8,X7)))&((X7!=cell45|X8!=cell55)|epred27_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_840])])])])])).
cnf(c_0_851,plain,(epred27_2(X1,X2)|X1=cell56|X1=cell57|X1=cell58|X1=cell59|X1=cell60|~epred28_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_831])).
cnf(c_0_852,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred28_2(cell2,esk2_0)|epred27_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_841, c_0_842]), c_0_843]), c_0_844]), c_0_845]), c_0_846]), c_0_847])).
cnf(c_0_853,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred28_2(cell99,esk1_0)|epred28_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_832, c_0_848]), c_0_797]), c_0_798]), c_0_799]), c_0_800]), c_0_801])).
cnf(c_0_854,plain,(cell46!=cell99), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_855,plain,(cell47!=cell99), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_856,plain,(cell48!=cell99), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_857,plain,(cell49!=cell99), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_858,plain,(cell50!=cell99), inference(split_conjunct,[status(thm)],[c_0_115])).
fof(c_0_859, plain, (![X4]:![X3]:(epred26_2(X3,X4)<=>(((((epred25_2(X3,X4)|(X4=cell36&X3=cell46))|(X4=cell37&X3=cell47))|(X4=cell38&X3=cell48))|(X4=cell39&X3=cell49))|(X4=cell40&X3=cell50)))), inference(apply_def,[status(thm)],[c_0_830, c_0_849])).
cnf(c_0_860,plain,(epred26_2(X1,X2)|X2=cell41|X2=cell42|X2=cell43|X2=cell44|X2=cell45|~epred27_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_850])).
cnf(c_0_861,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred27_2(cell2,esk2_0)|epred27_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_851, c_0_852]), c_0_805]), c_0_806]), c_0_807]), c_0_808]), c_0_809])).
cnf(c_0_862,plain,(cell41!=cell2), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_863,plain,(cell42!=cell2), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_864,plain,(cell43!=cell2), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_865,plain,(cell44!=cell2), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_866,plain,(cell45!=cell2), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_867,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred28_2(cell99,esk1_0)|epred27_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_841, c_0_853]), c_0_854]), c_0_855]), c_0_856]), c_0_857]), c_0_858])).
fof(c_0_868, plain, (![X4]:![X3]:(epred24_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35))|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40)))), introduced(definition)).
fof(c_0_869, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell40|(X5=cell39|(X5=cell38|(X5=cell37|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X5=cell39|(X5=cell38|(X5=cell37|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell40|(X6=cell49|(X5=cell38|(X5=cell37|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell38|(X5=cell37|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))&((((X5=cell40|(X5=cell39|(X6=cell48|(X5=cell37|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X5=cell39|(X6=cell48|(X5=cell37|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell40|(X6=cell49|(X6=cell48|(X5=cell37|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X5=cell37|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))))&(((((X5=cell40|(X5=cell39|(X5=cell38|(X6=cell47|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X5=cell39|(X5=cell38|(X6=cell47|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell40|(X6=cell49|(X5=cell38|(X6=cell47|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell38|(X6=cell47|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))&((((X5=cell40|(X5=cell39|(X6=cell48|(X6=cell47|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X5=cell39|(X6=cell48|(X6=cell47|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell40|(X6=cell49|(X6=cell48|(X6=cell47|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X6=cell47|(X5=cell36|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))))&((((((X5=cell40|(X5=cell39|(X5=cell38|(X5=cell37|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X5=cell39|(X5=cell38|(X5=cell37|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell40|(X6=cell49|(X5=cell38|(X5=cell37|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell38|(X5=cell37|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))&((((X5=cell40|(X5=cell39|(X6=cell48|(X5=cell37|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X5=cell39|(X6=cell48|(X5=cell37|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell40|(X6=cell49|(X6=cell48|(X5=cell37|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X5=cell37|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))))&(((((X5=cell40|(X5=cell39|(X5=cell38|(X6=cell47|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X5=cell39|(X5=cell38|(X6=cell47|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell40|(X6=cell49|(X5=cell38|(X6=cell47|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell38|(X6=cell47|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))&((((X5=cell40|(X5=cell39|(X6=cell48|(X6=cell47|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X5=cell39|(X6=cell48|(X6=cell47|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell40|(X6=cell49|(X6=cell48|(X6=cell47|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X6=cell47|(X6=cell46|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))))))&((((((~epred25_2(X8,X7)|epred26_2(X8,X7))&((X7!=cell36|X8!=cell46)|epred26_2(X8,X7)))&((X7!=cell37|X8!=cell47)|epred26_2(X8,X7)))&((X7!=cell38|X8!=cell48)|epred26_2(X8,X7)))&((X7!=cell39|X8!=cell49)|epred26_2(X8,X7)))&((X7!=cell40|X8!=cell50)|epred26_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_859])])])])])).
cnf(c_0_870,plain,(epred26_2(X1,X2)|X1=cell51|X1=cell52|X1=cell53|X1=cell54|X1=cell55|~epred27_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_850])).
cnf(c_0_871,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred27_2(cell2,esk2_0)|epred26_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_860, c_0_861]), c_0_862]), c_0_863]), c_0_864]), c_0_865]), c_0_866])).
cnf(c_0_872,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred27_2(cell99,esk1_0)|epred27_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_851, c_0_867]), c_0_816]), c_0_817]), c_0_818]), c_0_819]), c_0_820])).
cnf(c_0_873,plain,(cell41!=cell99), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_874,plain,(cell42!=cell99), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_875,plain,(cell43!=cell99), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_876,plain,(cell44!=cell99), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_877,plain,(cell45!=cell99), inference(split_conjunct,[status(thm)],[c_0_125])).
fof(c_0_878, plain, (![X4]:![X3]:(epred25_2(X3,X4)<=>(((((epred24_2(X3,X4)|(X4=cell31&X3=cell41))|(X4=cell32&X3=cell42))|(X4=cell33&X3=cell43))|(X4=cell34&X3=cell44))|(X4=cell35&X3=cell45)))), inference(apply_def,[status(thm)],[c_0_849, c_0_868])).
cnf(c_0_879,plain,(epred25_2(X1,X2)|X2=cell36|X2=cell37|X2=cell38|X2=cell39|X2=cell40|~epred26_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_869])).
cnf(c_0_880,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred26_2(cell2,esk2_0)|epred26_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_870, c_0_871]), c_0_824]), c_0_825]), c_0_826]), c_0_827]), c_0_828])).
cnf(c_0_881,plain,(cell36!=cell2), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_882,plain,(cell37!=cell2), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_883,plain,(cell38!=cell2), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_884,plain,(cell39!=cell2), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_885,plain,(cell40!=cell2), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_886,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred27_2(cell99,esk1_0)|epred26_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_860, c_0_872]), c_0_873]), c_0_874]), c_0_875]), c_0_876]), c_0_877])).
fof(c_0_887, plain, (![X4]:![X3]:(epred23_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30))|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35)))), introduced(definition)).
fof(c_0_888, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell35|(X5=cell34|(X5=cell33|(X5=cell32|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X5=cell34|(X5=cell33|(X5=cell32|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell35|(X6=cell44|(X5=cell33|(X5=cell32|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell33|(X5=cell32|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))&((((X5=cell35|(X5=cell34|(X6=cell43|(X5=cell32|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X5=cell34|(X6=cell43|(X5=cell32|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell35|(X6=cell44|(X6=cell43|(X5=cell32|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X5=cell32|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))))&(((((X5=cell35|(X5=cell34|(X5=cell33|(X6=cell42|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X5=cell34|(X5=cell33|(X6=cell42|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell35|(X6=cell44|(X5=cell33|(X6=cell42|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell33|(X6=cell42|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))&((((X5=cell35|(X5=cell34|(X6=cell43|(X6=cell42|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X5=cell34|(X6=cell43|(X6=cell42|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell35|(X6=cell44|(X6=cell43|(X6=cell42|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X6=cell42|(X5=cell31|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))))&((((((X5=cell35|(X5=cell34|(X5=cell33|(X5=cell32|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X5=cell34|(X5=cell33|(X5=cell32|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell35|(X6=cell44|(X5=cell33|(X5=cell32|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell33|(X5=cell32|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))&((((X5=cell35|(X5=cell34|(X6=cell43|(X5=cell32|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X5=cell34|(X6=cell43|(X5=cell32|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell35|(X6=cell44|(X6=cell43|(X5=cell32|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X5=cell32|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))))&(((((X5=cell35|(X5=cell34|(X5=cell33|(X6=cell42|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X5=cell34|(X5=cell33|(X6=cell42|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell35|(X6=cell44|(X5=cell33|(X6=cell42|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell33|(X6=cell42|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))&((((X5=cell35|(X5=cell34|(X6=cell43|(X6=cell42|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X5=cell34|(X6=cell43|(X6=cell42|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell35|(X6=cell44|(X6=cell43|(X6=cell42|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X6=cell42|(X6=cell41|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))))))&((((((~epred24_2(X8,X7)|epred25_2(X8,X7))&((X7!=cell31|X8!=cell41)|epred25_2(X8,X7)))&((X7!=cell32|X8!=cell42)|epred25_2(X8,X7)))&((X7!=cell33|X8!=cell43)|epred25_2(X8,X7)))&((X7!=cell34|X8!=cell44)|epred25_2(X8,X7)))&((X7!=cell35|X8!=cell45)|epred25_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_878])])])])])).
cnf(c_0_889,plain,(epred25_2(X1,X2)|X1=cell46|X1=cell47|X1=cell48|X1=cell49|X1=cell50|~epred26_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_869])).
cnf(c_0_890,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred26_2(cell2,esk2_0)|epred25_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_879, c_0_880]), c_0_881]), c_0_882]), c_0_883]), c_0_884]), c_0_885])).
cnf(c_0_891,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred26_2(cell99,esk1_0)|epred26_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_870, c_0_886]), c_0_835]), c_0_836]), c_0_837]), c_0_838]), c_0_839])).
cnf(c_0_892,plain,(cell36!=cell99), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_893,plain,(cell37!=cell99), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_894,plain,(cell38!=cell99), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_895,plain,(cell39!=cell99), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_896,plain,(cell40!=cell99), inference(split_conjunct,[status(thm)],[c_0_135])).
fof(c_0_897, plain, (![X4]:![X3]:(epred24_2(X3,X4)<=>(((((epred23_2(X3,X4)|(X4=cell26&X3=cell36))|(X4=cell27&X3=cell37))|(X4=cell28&X3=cell38))|(X4=cell29&X3=cell39))|(X4=cell30&X3=cell40)))), inference(apply_def,[status(thm)],[c_0_868, c_0_887])).
cnf(c_0_898,plain,(epred24_2(X1,X2)|X2=cell31|X2=cell32|X2=cell33|X2=cell34|X2=cell35|~epred25_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_888])).
cnf(c_0_899,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred25_2(cell2,esk2_0)|epred25_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_889, c_0_890]), c_0_843]), c_0_844]), c_0_845]), c_0_846]), c_0_847])).
cnf(c_0_900,plain,(cell31!=cell2), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_901,plain,(cell32!=cell2), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_902,plain,(cell33!=cell2), inference(split_conjunct,[status(thm)],[c_0_138])).
cnf(c_0_903,plain,(cell34!=cell2), inference(split_conjunct,[status(thm)],[c_0_139])).
cnf(c_0_904,plain,(cell35!=cell2), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_905,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred26_2(cell99,esk1_0)|epred25_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_879, c_0_891]), c_0_892]), c_0_893]), c_0_894]), c_0_895]), c_0_896])).
fof(c_0_906, plain, (![X4]:![X3]:(epred22_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25))|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30)))), introduced(definition)).
fof(c_0_907, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell30|(X5=cell29|(X5=cell28|(X5=cell27|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X5=cell29|(X5=cell28|(X5=cell27|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell30|(X6=cell39|(X5=cell28|(X5=cell27|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X6=cell39|(X5=cell28|(X5=cell27|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))&((((X5=cell30|(X5=cell29|(X6=cell38|(X5=cell27|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X5=cell29|(X6=cell38|(X5=cell27|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell30|(X6=cell39|(X6=cell38|(X5=cell27|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X6=cell39|(X6=cell38|(X5=cell27|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))))&(((((X5=cell30|(X5=cell29|(X5=cell28|(X6=cell37|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X5=cell29|(X5=cell28|(X6=cell37|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell30|(X6=cell39|(X5=cell28|(X6=cell37|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X6=cell39|(X5=cell28|(X6=cell37|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))&((((X5=cell30|(X5=cell29|(X6=cell38|(X6=cell37|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X5=cell29|(X6=cell38|(X6=cell37|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell30|(X6=cell39|(X6=cell38|(X6=cell37|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X6=cell39|(X6=cell38|(X6=cell37|(X5=cell26|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))))&((((((X5=cell30|(X5=cell29|(X5=cell28|(X5=cell27|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X5=cell29|(X5=cell28|(X5=cell27|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell30|(X6=cell39|(X5=cell28|(X5=cell27|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X6=cell39|(X5=cell28|(X5=cell27|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))&((((X5=cell30|(X5=cell29|(X6=cell38|(X5=cell27|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X5=cell29|(X6=cell38|(X5=cell27|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell30|(X6=cell39|(X6=cell38|(X5=cell27|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X6=cell39|(X6=cell38|(X5=cell27|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))))&(((((X5=cell30|(X5=cell29|(X5=cell28|(X6=cell37|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X5=cell29|(X5=cell28|(X6=cell37|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell30|(X6=cell39|(X5=cell28|(X6=cell37|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X6=cell39|(X5=cell28|(X6=cell37|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))&((((X5=cell30|(X5=cell29|(X6=cell38|(X6=cell37|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X5=cell29|(X6=cell38|(X6=cell37|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell30|(X6=cell39|(X6=cell38|(X6=cell37|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell40|(X6=cell39|(X6=cell38|(X6=cell37|(X6=cell36|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))))))&((((((~epred23_2(X8,X7)|epred24_2(X8,X7))&((X7!=cell26|X8!=cell36)|epred24_2(X8,X7)))&((X7!=cell27|X8!=cell37)|epred24_2(X8,X7)))&((X7!=cell28|X8!=cell38)|epred24_2(X8,X7)))&((X7!=cell29|X8!=cell39)|epred24_2(X8,X7)))&((X7!=cell30|X8!=cell40)|epred24_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_897])])])])])).
cnf(c_0_908,plain,(epred24_2(X1,X2)|X1=cell41|X1=cell42|X1=cell43|X1=cell44|X1=cell45|~epred25_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_888])).
cnf(c_0_909,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred25_2(cell2,esk2_0)|epred24_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_898, c_0_899]), c_0_900]), c_0_901]), c_0_902]), c_0_903]), c_0_904])).
cnf(c_0_910,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred25_2(cell99,esk1_0)|epred25_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_889, c_0_905]), c_0_854]), c_0_855]), c_0_856]), c_0_857]), c_0_858])).
cnf(c_0_911,plain,(cell31!=cell99), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_912,plain,(cell32!=cell99), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_913,plain,(cell33!=cell99), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_914,plain,(cell34!=cell99), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_915,plain,(cell35!=cell99), inference(split_conjunct,[status(thm)],[c_0_145])).
fof(c_0_916, plain, (![X4]:![X3]:(epred23_2(X3,X4)<=>(((((epred22_2(X3,X4)|(X4=cell21&X3=cell31))|(X4=cell22&X3=cell32))|(X4=cell23&X3=cell33))|(X4=cell24&X3=cell34))|(X4=cell25&X3=cell35)))), inference(apply_def,[status(thm)],[c_0_887, c_0_906])).
cnf(c_0_917,plain,(epred23_2(X1,X2)|X2=cell26|X2=cell27|X2=cell28|X2=cell29|X2=cell30|~epred24_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_907])).
cnf(c_0_918,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred24_2(cell2,esk2_0)|epred24_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_908, c_0_909]), c_0_862]), c_0_863]), c_0_864]), c_0_865]), c_0_866])).
cnf(c_0_919,plain,(cell26!=cell2), inference(split_conjunct,[status(thm)],[c_0_146])).
cnf(c_0_920,plain,(cell27!=cell2), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_921,plain,(cell28!=cell2), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_922,plain,(cell29!=cell2), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_923,plain,(cell30!=cell2), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_924,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred25_2(cell99,esk1_0)|epred24_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_898, c_0_910]), c_0_911]), c_0_912]), c_0_913]), c_0_914]), c_0_915])).
fof(c_0_925, plain, (![X4]:![X3]:(epred21_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20))|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25)))), introduced(definition)).
fof(c_0_926, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell25|(X5=cell24|(X5=cell23|(X5=cell22|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X5=cell24|(X5=cell23|(X5=cell22|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell25|(X6=cell34|(X5=cell23|(X5=cell22|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X6=cell34|(X5=cell23|(X5=cell22|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))&((((X5=cell25|(X5=cell24|(X6=cell33|(X5=cell22|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X5=cell24|(X6=cell33|(X5=cell22|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell25|(X6=cell34|(X6=cell33|(X5=cell22|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X6=cell34|(X6=cell33|(X5=cell22|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))))&(((((X5=cell25|(X5=cell24|(X5=cell23|(X6=cell32|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X5=cell24|(X5=cell23|(X6=cell32|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell25|(X6=cell34|(X5=cell23|(X6=cell32|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X6=cell34|(X5=cell23|(X6=cell32|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))&((((X5=cell25|(X5=cell24|(X6=cell33|(X6=cell32|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X5=cell24|(X6=cell33|(X6=cell32|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell25|(X6=cell34|(X6=cell33|(X6=cell32|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X6=cell34|(X6=cell33|(X6=cell32|(X5=cell21|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))))&((((((X5=cell25|(X5=cell24|(X5=cell23|(X5=cell22|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X5=cell24|(X5=cell23|(X5=cell22|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell25|(X6=cell34|(X5=cell23|(X5=cell22|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X6=cell34|(X5=cell23|(X5=cell22|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))&((((X5=cell25|(X5=cell24|(X6=cell33|(X5=cell22|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X5=cell24|(X6=cell33|(X5=cell22|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell25|(X6=cell34|(X6=cell33|(X5=cell22|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X6=cell34|(X6=cell33|(X5=cell22|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))))&(((((X5=cell25|(X5=cell24|(X5=cell23|(X6=cell32|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X5=cell24|(X5=cell23|(X6=cell32|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell25|(X6=cell34|(X5=cell23|(X6=cell32|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X6=cell34|(X5=cell23|(X6=cell32|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))&((((X5=cell25|(X5=cell24|(X6=cell33|(X6=cell32|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X5=cell24|(X6=cell33|(X6=cell32|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell25|(X6=cell34|(X6=cell33|(X6=cell32|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell35|(X6=cell34|(X6=cell33|(X6=cell32|(X6=cell31|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))))))&((((((~epred22_2(X8,X7)|epred23_2(X8,X7))&((X7!=cell21|X8!=cell31)|epred23_2(X8,X7)))&((X7!=cell22|X8!=cell32)|epred23_2(X8,X7)))&((X7!=cell23|X8!=cell33)|epred23_2(X8,X7)))&((X7!=cell24|X8!=cell34)|epred23_2(X8,X7)))&((X7!=cell25|X8!=cell35)|epred23_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_916])])])])])).
cnf(c_0_927,plain,(epred23_2(X1,X2)|X1=cell36|X1=cell37|X1=cell38|X1=cell39|X1=cell40|~epred24_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_907])).
cnf(c_0_928,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred24_2(cell2,esk2_0)|epred23_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_917, c_0_918]), c_0_919]), c_0_920]), c_0_921]), c_0_922]), c_0_923])).
cnf(c_0_929,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred24_2(cell99,esk1_0)|epred24_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_908, c_0_924]), c_0_873]), c_0_874]), c_0_875]), c_0_876]), c_0_877])).
cnf(c_0_930,plain,(cell26!=cell99), inference(split_conjunct,[status(thm)],[c_0_151])).
cnf(c_0_931,plain,(cell27!=cell99), inference(split_conjunct,[status(thm)],[c_0_152])).
cnf(c_0_932,plain,(cell28!=cell99), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_933,plain,(cell29!=cell99), inference(split_conjunct,[status(thm)],[c_0_154])).
cnf(c_0_934,plain,(cell30!=cell99), inference(split_conjunct,[status(thm)],[c_0_155])).
fof(c_0_935, plain, (![X4]:![X3]:(epred22_2(X3,X4)<=>(((((epred21_2(X3,X4)|(X4=cell16&X3=cell26))|(X4=cell17&X3=cell27))|(X4=cell18&X3=cell28))|(X4=cell19&X3=cell29))|(X4=cell20&X3=cell30)))), inference(apply_def,[status(thm)],[c_0_906, c_0_925])).
cnf(c_0_936,plain,(epred22_2(X1,X2)|X2=cell21|X2=cell22|X2=cell23|X2=cell24|X2=cell25|~epred23_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_926])).
cnf(c_0_937,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred23_2(cell2,esk2_0)|epred23_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_927, c_0_928]), c_0_881]), c_0_882]), c_0_883]), c_0_884]), c_0_885])).
cnf(c_0_938,plain,(cell2!=cell22), inference(split_conjunct,[status(thm)],[c_0_156])).
cnf(c_0_939,plain,(cell21!=cell2), inference(split_conjunct,[status(thm)],[c_0_157])).
cnf(c_0_940,plain,(cell23!=cell2), inference(split_conjunct,[status(thm)],[c_0_158])).
cnf(c_0_941,plain,(cell24!=cell2), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_942,plain,(cell25!=cell2), inference(split_conjunct,[status(thm)],[c_0_160])).
cnf(c_0_943,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred24_2(cell99,esk1_0)|epred23_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_917, c_0_929]), c_0_930]), c_0_931]), c_0_932]), c_0_933]), c_0_934])).
fof(c_0_944, plain, (![X4]:![X3]:(epred20_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15))|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20)))), introduced(definition)).
fof(c_0_945, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell20|(X5=cell19|(X5=cell18|(X5=cell17|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X5=cell19|(X5=cell18|(X5=cell17|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell20|(X6=cell29|(X5=cell18|(X5=cell17|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X6=cell29|(X5=cell18|(X5=cell17|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))&((((X5=cell20|(X5=cell19|(X6=cell28|(X5=cell17|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X5=cell19|(X6=cell28|(X5=cell17|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell20|(X6=cell29|(X6=cell28|(X5=cell17|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X6=cell29|(X6=cell28|(X5=cell17|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))))&(((((X5=cell20|(X5=cell19|(X5=cell18|(X6=cell27|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X5=cell19|(X5=cell18|(X6=cell27|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell20|(X6=cell29|(X5=cell18|(X6=cell27|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X6=cell29|(X5=cell18|(X6=cell27|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))&((((X5=cell20|(X5=cell19|(X6=cell28|(X6=cell27|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X5=cell19|(X6=cell28|(X6=cell27|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell20|(X6=cell29|(X6=cell28|(X6=cell27|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X6=cell29|(X6=cell28|(X6=cell27|(X5=cell16|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))))&((((((X5=cell20|(X5=cell19|(X5=cell18|(X5=cell17|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X5=cell19|(X5=cell18|(X5=cell17|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell20|(X6=cell29|(X5=cell18|(X5=cell17|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X6=cell29|(X5=cell18|(X5=cell17|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))&((((X5=cell20|(X5=cell19|(X6=cell28|(X5=cell17|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X5=cell19|(X6=cell28|(X5=cell17|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell20|(X6=cell29|(X6=cell28|(X5=cell17|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X6=cell29|(X6=cell28|(X5=cell17|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))))&(((((X5=cell20|(X5=cell19|(X5=cell18|(X6=cell27|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X5=cell19|(X5=cell18|(X6=cell27|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell20|(X6=cell29|(X5=cell18|(X6=cell27|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X6=cell29|(X5=cell18|(X6=cell27|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))&((((X5=cell20|(X5=cell19|(X6=cell28|(X6=cell27|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X5=cell19|(X6=cell28|(X6=cell27|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell20|(X6=cell29|(X6=cell28|(X6=cell27|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell30|(X6=cell29|(X6=cell28|(X6=cell27|(X6=cell26|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))))))&((((((~epred21_2(X8,X7)|epred22_2(X8,X7))&((X7!=cell16|X8!=cell26)|epred22_2(X8,X7)))&((X7!=cell17|X8!=cell27)|epred22_2(X8,X7)))&((X7!=cell18|X8!=cell28)|epred22_2(X8,X7)))&((X7!=cell19|X8!=cell29)|epred22_2(X8,X7)))&((X7!=cell20|X8!=cell30)|epred22_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_935])])])])])).
cnf(c_0_946,plain,(epred22_2(X1,X2)|X1=cell31|X1=cell32|X1=cell33|X1=cell34|X1=cell35|~epred23_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_926])).
cnf(c_0_947,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred23_2(cell2,esk2_0)|epred22_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_937]), c_0_938]), c_0_939]), c_0_940]), c_0_941]), c_0_942])).
cnf(c_0_948,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred23_2(cell99,esk1_0)|epred23_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_927, c_0_943]), c_0_892]), c_0_893]), c_0_894]), c_0_895]), c_0_896])).
cnf(c_0_949,plain,(cell99!=cell22), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_950,plain,(cell21!=cell99), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_951,plain,(cell23!=cell99), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_952,plain,(cell24!=cell99), inference(split_conjunct,[status(thm)],[c_0_164])).
cnf(c_0_953,plain,(cell25!=cell99), inference(split_conjunct,[status(thm)],[c_0_165])).
fof(c_0_954, plain, (![X4]:![X3]:(epred21_2(X3,X4)<=>(((((epred20_2(X3,X4)|(X4=cell11&X3=cell21))|(X4=cell12&X3=cell22))|(X4=cell13&X3=cell23))|(X4=cell14&X3=cell24))|(X4=cell15&X3=cell25)))), inference(apply_def,[status(thm)],[c_0_925, c_0_944])).
cnf(c_0_955,plain,(epred21_2(X1,X2)|X2=cell16|X2=cell17|X2=cell18|X2=cell19|X2=cell20|~epred22_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_945])).
cnf(c_0_956,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred22_2(cell2,esk2_0)|epred22_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_946, c_0_947]), c_0_900]), c_0_901]), c_0_902]), c_0_903]), c_0_904])).
cnf(c_0_957,plain,(cell16!=cell2), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_958,plain,(cell17!=cell2), inference(split_conjunct,[status(thm)],[c_0_167])).
cnf(c_0_959,plain,(cell18!=cell2), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_960,plain,(cell19!=cell2), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_961,plain,(cell20!=cell2), inference(split_conjunct,[status(thm)],[c_0_170])).
cnf(c_0_962,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred23_2(cell99,esk1_0)|epred22_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_948]), c_0_949]), c_0_950]), c_0_951]), c_0_952]), c_0_953])).
fof(c_0_963, plain, (![X4]:![X3]:(epred19_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100))|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15)))), introduced(definition)).
fof(c_0_964, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell15|(X5=cell14|(X5=cell13|(X5=cell12|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X5=cell14|(X5=cell13|(X5=cell12|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell15|(X6=cell24|(X5=cell13|(X5=cell12|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X6=cell24|(X5=cell13|(X5=cell12|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))&((((X5=cell15|(X5=cell14|(X6=cell23|(X5=cell12|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X5=cell14|(X6=cell23|(X5=cell12|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell15|(X6=cell24|(X6=cell23|(X5=cell12|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X6=cell24|(X6=cell23|(X5=cell12|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))))&(((((X5=cell15|(X5=cell14|(X5=cell13|(X6=cell22|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X5=cell14|(X5=cell13|(X6=cell22|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell15|(X6=cell24|(X5=cell13|(X6=cell22|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X6=cell24|(X5=cell13|(X6=cell22|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))&((((X5=cell15|(X5=cell14|(X6=cell23|(X6=cell22|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X5=cell14|(X6=cell23|(X6=cell22|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell15|(X6=cell24|(X6=cell23|(X6=cell22|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X6=cell24|(X6=cell23|(X6=cell22|(X5=cell11|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))))&((((((X5=cell15|(X5=cell14|(X5=cell13|(X5=cell12|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X5=cell14|(X5=cell13|(X5=cell12|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell15|(X6=cell24|(X5=cell13|(X5=cell12|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X6=cell24|(X5=cell13|(X5=cell12|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))&((((X5=cell15|(X5=cell14|(X6=cell23|(X5=cell12|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X5=cell14|(X6=cell23|(X5=cell12|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell15|(X6=cell24|(X6=cell23|(X5=cell12|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X6=cell24|(X6=cell23|(X5=cell12|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))))&(((((X5=cell15|(X5=cell14|(X5=cell13|(X6=cell22|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X5=cell14|(X5=cell13|(X6=cell22|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell15|(X6=cell24|(X5=cell13|(X6=cell22|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X6=cell24|(X5=cell13|(X6=cell22|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))&((((X5=cell15|(X5=cell14|(X6=cell23|(X6=cell22|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X5=cell14|(X6=cell23|(X6=cell22|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell15|(X6=cell24|(X6=cell23|(X6=cell22|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell25|(X6=cell24|(X6=cell23|(X6=cell22|(X6=cell21|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))))))&((((((~epred20_2(X8,X7)|epred21_2(X8,X7))&((X7!=cell11|X8!=cell21)|epred21_2(X8,X7)))&((X7!=cell12|X8!=cell22)|epred21_2(X8,X7)))&((X7!=cell13|X8!=cell23)|epred21_2(X8,X7)))&((X7!=cell14|X8!=cell24)|epred21_2(X8,X7)))&((X7!=cell15|X8!=cell25)|epred21_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_954])])])])])).
cnf(c_0_965,plain,(epred21_2(X1,X2)|X1=cell26|X1=cell27|X1=cell28|X1=cell29|X1=cell30|~epred22_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_945])).
cnf(c_0_966,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred22_2(cell2,esk2_0)|epred21_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_955, c_0_956]), c_0_957]), c_0_958]), c_0_959]), c_0_960]), c_0_961])).
cnf(c_0_967,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred22_2(cell99,esk1_0)|epred22_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_946, c_0_962]), c_0_911]), c_0_912]), c_0_913]), c_0_914]), c_0_915])).
cnf(c_0_968,plain,(cell16!=cell99), inference(split_conjunct,[status(thm)],[c_0_171])).
cnf(c_0_969,plain,(cell17!=cell99), inference(split_conjunct,[status(thm)],[c_0_172])).
cnf(c_0_970,plain,(cell18!=cell99), inference(split_conjunct,[status(thm)],[c_0_173])).
cnf(c_0_971,plain,(cell19!=cell99), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_972,plain,(cell20!=cell99), inference(split_conjunct,[status(thm)],[c_0_175])).
fof(c_0_973, plain, (![X4]:![X3]:(epred20_2(X3,X4)<=>(((((epred19_2(X3,X4)|(X4=cell6&X3=cell16))|(X4=cell7&X3=cell17))|(X4=cell8&X3=cell18))|(X4=cell9&X3=cell19))|(X4=cell10&X3=cell20)))), inference(apply_def,[status(thm)],[c_0_944, c_0_963])).
cnf(c_0_974,plain,(epred20_2(X1,X2)|X2=cell11|X2=cell12|X2=cell13|X2=cell14|X2=cell15|~epred21_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_964])).
cnf(c_0_975,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred21_2(cell2,esk2_0)|epred21_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_965, c_0_966]), c_0_919]), c_0_920]), c_0_921]), c_0_922]), c_0_923])).
cnf(c_0_976,plain,(cell12!=cell2), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_977,plain,(cell11!=cell2), inference(split_conjunct,[status(thm)],[c_0_177])).
cnf(c_0_978,plain,(cell2!=cell13), inference(split_conjunct,[status(thm)],[c_0_178])).
cnf(c_0_979,plain,(cell14!=cell2), inference(split_conjunct,[status(thm)],[c_0_179])).
cnf(c_0_980,plain,(cell15!=cell2), inference(split_conjunct,[status(thm)],[c_0_180])).
cnf(c_0_981,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred22_2(cell99,esk1_0)|epred21_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_955, c_0_967]), c_0_968]), c_0_969]), c_0_970]), c_0_971]), c_0_972])).
fof(c_0_982, plain, (![X4]:![X3]:(epred18_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100)))), introduced(definition)).
fof(c_0_983, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell10|(X5=cell9|(X5=cell8|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X5=cell9|(X5=cell8|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell10|(X6=cell19|(X5=cell8|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X6=cell19|(X5=cell8|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))&((((X5=cell10|(X5=cell9|(X6=cell18|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X5=cell9|(X6=cell18|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell10|(X6=cell19|(X6=cell18|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X6=cell19|(X6=cell18|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))))&(((((X5=cell10|(X5=cell9|(X5=cell8|(X6=cell17|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X5=cell9|(X5=cell8|(X6=cell17|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell10|(X6=cell19|(X5=cell8|(X6=cell17|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X6=cell19|(X5=cell8|(X6=cell17|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))&((((X5=cell10|(X5=cell9|(X6=cell18|(X6=cell17|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X5=cell9|(X6=cell18|(X6=cell17|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell10|(X6=cell19|(X6=cell18|(X6=cell17|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X6=cell19|(X6=cell18|(X6=cell17|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))))&((((((X5=cell10|(X5=cell9|(X5=cell8|(X5=cell7|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X5=cell9|(X5=cell8|(X5=cell7|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell10|(X6=cell19|(X5=cell8|(X5=cell7|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X6=cell19|(X5=cell8|(X5=cell7|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))&((((X5=cell10|(X5=cell9|(X6=cell18|(X5=cell7|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X5=cell9|(X6=cell18|(X5=cell7|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell10|(X6=cell19|(X6=cell18|(X5=cell7|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X6=cell19|(X6=cell18|(X5=cell7|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))))&(((((X5=cell10|(X5=cell9|(X5=cell8|(X6=cell17|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X5=cell9|(X5=cell8|(X6=cell17|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell10|(X6=cell19|(X5=cell8|(X6=cell17|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X6=cell19|(X5=cell8|(X6=cell17|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))&((((X5=cell10|(X5=cell9|(X6=cell18|(X6=cell17|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X5=cell9|(X6=cell18|(X6=cell17|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell10|(X6=cell19|(X6=cell18|(X6=cell17|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell20|(X6=cell19|(X6=cell18|(X6=cell17|(X6=cell16|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))))))&((((((~epred19_2(X8,X7)|epred20_2(X8,X7))&((X7!=cell6|X8!=cell16)|epred20_2(X8,X7)))&((X7!=cell7|X8!=cell17)|epred20_2(X8,X7)))&((X7!=cell8|X8!=cell18)|epred20_2(X8,X7)))&((X7!=cell9|X8!=cell19)|epred20_2(X8,X7)))&((X7!=cell10|X8!=cell20)|epred20_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_973])])])])])).
cnf(c_0_984,plain,(epred20_2(X1,X2)|X1=cell21|X1=cell22|X1=cell23|X1=cell24|X1=cell25|~epred21_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_964])).
cnf(c_0_985,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred21_2(cell2,esk2_0)|epred20_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_974, c_0_975]), c_0_976]), c_0_977]), c_0_978]), c_0_979]), c_0_980])).
cnf(c_0_986,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred21_2(cell99,esk1_0)|epred21_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_965, c_0_981]), c_0_930]), c_0_931]), c_0_932]), c_0_933]), c_0_934])).
cnf(c_0_987,plain,(cell99!=cell12), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_988,plain,(cell99!=cell11), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_989,plain,(cell99!=cell13), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_990,plain,(cell14!=cell99), inference(split_conjunct,[status(thm)],[c_0_184])).
cnf(c_0_991,plain,(cell15!=cell99), inference(split_conjunct,[status(thm)],[c_0_185])).
fof(c_0_992, plain, (![X4]:![X3]:(epred19_2(X3,X4)<=>(((((epred18_2(X3,X4)|(X4=cell1&X3=cell11))|(X4=cell2&X3=cell12))|(X4=cell3&X3=cell13))|(X4=cell4&X3=cell14))|(X4=cell5&X3=cell15)))), inference(apply_def,[status(thm)],[c_0_963, c_0_982])).
cnf(c_0_993,plain,(epred19_2(X1,X2)|X2=cell6|X2=cell7|X2=cell8|X2=cell9|X2=cell10|~epred20_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_983])).
cnf(c_0_994,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred20_2(cell2,esk2_0)|epred20_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_984, c_0_985]), c_0_938]), c_0_939]), c_0_940]), c_0_941]), c_0_942])).
cnf(c_0_995,plain,(cell6!=cell2), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_996,plain,(cell7!=cell2), inference(split_conjunct,[status(thm)],[c_0_187])).
cnf(c_0_997,plain,(cell8!=cell2), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_998,plain,(cell9!=cell2), inference(split_conjunct,[status(thm)],[c_0_189])).
cnf(c_0_999,plain,(cell10!=cell2), inference(split_conjunct,[status(thm)],[c_0_190])).
cnf(c_0_1000,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred21_2(cell99,esk1_0)|epred20_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_974, c_0_986]), c_0_987]), c_0_988]), c_0_989]), c_0_990]), c_0_991])).
fof(c_0_1001, plain, (![X4]:![X3]:(epred17_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95)))), introduced(definition)).
fof(c_0_1002, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell5|(X5=cell4|(X5=cell3|(X5=cell2|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X5=cell4|(X5=cell3|(X5=cell2|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell14|(X5=cell3|(X5=cell2|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X6=cell14|(X5=cell3|(X5=cell2|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell13|(X5=cell2|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X5=cell4|(X6=cell13|(X5=cell2|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell14|(X6=cell13|(X5=cell2|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X6=cell14|(X6=cell13|(X5=cell2|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))))&(((((X5=cell5|(X5=cell4|(X5=cell3|(X6=cell12|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X5=cell4|(X5=cell3|(X6=cell12|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell14|(X5=cell3|(X6=cell12|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X6=cell14|(X5=cell3|(X6=cell12|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell13|(X6=cell12|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X5=cell4|(X6=cell13|(X6=cell12|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell14|(X6=cell13|(X6=cell12|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X6=cell14|(X6=cell13|(X6=cell12|(X5=cell1|epred18_2(X6,X5))))))|~epred19_2(X6,X5))))))&((((((X5=cell5|(X5=cell4|(X5=cell3|(X5=cell2|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X5=cell4|(X5=cell3|(X5=cell2|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell14|(X5=cell3|(X5=cell2|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X6=cell14|(X5=cell3|(X5=cell2|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell13|(X5=cell2|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X5=cell4|(X6=cell13|(X5=cell2|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell14|(X6=cell13|(X5=cell2|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X6=cell14|(X6=cell13|(X5=cell2|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))))&(((((X5=cell5|(X5=cell4|(X5=cell3|(X6=cell12|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X5=cell4|(X5=cell3|(X6=cell12|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell14|(X5=cell3|(X6=cell12|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X6=cell14|(X5=cell3|(X6=cell12|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell13|(X6=cell12|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X5=cell4|(X6=cell13|(X6=cell12|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell14|(X6=cell13|(X6=cell12|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5))&((X6=cell15|(X6=cell14|(X6=cell13|(X6=cell12|(X6=cell11|epred18_2(X6,X5))))))|~epred19_2(X6,X5)))))))&((((((~epred18_2(X8,X7)|epred19_2(X8,X7))&((X7!=cell1|X8!=cell11)|epred19_2(X8,X7)))&((X7!=cell2|X8!=cell12)|epred19_2(X8,X7)))&((X7!=cell3|X8!=cell13)|epred19_2(X8,X7)))&((X7!=cell4|X8!=cell14)|epred19_2(X8,X7)))&((X7!=cell5|X8!=cell15)|epred19_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_992])])])])])).
cnf(c_0_1003,plain,(epred19_2(X1,X2)|X1=cell16|X1=cell17|X1=cell18|X1=cell19|X1=cell20|~epred20_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_983])).
cnf(c_0_1004,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred20_2(cell2,esk2_0)|epred19_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_993, c_0_994]), c_0_995]), c_0_996]), c_0_997]), c_0_998]), c_0_999])).
cnf(c_0_1005,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred20_2(cell99,esk1_0)|epred20_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_984, c_0_1000]), c_0_949]), c_0_950]), c_0_951]), c_0_952]), c_0_953])).
cnf(c_0_1006,plain,(cell6!=cell99), inference(split_conjunct,[status(thm)],[c_0_191])).
cnf(c_0_1007,plain,(cell7!=cell99), inference(split_conjunct,[status(thm)],[c_0_192])).
cnf(c_0_1008,plain,(cell8!=cell99), inference(split_conjunct,[status(thm)],[c_0_193])).
cnf(c_0_1009,plain,(cell9!=cell99), inference(split_conjunct,[status(thm)],[c_0_194])).
cnf(c_0_1010,plain,(cell10!=cell99), inference(split_conjunct,[status(thm)],[c_0_195])).
fof(c_0_1011, plain, (![X4]:![X3]:(epred18_2(X3,X4)<=>(((((epred17_2(X3,X4)|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100)))), inference(apply_def,[status(thm)],[c_0_982, c_0_1001])).
cnf(c_0_1012,plain,(epred18_2(X1,X2)|X2=cell1|X1=cell12|X2=cell3|X2=cell4|X2=cell5|~epred19_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1002])).
cnf(c_0_1013,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred19_2(cell2,esk2_0)|epred19_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1004]), c_0_957]), c_0_958]), c_0_959]), c_0_960]), c_0_961])).
cnf(c_0_1014,plain,(cell1!=cell2), inference(split_conjunct,[status(thm)],[c_0_196])).
cnf(c_0_1015,plain,(cell3!=cell2), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_1016,plain,(cell4!=cell2), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_1017,plain,(cell5!=cell2), inference(split_conjunct,[status(thm)],[c_0_199])).
cnf(c_0_1018,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred20_2(cell99,esk1_0)|epred19_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_993, c_0_1005]), c_0_1006]), c_0_1007]), c_0_1008]), c_0_1009]), c_0_1010])).
fof(c_0_1019, plain, (![X4]:![X3]:(epred16_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89)))), introduced(definition)).
fof(c_0_1020, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell99|(X5=cell98|(X5=cell97|(X5=cell96|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X5=cell98|(X5=cell97|(X5=cell96|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X5=cell97|(X5=cell96|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell97|(X5=cell96|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))))&((((X5=cell99|(X5=cell98|(X6=cell98|(X5=cell96|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X5=cell98|(X6=cell98|(X5=cell96|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X6=cell98|(X5=cell96|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X5=cell96|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))))&(((((X5=cell99|(X5=cell98|(X5=cell97|(X6=cell97|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X5=cell98|(X5=cell97|(X6=cell97|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X5=cell97|(X6=cell97|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell97|(X6=cell97|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))))&((((X5=cell99|(X5=cell98|(X6=cell98|(X6=cell97|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X5=cell98|(X6=cell98|(X6=cell97|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X6=cell98|(X6=cell97|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X6=cell97|(X5=cell95|epred17_2(X6,X5))))))|~epred18_2(X6,X5))))))&((((((X5=cell99|(X5=cell98|(X5=cell97|(X5=cell96|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X5=cell98|(X5=cell97|(X5=cell96|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X5=cell97|(X5=cell96|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell97|(X5=cell96|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))))&((((X5=cell99|(X5=cell98|(X6=cell98|(X5=cell96|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X5=cell98|(X6=cell98|(X5=cell96|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X6=cell98|(X5=cell96|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X5=cell96|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))))&(((((X5=cell99|(X5=cell98|(X5=cell97|(X6=cell97|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X5=cell98|(X5=cell97|(X6=cell97|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X5=cell97|(X6=cell97|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell97|(X6=cell97|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))))&((((X5=cell99|(X5=cell98|(X6=cell98|(X6=cell97|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X5=cell98|(X6=cell98|(X6=cell97|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X6=cell98|(X6=cell97|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X6=cell97|(X6=cell96|epred17_2(X6,X5))))))|~epred18_2(X6,X5)))))))&((((((~epred17_2(X8,X7)|epred18_2(X8,X7))&((X7!=cell95|X8!=cell96)|epred18_2(X8,X7)))&((X7!=cell96|X8!=cell97)|epred18_2(X8,X7)))&((X7!=cell97|X8!=cell98)|epred18_2(X8,X7)))&((X7!=cell98|X8!=cell99)|epred18_2(X8,X7)))&((X7!=cell99|X8!=cell100)|epred18_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1011])])])])])).
cnf(c_0_1021,plain,(epred18_2(X1,X2)|X1=cell11|X1=cell12|X1=cell13|X1=cell14|X1=cell15|~epred19_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1002])).
cnf(c_0_1022,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred19_2(cell2,esk2_0)|epred18_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1012, c_0_1013]), c_0_1014]), c_0_1015]), c_0_1016]), c_0_1017])).
cnf(c_0_1023,plain,(epred18_2(X1,X2)|X2=cell1|X2=cell2|X2=cell3|X2=cell4|X2=cell5|~epred19_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1002])).
cnf(c_0_1024,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred19_2(cell99,esk1_0)|epred19_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1018]), c_0_968]), c_0_969]), c_0_970]), c_0_971]), c_0_972])).
cnf(c_0_1025,plain,(cell1!=cell99), inference(split_conjunct,[status(thm)],[c_0_200])).
cnf(c_0_1026,plain,(cell3!=cell99), inference(split_conjunct,[status(thm)],[c_0_201])).
cnf(c_0_1027,plain,(cell4!=cell99), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_1028,plain,(cell5!=cell99), inference(split_conjunct,[status(thm)],[c_0_203])).
fof(c_0_1029, plain, (![X4]:![X3]:(epred17_2(X3,X4)<=>(((((epred16_2(X3,X4)|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95)))), inference(apply_def,[status(thm)],[c_0_1001, c_0_1019])).
cnf(c_0_1030,plain,(epred17_2(X1,X2)|X2=cell95|X2=cell96|X2=cell97|X2=cell98|X2=cell99|~epred18_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1020])).
cnf(c_0_1031,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred18_2(cell2,esk2_0)|epred18_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1021, c_0_1022]), c_0_976]), c_0_977]), c_0_978]), c_0_979]), c_0_980])).
cnf(c_0_1032,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred19_2(cell99,esk1_0)|epred18_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1023, c_0_1024]), c_0_669]), c_0_1025]), c_0_1026]), c_0_1027]), c_0_1028])).
fof(c_0_1033, plain, (![X4]:![X3]:(epred15_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84)))), introduced(definition)).
fof(c_0_1034, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell94|(X5=cell93|(X5=cell92|(X5=cell91|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X5=cell93|(X5=cell92|(X5=cell91|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X5=cell92|(X5=cell91|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell92|(X5=cell91|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))))&((((X5=cell94|(X5=cell93|(X6=cell93|(X5=cell91|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X5=cell93|(X6=cell93|(X5=cell91|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X6=cell93|(X5=cell91|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X5=cell91|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))))&(((((X5=cell94|(X5=cell93|(X5=cell92|(X6=cell92|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X5=cell93|(X5=cell92|(X6=cell92|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X5=cell92|(X6=cell92|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell92|(X6=cell92|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))))&((((X5=cell94|(X5=cell93|(X6=cell93|(X6=cell92|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X5=cell93|(X6=cell93|(X6=cell92|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X6=cell93|(X6=cell92|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X6=cell92|(X5=cell89|epred16_2(X6,X5))))))|~epred17_2(X6,X5))))))&((((((X5=cell94|(X5=cell93|(X5=cell92|(X5=cell91|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X5=cell93|(X5=cell92|(X5=cell91|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X5=cell92|(X5=cell91|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell92|(X5=cell91|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))))&((((X5=cell94|(X5=cell93|(X6=cell93|(X5=cell91|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X5=cell93|(X6=cell93|(X5=cell91|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X6=cell93|(X5=cell91|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X5=cell91|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))))&(((((X5=cell94|(X5=cell93|(X5=cell92|(X6=cell92|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X5=cell93|(X5=cell92|(X6=cell92|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X5=cell92|(X6=cell92|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell92|(X6=cell92|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))))&((((X5=cell94|(X5=cell93|(X6=cell93|(X6=cell92|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X5=cell93|(X6=cell93|(X6=cell92|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X6=cell93|(X6=cell92|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X6=cell92|(X6=cell90|epred16_2(X6,X5))))))|~epred17_2(X6,X5)))))))&((((((~epred16_2(X8,X7)|epred17_2(X8,X7))&((X7!=cell89|X8!=cell90)|epred17_2(X8,X7)))&((X7!=cell91|X8!=cell92)|epred17_2(X8,X7)))&((X7!=cell92|X8!=cell93)|epred17_2(X8,X7)))&((X7!=cell93|X8!=cell94)|epred17_2(X8,X7)))&((X7!=cell94|X8!=cell95)|epred17_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1029])])])])])).
cnf(c_0_1035,plain,(epred17_2(X1,X2)|X1=cell96|X1=cell97|X1=cell98|X1=cell99|X1=cell100|~epred18_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1020])).
cnf(c_0_1036,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred18_2(cell2,esk2_0)|epred17_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1030, c_0_1031]), c_0_669]), c_0_713]), c_0_670]), c_0_671]), c_0_672])).
cnf(c_0_1037,plain,(epred17_2(X1,X2)|X2=cell95|X2=cell96|X2=cell97|X2=cell98|X1=cell100|~epred18_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1020])).
cnf(c_0_1038,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred18_2(cell99,esk1_0)|epred18_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1021, c_0_1032]), c_0_987]), c_0_988]), c_0_989]), c_0_990]), c_0_991])).
fof(c_0_1039, plain, (![X4]:![X3]:(epred16_2(X3,X4)<=>(((((epred15_2(X3,X4)|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89)))), inference(apply_def,[status(thm)],[c_0_1019, c_0_1033])).
cnf(c_0_1040,plain,(epred16_2(X1,X2)|X2=cell89|X2=cell91|X2=cell92|X2=cell93|X2=cell94|~epred17_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1034])).
cnf(c_0_1041,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred17_2(cell2,esk2_0)|epred17_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1035, c_0_1036]), c_0_669]), c_0_670]), c_0_671]), c_0_672]), c_0_673])).
cnf(c_0_1042,plain,(epred17_2(X1,X2)|X1=cell96|X1=cell97|X1=cell98|X2=cell98|X1=cell100|~epred18_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1020])).
cnf(c_0_1043,plain,(esk1_0=cell100|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred18_2(cell99,esk1_0)|epred17_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1037, c_0_1038]), c_0_733]), c_0_701]), c_0_702]), c_0_703])).
fof(c_0_1044, plain, (![X4]:![X3]:(epred14_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78)))), introduced(definition)).
fof(c_0_1045, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell88|(X5=cell87|(X5=cell86|(X5=cell85|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X5=cell87|(X5=cell86|(X5=cell85|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X5=cell86|(X5=cell85|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X6=cell88|(X5=cell86|(X5=cell85|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))))&((((X5=cell88|(X5=cell87|(X6=cell87|(X5=cell85|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X5=cell87|(X6=cell87|(X5=cell85|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X6=cell87|(X5=cell85|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X6=cell88|(X6=cell87|(X5=cell85|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))))&(((((X5=cell88|(X5=cell87|(X5=cell86|(X6=cell86|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X5=cell87|(X5=cell86|(X6=cell86|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X5=cell86|(X6=cell86|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X6=cell88|(X5=cell86|(X6=cell86|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))))&((((X5=cell88|(X5=cell87|(X6=cell87|(X6=cell86|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X5=cell87|(X6=cell87|(X6=cell86|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X6=cell87|(X6=cell86|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X6=cell88|(X6=cell87|(X6=cell86|(X5=cell84|epred15_2(X6,X5))))))|~epred16_2(X6,X5))))))&((((((X5=cell88|(X5=cell87|(X5=cell86|(X5=cell85|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X5=cell87|(X5=cell86|(X5=cell85|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X5=cell86|(X5=cell85|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X6=cell88|(X5=cell86|(X5=cell85|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))))&((((X5=cell88|(X5=cell87|(X6=cell87|(X5=cell85|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X5=cell87|(X6=cell87|(X5=cell85|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X6=cell87|(X5=cell85|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X6=cell88|(X6=cell87|(X5=cell85|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))))&(((((X5=cell88|(X5=cell87|(X5=cell86|(X6=cell86|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X5=cell87|(X5=cell86|(X6=cell86|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X5=cell86|(X6=cell86|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X6=cell88|(X5=cell86|(X6=cell86|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))))&((((X5=cell88|(X5=cell87|(X6=cell87|(X6=cell86|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X5=cell87|(X6=cell87|(X6=cell86|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X6=cell87|(X6=cell86|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5))&((X6=cell89|(X6=cell88|(X6=cell87|(X6=cell86|(X6=cell85|epred15_2(X6,X5))))))|~epred16_2(X6,X5)))))))&((((((~epred15_2(X8,X7)|epred16_2(X8,X7))&((X7!=cell84|X8!=cell85)|epred16_2(X8,X7)))&((X7!=cell85|X8!=cell86)|epred16_2(X8,X7)))&((X7!=cell86|X8!=cell87)|epred16_2(X8,X7)))&((X7!=cell87|X8!=cell88)|epred16_2(X8,X7)))&((X7!=cell88|X8!=cell89)|epred16_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1039])])])])])).
cnf(c_0_1046,plain,(epred16_2(X1,X2)|X1=cell90|X1=cell92|X1=cell93|X1=cell94|X1=cell95|~epred17_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1034])).
cnf(c_0_1047,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred17_2(cell2,esk2_0)|epred16_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1040, c_0_1041]), c_0_679]), c_0_709]), c_0_710]), c_0_711]), c_0_712])).
cnf(c_0_1048,plain,(esk1_0=cell98|esk1_0=cell100|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred17_2(cell99,esk1_0)|epred17_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1042, c_0_1043]), c_0_701]), c_0_702]), c_0_703]), c_0_704])).
fof(c_0_1049, plain, (![X4]:![X3]:(epred15_2(X3,X4)<=>(((((epred14_2(X3,X4)|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84)))), inference(apply_def,[status(thm)],[c_0_1033, c_0_1044])).
cnf(c_0_1050,plain,(epred15_2(X1,X2)|X2=cell84|X2=cell85|X2=cell86|X2=cell87|X2=cell88|~epred16_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1045])).
cnf(c_0_1051,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred16_2(cell2,esk2_0)|epred16_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1046, c_0_1047]), c_0_681]), c_0_710]), c_0_711]), c_0_712]), c_0_713])).
cnf(c_0_1052,plain,(esk1_0=cell98|esk1_0=cell100|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred17_2(cell99,esk1_0)|epred16_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1040, c_0_1048]), c_0_686]), c_0_729]), c_0_730]), c_0_731]), c_0_732])).
fof(c_0_1053, plain, (![X4]:![X3]:(epred13_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73)))), introduced(definition)).
fof(c_0_1054, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell83|(X5=cell82|(X5=cell81|(X5=cell79|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X5=cell82|(X5=cell81|(X5=cell79|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X5=cell81|(X5=cell79|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X6=cell83|(X5=cell81|(X5=cell79|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))))&((((X5=cell83|(X5=cell82|(X6=cell82|(X5=cell79|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X5=cell82|(X6=cell82|(X5=cell79|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X6=cell82|(X5=cell79|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X6=cell83|(X6=cell82|(X5=cell79|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))))&(((((X5=cell83|(X5=cell82|(X5=cell81|(X6=cell80|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X5=cell82|(X5=cell81|(X6=cell80|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X5=cell81|(X6=cell80|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X6=cell83|(X5=cell81|(X6=cell80|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))))&((((X5=cell83|(X5=cell82|(X6=cell82|(X6=cell80|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X5=cell82|(X6=cell82|(X6=cell80|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X6=cell82|(X6=cell80|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X6=cell83|(X6=cell82|(X6=cell80|(X5=cell78|epred14_2(X6,X5))))))|~epred15_2(X6,X5))))))&((((((X5=cell83|(X5=cell82|(X5=cell81|(X5=cell79|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X5=cell82|(X5=cell81|(X5=cell79|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X5=cell81|(X5=cell79|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X6=cell83|(X5=cell81|(X5=cell79|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))))&((((X5=cell83|(X5=cell82|(X6=cell82|(X5=cell79|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X5=cell82|(X6=cell82|(X5=cell79|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X6=cell82|(X5=cell79|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X6=cell83|(X6=cell82|(X5=cell79|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))))&(((((X5=cell83|(X5=cell82|(X5=cell81|(X6=cell80|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X5=cell82|(X5=cell81|(X6=cell80|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X5=cell81|(X6=cell80|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X6=cell83|(X5=cell81|(X6=cell80|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))))&((((X5=cell83|(X5=cell82|(X6=cell82|(X6=cell80|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X5=cell82|(X6=cell82|(X6=cell80|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X6=cell82|(X6=cell80|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5))&((X6=cell84|(X6=cell83|(X6=cell82|(X6=cell80|(X6=cell79|epred14_2(X6,X5))))))|~epred15_2(X6,X5)))))))&((((((~epred14_2(X8,X7)|epred15_2(X8,X7))&((X7!=cell78|X8!=cell79)|epred15_2(X8,X7)))&((X7!=cell79|X8!=cell80)|epred15_2(X8,X7)))&((X7!=cell81|X8!=cell82)|epred15_2(X8,X7)))&((X7!=cell82|X8!=cell83)|epred15_2(X8,X7)))&((X7!=cell83|X8!=cell84)|epred15_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1049])])])])])).
cnf(c_0_1055,plain,(epred15_2(X1,X2)|X1=cell85|X1=cell86|X1=cell87|X1=cell88|X1=cell89|~epred16_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1045])).
cnf(c_0_1056,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred16_2(cell2,esk2_0)|epred15_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1050, c_0_1051]), c_0_680]), c_0_697]), c_0_698]), c_0_682]), c_0_683])).
cnf(c_0_1057,plain,(epred15_2(X1,X2)|X2=cell84|X2=cell85|X2=cell86|X2=cell87|X1=cell89|~epred16_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1045])).
cnf(c_0_1058,plain,(esk1_0=cell98|esk1_0=cell100|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred16_2(cell99,esk1_0)|epred16_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1046, c_0_1052]), c_0_688]), c_0_730]), c_0_731]), c_0_732]), c_0_733])).
fof(c_0_1059, plain, (![X4]:![X3]:(epred14_2(X3,X4)<=>(((((epred13_2(X3,X4)|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78)))), inference(apply_def,[status(thm)],[c_0_1044, c_0_1053])).
cnf(c_0_1060,plain,(epred14_2(X1,X2)|X2=cell78|X2=cell79|X2=cell81|X2=cell82|X2=cell83|~epred15_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1054])).
cnf(c_0_1061,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred15_2(cell2,esk2_0)|epred15_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1055, c_0_1056]), c_0_679]), c_0_680]), c_0_698]), c_0_682]), c_0_683])).
cnf(c_0_1062,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred16_2(cell99,esk1_0)|epred15_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1057, c_0_1058]), c_0_718]), c_0_719]), c_0_689]), c_0_690])).
fof(c_0_1063, plain, (![X4]:![X3]:(epred12_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67)))), introduced(definition)).
fof(c_0_1064, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell77|(X5=cell76|(X5=cell75|(X5=cell74|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X5=cell76|(X5=cell75|(X5=cell74|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X5=cell75|(X5=cell74|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X6=cell77|(X5=cell75|(X5=cell74|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))))&((((X5=cell77|(X5=cell76|(X6=cell76|(X5=cell74|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X5=cell76|(X6=cell76|(X5=cell74|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X6=cell76|(X5=cell74|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X6=cell77|(X6=cell76|(X5=cell74|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))))&(((((X5=cell77|(X5=cell76|(X5=cell75|(X6=cell75|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X5=cell76|(X5=cell75|(X6=cell75|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X5=cell75|(X6=cell75|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X6=cell77|(X5=cell75|(X6=cell75|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))))&((((X5=cell77|(X5=cell76|(X6=cell76|(X6=cell75|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X5=cell76|(X6=cell76|(X6=cell75|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X6=cell76|(X6=cell75|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X6=cell77|(X6=cell76|(X6=cell75|(X5=cell73|epred13_2(X6,X5))))))|~epred14_2(X6,X5))))))&((((((X5=cell77|(X5=cell76|(X5=cell75|(X5=cell74|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X5=cell76|(X5=cell75|(X5=cell74|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X5=cell75|(X5=cell74|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X6=cell77|(X5=cell75|(X5=cell74|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))))&((((X5=cell77|(X5=cell76|(X6=cell76|(X5=cell74|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X5=cell76|(X6=cell76|(X5=cell74|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X6=cell76|(X5=cell74|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X6=cell77|(X6=cell76|(X5=cell74|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))))&(((((X5=cell77|(X5=cell76|(X5=cell75|(X6=cell75|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X5=cell76|(X5=cell75|(X6=cell75|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X5=cell75|(X6=cell75|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X6=cell77|(X5=cell75|(X6=cell75|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))))&((((X5=cell77|(X5=cell76|(X6=cell76|(X6=cell75|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X5=cell76|(X6=cell76|(X6=cell75|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X6=cell76|(X6=cell75|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5))&((X6=cell78|(X6=cell77|(X6=cell76|(X6=cell75|(X6=cell74|epred13_2(X6,X5))))))|~epred14_2(X6,X5)))))))&((((((~epred13_2(X8,X7)|epred14_2(X8,X7))&((X7!=cell73|X8!=cell74)|epred14_2(X8,X7)))&((X7!=cell74|X8!=cell75)|epred14_2(X8,X7)))&((X7!=cell75|X8!=cell76)|epred14_2(X8,X7)))&((X7!=cell76|X8!=cell77)|epred14_2(X8,X7)))&((X7!=cell77|X8!=cell78)|epred14_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1059])])])])])).
cnf(c_0_1065,plain,(epred14_2(X1,X2)|X1=cell79|X1=cell80|X1=cell82|X1=cell83|X1=cell84|~epred15_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1054])).
cnf(c_0_1066,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred15_2(cell2,esk2_0)|epred14_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1060, c_0_1061]), c_0_723]), c_0_726]), c_0_694]), c_0_695]), c_0_696])).
cnf(c_0_1067,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred15_2(cell99,esk1_0)|epred15_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1055, c_0_1062]), c_0_686]), c_0_687]), c_0_719]), c_0_689]), c_0_690])).
fof(c_0_1068, plain, (![X4]:![X3]:(epred13_2(X3,X4)<=>(((((epred12_2(X3,X4)|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73)))), inference(apply_def,[status(thm)],[c_0_1053, c_0_1063])).
cnf(c_0_1069,plain,(epred13_2(X1,X2)|X2=cell73|X2=cell74|X2=cell75|X2=cell76|X2=cell77|~epred14_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1064])).
cnf(c_0_1070,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred14_2(cell2,esk2_0)|epred14_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1065, c_0_1066]), c_0_723]), c_0_727]), c_0_695]), c_0_696]), c_0_697])).
cnf(c_0_1071,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred15_2(cell99,esk1_0)|epred14_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1060, c_0_1067]), c_0_791]), c_0_742]), c_0_715]), c_0_716]), c_0_717])).
fof(c_0_1072, plain, (![X4]:![X3]:(epred11_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62)))), introduced(definition)).
fof(c_0_1073, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell72|(X5=cell71|(X5=cell69|(X5=cell68|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X5=cell71|(X5=cell69|(X5=cell68|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X5=cell69|(X5=cell68|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X6=cell72|(X5=cell69|(X5=cell68|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))))&((((X5=cell72|(X5=cell71|(X6=cell70|(X5=cell68|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X5=cell71|(X6=cell70|(X5=cell68|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X6=cell70|(X5=cell68|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X6=cell72|(X6=cell70|(X5=cell68|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))))&(((((X5=cell72|(X5=cell71|(X5=cell69|(X6=cell69|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X5=cell71|(X5=cell69|(X6=cell69|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X5=cell69|(X6=cell69|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X6=cell72|(X5=cell69|(X6=cell69|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))))&((((X5=cell72|(X5=cell71|(X6=cell70|(X6=cell69|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X5=cell71|(X6=cell70|(X6=cell69|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X6=cell70|(X6=cell69|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X6=cell72|(X6=cell70|(X6=cell69|(X5=cell67|epred12_2(X6,X5))))))|~epred13_2(X6,X5))))))&((((((X5=cell72|(X5=cell71|(X5=cell69|(X5=cell68|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X5=cell71|(X5=cell69|(X5=cell68|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X5=cell69|(X5=cell68|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X6=cell72|(X5=cell69|(X5=cell68|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))))&((((X5=cell72|(X5=cell71|(X6=cell70|(X5=cell68|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X5=cell71|(X6=cell70|(X5=cell68|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X6=cell70|(X5=cell68|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X6=cell72|(X6=cell70|(X5=cell68|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))))&(((((X5=cell72|(X5=cell71|(X5=cell69|(X6=cell69|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X5=cell71|(X5=cell69|(X6=cell69|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X5=cell69|(X6=cell69|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X6=cell72|(X5=cell69|(X6=cell69|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))))&((((X5=cell72|(X5=cell71|(X6=cell70|(X6=cell69|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X5=cell71|(X6=cell70|(X6=cell69|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X6=cell70|(X6=cell69|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5))&((X6=cell73|(X6=cell72|(X6=cell70|(X6=cell69|(X6=cell68|epred12_2(X6,X5))))))|~epred13_2(X6,X5)))))))&((((((~epred12_2(X8,X7)|epred13_2(X8,X7))&((X7!=cell67|X8!=cell68)|epred13_2(X8,X7)))&((X7!=cell68|X8!=cell69)|epred13_2(X8,X7)))&((X7!=cell69|X8!=cell70)|epred13_2(X8,X7)))&((X7!=cell71|X8!=cell72)|epred13_2(X8,X7)))&((X7!=cell72|X8!=cell73)|epred13_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1068])])])])])).
cnf(c_0_1074,plain,(epred13_2(X1,X2)|X1=cell74|X1=cell75|X1=cell76|X1=cell77|X1=cell78|~epred14_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1064])).
cnf(c_0_1075,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred14_2(cell2,esk2_0)|epred13_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1069, c_0_1070]), c_0_749]), c_0_750]), c_0_751]), c_0_724]), c_0_725])).
cnf(c_0_1076,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred14_2(cell99,esk1_0)|epred14_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1065, c_0_1071]), c_0_791]), c_0_743]), c_0_716]), c_0_717]), c_0_718])).
fof(c_0_1077, plain, (![X4]:![X3]:(epred12_2(X3,X4)<=>(((((epred11_2(X3,X4)|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67)))), inference(apply_def,[status(thm)],[c_0_1063, c_0_1072])).
cnf(c_0_1078,plain,(epred12_2(X1,X2)|X2=cell67|X2=cell68|X2=cell69|X2=cell71|X2=cell72|~epred13_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1073])).
cnf(c_0_1079,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred13_2(cell2,esk2_0)|epred13_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1074, c_0_1075]), c_0_750]), c_0_751]), c_0_724]), c_0_725]), c_0_726])).
cnf(c_0_1080,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred14_2(cell99,esk1_0)|epred13_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1069, c_0_1076]), c_0_760]), c_0_761]), c_0_762]), c_0_740]), c_0_741])).
fof(c_0_1081, plain, (![X4]:![X3]:(epred10_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56)))), introduced(definition)).
fof(c_0_1082, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell66|(X5=cell65|(X5=cell64|(X5=cell63|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X5=cell65|(X5=cell64|(X5=cell63|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X5=cell64|(X5=cell63|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X6=cell66|(X5=cell64|(X5=cell63|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))))&((((X5=cell66|(X5=cell65|(X6=cell65|(X5=cell63|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X5=cell65|(X6=cell65|(X5=cell63|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X6=cell65|(X5=cell63|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X6=cell66|(X6=cell65|(X5=cell63|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))))&(((((X5=cell66|(X5=cell65|(X5=cell64|(X6=cell64|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X5=cell65|(X5=cell64|(X6=cell64|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X5=cell64|(X6=cell64|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X6=cell66|(X5=cell64|(X6=cell64|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))))&((((X5=cell66|(X5=cell65|(X6=cell65|(X6=cell64|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X5=cell65|(X6=cell65|(X6=cell64|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X6=cell65|(X6=cell64|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X6=cell66|(X6=cell65|(X6=cell64|(X5=cell62|epred11_2(X6,X5))))))|~epred12_2(X6,X5))))))&((((((X5=cell66|(X5=cell65|(X5=cell64|(X5=cell63|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X5=cell65|(X5=cell64|(X5=cell63|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X5=cell64|(X5=cell63|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X6=cell66|(X5=cell64|(X5=cell63|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))))&((((X5=cell66|(X5=cell65|(X6=cell65|(X5=cell63|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X5=cell65|(X6=cell65|(X5=cell63|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X6=cell65|(X5=cell63|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X6=cell66|(X6=cell65|(X5=cell63|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))))&(((((X5=cell66|(X5=cell65|(X5=cell64|(X6=cell64|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X5=cell65|(X5=cell64|(X6=cell64|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X5=cell64|(X6=cell64|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X6=cell66|(X5=cell64|(X6=cell64|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))))&((((X5=cell66|(X5=cell65|(X6=cell65|(X6=cell64|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X5=cell65|(X6=cell65|(X6=cell64|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X6=cell65|(X6=cell64|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5))&((X6=cell67|(X6=cell66|(X6=cell65|(X6=cell64|(X6=cell63|epred11_2(X6,X5))))))|~epred12_2(X6,X5)))))))&((((((~epred11_2(X8,X7)|epred12_2(X8,X7))&((X7!=cell62|X8!=cell63)|epred12_2(X8,X7)))&((X7!=cell63|X8!=cell64)|epred12_2(X8,X7)))&((X7!=cell64|X8!=cell65)|epred12_2(X8,X7)))&((X7!=cell65|X8!=cell66)|epred12_2(X8,X7)))&((X7!=cell66|X8!=cell67)|epred12_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1077])])])])])).
cnf(c_0_1083,plain,(epred12_2(X1,X2)|X1=cell68|X1=cell69|X1=cell70|X1=cell72|X1=cell73|~epred13_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1073])).
cnf(c_0_1084,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred13_2(cell2,esk2_0)|epred12_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1078, c_0_1079]), c_0_767]), c_0_768]), c_0_769]), c_0_747]), c_0_748])).
cnf(c_0_1085,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred13_2(cell99,esk1_0)|epred13_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1074, c_0_1080]), c_0_761]), c_0_762]), c_0_740]), c_0_741]), c_0_742])).
fof(c_0_1086, plain, (![X4]:![X3]:(epred11_2(X3,X4)<=>(((((epred10_2(X3,X4)|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62)))), inference(apply_def,[status(thm)],[c_0_1072, c_0_1081])).
cnf(c_0_1087,plain,(epred11_2(X1,X2)|X2=cell62|X2=cell63|X2=cell64|X2=cell65|X2=cell66|~epred12_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1082])).
cnf(c_0_1088,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred12_2(cell2,esk2_0)|epred12_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1083, c_0_1084]), c_0_768]), c_0_769]), c_0_770]), c_0_748]), c_0_749])).
cnf(c_0_1089,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred13_2(cell99,esk1_0)|epred12_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1078, c_0_1085]), c_0_778]), c_0_779]), c_0_780]), c_0_758]), c_0_759])).
fof(c_0_1090, plain, (![X4]:![X3]:(epred9_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50)))), introduced(definition)).
fof(c_0_1091, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell61|(X5=cell59|(X5=cell58|(X5=cell57|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X5=cell59|(X5=cell58|(X5=cell57|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X5=cell58|(X5=cell57|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X6=cell60|(X5=cell58|(X5=cell57|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))))&((((X5=cell61|(X5=cell59|(X6=cell59|(X5=cell57|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X5=cell59|(X6=cell59|(X5=cell57|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X6=cell59|(X5=cell57|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X6=cell60|(X6=cell59|(X5=cell57|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))))&(((((X5=cell61|(X5=cell59|(X5=cell58|(X6=cell58|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X5=cell59|(X5=cell58|(X6=cell58|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X5=cell58|(X6=cell58|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X6=cell60|(X5=cell58|(X6=cell58|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))))&((((X5=cell61|(X5=cell59|(X6=cell59|(X6=cell58|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X5=cell59|(X6=cell59|(X6=cell58|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X6=cell59|(X6=cell58|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X6=cell60|(X6=cell59|(X6=cell58|(X5=cell56|epred10_2(X6,X5))))))|~epred11_2(X6,X5))))))&((((((X5=cell61|(X5=cell59|(X5=cell58|(X5=cell57|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X5=cell59|(X5=cell58|(X5=cell57|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X5=cell58|(X5=cell57|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X6=cell60|(X5=cell58|(X5=cell57|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))))&((((X5=cell61|(X5=cell59|(X6=cell59|(X5=cell57|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X5=cell59|(X6=cell59|(X5=cell57|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X6=cell59|(X5=cell57|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X6=cell60|(X6=cell59|(X5=cell57|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))))&(((((X5=cell61|(X5=cell59|(X5=cell58|(X6=cell58|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X5=cell59|(X5=cell58|(X6=cell58|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X5=cell58|(X6=cell58|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X6=cell60|(X5=cell58|(X6=cell58|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))))&((((X5=cell61|(X5=cell59|(X6=cell59|(X6=cell58|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X5=cell59|(X6=cell59|(X6=cell58|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X6=cell59|(X6=cell58|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5))&((X6=cell62|(X6=cell60|(X6=cell59|(X6=cell58|(X6=cell57|epred10_2(X6,X5))))))|~epred11_2(X6,X5)))))))&((((((~epred10_2(X8,X7)|epred11_2(X8,X7))&((X7!=cell56|X8!=cell57)|epred11_2(X8,X7)))&((X7!=cell57|X8!=cell58)|epred11_2(X8,X7)))&((X7!=cell58|X8!=cell59)|epred11_2(X8,X7)))&((X7!=cell59|X8!=cell60)|epred11_2(X8,X7)))&((X7!=cell61|X8!=cell62)|epred11_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1086])])])])])).
cnf(c_0_1092,plain,(epred11_2(X1,X2)|X1=cell63|X1=cell64|X1=cell65|X1=cell66|X1=cell67|~epred12_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1082])).
cnf(c_0_1093,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred12_2(cell2,esk2_0)|epred11_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1087, c_0_1088]), c_0_786]), c_0_787]), c_0_788]), c_0_789]), c_0_766])).
cnf(c_0_1094,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred12_2(cell99,esk1_0)|epred12_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1083, c_0_1089]), c_0_779]), c_0_780]), c_0_781]), c_0_759]), c_0_760])).
fof(c_0_1095, plain, (![X4]:![X3]:(epred10_2(X3,X4)<=>(((((epred9_2(X3,X4)|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56)))), inference(apply_def,[status(thm)],[c_0_1081, c_0_1090])).
cnf(c_0_1096,plain,(epred10_2(X1,X2)|X2=cell56|X2=cell57|X2=cell58|X2=cell59|X2=cell61|~epred11_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1091])).
cnf(c_0_1097,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred11_2(cell2,esk2_0)|epred11_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1092, c_0_1093]), c_0_787]), c_0_788]), c_0_789]), c_0_766]), c_0_767])).
cnf(c_0_1098,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred12_2(cell99,esk1_0)|epred11_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1087, c_0_1094]), c_0_798]), c_0_799]), c_0_800]), c_0_801]), c_0_777])).
fof(c_0_1099, plain, (![X4]:![X3]:(epred8_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45)))), introduced(definition)).
fof(c_0_1100, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell55|(X5=cell54|(X5=cell53|(X5=cell52|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X5=cell54|(X5=cell53|(X5=cell52|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X5=cell53|(X5=cell52|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X6=cell55|(X5=cell53|(X5=cell52|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell54|(X5=cell52|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X5=cell54|(X6=cell54|(X5=cell52|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X6=cell54|(X5=cell52|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X6=cell55|(X6=cell54|(X5=cell52|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))))&(((((X5=cell55|(X5=cell54|(X5=cell53|(X6=cell53|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X5=cell54|(X5=cell53|(X6=cell53|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X5=cell53|(X6=cell53|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X6=cell55|(X5=cell53|(X6=cell53|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell54|(X6=cell53|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X5=cell54|(X6=cell54|(X6=cell53|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X6=cell54|(X6=cell53|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X6=cell55|(X6=cell54|(X6=cell53|(X5=cell51|epred9_2(X6,X5))))))|~epred10_2(X6,X5))))))&((((((X5=cell55|(X5=cell54|(X5=cell53|(X5=cell52|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X5=cell54|(X5=cell53|(X5=cell52|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X5=cell53|(X5=cell52|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X6=cell55|(X5=cell53|(X5=cell52|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell54|(X5=cell52|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X5=cell54|(X6=cell54|(X5=cell52|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X6=cell54|(X5=cell52|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X6=cell55|(X6=cell54|(X5=cell52|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))))&(((((X5=cell55|(X5=cell54|(X5=cell53|(X6=cell53|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X5=cell54|(X5=cell53|(X6=cell53|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X5=cell53|(X6=cell53|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X6=cell55|(X5=cell53|(X6=cell53|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell54|(X6=cell53|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X5=cell54|(X6=cell54|(X6=cell53|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X6=cell54|(X6=cell53|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5))&((X6=cell56|(X6=cell55|(X6=cell54|(X6=cell53|(X6=cell52|epred9_2(X6,X5))))))|~epred10_2(X6,X5)))))))&((((((~epred9_2(X8,X7)|epred10_2(X8,X7))&((X7!=cell51|X8!=cell52)|epred10_2(X8,X7)))&((X7!=cell52|X8!=cell53)|epred10_2(X8,X7)))&((X7!=cell53|X8!=cell54)|epred10_2(X8,X7)))&((X7!=cell54|X8!=cell55)|epred10_2(X8,X7)))&((X7!=cell55|X8!=cell56)|epred10_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1095])])])])])).
cnf(c_0_1101,plain,(epred10_2(X1,X2)|X1=cell57|X1=cell58|X1=cell59|X1=cell60|X1=cell62|~epred11_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1091])).
cnf(c_0_1102,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred11_2(cell2,esk2_0)|epred10_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1096, c_0_1097]), c_0_805]), c_0_806]), c_0_807]), c_0_808]), c_0_785])).
cnf(c_0_1103,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred11_2(cell99,esk1_0)|epred11_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1092, c_0_1098]), c_0_799]), c_0_800]), c_0_801]), c_0_777]), c_0_778])).
fof(c_0_1104, plain, (![X4]:![X3]:(epred9_2(X3,X4)<=>(((((epred8_2(X3,X4)|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50)))), inference(apply_def,[status(thm)],[c_0_1090, c_0_1099])).
cnf(c_0_1105,plain,(epred9_2(X1,X2)|X2=cell51|X2=cell52|X2=cell53|X2=cell54|X2=cell55|~epred10_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1100])).
cnf(c_0_1106,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred10_2(cell2,esk2_0)|epred10_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1101, c_0_1102]), c_0_806]), c_0_807]), c_0_808]), c_0_809]), c_0_786])).
cnf(c_0_1107,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred11_2(cell99,esk1_0)|epred10_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1096, c_0_1103]), c_0_816]), c_0_817]), c_0_818]), c_0_819]), c_0_797])).
fof(c_0_1108, plain, (![X4]:![X3]:(epred7_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39)))), introduced(definition)).
fof(c_0_1109, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell49|(X5=cell48|(X5=cell47|(X5=cell46|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X5=cell48|(X5=cell47|(X5=cell46|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X5=cell47|(X5=cell46|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell47|(X5=cell46|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))))&((((X5=cell49|(X5=cell48|(X6=cell48|(X5=cell46|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X5=cell48|(X6=cell48|(X5=cell46|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X6=cell48|(X5=cell46|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X5=cell46|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))))&(((((X5=cell49|(X5=cell48|(X5=cell47|(X6=cell47|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X5=cell48|(X5=cell47|(X6=cell47|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X5=cell47|(X6=cell47|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell47|(X6=cell47|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))))&((((X5=cell49|(X5=cell48|(X6=cell48|(X6=cell47|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X5=cell48|(X6=cell48|(X6=cell47|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X6=cell48|(X6=cell47|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X6=cell47|(X5=cell45|epred8_2(X6,X5))))))|~epred9_2(X6,X5))))))&((((((X5=cell49|(X5=cell48|(X5=cell47|(X5=cell46|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X5=cell48|(X5=cell47|(X5=cell46|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X5=cell47|(X5=cell46|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell47|(X5=cell46|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))))&((((X5=cell49|(X5=cell48|(X6=cell48|(X5=cell46|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X5=cell48|(X6=cell48|(X5=cell46|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X6=cell48|(X5=cell46|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X5=cell46|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))))&(((((X5=cell49|(X5=cell48|(X5=cell47|(X6=cell47|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X5=cell48|(X5=cell47|(X6=cell47|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X5=cell47|(X6=cell47|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell47|(X6=cell47|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))))&((((X5=cell49|(X5=cell48|(X6=cell48|(X6=cell47|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X5=cell48|(X6=cell48|(X6=cell47|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X6=cell48|(X6=cell47|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X6=cell47|(X6=cell46|epred8_2(X6,X5))))))|~epred9_2(X6,X5)))))))&((((((~epred8_2(X8,X7)|epred9_2(X8,X7))&((X7!=cell45|X8!=cell46)|epred9_2(X8,X7)))&((X7!=cell46|X8!=cell47)|epred9_2(X8,X7)))&((X7!=cell47|X8!=cell48)|epred9_2(X8,X7)))&((X7!=cell48|X8!=cell49)|epred9_2(X8,X7)))&((X7!=cell49|X8!=cell50)|epred9_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1104])])])])])).
cnf(c_0_1110,plain,(epred9_2(X1,X2)|X1=cell52|X1=cell53|X1=cell54|X1=cell55|X1=cell56|~epred10_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1100])).
cnf(c_0_1111,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred10_2(cell2,esk2_0)|epred9_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1105, c_0_1106]), c_0_824]), c_0_825]), c_0_826]), c_0_827]), c_0_828])).
cnf(c_0_1112,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred10_2(cell99,esk1_0)|epred10_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1101, c_0_1107]), c_0_817]), c_0_818]), c_0_819]), c_0_820]), c_0_798])).
fof(c_0_1113, plain, (![X4]:![X3]:(epred8_2(X3,X4)<=>(((((epred7_2(X3,X4)|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45)))), inference(apply_def,[status(thm)],[c_0_1099, c_0_1108])).
cnf(c_0_1114,plain,(epred8_2(X1,X2)|X2=cell45|X2=cell46|X2=cell47|X2=cell48|X2=cell49|~epred9_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1109])).
cnf(c_0_1115,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred9_2(cell2,esk2_0)|epred9_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1110, c_0_1111]), c_0_825]), c_0_826]), c_0_827]), c_0_828]), c_0_805])).
cnf(c_0_1116,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred10_2(cell99,esk1_0)|epred9_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1105, c_0_1112]), c_0_835]), c_0_836]), c_0_837]), c_0_838]), c_0_839])).
fof(c_0_1117, plain, (![X4]:![X3]:(epred6_2(X3,X4)<=>((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34)))), introduced(definition)).
fof(c_0_1118, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell44|(X5=cell43|(X5=cell42|(X5=cell41|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X5=cell43|(X5=cell42|(X5=cell41|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X5=cell42|(X5=cell41|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell42|(X5=cell41|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))))&((((X5=cell44|(X5=cell43|(X6=cell43|(X5=cell41|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X5=cell43|(X6=cell43|(X5=cell41|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X6=cell43|(X5=cell41|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X5=cell41|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))))&(((((X5=cell44|(X5=cell43|(X5=cell42|(X6=cell42|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X5=cell43|(X5=cell42|(X6=cell42|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X5=cell42|(X6=cell42|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell42|(X6=cell42|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))))&((((X5=cell44|(X5=cell43|(X6=cell43|(X6=cell42|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X5=cell43|(X6=cell43|(X6=cell42|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X6=cell43|(X6=cell42|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X6=cell42|(X5=cell39|epred7_2(X6,X5))))))|~epred8_2(X6,X5))))))&((((((X5=cell44|(X5=cell43|(X5=cell42|(X5=cell41|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X5=cell43|(X5=cell42|(X5=cell41|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X5=cell42|(X5=cell41|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell42|(X5=cell41|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))))&((((X5=cell44|(X5=cell43|(X6=cell43|(X5=cell41|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X5=cell43|(X6=cell43|(X5=cell41|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X6=cell43|(X5=cell41|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X5=cell41|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))))&(((((X5=cell44|(X5=cell43|(X5=cell42|(X6=cell42|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X5=cell43|(X5=cell42|(X6=cell42|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X5=cell42|(X6=cell42|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell42|(X6=cell42|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))))&((((X5=cell44|(X5=cell43|(X6=cell43|(X6=cell42|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X5=cell43|(X6=cell43|(X6=cell42|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X6=cell43|(X6=cell42|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X6=cell42|(X6=cell40|epred7_2(X6,X5))))))|~epred8_2(X6,X5)))))))&((((((~epred7_2(X8,X7)|epred8_2(X8,X7))&((X7!=cell39|X8!=cell40)|epred8_2(X8,X7)))&((X7!=cell41|X8!=cell42)|epred8_2(X8,X7)))&((X7!=cell42|X8!=cell43)|epred8_2(X8,X7)))&((X7!=cell43|X8!=cell44)|epred8_2(X8,X7)))&((X7!=cell44|X8!=cell45)|epred8_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1113])])])])])).
cnf(c_0_1119,plain,(epred8_2(X1,X2)|X1=cell46|X1=cell47|X1=cell48|X1=cell49|X1=cell50|~epred9_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1109])).
cnf(c_0_1120,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred9_2(cell2,esk2_0)|epred8_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1114, c_0_1115]), c_0_866]), c_0_843]), c_0_844]), c_0_845]), c_0_846])).
cnf(c_0_1121,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred9_2(cell99,esk1_0)|epred9_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1110, c_0_1116]), c_0_836]), c_0_837]), c_0_838]), c_0_839]), c_0_816])).
fof(c_0_1122, plain, (![X4]:![X3]:(epred7_2(X3,X4)<=>(((((epred6_2(X3,X4)|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39)))), inference(apply_def,[status(thm)],[c_0_1108, c_0_1117])).
cnf(c_0_1123,plain,(epred7_2(X1,X2)|X2=cell39|X2=cell41|X2=cell42|X2=cell43|X2=cell44|~epred8_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1118])).
cnf(c_0_1124,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred8_2(cell2,esk2_0)|epred8_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1119, c_0_1120]), c_0_843]), c_0_844]), c_0_845]), c_0_846]), c_0_847])).
cnf(c_0_1125,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred9_2(cell99,esk1_0)|epred8_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1114, c_0_1121]), c_0_877]), c_0_854]), c_0_855]), c_0_856]), c_0_857])).
fof(c_0_1126, plain, (![X4]:![X3]:(epred5_2(X3,X4)<=>(((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28)))), introduced(definition)).
fof(c_0_1127, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell38|(X5=cell37|(X5=cell36|(X5=cell35|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X5=cell37|(X5=cell36|(X5=cell35|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X5=cell36|(X5=cell35|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X6=cell38|(X5=cell36|(X5=cell35|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))))&((((X5=cell38|(X5=cell37|(X6=cell37|(X5=cell35|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X5=cell37|(X6=cell37|(X5=cell35|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X6=cell37|(X5=cell35|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X6=cell38|(X6=cell37|(X5=cell35|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))))&(((((X5=cell38|(X5=cell37|(X5=cell36|(X6=cell36|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X5=cell37|(X5=cell36|(X6=cell36|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X5=cell36|(X6=cell36|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X6=cell38|(X5=cell36|(X6=cell36|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))))&((((X5=cell38|(X5=cell37|(X6=cell37|(X6=cell36|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X5=cell37|(X6=cell37|(X6=cell36|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X6=cell37|(X6=cell36|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X6=cell38|(X6=cell37|(X6=cell36|(X5=cell34|epred6_2(X6,X5))))))|~epred7_2(X6,X5))))))&((((((X5=cell38|(X5=cell37|(X5=cell36|(X5=cell35|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X5=cell37|(X5=cell36|(X5=cell35|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X5=cell36|(X5=cell35|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X6=cell38|(X5=cell36|(X5=cell35|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))))&((((X5=cell38|(X5=cell37|(X6=cell37|(X5=cell35|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X5=cell37|(X6=cell37|(X5=cell35|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X6=cell37|(X5=cell35|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X6=cell38|(X6=cell37|(X5=cell35|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))))&(((((X5=cell38|(X5=cell37|(X5=cell36|(X6=cell36|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X5=cell37|(X5=cell36|(X6=cell36|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X5=cell36|(X6=cell36|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X6=cell38|(X5=cell36|(X6=cell36|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))))&((((X5=cell38|(X5=cell37|(X6=cell37|(X6=cell36|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X5=cell37|(X6=cell37|(X6=cell36|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X6=cell37|(X6=cell36|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5))&((X6=cell39|(X6=cell38|(X6=cell37|(X6=cell36|(X6=cell35|epred6_2(X6,X5))))))|~epred7_2(X6,X5)))))))&((((((~epred6_2(X8,X7)|epred7_2(X8,X7))&((X7!=cell34|X8!=cell35)|epred7_2(X8,X7)))&((X7!=cell35|X8!=cell36)|epred7_2(X8,X7)))&((X7!=cell36|X8!=cell37)|epred7_2(X8,X7)))&((X7!=cell37|X8!=cell38)|epred7_2(X8,X7)))&((X7!=cell38|X8!=cell39)|epred7_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1122])])])])])).
cnf(c_0_1128,plain,(epred7_2(X1,X2)|X1=cell40|X1=cell42|X1=cell43|X1=cell44|X1=cell45|~epred8_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1118])).
cnf(c_0_1129,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred8_2(cell2,esk2_0)|epred7_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1123, c_0_1124]), c_0_884]), c_0_862]), c_0_863]), c_0_864]), c_0_865])).
cnf(c_0_1130,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred8_2(cell99,esk1_0)|epred8_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1119, c_0_1125]), c_0_854]), c_0_855]), c_0_856]), c_0_857]), c_0_858])).
fof(c_0_1131, plain, (![X4]:![X3]:(epred6_2(X3,X4)<=>(((((epred5_2(X3,X4)|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34)))), inference(apply_def,[status(thm)],[c_0_1117, c_0_1126])).
cnf(c_0_1132,plain,(epred6_2(X1,X2)|X2=cell34|X2=cell35|X2=cell36|X2=cell37|X2=cell38|~epred7_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1127])).
cnf(c_0_1133,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred7_2(cell2,esk2_0)|epred7_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1128, c_0_1129]), c_0_885]), c_0_863]), c_0_864]), c_0_865]), c_0_866])).
cnf(c_0_1134,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred8_2(cell99,esk1_0)|epred7_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1123, c_0_1130]), c_0_895]), c_0_873]), c_0_874]), c_0_875]), c_0_876])).
fof(c_0_1135, plain, (![X4]:![X3]:(epred4_2(X3,X4)<=>((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23)))), introduced(definition)).
fof(c_0_1136, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell33|(X5=cell32|(X5=cell31|(X5=cell29|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X5=cell32|(X5=cell31|(X5=cell29|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X5=cell31|(X5=cell29|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X6=cell33|(X5=cell31|(X5=cell29|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))))&((((X5=cell33|(X5=cell32|(X6=cell32|(X5=cell29|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X5=cell32|(X6=cell32|(X5=cell29|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X6=cell32|(X5=cell29|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X6=cell33|(X6=cell32|(X5=cell29|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))))&(((((X5=cell33|(X5=cell32|(X5=cell31|(X6=cell30|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X5=cell32|(X5=cell31|(X6=cell30|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X5=cell31|(X6=cell30|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X6=cell33|(X5=cell31|(X6=cell30|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))))&((((X5=cell33|(X5=cell32|(X6=cell32|(X6=cell30|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X5=cell32|(X6=cell32|(X6=cell30|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X6=cell32|(X6=cell30|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X6=cell33|(X6=cell32|(X6=cell30|(X5=cell28|epred5_2(X6,X5))))))|~epred6_2(X6,X5))))))&((((((X5=cell33|(X5=cell32|(X5=cell31|(X5=cell29|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X5=cell32|(X5=cell31|(X5=cell29|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X5=cell31|(X5=cell29|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X6=cell33|(X5=cell31|(X5=cell29|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))))&((((X5=cell33|(X5=cell32|(X6=cell32|(X5=cell29|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X5=cell32|(X6=cell32|(X5=cell29|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X6=cell32|(X5=cell29|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X6=cell33|(X6=cell32|(X5=cell29|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))))&(((((X5=cell33|(X5=cell32|(X5=cell31|(X6=cell30|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X5=cell32|(X5=cell31|(X6=cell30|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X5=cell31|(X6=cell30|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X6=cell33|(X5=cell31|(X6=cell30|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))))&((((X5=cell33|(X5=cell32|(X6=cell32|(X6=cell30|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X5=cell32|(X6=cell32|(X6=cell30|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X6=cell32|(X6=cell30|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5))&((X6=cell34|(X6=cell33|(X6=cell32|(X6=cell30|(X6=cell29|epred5_2(X6,X5))))))|~epred6_2(X6,X5)))))))&((((((~epred5_2(X8,X7)|epred6_2(X8,X7))&((X7!=cell28|X8!=cell29)|epred6_2(X8,X7)))&((X7!=cell29|X8!=cell30)|epred6_2(X8,X7)))&((X7!=cell31|X8!=cell32)|epred6_2(X8,X7)))&((X7!=cell32|X8!=cell33)|epred6_2(X8,X7)))&((X7!=cell33|X8!=cell34)|epred6_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1131])])])])])).
cnf(c_0_1137,plain,(epred6_2(X1,X2)|X1=cell35|X1=cell36|X1=cell37|X1=cell38|X1=cell39|~epred7_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1127])).
cnf(c_0_1138,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred7_2(cell2,esk2_0)|epred6_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1132, c_0_1133]), c_0_903]), c_0_904]), c_0_881]), c_0_882]), c_0_883])).
cnf(c_0_1139,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred7_2(cell99,esk1_0)|epred7_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1128, c_0_1134]), c_0_896]), c_0_874]), c_0_875]), c_0_876]), c_0_877])).
fof(c_0_1140, plain, (![X4]:![X3]:(epred5_2(X3,X4)<=>(((((epred4_2(X3,X4)|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28)))), inference(apply_def,[status(thm)],[c_0_1126, c_0_1135])).
cnf(c_0_1141,plain,(epred5_2(X1,X2)|X2=cell28|X2=cell29|X2=cell31|X2=cell32|X2=cell33|~epred6_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1136])).
cnf(c_0_1142,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred6_2(cell2,esk2_0)|epred6_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1137, c_0_1138]), c_0_904]), c_0_881]), c_0_882]), c_0_883]), c_0_884])).
cnf(c_0_1143,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred7_2(cell99,esk1_0)|epred6_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1132, c_0_1139]), c_0_914]), c_0_915]), c_0_892]), c_0_893]), c_0_894])).
fof(c_0_1144, plain, (![X4]:![X3]:(epred3_2(X3,X4)<=>(((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17)))), introduced(definition)).
fof(c_0_1145, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell27|(X5=cell26|(X5=cell25|(X5=cell24|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X5=cell26|(X5=cell25|(X5=cell24|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X5=cell25|(X5=cell24|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X6=cell27|(X5=cell25|(X5=cell24|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))))&((((X5=cell27|(X5=cell26|(X6=cell26|(X5=cell24|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X5=cell26|(X6=cell26|(X5=cell24|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X6=cell26|(X5=cell24|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X6=cell27|(X6=cell26|(X5=cell24|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))))&(((((X5=cell27|(X5=cell26|(X5=cell25|(X6=cell25|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X5=cell26|(X5=cell25|(X6=cell25|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X5=cell25|(X6=cell25|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X6=cell27|(X5=cell25|(X6=cell25|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))))&((((X5=cell27|(X5=cell26|(X6=cell26|(X6=cell25|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X5=cell26|(X6=cell26|(X6=cell25|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X6=cell26|(X6=cell25|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X6=cell27|(X6=cell26|(X6=cell25|(X5=cell23|epred4_2(X6,X5))))))|~epred5_2(X6,X5))))))&((((((X5=cell27|(X5=cell26|(X5=cell25|(X5=cell24|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X5=cell26|(X5=cell25|(X5=cell24|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X5=cell25|(X5=cell24|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X6=cell27|(X5=cell25|(X5=cell24|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))))&((((X5=cell27|(X5=cell26|(X6=cell26|(X5=cell24|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X5=cell26|(X6=cell26|(X5=cell24|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X6=cell26|(X5=cell24|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X6=cell27|(X6=cell26|(X5=cell24|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))))&(((((X5=cell27|(X5=cell26|(X5=cell25|(X6=cell25|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X5=cell26|(X5=cell25|(X6=cell25|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X5=cell25|(X6=cell25|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X6=cell27|(X5=cell25|(X6=cell25|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))))&((((X5=cell27|(X5=cell26|(X6=cell26|(X6=cell25|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X5=cell26|(X6=cell26|(X6=cell25|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X6=cell26|(X6=cell25|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5))&((X6=cell28|(X6=cell27|(X6=cell26|(X6=cell25|(X6=cell24|epred4_2(X6,X5))))))|~epred5_2(X6,X5)))))))&((((((~epred4_2(X8,X7)|epred5_2(X8,X7))&((X7!=cell23|X8!=cell24)|epred5_2(X8,X7)))&((X7!=cell24|X8!=cell25)|epred5_2(X8,X7)))&((X7!=cell25|X8!=cell26)|epred5_2(X8,X7)))&((X7!=cell26|X8!=cell27)|epred5_2(X8,X7)))&((X7!=cell27|X8!=cell28)|epred5_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1140])])])])])).
cnf(c_0_1146,plain,(epred5_2(X1,X2)|X1=cell29|X1=cell30|X1=cell32|X1=cell33|X1=cell34|~epred6_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1136])).
cnf(c_0_1147,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred6_2(cell2,esk2_0)|epred5_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1141, c_0_1142]), c_0_921]), c_0_922]), c_0_900]), c_0_901]), c_0_902])).
cnf(c_0_1148,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred6_2(cell99,esk1_0)|epred6_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1137, c_0_1143]), c_0_915]), c_0_892]), c_0_893]), c_0_894]), c_0_895])).
fof(c_0_1149, plain, (![X4]:![X3]:(epred4_2(X3,X4)<=>(((((epred3_2(X3,X4)|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23)))), inference(apply_def,[status(thm)],[c_0_1135, c_0_1144])).
cnf(c_0_1150,plain,(epred4_2(X1,X2)|X2=cell23|X2=cell24|X2=cell25|X2=cell26|X2=cell27|~epred5_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1145])).
cnf(c_0_1151,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred5_2(cell2,esk2_0)|epred5_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1146, c_0_1147]), c_0_922]), c_0_923]), c_0_901]), c_0_902]), c_0_903])).
cnf(c_0_1152,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred6_2(cell99,esk1_0)|epred5_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1141, c_0_1148]), c_0_932]), c_0_933]), c_0_911]), c_0_912]), c_0_913])).
fof(c_0_1153, plain, (![X4]:![X3]:(epred2_2(X3,X4)<=>((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12)))), introduced(definition)).
fof(c_0_1154, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell22|(X5=cell21|(X5=cell19|(X5=cell18|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X5=cell21|(X5=cell19|(X5=cell18|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X5=cell19|(X5=cell18|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X6=cell22|(X5=cell19|(X5=cell18|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))))&((((X5=cell22|(X5=cell21|(X6=cell20|(X5=cell18|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X5=cell21|(X6=cell20|(X5=cell18|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X6=cell20|(X5=cell18|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X6=cell22|(X6=cell20|(X5=cell18|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))))&(((((X5=cell22|(X5=cell21|(X5=cell19|(X6=cell19|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X5=cell21|(X5=cell19|(X6=cell19|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X5=cell19|(X6=cell19|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X6=cell22|(X5=cell19|(X6=cell19|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))))&((((X5=cell22|(X5=cell21|(X6=cell20|(X6=cell19|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X5=cell21|(X6=cell20|(X6=cell19|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X6=cell20|(X6=cell19|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X6=cell22|(X6=cell20|(X6=cell19|(X5=cell17|epred3_2(X6,X5))))))|~epred4_2(X6,X5))))))&((((((X5=cell22|(X5=cell21|(X5=cell19|(X5=cell18|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X5=cell21|(X5=cell19|(X5=cell18|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X5=cell19|(X5=cell18|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X6=cell22|(X5=cell19|(X5=cell18|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))))&((((X5=cell22|(X5=cell21|(X6=cell20|(X5=cell18|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X5=cell21|(X6=cell20|(X5=cell18|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X6=cell20|(X5=cell18|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X6=cell22|(X6=cell20|(X5=cell18|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))))&(((((X5=cell22|(X5=cell21|(X5=cell19|(X6=cell19|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X5=cell21|(X5=cell19|(X6=cell19|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X5=cell19|(X6=cell19|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X6=cell22|(X5=cell19|(X6=cell19|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))))&((((X5=cell22|(X5=cell21|(X6=cell20|(X6=cell19|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X5=cell21|(X6=cell20|(X6=cell19|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X6=cell20|(X6=cell19|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5))&((X6=cell23|(X6=cell22|(X6=cell20|(X6=cell19|(X6=cell18|epred3_2(X6,X5))))))|~epred4_2(X6,X5)))))))&((((((~epred3_2(X8,X7)|epred4_2(X8,X7))&((X7!=cell17|X8!=cell18)|epred4_2(X8,X7)))&((X7!=cell18|X8!=cell19)|epred4_2(X8,X7)))&((X7!=cell19|X8!=cell20)|epred4_2(X8,X7)))&((X7!=cell21|X8!=cell22)|epred4_2(X8,X7)))&((X7!=cell22|X8!=cell23)|epred4_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1149])])])])])).
cnf(c_0_1155,plain,(epred4_2(X1,X2)|X1=cell24|X1=cell25|X1=cell26|X1=cell27|X1=cell28|~epred5_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1145])).
cnf(c_0_1156,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred5_2(cell2,esk2_0)|epred4_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1150, c_0_1151]), c_0_940]), c_0_941]), c_0_942]), c_0_919]), c_0_920])).
cnf(c_0_1157,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred5_2(cell99,esk1_0)|epred5_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1146, c_0_1152]), c_0_933]), c_0_934]), c_0_912]), c_0_913]), c_0_914])).
fof(c_0_1158, plain, (![X4]:![X3]:(epred1_2(X3,X4)<=>(((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6)))), introduced(definition)).
fof(c_0_1159, plain, (![X4]:![X3]:(epred3_2(X3,X4)<=>(((((epred2_2(X3,X4)|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17)))), inference(apply_def,[status(thm)],[c_0_1144, c_0_1153])).
cnf(c_0_1160,plain,(epred3_2(X1,X2)|X2=cell17|X2=cell18|X2=cell19|X2=cell21|X2=cell22|~epred4_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1154])).
cnf(c_0_1161,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred4_2(cell2,esk2_0)|epred4_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1155, c_0_1156]), c_0_941]), c_0_942]), c_0_919]), c_0_920]), c_0_921])).
cnf(c_0_1162,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred5_2(cell99,esk1_0)|epred4_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1150, c_0_1157]), c_0_951]), c_0_952]), c_0_953]), c_0_930]), c_0_931])).
fof(c_0_1163, plain, (![X4]:![X3]:(epred2_2(X3,X4)<=>(((((epred1_2(X3,X4)|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12)))), inference(apply_def,[status(thm)],[c_0_1153, c_0_1158])).
fof(c_0_1164, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell16|(X5=cell15|(X5=cell14|(X5=cell13|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X5=cell15|(X5=cell14|(X5=cell13|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X5=cell14|(X5=cell13|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X6=cell16|(X5=cell14|(X5=cell13|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))))&((((X5=cell16|(X5=cell15|(X6=cell15|(X5=cell13|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X5=cell15|(X6=cell15|(X5=cell13|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X6=cell15|(X5=cell13|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X6=cell16|(X6=cell15|(X5=cell13|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))))&(((((X5=cell16|(X5=cell15|(X5=cell14|(X6=cell14|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X5=cell15|(X5=cell14|(X6=cell14|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X5=cell14|(X6=cell14|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X6=cell16|(X5=cell14|(X6=cell14|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))))&((((X5=cell16|(X5=cell15|(X6=cell15|(X6=cell14|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X5=cell15|(X6=cell15|(X6=cell14|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X6=cell15|(X6=cell14|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X6=cell16|(X6=cell15|(X6=cell14|(X5=cell12|epred2_2(X6,X5))))))|~epred3_2(X6,X5))))))&((((((X5=cell16|(X5=cell15|(X5=cell14|(X5=cell13|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X5=cell15|(X5=cell14|(X5=cell13|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X5=cell14|(X5=cell13|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X6=cell16|(X5=cell14|(X5=cell13|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))))&((((X5=cell16|(X5=cell15|(X6=cell15|(X5=cell13|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X5=cell15|(X6=cell15|(X5=cell13|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X6=cell15|(X5=cell13|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X6=cell16|(X6=cell15|(X5=cell13|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))))&(((((X5=cell16|(X5=cell15|(X5=cell14|(X6=cell14|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X5=cell15|(X5=cell14|(X6=cell14|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X5=cell14|(X6=cell14|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X6=cell16|(X5=cell14|(X6=cell14|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))))&((((X5=cell16|(X5=cell15|(X6=cell15|(X6=cell14|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X5=cell15|(X6=cell15|(X6=cell14|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X6=cell15|(X6=cell14|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5))&((X6=cell17|(X6=cell16|(X6=cell15|(X6=cell14|(X6=cell13|epred2_2(X6,X5))))))|~epred3_2(X6,X5)))))))&((((((~epred2_2(X8,X7)|epred3_2(X8,X7))&((X7!=cell12|X8!=cell13)|epred3_2(X8,X7)))&((X7!=cell13|X8!=cell14)|epred3_2(X8,X7)))&((X7!=cell14|X8!=cell15)|epred3_2(X8,X7)))&((X7!=cell15|X8!=cell16)|epred3_2(X8,X7)))&((X7!=cell16|X8!=cell17)|epred3_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1159])])])])])).
cnf(c_0_1165,plain,(epred3_2(X1,X2)|X1=cell18|X1=cell19|X1=cell20|X1=cell22|X1=cell23|~epred4_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1154])).
cnf(c_0_1166,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred4_2(cell2,esk2_0)|epred3_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1160, c_0_1161]), c_0_938]), c_0_958]), c_0_959]), c_0_960]), c_0_939])).
cnf(c_0_1167,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred4_2(cell99,esk1_0)|epred4_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1155, c_0_1162]), c_0_952]), c_0_953]), c_0_930]), c_0_931]), c_0_932])).
fof(c_0_1168, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell5|(X5=cell4|(X5=cell3|(X5=cell2|X5=cell1))))|~epred1_2(X6,X5))&((X6=cell6|(X5=cell4|(X5=cell3|(X5=cell2|X5=cell1))))|~epred1_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X5=cell3|(X5=cell2|X5=cell1))))|~epred1_2(X6,X5))&((X6=cell6|(X6=cell5|(X5=cell3|(X5=cell2|X5=cell1))))|~epred1_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell4|(X5=cell2|X5=cell1))))|~epred1_2(X6,X5))&((X6=cell6|(X5=cell4|(X6=cell4|(X5=cell2|X5=cell1))))|~epred1_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X6=cell4|(X5=cell2|X5=cell1))))|~epred1_2(X6,X5))&((X6=cell6|(X6=cell5|(X6=cell4|(X5=cell2|X5=cell1))))|~epred1_2(X6,X5)))))&(((((X5=cell5|(X5=cell4|(X5=cell3|(X6=cell3|X5=cell1))))|~epred1_2(X6,X5))&((X6=cell6|(X5=cell4|(X5=cell3|(X6=cell3|X5=cell1))))|~epred1_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X5=cell3|(X6=cell3|X5=cell1))))|~epred1_2(X6,X5))&((X6=cell6|(X6=cell5|(X5=cell3|(X6=cell3|X5=cell1))))|~epred1_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell4|(X6=cell3|X5=cell1))))|~epred1_2(X6,X5))&((X6=cell6|(X5=cell4|(X6=cell4|(X6=cell3|X5=cell1))))|~epred1_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X6=cell4|(X6=cell3|X5=cell1))))|~epred1_2(X6,X5))&((X6=cell6|(X6=cell5|(X6=cell4|(X6=cell3|X5=cell1))))|~epred1_2(X6,X5))))))&((((((X5=cell5|(X5=cell4|(X5=cell3|(X5=cell2|X6=cell2))))|~epred1_2(X6,X5))&((X6=cell6|(X5=cell4|(X5=cell3|(X5=cell2|X6=cell2))))|~epred1_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X5=cell3|(X5=cell2|X6=cell2))))|~epred1_2(X6,X5))&((X6=cell6|(X6=cell5|(X5=cell3|(X5=cell2|X6=cell2))))|~epred1_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell4|(X5=cell2|X6=cell2))))|~epred1_2(X6,X5))&((X6=cell6|(X5=cell4|(X6=cell4|(X5=cell2|X6=cell2))))|~epred1_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X6=cell4|(X5=cell2|X6=cell2))))|~epred1_2(X6,X5))&((X6=cell6|(X6=cell5|(X6=cell4|(X5=cell2|X6=cell2))))|~epred1_2(X6,X5)))))&(((((X5=cell5|(X5=cell4|(X5=cell3|(X6=cell3|X6=cell2))))|~epred1_2(X6,X5))&((X6=cell6|(X5=cell4|(X5=cell3|(X6=cell3|X6=cell2))))|~epred1_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X5=cell3|(X6=cell3|X6=cell2))))|~epred1_2(X6,X5))&((X6=cell6|(X6=cell5|(X5=cell3|(X6=cell3|X6=cell2))))|~epred1_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell4|(X6=cell3|X6=cell2))))|~epred1_2(X6,X5))&((X6=cell6|(X5=cell4|(X6=cell4|(X6=cell3|X6=cell2))))|~epred1_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X6=cell4|(X6=cell3|X6=cell2))))|~epred1_2(X6,X5))&((X6=cell6|(X6=cell5|(X6=cell4|(X6=cell3|X6=cell2))))|~epred1_2(X6,X5)))))))&((((((X7!=cell1|X8!=cell2)|epred1_2(X8,X7))&((X7!=cell2|X8!=cell3)|epred1_2(X8,X7)))&((X7!=cell3|X8!=cell4)|epred1_2(X8,X7)))&((X7!=cell4|X8!=cell5)|epred1_2(X8,X7)))&((X7!=cell5|X8!=cell6)|epred1_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1158])])])])])).
fof(c_0_1169, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell11|(X5=cell9|(X5=cell8|(X5=cell7|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X5=cell9|(X5=cell8|(X5=cell7|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X5=cell8|(X5=cell7|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X6=cell10|(X5=cell8|(X5=cell7|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))))&((((X5=cell11|(X5=cell9|(X6=cell9|(X5=cell7|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X5=cell9|(X6=cell9|(X5=cell7|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X6=cell9|(X5=cell7|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X6=cell10|(X6=cell9|(X5=cell7|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))))&(((((X5=cell11|(X5=cell9|(X5=cell8|(X6=cell8|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X5=cell9|(X5=cell8|(X6=cell8|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X5=cell8|(X6=cell8|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X6=cell10|(X5=cell8|(X6=cell8|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))))&((((X5=cell11|(X5=cell9|(X6=cell9|(X6=cell8|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X5=cell9|(X6=cell9|(X6=cell8|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X6=cell9|(X6=cell8|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X6=cell10|(X6=cell9|(X6=cell8|(X5=cell6|epred1_2(X6,X5))))))|~epred2_2(X6,X5))))))&((((((X5=cell11|(X5=cell9|(X5=cell8|(X5=cell7|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X5=cell9|(X5=cell8|(X5=cell7|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X5=cell8|(X5=cell7|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X6=cell10|(X5=cell8|(X5=cell7|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))))&((((X5=cell11|(X5=cell9|(X6=cell9|(X5=cell7|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X5=cell9|(X6=cell9|(X5=cell7|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X6=cell9|(X5=cell7|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X6=cell10|(X6=cell9|(X5=cell7|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))))&(((((X5=cell11|(X5=cell9|(X5=cell8|(X6=cell8|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X5=cell9|(X5=cell8|(X6=cell8|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X5=cell8|(X6=cell8|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X6=cell10|(X5=cell8|(X6=cell8|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))))&((((X5=cell11|(X5=cell9|(X6=cell9|(X6=cell8|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X5=cell9|(X6=cell9|(X6=cell8|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X6=cell9|(X6=cell8|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5))&((X6=cell12|(X6=cell10|(X6=cell9|(X6=cell8|(X6=cell7|epred1_2(X6,X5))))))|~epred2_2(X6,X5)))))))&((((((~epred1_2(X8,X7)|epred2_2(X8,X7))&((X7!=cell6|X8!=cell7)|epred2_2(X8,X7)))&((X7!=cell7|X8!=cell8)|epred2_2(X8,X7)))&((X7!=cell8|X8!=cell9)|epred2_2(X8,X7)))&((X7!=cell9|X8!=cell10)|epred2_2(X8,X7)))&((X7!=cell11|X8!=cell12)|epred2_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1163])])])])])).
cnf(c_0_1170,plain,(epred2_2(X1,X2)|X2=cell12|X2=cell13|X2=cell14|X2=cell15|X2=cell16|~epred3_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1164])).
cnf(c_0_1171,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred3_2(cell2,esk2_0)|epred3_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1165, c_0_1166]), c_0_938]), c_0_959]), c_0_960]), c_0_961]), c_0_940])).
cnf(c_0_1172,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred4_2(cell99,esk1_0)|epred3_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1160, c_0_1167]), c_0_949]), c_0_969]), c_0_970]), c_0_971]), c_0_950])).
cnf(c_0_1173,plain,(X2=cell1|X1=cell3|X2=cell3|X2=cell4|X2=cell5|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1168])).
cnf(c_0_1174,plain,(epred1_2(X1,X2)|X2=cell6|X2=cell7|X2=cell8|X2=cell9|X1=cell12|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1169])).
cnf(c_0_1175,plain,(epred2_2(X1,X2)|X1=cell13|X1=cell14|X1=cell15|X1=cell16|X1=cell17|~epred3_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1164])).
cnf(c_0_1176,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred3_2(cell2,esk2_0)|epred2_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1170, c_0_1171]), c_0_976]), c_0_978]), c_0_979]), c_0_980]), c_0_957])).
cnf(c_0_1177,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred3_2(cell99,esk1_0)|epred3_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1165, c_0_1172]), c_0_949]), c_0_970]), c_0_971]), c_0_972]), c_0_951])).
cnf(c_0_1178,plain,(cell1=X1|cell3=X1|cell3=X2|cell4=X1|cell5=X1|cell12=X2|cell6=X1|cell7=X1|cell8=X1|cell9=X1|~epred2_2(X2,X1)), inference(spm,[status(thm)],[c_0_1173, c_0_1174])).
cnf(c_0_1179,plain,(esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred2_2(cell2,esk2_0)|epred2_2(esk2_0,cell2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1175, c_0_1176]), c_0_978]), c_0_979]), c_0_980]), c_0_957]), c_0_958])).
cnf(c_0_1180,plain,(X2=cell1|X2=cell2|X2=cell3|X2=cell4|X2=cell5|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1168])).
cnf(c_0_1181,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred3_2(cell99,esk1_0)|epred2_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1170, c_0_1177]), c_0_987]), c_0_989]), c_0_990]), c_0_991]), c_0_968])).
cnf(c_0_1182,plain,(esk2_0=cell3|esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|epred2_2(cell2,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1178, c_0_1179]), c_0_1014]), c_0_1015]), c_0_1016]), c_0_1017]), c_0_995]), c_0_996]), c_0_997]), c_0_998])).
cnf(c_0_1183,plain,(cell2=X1|cell1=X1|cell3=X1|cell4=X1|cell5=X1|cell12=X2|cell6=X1|cell7=X1|cell8=X1|cell9=X1|~epred2_2(X2,X1)), inference(spm,[status(thm)],[c_0_1180, c_0_1174])).
cnf(c_0_1184,plain,(esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred2_2(cell99,esk1_0)|epred2_2(esk1_0,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1175, c_0_1181]), c_0_989]), c_0_990]), c_0_991]), c_0_968]), c_0_969])).
cnf(c_0_1185,negated_conjecture,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred2_2(cell2,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1182, c_0_685]), c_0_1026]), c_0_987])).
cnf(c_0_1186,plain,(X1=cell2|X1=cell3|X2=cell3|X2=cell4|X1=cell6|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1168])).
cnf(c_0_1187,plain,(esk1_0=cell12|esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred2_2(cell99,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1183, c_0_1184]), c_0_669]), c_0_1025]), c_0_1026]), c_0_1027]), c_0_1028]), c_0_1006]), c_0_1007]), c_0_1008]), c_0_1009])).
cnf(c_0_1188,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1183, c_0_1185]), c_0_669]), c_0_1025]), c_0_1026]), c_0_1027]), c_0_1028]), c_0_976]), c_0_1006]), c_0_1007]), c_0_1008]), c_0_1009])).
cnf(c_0_1189,plain,(cell12!=cell11), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_1190,plain,(cell89!=cell11), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_1191,plain,(cell98!=cell11), inference(split_conjunct,[status(thm)],[c_0_206])).
cnf(c_0_1192,plain,(cell100!=cell11), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_1193,plain,(cell2=X1|cell3=X2|cell3=X1|cell4=X2|cell6=X1|cell12=X1|cell6=X2|cell7=X2|cell8=X2|cell9=X2|~epred2_2(X1,X2)), inference(spm,[status(thm)],[c_0_1186, c_0_1174])).
cnf(c_0_1194,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|epred2_2(cell99,cell11)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1187, c_0_1188]), c_0_1189]), c_0_1190]), c_0_1191]), c_0_1192])).
cnf(c_0_1195,plain,(cell3!=cell11), inference(split_conjunct,[status(thm)],[c_0_208])).
cnf(c_0_1196,plain,(cell4!=cell11), inference(split_conjunct,[status(thm)],[c_0_209])).
cnf(c_0_1197,plain,(cell6!=cell11), inference(split_conjunct,[status(thm)],[c_0_210])).
cnf(c_0_1198,plain,(cell7!=cell11), inference(split_conjunct,[status(thm)],[c_0_211])).
cnf(c_0_1199,plain,(cell8!=cell11), inference(split_conjunct,[status(thm)],[c_0_212])).
cnf(c_0_1200,plain,(cell9!=cell11), inference(split_conjunct,[status(thm)],[c_0_213])).
cnf(c_0_1201,plain,(X1=cell2|X1=cell3|X1=cell4|X1=cell5|X1=cell6|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1168])).
cnf(c_0_1202,plain,(epred1_2(X1,X2)|X2=cell6|X1=cell8|X1=cell9|X1=cell10|X1=cell12|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1169])).
cnf(c_0_1203,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1193, c_0_1194]), c_0_669]), c_0_1195]), c_0_1026]), c_0_1196]), c_0_1006]), c_0_987]), c_0_1197]), c_0_1198]), c_0_1199]), c_0_1200])).
cnf(c_0_1204,plain,(cell12!=cell22), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_1205,plain,(cell89!=cell22), inference(split_conjunct,[status(thm)],[c_0_215])).
cnf(c_0_1206,plain,(cell98!=cell22), inference(split_conjunct,[status(thm)],[c_0_216])).
cnf(c_0_1207,plain,(cell100!=cell22), inference(split_conjunct,[status(thm)],[c_0_217])).
cnf(c_0_1208,plain,(cell2=X1|cell3=X1|cell4=X1|cell5=X1|cell6=X1|cell12=X1|cell6=X2|cell8=X1|cell9=X1|cell10=X1|~epred2_2(X1,X2)), inference(spm,[status(thm)],[c_0_1201, c_0_1202])).
cnf(c_0_1209,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|esk1_0=cell12|esk1_0=cell13|epred2_2(cell99,cell22)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1187, c_0_1203]), c_0_1204]), c_0_1205]), c_0_1206]), c_0_1207])).
cnf(c_0_1210,plain,(cell3!=cell22), inference(split_conjunct,[status(thm)],[c_0_218])).
cnf(c_0_1211,plain,(cell4!=cell22), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_1212,plain,(cell6!=cell22), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_1213,plain,(cell7!=cell22), inference(split_conjunct,[status(thm)],[c_0_221])).
cnf(c_0_1214,plain,(cell8!=cell22), inference(split_conjunct,[status(thm)],[c_0_222])).
cnf(c_0_1215,plain,(cell9!=cell22), inference(split_conjunct,[status(thm)],[c_0_223])).
cnf(c_0_1216,plain,(esk1_0=cell6|esk1_0=cell12|esk1_0=cell89|esk1_0=cell98|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1208, c_0_1187]), c_0_669]), c_0_1026]), c_0_1027]), c_0_1028]), c_0_1006]), c_0_987]), c_0_1008]), c_0_1009]), c_0_1010])).
cnf(c_0_1217,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|esk1_0=cell12|esk1_0=cell13), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1193, c_0_1209]), c_0_669]), c_0_1210]), c_0_1026]), c_0_1211]), c_0_1006]), c_0_987]), c_0_1212]), c_0_1213]), c_0_1214]), c_0_1215])).
cnf(c_0_1218,plain,(cell6!=cell13), inference(split_conjunct,[status(thm)],[c_0_224])).
cnf(c_0_1219,plain,(cell12!=cell13), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_1220,plain,(cell89!=cell13), inference(split_conjunct,[status(thm)],[c_0_226])).
cnf(c_0_1221,plain,(cell98!=cell13), inference(split_conjunct,[status(thm)],[c_0_227])).
cnf(c_0_1222,plain,(cell100!=cell13), inference(split_conjunct,[status(thm)],[c_0_228])).
cnf(c_0_1223,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|esk1_0=cell12), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1216, c_0_1217]), c_0_1218]), c_0_1219]), c_0_1220]), c_0_1221]), c_0_1222])).
cnf(c_0_1224,plain,(cell89!=cell12), inference(split_conjunct,[status(thm)],[c_0_229])).
cnf(c_0_1225,plain,(cell98!=cell12), inference(split_conjunct,[status(thm)],[c_0_230])).
cnf(c_0_1226,plain,(cell100!=cell12), inference(split_conjunct,[status(thm)],[c_0_231])).
cnf(c_0_1227,plain,(X1=cell2|X1=cell3|X2=cell3|X2=cell4|X2=cell5|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1168])).
cnf(c_0_1228,plain,(epred1_2(X1,X2)|X2=cell6|X2=cell7|X2=cell8|X2=cell9|X2=cell11|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1169])).
cnf(c_0_1229,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred2_2(cell99,cell12)|epred2_2(cell12,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1184, c_0_1223]), c_0_1224]), c_0_1225]), c_0_1226])).
cnf(c_0_1230,plain,(cell3!=cell12), inference(split_conjunct,[status(thm)],[c_0_232])).
cnf(c_0_1231,plain,(cell4!=cell12), inference(split_conjunct,[status(thm)],[c_0_233])).
cnf(c_0_1232,plain,(cell6!=cell12), inference(split_conjunct,[status(thm)],[c_0_234])).
cnf(c_0_1233,plain,(cell7!=cell12), inference(split_conjunct,[status(thm)],[c_0_235])).
cnf(c_0_1234,plain,(cell8!=cell12), inference(split_conjunct,[status(thm)],[c_0_236])).
cnf(c_0_1235,plain,(cell9!=cell12), inference(split_conjunct,[status(thm)],[c_0_237])).
cnf(c_0_1236,plain,(epred35_2(X1,X2)|X2=cell86|X2=cell87|X2=cell88|X2=cell89|X1=cell100|~epred36_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_657])).
cnf(c_0_1237,plain,(cell2=X1|cell3=X2|cell3=X1|cell4=X2|cell5=X2|cell11=X2|cell6=X2|cell7=X2|cell8=X2|cell9=X2|~epred2_2(X1,X2)), inference(spm,[status(thm)],[c_0_1227, c_0_1228])).
cnf(c_0_1238,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90|epred2_2(cell12,cell99)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1193, c_0_1229]), c_0_669]), c_0_1230]), c_0_1026]), c_0_1231]), c_0_1006]), c_0_987]), c_0_1232]), c_0_1233]), c_0_1234]), c_0_1235])).
cnf(c_0_1239,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell86|esk2_0=cell87|esk1_0=cell100|epred36_2(esk2_0,esk1_0)|epred35_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_1236, c_0_663])).
cnf(c_0_1240,plain,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|esk2_0=cell90), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1237, c_0_1238]), c_0_976]), c_0_1026]), c_0_1230]), c_0_1027]), c_0_1028]), c_0_988]), c_0_1006]), c_0_1007]), c_0_1008]), c_0_1009])).
cnf(c_0_1241,plain,(cell89!=cell90), inference(split_conjunct,[status(thm)],[c_0_238])).
cnf(c_0_1242,plain,(cell88!=cell90), inference(split_conjunct,[status(thm)],[c_0_239])).
cnf(c_0_1243,plain,(cell86!=cell90), inference(split_conjunct,[status(thm)],[c_0_240])).
cnf(c_0_1244,plain,(cell87!=cell90), inference(split_conjunct,[status(thm)],[c_0_241])).
cnf(c_0_1245,plain,(esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred36_2(cell90,esk1_0)|epred35_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1239, c_0_1240]), c_0_1241]), c_0_1242]), c_0_1243]), c_0_1244])).
cnf(c_0_1246,plain,(cell96!=cell90), inference(split_conjunct,[status(thm)],[c_0_242])).
cnf(c_0_1247,plain,(cell97!=cell90), inference(split_conjunct,[status(thm)],[c_0_243])).
cnf(c_0_1248,plain,(cell98!=cell90), inference(split_conjunct,[status(thm)],[c_0_244])).
cnf(c_0_1249,plain,(cell100!=cell90), inference(split_conjunct,[status(thm)],[c_0_245])).
cnf(c_0_1250,plain,(esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred35_2(cell90,esk1_0)|epred35_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_662, c_0_1245]), c_0_688]), c_0_1246]), c_0_1247]), c_0_1248]), c_0_1249])).
cnf(c_0_1251,plain,(cell81!=cell90), inference(split_conjunct,[status(thm)],[c_0_246])).
cnf(c_0_1252,plain,(cell82!=cell90), inference(split_conjunct,[status(thm)],[c_0_247])).
cnf(c_0_1253,plain,(cell83!=cell90), inference(split_conjunct,[status(thm)],[c_0_248])).
cnf(c_0_1254,plain,(cell84!=cell90), inference(split_conjunct,[status(thm)],[c_0_249])).
cnf(c_0_1255,plain,(cell85!=cell90), inference(split_conjunct,[status(thm)],[c_0_250])).
cnf(c_0_1256,plain,(esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred35_2(cell90,esk1_0)|epred34_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_1250]), c_0_1251]), c_0_1252]), c_0_1253]), c_0_1254]), c_0_1255])).
cnf(c_0_1257,plain,(cell91!=cell90), inference(split_conjunct,[status(thm)],[c_0_251])).
cnf(c_0_1258,plain,(cell92!=cell90), inference(split_conjunct,[status(thm)],[c_0_252])).
cnf(c_0_1259,plain,(cell93!=cell90), inference(split_conjunct,[status(thm)],[c_0_253])).
cnf(c_0_1260,plain,(cell94!=cell90), inference(split_conjunct,[status(thm)],[c_0_254])).
cnf(c_0_1261,plain,(cell95!=cell90), inference(split_conjunct,[status(thm)],[c_0_255])).
cnf(c_0_1262,plain,(esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred34_2(cell90,esk1_0)|epred34_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_1256]), c_0_1257]), c_0_1258]), c_0_1259]), c_0_1260]), c_0_1261])).
cnf(c_0_1263,plain,(cell79!=cell90), inference(split_conjunct,[status(thm)],[c_0_256])).
cnf(c_0_1264,plain,(cell76!=cell90), inference(split_conjunct,[status(thm)],[c_0_257])).
cnf(c_0_1265,plain,(cell77!=cell90), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_1266,plain,(cell78!=cell90), inference(split_conjunct,[status(thm)],[c_0_259])).
cnf(c_0_1267,plain,(cell80!=cell90), inference(split_conjunct,[status(thm)],[c_0_260])).
cnf(c_0_1268,plain,(epred33_2(X1,X2)|X1=cell86|X1=cell87|X1=cell88|X1=cell89|X2=cell80|~epred34_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_706])).
cnf(c_0_1269,plain,(esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred34_2(cell90,esk1_0)|epred33_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_721, c_0_1262]), c_0_1263]), c_0_1264]), c_0_1265]), c_0_1266]), c_0_1267])).
cnf(c_0_1270,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred33_2(cell90,esk1_0)|epred33_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1268, c_0_1269]), c_0_1241]), c_0_1242]), c_0_1243]), c_0_1244])).
cnf(c_0_1271,plain,(cell71!=cell90), inference(split_conjunct,[status(thm)],[c_0_261])).
cnf(c_0_1272,plain,(cell72!=cell90), inference(split_conjunct,[status(thm)],[c_0_262])).
cnf(c_0_1273,plain,(cell73!=cell90), inference(split_conjunct,[status(thm)],[c_0_263])).
cnf(c_0_1274,plain,(cell74!=cell90), inference(split_conjunct,[status(thm)],[c_0_264])).
cnf(c_0_1275,plain,(cell75!=cell90), inference(split_conjunct,[status(thm)],[c_0_265])).
cnf(c_0_1276,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred33_2(cell90,esk1_0)|epred32_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_745, c_0_1270]), c_0_1271]), c_0_1272]), c_0_1273]), c_0_1274]), c_0_1275])).
cnf(c_0_1277,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred32_2(cell90,esk1_0)|epred32_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_755, c_0_1276]), c_0_1251]), c_0_1252]), c_0_1253]), c_0_1254]), c_0_1255])).
cnf(c_0_1278,plain,(cell66!=cell90), inference(split_conjunct,[status(thm)],[c_0_266])).
cnf(c_0_1279,plain,(cell67!=cell90), inference(split_conjunct,[status(thm)],[c_0_267])).
cnf(c_0_1280,plain,(cell68!=cell90), inference(split_conjunct,[status(thm)],[c_0_268])).
cnf(c_0_1281,plain,(cell69!=cell90), inference(split_conjunct,[status(thm)],[c_0_269])).
cnf(c_0_1282,plain,(cell70!=cell90), inference(split_conjunct,[status(thm)],[c_0_270])).
cnf(c_0_1283,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred32_2(cell90,esk1_0)|epred31_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1277]), c_0_1278]), c_0_1279]), c_0_1280]), c_0_1281]), c_0_1282])).
cnf(c_0_1284,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred31_2(cell90,esk1_0)|epred31_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_774, c_0_1283]), c_0_1263]), c_0_1264]), c_0_1265]), c_0_1266]), c_0_1267])).
cnf(c_0_1285,plain,(cell61!=cell90), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_1286,plain,(cell62!=cell90), inference(split_conjunct,[status(thm)],[c_0_272])).
cnf(c_0_1287,plain,(cell63!=cell90), inference(split_conjunct,[status(thm)],[c_0_273])).
cnf(c_0_1288,plain,(cell64!=cell90), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_1289,plain,(cell65!=cell90), inference(split_conjunct,[status(thm)],[c_0_275])).
cnf(c_0_1290,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred31_2(cell90,esk1_0)|epred30_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_783, c_0_1284]), c_0_1285]), c_0_1286]), c_0_1287]), c_0_1288]), c_0_1289])).
cnf(c_0_1291,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred30_2(cell90,esk1_0)|epred30_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_1290]), c_0_1271]), c_0_1272]), c_0_1273]), c_0_1274]), c_0_1275])).
cnf(c_0_1292,plain,(cell56!=cell90), inference(split_conjunct,[status(thm)],[c_0_276])).
cnf(c_0_1293,plain,(cell57!=cell90), inference(split_conjunct,[status(thm)],[c_0_277])).
cnf(c_0_1294,plain,(cell58!=cell90), inference(split_conjunct,[status(thm)],[c_0_278])).
cnf(c_0_1295,plain,(cell59!=cell90), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_1296,plain,(cell60!=cell90), inference(split_conjunct,[status(thm)],[c_0_280])).
cnf(c_0_1297,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred30_2(cell90,esk1_0)|epred29_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_803, c_0_1291]), c_0_1292]), c_0_1293]), c_0_1294]), c_0_1295]), c_0_1296])).
cnf(c_0_1298,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred29_2(cell90,esk1_0)|epred29_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_813, c_0_1297]), c_0_1278]), c_0_1279]), c_0_1280]), c_0_1281]), c_0_1282])).
cnf(c_0_1299,plain,(cell51!=cell90), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_1300,plain,(cell52!=cell90), inference(split_conjunct,[status(thm)],[c_0_282])).
cnf(c_0_1301,plain,(cell53!=cell90), inference(split_conjunct,[status(thm)],[c_0_283])).
cnf(c_0_1302,plain,(cell54!=cell90), inference(split_conjunct,[status(thm)],[c_0_284])).
cnf(c_0_1303,plain,(cell55!=cell90), inference(split_conjunct,[status(thm)],[c_0_285])).
cnf(c_0_1304,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred29_2(cell90,esk1_0)|epred28_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_822, c_0_1298]), c_0_1299]), c_0_1300]), c_0_1301]), c_0_1302]), c_0_1303])).
cnf(c_0_1305,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred28_2(cell90,esk1_0)|epred28_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_832, c_0_1304]), c_0_1285]), c_0_1286]), c_0_1287]), c_0_1288]), c_0_1289])).
cnf(c_0_1306,plain,(cell46!=cell90), inference(split_conjunct,[status(thm)],[c_0_286])).
cnf(c_0_1307,plain,(cell47!=cell90), inference(split_conjunct,[status(thm)],[c_0_287])).
cnf(c_0_1308,plain,(cell48!=cell90), inference(split_conjunct,[status(thm)],[c_0_288])).
cnf(c_0_1309,plain,(cell49!=cell90), inference(split_conjunct,[status(thm)],[c_0_289])).
cnf(c_0_1310,plain,(cell50!=cell90), inference(split_conjunct,[status(thm)],[c_0_290])).
cnf(c_0_1311,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred28_2(cell90,esk1_0)|epred27_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_841, c_0_1305]), c_0_1306]), c_0_1307]), c_0_1308]), c_0_1309]), c_0_1310])).
cnf(c_0_1312,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred27_2(cell90,esk1_0)|epred27_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_851, c_0_1311]), c_0_1292]), c_0_1293]), c_0_1294]), c_0_1295]), c_0_1296])).
cnf(c_0_1313,plain,(cell41!=cell90), inference(split_conjunct,[status(thm)],[c_0_291])).
cnf(c_0_1314,plain,(cell42!=cell90), inference(split_conjunct,[status(thm)],[c_0_292])).
cnf(c_0_1315,plain,(cell43!=cell90), inference(split_conjunct,[status(thm)],[c_0_293])).
cnf(c_0_1316,plain,(cell44!=cell90), inference(split_conjunct,[status(thm)],[c_0_294])).
cnf(c_0_1317,plain,(cell45!=cell90), inference(split_conjunct,[status(thm)],[c_0_295])).
cnf(c_0_1318,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred27_2(cell90,esk1_0)|epred26_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_860, c_0_1312]), c_0_1313]), c_0_1314]), c_0_1315]), c_0_1316]), c_0_1317])).
cnf(c_0_1319,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred26_2(cell90,esk1_0)|epred26_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_870, c_0_1318]), c_0_1299]), c_0_1300]), c_0_1301]), c_0_1302]), c_0_1303])).
cnf(c_0_1320,plain,(cell36!=cell90), inference(split_conjunct,[status(thm)],[c_0_296])).
cnf(c_0_1321,plain,(cell37!=cell90), inference(split_conjunct,[status(thm)],[c_0_297])).
cnf(c_0_1322,plain,(cell38!=cell90), inference(split_conjunct,[status(thm)],[c_0_298])).
cnf(c_0_1323,plain,(cell39!=cell90), inference(split_conjunct,[status(thm)],[c_0_299])).
cnf(c_0_1324,plain,(cell40!=cell90), inference(split_conjunct,[status(thm)],[c_0_300])).
cnf(c_0_1325,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred26_2(cell90,esk1_0)|epred25_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_879, c_0_1319]), c_0_1320]), c_0_1321]), c_0_1322]), c_0_1323]), c_0_1324])).
cnf(c_0_1326,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred25_2(cell90,esk1_0)|epred25_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_889, c_0_1325]), c_0_1306]), c_0_1307]), c_0_1308]), c_0_1309]), c_0_1310])).
cnf(c_0_1327,plain,(cell31!=cell90), inference(split_conjunct,[status(thm)],[c_0_301])).
cnf(c_0_1328,plain,(cell32!=cell90), inference(split_conjunct,[status(thm)],[c_0_302])).
cnf(c_0_1329,plain,(cell33!=cell90), inference(split_conjunct,[status(thm)],[c_0_303])).
cnf(c_0_1330,plain,(cell34!=cell90), inference(split_conjunct,[status(thm)],[c_0_304])).
cnf(c_0_1331,plain,(cell35!=cell90), inference(split_conjunct,[status(thm)],[c_0_305])).
cnf(c_0_1332,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred25_2(cell90,esk1_0)|epred24_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_898, c_0_1326]), c_0_1327]), c_0_1328]), c_0_1329]), c_0_1330]), c_0_1331])).
cnf(c_0_1333,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred24_2(cell90,esk1_0)|epred24_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_908, c_0_1332]), c_0_1313]), c_0_1314]), c_0_1315]), c_0_1316]), c_0_1317])).
cnf(c_0_1334,plain,(cell26!=cell90), inference(split_conjunct,[status(thm)],[c_0_306])).
cnf(c_0_1335,plain,(cell27!=cell90), inference(split_conjunct,[status(thm)],[c_0_307])).
cnf(c_0_1336,plain,(cell28!=cell90), inference(split_conjunct,[status(thm)],[c_0_308])).
cnf(c_0_1337,plain,(cell29!=cell90), inference(split_conjunct,[status(thm)],[c_0_309])).
cnf(c_0_1338,plain,(cell30!=cell90), inference(split_conjunct,[status(thm)],[c_0_310])).
cnf(c_0_1339,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred24_2(cell90,esk1_0)|epred23_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_917, c_0_1333]), c_0_1334]), c_0_1335]), c_0_1336]), c_0_1337]), c_0_1338])).
cnf(c_0_1340,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred23_2(cell90,esk1_0)|epred23_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_927, c_0_1339]), c_0_1320]), c_0_1321]), c_0_1322]), c_0_1323]), c_0_1324])).
cnf(c_0_1341,plain,(cell90!=cell22), inference(split_conjunct,[status(thm)],[c_0_311])).
cnf(c_0_1342,plain,(cell21!=cell90), inference(split_conjunct,[status(thm)],[c_0_312])).
cnf(c_0_1343,plain,(cell23!=cell90), inference(split_conjunct,[status(thm)],[c_0_313])).
cnf(c_0_1344,plain,(cell24!=cell90), inference(split_conjunct,[status(thm)],[c_0_314])).
cnf(c_0_1345,plain,(cell25!=cell90), inference(split_conjunct,[status(thm)],[c_0_315])).
cnf(c_0_1346,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred23_2(cell90,esk1_0)|epred22_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_1340]), c_0_1341]), c_0_1342]), c_0_1343]), c_0_1344]), c_0_1345])).
cnf(c_0_1347,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred22_2(cell90,esk1_0)|epred22_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_946, c_0_1346]), c_0_1327]), c_0_1328]), c_0_1329]), c_0_1330]), c_0_1331])).
cnf(c_0_1348,plain,(cell16!=cell90), inference(split_conjunct,[status(thm)],[c_0_316])).
cnf(c_0_1349,plain,(cell17!=cell90), inference(split_conjunct,[status(thm)],[c_0_317])).
cnf(c_0_1350,plain,(cell18!=cell90), inference(split_conjunct,[status(thm)],[c_0_318])).
cnf(c_0_1351,plain,(cell19!=cell90), inference(split_conjunct,[status(thm)],[c_0_319])).
cnf(c_0_1352,plain,(cell20!=cell90), inference(split_conjunct,[status(thm)],[c_0_320])).
cnf(c_0_1353,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred22_2(cell90,esk1_0)|epred21_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_955, c_0_1347]), c_0_1348]), c_0_1349]), c_0_1350]), c_0_1351]), c_0_1352])).
cnf(c_0_1354,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred21_2(cell90,esk1_0)|epred21_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_965, c_0_1353]), c_0_1334]), c_0_1335]), c_0_1336]), c_0_1337]), c_0_1338])).
cnf(c_0_1355,plain,(cell90!=cell12), inference(split_conjunct,[status(thm)],[c_0_321])).
cnf(c_0_1356,plain,(cell90!=cell11), inference(split_conjunct,[status(thm)],[c_0_322])).
cnf(c_0_1357,plain,(cell90!=cell13), inference(split_conjunct,[status(thm)],[c_0_323])).
cnf(c_0_1358,plain,(cell14!=cell90), inference(split_conjunct,[status(thm)],[c_0_324])).
cnf(c_0_1359,plain,(cell15!=cell90), inference(split_conjunct,[status(thm)],[c_0_325])).
cnf(c_0_1360,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred21_2(cell90,esk1_0)|epred20_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_974, c_0_1354]), c_0_1355]), c_0_1356]), c_0_1357]), c_0_1358]), c_0_1359])).
cnf(c_0_1361,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred20_2(cell90,esk1_0)|epred20_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_984, c_0_1360]), c_0_1341]), c_0_1342]), c_0_1343]), c_0_1344]), c_0_1345])).
cnf(c_0_1362,plain,(cell6!=cell90), inference(split_conjunct,[status(thm)],[c_0_326])).
cnf(c_0_1363,plain,(cell7!=cell90), inference(split_conjunct,[status(thm)],[c_0_327])).
cnf(c_0_1364,plain,(cell8!=cell90), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_1365,plain,(cell9!=cell90), inference(split_conjunct,[status(thm)],[c_0_329])).
cnf(c_0_1366,plain,(cell10!=cell90), inference(split_conjunct,[status(thm)],[c_0_330])).
cnf(c_0_1367,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred20_2(cell90,esk1_0)|epred19_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_993, c_0_1361]), c_0_1362]), c_0_1363]), c_0_1364]), c_0_1365]), c_0_1366])).
cnf(c_0_1368,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred19_2(cell90,esk1_0)|epred19_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1367]), c_0_1348]), c_0_1349]), c_0_1350]), c_0_1351]), c_0_1352])).
cnf(c_0_1369,plain,(cell1!=cell90), inference(split_conjunct,[status(thm)],[c_0_331])).
cnf(c_0_1370,plain,(cell3!=cell90), inference(split_conjunct,[status(thm)],[c_0_332])).
cnf(c_0_1371,plain,(cell4!=cell90), inference(split_conjunct,[status(thm)],[c_0_333])).
cnf(c_0_1372,plain,(cell5!=cell90), inference(split_conjunct,[status(thm)],[c_0_334])).
cnf(c_0_1373,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred19_2(cell90,esk1_0)|epred18_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1023, c_0_1368]), c_0_681]), c_0_1369]), c_0_1370]), c_0_1371]), c_0_1372])).
cnf(c_0_1374,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred18_2(cell90,esk1_0)|epred18_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1021, c_0_1373]), c_0_1355]), c_0_1356]), c_0_1357]), c_0_1358]), c_0_1359])).
cnf(c_0_1375,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred18_2(cell90,esk1_0)|epred17_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1030, c_0_1374]), c_0_688]), c_0_1261]), c_0_1246]), c_0_1247]), c_0_1248])).
cnf(c_0_1376,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred17_2(cell90,esk1_0)|epred17_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1035, c_0_1375]), c_0_688]), c_0_1246]), c_0_1247]), c_0_1248]), c_0_1249])).
cnf(c_0_1377,plain,(epred16_2(X1,X2)|X2=cell89|X1=cell92|X1=cell93|X1=cell94|X1=cell95|~epred17_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1034])).
cnf(c_0_1378,plain,(esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred17_2(cell90,esk1_0)|epred16_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1040, c_0_1376]), c_0_1241]), c_0_1257]), c_0_1258]), c_0_1259]), c_0_1260])).
cnf(c_0_1379,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred16_2(cell90,esk1_0)|epred16_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1377, c_0_1378]), c_0_1258]), c_0_1259]), c_0_1260]), c_0_1261])).
cnf(c_0_1380,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred16_2(cell90,esk1_0)|epred15_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1057, c_0_1379]), c_0_1254]), c_0_1255]), c_0_1243]), c_0_1244])).
cnf(c_0_1381,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred15_2(cell90,esk1_0)|epred15_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1055, c_0_1380]), c_0_1241]), c_0_1242]), c_0_1255]), c_0_1243]), c_0_1244])).
cnf(c_0_1382,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred15_2(cell90,esk1_0)|epred14_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1060, c_0_1381]), c_0_1263]), c_0_1266]), c_0_1251]), c_0_1252]), c_0_1253])).
cnf(c_0_1383,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred14_2(cell90,esk1_0)|epred14_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1065, c_0_1382]), c_0_1263]), c_0_1267]), c_0_1252]), c_0_1253]), c_0_1254])).
cnf(c_0_1384,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred14_2(cell90,esk1_0)|epred13_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1069, c_0_1383]), c_0_1273]), c_0_1274]), c_0_1275]), c_0_1264]), c_0_1265])).
cnf(c_0_1385,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred13_2(cell90,esk1_0)|epred13_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1074, c_0_1384]), c_0_1274]), c_0_1275]), c_0_1264]), c_0_1265]), c_0_1266])).
cnf(c_0_1386,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred13_2(cell90,esk1_0)|epred12_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1078, c_0_1385]), c_0_1279]), c_0_1280]), c_0_1281]), c_0_1271]), c_0_1272])).
cnf(c_0_1387,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred12_2(cell90,esk1_0)|epred12_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1083, c_0_1386]), c_0_1280]), c_0_1281]), c_0_1282]), c_0_1272]), c_0_1273])).
cnf(c_0_1388,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred12_2(cell90,esk1_0)|epred11_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1087, c_0_1387]), c_0_1286]), c_0_1287]), c_0_1288]), c_0_1289]), c_0_1278])).
cnf(c_0_1389,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred11_2(cell90,esk1_0)|epred11_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1092, c_0_1388]), c_0_1287]), c_0_1288]), c_0_1289]), c_0_1278]), c_0_1279])).
cnf(c_0_1390,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred11_2(cell90,esk1_0)|epred10_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1096, c_0_1389]), c_0_1292]), c_0_1293]), c_0_1294]), c_0_1295]), c_0_1285])).
cnf(c_0_1391,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred10_2(cell90,esk1_0)|epred10_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1101, c_0_1390]), c_0_1293]), c_0_1294]), c_0_1295]), c_0_1296]), c_0_1286])).
cnf(c_0_1392,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred10_2(cell90,esk1_0)|epred9_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1105, c_0_1391]), c_0_1299]), c_0_1300]), c_0_1301]), c_0_1302]), c_0_1303])).
cnf(c_0_1393,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred9_2(cell90,esk1_0)|epred9_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1110, c_0_1392]), c_0_1300]), c_0_1301]), c_0_1302]), c_0_1303]), c_0_1292])).
cnf(c_0_1394,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred9_2(cell90,esk1_0)|epred8_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1114, c_0_1393]), c_0_1317]), c_0_1306]), c_0_1307]), c_0_1308]), c_0_1309])).
cnf(c_0_1395,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred8_2(cell90,esk1_0)|epred8_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1119, c_0_1394]), c_0_1306]), c_0_1307]), c_0_1308]), c_0_1309]), c_0_1310])).
cnf(c_0_1396,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred8_2(cell90,esk1_0)|epred7_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1123, c_0_1395]), c_0_1323]), c_0_1313]), c_0_1314]), c_0_1315]), c_0_1316])).
cnf(c_0_1397,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred7_2(cell90,esk1_0)|epred7_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1128, c_0_1396]), c_0_1324]), c_0_1314]), c_0_1315]), c_0_1316]), c_0_1317])).
cnf(c_0_1398,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred7_2(cell90,esk1_0)|epred6_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1132, c_0_1397]), c_0_1330]), c_0_1331]), c_0_1320]), c_0_1321]), c_0_1322])).
cnf(c_0_1399,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred6_2(cell90,esk1_0)|epred6_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1137, c_0_1398]), c_0_1331]), c_0_1320]), c_0_1321]), c_0_1322]), c_0_1323])).
cnf(c_0_1400,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred6_2(cell90,esk1_0)|epred5_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1141, c_0_1399]), c_0_1336]), c_0_1337]), c_0_1327]), c_0_1328]), c_0_1329])).
cnf(c_0_1401,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred5_2(cell90,esk1_0)|epred5_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1146, c_0_1400]), c_0_1337]), c_0_1338]), c_0_1328]), c_0_1329]), c_0_1330])).
cnf(c_0_1402,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred5_2(cell90,esk1_0)|epred4_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1150, c_0_1401]), c_0_1343]), c_0_1344]), c_0_1345]), c_0_1334]), c_0_1335])).
cnf(c_0_1403,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred4_2(cell90,esk1_0)|epred4_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1155, c_0_1402]), c_0_1344]), c_0_1345]), c_0_1334]), c_0_1335]), c_0_1336])).
cnf(c_0_1404,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred4_2(cell90,esk1_0)|epred3_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1160, c_0_1403]), c_0_1341]), c_0_1349]), c_0_1350]), c_0_1351]), c_0_1342])).
cnf(c_0_1405,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred3_2(cell90,esk1_0)|epred3_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1165, c_0_1404]), c_0_1341]), c_0_1350]), c_0_1351]), c_0_1352]), c_0_1343])).
cnf(c_0_1406,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred3_2(cell90,esk1_0)|epred2_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1170, c_0_1405]), c_0_1355]), c_0_1357]), c_0_1358]), c_0_1359]), c_0_1348])).
cnf(c_0_1407,plain,(epred1_2(X1,X2)|X1=cell7|X1=cell8|X1=cell9|X1=cell10|X1=cell12|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1169])).
cnf(c_0_1408,plain,(cell2=X1|cell1=X1|cell3=X1|cell4=X1|cell5=X1|cell11=X1|cell6=X1|cell7=X1|cell8=X1|cell9=X1|~epred2_2(X2,X1)), inference(spm,[status(thm)],[c_0_1180, c_0_1228])).
cnf(c_0_1409,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred2_2(cell90,esk1_0)|epred2_2(esk1_0,cell90)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1175, c_0_1406]), c_0_1357]), c_0_1358]), c_0_1359]), c_0_1348]), c_0_1349])).
cnf(c_0_1410,plain,(cell2=X1|cell3=X1|cell4=X1|cell5=X1|cell6=X1|cell12=X1|cell7=X1|cell8=X1|cell9=X1|cell10=X1|~epred2_2(X1,X2)), inference(spm,[status(thm)],[c_0_1201, c_0_1407])).
cnf(c_0_1411,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79|epred2_2(cell90,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1408, c_0_1409]), c_0_681]), c_0_1369]), c_0_1370]), c_0_1371]), c_0_1372]), c_0_1356]), c_0_1362]), c_0_1363]), c_0_1364]), c_0_1365])).
cnf(c_0_1412,plain,(esk1_0=cell89|esk1_0=cell80|esk1_0=cell100|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1410, c_0_1411]), c_0_681]), c_0_1370]), c_0_1371]), c_0_1372]), c_0_1362]), c_0_1355]), c_0_1363]), c_0_1364]), c_0_1365]), c_0_1366])).
cnf(c_0_1413,negated_conjecture,(esk1_0=cell89|esk1_0=cell80|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1412]), c_0_1226]), c_0_1222]), c_0_1207]), c_0_1192]), c_0_673])).
cnf(c_0_1414,plain,(cell80!=cell12), inference(split_conjunct,[status(thm)],[c_0_335])).
cnf(c_0_1415,plain,(cell80!=cell13), inference(split_conjunct,[status(thm)],[c_0_336])).
cnf(c_0_1416,plain,(cell80!=cell22), inference(split_conjunct,[status(thm)],[c_0_337])).
cnf(c_0_1417,plain,(cell80!=cell11), inference(split_conjunct,[status(thm)],[c_0_338])).
cnf(c_0_1418,negated_conjecture,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|esk2_0=cell79), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1413]), c_0_1414]), c_0_1415]), c_0_1416]), c_0_1417]), c_0_727])).
cnf(c_0_1419,negated_conjecture,(esk2_0=cell89|esk2_0=cell88|esk2_0=cell79), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1418]), c_0_1224]), c_0_1220]), c_0_1205]), c_0_1190]), c_0_679])).
cnf(c_0_1420,plain,(cell89!=cell79), inference(split_conjunct,[status(thm)],[c_0_339])).
cnf(c_0_1421,plain,(cell88!=cell79), inference(split_conjunct,[status(thm)],[c_0_340])).
cnf(c_0_1422,plain,(cell86!=cell79), inference(split_conjunct,[status(thm)],[c_0_341])).
cnf(c_0_1423,plain,(cell87!=cell79), inference(split_conjunct,[status(thm)],[c_0_342])).
cnf(c_0_1424,plain,(esk2_0=cell89|esk2_0=cell88|epred36_2(cell79,esk1_0)|epred35_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_684, c_0_1419]), c_0_1420]), c_0_1421]), c_0_1263]), c_0_1422]), c_0_1423])).
cnf(c_0_1425,plain,(cell96!=cell79), inference(split_conjunct,[status(thm)],[c_0_343])).
cnf(c_0_1426,plain,(cell97!=cell79), inference(split_conjunct,[status(thm)],[c_0_344])).
cnf(c_0_1427,plain,(cell98!=cell79), inference(split_conjunct,[status(thm)],[c_0_345])).
cnf(c_0_1428,plain,(cell100!=cell79), inference(split_conjunct,[status(thm)],[c_0_346])).
cnf(c_0_1429,plain,(esk2_0=cell89|esk2_0=cell88|epred35_2(cell79,esk1_0)|epred35_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_662, c_0_1424]), c_0_791]), c_0_1425]), c_0_1426]), c_0_1427]), c_0_1428])).
cnf(c_0_1430,plain,(cell81!=cell79), inference(split_conjunct,[status(thm)],[c_0_347])).
cnf(c_0_1431,plain,(cell82!=cell79), inference(split_conjunct,[status(thm)],[c_0_348])).
cnf(c_0_1432,plain,(cell83!=cell79), inference(split_conjunct,[status(thm)],[c_0_349])).
cnf(c_0_1433,plain,(cell84!=cell79), inference(split_conjunct,[status(thm)],[c_0_350])).
cnf(c_0_1434,plain,(cell85!=cell79), inference(split_conjunct,[status(thm)],[c_0_351])).
cnf(c_0_1435,plain,(esk2_0=cell89|esk2_0=cell88|epred35_2(cell79,esk1_0)|epred34_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_1429]), c_0_1430]), c_0_1431]), c_0_1432]), c_0_1433]), c_0_1434])).
cnf(c_0_1436,plain,(cell91!=cell79), inference(split_conjunct,[status(thm)],[c_0_352])).
cnf(c_0_1437,plain,(cell92!=cell79), inference(split_conjunct,[status(thm)],[c_0_353])).
cnf(c_0_1438,plain,(cell93!=cell79), inference(split_conjunct,[status(thm)],[c_0_354])).
cnf(c_0_1439,plain,(cell94!=cell79), inference(split_conjunct,[status(thm)],[c_0_355])).
cnf(c_0_1440,plain,(cell95!=cell79), inference(split_conjunct,[status(thm)],[c_0_356])).
cnf(c_0_1441,plain,(esk2_0=cell89|esk2_0=cell88|epred34_2(cell79,esk1_0)|epred34_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_1435]), c_0_1436]), c_0_1437]), c_0_1438]), c_0_1439]), c_0_1440])).
cnf(c_0_1442,plain,(cell76!=cell79), inference(split_conjunct,[status(thm)],[c_0_357])).
cnf(c_0_1443,plain,(cell77!=cell79), inference(split_conjunct,[status(thm)],[c_0_358])).
cnf(c_0_1444,plain,(cell78!=cell79), inference(split_conjunct,[status(thm)],[c_0_359])).
cnf(c_0_1445,plain,(cell80!=cell79), inference(split_conjunct,[status(thm)],[c_0_360])).
cnf(c_0_1446,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred34_2(cell79,esk1_0)|epred33_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_738, c_0_1441]), c_0_1442]), c_0_1443]), c_0_1444]), c_0_1445])).
cnf(c_0_1447,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred33_2(cell79,esk1_0)|epred33_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_736, c_0_1446]), c_0_1420]), c_0_1421]), c_0_1263]), c_0_1422]), c_0_1423])).
cnf(c_0_1448,plain,(cell71!=cell79), inference(split_conjunct,[status(thm)],[c_0_361])).
cnf(c_0_1449,plain,(cell72!=cell79), inference(split_conjunct,[status(thm)],[c_0_362])).
cnf(c_0_1450,plain,(cell73!=cell79), inference(split_conjunct,[status(thm)],[c_0_363])).
cnf(c_0_1451,plain,(cell74!=cell79), inference(split_conjunct,[status(thm)],[c_0_364])).
cnf(c_0_1452,plain,(cell75!=cell79), inference(split_conjunct,[status(thm)],[c_0_365])).
cnf(c_0_1453,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred33_2(cell79,esk1_0)|epred32_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_745, c_0_1447]), c_0_1448]), c_0_1449]), c_0_1450]), c_0_1451]), c_0_1452])).
cnf(c_0_1454,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred32_2(cell79,esk1_0)|epred32_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_755, c_0_1453]), c_0_1430]), c_0_1431]), c_0_1432]), c_0_1433]), c_0_1434])).
cnf(c_0_1455,plain,(cell66!=cell79), inference(split_conjunct,[status(thm)],[c_0_366])).
cnf(c_0_1456,plain,(cell67!=cell79), inference(split_conjunct,[status(thm)],[c_0_367])).
cnf(c_0_1457,plain,(cell68!=cell79), inference(split_conjunct,[status(thm)],[c_0_368])).
cnf(c_0_1458,plain,(cell69!=cell79), inference(split_conjunct,[status(thm)],[c_0_369])).
cnf(c_0_1459,plain,(cell70!=cell79), inference(split_conjunct,[status(thm)],[c_0_370])).
cnf(c_0_1460,plain,(epred31_2(X1,X2)|X1=cell76|X1=cell77|X1=cell78|X2=cell69|X1=cell80|~epred32_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_754])).
cnf(c_0_1461,plain,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred32_2(cell79,esk1_0)|epred31_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1454]), c_0_1455]), c_0_1456]), c_0_1457]), c_0_1458]), c_0_1459])).
cnf(c_0_1462,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred31_2(cell79,esk1_0)|epred31_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1460, c_0_1461]), c_0_1442]), c_0_1443]), c_0_1444]), c_0_1445])).
cnf(c_0_1463,plain,(cell61!=cell79), inference(split_conjunct,[status(thm)],[c_0_371])).
cnf(c_0_1464,plain,(cell62!=cell79), inference(split_conjunct,[status(thm)],[c_0_372])).
cnf(c_0_1465,plain,(cell63!=cell79), inference(split_conjunct,[status(thm)],[c_0_373])).
cnf(c_0_1466,plain,(cell64!=cell79), inference(split_conjunct,[status(thm)],[c_0_374])).
cnf(c_0_1467,plain,(cell65!=cell79), inference(split_conjunct,[status(thm)],[c_0_375])).
cnf(c_0_1468,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred31_2(cell79,esk1_0)|epred30_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_783, c_0_1462]), c_0_1463]), c_0_1464]), c_0_1465]), c_0_1466]), c_0_1467])).
cnf(c_0_1469,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred30_2(cell79,esk1_0)|epred30_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_1468]), c_0_1448]), c_0_1449]), c_0_1450]), c_0_1451]), c_0_1452])).
cnf(c_0_1470,plain,(cell56!=cell79), inference(split_conjunct,[status(thm)],[c_0_376])).
cnf(c_0_1471,plain,(cell57!=cell79), inference(split_conjunct,[status(thm)],[c_0_377])).
cnf(c_0_1472,plain,(cell58!=cell79), inference(split_conjunct,[status(thm)],[c_0_378])).
cnf(c_0_1473,plain,(cell59!=cell79), inference(split_conjunct,[status(thm)],[c_0_379])).
cnf(c_0_1474,plain,(cell60!=cell79), inference(split_conjunct,[status(thm)],[c_0_380])).
cnf(c_0_1475,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred30_2(cell79,esk1_0)|epred29_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_803, c_0_1469]), c_0_1470]), c_0_1471]), c_0_1472]), c_0_1473]), c_0_1474])).
cnf(c_0_1476,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred29_2(cell79,esk1_0)|epred29_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_813, c_0_1475]), c_0_1455]), c_0_1456]), c_0_1457]), c_0_1458]), c_0_1459])).
cnf(c_0_1477,plain,(cell51!=cell79), inference(split_conjunct,[status(thm)],[c_0_381])).
cnf(c_0_1478,plain,(cell52!=cell79), inference(split_conjunct,[status(thm)],[c_0_382])).
cnf(c_0_1479,plain,(cell53!=cell79), inference(split_conjunct,[status(thm)],[c_0_383])).
cnf(c_0_1480,plain,(cell54!=cell79), inference(split_conjunct,[status(thm)],[c_0_384])).
cnf(c_0_1481,plain,(cell55!=cell79), inference(split_conjunct,[status(thm)],[c_0_385])).
cnf(c_0_1482,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred29_2(cell79,esk1_0)|epred28_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_822, c_0_1476]), c_0_1477]), c_0_1478]), c_0_1479]), c_0_1480]), c_0_1481])).
cnf(c_0_1483,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred28_2(cell79,esk1_0)|epred28_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_832, c_0_1482]), c_0_1463]), c_0_1464]), c_0_1465]), c_0_1466]), c_0_1467])).
cnf(c_0_1484,plain,(cell46!=cell79), inference(split_conjunct,[status(thm)],[c_0_386])).
cnf(c_0_1485,plain,(cell47!=cell79), inference(split_conjunct,[status(thm)],[c_0_387])).
cnf(c_0_1486,plain,(cell48!=cell79), inference(split_conjunct,[status(thm)],[c_0_388])).
cnf(c_0_1487,plain,(cell49!=cell79), inference(split_conjunct,[status(thm)],[c_0_389])).
cnf(c_0_1488,plain,(cell50!=cell79), inference(split_conjunct,[status(thm)],[c_0_390])).
cnf(c_0_1489,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred28_2(cell79,esk1_0)|epred27_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_841, c_0_1483]), c_0_1484]), c_0_1485]), c_0_1486]), c_0_1487]), c_0_1488])).
cnf(c_0_1490,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred27_2(cell79,esk1_0)|epred27_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_851, c_0_1489]), c_0_1470]), c_0_1471]), c_0_1472]), c_0_1473]), c_0_1474])).
cnf(c_0_1491,plain,(cell41!=cell79), inference(split_conjunct,[status(thm)],[c_0_391])).
cnf(c_0_1492,plain,(cell42!=cell79), inference(split_conjunct,[status(thm)],[c_0_392])).
cnf(c_0_1493,plain,(cell43!=cell79), inference(split_conjunct,[status(thm)],[c_0_393])).
cnf(c_0_1494,plain,(cell44!=cell79), inference(split_conjunct,[status(thm)],[c_0_394])).
cnf(c_0_1495,plain,(cell45!=cell79), inference(split_conjunct,[status(thm)],[c_0_395])).
cnf(c_0_1496,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred27_2(cell79,esk1_0)|epred26_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_860, c_0_1490]), c_0_1491]), c_0_1492]), c_0_1493]), c_0_1494]), c_0_1495])).
cnf(c_0_1497,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred26_2(cell79,esk1_0)|epred26_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_870, c_0_1496]), c_0_1477]), c_0_1478]), c_0_1479]), c_0_1480]), c_0_1481])).
cnf(c_0_1498,plain,(cell36!=cell79), inference(split_conjunct,[status(thm)],[c_0_396])).
cnf(c_0_1499,plain,(cell37!=cell79), inference(split_conjunct,[status(thm)],[c_0_397])).
cnf(c_0_1500,plain,(cell38!=cell79), inference(split_conjunct,[status(thm)],[c_0_398])).
cnf(c_0_1501,plain,(cell39!=cell79), inference(split_conjunct,[status(thm)],[c_0_399])).
cnf(c_0_1502,plain,(cell40!=cell79), inference(split_conjunct,[status(thm)],[c_0_400])).
cnf(c_0_1503,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred26_2(cell79,esk1_0)|epred25_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_879, c_0_1497]), c_0_1498]), c_0_1499]), c_0_1500]), c_0_1501]), c_0_1502])).
cnf(c_0_1504,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred25_2(cell79,esk1_0)|epred25_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_889, c_0_1503]), c_0_1484]), c_0_1485]), c_0_1486]), c_0_1487]), c_0_1488])).
cnf(c_0_1505,plain,(cell31!=cell79), inference(split_conjunct,[status(thm)],[c_0_401])).
cnf(c_0_1506,plain,(cell32!=cell79), inference(split_conjunct,[status(thm)],[c_0_402])).
cnf(c_0_1507,plain,(cell33!=cell79), inference(split_conjunct,[status(thm)],[c_0_403])).
cnf(c_0_1508,plain,(cell34!=cell79), inference(split_conjunct,[status(thm)],[c_0_404])).
cnf(c_0_1509,plain,(cell35!=cell79), inference(split_conjunct,[status(thm)],[c_0_405])).
cnf(c_0_1510,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred25_2(cell79,esk1_0)|epred24_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_898, c_0_1504]), c_0_1505]), c_0_1506]), c_0_1507]), c_0_1508]), c_0_1509])).
cnf(c_0_1511,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred24_2(cell79,esk1_0)|epred24_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_908, c_0_1510]), c_0_1491]), c_0_1492]), c_0_1493]), c_0_1494]), c_0_1495])).
cnf(c_0_1512,plain,(cell26!=cell79), inference(split_conjunct,[status(thm)],[c_0_406])).
cnf(c_0_1513,plain,(cell27!=cell79), inference(split_conjunct,[status(thm)],[c_0_407])).
cnf(c_0_1514,plain,(cell28!=cell79), inference(split_conjunct,[status(thm)],[c_0_408])).
cnf(c_0_1515,plain,(cell29!=cell79), inference(split_conjunct,[status(thm)],[c_0_409])).
cnf(c_0_1516,plain,(cell30!=cell79), inference(split_conjunct,[status(thm)],[c_0_410])).
cnf(c_0_1517,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred24_2(cell79,esk1_0)|epred23_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_917, c_0_1511]), c_0_1512]), c_0_1513]), c_0_1514]), c_0_1515]), c_0_1516])).
cnf(c_0_1518,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred23_2(cell79,esk1_0)|epred23_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_927, c_0_1517]), c_0_1498]), c_0_1499]), c_0_1500]), c_0_1501]), c_0_1502])).
cnf(c_0_1519,plain,(cell79!=cell22), inference(split_conjunct,[status(thm)],[c_0_411])).
cnf(c_0_1520,plain,(cell21!=cell79), inference(split_conjunct,[status(thm)],[c_0_412])).
cnf(c_0_1521,plain,(cell23!=cell79), inference(split_conjunct,[status(thm)],[c_0_413])).
cnf(c_0_1522,plain,(cell24!=cell79), inference(split_conjunct,[status(thm)],[c_0_414])).
cnf(c_0_1523,plain,(cell25!=cell79), inference(split_conjunct,[status(thm)],[c_0_415])).
cnf(c_0_1524,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred23_2(cell79,esk1_0)|epred22_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_1518]), c_0_1519]), c_0_1520]), c_0_1521]), c_0_1522]), c_0_1523])).
cnf(c_0_1525,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred22_2(cell79,esk1_0)|epred22_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_946, c_0_1524]), c_0_1505]), c_0_1506]), c_0_1507]), c_0_1508]), c_0_1509])).
cnf(c_0_1526,plain,(cell16!=cell79), inference(split_conjunct,[status(thm)],[c_0_416])).
cnf(c_0_1527,plain,(cell17!=cell79), inference(split_conjunct,[status(thm)],[c_0_417])).
cnf(c_0_1528,plain,(cell18!=cell79), inference(split_conjunct,[status(thm)],[c_0_418])).
cnf(c_0_1529,plain,(cell19!=cell79), inference(split_conjunct,[status(thm)],[c_0_419])).
cnf(c_0_1530,plain,(cell20!=cell79), inference(split_conjunct,[status(thm)],[c_0_420])).
cnf(c_0_1531,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred22_2(cell79,esk1_0)|epred21_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_955, c_0_1525]), c_0_1526]), c_0_1527]), c_0_1528]), c_0_1529]), c_0_1530])).
cnf(c_0_1532,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred21_2(cell79,esk1_0)|epred21_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_965, c_0_1531]), c_0_1512]), c_0_1513]), c_0_1514]), c_0_1515]), c_0_1516])).
cnf(c_0_1533,plain,(cell79!=cell12), inference(split_conjunct,[status(thm)],[c_0_421])).
cnf(c_0_1534,plain,(cell79!=cell11), inference(split_conjunct,[status(thm)],[c_0_422])).
cnf(c_0_1535,plain,(cell79!=cell13), inference(split_conjunct,[status(thm)],[c_0_423])).
cnf(c_0_1536,plain,(cell14!=cell79), inference(split_conjunct,[status(thm)],[c_0_424])).
cnf(c_0_1537,plain,(cell15!=cell79), inference(split_conjunct,[status(thm)],[c_0_425])).
cnf(c_0_1538,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred21_2(cell79,esk1_0)|epred20_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_974, c_0_1532]), c_0_1533]), c_0_1534]), c_0_1535]), c_0_1536]), c_0_1537])).
cnf(c_0_1539,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred20_2(cell79,esk1_0)|epred20_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_984, c_0_1538]), c_0_1519]), c_0_1520]), c_0_1521]), c_0_1522]), c_0_1523])).
cnf(c_0_1540,plain,(cell6!=cell79), inference(split_conjunct,[status(thm)],[c_0_426])).
cnf(c_0_1541,plain,(cell7!=cell79), inference(split_conjunct,[status(thm)],[c_0_427])).
cnf(c_0_1542,plain,(cell8!=cell79), inference(split_conjunct,[status(thm)],[c_0_428])).
cnf(c_0_1543,plain,(cell9!=cell79), inference(split_conjunct,[status(thm)],[c_0_429])).
cnf(c_0_1544,plain,(cell10!=cell79), inference(split_conjunct,[status(thm)],[c_0_430])).
cnf(c_0_1545,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred20_2(cell79,esk1_0)|epred19_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_993, c_0_1539]), c_0_1540]), c_0_1541]), c_0_1542]), c_0_1543]), c_0_1544])).
cnf(c_0_1546,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred19_2(cell79,esk1_0)|epred19_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1545]), c_0_1526]), c_0_1527]), c_0_1528]), c_0_1529]), c_0_1530])).
cnf(c_0_1547,plain,(cell1!=cell79), inference(split_conjunct,[status(thm)],[c_0_431])).
cnf(c_0_1548,plain,(cell3!=cell79), inference(split_conjunct,[status(thm)],[c_0_432])).
cnf(c_0_1549,plain,(cell4!=cell79), inference(split_conjunct,[status(thm)],[c_0_433])).
cnf(c_0_1550,plain,(cell5!=cell79), inference(split_conjunct,[status(thm)],[c_0_434])).
cnf(c_0_1551,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred19_2(cell79,esk1_0)|epred18_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1023, c_0_1546]), c_0_723]), c_0_1547]), c_0_1548]), c_0_1549]), c_0_1550])).
cnf(c_0_1552,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred18_2(cell79,esk1_0)|epred18_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1021, c_0_1551]), c_0_1533]), c_0_1534]), c_0_1535]), c_0_1536]), c_0_1537])).
cnf(c_0_1553,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred18_2(cell79,esk1_0)|epred17_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1030, c_0_1552]), c_0_791]), c_0_1440]), c_0_1425]), c_0_1426]), c_0_1427])).
cnf(c_0_1554,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred17_2(cell79,esk1_0)|epred17_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1035, c_0_1553]), c_0_791]), c_0_1425]), c_0_1426]), c_0_1427]), c_0_1428])).
cnf(c_0_1555,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred17_2(cell79,esk1_0)|epred16_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1040, c_0_1554]), c_0_1420]), c_0_1436]), c_0_1437]), c_0_1438]), c_0_1439])).
cnf(c_0_1556,plain,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred16_2(cell79,esk1_0)|epred16_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1046, c_0_1555]), c_0_1263]), c_0_1437]), c_0_1438]), c_0_1439]), c_0_1440])).
cnf(c_0_1557,plain,(esk1_0=cell89|esk1_0=cell69|esk2_0=cell89|esk2_0=cell88|epred16_2(cell79,esk1_0)|epred15_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1057, c_0_1556]), c_0_1433]), c_0_1434]), c_0_1422]), c_0_1423])).
cnf(c_0_1558,plain,(epred14_2(X1,X2)|X2=cell78|X1=cell80|X2=cell81|X2=cell82|X2=cell83|~epred15_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1054])).
cnf(c_0_1559,plain,(esk1_0=cell89|esk1_0=cell69|esk2_0=cell89|esk2_0=cell88|epred15_2(cell79,esk1_0)|epred15_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1055, c_0_1557]), c_0_1420]), c_0_1421]), c_0_1434]), c_0_1422]), c_0_1423])).
cnf(c_0_1560,plain,(epred14_2(X1,X2)|X2=cell78|X1=cell80|X1=cell82|X1=cell83|X1=cell84|~epred15_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1054])).
cnf(c_0_1561,plain,(esk1_0=cell80|esk1_0=cell89|esk1_0=cell69|esk2_0=cell89|esk2_0=cell88|epred15_2(cell79,esk1_0)|epred14_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1558, c_0_1559]), c_0_1444]), c_0_1430]), c_0_1431]), c_0_1432])).
cnf(c_0_1562,plain,(esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk1_0=cell69|esk2_0=cell89|esk2_0=cell88|epred14_2(cell79,esk1_0)|epred14_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1560, c_0_1561]), c_0_1445]), c_0_1431]), c_0_1432]), c_0_1433])).
cnf(c_0_1563,plain,(esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk1_0=cell69|esk2_0=cell89|esk2_0=cell88|epred14_2(cell79,esk1_0)|epred13_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1069, c_0_1562]), c_0_1450]), c_0_1451]), c_0_1452]), c_0_1442]), c_0_1443])).
cnf(c_0_1564,plain,(esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk1_0=cell69|esk2_0=cell89|esk2_0=cell88|epred13_2(cell79,esk1_0)|epred13_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1074, c_0_1563]), c_0_1451]), c_0_1452]), c_0_1442]), c_0_1443]), c_0_1444])).
cnf(c_0_1565,plain,(epred12_2(X1,X2)|X1=cell68|X1=cell69|X2=cell69|X1=cell72|X1=cell73|~epred13_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1073])).
cnf(c_0_1566,plain,(esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk1_0=cell69|esk2_0=cell89|esk2_0=cell88|epred13_2(cell79,esk1_0)|epred12_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1078, c_0_1564]), c_0_1456]), c_0_1457]), c_0_1458]), c_0_1448]), c_0_1449])).
cnf(c_0_1567,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred12_2(cell79,esk1_0)|epred12_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1565, c_0_1566]), c_0_1457]), c_0_1458]), c_0_1449]), c_0_1450])).
cnf(c_0_1568,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred12_2(cell79,esk1_0)|epred11_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1087, c_0_1567]), c_0_1464]), c_0_1465]), c_0_1466]), c_0_1467]), c_0_1455])).
cnf(c_0_1569,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred11_2(cell79,esk1_0)|epred11_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1092, c_0_1568]), c_0_1465]), c_0_1466]), c_0_1467]), c_0_1455]), c_0_1456])).
cnf(c_0_1570,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred11_2(cell79,esk1_0)|epred10_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1096, c_0_1569]), c_0_1470]), c_0_1471]), c_0_1472]), c_0_1473]), c_0_1463])).
cnf(c_0_1571,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred10_2(cell79,esk1_0)|epred10_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1101, c_0_1570]), c_0_1471]), c_0_1472]), c_0_1473]), c_0_1474]), c_0_1464])).
cnf(c_0_1572,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred10_2(cell79,esk1_0)|epred9_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1105, c_0_1571]), c_0_1477]), c_0_1478]), c_0_1479]), c_0_1480]), c_0_1481])).
cnf(c_0_1573,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred9_2(cell79,esk1_0)|epred9_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1110, c_0_1572]), c_0_1478]), c_0_1479]), c_0_1480]), c_0_1481]), c_0_1470])).
cnf(c_0_1574,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred9_2(cell79,esk1_0)|epred8_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1114, c_0_1573]), c_0_1495]), c_0_1484]), c_0_1485]), c_0_1486]), c_0_1487])).
cnf(c_0_1575,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred8_2(cell79,esk1_0)|epred8_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1119, c_0_1574]), c_0_1484]), c_0_1485]), c_0_1486]), c_0_1487]), c_0_1488])).
cnf(c_0_1576,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred8_2(cell79,esk1_0)|epred7_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1123, c_0_1575]), c_0_1501]), c_0_1491]), c_0_1492]), c_0_1493]), c_0_1494])).
cnf(c_0_1577,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred7_2(cell79,esk1_0)|epred7_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1128, c_0_1576]), c_0_1502]), c_0_1492]), c_0_1493]), c_0_1494]), c_0_1495])).
cnf(c_0_1578,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred7_2(cell79,esk1_0)|epred6_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1132, c_0_1577]), c_0_1508]), c_0_1509]), c_0_1498]), c_0_1499]), c_0_1500])).
cnf(c_0_1579,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred6_2(cell79,esk1_0)|epred6_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1137, c_0_1578]), c_0_1509]), c_0_1498]), c_0_1499]), c_0_1500]), c_0_1501])).
cnf(c_0_1580,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred6_2(cell79,esk1_0)|epred5_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1141, c_0_1579]), c_0_1514]), c_0_1515]), c_0_1505]), c_0_1506]), c_0_1507])).
cnf(c_0_1581,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred5_2(cell79,esk1_0)|epred5_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1146, c_0_1580]), c_0_1515]), c_0_1516]), c_0_1506]), c_0_1507]), c_0_1508])).
cnf(c_0_1582,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred5_2(cell79,esk1_0)|epred4_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1150, c_0_1581]), c_0_1521]), c_0_1522]), c_0_1523]), c_0_1512]), c_0_1513])).
cnf(c_0_1583,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred4_2(cell79,esk1_0)|epred4_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1155, c_0_1582]), c_0_1522]), c_0_1523]), c_0_1512]), c_0_1513]), c_0_1514])).
cnf(c_0_1584,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred4_2(cell79,esk1_0)|epred3_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1160, c_0_1583]), c_0_1519]), c_0_1527]), c_0_1528]), c_0_1529]), c_0_1520])).
cnf(c_0_1585,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred3_2(cell79,esk1_0)|epred3_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1165, c_0_1584]), c_0_1519]), c_0_1528]), c_0_1529]), c_0_1530]), c_0_1521])).
cnf(c_0_1586,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred3_2(cell79,esk1_0)|epred2_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1170, c_0_1585]), c_0_1533]), c_0_1535]), c_0_1536]), c_0_1537]), c_0_1526])).
cnf(c_0_1587,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred2_2(cell79,esk1_0)|epred2_2(esk1_0,cell79)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1175, c_0_1586]), c_0_1535]), c_0_1536]), c_0_1537]), c_0_1526]), c_0_1527])).
cnf(c_0_1588,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88|epred2_2(cell79,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1408, c_0_1587]), c_0_723]), c_0_1547]), c_0_1548]), c_0_1549]), c_0_1550]), c_0_1534]), c_0_1540]), c_0_1541]), c_0_1542]), c_0_1543])).
cnf(c_0_1589,plain,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell80|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1410, c_0_1588]), c_0_723]), c_0_1548]), c_0_1549]), c_0_1550]), c_0_1540]), c_0_1533]), c_0_1541]), c_0_1542]), c_0_1543]), c_0_1544])).
cnf(c_0_1590,negated_conjecture,(esk1_0=cell69|esk1_0=cell78|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1589]), c_0_1414]), c_0_1415]), c_0_1416]), c_0_1417]), c_0_727])).
cnf(c_0_1591,plain,(cell78!=cell12), inference(split_conjunct,[status(thm)],[c_0_435])).
cnf(c_0_1592,plain,(cell78!=cell13), inference(split_conjunct,[status(thm)],[c_0_436])).
cnf(c_0_1593,plain,(cell78!=cell22), inference(split_conjunct,[status(thm)],[c_0_437])).
cnf(c_0_1594,plain,(cell78!=cell11), inference(split_conjunct,[status(thm)],[c_0_438])).
cnf(c_0_1595,negated_conjecture,(esk1_0=cell69|esk1_0=cell89|esk2_0=cell89|esk2_0=cell88), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1590]), c_0_1591]), c_0_1592]), c_0_1593]), c_0_1594]), c_0_726])).
cnf(c_0_1596,plain,(cell69!=cell12), inference(split_conjunct,[status(thm)],[c_0_439])).
cnf(c_0_1597,plain,(cell69!=cell13), inference(split_conjunct,[status(thm)],[c_0_440])).
cnf(c_0_1598,plain,(cell69!=cell22), inference(split_conjunct,[status(thm)],[c_0_441])).
cnf(c_0_1599,plain,(cell69!=cell11), inference(split_conjunct,[status(thm)],[c_0_442])).
cnf(c_0_1600,plain,(epred35_2(X1,X2)|X2=cell86|X2=cell87|X1=cell98|X2=cell89|X2=cell90|~epred36_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_657])).
cnf(c_0_1601,negated_conjecture,(esk1_0=cell89|esk2_0=cell89|esk2_0=cell88), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1595]), c_0_1596]), c_0_1597]), c_0_1598]), c_0_1599]), c_0_769])).
cnf(c_0_1602,plain,(esk2_0=cell89|esk2_0=cell90|esk2_0=cell86|esk2_0=cell87|esk1_0=cell98|epred36_2(esk2_0,esk1_0)|epred35_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_1600, c_0_663])).
cnf(c_0_1603,negated_conjecture,(esk2_0=cell89|esk2_0=cell88), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1601]), c_0_1224]), c_0_1220]), c_0_1205]), c_0_1190]), c_0_679])).
cnf(c_0_1604,plain,(cell89!=cell88), inference(split_conjunct,[status(thm)],[c_0_443])).
cnf(c_0_1605,plain,(cell86!=cell88), inference(split_conjunct,[status(thm)],[c_0_444])).
cnf(c_0_1606,plain,(cell87!=cell88), inference(split_conjunct,[status(thm)],[c_0_445])).
cnf(c_0_1607,plain,(esk1_0=cell98|esk2_0=cell89|epred36_2(cell88,esk1_0)|epred35_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1602, c_0_1603]), c_0_1604]), c_0_1242]), c_0_1605]), c_0_1606])).
cnf(c_0_1608,plain,(cell96!=cell88), inference(split_conjunct,[status(thm)],[c_0_446])).
cnf(c_0_1609,plain,(cell97!=cell88), inference(split_conjunct,[status(thm)],[c_0_447])).
cnf(c_0_1610,plain,(cell98!=cell88), inference(split_conjunct,[status(thm)],[c_0_448])).
cnf(c_0_1611,plain,(cell100!=cell88), inference(split_conjunct,[status(thm)],[c_0_449])).
cnf(c_0_1612,plain,(esk1_0=cell98|esk2_0=cell89|epred35_2(cell88,esk1_0)|epred35_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_662, c_0_1607]), c_0_687]), c_0_1608]), c_0_1609]), c_0_1610]), c_0_1611])).
cnf(c_0_1613,plain,(cell81!=cell88), inference(split_conjunct,[status(thm)],[c_0_450])).
cnf(c_0_1614,plain,(cell82!=cell88), inference(split_conjunct,[status(thm)],[c_0_451])).
cnf(c_0_1615,plain,(cell83!=cell88), inference(split_conjunct,[status(thm)],[c_0_452])).
cnf(c_0_1616,plain,(cell84!=cell88), inference(split_conjunct,[status(thm)],[c_0_453])).
cnf(c_0_1617,plain,(cell85!=cell88), inference(split_conjunct,[status(thm)],[c_0_454])).
cnf(c_0_1618,plain,(esk1_0=cell98|esk2_0=cell89|epred35_2(cell88,esk1_0)|epred34_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_1612]), c_0_1613]), c_0_1614]), c_0_1615]), c_0_1616]), c_0_1617])).
cnf(c_0_1619,plain,(cell91!=cell88), inference(split_conjunct,[status(thm)],[c_0_455])).
cnf(c_0_1620,plain,(cell92!=cell88), inference(split_conjunct,[status(thm)],[c_0_456])).
cnf(c_0_1621,plain,(cell93!=cell88), inference(split_conjunct,[status(thm)],[c_0_457])).
cnf(c_0_1622,plain,(cell94!=cell88), inference(split_conjunct,[status(thm)],[c_0_458])).
cnf(c_0_1623,plain,(cell95!=cell88), inference(split_conjunct,[status(thm)],[c_0_459])).
cnf(c_0_1624,plain,(esk1_0=cell98|esk2_0=cell89|epred34_2(cell88,esk1_0)|epred34_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_1618]), c_0_1619]), c_0_1620]), c_0_1621]), c_0_1622]), c_0_1623])).
cnf(c_0_1625,plain,(cell76!=cell88), inference(split_conjunct,[status(thm)],[c_0_460])).
cnf(c_0_1626,plain,(cell77!=cell88), inference(split_conjunct,[status(thm)],[c_0_461])).
cnf(c_0_1627,plain,(cell78!=cell88), inference(split_conjunct,[status(thm)],[c_0_462])).
cnf(c_0_1628,plain,(cell80!=cell88), inference(split_conjunct,[status(thm)],[c_0_463])).
cnf(c_0_1629,plain,(epred33_2(X1,X2)|X1=cell86|X1=cell87|X2=cell78|X1=cell89|X1=cell90|~epred34_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_706])).
cnf(c_0_1630,plain,(esk1_0=cell98|esk2_0=cell89|epred34_2(cell88,esk1_0)|epred33_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_721, c_0_1624]), c_0_1421]), c_0_1625]), c_0_1626]), c_0_1627]), c_0_1628])).
cnf(c_0_1631,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred33_2(cell88,esk1_0)|epred33_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1629, c_0_1630]), c_0_1604]), c_0_1242]), c_0_1605]), c_0_1606])).
cnf(c_0_1632,plain,(cell71!=cell88), inference(split_conjunct,[status(thm)],[c_0_464])).
cnf(c_0_1633,plain,(cell72!=cell88), inference(split_conjunct,[status(thm)],[c_0_465])).
cnf(c_0_1634,plain,(cell73!=cell88), inference(split_conjunct,[status(thm)],[c_0_466])).
cnf(c_0_1635,plain,(cell74!=cell88), inference(split_conjunct,[status(thm)],[c_0_467])).
cnf(c_0_1636,plain,(cell75!=cell88), inference(split_conjunct,[status(thm)],[c_0_468])).
cnf(c_0_1637,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred33_2(cell88,esk1_0)|epred32_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_745, c_0_1631]), c_0_1632]), c_0_1633]), c_0_1634]), c_0_1635]), c_0_1636])).
cnf(c_0_1638,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred32_2(cell88,esk1_0)|epred32_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_755, c_0_1637]), c_0_1613]), c_0_1614]), c_0_1615]), c_0_1616]), c_0_1617])).
cnf(c_0_1639,plain,(cell66!=cell88), inference(split_conjunct,[status(thm)],[c_0_469])).
cnf(c_0_1640,plain,(cell67!=cell88), inference(split_conjunct,[status(thm)],[c_0_470])).
cnf(c_0_1641,plain,(cell68!=cell88), inference(split_conjunct,[status(thm)],[c_0_471])).
cnf(c_0_1642,plain,(cell69!=cell88), inference(split_conjunct,[status(thm)],[c_0_472])).
cnf(c_0_1643,plain,(cell70!=cell88), inference(split_conjunct,[status(thm)],[c_0_473])).
cnf(c_0_1644,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred32_2(cell88,esk1_0)|epred31_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1638]), c_0_1639]), c_0_1640]), c_0_1641]), c_0_1642]), c_0_1643])).
cnf(c_0_1645,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred31_2(cell88,esk1_0)|epred31_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_774, c_0_1644]), c_0_1421]), c_0_1625]), c_0_1626]), c_0_1627]), c_0_1628])).
cnf(c_0_1646,plain,(cell61!=cell88), inference(split_conjunct,[status(thm)],[c_0_474])).
cnf(c_0_1647,plain,(cell62!=cell88), inference(split_conjunct,[status(thm)],[c_0_475])).
cnf(c_0_1648,plain,(cell63!=cell88), inference(split_conjunct,[status(thm)],[c_0_476])).
cnf(c_0_1649,plain,(cell64!=cell88), inference(split_conjunct,[status(thm)],[c_0_477])).
cnf(c_0_1650,plain,(cell65!=cell88), inference(split_conjunct,[status(thm)],[c_0_478])).
cnf(c_0_1651,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred31_2(cell88,esk1_0)|epred30_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_783, c_0_1645]), c_0_1646]), c_0_1647]), c_0_1648]), c_0_1649]), c_0_1650])).
cnf(c_0_1652,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred30_2(cell88,esk1_0)|epred30_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_1651]), c_0_1632]), c_0_1633]), c_0_1634]), c_0_1635]), c_0_1636])).
cnf(c_0_1653,plain,(cell56!=cell88), inference(split_conjunct,[status(thm)],[c_0_479])).
cnf(c_0_1654,plain,(cell57!=cell88), inference(split_conjunct,[status(thm)],[c_0_480])).
cnf(c_0_1655,plain,(cell58!=cell88), inference(split_conjunct,[status(thm)],[c_0_481])).
cnf(c_0_1656,plain,(cell59!=cell88), inference(split_conjunct,[status(thm)],[c_0_482])).
cnf(c_0_1657,plain,(cell60!=cell88), inference(split_conjunct,[status(thm)],[c_0_483])).
cnf(c_0_1658,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred30_2(cell88,esk1_0)|epred29_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_803, c_0_1652]), c_0_1653]), c_0_1654]), c_0_1655]), c_0_1656]), c_0_1657])).
cnf(c_0_1659,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred29_2(cell88,esk1_0)|epred29_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_813, c_0_1658]), c_0_1639]), c_0_1640]), c_0_1641]), c_0_1642]), c_0_1643])).
cnf(c_0_1660,plain,(cell51!=cell88), inference(split_conjunct,[status(thm)],[c_0_484])).
cnf(c_0_1661,plain,(cell52!=cell88), inference(split_conjunct,[status(thm)],[c_0_485])).
cnf(c_0_1662,plain,(cell53!=cell88), inference(split_conjunct,[status(thm)],[c_0_486])).
cnf(c_0_1663,plain,(cell54!=cell88), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_1664,plain,(cell55!=cell88), inference(split_conjunct,[status(thm)],[c_0_488])).
cnf(c_0_1665,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred29_2(cell88,esk1_0)|epred28_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_822, c_0_1659]), c_0_1660]), c_0_1661]), c_0_1662]), c_0_1663]), c_0_1664])).
cnf(c_0_1666,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred28_2(cell88,esk1_0)|epred28_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_832, c_0_1665]), c_0_1646]), c_0_1647]), c_0_1648]), c_0_1649]), c_0_1650])).
cnf(c_0_1667,plain,(cell46!=cell88), inference(split_conjunct,[status(thm)],[c_0_489])).
cnf(c_0_1668,plain,(cell47!=cell88), inference(split_conjunct,[status(thm)],[c_0_490])).
cnf(c_0_1669,plain,(cell48!=cell88), inference(split_conjunct,[status(thm)],[c_0_491])).
cnf(c_0_1670,plain,(cell49!=cell88), inference(split_conjunct,[status(thm)],[c_0_492])).
cnf(c_0_1671,plain,(cell50!=cell88), inference(split_conjunct,[status(thm)],[c_0_493])).
cnf(c_0_1672,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred28_2(cell88,esk1_0)|epred27_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_841, c_0_1666]), c_0_1667]), c_0_1668]), c_0_1669]), c_0_1670]), c_0_1671])).
cnf(c_0_1673,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred27_2(cell88,esk1_0)|epred27_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_851, c_0_1672]), c_0_1653]), c_0_1654]), c_0_1655]), c_0_1656]), c_0_1657])).
cnf(c_0_1674,plain,(cell41!=cell88), inference(split_conjunct,[status(thm)],[c_0_494])).
cnf(c_0_1675,plain,(cell42!=cell88), inference(split_conjunct,[status(thm)],[c_0_495])).
cnf(c_0_1676,plain,(cell43!=cell88), inference(split_conjunct,[status(thm)],[c_0_496])).
cnf(c_0_1677,plain,(cell44!=cell88), inference(split_conjunct,[status(thm)],[c_0_497])).
cnf(c_0_1678,plain,(cell45!=cell88), inference(split_conjunct,[status(thm)],[c_0_498])).
cnf(c_0_1679,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred27_2(cell88,esk1_0)|epred26_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_860, c_0_1673]), c_0_1674]), c_0_1675]), c_0_1676]), c_0_1677]), c_0_1678])).
cnf(c_0_1680,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred26_2(cell88,esk1_0)|epred26_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_870, c_0_1679]), c_0_1660]), c_0_1661]), c_0_1662]), c_0_1663]), c_0_1664])).
cnf(c_0_1681,plain,(cell36!=cell88), inference(split_conjunct,[status(thm)],[c_0_499])).
cnf(c_0_1682,plain,(cell37!=cell88), inference(split_conjunct,[status(thm)],[c_0_500])).
cnf(c_0_1683,plain,(cell38!=cell88), inference(split_conjunct,[status(thm)],[c_0_501])).
cnf(c_0_1684,plain,(cell39!=cell88), inference(split_conjunct,[status(thm)],[c_0_502])).
cnf(c_0_1685,plain,(cell40!=cell88), inference(split_conjunct,[status(thm)],[c_0_503])).
cnf(c_0_1686,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred26_2(cell88,esk1_0)|epred25_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_879, c_0_1680]), c_0_1681]), c_0_1682]), c_0_1683]), c_0_1684]), c_0_1685])).
cnf(c_0_1687,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred25_2(cell88,esk1_0)|epred25_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_889, c_0_1686]), c_0_1667]), c_0_1668]), c_0_1669]), c_0_1670]), c_0_1671])).
cnf(c_0_1688,plain,(cell31!=cell88), inference(split_conjunct,[status(thm)],[c_0_504])).
cnf(c_0_1689,plain,(cell32!=cell88), inference(split_conjunct,[status(thm)],[c_0_505])).
cnf(c_0_1690,plain,(cell33!=cell88), inference(split_conjunct,[status(thm)],[c_0_506])).
cnf(c_0_1691,plain,(cell34!=cell88), inference(split_conjunct,[status(thm)],[c_0_507])).
cnf(c_0_1692,plain,(cell35!=cell88), inference(split_conjunct,[status(thm)],[c_0_508])).
cnf(c_0_1693,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred25_2(cell88,esk1_0)|epred24_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_898, c_0_1687]), c_0_1688]), c_0_1689]), c_0_1690]), c_0_1691]), c_0_1692])).
cnf(c_0_1694,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred24_2(cell88,esk1_0)|epred24_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_908, c_0_1693]), c_0_1674]), c_0_1675]), c_0_1676]), c_0_1677]), c_0_1678])).
cnf(c_0_1695,plain,(cell26!=cell88), inference(split_conjunct,[status(thm)],[c_0_509])).
cnf(c_0_1696,plain,(cell27!=cell88), inference(split_conjunct,[status(thm)],[c_0_510])).
cnf(c_0_1697,plain,(cell28!=cell88), inference(split_conjunct,[status(thm)],[c_0_511])).
cnf(c_0_1698,plain,(cell29!=cell88), inference(split_conjunct,[status(thm)],[c_0_512])).
cnf(c_0_1699,plain,(cell30!=cell88), inference(split_conjunct,[status(thm)],[c_0_513])).
cnf(c_0_1700,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred24_2(cell88,esk1_0)|epred23_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_917, c_0_1694]), c_0_1695]), c_0_1696]), c_0_1697]), c_0_1698]), c_0_1699])).
cnf(c_0_1701,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred23_2(cell88,esk1_0)|epred23_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_927, c_0_1700]), c_0_1681]), c_0_1682]), c_0_1683]), c_0_1684]), c_0_1685])).
cnf(c_0_1702,plain,(cell88!=cell22), inference(split_conjunct,[status(thm)],[c_0_514])).
cnf(c_0_1703,plain,(cell21!=cell88), inference(split_conjunct,[status(thm)],[c_0_515])).
cnf(c_0_1704,plain,(cell23!=cell88), inference(split_conjunct,[status(thm)],[c_0_516])).
cnf(c_0_1705,plain,(cell24!=cell88), inference(split_conjunct,[status(thm)],[c_0_517])).
cnf(c_0_1706,plain,(cell25!=cell88), inference(split_conjunct,[status(thm)],[c_0_518])).
cnf(c_0_1707,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred23_2(cell88,esk1_0)|epred22_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_1701]), c_0_1702]), c_0_1703]), c_0_1704]), c_0_1705]), c_0_1706])).
cnf(c_0_1708,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred22_2(cell88,esk1_0)|epred22_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_946, c_0_1707]), c_0_1688]), c_0_1689]), c_0_1690]), c_0_1691]), c_0_1692])).
cnf(c_0_1709,plain,(cell16!=cell88), inference(split_conjunct,[status(thm)],[c_0_519])).
cnf(c_0_1710,plain,(cell17!=cell88), inference(split_conjunct,[status(thm)],[c_0_520])).
cnf(c_0_1711,plain,(cell18!=cell88), inference(split_conjunct,[status(thm)],[c_0_521])).
cnf(c_0_1712,plain,(cell19!=cell88), inference(split_conjunct,[status(thm)],[c_0_522])).
cnf(c_0_1713,plain,(cell20!=cell88), inference(split_conjunct,[status(thm)],[c_0_523])).
cnf(c_0_1714,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred22_2(cell88,esk1_0)|epred21_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_955, c_0_1708]), c_0_1709]), c_0_1710]), c_0_1711]), c_0_1712]), c_0_1713])).
cnf(c_0_1715,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred21_2(cell88,esk1_0)|epred21_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_965, c_0_1714]), c_0_1695]), c_0_1696]), c_0_1697]), c_0_1698]), c_0_1699])).
cnf(c_0_1716,plain,(cell88!=cell12), inference(split_conjunct,[status(thm)],[c_0_524])).
cnf(c_0_1717,plain,(cell88!=cell11), inference(split_conjunct,[status(thm)],[c_0_525])).
cnf(c_0_1718,plain,(cell88!=cell13), inference(split_conjunct,[status(thm)],[c_0_526])).
cnf(c_0_1719,plain,(cell14!=cell88), inference(split_conjunct,[status(thm)],[c_0_527])).
cnf(c_0_1720,plain,(cell15!=cell88), inference(split_conjunct,[status(thm)],[c_0_528])).
cnf(c_0_1721,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred21_2(cell88,esk1_0)|epred20_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_974, c_0_1715]), c_0_1716]), c_0_1717]), c_0_1718]), c_0_1719]), c_0_1720])).
cnf(c_0_1722,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred20_2(cell88,esk1_0)|epred20_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_984, c_0_1721]), c_0_1702]), c_0_1703]), c_0_1704]), c_0_1705]), c_0_1706])).
cnf(c_0_1723,plain,(cell6!=cell88), inference(split_conjunct,[status(thm)],[c_0_529])).
cnf(c_0_1724,plain,(cell7!=cell88), inference(split_conjunct,[status(thm)],[c_0_530])).
cnf(c_0_1725,plain,(cell8!=cell88), inference(split_conjunct,[status(thm)],[c_0_531])).
cnf(c_0_1726,plain,(cell9!=cell88), inference(split_conjunct,[status(thm)],[c_0_532])).
cnf(c_0_1727,plain,(cell10!=cell88), inference(split_conjunct,[status(thm)],[c_0_533])).
cnf(c_0_1728,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred20_2(cell88,esk1_0)|epred19_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_993, c_0_1722]), c_0_1723]), c_0_1724]), c_0_1725]), c_0_1726]), c_0_1727])).
cnf(c_0_1729,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred19_2(cell88,esk1_0)|epred19_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1728]), c_0_1709]), c_0_1710]), c_0_1711]), c_0_1712]), c_0_1713])).
cnf(c_0_1730,plain,(cell1!=cell88), inference(split_conjunct,[status(thm)],[c_0_534])).
cnf(c_0_1731,plain,(cell3!=cell88), inference(split_conjunct,[status(thm)],[c_0_535])).
cnf(c_0_1732,plain,(cell4!=cell88), inference(split_conjunct,[status(thm)],[c_0_536])).
cnf(c_0_1733,plain,(cell5!=cell88), inference(split_conjunct,[status(thm)],[c_0_537])).
cnf(c_0_1734,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred19_2(cell88,esk1_0)|epred18_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1023, c_0_1729]), c_0_680]), c_0_1730]), c_0_1731]), c_0_1732]), c_0_1733])).
cnf(c_0_1735,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred18_2(cell88,esk1_0)|epred18_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1021, c_0_1734]), c_0_1716]), c_0_1717]), c_0_1718]), c_0_1719]), c_0_1720])).
cnf(c_0_1736,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred18_2(cell88,esk1_0)|epred17_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1030, c_0_1735]), c_0_687]), c_0_1623]), c_0_1608]), c_0_1609]), c_0_1610])).
cnf(c_0_1737,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred17_2(cell88,esk1_0)|epred17_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1035, c_0_1736]), c_0_687]), c_0_1608]), c_0_1609]), c_0_1610]), c_0_1611])).
cnf(c_0_1738,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred17_2(cell88,esk1_0)|epred16_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1040, c_0_1737]), c_0_1604]), c_0_1619]), c_0_1620]), c_0_1621]), c_0_1622])).
cnf(c_0_1739,plain,(esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred16_2(cell88,esk1_0)|epred16_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1046, c_0_1738]), c_0_1242]), c_0_1620]), c_0_1621]), c_0_1622]), c_0_1623])).
cnf(c_0_1740,plain,(epred15_2(X1,X2)|X1=cell85|X1=cell86|X1=cell87|X2=cell87|X1=cell89|~epred16_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1045])).
cnf(c_0_1741,plain,(esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred16_2(cell88,esk1_0)|epred15_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1057, c_0_1739]), c_0_1616]), c_0_1617]), c_0_1605]), c_0_1606])).
cnf(c_0_1742,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred15_2(cell88,esk1_0)|epred15_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1740, c_0_1741]), c_0_1604]), c_0_1617]), c_0_1605]), c_0_1606])).
cnf(c_0_1743,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred15_2(cell88,esk1_0)|epred14_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1060, c_0_1742]), c_0_1421]), c_0_1627]), c_0_1613]), c_0_1614]), c_0_1615])).
cnf(c_0_1744,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred14_2(cell88,esk1_0)|epred14_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1065, c_0_1743]), c_0_1421]), c_0_1628]), c_0_1614]), c_0_1615]), c_0_1616])).
cnf(c_0_1745,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred14_2(cell88,esk1_0)|epred13_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1069, c_0_1744]), c_0_1634]), c_0_1635]), c_0_1636]), c_0_1625]), c_0_1626])).
cnf(c_0_1746,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred13_2(cell88,esk1_0)|epred13_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1074, c_0_1745]), c_0_1635]), c_0_1636]), c_0_1625]), c_0_1626]), c_0_1627])).
cnf(c_0_1747,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred13_2(cell88,esk1_0)|epred12_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1078, c_0_1746]), c_0_1640]), c_0_1641]), c_0_1642]), c_0_1632]), c_0_1633])).
cnf(c_0_1748,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred12_2(cell88,esk1_0)|epred12_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1083, c_0_1747]), c_0_1641]), c_0_1642]), c_0_1643]), c_0_1633]), c_0_1634])).
cnf(c_0_1749,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred12_2(cell88,esk1_0)|epred11_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1087, c_0_1748]), c_0_1647]), c_0_1648]), c_0_1649]), c_0_1650]), c_0_1639])).
cnf(c_0_1750,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred11_2(cell88,esk1_0)|epred11_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1092, c_0_1749]), c_0_1648]), c_0_1649]), c_0_1650]), c_0_1639]), c_0_1640])).
cnf(c_0_1751,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred11_2(cell88,esk1_0)|epred10_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1096, c_0_1750]), c_0_1653]), c_0_1654]), c_0_1655]), c_0_1656]), c_0_1646])).
cnf(c_0_1752,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred10_2(cell88,esk1_0)|epred10_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1101, c_0_1751]), c_0_1654]), c_0_1655]), c_0_1656]), c_0_1657]), c_0_1647])).
cnf(c_0_1753,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred10_2(cell88,esk1_0)|epred9_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1105, c_0_1752]), c_0_1660]), c_0_1661]), c_0_1662]), c_0_1663]), c_0_1664])).
cnf(c_0_1754,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred9_2(cell88,esk1_0)|epred9_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1110, c_0_1753]), c_0_1661]), c_0_1662]), c_0_1663]), c_0_1664]), c_0_1653])).
cnf(c_0_1755,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred9_2(cell88,esk1_0)|epred8_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1114, c_0_1754]), c_0_1678]), c_0_1667]), c_0_1668]), c_0_1669]), c_0_1670])).
cnf(c_0_1756,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred8_2(cell88,esk1_0)|epred8_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1119, c_0_1755]), c_0_1667]), c_0_1668]), c_0_1669]), c_0_1670]), c_0_1671])).
cnf(c_0_1757,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred8_2(cell88,esk1_0)|epred7_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1123, c_0_1756]), c_0_1684]), c_0_1674]), c_0_1675]), c_0_1676]), c_0_1677])).
cnf(c_0_1758,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred7_2(cell88,esk1_0)|epred7_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1128, c_0_1757]), c_0_1685]), c_0_1675]), c_0_1676]), c_0_1677]), c_0_1678])).
cnf(c_0_1759,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred7_2(cell88,esk1_0)|epred6_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1132, c_0_1758]), c_0_1691]), c_0_1692]), c_0_1681]), c_0_1682]), c_0_1683])).
cnf(c_0_1760,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred6_2(cell88,esk1_0)|epred6_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1137, c_0_1759]), c_0_1692]), c_0_1681]), c_0_1682]), c_0_1683]), c_0_1684])).
cnf(c_0_1761,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred6_2(cell88,esk1_0)|epred5_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1141, c_0_1760]), c_0_1697]), c_0_1698]), c_0_1688]), c_0_1689]), c_0_1690])).
cnf(c_0_1762,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred5_2(cell88,esk1_0)|epred5_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1146, c_0_1761]), c_0_1698]), c_0_1699]), c_0_1689]), c_0_1690]), c_0_1691])).
cnf(c_0_1763,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred5_2(cell88,esk1_0)|epred4_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1150, c_0_1762]), c_0_1704]), c_0_1705]), c_0_1706]), c_0_1695]), c_0_1696])).
cnf(c_0_1764,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred4_2(cell88,esk1_0)|epred4_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1155, c_0_1763]), c_0_1705]), c_0_1706]), c_0_1695]), c_0_1696]), c_0_1697])).
cnf(c_0_1765,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred4_2(cell88,esk1_0)|epred3_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1160, c_0_1764]), c_0_1702]), c_0_1710]), c_0_1711]), c_0_1712]), c_0_1703])).
cnf(c_0_1766,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred3_2(cell88,esk1_0)|epred3_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1165, c_0_1765]), c_0_1702]), c_0_1711]), c_0_1712]), c_0_1713]), c_0_1704])).
cnf(c_0_1767,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred3_2(cell88,esk1_0)|epred2_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1170, c_0_1766]), c_0_1716]), c_0_1718]), c_0_1719]), c_0_1720]), c_0_1709])).
cnf(c_0_1768,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred2_2(cell88,esk1_0)|epred2_2(esk1_0,cell88)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1175, c_0_1767]), c_0_1718]), c_0_1719]), c_0_1720]), c_0_1709]), c_0_1710])).
cnf(c_0_1769,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89|epred2_2(cell88,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1408, c_0_1768]), c_0_680]), c_0_1730]), c_0_1731]), c_0_1732]), c_0_1733]), c_0_1717]), c_0_1723]), c_0_1724]), c_0_1725]), c_0_1726])).
cnf(c_0_1770,plain,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk1_0=cell98|esk2_0=cell89), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1410, c_0_1769]), c_0_680]), c_0_1731]), c_0_1732]), c_0_1733]), c_0_1723]), c_0_1716]), c_0_1724]), c_0_1725]), c_0_1726]), c_0_1727])).
cnf(c_0_1771,negated_conjecture,(esk1_0=cell87|esk1_0=cell89|esk1_0=cell78|esk2_0=cell89), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1770]), c_0_1225]), c_0_1221]), c_0_1206]), c_0_1191]), c_0_672])).
cnf(c_0_1772,plain,(cell87!=cell12), inference(split_conjunct,[status(thm)],[c_0_538])).
cnf(c_0_1773,plain,(cell87!=cell13), inference(split_conjunct,[status(thm)],[c_0_539])).
cnf(c_0_1774,plain,(cell87!=cell22), inference(split_conjunct,[status(thm)],[c_0_540])).
cnf(c_0_1775,plain,(cell87!=cell11), inference(split_conjunct,[status(thm)],[c_0_541])).
cnf(c_0_1776,negated_conjecture,(esk1_0=cell89|esk1_0=cell78|esk2_0=cell89), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1771]), c_0_1772]), c_0_1773]), c_0_1774]), c_0_1775]), c_0_683])).
cnf(c_0_1777,plain,(epred35_2(X1,X2)|X2=cell86|X2=cell87|X2=cell88|X1=cell99|X2=cell90|~epred36_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_657])).
cnf(c_0_1778,negated_conjecture,(esk1_0=cell89|esk2_0=cell89), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1776]), c_0_1591]), c_0_1592]), c_0_1593]), c_0_1594]), c_0_726])).
cnf(c_0_1779,plain,(esk2_0=cell88|esk2_0=cell90|esk1_0=cell99|esk2_0=cell86|esk2_0=cell87|epred36_2(esk2_0,esk1_0)|epred35_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_1777, c_0_663])).
cnf(c_0_1780,negated_conjecture,(esk2_0=cell89), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_1778]), c_0_1224]), c_0_1220]), c_0_1205]), c_0_1190]), c_0_679])).
cnf(c_0_1781,plain,(cell86!=cell89), inference(split_conjunct,[status(thm)],[c_0_542])).
cnf(c_0_1782,plain,(cell87!=cell89), inference(split_conjunct,[status(thm)],[c_0_543])).
cnf(c_0_1783,plain,(esk1_0=cell99|epred36_2(cell89,esk1_0)|epred35_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_1779, c_0_1780]), c_0_1780]), c_0_1780]), c_0_1780]), c_0_1780]), c_0_1780]), c_0_1604]), c_0_1241]), c_0_1781]), c_0_1782])).
cnf(c_0_1784,plain,(cell96!=cell89), inference(split_conjunct,[status(thm)],[c_0_544])).
cnf(c_0_1785,plain,(cell97!=cell89), inference(split_conjunct,[status(thm)],[c_0_545])).
cnf(c_0_1786,plain,(cell98!=cell89), inference(split_conjunct,[status(thm)],[c_0_546])).
cnf(c_0_1787,plain,(cell100!=cell89), inference(split_conjunct,[status(thm)],[c_0_547])).
cnf(c_0_1788,plain,(esk1_0=cell99|epred35_2(cell89,esk1_0)|epred35_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_662, c_0_1783]), c_0_686]), c_0_1784]), c_0_1785]), c_0_1786]), c_0_1787])).
cnf(c_0_1789,plain,(cell81!=cell89), inference(split_conjunct,[status(thm)],[c_0_548])).
cnf(c_0_1790,plain,(cell82!=cell89), inference(split_conjunct,[status(thm)],[c_0_549])).
cnf(c_0_1791,plain,(cell83!=cell89), inference(split_conjunct,[status(thm)],[c_0_550])).
cnf(c_0_1792,plain,(cell84!=cell89), inference(split_conjunct,[status(thm)],[c_0_551])).
cnf(c_0_1793,plain,(cell85!=cell89), inference(split_conjunct,[status(thm)],[c_0_552])).
cnf(c_0_1794,plain,(esk1_0=cell99|epred35_2(cell89,esk1_0)|epred34_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_1788]), c_0_1789]), c_0_1790]), c_0_1791]), c_0_1792]), c_0_1793])).
cnf(c_0_1795,plain,(cell91!=cell89), inference(split_conjunct,[status(thm)],[c_0_553])).
cnf(c_0_1796,plain,(cell92!=cell89), inference(split_conjunct,[status(thm)],[c_0_554])).
cnf(c_0_1797,plain,(cell93!=cell89), inference(split_conjunct,[status(thm)],[c_0_555])).
cnf(c_0_1798,plain,(cell94!=cell89), inference(split_conjunct,[status(thm)],[c_0_556])).
cnf(c_0_1799,plain,(cell95!=cell89), inference(split_conjunct,[status(thm)],[c_0_557])).
cnf(c_0_1800,plain,(esk1_0=cell99|epred34_2(cell89,esk1_0)|epred34_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_1794]), c_0_1795]), c_0_1796]), c_0_1797]), c_0_1798]), c_0_1799])).
cnf(c_0_1801,plain,(cell76!=cell89), inference(split_conjunct,[status(thm)],[c_0_558])).
cnf(c_0_1802,plain,(cell77!=cell89), inference(split_conjunct,[status(thm)],[c_0_559])).
cnf(c_0_1803,plain,(cell78!=cell89), inference(split_conjunct,[status(thm)],[c_0_560])).
cnf(c_0_1804,plain,(cell80!=cell89), inference(split_conjunct,[status(thm)],[c_0_561])).
cnf(c_0_1805,plain,(epred33_2(X1,X2)|X1=cell86|X1=cell87|X1=cell88|X2=cell79|X1=cell90|~epred34_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_706])).
cnf(c_0_1806,plain,(esk1_0=cell99|epred34_2(cell89,esk1_0)|epred33_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_721, c_0_1800]), c_0_1420]), c_0_1801]), c_0_1802]), c_0_1803]), c_0_1804])).
cnf(c_0_1807,plain,(esk1_0=cell79|esk1_0=cell99|epred33_2(cell89,esk1_0)|epred33_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1805, c_0_1806]), c_0_1604]), c_0_1241]), c_0_1781]), c_0_1782])).
cnf(c_0_1808,plain,(cell71!=cell89), inference(split_conjunct,[status(thm)],[c_0_562])).
cnf(c_0_1809,plain,(cell72!=cell89), inference(split_conjunct,[status(thm)],[c_0_563])).
cnf(c_0_1810,plain,(cell73!=cell89), inference(split_conjunct,[status(thm)],[c_0_564])).
cnf(c_0_1811,plain,(cell74!=cell89), inference(split_conjunct,[status(thm)],[c_0_565])).
cnf(c_0_1812,plain,(cell75!=cell89), inference(split_conjunct,[status(thm)],[c_0_566])).
cnf(c_0_1813,plain,(esk1_0=cell79|esk1_0=cell99|epred33_2(cell89,esk1_0)|epred32_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_745, c_0_1807]), c_0_1808]), c_0_1809]), c_0_1810]), c_0_1811]), c_0_1812])).
cnf(c_0_1814,plain,(epred31_2(X1,X2)|X2=cell66|X2=cell67|X2=cell68|X1=cell79|X2=cell70|~epred32_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_754])).
cnf(c_0_1815,plain,(esk1_0=cell79|esk1_0=cell99|epred32_2(cell89,esk1_0)|epred32_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_755, c_0_1813]), c_0_1789]), c_0_1790]), c_0_1791]), c_0_1792]), c_0_1793])).
cnf(c_0_1816,plain,(cell66!=cell89), inference(split_conjunct,[status(thm)],[c_0_567])).
cnf(c_0_1817,plain,(cell67!=cell89), inference(split_conjunct,[status(thm)],[c_0_568])).
cnf(c_0_1818,plain,(cell68!=cell89), inference(split_conjunct,[status(thm)],[c_0_569])).
cnf(c_0_1819,plain,(cell70!=cell89), inference(split_conjunct,[status(thm)],[c_0_570])).
cnf(c_0_1820,plain,(esk1_0=cell79|esk1_0=cell99|epred32_2(cell89,esk1_0)|epred31_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1814, c_0_1815]), c_0_1816]), c_0_1817]), c_0_1818]), c_0_1819])).
cnf(c_0_1821,plain,(esk1_0=cell79|esk1_0=cell99|epred31_2(cell89,esk1_0)|epred31_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_774, c_0_1820]), c_0_1420]), c_0_1801]), c_0_1802]), c_0_1803]), c_0_1804])).
cnf(c_0_1822,plain,(cell61!=cell89), inference(split_conjunct,[status(thm)],[c_0_571])).
cnf(c_0_1823,plain,(cell62!=cell89), inference(split_conjunct,[status(thm)],[c_0_572])).
cnf(c_0_1824,plain,(cell63!=cell89), inference(split_conjunct,[status(thm)],[c_0_573])).
cnf(c_0_1825,plain,(cell64!=cell89), inference(split_conjunct,[status(thm)],[c_0_574])).
cnf(c_0_1826,plain,(cell65!=cell89), inference(split_conjunct,[status(thm)],[c_0_575])).
cnf(c_0_1827,plain,(esk1_0=cell79|esk1_0=cell99|epred31_2(cell89,esk1_0)|epred30_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_783, c_0_1821]), c_0_1822]), c_0_1823]), c_0_1824]), c_0_1825]), c_0_1826])).
cnf(c_0_1828,plain,(esk1_0=cell79|esk1_0=cell99|epred30_2(cell89,esk1_0)|epred30_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_794, c_0_1827]), c_0_1808]), c_0_1809]), c_0_1810]), c_0_1811]), c_0_1812])).
cnf(c_0_1829,plain,(cell56!=cell89), inference(split_conjunct,[status(thm)],[c_0_576])).
cnf(c_0_1830,plain,(cell57!=cell89), inference(split_conjunct,[status(thm)],[c_0_577])).
cnf(c_0_1831,plain,(cell58!=cell89), inference(split_conjunct,[status(thm)],[c_0_578])).
cnf(c_0_1832,plain,(cell59!=cell89), inference(split_conjunct,[status(thm)],[c_0_579])).
cnf(c_0_1833,plain,(cell60!=cell89), inference(split_conjunct,[status(thm)],[c_0_580])).
cnf(c_0_1834,plain,(esk1_0=cell79|esk1_0=cell99|epred30_2(cell89,esk1_0)|epred29_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_803, c_0_1828]), c_0_1829]), c_0_1830]), c_0_1831]), c_0_1832]), c_0_1833])).
cnf(c_0_1835,plain,(cell69!=cell89), inference(split_conjunct,[status(thm)],[c_0_581])).
cnf(c_0_1836,plain,(esk1_0=cell79|esk1_0=cell99|epred29_2(cell89,esk1_0)|epred29_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_813, c_0_1834]), c_0_1816]), c_0_1817]), c_0_1818]), c_0_1835]), c_0_1819])).
cnf(c_0_1837,plain,(cell51!=cell89), inference(split_conjunct,[status(thm)],[c_0_582])).
cnf(c_0_1838,plain,(cell52!=cell89), inference(split_conjunct,[status(thm)],[c_0_583])).
cnf(c_0_1839,plain,(cell53!=cell89), inference(split_conjunct,[status(thm)],[c_0_584])).
cnf(c_0_1840,plain,(cell54!=cell89), inference(split_conjunct,[status(thm)],[c_0_585])).
cnf(c_0_1841,plain,(cell55!=cell89), inference(split_conjunct,[status(thm)],[c_0_586])).
cnf(c_0_1842,plain,(esk1_0=cell79|esk1_0=cell99|epred29_2(cell89,esk1_0)|epred28_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_822, c_0_1836]), c_0_1837]), c_0_1838]), c_0_1839]), c_0_1840]), c_0_1841])).
cnf(c_0_1843,plain,(esk1_0=cell79|esk1_0=cell99|epred28_2(cell89,esk1_0)|epred28_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_832, c_0_1842]), c_0_1822]), c_0_1823]), c_0_1824]), c_0_1825]), c_0_1826])).
cnf(c_0_1844,plain,(cell46!=cell89), inference(split_conjunct,[status(thm)],[c_0_587])).
cnf(c_0_1845,plain,(cell47!=cell89), inference(split_conjunct,[status(thm)],[c_0_588])).
cnf(c_0_1846,plain,(cell48!=cell89), inference(split_conjunct,[status(thm)],[c_0_589])).
cnf(c_0_1847,plain,(cell49!=cell89), inference(split_conjunct,[status(thm)],[c_0_590])).
cnf(c_0_1848,plain,(cell50!=cell89), inference(split_conjunct,[status(thm)],[c_0_591])).
cnf(c_0_1849,plain,(esk1_0=cell79|esk1_0=cell99|epred28_2(cell89,esk1_0)|epred27_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_841, c_0_1843]), c_0_1844]), c_0_1845]), c_0_1846]), c_0_1847]), c_0_1848])).
cnf(c_0_1850,plain,(esk1_0=cell79|esk1_0=cell99|epred27_2(cell89,esk1_0)|epred27_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_851, c_0_1849]), c_0_1829]), c_0_1830]), c_0_1831]), c_0_1832]), c_0_1833])).
cnf(c_0_1851,plain,(cell41!=cell89), inference(split_conjunct,[status(thm)],[c_0_592])).
cnf(c_0_1852,plain,(cell42!=cell89), inference(split_conjunct,[status(thm)],[c_0_593])).
cnf(c_0_1853,plain,(cell43!=cell89), inference(split_conjunct,[status(thm)],[c_0_594])).
cnf(c_0_1854,plain,(cell44!=cell89), inference(split_conjunct,[status(thm)],[c_0_595])).
cnf(c_0_1855,plain,(cell45!=cell89), inference(split_conjunct,[status(thm)],[c_0_596])).
cnf(c_0_1856,plain,(esk1_0=cell79|esk1_0=cell99|epred27_2(cell89,esk1_0)|epred26_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_860, c_0_1850]), c_0_1851]), c_0_1852]), c_0_1853]), c_0_1854]), c_0_1855])).
cnf(c_0_1857,plain,(esk1_0=cell79|esk1_0=cell99|epred26_2(cell89,esk1_0)|epred26_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_870, c_0_1856]), c_0_1837]), c_0_1838]), c_0_1839]), c_0_1840]), c_0_1841])).
cnf(c_0_1858,plain,(cell36!=cell89), inference(split_conjunct,[status(thm)],[c_0_597])).
cnf(c_0_1859,plain,(cell37!=cell89), inference(split_conjunct,[status(thm)],[c_0_598])).
cnf(c_0_1860,plain,(cell38!=cell89), inference(split_conjunct,[status(thm)],[c_0_599])).
cnf(c_0_1861,plain,(cell39!=cell89), inference(split_conjunct,[status(thm)],[c_0_600])).
cnf(c_0_1862,plain,(cell40!=cell89), inference(split_conjunct,[status(thm)],[c_0_601])).
cnf(c_0_1863,plain,(esk1_0=cell79|esk1_0=cell99|epred26_2(cell89,esk1_0)|epred25_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_879, c_0_1857]), c_0_1858]), c_0_1859]), c_0_1860]), c_0_1861]), c_0_1862])).
cnf(c_0_1864,plain,(esk1_0=cell79|esk1_0=cell99|epred25_2(cell89,esk1_0)|epred25_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_889, c_0_1863]), c_0_1844]), c_0_1845]), c_0_1846]), c_0_1847]), c_0_1848])).
cnf(c_0_1865,plain,(cell31!=cell89), inference(split_conjunct,[status(thm)],[c_0_602])).
cnf(c_0_1866,plain,(cell32!=cell89), inference(split_conjunct,[status(thm)],[c_0_603])).
cnf(c_0_1867,plain,(cell33!=cell89), inference(split_conjunct,[status(thm)],[c_0_604])).
cnf(c_0_1868,plain,(cell34!=cell89), inference(split_conjunct,[status(thm)],[c_0_605])).
cnf(c_0_1869,plain,(cell35!=cell89), inference(split_conjunct,[status(thm)],[c_0_606])).
cnf(c_0_1870,plain,(esk1_0=cell79|esk1_0=cell99|epred25_2(cell89,esk1_0)|epred24_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_898, c_0_1864]), c_0_1865]), c_0_1866]), c_0_1867]), c_0_1868]), c_0_1869])).
cnf(c_0_1871,plain,(esk1_0=cell79|esk1_0=cell99|epred24_2(cell89,esk1_0)|epred24_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_908, c_0_1870]), c_0_1851]), c_0_1852]), c_0_1853]), c_0_1854]), c_0_1855])).
cnf(c_0_1872,plain,(cell26!=cell89), inference(split_conjunct,[status(thm)],[c_0_607])).
cnf(c_0_1873,plain,(cell27!=cell89), inference(split_conjunct,[status(thm)],[c_0_608])).
cnf(c_0_1874,plain,(cell28!=cell89), inference(split_conjunct,[status(thm)],[c_0_609])).
cnf(c_0_1875,plain,(cell29!=cell89), inference(split_conjunct,[status(thm)],[c_0_610])).
cnf(c_0_1876,plain,(cell30!=cell89), inference(split_conjunct,[status(thm)],[c_0_611])).
cnf(c_0_1877,plain,(esk1_0=cell79|esk1_0=cell99|epred24_2(cell89,esk1_0)|epred23_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_917, c_0_1871]), c_0_1872]), c_0_1873]), c_0_1874]), c_0_1875]), c_0_1876])).
cnf(c_0_1878,plain,(esk1_0=cell79|esk1_0=cell99|epred23_2(cell89,esk1_0)|epred23_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_927, c_0_1877]), c_0_1858]), c_0_1859]), c_0_1860]), c_0_1861]), c_0_1862])).
cnf(c_0_1879,plain,(cell21!=cell89), inference(split_conjunct,[status(thm)],[c_0_612])).
cnf(c_0_1880,plain,(cell23!=cell89), inference(split_conjunct,[status(thm)],[c_0_613])).
cnf(c_0_1881,plain,(cell24!=cell89), inference(split_conjunct,[status(thm)],[c_0_614])).
cnf(c_0_1882,plain,(cell25!=cell89), inference(split_conjunct,[status(thm)],[c_0_615])).
cnf(c_0_1883,plain,(esk1_0=cell79|esk1_0=cell99|epred23_2(cell89,esk1_0)|epred22_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_1878]), c_0_1205]), c_0_1879]), c_0_1880]), c_0_1881]), c_0_1882])).
cnf(c_0_1884,plain,(esk1_0=cell79|esk1_0=cell99|epred22_2(cell89,esk1_0)|epred22_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_946, c_0_1883]), c_0_1865]), c_0_1866]), c_0_1867]), c_0_1868]), c_0_1869])).
cnf(c_0_1885,plain,(cell16!=cell89), inference(split_conjunct,[status(thm)],[c_0_616])).
cnf(c_0_1886,plain,(cell17!=cell89), inference(split_conjunct,[status(thm)],[c_0_617])).
cnf(c_0_1887,plain,(cell18!=cell89), inference(split_conjunct,[status(thm)],[c_0_618])).
cnf(c_0_1888,plain,(cell19!=cell89), inference(split_conjunct,[status(thm)],[c_0_619])).
cnf(c_0_1889,plain,(cell20!=cell89), inference(split_conjunct,[status(thm)],[c_0_620])).
cnf(c_0_1890,plain,(esk1_0=cell79|esk1_0=cell99|epred22_2(cell89,esk1_0)|epred21_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_955, c_0_1884]), c_0_1885]), c_0_1886]), c_0_1887]), c_0_1888]), c_0_1889])).
cnf(c_0_1891,plain,(esk1_0=cell79|esk1_0=cell99|epred21_2(cell89,esk1_0)|epred21_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_965, c_0_1890]), c_0_1872]), c_0_1873]), c_0_1874]), c_0_1875]), c_0_1876])).
cnf(c_0_1892,plain,(cell14!=cell89), inference(split_conjunct,[status(thm)],[c_0_621])).
cnf(c_0_1893,plain,(cell15!=cell89), inference(split_conjunct,[status(thm)],[c_0_622])).
cnf(c_0_1894,plain,(esk1_0=cell79|esk1_0=cell99|epred21_2(cell89,esk1_0)|epred20_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_974, c_0_1891]), c_0_1224]), c_0_1190]), c_0_1220]), c_0_1892]), c_0_1893])).
cnf(c_0_1895,plain,(esk1_0=cell79|esk1_0=cell99|epred20_2(cell89,esk1_0)|epred20_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_984, c_0_1894]), c_0_1205]), c_0_1879]), c_0_1880]), c_0_1881]), c_0_1882])).
cnf(c_0_1896,plain,(cell6!=cell89), inference(split_conjunct,[status(thm)],[c_0_623])).
cnf(c_0_1897,plain,(cell7!=cell89), inference(split_conjunct,[status(thm)],[c_0_624])).
cnf(c_0_1898,plain,(cell8!=cell89), inference(split_conjunct,[status(thm)],[c_0_625])).
cnf(c_0_1899,plain,(cell9!=cell89), inference(split_conjunct,[status(thm)],[c_0_626])).
cnf(c_0_1900,plain,(cell10!=cell89), inference(split_conjunct,[status(thm)],[c_0_627])).
cnf(c_0_1901,plain,(esk1_0=cell79|esk1_0=cell99|epred20_2(cell89,esk1_0)|epred19_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_993, c_0_1895]), c_0_1896]), c_0_1897]), c_0_1898]), c_0_1899]), c_0_1900])).
cnf(c_0_1902,plain,(esk1_0=cell79|esk1_0=cell99|epred19_2(cell89,esk1_0)|epred19_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1003, c_0_1901]), c_0_1885]), c_0_1886]), c_0_1887]), c_0_1888]), c_0_1889])).
cnf(c_0_1903,plain,(cell1!=cell89), inference(split_conjunct,[status(thm)],[c_0_628])).
cnf(c_0_1904,plain,(cell3!=cell89), inference(split_conjunct,[status(thm)],[c_0_629])).
cnf(c_0_1905,plain,(cell4!=cell89), inference(split_conjunct,[status(thm)],[c_0_630])).
cnf(c_0_1906,plain,(cell5!=cell89), inference(split_conjunct,[status(thm)],[c_0_631])).
cnf(c_0_1907,plain,(esk1_0=cell79|esk1_0=cell99|epred19_2(cell89,esk1_0)|epred18_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1023, c_0_1902]), c_0_679]), c_0_1903]), c_0_1904]), c_0_1905]), c_0_1906])).
cnf(c_0_1908,plain,(epred17_2(X1,X2)|X2=cell95|X2=cell96|X2=cell97|X1=cell99|X2=cell99|~epred18_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1020])).
cnf(c_0_1909,plain,(esk1_0=cell79|esk1_0=cell99|epred18_2(cell89,esk1_0)|epred18_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1021, c_0_1907]), c_0_1224]), c_0_1190]), c_0_1220]), c_0_1892]), c_0_1893])).
cnf(c_0_1910,plain,(esk1_0=cell99|esk1_0=cell79|epred18_2(cell89,esk1_0)|epred17_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1908, c_0_1909]), c_0_686]), c_0_1799]), c_0_1784]), c_0_1785])).
cnf(c_0_1911,plain,(epred16_2(X1,X2)|X1=cell90|X2=cell91|X2=cell92|X2=cell93|X2=cell94|~epred17_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1034])).
cnf(c_0_1912,plain,(esk1_0=cell99|esk1_0=cell79|epred17_2(cell89,esk1_0)|epred17_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1035, c_0_1910]), c_0_686]), c_0_1784]), c_0_1785]), c_0_1786]), c_0_1787])).
cnf(c_0_1913,plain,(esk1_0=cell90|esk1_0=cell99|esk1_0=cell79|epred17_2(cell89,esk1_0)|epred16_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1911, c_0_1912]), c_0_1795]), c_0_1796]), c_0_1797]), c_0_1798])).
cnf(c_0_1914,plain,(esk1_0=cell90|esk1_0=cell99|esk1_0=cell79|epred16_2(cell89,esk1_0)|epred16_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1046, c_0_1913]), c_0_1241]), c_0_1796]), c_0_1797]), c_0_1798]), c_0_1799])).
cnf(c_0_1915,plain,(epred15_2(X1,X2)|X1=cell85|X1=cell86|X1=cell87|X1=cell88|X2=cell88|~epred16_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1045])).
cnf(c_0_1916,plain,(esk1_0=cell90|esk1_0=cell99|esk1_0=cell79|epred16_2(cell89,esk1_0)|epred15_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1050, c_0_1914]), c_0_1604]), c_0_1792]), c_0_1793]), c_0_1781]), c_0_1782])).
cnf(c_0_1917,plain,(epred14_2(X1,X2)|X1=cell79|X2=cell79|X2=cell81|X2=cell82|X2=cell83|~epred15_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1054])).
cnf(c_0_1918,plain,(esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|esk1_0=cell79|epred15_2(cell89,esk1_0)|epred15_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1915, c_0_1916]), c_0_1604]), c_0_1793]), c_0_1781]), c_0_1782])).
cnf(c_0_1919,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred15_2(cell89,esk1_0)|epred14_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1917, c_0_1918]), c_0_1420]), c_0_1789]), c_0_1790]), c_0_1791])).
cnf(c_0_1920,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred14_2(cell89,esk1_0)|epred14_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1065, c_0_1919]), c_0_1420]), c_0_1804]), c_0_1790]), c_0_1791]), c_0_1792])).
cnf(c_0_1921,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred14_2(cell89,esk1_0)|epred13_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1069, c_0_1920]), c_0_1810]), c_0_1811]), c_0_1812]), c_0_1801]), c_0_1802])).
cnf(c_0_1922,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred13_2(cell89,esk1_0)|epred13_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1074, c_0_1921]), c_0_1811]), c_0_1812]), c_0_1801]), c_0_1802]), c_0_1803])).
cnf(c_0_1923,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred13_2(cell89,esk1_0)|epred12_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1078, c_0_1922]), c_0_1817]), c_0_1818]), c_0_1835]), c_0_1808]), c_0_1809])).
cnf(c_0_1924,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred12_2(cell89,esk1_0)|epred12_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1083, c_0_1923]), c_0_1818]), c_0_1835]), c_0_1819]), c_0_1809]), c_0_1810])).
cnf(c_0_1925,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred12_2(cell89,esk1_0)|epred11_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1087, c_0_1924]), c_0_1823]), c_0_1824]), c_0_1825]), c_0_1826]), c_0_1816])).
cnf(c_0_1926,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred11_2(cell89,esk1_0)|epred11_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1092, c_0_1925]), c_0_1824]), c_0_1825]), c_0_1826]), c_0_1816]), c_0_1817])).
cnf(c_0_1927,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred11_2(cell89,esk1_0)|epred10_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1096, c_0_1926]), c_0_1829]), c_0_1830]), c_0_1831]), c_0_1832]), c_0_1822])).
cnf(c_0_1928,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred10_2(cell89,esk1_0)|epred10_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1101, c_0_1927]), c_0_1830]), c_0_1831]), c_0_1832]), c_0_1833]), c_0_1823])).
cnf(c_0_1929,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred10_2(cell89,esk1_0)|epred9_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1105, c_0_1928]), c_0_1837]), c_0_1838]), c_0_1839]), c_0_1840]), c_0_1841])).
cnf(c_0_1930,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred9_2(cell89,esk1_0)|epred9_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1110, c_0_1929]), c_0_1838]), c_0_1839]), c_0_1840]), c_0_1841]), c_0_1829])).
cnf(c_0_1931,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred9_2(cell89,esk1_0)|epred8_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1114, c_0_1930]), c_0_1855]), c_0_1844]), c_0_1845]), c_0_1846]), c_0_1847])).
cnf(c_0_1932,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred8_2(cell89,esk1_0)|epred8_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1119, c_0_1931]), c_0_1844]), c_0_1845]), c_0_1846]), c_0_1847]), c_0_1848])).
cnf(c_0_1933,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred8_2(cell89,esk1_0)|epred7_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1123, c_0_1932]), c_0_1861]), c_0_1851]), c_0_1852]), c_0_1853]), c_0_1854])).
cnf(c_0_1934,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred7_2(cell89,esk1_0)|epred7_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1128, c_0_1933]), c_0_1862]), c_0_1852]), c_0_1853]), c_0_1854]), c_0_1855])).
cnf(c_0_1935,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred7_2(cell89,esk1_0)|epred6_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1132, c_0_1934]), c_0_1868]), c_0_1869]), c_0_1858]), c_0_1859]), c_0_1860])).
cnf(c_0_1936,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred6_2(cell89,esk1_0)|epred6_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1137, c_0_1935]), c_0_1869]), c_0_1858]), c_0_1859]), c_0_1860]), c_0_1861])).
cnf(c_0_1937,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred6_2(cell89,esk1_0)|epred5_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1141, c_0_1936]), c_0_1874]), c_0_1875]), c_0_1865]), c_0_1866]), c_0_1867])).
cnf(c_0_1938,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred5_2(cell89,esk1_0)|epred5_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1146, c_0_1937]), c_0_1875]), c_0_1876]), c_0_1866]), c_0_1867]), c_0_1868])).
cnf(c_0_1939,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred5_2(cell89,esk1_0)|epred4_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1150, c_0_1938]), c_0_1880]), c_0_1881]), c_0_1882]), c_0_1872]), c_0_1873])).
cnf(c_0_1940,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred4_2(cell89,esk1_0)|epred4_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1155, c_0_1939]), c_0_1881]), c_0_1882]), c_0_1872]), c_0_1873]), c_0_1874])).
cnf(c_0_1941,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred4_2(cell89,esk1_0)|epred3_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1160, c_0_1940]), c_0_1205]), c_0_1886]), c_0_1887]), c_0_1888]), c_0_1879])).
cnf(c_0_1942,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred3_2(cell89,esk1_0)|epred3_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1165, c_0_1941]), c_0_1205]), c_0_1887]), c_0_1888]), c_0_1889]), c_0_1880])).
cnf(c_0_1943,plain,(X2=cell1|X1=cell3|X2=cell3|X1=cell5|X1=cell6|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_1168])).
cnf(c_0_1944,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred3_2(cell89,esk1_0)|epred2_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1170, c_0_1942]), c_0_1224]), c_0_1220]), c_0_1892]), c_0_1893]), c_0_1885])).
cnf(c_0_1945,plain,(cell1=X1|cell3=X1|cell3=X2|cell5=X2|cell6=X2|cell12=X2|cell7=X2|cell8=X2|cell9=X2|cell10=X2|~epred2_2(X2,X1)), inference(spm,[status(thm)],[c_0_1943, c_0_1407])).
cnf(c_0_1946,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred2_2(cell89,esk1_0)|epred2_2(esk1_0,cell89)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1175, c_0_1944]), c_0_1220]), c_0_1892]), c_0_1893]), c_0_1885]), c_0_1886])).
cnf(c_0_1947,plain,(esk2_0=cell1|esk2_0=cell3|esk2_0=cell12|esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1945, c_0_1182]), c_0_1015]), c_0_1017]), c_0_995]), c_0_976]), c_0_996]), c_0_997]), c_0_998]), c_0_999])).
cnf(c_0_1948,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99|epred2_2(cell89,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1408, c_0_1946]), c_0_679]), c_0_1903]), c_0_1904]), c_0_1905]), c_0_1906]), c_0_1190]), c_0_1896]), c_0_1897]), c_0_1898]), c_0_1899])).
cnf(c_0_1949,plain,(esk1_0=cell12|esk1_0=cell13|esk1_0=cell22|esk1_0=cell11), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_1947, c_0_1780]), c_0_1780]), c_0_1780]), c_0_1903]), c_0_1904]), c_0_1224])).
cnf(c_0_1950,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90|esk1_0=cell99), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1410, c_0_1948]), c_0_679]), c_0_1904]), c_0_1905]), c_0_1906]), c_0_1896]), c_0_1224]), c_0_1897]), c_0_1898]), c_0_1899]), c_0_1900])).
cnf(c_0_1951,plain,(esk1_0=cell79|esk1_0=cell88|esk1_0=cell90), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1949, c_0_1950]), c_0_987]), c_0_989]), c_0_949]), c_0_988])).
cnf(c_0_1952,plain,(esk1_0=cell79|esk1_0=cell88), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1949, c_0_1951]), c_0_1355]), c_0_1357]), c_0_1341]), c_0_1356])).
cnf(c_0_1953,plain,(esk1_0=cell88), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1949, c_0_1952]), c_0_1533]), c_0_1535]), c_0_1519]), c_0_1534])).
cnf(c_0_1954,plain,($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_1949, c_0_1953]), c_0_1953]), c_0_1953]), c_0_1716]), c_0_1718]), c_0_1702]), c_0_1953]), c_0_1717]), ['proof']).
# SZS output end CNFRefutation.
EOF
