% Problem    : PUZ077+1 : TPTP v6.1.0. Released v3.5.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n185.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 06:02:11 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, axiom, (![X3]:![X4]:(north(X4,X3)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof34976)).
fof(c_0_1, axiom, (![X1]:![X2]:(west(X2,X1)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60))|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65))|(X2=cell56&X1=cell66))|(X2=cell57&X1=cell67))|(X2=cell58&X1=cell68))|(X2=cell59&X1=cell69))|(X2=cell60&X1=cell70))|(X2=cell61&X1=cell71))|(X2=cell62&X1=cell72))|(X2=cell63&X1=cell73))|(X2=cell64&X1=cell74))|(X2=cell65&X1=cell75))|(X2=cell66&X1=cell76))|(X2=cell67&X1=cell77))|(X2=cell68&X1=cell78))|(X2=cell69&X1=cell79))|(X2=cell70&X1=cell80))|(X2=cell71&X1=cell81))|(X2=cell72&X1=cell82))|(X2=cell73&X1=cell83))|(X2=cell74&X1=cell84))|(X2=cell75&X1=cell85))|(X2=cell76&X1=cell86))|(X2=cell77&X1=cell87))|(X2=cell78&X1=cell88))|(X2=cell79&X1=cell89))|(X2=cell80&X1=cell90))|(X2=cell81&X1=cell91))|(X2=cell82&X1=cell92))|(X2=cell83&X1=cell93))|(X2=cell84&X1=cell94))|(X2=cell85&X1=cell95))|(X2=cell86&X1=cell96))|(X2=cell87&X1=cell97))|(X2=cell88&X1=cell98))|(X2=cell89&X1=cell99))|(X2=cell90&X1=cell100)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof34975)).
fof(c_0_2, axiom, (![X43]:![X44]:(duenorth(X44,X43)<=>((((((((north(X44,X43)|?[X45]:(north(X44,X45)&north(X45,X43)))|?[X46]:?[X47]:((north(X44,X46)&north(X46,X47))&north(X47,X43)))|?[X48]:?[X49]:?[X50]:(((north(X44,X48)&north(X48,X49))&north(X49,X50))&north(X50,X43)))|?[X51]:?[X52]:?[X53]:?[X54]:((((north(X44,X51)&north(X51,X52))&north(X52,X53))&north(X53,X54))&north(X54,X43)))|?[X55]:?[X56]:?[X57]:?[X58]:?[X59]:(((((north(X44,X55)&north(X55,X56))&north(X56,X57))&north(X57,X58))&north(X58,X59))&north(X59,X43)))|?[X60]:?[X61]:?[X62]:?[X63]:?[X64]:?[X65]:((((((north(X44,X60)&north(X60,X61))&north(X61,X62))&north(X62,X63))&north(X63,X64))&north(X64,X65))&north(X65,X43)))|?[X66]:?[X67]:?[X68]:?[X69]:?[X70]:?[X71]:?[X72]:(((((((north(X44,X66)&north(X66,X67))&north(X67,X68))&north(X68,X69))&north(X69,X70))&north(X70,X71))&north(X71,X72))&north(X72,X43)))|?[X73]:?[X74]:?[X75]:?[X76]:?[X77]:?[X78]:?[X79]:?[X80]:((((((((north(X44,X73)&north(X73,X74))&north(X74,X75))&north(X75,X76))&north(X76,X77))&north(X77,X78))&north(X78,X79))&north(X79,X80))&north(X80,X43))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof34978)).
fof(c_0_3, axiom, (![X81]:![X82]:(vert(X82,X81)<=>(duenorth(X82,X81)|duenorth(X81,X82)))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof34979)).
fof(c_0_4, axiom, (![X5]:![X6]:(duewest(X6,X5)<=>((((((((west(X6,X5)|?[X7]:(west(X6,X7)&west(X7,X5)))|?[X8]:?[X9]:((west(X6,X8)&west(X8,X9))&west(X9,X5)))|?[X10]:?[X11]:?[X12]:(((west(X6,X10)&west(X10,X11))&west(X11,X12))&west(X12,X5)))|?[X13]:?[X14]:?[X15]:?[X16]:((((west(X6,X13)&west(X13,X14))&west(X14,X15))&west(X15,X16))&west(X16,X5)))|?[X17]:?[X18]:?[X19]:?[X20]:?[X21]:(((((west(X6,X17)&west(X17,X18))&west(X18,X19))&west(X19,X20))&west(X20,X21))&west(X21,X5)))|?[X22]:?[X23]:?[X24]:?[X25]:?[X26]:?[X27]:((((((west(X6,X22)&west(X22,X23))&west(X23,X24))&west(X24,X25))&west(X25,X26))&west(X26,X27))&west(X27,X5)))|?[X28]:?[X29]:?[X30]:?[X31]:?[X32]:?[X33]:?[X34]:(((((((west(X6,X28)&west(X28,X29))&west(X29,X30))&west(X30,X31))&west(X31,X32))&west(X32,X33))&west(X33,X34))&west(X34,X5)))|?[X35]:?[X36]:?[X37]:?[X38]:?[X39]:?[X40]:?[X41]:?[X42]:((((((((west(X6,X35)&west(X35,X36))&west(X36,X37))&west(X37,X38))&west(X38,X39))&west(X39,X40))&west(X40,X41))&west(X41,X42))&west(X42,X5))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof34977)).
fof(c_0_5, axiom, (![X81]:![X82]:(westof(X82,X81)<=>(duewest(X82,X81)|?[X83]:(vert(X82,X83)&duewest(X83,X81))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof34980)).
fof(c_0_6, conjecture, (westof(cell1,cell100)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', tlhfof34981)).
fof(c_0_7, plain, (![X4]:![X3]:(epred19_2(X3,X4)<=>(((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6)))), introduced(definition)).
fof(c_0_8, plain, (![X4]:![X3]:(epred20_2(X3,X4)<=>((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12)))), introduced(definition)).
fof(c_0_9, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell5|(X5=cell4|(X5=cell3|(X5=cell2|X5=cell1))))|~epred19_2(X6,X5))&((X6=cell6|(X5=cell4|(X5=cell3|(X5=cell2|X5=cell1))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X5=cell3|(X5=cell2|X5=cell1))))|~epred19_2(X6,X5))&((X6=cell6|(X6=cell5|(X5=cell3|(X5=cell2|X5=cell1))))|~epred19_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell4|(X5=cell2|X5=cell1))))|~epred19_2(X6,X5))&((X6=cell6|(X5=cell4|(X6=cell4|(X5=cell2|X5=cell1))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X6=cell4|(X5=cell2|X5=cell1))))|~epred19_2(X6,X5))&((X6=cell6|(X6=cell5|(X6=cell4|(X5=cell2|X5=cell1))))|~epred19_2(X6,X5)))))&(((((X5=cell5|(X5=cell4|(X5=cell3|(X6=cell3|X5=cell1))))|~epred19_2(X6,X5))&((X6=cell6|(X5=cell4|(X5=cell3|(X6=cell3|X5=cell1))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X5=cell3|(X6=cell3|X5=cell1))))|~epred19_2(X6,X5))&((X6=cell6|(X6=cell5|(X5=cell3|(X6=cell3|X5=cell1))))|~epred19_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell4|(X6=cell3|X5=cell1))))|~epred19_2(X6,X5))&((X6=cell6|(X5=cell4|(X6=cell4|(X6=cell3|X5=cell1))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X6=cell4|(X6=cell3|X5=cell1))))|~epred19_2(X6,X5))&((X6=cell6|(X6=cell5|(X6=cell4|(X6=cell3|X5=cell1))))|~epred19_2(X6,X5))))))&((((((X5=cell5|(X5=cell4|(X5=cell3|(X5=cell2|X6=cell2))))|~epred19_2(X6,X5))&((X6=cell6|(X5=cell4|(X5=cell3|(X5=cell2|X6=cell2))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X5=cell3|(X5=cell2|X6=cell2))))|~epred19_2(X6,X5))&((X6=cell6|(X6=cell5|(X5=cell3|(X5=cell2|X6=cell2))))|~epred19_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell4|(X5=cell2|X6=cell2))))|~epred19_2(X6,X5))&((X6=cell6|(X5=cell4|(X6=cell4|(X5=cell2|X6=cell2))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X6=cell4|(X5=cell2|X6=cell2))))|~epred19_2(X6,X5))&((X6=cell6|(X6=cell5|(X6=cell4|(X5=cell2|X6=cell2))))|~epred19_2(X6,X5)))))&(((((X5=cell5|(X5=cell4|(X5=cell3|(X6=cell3|X6=cell2))))|~epred19_2(X6,X5))&((X6=cell6|(X5=cell4|(X5=cell3|(X6=cell3|X6=cell2))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X5=cell3|(X6=cell3|X6=cell2))))|~epred19_2(X6,X5))&((X6=cell6|(X6=cell5|(X5=cell3|(X6=cell3|X6=cell2))))|~epred19_2(X6,X5))))&((((X5=cell5|(X5=cell4|(X6=cell4|(X6=cell3|X6=cell2))))|~epred19_2(X6,X5))&((X6=cell6|(X5=cell4|(X6=cell4|(X6=cell3|X6=cell2))))|~epred19_2(X6,X5)))&(((X5=cell5|(X6=cell5|(X6=cell4|(X6=cell3|X6=cell2))))|~epred19_2(X6,X5))&((X6=cell6|(X6=cell5|(X6=cell4|(X6=cell3|X6=cell2))))|~epred19_2(X6,X5)))))))&((((((X7!=cell1|X8!=cell2)|epred19_2(X8,X7))&((X7!=cell2|X8!=cell3)|epred19_2(X8,X7)))&((X7!=cell3|X8!=cell4)|epred19_2(X8,X7)))&((X7!=cell4|X8!=cell5)|epred19_2(X8,X7)))&((X7!=cell5|X8!=cell6)|epred19_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
fof(c_0_10, plain, (![X4]:![X3]:(epred21_2(X3,X4)<=>(((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17)))), introduced(definition)).
fof(c_0_11, plain, (![X4]:![X3]:(epred20_2(X3,X4)<=>(((((epred19_2(X3,X4)|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12)))), inference(apply_def,[status(thm)],[c_0_8, c_0_7])).
cnf(c_0_12,plain,(epred19_2(X1,X2)|X1!=cell3|X2!=cell2), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, plain, (![X4]:![X3]:(epred22_2(X3,X4)<=>((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23)))), introduced(definition)).
fof(c_0_14, plain, (![X4]:![X3]:(epred21_2(X3,X4)<=>(((((epred20_2(X3,X4)|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17)))), inference(apply_def,[status(thm)],[c_0_10, c_0_8])).
fof(c_0_15, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell11|(X5=cell9|(X5=cell8|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X5=cell9|(X5=cell8|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X5=cell8|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X6=cell10|(X5=cell8|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))&((((X5=cell11|(X5=cell9|(X6=cell9|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X5=cell9|(X6=cell9|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X6=cell9|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X6=cell10|(X6=cell9|(X5=cell7|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))))&(((((X5=cell11|(X5=cell9|(X5=cell8|(X6=cell8|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X5=cell9|(X5=cell8|(X6=cell8|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X5=cell8|(X6=cell8|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X6=cell10|(X5=cell8|(X6=cell8|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))&((((X5=cell11|(X5=cell9|(X6=cell9|(X6=cell8|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X5=cell9|(X6=cell9|(X6=cell8|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X6=cell9|(X6=cell8|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X6=cell10|(X6=cell9|(X6=cell8|(X5=cell6|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))))&((((((X5=cell11|(X5=cell9|(X5=cell8|(X5=cell7|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X5=cell9|(X5=cell8|(X5=cell7|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X5=cell8|(X5=cell7|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X6=cell10|(X5=cell8|(X5=cell7|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))&((((X5=cell11|(X5=cell9|(X6=cell9|(X5=cell7|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X5=cell9|(X6=cell9|(X5=cell7|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X6=cell9|(X5=cell7|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X6=cell10|(X6=cell9|(X5=cell7|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))))&(((((X5=cell11|(X5=cell9|(X5=cell8|(X6=cell8|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X5=cell9|(X5=cell8|(X6=cell8|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X5=cell8|(X6=cell8|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X6=cell10|(X5=cell8|(X6=cell8|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))))&((((X5=cell11|(X5=cell9|(X6=cell9|(X6=cell8|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X5=cell9|(X6=cell9|(X6=cell8|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))&(((X5=cell11|(X6=cell10|(X6=cell9|(X6=cell8|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5))&((X6=cell12|(X6=cell10|(X6=cell9|(X6=cell8|(X6=cell7|epred19_2(X6,X5))))))|~epred20_2(X6,X5)))))))&((((((~epred19_2(X8,X7)|epred20_2(X8,X7))&((X7!=cell6|X8!=cell7)|epred20_2(X8,X7)))&((X7!=cell7|X8!=cell8)|epred20_2(X8,X7)))&((X7!=cell8|X8!=cell9)|epred20_2(X8,X7)))&((X7!=cell9|X8!=cell10)|epred20_2(X8,X7)))&((X7!=cell11|X8!=cell12)|epred20_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
cnf(c_0_16,plain,(epred19_2(X1,cell2)|cell3!=X1), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_17,plain,(epred19_2(X1,X2)|X1!=cell4|X2!=cell3), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_18, plain, (![X4]:![X3]:(epred23_2(X3,X4)<=>(((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28)))), introduced(definition)).
fof(c_0_19, plain, (![X4]:![X3]:(epred22_2(X3,X4)<=>(((((epred21_2(X3,X4)|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23)))), inference(apply_def,[status(thm)],[c_0_13, c_0_10])).
fof(c_0_20, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell16|(X5=cell15|(X5=cell14|(X5=cell13|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X5=cell15|(X5=cell14|(X5=cell13|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X5=cell14|(X5=cell13|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X6=cell16|(X5=cell14|(X5=cell13|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))&((((X5=cell16|(X5=cell15|(X6=cell15|(X5=cell13|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X5=cell15|(X6=cell15|(X5=cell13|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X6=cell15|(X5=cell13|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X6=cell16|(X6=cell15|(X5=cell13|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))))&(((((X5=cell16|(X5=cell15|(X5=cell14|(X6=cell14|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X5=cell15|(X5=cell14|(X6=cell14|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X5=cell14|(X6=cell14|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X6=cell16|(X5=cell14|(X6=cell14|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))&((((X5=cell16|(X5=cell15|(X6=cell15|(X6=cell14|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X5=cell15|(X6=cell15|(X6=cell14|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X6=cell15|(X6=cell14|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X6=cell16|(X6=cell15|(X6=cell14|(X5=cell12|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))))&((((((X5=cell16|(X5=cell15|(X5=cell14|(X5=cell13|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X5=cell15|(X5=cell14|(X5=cell13|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X5=cell14|(X5=cell13|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X6=cell16|(X5=cell14|(X5=cell13|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))&((((X5=cell16|(X5=cell15|(X6=cell15|(X5=cell13|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X5=cell15|(X6=cell15|(X5=cell13|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X6=cell15|(X5=cell13|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X6=cell16|(X6=cell15|(X5=cell13|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))))&(((((X5=cell16|(X5=cell15|(X5=cell14|(X6=cell14|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X5=cell15|(X5=cell14|(X6=cell14|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X5=cell14|(X6=cell14|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X6=cell16|(X5=cell14|(X6=cell14|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))))&((((X5=cell16|(X5=cell15|(X6=cell15|(X6=cell14|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X5=cell15|(X6=cell15|(X6=cell14|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))&(((X5=cell16|(X6=cell16|(X6=cell15|(X6=cell14|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5))&((X6=cell17|(X6=cell16|(X6=cell15|(X6=cell14|(X6=cell13|epred20_2(X6,X5))))))|~epred21_2(X6,X5)))))))&((((((~epred20_2(X8,X7)|epred21_2(X8,X7))&((X7!=cell12|X8!=cell13)|epred21_2(X8,X7)))&((X7!=cell13|X8!=cell14)|epred21_2(X8,X7)))&((X7!=cell14|X8!=cell15)|epred21_2(X8,X7)))&((X7!=cell15|X8!=cell16)|epred21_2(X8,X7)))&((X7!=cell16|X8!=cell17)|epred21_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])])).
cnf(c_0_21,plain,(epred20_2(X1,X2)|~epred19_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22,plain,(epred19_2(cell3,cell2)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_23,plain,(epred19_2(X1,cell3)|cell4!=X1), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_24,plain,(epred19_2(X1,X2)|X1!=cell2|X2!=cell1), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_25, plain, (![X2]:![X1]:(epred4_2(X1,X2)<=>((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30)))), introduced(definition)).
fof(c_0_26, plain, (![X2]:![X1]:(epred3_2(X1,X2)<=>(((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25)))), introduced(definition)).
fof(c_0_27, plain, (![X2]:![X1]:(epred2_2(X1,X2)<=>((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20)))), introduced(definition)).
fof(c_0_28, plain, (![X2]:![X1]:(epred1_2(X1,X2)<=>(((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15)))), introduced(definition)).
fof(c_0_29, plain, (![X4]:![X3]:(epred24_2(X3,X4)<=>((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34)))), introduced(definition)).
fof(c_0_30, plain, (![X4]:![X3]:(epred23_2(X3,X4)<=>(((((epred22_2(X3,X4)|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28)))), inference(apply_def,[status(thm)],[c_0_18, c_0_13])).
fof(c_0_31, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell22|(X5=cell21|(X5=cell19|(X5=cell18|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X5=cell21|(X5=cell19|(X5=cell18|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X5=cell19|(X5=cell18|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X6=cell22|(X5=cell19|(X5=cell18|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))&((((X5=cell22|(X5=cell21|(X6=cell20|(X5=cell18|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X5=cell21|(X6=cell20|(X5=cell18|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X6=cell20|(X5=cell18|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X6=cell22|(X6=cell20|(X5=cell18|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))))&(((((X5=cell22|(X5=cell21|(X5=cell19|(X6=cell19|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X5=cell21|(X5=cell19|(X6=cell19|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X5=cell19|(X6=cell19|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X6=cell22|(X5=cell19|(X6=cell19|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))&((((X5=cell22|(X5=cell21|(X6=cell20|(X6=cell19|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X5=cell21|(X6=cell20|(X6=cell19|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X6=cell20|(X6=cell19|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X6=cell22|(X6=cell20|(X6=cell19|(X5=cell17|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))))&((((((X5=cell22|(X5=cell21|(X5=cell19|(X5=cell18|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X5=cell21|(X5=cell19|(X5=cell18|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X5=cell19|(X5=cell18|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X6=cell22|(X5=cell19|(X5=cell18|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))&((((X5=cell22|(X5=cell21|(X6=cell20|(X5=cell18|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X5=cell21|(X6=cell20|(X5=cell18|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X6=cell20|(X5=cell18|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X6=cell22|(X6=cell20|(X5=cell18|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))))&(((((X5=cell22|(X5=cell21|(X5=cell19|(X6=cell19|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X5=cell21|(X5=cell19|(X6=cell19|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X5=cell19|(X6=cell19|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X6=cell22|(X5=cell19|(X6=cell19|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))))&((((X5=cell22|(X5=cell21|(X6=cell20|(X6=cell19|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X5=cell21|(X6=cell20|(X6=cell19|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))&(((X5=cell22|(X6=cell22|(X6=cell20|(X6=cell19|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5))&((X6=cell23|(X6=cell22|(X6=cell20|(X6=cell19|(X6=cell18|epred21_2(X6,X5))))))|~epred22_2(X6,X5)))))))&((((((~epred21_2(X8,X7)|epred22_2(X8,X7))&((X7!=cell17|X8!=cell18)|epred22_2(X8,X7)))&((X7!=cell18|X8!=cell19)|epred22_2(X8,X7)))&((X7!=cell19|X8!=cell20)|epred22_2(X8,X7)))&((X7!=cell21|X8!=cell22)|epred22_2(X8,X7)))&((X7!=cell22|X8!=cell23)|epred22_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])])).
cnf(c_0_32,plain,(epred21_2(X1,X2)|~epred20_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33,plain,(epred20_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_34,plain,(epred19_2(cell4,cell3)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_35,plain,(epred19_2(X1,cell1)|cell2!=X1), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_36,plain,(epred19_2(X1,X2)|X1!=cell5|X2!=cell4), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_37, plain, (![X2]:![X1]:(epred4_2(X1,X2)<=>(((((epred3_2(X1,X2)|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30)))), inference(apply_def,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_38, plain, (![X2]:![X1]:(epred2_2(X1,X2)<=>(((((epred1_2(X1,X2)|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20)))), inference(apply_def,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_39, plain, (![X4]:![X3]:(epred25_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39)))), introduced(definition)).
fof(c_0_40, plain, (![X4]:![X3]:(epred24_2(X3,X4)<=>(((((epred23_2(X3,X4)|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34)))), inference(apply_def,[status(thm)],[c_0_29, c_0_18])).
fof(c_0_41, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell27|(X5=cell26|(X5=cell25|(X5=cell24|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X5=cell26|(X5=cell25|(X5=cell24|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X5=cell25|(X5=cell24|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X6=cell27|(X5=cell25|(X5=cell24|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))&((((X5=cell27|(X5=cell26|(X6=cell26|(X5=cell24|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X5=cell26|(X6=cell26|(X5=cell24|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X6=cell26|(X5=cell24|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X6=cell27|(X6=cell26|(X5=cell24|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))))&(((((X5=cell27|(X5=cell26|(X5=cell25|(X6=cell25|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X5=cell26|(X5=cell25|(X6=cell25|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X5=cell25|(X6=cell25|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X6=cell27|(X5=cell25|(X6=cell25|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))&((((X5=cell27|(X5=cell26|(X6=cell26|(X6=cell25|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X5=cell26|(X6=cell26|(X6=cell25|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X6=cell26|(X6=cell25|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X6=cell27|(X6=cell26|(X6=cell25|(X5=cell23|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))))&((((((X5=cell27|(X5=cell26|(X5=cell25|(X5=cell24|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X5=cell26|(X5=cell25|(X5=cell24|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X5=cell25|(X5=cell24|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X6=cell27|(X5=cell25|(X5=cell24|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))&((((X5=cell27|(X5=cell26|(X6=cell26|(X5=cell24|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X5=cell26|(X6=cell26|(X5=cell24|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X6=cell26|(X5=cell24|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X6=cell27|(X6=cell26|(X5=cell24|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))))&(((((X5=cell27|(X5=cell26|(X5=cell25|(X6=cell25|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X5=cell26|(X5=cell25|(X6=cell25|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X5=cell25|(X6=cell25|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X6=cell27|(X5=cell25|(X6=cell25|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))))&((((X5=cell27|(X5=cell26|(X6=cell26|(X6=cell25|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X5=cell26|(X6=cell26|(X6=cell25|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))&(((X5=cell27|(X6=cell27|(X6=cell26|(X6=cell25|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5))&((X6=cell28|(X6=cell27|(X6=cell26|(X6=cell25|(X6=cell24|epred22_2(X6,X5))))))|~epred23_2(X6,X5)))))))&((((((~epred22_2(X8,X7)|epred23_2(X8,X7))&((X7!=cell23|X8!=cell24)|epred23_2(X8,X7)))&((X7!=cell24|X8!=cell25)|epred23_2(X8,X7)))&((X7!=cell25|X8!=cell26)|epred23_2(X8,X7)))&((X7!=cell26|X8!=cell27)|epred23_2(X8,X7)))&((X7!=cell27|X8!=cell28)|epred23_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])])).
cnf(c_0_42,plain,(epred22_2(X1,X2)|~epred21_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43,plain,(epred21_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_44,plain,(epred20_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_21, c_0_34])).
cnf(c_0_45,plain,(epred19_2(cell2,cell1)), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_46,plain,(epred19_2(X1,cell4)|cell5!=X1), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_47,plain,(epred19_2(X1,X2)|X1!=cell6|X2!=cell5), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_48, plain, (![X2]:![X1]:(epred5_2(X1,X2)<=>(((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35)))), introduced(definition)).
fof(c_0_49, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell20|(X3=cell19|(X3=cell18|(X3=cell17|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X3=cell19|(X3=cell18|(X3=cell17|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))&(((X3=cell20|(X4=cell29|(X3=cell18|(X3=cell17|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X4=cell29|(X3=cell18|(X3=cell17|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))))&((((X3=cell20|(X3=cell19|(X4=cell28|(X3=cell17|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X3=cell19|(X4=cell28|(X3=cell17|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))&(((X3=cell20|(X4=cell29|(X4=cell28|(X3=cell17|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X4=cell29|(X4=cell28|(X3=cell17|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))))&(((((X3=cell20|(X3=cell19|(X3=cell18|(X4=cell27|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X3=cell19|(X3=cell18|(X4=cell27|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))&(((X3=cell20|(X4=cell29|(X3=cell18|(X4=cell27|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X4=cell29|(X3=cell18|(X4=cell27|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))))&((((X3=cell20|(X3=cell19|(X4=cell28|(X4=cell27|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X3=cell19|(X4=cell28|(X4=cell27|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))&(((X3=cell20|(X4=cell29|(X4=cell28|(X4=cell27|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X4=cell29|(X4=cell28|(X4=cell27|(X3=cell16|epred3_2(X4,X3))))))|~epred4_2(X4,X3))))))&((((((X3=cell20|(X3=cell19|(X3=cell18|(X3=cell17|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X3=cell19|(X3=cell18|(X3=cell17|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))&(((X3=cell20|(X4=cell29|(X3=cell18|(X3=cell17|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X4=cell29|(X3=cell18|(X3=cell17|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))))&((((X3=cell20|(X3=cell19|(X4=cell28|(X3=cell17|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X3=cell19|(X4=cell28|(X3=cell17|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))&(((X3=cell20|(X4=cell29|(X4=cell28|(X3=cell17|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X4=cell29|(X4=cell28|(X3=cell17|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))))&(((((X3=cell20|(X3=cell19|(X3=cell18|(X4=cell27|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X3=cell19|(X3=cell18|(X4=cell27|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))&(((X3=cell20|(X4=cell29|(X3=cell18|(X4=cell27|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X4=cell29|(X3=cell18|(X4=cell27|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))))&((((X3=cell20|(X3=cell19|(X4=cell28|(X4=cell27|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X3=cell19|(X4=cell28|(X4=cell27|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))&(((X3=cell20|(X4=cell29|(X4=cell28|(X4=cell27|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3))&((X4=cell30|(X4=cell29|(X4=cell28|(X4=cell27|(X4=cell26|epred3_2(X4,X3))))))|~epred4_2(X4,X3)))))))&((((((~epred3_2(X6,X5)|epred4_2(X6,X5))&((X5!=cell16|X6!=cell26)|epred4_2(X6,X5)))&((X5!=cell17|X6!=cell27)|epred4_2(X6,X5)))&((X5!=cell18|X6!=cell28)|epred4_2(X6,X5)))&((X5!=cell19|X6!=cell29)|epred4_2(X6,X5)))&((X5!=cell20|X6!=cell30)|epred4_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])])])).
fof(c_0_50, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell10|(X3=cell9|(X3=cell8|(X3=cell7|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X3=cell9|(X3=cell8|(X3=cell7|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))&(((X3=cell10|(X4=cell19|(X3=cell8|(X3=cell7|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X4=cell19|(X3=cell8|(X3=cell7|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))))&((((X3=cell10|(X3=cell9|(X4=cell18|(X3=cell7|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X3=cell9|(X4=cell18|(X3=cell7|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))&(((X3=cell10|(X4=cell19|(X4=cell18|(X3=cell7|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X4=cell19|(X4=cell18|(X3=cell7|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))))&(((((X3=cell10|(X3=cell9|(X3=cell8|(X4=cell17|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X3=cell9|(X3=cell8|(X4=cell17|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))&(((X3=cell10|(X4=cell19|(X3=cell8|(X4=cell17|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X4=cell19|(X3=cell8|(X4=cell17|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))))&((((X3=cell10|(X3=cell9|(X4=cell18|(X4=cell17|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X3=cell9|(X4=cell18|(X4=cell17|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))&(((X3=cell10|(X4=cell19|(X4=cell18|(X4=cell17|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X4=cell19|(X4=cell18|(X4=cell17|(X3=cell6|epred1_2(X4,X3))))))|~epred2_2(X4,X3))))))&((((((X3=cell10|(X3=cell9|(X3=cell8|(X3=cell7|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X3=cell9|(X3=cell8|(X3=cell7|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))&(((X3=cell10|(X4=cell19|(X3=cell8|(X3=cell7|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X4=cell19|(X3=cell8|(X3=cell7|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))))&((((X3=cell10|(X3=cell9|(X4=cell18|(X3=cell7|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X3=cell9|(X4=cell18|(X3=cell7|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))&(((X3=cell10|(X4=cell19|(X4=cell18|(X3=cell7|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X4=cell19|(X4=cell18|(X3=cell7|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))))&(((((X3=cell10|(X3=cell9|(X3=cell8|(X4=cell17|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X3=cell9|(X3=cell8|(X4=cell17|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))&(((X3=cell10|(X4=cell19|(X3=cell8|(X4=cell17|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X4=cell19|(X3=cell8|(X4=cell17|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))))&((((X3=cell10|(X3=cell9|(X4=cell18|(X4=cell17|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X3=cell9|(X4=cell18|(X4=cell17|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))&(((X3=cell10|(X4=cell19|(X4=cell18|(X4=cell17|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3))&((X4=cell20|(X4=cell19|(X4=cell18|(X4=cell17|(X4=cell16|epred1_2(X4,X3))))))|~epred2_2(X4,X3)))))))&((((((~epred1_2(X6,X5)|epred2_2(X6,X5))&((X5!=cell6|X6!=cell16)|epred2_2(X6,X5)))&((X5!=cell7|X6!=cell17)|epred2_2(X6,X5)))&((X5!=cell8|X6!=cell18)|epred2_2(X6,X5)))&((X5!=cell9|X6!=cell19)|epred2_2(X6,X5)))&((X5!=cell10|X6!=cell20)|epred2_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])])])).
fof(c_0_51, plain, (![X4]:![X3]:(epred26_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45)))), introduced(definition)).
fof(c_0_52, plain, (![X4]:![X3]:(epred25_2(X3,X4)<=>(((((epred24_2(X3,X4)|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39)))), inference(apply_def,[status(thm)],[c_0_39, c_0_29])).
fof(c_0_53, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell33|(X5=cell32|(X5=cell31|(X5=cell29|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X5=cell32|(X5=cell31|(X5=cell29|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X5=cell31|(X5=cell29|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X6=cell33|(X5=cell31|(X5=cell29|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))&((((X5=cell33|(X5=cell32|(X6=cell32|(X5=cell29|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X5=cell32|(X6=cell32|(X5=cell29|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X6=cell32|(X5=cell29|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X6=cell33|(X6=cell32|(X5=cell29|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))))&(((((X5=cell33|(X5=cell32|(X5=cell31|(X6=cell30|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X5=cell32|(X5=cell31|(X6=cell30|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X5=cell31|(X6=cell30|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X6=cell33|(X5=cell31|(X6=cell30|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))&((((X5=cell33|(X5=cell32|(X6=cell32|(X6=cell30|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X5=cell32|(X6=cell32|(X6=cell30|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X6=cell32|(X6=cell30|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X6=cell33|(X6=cell32|(X6=cell30|(X5=cell28|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))))&((((((X5=cell33|(X5=cell32|(X5=cell31|(X5=cell29|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X5=cell32|(X5=cell31|(X5=cell29|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X5=cell31|(X5=cell29|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X6=cell33|(X5=cell31|(X5=cell29|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))&((((X5=cell33|(X5=cell32|(X6=cell32|(X5=cell29|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X5=cell32|(X6=cell32|(X5=cell29|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X6=cell32|(X5=cell29|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X6=cell33|(X6=cell32|(X5=cell29|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))))&(((((X5=cell33|(X5=cell32|(X5=cell31|(X6=cell30|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X5=cell32|(X5=cell31|(X6=cell30|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X5=cell31|(X6=cell30|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X6=cell33|(X5=cell31|(X6=cell30|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))))&((((X5=cell33|(X5=cell32|(X6=cell32|(X6=cell30|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X5=cell32|(X6=cell32|(X6=cell30|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))&(((X5=cell33|(X6=cell33|(X6=cell32|(X6=cell30|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5))&((X6=cell34|(X6=cell33|(X6=cell32|(X6=cell30|(X6=cell29|epred23_2(X6,X5))))))|~epred24_2(X6,X5)))))))&((((((~epred23_2(X8,X7)|epred24_2(X8,X7))&((X7!=cell28|X8!=cell29)|epred24_2(X8,X7)))&((X7!=cell29|X8!=cell30)|epred24_2(X8,X7)))&((X7!=cell31|X8!=cell32)|epred24_2(X8,X7)))&((X7!=cell32|X8!=cell33)|epred24_2(X8,X7)))&((X7!=cell33|X8!=cell34)|epred24_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])])])).
cnf(c_0_54,plain,(epred23_2(X1,X2)|~epred22_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55,plain,(epred22_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_56,plain,(epred21_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_32, c_0_44])).
cnf(c_0_57,plain,(epred20_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_21, c_0_45])).
cnf(c_0_58,plain,(epred19_2(cell5,cell4)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_59,plain,(epred19_2(X1,cell5)|cell6!=X1), inference(er,[status(thm)],[c_0_47])).
fof(c_0_60, plain, (![X2]:![X1]:(epred6_2(X1,X2)<=>((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40)))), introduced(definition)).
fof(c_0_61, plain, (![X2]:![X1]:(epred5_2(X1,X2)<=>(((((epred4_2(X1,X2)|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35)))), inference(apply_def,[status(thm)],[c_0_48, c_0_25])).
cnf(c_0_62,plain,(epred4_2(X1,X2)|X1!=cell30|X2!=cell20), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_63, plain, (![X2]:![X1]:(epred3_2(X1,X2)<=>(((((epred2_2(X1,X2)|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25)))), inference(apply_def,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_64,plain,(epred2_2(X1,X2)|X1!=cell20|X2!=cell10), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_65, plain, (![X4]:![X3]:(epred27_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50)))), introduced(definition)).
fof(c_0_66, plain, (![X4]:![X3]:(epred26_2(X3,X4)<=>(((((epred25_2(X3,X4)|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45)))), inference(apply_def,[status(thm)],[c_0_51, c_0_39])).
fof(c_0_67, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell38|(X5=cell37|(X5=cell36|(X5=cell35|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X5=cell37|(X5=cell36|(X5=cell35|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X5=cell36|(X5=cell35|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X6=cell38|(X5=cell36|(X5=cell35|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))&((((X5=cell38|(X5=cell37|(X6=cell37|(X5=cell35|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X5=cell37|(X6=cell37|(X5=cell35|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X6=cell37|(X5=cell35|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X6=cell38|(X6=cell37|(X5=cell35|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))))&(((((X5=cell38|(X5=cell37|(X5=cell36|(X6=cell36|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X5=cell37|(X5=cell36|(X6=cell36|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X5=cell36|(X6=cell36|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X6=cell38|(X5=cell36|(X6=cell36|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))&((((X5=cell38|(X5=cell37|(X6=cell37|(X6=cell36|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X5=cell37|(X6=cell37|(X6=cell36|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X6=cell37|(X6=cell36|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X6=cell38|(X6=cell37|(X6=cell36|(X5=cell34|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))))&((((((X5=cell38|(X5=cell37|(X5=cell36|(X5=cell35|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X5=cell37|(X5=cell36|(X5=cell35|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X5=cell36|(X5=cell35|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X6=cell38|(X5=cell36|(X5=cell35|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))&((((X5=cell38|(X5=cell37|(X6=cell37|(X5=cell35|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X5=cell37|(X6=cell37|(X5=cell35|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X6=cell37|(X5=cell35|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X6=cell38|(X6=cell37|(X5=cell35|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))))&(((((X5=cell38|(X5=cell37|(X5=cell36|(X6=cell36|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X5=cell37|(X5=cell36|(X6=cell36|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X5=cell36|(X6=cell36|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X6=cell38|(X5=cell36|(X6=cell36|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))))&((((X5=cell38|(X5=cell37|(X6=cell37|(X6=cell36|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X5=cell37|(X6=cell37|(X6=cell36|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))&(((X5=cell38|(X6=cell38|(X6=cell37|(X6=cell36|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5))&((X6=cell39|(X6=cell38|(X6=cell37|(X6=cell36|(X6=cell35|epred24_2(X6,X5))))))|~epred25_2(X6,X5)))))))&((((((~epred24_2(X8,X7)|epred25_2(X8,X7))&((X7!=cell34|X8!=cell35)|epred25_2(X8,X7)))&((X7!=cell35|X8!=cell36)|epred25_2(X8,X7)))&((X7!=cell36|X8!=cell37)|epred25_2(X8,X7)))&((X7!=cell37|X8!=cell38)|epred25_2(X8,X7)))&((X7!=cell38|X8!=cell39)|epred25_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_52])])])])])).
cnf(c_0_68,plain,(epred24_2(X1,X2)|~epred23_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_69,plain,(epred23_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_70,plain,(epred22_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_42, c_0_56])).
cnf(c_0_71,plain,(epred21_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_32, c_0_57])).
cnf(c_0_72,plain,(epred20_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_21, c_0_58])).
cnf(c_0_73,plain,(epred19_2(cell6,cell5)), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_74,plain,(epred20_2(X1,X2)|X1!=cell7|X2!=cell6), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_75, plain, (![X2]:![X1]:(epred7_2(X1,X2)<=>(((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45)))), introduced(definition)).
fof(c_0_76, plain, (![X2]:![X1]:(epred6_2(X1,X2)<=>(((((epred5_2(X1,X2)|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40)))), inference(apply_def,[status(thm)],[c_0_60, c_0_48])).
fof(c_0_77, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell25|(X3=cell24|(X3=cell23|(X3=cell22|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X3=cell24|(X3=cell23|(X3=cell22|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))&(((X3=cell25|(X4=cell34|(X3=cell23|(X3=cell22|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X4=cell34|(X3=cell23|(X3=cell22|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))))&((((X3=cell25|(X3=cell24|(X4=cell33|(X3=cell22|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X3=cell24|(X4=cell33|(X3=cell22|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))&(((X3=cell25|(X4=cell34|(X4=cell33|(X3=cell22|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X4=cell34|(X4=cell33|(X3=cell22|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))))&(((((X3=cell25|(X3=cell24|(X3=cell23|(X4=cell32|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X3=cell24|(X3=cell23|(X4=cell32|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))&(((X3=cell25|(X4=cell34|(X3=cell23|(X4=cell32|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X4=cell34|(X3=cell23|(X4=cell32|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))))&((((X3=cell25|(X3=cell24|(X4=cell33|(X4=cell32|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X3=cell24|(X4=cell33|(X4=cell32|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))&(((X3=cell25|(X4=cell34|(X4=cell33|(X4=cell32|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X4=cell34|(X4=cell33|(X4=cell32|(X3=cell21|epred4_2(X4,X3))))))|~epred5_2(X4,X3))))))&((((((X3=cell25|(X3=cell24|(X3=cell23|(X3=cell22|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X3=cell24|(X3=cell23|(X3=cell22|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))&(((X3=cell25|(X4=cell34|(X3=cell23|(X3=cell22|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X4=cell34|(X3=cell23|(X3=cell22|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))))&((((X3=cell25|(X3=cell24|(X4=cell33|(X3=cell22|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X3=cell24|(X4=cell33|(X3=cell22|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))&(((X3=cell25|(X4=cell34|(X4=cell33|(X3=cell22|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X4=cell34|(X4=cell33|(X3=cell22|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))))&(((((X3=cell25|(X3=cell24|(X3=cell23|(X4=cell32|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X3=cell24|(X3=cell23|(X4=cell32|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))&(((X3=cell25|(X4=cell34|(X3=cell23|(X4=cell32|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X4=cell34|(X3=cell23|(X4=cell32|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))))&((((X3=cell25|(X3=cell24|(X4=cell33|(X4=cell32|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X3=cell24|(X4=cell33|(X4=cell32|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))&(((X3=cell25|(X4=cell34|(X4=cell33|(X4=cell32|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3))&((X4=cell35|(X4=cell34|(X4=cell33|(X4=cell32|(X4=cell31|epred4_2(X4,X3))))))|~epred5_2(X4,X3)))))))&((((((~epred4_2(X6,X5)|epred5_2(X6,X5))&((X5!=cell21|X6!=cell31)|epred5_2(X6,X5)))&((X5!=cell22|X6!=cell32)|epred5_2(X6,X5)))&((X5!=cell23|X6!=cell33)|epred5_2(X6,X5)))&((X5!=cell24|X6!=cell34)|epred5_2(X6,X5)))&((X5!=cell25|X6!=cell35)|epred5_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_61])])])])])).
cnf(c_0_78,plain,(epred4_2(X1,cell20)|cell30!=X1), inference(er,[status(thm)],[c_0_62])).
fof(c_0_79, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell15|(X3=cell14|(X3=cell13|(X3=cell12|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X3=cell14|(X3=cell13|(X3=cell12|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))&(((X3=cell15|(X4=cell24|(X3=cell13|(X3=cell12|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X4=cell24|(X3=cell13|(X3=cell12|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))))&((((X3=cell15|(X3=cell14|(X4=cell23|(X3=cell12|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X3=cell14|(X4=cell23|(X3=cell12|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))&(((X3=cell15|(X4=cell24|(X4=cell23|(X3=cell12|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X4=cell24|(X4=cell23|(X3=cell12|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))))&(((((X3=cell15|(X3=cell14|(X3=cell13|(X4=cell22|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X3=cell14|(X3=cell13|(X4=cell22|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))&(((X3=cell15|(X4=cell24|(X3=cell13|(X4=cell22|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X4=cell24|(X3=cell13|(X4=cell22|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))))&((((X3=cell15|(X3=cell14|(X4=cell23|(X4=cell22|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X3=cell14|(X4=cell23|(X4=cell22|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))&(((X3=cell15|(X4=cell24|(X4=cell23|(X4=cell22|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X4=cell24|(X4=cell23|(X4=cell22|(X3=cell11|epred2_2(X4,X3))))))|~epred3_2(X4,X3))))))&((((((X3=cell15|(X3=cell14|(X3=cell13|(X3=cell12|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X3=cell14|(X3=cell13|(X3=cell12|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))&(((X3=cell15|(X4=cell24|(X3=cell13|(X3=cell12|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X4=cell24|(X3=cell13|(X3=cell12|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))))&((((X3=cell15|(X3=cell14|(X4=cell23|(X3=cell12|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X3=cell14|(X4=cell23|(X3=cell12|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))&(((X3=cell15|(X4=cell24|(X4=cell23|(X3=cell12|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X4=cell24|(X4=cell23|(X3=cell12|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))))&(((((X3=cell15|(X3=cell14|(X3=cell13|(X4=cell22|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X3=cell14|(X3=cell13|(X4=cell22|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))&(((X3=cell15|(X4=cell24|(X3=cell13|(X4=cell22|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X4=cell24|(X3=cell13|(X4=cell22|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))))&((((X3=cell15|(X3=cell14|(X4=cell23|(X4=cell22|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X3=cell14|(X4=cell23|(X4=cell22|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))&(((X3=cell15|(X4=cell24|(X4=cell23|(X4=cell22|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3))&((X4=cell25|(X4=cell24|(X4=cell23|(X4=cell22|(X4=cell21|epred2_2(X4,X3))))))|~epred3_2(X4,X3)))))))&((((((~epred2_2(X6,X5)|epred3_2(X6,X5))&((X5!=cell11|X6!=cell21)|epred3_2(X6,X5)))&((X5!=cell12|X6!=cell22)|epred3_2(X6,X5)))&((X5!=cell13|X6!=cell23)|epred3_2(X6,X5)))&((X5!=cell14|X6!=cell24)|epred3_2(X6,X5)))&((X5!=cell15|X6!=cell25)|epred3_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_63])])])])])).
cnf(c_0_80,plain,(epred2_2(X1,cell10)|cell20!=X1), inference(er,[status(thm)],[c_0_64])).
fof(c_0_81, plain, (![X4]:![X3]:(epred28_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56)))), introduced(definition)).
fof(c_0_82, plain, (![X4]:![X3]:(epred27_2(X3,X4)<=>(((((epred26_2(X3,X4)|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50)))), inference(apply_def,[status(thm)],[c_0_65, c_0_51])).
fof(c_0_83, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell44|(X5=cell43|(X5=cell42|(X5=cell41|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X5=cell43|(X5=cell42|(X5=cell41|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X5=cell42|(X5=cell41|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell42|(X5=cell41|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))&((((X5=cell44|(X5=cell43|(X6=cell43|(X5=cell41|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X5=cell43|(X6=cell43|(X5=cell41|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X6=cell43|(X5=cell41|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X5=cell41|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))))&(((((X5=cell44|(X5=cell43|(X5=cell42|(X6=cell42|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X5=cell43|(X5=cell42|(X6=cell42|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X5=cell42|(X6=cell42|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell42|(X6=cell42|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))&((((X5=cell44|(X5=cell43|(X6=cell43|(X6=cell42|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X5=cell43|(X6=cell43|(X6=cell42|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X6=cell43|(X6=cell42|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X6=cell42|(X5=cell39|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))))&((((((X5=cell44|(X5=cell43|(X5=cell42|(X5=cell41|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X5=cell43|(X5=cell42|(X5=cell41|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X5=cell42|(X5=cell41|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell42|(X5=cell41|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))&((((X5=cell44|(X5=cell43|(X6=cell43|(X5=cell41|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X5=cell43|(X6=cell43|(X5=cell41|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X6=cell43|(X5=cell41|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X5=cell41|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))))&(((((X5=cell44|(X5=cell43|(X5=cell42|(X6=cell42|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X5=cell43|(X5=cell42|(X6=cell42|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X5=cell42|(X6=cell42|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X6=cell44|(X5=cell42|(X6=cell42|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))))&((((X5=cell44|(X5=cell43|(X6=cell43|(X6=cell42|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X5=cell43|(X6=cell43|(X6=cell42|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))&(((X5=cell44|(X6=cell44|(X6=cell43|(X6=cell42|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5))&((X6=cell45|(X6=cell44|(X6=cell43|(X6=cell42|(X6=cell40|epred25_2(X6,X5))))))|~epred26_2(X6,X5)))))))&((((((~epred25_2(X8,X7)|epred26_2(X8,X7))&((X7!=cell39|X8!=cell40)|epred26_2(X8,X7)))&((X7!=cell41|X8!=cell42)|epred26_2(X8,X7)))&((X7!=cell42|X8!=cell43)|epred26_2(X8,X7)))&((X7!=cell43|X8!=cell44)|epred26_2(X8,X7)))&((X7!=cell44|X8!=cell45)|epred26_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_66])])])])])).
cnf(c_0_84,plain,(epred25_2(X1,X2)|~epred24_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_85,plain,(epred24_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_86,plain,(epred23_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_54, c_0_70])).
cnf(c_0_87,plain,(epred22_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_42, c_0_71])).
cnf(c_0_88,plain,(epred21_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_32, c_0_72])).
cnf(c_0_89,plain,(epred20_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_21, c_0_73])).
cnf(c_0_90,plain,(epred20_2(X1,cell6)|cell7!=X1), inference(er,[status(thm)],[c_0_74])).
cnf(c_0_91,plain,(epred20_2(X1,X2)|X1!=cell8|X2!=cell7), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_92, plain, (![X2]:![X1]:(epred8_2(X1,X2)<=>((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50)))), introduced(definition)).
fof(c_0_93, plain, (![X2]:![X1]:(epred7_2(X1,X2)<=>(((((epred6_2(X1,X2)|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45)))), inference(apply_def,[status(thm)],[c_0_75, c_0_60])).
fof(c_0_94, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell30|(X3=cell29|(X3=cell28|(X3=cell27|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X3=cell29|(X3=cell28|(X3=cell27|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))&(((X3=cell30|(X4=cell39|(X3=cell28|(X3=cell27|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X4=cell39|(X3=cell28|(X3=cell27|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))))&((((X3=cell30|(X3=cell29|(X4=cell38|(X3=cell27|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X3=cell29|(X4=cell38|(X3=cell27|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))&(((X3=cell30|(X4=cell39|(X4=cell38|(X3=cell27|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X4=cell39|(X4=cell38|(X3=cell27|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))))&(((((X3=cell30|(X3=cell29|(X3=cell28|(X4=cell37|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X3=cell29|(X3=cell28|(X4=cell37|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))&(((X3=cell30|(X4=cell39|(X3=cell28|(X4=cell37|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X4=cell39|(X3=cell28|(X4=cell37|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))))&((((X3=cell30|(X3=cell29|(X4=cell38|(X4=cell37|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X3=cell29|(X4=cell38|(X4=cell37|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))&(((X3=cell30|(X4=cell39|(X4=cell38|(X4=cell37|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X4=cell39|(X4=cell38|(X4=cell37|(X3=cell26|epred5_2(X4,X3))))))|~epred6_2(X4,X3))))))&((((((X3=cell30|(X3=cell29|(X3=cell28|(X3=cell27|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X3=cell29|(X3=cell28|(X3=cell27|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))&(((X3=cell30|(X4=cell39|(X3=cell28|(X3=cell27|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X4=cell39|(X3=cell28|(X3=cell27|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))))&((((X3=cell30|(X3=cell29|(X4=cell38|(X3=cell27|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X3=cell29|(X4=cell38|(X3=cell27|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))&(((X3=cell30|(X4=cell39|(X4=cell38|(X3=cell27|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X4=cell39|(X4=cell38|(X3=cell27|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))))&(((((X3=cell30|(X3=cell29|(X3=cell28|(X4=cell37|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X3=cell29|(X3=cell28|(X4=cell37|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))&(((X3=cell30|(X4=cell39|(X3=cell28|(X4=cell37|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X4=cell39|(X3=cell28|(X4=cell37|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))))&((((X3=cell30|(X3=cell29|(X4=cell38|(X4=cell37|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X3=cell29|(X4=cell38|(X4=cell37|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))&(((X3=cell30|(X4=cell39|(X4=cell38|(X4=cell37|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3))&((X4=cell40|(X4=cell39|(X4=cell38|(X4=cell37|(X4=cell36|epred5_2(X4,X3))))))|~epred6_2(X4,X3)))))))&((((((~epred5_2(X6,X5)|epred6_2(X6,X5))&((X5!=cell26|X6!=cell36)|epred6_2(X6,X5)))&((X5!=cell27|X6!=cell37)|epred6_2(X6,X5)))&((X5!=cell28|X6!=cell38)|epred6_2(X6,X5)))&((X5!=cell29|X6!=cell39)|epred6_2(X6,X5)))&((X5!=cell30|X6!=cell40)|epred6_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_76])])])])])).
cnf(c_0_95,plain,(epred5_2(X1,X2)|~epred4_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_96,plain,(epred4_2(cell30,cell20)), inference(er,[status(thm)],[c_0_78])).
cnf(c_0_97,plain,(epred3_2(X1,X2)|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_98,plain,(epred2_2(cell20,cell10)), inference(er,[status(thm)],[c_0_80])).
fof(c_0_99, plain, (![X4]:![X3]:(epred29_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62)))), introduced(definition)).
fof(c_0_100, plain, (![X4]:![X3]:(epred28_2(X3,X4)<=>(((((epred27_2(X3,X4)|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56)))), inference(apply_def,[status(thm)],[c_0_81, c_0_65])).
fof(c_0_101, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell49|(X5=cell48|(X5=cell47|(X5=cell46|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X5=cell48|(X5=cell47|(X5=cell46|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X5=cell47|(X5=cell46|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell47|(X5=cell46|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))&((((X5=cell49|(X5=cell48|(X6=cell48|(X5=cell46|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X5=cell48|(X6=cell48|(X5=cell46|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X6=cell48|(X5=cell46|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X5=cell46|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))))&(((((X5=cell49|(X5=cell48|(X5=cell47|(X6=cell47|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X5=cell48|(X5=cell47|(X6=cell47|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X5=cell47|(X6=cell47|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell47|(X6=cell47|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))&((((X5=cell49|(X5=cell48|(X6=cell48|(X6=cell47|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X5=cell48|(X6=cell48|(X6=cell47|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X6=cell48|(X6=cell47|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X6=cell47|(X5=cell45|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))))&((((((X5=cell49|(X5=cell48|(X5=cell47|(X5=cell46|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X5=cell48|(X5=cell47|(X5=cell46|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X5=cell47|(X5=cell46|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell47|(X5=cell46|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))&((((X5=cell49|(X5=cell48|(X6=cell48|(X5=cell46|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X5=cell48|(X6=cell48|(X5=cell46|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X6=cell48|(X5=cell46|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X5=cell46|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))))&(((((X5=cell49|(X5=cell48|(X5=cell47|(X6=cell47|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X5=cell48|(X5=cell47|(X6=cell47|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X5=cell47|(X6=cell47|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X6=cell49|(X5=cell47|(X6=cell47|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))))&((((X5=cell49|(X5=cell48|(X6=cell48|(X6=cell47|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X5=cell48|(X6=cell48|(X6=cell47|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))&(((X5=cell49|(X6=cell49|(X6=cell48|(X6=cell47|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5))&((X6=cell50|(X6=cell49|(X6=cell48|(X6=cell47|(X6=cell46|epred26_2(X6,X5))))))|~epred27_2(X6,X5)))))))&((((((~epred26_2(X8,X7)|epred27_2(X8,X7))&((X7!=cell45|X8!=cell46)|epred27_2(X8,X7)))&((X7!=cell46|X8!=cell47)|epred27_2(X8,X7)))&((X7!=cell47|X8!=cell48)|epred27_2(X8,X7)))&((X7!=cell48|X8!=cell49)|epred27_2(X8,X7)))&((X7!=cell49|X8!=cell50)|epred27_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_82])])])])])).
cnf(c_0_102,plain,(epred26_2(X1,X2)|~epred25_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_103,plain,(epred25_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_104,plain,(epred24_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_68, c_0_86])).
cnf(c_0_105,plain,(epred23_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_54, c_0_87])).
cnf(c_0_106,plain,(epred22_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_42, c_0_88])).
cnf(c_0_107,plain,(epred21_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_32, c_0_89])).
cnf(c_0_108,plain,(epred20_2(cell7,cell6)), inference(er,[status(thm)],[c_0_90])).
cnf(c_0_109,plain,(epred20_2(X1,cell7)|cell8!=X1), inference(er,[status(thm)],[c_0_91])).
cnf(c_0_110,plain,(epred20_2(X1,X2)|X1!=cell9|X2!=cell8), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_111, plain, (![X2]:![X1]:(epred9_2(X1,X2)<=>(((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55)))), introduced(definition)).
fof(c_0_112, plain, (![X2]:![X1]:(epred8_2(X1,X2)<=>(((((epred7_2(X1,X2)|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50)))), inference(apply_def,[status(thm)],[c_0_92, c_0_75])).
fof(c_0_113, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell35|(X3=cell34|(X3=cell33|(X3=cell32|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X3=cell34|(X3=cell33|(X3=cell32|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))&(((X3=cell35|(X4=cell44|(X3=cell33|(X3=cell32|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X4=cell44|(X3=cell33|(X3=cell32|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))))&((((X3=cell35|(X3=cell34|(X4=cell43|(X3=cell32|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X3=cell34|(X4=cell43|(X3=cell32|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))&(((X3=cell35|(X4=cell44|(X4=cell43|(X3=cell32|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X4=cell44|(X4=cell43|(X3=cell32|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))))&(((((X3=cell35|(X3=cell34|(X3=cell33|(X4=cell42|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X3=cell34|(X3=cell33|(X4=cell42|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))&(((X3=cell35|(X4=cell44|(X3=cell33|(X4=cell42|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X4=cell44|(X3=cell33|(X4=cell42|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))))&((((X3=cell35|(X3=cell34|(X4=cell43|(X4=cell42|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X3=cell34|(X4=cell43|(X4=cell42|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))&(((X3=cell35|(X4=cell44|(X4=cell43|(X4=cell42|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X4=cell44|(X4=cell43|(X4=cell42|(X3=cell31|epred6_2(X4,X3))))))|~epred7_2(X4,X3))))))&((((((X3=cell35|(X3=cell34|(X3=cell33|(X3=cell32|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X3=cell34|(X3=cell33|(X3=cell32|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))&(((X3=cell35|(X4=cell44|(X3=cell33|(X3=cell32|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X4=cell44|(X3=cell33|(X3=cell32|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))))&((((X3=cell35|(X3=cell34|(X4=cell43|(X3=cell32|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X3=cell34|(X4=cell43|(X3=cell32|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))&(((X3=cell35|(X4=cell44|(X4=cell43|(X3=cell32|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X4=cell44|(X4=cell43|(X3=cell32|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))))&(((((X3=cell35|(X3=cell34|(X3=cell33|(X4=cell42|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X3=cell34|(X3=cell33|(X4=cell42|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))&(((X3=cell35|(X4=cell44|(X3=cell33|(X4=cell42|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X4=cell44|(X3=cell33|(X4=cell42|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))))&((((X3=cell35|(X3=cell34|(X4=cell43|(X4=cell42|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X3=cell34|(X4=cell43|(X4=cell42|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))&(((X3=cell35|(X4=cell44|(X4=cell43|(X4=cell42|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3))&((X4=cell45|(X4=cell44|(X4=cell43|(X4=cell42|(X4=cell41|epred6_2(X4,X3))))))|~epred7_2(X4,X3)))))))&((((((~epred6_2(X6,X5)|epred7_2(X6,X5))&((X5!=cell31|X6!=cell41)|epred7_2(X6,X5)))&((X5!=cell32|X6!=cell42)|epred7_2(X6,X5)))&((X5!=cell33|X6!=cell43)|epred7_2(X6,X5)))&((X5!=cell34|X6!=cell44)|epred7_2(X6,X5)))&((X5!=cell35|X6!=cell45)|epred7_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_93])])])])])).
cnf(c_0_114,plain,(epred6_2(X1,X2)|~epred5_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_115,plain,(epred5_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_116,plain,(epred6_2(X1,X2)|X1!=cell40|X2!=cell30), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_117,plain,(epred4_2(X1,X2)|~epred3_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_118,plain,(epred3_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
fof(c_0_119, plain, (![X4]:![X3]:(epred30_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67)))), introduced(definition)).
fof(c_0_120, plain, (![X4]:![X3]:(epred29_2(X3,X4)<=>(((((epred28_2(X3,X4)|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62)))), inference(apply_def,[status(thm)],[c_0_99, c_0_81])).
fof(c_0_121, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell55|(X5=cell54|(X5=cell53|(X5=cell52|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X5=cell54|(X5=cell53|(X5=cell52|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X5=cell53|(X5=cell52|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X6=cell55|(X5=cell53|(X5=cell52|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell54|(X5=cell52|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X5=cell54|(X6=cell54|(X5=cell52|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X6=cell54|(X5=cell52|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X6=cell55|(X6=cell54|(X5=cell52|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))))&(((((X5=cell55|(X5=cell54|(X5=cell53|(X6=cell53|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X5=cell54|(X5=cell53|(X6=cell53|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X5=cell53|(X6=cell53|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X6=cell55|(X5=cell53|(X6=cell53|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell54|(X6=cell53|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X5=cell54|(X6=cell54|(X6=cell53|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X6=cell54|(X6=cell53|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X6=cell55|(X6=cell54|(X6=cell53|(X5=cell51|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))))&((((((X5=cell55|(X5=cell54|(X5=cell53|(X5=cell52|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X5=cell54|(X5=cell53|(X5=cell52|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X5=cell53|(X5=cell52|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X6=cell55|(X5=cell53|(X5=cell52|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell54|(X5=cell52|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X5=cell54|(X6=cell54|(X5=cell52|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X6=cell54|(X5=cell52|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X6=cell55|(X6=cell54|(X5=cell52|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))))&(((((X5=cell55|(X5=cell54|(X5=cell53|(X6=cell53|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X5=cell54|(X5=cell53|(X6=cell53|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X5=cell53|(X6=cell53|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X6=cell55|(X5=cell53|(X6=cell53|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))))&((((X5=cell55|(X5=cell54|(X6=cell54|(X6=cell53|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X5=cell54|(X6=cell54|(X6=cell53|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))&(((X5=cell55|(X6=cell55|(X6=cell54|(X6=cell53|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5))&((X6=cell56|(X6=cell55|(X6=cell54|(X6=cell53|(X6=cell52|epred27_2(X6,X5))))))|~epred28_2(X6,X5)))))))&((((((~epred27_2(X8,X7)|epred28_2(X8,X7))&((X7!=cell51|X8!=cell52)|epred28_2(X8,X7)))&((X7!=cell52|X8!=cell53)|epred28_2(X8,X7)))&((X7!=cell53|X8!=cell54)|epred28_2(X8,X7)))&((X7!=cell54|X8!=cell55)|epred28_2(X8,X7)))&((X7!=cell55|X8!=cell56)|epred28_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_100])])])])])).
cnf(c_0_122,plain,(epred27_2(X1,X2)|~epred26_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_123,plain,(epred26_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_124,plain,(epred25_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_84, c_0_104])).
cnf(c_0_125,plain,(epred24_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_68, c_0_105])).
cnf(c_0_126,plain,(epred23_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_54, c_0_106])).
cnf(c_0_127,plain,(epred22_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_42, c_0_107])).
cnf(c_0_128,plain,(epred21_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_32, c_0_108])).
cnf(c_0_129,plain,(epred20_2(cell8,cell7)), inference(er,[status(thm)],[c_0_109])).
cnf(c_0_130,plain,(epred20_2(X1,cell8)|cell9!=X1), inference(er,[status(thm)],[c_0_110])).
cnf(c_0_131,plain,(epred20_2(X1,X2)|X1!=cell10|X2!=cell9), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_132, plain, (![X2]:![X1]:(epred10_2(X1,X2)<=>((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60)))), introduced(definition)).
fof(c_0_133, plain, (![X2]:![X1]:(epred9_2(X1,X2)<=>(((((epred8_2(X1,X2)|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55)))), inference(apply_def,[status(thm)],[c_0_111, c_0_92])).
fof(c_0_134, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell40|(X3=cell39|(X3=cell38|(X3=cell37|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X3=cell39|(X3=cell38|(X3=cell37|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))&(((X3=cell40|(X4=cell49|(X3=cell38|(X3=cell37|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X4=cell49|(X3=cell38|(X3=cell37|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))))&((((X3=cell40|(X3=cell39|(X4=cell48|(X3=cell37|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X3=cell39|(X4=cell48|(X3=cell37|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))&(((X3=cell40|(X4=cell49|(X4=cell48|(X3=cell37|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X4=cell49|(X4=cell48|(X3=cell37|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))))&(((((X3=cell40|(X3=cell39|(X3=cell38|(X4=cell47|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X3=cell39|(X3=cell38|(X4=cell47|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))&(((X3=cell40|(X4=cell49|(X3=cell38|(X4=cell47|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X4=cell49|(X3=cell38|(X4=cell47|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))))&((((X3=cell40|(X3=cell39|(X4=cell48|(X4=cell47|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X3=cell39|(X4=cell48|(X4=cell47|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))&(((X3=cell40|(X4=cell49|(X4=cell48|(X4=cell47|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X4=cell49|(X4=cell48|(X4=cell47|(X3=cell36|epred7_2(X4,X3))))))|~epred8_2(X4,X3))))))&((((((X3=cell40|(X3=cell39|(X3=cell38|(X3=cell37|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X3=cell39|(X3=cell38|(X3=cell37|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))&(((X3=cell40|(X4=cell49|(X3=cell38|(X3=cell37|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X4=cell49|(X3=cell38|(X3=cell37|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))))&((((X3=cell40|(X3=cell39|(X4=cell48|(X3=cell37|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X3=cell39|(X4=cell48|(X3=cell37|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))&(((X3=cell40|(X4=cell49|(X4=cell48|(X3=cell37|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X4=cell49|(X4=cell48|(X3=cell37|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))))&(((((X3=cell40|(X3=cell39|(X3=cell38|(X4=cell47|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X3=cell39|(X3=cell38|(X4=cell47|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))&(((X3=cell40|(X4=cell49|(X3=cell38|(X4=cell47|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X4=cell49|(X3=cell38|(X4=cell47|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))))&((((X3=cell40|(X3=cell39|(X4=cell48|(X4=cell47|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X3=cell39|(X4=cell48|(X4=cell47|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))&(((X3=cell40|(X4=cell49|(X4=cell48|(X4=cell47|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3))&((X4=cell50|(X4=cell49|(X4=cell48|(X4=cell47|(X4=cell46|epred7_2(X4,X3))))))|~epred8_2(X4,X3)))))))&((((((~epred7_2(X6,X5)|epred8_2(X6,X5))&((X5!=cell36|X6!=cell46)|epred8_2(X6,X5)))&((X5!=cell37|X6!=cell47)|epred8_2(X6,X5)))&((X5!=cell38|X6!=cell48)|epred8_2(X6,X5)))&((X5!=cell39|X6!=cell49)|epred8_2(X6,X5)))&((X5!=cell40|X6!=cell50)|epred8_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_112])])])])])).
cnf(c_0_135,plain,(epred7_2(X1,X2)|~epred6_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_136,plain,(epred6_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_137,plain,(epred6_2(X1,cell30)|cell40!=X1), inference(er,[status(thm)],[c_0_116])).
cnf(c_0_138,plain,(epred4_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_117, c_0_118])).
fof(c_0_139, plain, (![X4]:![X3]:(epred31_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73)))), introduced(definition)).
fof(c_0_140, plain, (![X4]:![X3]:(epred30_2(X3,X4)<=>(((((epred29_2(X3,X4)|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67)))), inference(apply_def,[status(thm)],[c_0_119, c_0_99])).
fof(c_0_141, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell61|(X5=cell59|(X5=cell58|(X5=cell57|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X5=cell59|(X5=cell58|(X5=cell57|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X5=cell58|(X5=cell57|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X6=cell60|(X5=cell58|(X5=cell57|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))&((((X5=cell61|(X5=cell59|(X6=cell59|(X5=cell57|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X5=cell59|(X6=cell59|(X5=cell57|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X6=cell59|(X5=cell57|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X6=cell60|(X6=cell59|(X5=cell57|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))))&(((((X5=cell61|(X5=cell59|(X5=cell58|(X6=cell58|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X5=cell59|(X5=cell58|(X6=cell58|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X5=cell58|(X6=cell58|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X6=cell60|(X5=cell58|(X6=cell58|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))&((((X5=cell61|(X5=cell59|(X6=cell59|(X6=cell58|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X5=cell59|(X6=cell59|(X6=cell58|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X6=cell59|(X6=cell58|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X6=cell60|(X6=cell59|(X6=cell58|(X5=cell56|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))))&((((((X5=cell61|(X5=cell59|(X5=cell58|(X5=cell57|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X5=cell59|(X5=cell58|(X5=cell57|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X5=cell58|(X5=cell57|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X6=cell60|(X5=cell58|(X5=cell57|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))&((((X5=cell61|(X5=cell59|(X6=cell59|(X5=cell57|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X5=cell59|(X6=cell59|(X5=cell57|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X6=cell59|(X5=cell57|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X6=cell60|(X6=cell59|(X5=cell57|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))))&(((((X5=cell61|(X5=cell59|(X5=cell58|(X6=cell58|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X5=cell59|(X5=cell58|(X6=cell58|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X5=cell58|(X6=cell58|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X6=cell60|(X5=cell58|(X6=cell58|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))))&((((X5=cell61|(X5=cell59|(X6=cell59|(X6=cell58|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X5=cell59|(X6=cell59|(X6=cell58|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))&(((X5=cell61|(X6=cell60|(X6=cell59|(X6=cell58|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5))&((X6=cell62|(X6=cell60|(X6=cell59|(X6=cell58|(X6=cell57|epred28_2(X6,X5))))))|~epred29_2(X6,X5)))))))&((((((~epred28_2(X8,X7)|epred29_2(X8,X7))&((X7!=cell56|X8!=cell57)|epred29_2(X8,X7)))&((X7!=cell57|X8!=cell58)|epred29_2(X8,X7)))&((X7!=cell58|X8!=cell59)|epred29_2(X8,X7)))&((X7!=cell59|X8!=cell60)|epred29_2(X8,X7)))&((X7!=cell61|X8!=cell62)|epred29_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_120])])])])])).
cnf(c_0_142,plain,(epred28_2(X1,X2)|~epred27_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_143,plain,(epred27_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_144,plain,(epred26_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_102, c_0_124])).
cnf(c_0_145,plain,(epred25_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_84, c_0_125])).
cnf(c_0_146,plain,(epred24_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_68, c_0_126])).
cnf(c_0_147,plain,(epred23_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_54, c_0_127])).
cnf(c_0_148,plain,(epred22_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_42, c_0_128])).
cnf(c_0_149,plain,(epred21_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_32, c_0_129])).
cnf(c_0_150,plain,(epred20_2(cell9,cell8)), inference(er,[status(thm)],[c_0_130])).
cnf(c_0_151,plain,(epred20_2(X1,cell9)|cell10!=X1), inference(er,[status(thm)],[c_0_131])).
fof(c_0_152, plain, (![X2]:![X1]:(epred11_2(X1,X2)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60))|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65)))), introduced(definition)).
fof(c_0_153, plain, (![X2]:![X1]:(epred10_2(X1,X2)<=>(((((epred9_2(X1,X2)|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60)))), inference(apply_def,[status(thm)],[c_0_132, c_0_111])).
fof(c_0_154, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell45|(X3=cell44|(X3=cell43|(X3=cell42|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X3=cell44|(X3=cell43|(X3=cell42|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))&(((X3=cell45|(X4=cell54|(X3=cell43|(X3=cell42|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X4=cell54|(X3=cell43|(X3=cell42|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))))&((((X3=cell45|(X3=cell44|(X4=cell53|(X3=cell42|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X3=cell44|(X4=cell53|(X3=cell42|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))&(((X3=cell45|(X4=cell54|(X4=cell53|(X3=cell42|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X4=cell54|(X4=cell53|(X3=cell42|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))))&(((((X3=cell45|(X3=cell44|(X3=cell43|(X4=cell52|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X3=cell44|(X3=cell43|(X4=cell52|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))&(((X3=cell45|(X4=cell54|(X3=cell43|(X4=cell52|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X4=cell54|(X3=cell43|(X4=cell52|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))))&((((X3=cell45|(X3=cell44|(X4=cell53|(X4=cell52|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X3=cell44|(X4=cell53|(X4=cell52|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))&(((X3=cell45|(X4=cell54|(X4=cell53|(X4=cell52|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X4=cell54|(X4=cell53|(X4=cell52|(X3=cell41|epred8_2(X4,X3))))))|~epred9_2(X4,X3))))))&((((((X3=cell45|(X3=cell44|(X3=cell43|(X3=cell42|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X3=cell44|(X3=cell43|(X3=cell42|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))&(((X3=cell45|(X4=cell54|(X3=cell43|(X3=cell42|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X4=cell54|(X3=cell43|(X3=cell42|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))))&((((X3=cell45|(X3=cell44|(X4=cell53|(X3=cell42|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X3=cell44|(X4=cell53|(X3=cell42|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))&(((X3=cell45|(X4=cell54|(X4=cell53|(X3=cell42|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X4=cell54|(X4=cell53|(X3=cell42|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))))&(((((X3=cell45|(X3=cell44|(X3=cell43|(X4=cell52|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X3=cell44|(X3=cell43|(X4=cell52|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))&(((X3=cell45|(X4=cell54|(X3=cell43|(X4=cell52|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X4=cell54|(X3=cell43|(X4=cell52|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))))&((((X3=cell45|(X3=cell44|(X4=cell53|(X4=cell52|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X3=cell44|(X4=cell53|(X4=cell52|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))&(((X3=cell45|(X4=cell54|(X4=cell53|(X4=cell52|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3))&((X4=cell55|(X4=cell54|(X4=cell53|(X4=cell52|(X4=cell51|epred8_2(X4,X3))))))|~epred9_2(X4,X3)))))))&((((((~epred8_2(X6,X5)|epred9_2(X6,X5))&((X5!=cell41|X6!=cell51)|epred9_2(X6,X5)))&((X5!=cell42|X6!=cell52)|epred9_2(X6,X5)))&((X5!=cell43|X6!=cell53)|epred9_2(X6,X5)))&((X5!=cell44|X6!=cell54)|epred9_2(X6,X5)))&((X5!=cell45|X6!=cell55)|epred9_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_133])])])])])).
cnf(c_0_155,plain,(epred8_2(X1,X2)|~epred7_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_156,plain,(epred7_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_157,plain,(epred6_2(cell40,cell30)), inference(er,[status(thm)],[c_0_137])).
cnf(c_0_158,plain,(epred5_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_95, c_0_138])).
fof(c_0_159, plain, (![X4]:![X3]:(epred32_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78)))), introduced(definition)).
fof(c_0_160, plain, (![X4]:![X3]:(epred31_2(X3,X4)<=>(((((epred30_2(X3,X4)|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73)))), inference(apply_def,[status(thm)],[c_0_139, c_0_119])).
fof(c_0_161, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell66|(X5=cell65|(X5=cell64|(X5=cell63|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X5=cell65|(X5=cell64|(X5=cell63|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X5=cell64|(X5=cell63|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X6=cell66|(X5=cell64|(X5=cell63|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))&((((X5=cell66|(X5=cell65|(X6=cell65|(X5=cell63|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X5=cell65|(X6=cell65|(X5=cell63|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X6=cell65|(X5=cell63|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X6=cell66|(X6=cell65|(X5=cell63|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))))&(((((X5=cell66|(X5=cell65|(X5=cell64|(X6=cell64|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X5=cell65|(X5=cell64|(X6=cell64|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X5=cell64|(X6=cell64|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X6=cell66|(X5=cell64|(X6=cell64|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))&((((X5=cell66|(X5=cell65|(X6=cell65|(X6=cell64|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X5=cell65|(X6=cell65|(X6=cell64|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X6=cell65|(X6=cell64|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X6=cell66|(X6=cell65|(X6=cell64|(X5=cell62|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))))&((((((X5=cell66|(X5=cell65|(X5=cell64|(X5=cell63|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X5=cell65|(X5=cell64|(X5=cell63|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X5=cell64|(X5=cell63|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X6=cell66|(X5=cell64|(X5=cell63|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))&((((X5=cell66|(X5=cell65|(X6=cell65|(X5=cell63|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X5=cell65|(X6=cell65|(X5=cell63|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X6=cell65|(X5=cell63|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X6=cell66|(X6=cell65|(X5=cell63|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))))&(((((X5=cell66|(X5=cell65|(X5=cell64|(X6=cell64|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X5=cell65|(X5=cell64|(X6=cell64|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X5=cell64|(X6=cell64|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X6=cell66|(X5=cell64|(X6=cell64|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))))&((((X5=cell66|(X5=cell65|(X6=cell65|(X6=cell64|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X5=cell65|(X6=cell65|(X6=cell64|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))&(((X5=cell66|(X6=cell66|(X6=cell65|(X6=cell64|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5))&((X6=cell67|(X6=cell66|(X6=cell65|(X6=cell64|(X6=cell63|epred29_2(X6,X5))))))|~epred30_2(X6,X5)))))))&((((((~epred29_2(X8,X7)|epred30_2(X8,X7))&((X7!=cell62|X8!=cell63)|epred30_2(X8,X7)))&((X7!=cell63|X8!=cell64)|epred30_2(X8,X7)))&((X7!=cell64|X8!=cell65)|epred30_2(X8,X7)))&((X7!=cell65|X8!=cell66)|epred30_2(X8,X7)))&((X7!=cell66|X8!=cell67)|epred30_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_140])])])])])).
cnf(c_0_162,plain,(epred29_2(X1,X2)|~epred28_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_163,plain,(epred28_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_164,plain,(epred27_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_122, c_0_144])).
cnf(c_0_165,plain,(epred26_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_102, c_0_145])).
cnf(c_0_166,plain,(epred25_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_84, c_0_146])).
cnf(c_0_167,plain,(epred24_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_68, c_0_147])).
cnf(c_0_168,plain,(epred23_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_54, c_0_148])).
cnf(c_0_169,plain,(epred22_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_42, c_0_149])).
cnf(c_0_170,plain,(epred21_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_32, c_0_150])).
cnf(c_0_171,plain,(epred20_2(cell10,cell9)), inference(er,[status(thm)],[c_0_151])).
fof(c_0_172, plain, (![X2]:![X1]:(epred12_2(X1,X2)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60))|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65))|(X2=cell56&X1=cell66))|(X2=cell57&X1=cell67))|(X2=cell58&X1=cell68))|(X2=cell59&X1=cell69))|(X2=cell60&X1=cell70)))), introduced(definition)).
fof(c_0_173, plain, (![X2]:![X1]:(epred11_2(X1,X2)<=>(((((epred10_2(X1,X2)|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65)))), inference(apply_def,[status(thm)],[c_0_152, c_0_132])).
fof(c_0_174, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell50|(X3=cell49|(X3=cell48|(X3=cell47|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X3=cell49|(X3=cell48|(X3=cell47|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))&(((X3=cell50|(X4=cell59|(X3=cell48|(X3=cell47|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X4=cell59|(X3=cell48|(X3=cell47|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))))&((((X3=cell50|(X3=cell49|(X4=cell58|(X3=cell47|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X3=cell49|(X4=cell58|(X3=cell47|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))&(((X3=cell50|(X4=cell59|(X4=cell58|(X3=cell47|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X4=cell59|(X4=cell58|(X3=cell47|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))))&(((((X3=cell50|(X3=cell49|(X3=cell48|(X4=cell57|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X3=cell49|(X3=cell48|(X4=cell57|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))&(((X3=cell50|(X4=cell59|(X3=cell48|(X4=cell57|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X4=cell59|(X3=cell48|(X4=cell57|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))))&((((X3=cell50|(X3=cell49|(X4=cell58|(X4=cell57|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X3=cell49|(X4=cell58|(X4=cell57|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))&(((X3=cell50|(X4=cell59|(X4=cell58|(X4=cell57|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X4=cell59|(X4=cell58|(X4=cell57|(X3=cell46|epred9_2(X4,X3))))))|~epred10_2(X4,X3))))))&((((((X3=cell50|(X3=cell49|(X3=cell48|(X3=cell47|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X3=cell49|(X3=cell48|(X3=cell47|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))&(((X3=cell50|(X4=cell59|(X3=cell48|(X3=cell47|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X4=cell59|(X3=cell48|(X3=cell47|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))))&((((X3=cell50|(X3=cell49|(X4=cell58|(X3=cell47|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X3=cell49|(X4=cell58|(X3=cell47|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))&(((X3=cell50|(X4=cell59|(X4=cell58|(X3=cell47|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X4=cell59|(X4=cell58|(X3=cell47|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))))&(((((X3=cell50|(X3=cell49|(X3=cell48|(X4=cell57|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X3=cell49|(X3=cell48|(X4=cell57|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))&(((X3=cell50|(X4=cell59|(X3=cell48|(X4=cell57|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X4=cell59|(X3=cell48|(X4=cell57|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))))&((((X3=cell50|(X3=cell49|(X4=cell58|(X4=cell57|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X3=cell49|(X4=cell58|(X4=cell57|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))&(((X3=cell50|(X4=cell59|(X4=cell58|(X4=cell57|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3))&((X4=cell60|(X4=cell59|(X4=cell58|(X4=cell57|(X4=cell56|epred9_2(X4,X3))))))|~epred10_2(X4,X3)))))))&((((((~epred9_2(X6,X5)|epred10_2(X6,X5))&((X5!=cell46|X6!=cell56)|epred10_2(X6,X5)))&((X5!=cell47|X6!=cell57)|epred10_2(X6,X5)))&((X5!=cell48|X6!=cell58)|epred10_2(X6,X5)))&((X5!=cell49|X6!=cell59)|epred10_2(X6,X5)))&((X5!=cell50|X6!=cell60)|epred10_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_153])])])])])).
cnf(c_0_175,plain,(epred9_2(X1,X2)|~epred8_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_154])).
cnf(c_0_176,plain,(epred8_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_177,plain,(epred7_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_135, c_0_157])).
cnf(c_0_178,plain,(epred6_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_114, c_0_158])).
fof(c_0_179, plain, (![X4]:![X3]:(epred33_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84)))), introduced(definition)).
fof(c_0_180, plain, (![X4]:![X3]:(epred32_2(X3,X4)<=>(((((epred31_2(X3,X4)|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78)))), inference(apply_def,[status(thm)],[c_0_159, c_0_139])).
fof(c_0_181, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell72|(X5=cell71|(X5=cell69|(X5=cell68|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X5=cell71|(X5=cell69|(X5=cell68|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X5=cell69|(X5=cell68|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X6=cell72|(X5=cell69|(X5=cell68|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))&((((X5=cell72|(X5=cell71|(X6=cell70|(X5=cell68|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X5=cell71|(X6=cell70|(X5=cell68|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X6=cell70|(X5=cell68|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X6=cell72|(X6=cell70|(X5=cell68|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))))&(((((X5=cell72|(X5=cell71|(X5=cell69|(X6=cell69|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X5=cell71|(X5=cell69|(X6=cell69|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X5=cell69|(X6=cell69|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X6=cell72|(X5=cell69|(X6=cell69|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))&((((X5=cell72|(X5=cell71|(X6=cell70|(X6=cell69|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X5=cell71|(X6=cell70|(X6=cell69|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X6=cell70|(X6=cell69|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X6=cell72|(X6=cell70|(X6=cell69|(X5=cell67|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))))&((((((X5=cell72|(X5=cell71|(X5=cell69|(X5=cell68|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X5=cell71|(X5=cell69|(X5=cell68|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X5=cell69|(X5=cell68|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X6=cell72|(X5=cell69|(X5=cell68|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))&((((X5=cell72|(X5=cell71|(X6=cell70|(X5=cell68|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X5=cell71|(X6=cell70|(X5=cell68|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X6=cell70|(X5=cell68|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X6=cell72|(X6=cell70|(X5=cell68|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))))&(((((X5=cell72|(X5=cell71|(X5=cell69|(X6=cell69|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X5=cell71|(X5=cell69|(X6=cell69|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X5=cell69|(X6=cell69|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X6=cell72|(X5=cell69|(X6=cell69|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))))&((((X5=cell72|(X5=cell71|(X6=cell70|(X6=cell69|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X5=cell71|(X6=cell70|(X6=cell69|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))&(((X5=cell72|(X6=cell72|(X6=cell70|(X6=cell69|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5))&((X6=cell73|(X6=cell72|(X6=cell70|(X6=cell69|(X6=cell68|epred30_2(X6,X5))))))|~epred31_2(X6,X5)))))))&((((((~epred30_2(X8,X7)|epred31_2(X8,X7))&((X7!=cell67|X8!=cell68)|epred31_2(X8,X7)))&((X7!=cell68|X8!=cell69)|epred31_2(X8,X7)))&((X7!=cell69|X8!=cell70)|epred31_2(X8,X7)))&((X7!=cell71|X8!=cell72)|epred31_2(X8,X7)))&((X7!=cell72|X8!=cell73)|epred31_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_160])])])])])).
cnf(c_0_182,plain,(epred30_2(X1,X2)|~epred29_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_183,plain,(epred29_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_162, c_0_163])).
cnf(c_0_184,plain,(epred28_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_142, c_0_164])).
cnf(c_0_185,plain,(epred27_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_122, c_0_165])).
cnf(c_0_186,plain,(epred26_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_102, c_0_166])).
cnf(c_0_187,plain,(epred25_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_84, c_0_167])).
cnf(c_0_188,plain,(epred24_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_68, c_0_168])).
cnf(c_0_189,plain,(epred23_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_54, c_0_169])).
cnf(c_0_190,plain,(epred22_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_42, c_0_170])).
cnf(c_0_191,plain,(epred21_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_32, c_0_171])).
fof(c_0_192, plain, (![X2]:![X1]:(epred13_2(X1,X2)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60))|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65))|(X2=cell56&X1=cell66))|(X2=cell57&X1=cell67))|(X2=cell58&X1=cell68))|(X2=cell59&X1=cell69))|(X2=cell60&X1=cell70))|(X2=cell61&X1=cell71))|(X2=cell62&X1=cell72))|(X2=cell63&X1=cell73))|(X2=cell64&X1=cell74))|(X2=cell65&X1=cell75)))), introduced(definition)).
fof(c_0_193, plain, (![X2]:![X1]:(epred12_2(X1,X2)<=>(((((epred11_2(X1,X2)|(X2=cell56&X1=cell66))|(X2=cell57&X1=cell67))|(X2=cell58&X1=cell68))|(X2=cell59&X1=cell69))|(X2=cell60&X1=cell70)))), inference(apply_def,[status(thm)],[c_0_172, c_0_152])).
fof(c_0_194, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell55|(X3=cell54|(X3=cell53|(X3=cell52|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X3=cell54|(X3=cell53|(X3=cell52|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))&(((X3=cell55|(X4=cell64|(X3=cell53|(X3=cell52|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X4=cell64|(X3=cell53|(X3=cell52|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))))&((((X3=cell55|(X3=cell54|(X4=cell63|(X3=cell52|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X3=cell54|(X4=cell63|(X3=cell52|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))&(((X3=cell55|(X4=cell64|(X4=cell63|(X3=cell52|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X4=cell64|(X4=cell63|(X3=cell52|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))))&(((((X3=cell55|(X3=cell54|(X3=cell53|(X4=cell62|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X3=cell54|(X3=cell53|(X4=cell62|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))&(((X3=cell55|(X4=cell64|(X3=cell53|(X4=cell62|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X4=cell64|(X3=cell53|(X4=cell62|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))))&((((X3=cell55|(X3=cell54|(X4=cell63|(X4=cell62|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X3=cell54|(X4=cell63|(X4=cell62|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))&(((X3=cell55|(X4=cell64|(X4=cell63|(X4=cell62|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X4=cell64|(X4=cell63|(X4=cell62|(X3=cell51|epred10_2(X4,X3))))))|~epred11_2(X4,X3))))))&((((((X3=cell55|(X3=cell54|(X3=cell53|(X3=cell52|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X3=cell54|(X3=cell53|(X3=cell52|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))&(((X3=cell55|(X4=cell64|(X3=cell53|(X3=cell52|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X4=cell64|(X3=cell53|(X3=cell52|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))))&((((X3=cell55|(X3=cell54|(X4=cell63|(X3=cell52|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X3=cell54|(X4=cell63|(X3=cell52|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))&(((X3=cell55|(X4=cell64|(X4=cell63|(X3=cell52|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X4=cell64|(X4=cell63|(X3=cell52|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))))&(((((X3=cell55|(X3=cell54|(X3=cell53|(X4=cell62|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X3=cell54|(X3=cell53|(X4=cell62|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))&(((X3=cell55|(X4=cell64|(X3=cell53|(X4=cell62|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X4=cell64|(X3=cell53|(X4=cell62|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))))&((((X3=cell55|(X3=cell54|(X4=cell63|(X4=cell62|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X3=cell54|(X4=cell63|(X4=cell62|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))&(((X3=cell55|(X4=cell64|(X4=cell63|(X4=cell62|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3))&((X4=cell65|(X4=cell64|(X4=cell63|(X4=cell62|(X4=cell61|epred10_2(X4,X3))))))|~epred11_2(X4,X3)))))))&((((((~epred10_2(X6,X5)|epred11_2(X6,X5))&((X5!=cell51|X6!=cell61)|epred11_2(X6,X5)))&((X5!=cell52|X6!=cell62)|epred11_2(X6,X5)))&((X5!=cell53|X6!=cell63)|epred11_2(X6,X5)))&((X5!=cell54|X6!=cell64)|epred11_2(X6,X5)))&((X5!=cell55|X6!=cell65)|epred11_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_173])])])])])).
cnf(c_0_195,plain,(epred10_2(X1,X2)|~epred9_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_196,plain,(epred9_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_175, c_0_176])).
cnf(c_0_197,plain,(epred8_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_155, c_0_177])).
cnf(c_0_198,plain,(epred7_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_135, c_0_178])).
cnf(c_0_199,plain,(epred8_2(X1,X2)|X1!=cell50|X2!=cell40), inference(split_conjunct,[status(thm)],[c_0_134])).
fof(c_0_200, plain, (![X4]:![X3]:(epred34_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89)))), introduced(definition)).
fof(c_0_201, plain, (![X4]:![X3]:(epred33_2(X3,X4)<=>(((((epred32_2(X3,X4)|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84)))), inference(apply_def,[status(thm)],[c_0_179, c_0_159])).
fof(c_0_202, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell77|(X5=cell76|(X5=cell75|(X5=cell74|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X5=cell76|(X5=cell75|(X5=cell74|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X5=cell75|(X5=cell74|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X6=cell77|(X5=cell75|(X5=cell74|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))&((((X5=cell77|(X5=cell76|(X6=cell76|(X5=cell74|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X5=cell76|(X6=cell76|(X5=cell74|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X6=cell76|(X5=cell74|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X6=cell77|(X6=cell76|(X5=cell74|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))))&(((((X5=cell77|(X5=cell76|(X5=cell75|(X6=cell75|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X5=cell76|(X5=cell75|(X6=cell75|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X5=cell75|(X6=cell75|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X6=cell77|(X5=cell75|(X6=cell75|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))&((((X5=cell77|(X5=cell76|(X6=cell76|(X6=cell75|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X5=cell76|(X6=cell76|(X6=cell75|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X6=cell76|(X6=cell75|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X6=cell77|(X6=cell76|(X6=cell75|(X5=cell73|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))))&((((((X5=cell77|(X5=cell76|(X5=cell75|(X5=cell74|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X5=cell76|(X5=cell75|(X5=cell74|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X5=cell75|(X5=cell74|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X6=cell77|(X5=cell75|(X5=cell74|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))&((((X5=cell77|(X5=cell76|(X6=cell76|(X5=cell74|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X5=cell76|(X6=cell76|(X5=cell74|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X6=cell76|(X5=cell74|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X6=cell77|(X6=cell76|(X5=cell74|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))))&(((((X5=cell77|(X5=cell76|(X5=cell75|(X6=cell75|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X5=cell76|(X5=cell75|(X6=cell75|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X5=cell75|(X6=cell75|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X6=cell77|(X5=cell75|(X6=cell75|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))))&((((X5=cell77|(X5=cell76|(X6=cell76|(X6=cell75|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X5=cell76|(X6=cell76|(X6=cell75|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))&(((X5=cell77|(X6=cell77|(X6=cell76|(X6=cell75|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5))&((X6=cell78|(X6=cell77|(X6=cell76|(X6=cell75|(X6=cell74|epred31_2(X6,X5))))))|~epred32_2(X6,X5)))))))&((((((~epred31_2(X8,X7)|epred32_2(X8,X7))&((X7!=cell73|X8!=cell74)|epred32_2(X8,X7)))&((X7!=cell74|X8!=cell75)|epred32_2(X8,X7)))&((X7!=cell75|X8!=cell76)|epred32_2(X8,X7)))&((X7!=cell76|X8!=cell77)|epred32_2(X8,X7)))&((X7!=cell77|X8!=cell78)|epred32_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_180])])])])])).
cnf(c_0_203,plain,(epred31_2(X1,X2)|~epred30_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_204,plain,(epred30_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_182, c_0_183])).
cnf(c_0_205,plain,(epred29_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_162, c_0_184])).
cnf(c_0_206,plain,(epred28_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_142, c_0_185])).
cnf(c_0_207,plain,(epred27_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_122, c_0_186])).
cnf(c_0_208,plain,(epred26_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_102, c_0_187])).
cnf(c_0_209,plain,(epred25_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_84, c_0_188])).
cnf(c_0_210,plain,(epred24_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_68, c_0_189])).
cnf(c_0_211,plain,(epred23_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_54, c_0_190])).
cnf(c_0_212,plain,(epred22_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_42, c_0_191])).
fof(c_0_213, plain, (![X2]:![X1]:(epred14_2(X1,X2)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60))|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65))|(X2=cell56&X1=cell66))|(X2=cell57&X1=cell67))|(X2=cell58&X1=cell68))|(X2=cell59&X1=cell69))|(X2=cell60&X1=cell70))|(X2=cell61&X1=cell71))|(X2=cell62&X1=cell72))|(X2=cell63&X1=cell73))|(X2=cell64&X1=cell74))|(X2=cell65&X1=cell75))|(X2=cell66&X1=cell76))|(X2=cell67&X1=cell77))|(X2=cell68&X1=cell78))|(X2=cell69&X1=cell79))|(X2=cell70&X1=cell80)))), introduced(definition)).
fof(c_0_214, plain, (![X2]:![X1]:(epred13_2(X1,X2)<=>(((((epred12_2(X1,X2)|(X2=cell61&X1=cell71))|(X2=cell62&X1=cell72))|(X2=cell63&X1=cell73))|(X2=cell64&X1=cell74))|(X2=cell65&X1=cell75)))), inference(apply_def,[status(thm)],[c_0_192, c_0_172])).
fof(c_0_215, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell60|(X3=cell59|(X3=cell58|(X3=cell57|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X3=cell59|(X3=cell58|(X3=cell57|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))&(((X3=cell60|(X4=cell69|(X3=cell58|(X3=cell57|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X4=cell69|(X3=cell58|(X3=cell57|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))))&((((X3=cell60|(X3=cell59|(X4=cell68|(X3=cell57|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X3=cell59|(X4=cell68|(X3=cell57|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))&(((X3=cell60|(X4=cell69|(X4=cell68|(X3=cell57|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X4=cell69|(X4=cell68|(X3=cell57|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))))&(((((X3=cell60|(X3=cell59|(X3=cell58|(X4=cell67|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X3=cell59|(X3=cell58|(X4=cell67|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))&(((X3=cell60|(X4=cell69|(X3=cell58|(X4=cell67|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X4=cell69|(X3=cell58|(X4=cell67|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))))&((((X3=cell60|(X3=cell59|(X4=cell68|(X4=cell67|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X3=cell59|(X4=cell68|(X4=cell67|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))&(((X3=cell60|(X4=cell69|(X4=cell68|(X4=cell67|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X4=cell69|(X4=cell68|(X4=cell67|(X3=cell56|epred11_2(X4,X3))))))|~epred12_2(X4,X3))))))&((((((X3=cell60|(X3=cell59|(X3=cell58|(X3=cell57|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X3=cell59|(X3=cell58|(X3=cell57|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))&(((X3=cell60|(X4=cell69|(X3=cell58|(X3=cell57|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X4=cell69|(X3=cell58|(X3=cell57|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))))&((((X3=cell60|(X3=cell59|(X4=cell68|(X3=cell57|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X3=cell59|(X4=cell68|(X3=cell57|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))&(((X3=cell60|(X4=cell69|(X4=cell68|(X3=cell57|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X4=cell69|(X4=cell68|(X3=cell57|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))))&(((((X3=cell60|(X3=cell59|(X3=cell58|(X4=cell67|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X3=cell59|(X3=cell58|(X4=cell67|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))&(((X3=cell60|(X4=cell69|(X3=cell58|(X4=cell67|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X4=cell69|(X3=cell58|(X4=cell67|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))))&((((X3=cell60|(X3=cell59|(X4=cell68|(X4=cell67|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X3=cell59|(X4=cell68|(X4=cell67|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))&(((X3=cell60|(X4=cell69|(X4=cell68|(X4=cell67|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3))&((X4=cell70|(X4=cell69|(X4=cell68|(X4=cell67|(X4=cell66|epred11_2(X4,X3))))))|~epred12_2(X4,X3)))))))&((((((~epred11_2(X6,X5)|epred12_2(X6,X5))&((X5!=cell56|X6!=cell66)|epred12_2(X6,X5)))&((X5!=cell57|X6!=cell67)|epred12_2(X6,X5)))&((X5!=cell58|X6!=cell68)|epred12_2(X6,X5)))&((X5!=cell59|X6!=cell69)|epred12_2(X6,X5)))&((X5!=cell60|X6!=cell70)|epred12_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_193])])])])])).
cnf(c_0_216,plain,(epred11_2(X1,X2)|~epred10_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_194])).
cnf(c_0_217,plain,(epred10_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_195, c_0_196])).
cnf(c_0_218,plain,(epred9_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_175, c_0_197])).
cnf(c_0_219,plain,(epred8_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_155, c_0_198])).
cnf(c_0_220,plain,(epred8_2(X1,cell40)|cell50!=X1), inference(er,[status(thm)],[c_0_199])).
fof(c_0_221, plain, (![X4]:![X3]:(epred35_2(X3,X4)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95)))), introduced(definition)).
fof(c_0_222, plain, (![X4]:![X3]:(epred34_2(X3,X4)<=>(((((epred33_2(X3,X4)|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89)))), inference(apply_def,[status(thm)],[c_0_200, c_0_179])).
fof(c_0_223, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell83|(X5=cell82|(X5=cell81|(X5=cell79|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X5=cell82|(X5=cell81|(X5=cell79|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X5=cell81|(X5=cell79|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X6=cell83|(X5=cell81|(X5=cell79|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))&((((X5=cell83|(X5=cell82|(X6=cell82|(X5=cell79|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X5=cell82|(X6=cell82|(X5=cell79|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X6=cell82|(X5=cell79|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X6=cell83|(X6=cell82|(X5=cell79|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))))&(((((X5=cell83|(X5=cell82|(X5=cell81|(X6=cell80|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X5=cell82|(X5=cell81|(X6=cell80|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X5=cell81|(X6=cell80|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X6=cell83|(X5=cell81|(X6=cell80|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))&((((X5=cell83|(X5=cell82|(X6=cell82|(X6=cell80|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X5=cell82|(X6=cell82|(X6=cell80|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X6=cell82|(X6=cell80|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X6=cell83|(X6=cell82|(X6=cell80|(X5=cell78|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))))&((((((X5=cell83|(X5=cell82|(X5=cell81|(X5=cell79|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X5=cell82|(X5=cell81|(X5=cell79|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X5=cell81|(X5=cell79|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X6=cell83|(X5=cell81|(X5=cell79|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))&((((X5=cell83|(X5=cell82|(X6=cell82|(X5=cell79|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X5=cell82|(X6=cell82|(X5=cell79|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X6=cell82|(X5=cell79|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X6=cell83|(X6=cell82|(X5=cell79|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))))&(((((X5=cell83|(X5=cell82|(X5=cell81|(X6=cell80|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X5=cell82|(X5=cell81|(X6=cell80|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X5=cell81|(X6=cell80|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X6=cell83|(X5=cell81|(X6=cell80|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))))&((((X5=cell83|(X5=cell82|(X6=cell82|(X6=cell80|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X5=cell82|(X6=cell82|(X6=cell80|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))&(((X5=cell83|(X6=cell83|(X6=cell82|(X6=cell80|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5))&((X6=cell84|(X6=cell83|(X6=cell82|(X6=cell80|(X6=cell79|epred32_2(X6,X5))))))|~epred33_2(X6,X5)))))))&((((((~epred32_2(X8,X7)|epred33_2(X8,X7))&((X7!=cell78|X8!=cell79)|epred33_2(X8,X7)))&((X7!=cell79|X8!=cell80)|epred33_2(X8,X7)))&((X7!=cell81|X8!=cell82)|epred33_2(X8,X7)))&((X7!=cell82|X8!=cell83)|epred33_2(X8,X7)))&((X7!=cell83|X8!=cell84)|epred33_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_201])])])])])).
cnf(c_0_224,plain,(epred32_2(X1,X2)|~epred31_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_225,plain,(epred31_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_203, c_0_204])).
cnf(c_0_226,plain,(epred30_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_182, c_0_205])).
cnf(c_0_227,plain,(epred29_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_162, c_0_206])).
cnf(c_0_228,plain,(epred28_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_142, c_0_207])).
cnf(c_0_229,plain,(epred27_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_122, c_0_208])).
cnf(c_0_230,plain,(epred26_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_102, c_0_209])).
cnf(c_0_231,plain,(epred25_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_84, c_0_210])).
cnf(c_0_232,plain,(epred24_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_68, c_0_211])).
cnf(c_0_233,plain,(epred23_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_54, c_0_212])).
fof(c_0_234, plain, (![X2]:![X1]:(epred15_2(X1,X2)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60))|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65))|(X2=cell56&X1=cell66))|(X2=cell57&X1=cell67))|(X2=cell58&X1=cell68))|(X2=cell59&X1=cell69))|(X2=cell60&X1=cell70))|(X2=cell61&X1=cell71))|(X2=cell62&X1=cell72))|(X2=cell63&X1=cell73))|(X2=cell64&X1=cell74))|(X2=cell65&X1=cell75))|(X2=cell66&X1=cell76))|(X2=cell67&X1=cell77))|(X2=cell68&X1=cell78))|(X2=cell69&X1=cell79))|(X2=cell70&X1=cell80))|(X2=cell71&X1=cell81))|(X2=cell72&X1=cell82))|(X2=cell73&X1=cell83))|(X2=cell74&X1=cell84))|(X2=cell75&X1=cell85)))), introduced(definition)).
fof(c_0_235, plain, (![X2]:![X1]:(epred14_2(X1,X2)<=>(((((epred13_2(X1,X2)|(X2=cell66&X1=cell76))|(X2=cell67&X1=cell77))|(X2=cell68&X1=cell78))|(X2=cell69&X1=cell79))|(X2=cell70&X1=cell80)))), inference(apply_def,[status(thm)],[c_0_213, c_0_192])).
fof(c_0_236, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell65|(X3=cell64|(X3=cell63|(X3=cell62|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X3=cell64|(X3=cell63|(X3=cell62|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))&(((X3=cell65|(X4=cell74|(X3=cell63|(X3=cell62|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X4=cell74|(X3=cell63|(X3=cell62|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))))&((((X3=cell65|(X3=cell64|(X4=cell73|(X3=cell62|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X3=cell64|(X4=cell73|(X3=cell62|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))&(((X3=cell65|(X4=cell74|(X4=cell73|(X3=cell62|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X4=cell74|(X4=cell73|(X3=cell62|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))))&(((((X3=cell65|(X3=cell64|(X3=cell63|(X4=cell72|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X3=cell64|(X3=cell63|(X4=cell72|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))&(((X3=cell65|(X4=cell74|(X3=cell63|(X4=cell72|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X4=cell74|(X3=cell63|(X4=cell72|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))))&((((X3=cell65|(X3=cell64|(X4=cell73|(X4=cell72|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X3=cell64|(X4=cell73|(X4=cell72|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))&(((X3=cell65|(X4=cell74|(X4=cell73|(X4=cell72|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X4=cell74|(X4=cell73|(X4=cell72|(X3=cell61|epred12_2(X4,X3))))))|~epred13_2(X4,X3))))))&((((((X3=cell65|(X3=cell64|(X3=cell63|(X3=cell62|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X3=cell64|(X3=cell63|(X3=cell62|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))&(((X3=cell65|(X4=cell74|(X3=cell63|(X3=cell62|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X4=cell74|(X3=cell63|(X3=cell62|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))))&((((X3=cell65|(X3=cell64|(X4=cell73|(X3=cell62|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X3=cell64|(X4=cell73|(X3=cell62|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))&(((X3=cell65|(X4=cell74|(X4=cell73|(X3=cell62|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X4=cell74|(X4=cell73|(X3=cell62|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))))&(((((X3=cell65|(X3=cell64|(X3=cell63|(X4=cell72|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X3=cell64|(X3=cell63|(X4=cell72|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))&(((X3=cell65|(X4=cell74|(X3=cell63|(X4=cell72|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X4=cell74|(X3=cell63|(X4=cell72|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))))&((((X3=cell65|(X3=cell64|(X4=cell73|(X4=cell72|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X3=cell64|(X4=cell73|(X4=cell72|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))&(((X3=cell65|(X4=cell74|(X4=cell73|(X4=cell72|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3))&((X4=cell75|(X4=cell74|(X4=cell73|(X4=cell72|(X4=cell71|epred12_2(X4,X3))))))|~epred13_2(X4,X3)))))))&((((((~epred12_2(X6,X5)|epred13_2(X6,X5))&((X5!=cell61|X6!=cell71)|epred13_2(X6,X5)))&((X5!=cell62|X6!=cell72)|epred13_2(X6,X5)))&((X5!=cell63|X6!=cell73)|epred13_2(X6,X5)))&((X5!=cell64|X6!=cell74)|epred13_2(X6,X5)))&((X5!=cell65|X6!=cell75)|epred13_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_214])])])])])).
cnf(c_0_237,plain,(epred12_2(X1,X2)|~epred11_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_215])).
cnf(c_0_238,plain,(epred11_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_216, c_0_217])).
cnf(c_0_239,plain,(epred10_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_195, c_0_218])).
cnf(c_0_240,plain,(epred9_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_175, c_0_219])).
cnf(c_0_241,plain,(epred8_2(cell50,cell40)), inference(er,[status(thm)],[c_0_220])).
fof(c_0_242, plain, (![X4]:![X3]:(epred36_2(X3,X4)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X4=cell1&X3=cell2)|(X4=cell2&X3=cell3))|(X4=cell3&X3=cell4))|(X4=cell4&X3=cell5))|(X4=cell5&X3=cell6))|(X4=cell6&X3=cell7))|(X4=cell7&X3=cell8))|(X4=cell8&X3=cell9))|(X4=cell9&X3=cell10))|(X4=cell11&X3=cell12))|(X4=cell12&X3=cell13))|(X4=cell13&X3=cell14))|(X4=cell14&X3=cell15))|(X4=cell15&X3=cell16))|(X4=cell16&X3=cell17))|(X4=cell17&X3=cell18))|(X4=cell18&X3=cell19))|(X4=cell19&X3=cell20))|(X4=cell21&X3=cell22))|(X4=cell22&X3=cell23))|(X4=cell23&X3=cell24))|(X4=cell24&X3=cell25))|(X4=cell25&X3=cell26))|(X4=cell26&X3=cell27))|(X4=cell27&X3=cell28))|(X4=cell28&X3=cell29))|(X4=cell29&X3=cell30))|(X4=cell31&X3=cell32))|(X4=cell32&X3=cell33))|(X4=cell33&X3=cell34))|(X4=cell34&X3=cell35))|(X4=cell35&X3=cell36))|(X4=cell36&X3=cell37))|(X4=cell37&X3=cell38))|(X4=cell38&X3=cell39))|(X4=cell39&X3=cell40))|(X4=cell41&X3=cell42))|(X4=cell42&X3=cell43))|(X4=cell43&X3=cell44))|(X4=cell44&X3=cell45))|(X4=cell45&X3=cell46))|(X4=cell46&X3=cell47))|(X4=cell47&X3=cell48))|(X4=cell48&X3=cell49))|(X4=cell49&X3=cell50))|(X4=cell51&X3=cell52))|(X4=cell52&X3=cell53))|(X4=cell53&X3=cell54))|(X4=cell54&X3=cell55))|(X4=cell55&X3=cell56))|(X4=cell56&X3=cell57))|(X4=cell57&X3=cell58))|(X4=cell58&X3=cell59))|(X4=cell59&X3=cell60))|(X4=cell61&X3=cell62))|(X4=cell62&X3=cell63))|(X4=cell63&X3=cell64))|(X4=cell64&X3=cell65))|(X4=cell65&X3=cell66))|(X4=cell66&X3=cell67))|(X4=cell67&X3=cell68))|(X4=cell68&X3=cell69))|(X4=cell69&X3=cell70))|(X4=cell71&X3=cell72))|(X4=cell72&X3=cell73))|(X4=cell73&X3=cell74))|(X4=cell74&X3=cell75))|(X4=cell75&X3=cell76))|(X4=cell76&X3=cell77))|(X4=cell77&X3=cell78))|(X4=cell78&X3=cell79))|(X4=cell79&X3=cell80))|(X4=cell81&X3=cell82))|(X4=cell82&X3=cell83))|(X4=cell83&X3=cell84))|(X4=cell84&X3=cell85))|(X4=cell85&X3=cell86))|(X4=cell86&X3=cell87))|(X4=cell87&X3=cell88))|(X4=cell88&X3=cell89))|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95))|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100)))), introduced(definition)).
fof(c_0_243, plain, (![X4]:![X3]:(epred35_2(X3,X4)<=>(((((epred34_2(X3,X4)|(X4=cell89&X3=cell90))|(X4=cell91&X3=cell92))|(X4=cell92&X3=cell93))|(X4=cell93&X3=cell94))|(X4=cell94&X3=cell95)))), inference(apply_def,[status(thm)],[c_0_221, c_0_200])).
fof(c_0_244, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell88|(X5=cell87|(X5=cell86|(X5=cell85|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X5=cell87|(X5=cell86|(X5=cell85|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X5=cell86|(X5=cell85|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X6=cell88|(X5=cell86|(X5=cell85|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))&((((X5=cell88|(X5=cell87|(X6=cell87|(X5=cell85|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X5=cell87|(X6=cell87|(X5=cell85|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X6=cell87|(X5=cell85|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X6=cell88|(X6=cell87|(X5=cell85|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))))&(((((X5=cell88|(X5=cell87|(X5=cell86|(X6=cell86|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X5=cell87|(X5=cell86|(X6=cell86|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X5=cell86|(X6=cell86|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X6=cell88|(X5=cell86|(X6=cell86|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))&((((X5=cell88|(X5=cell87|(X6=cell87|(X6=cell86|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X5=cell87|(X6=cell87|(X6=cell86|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X6=cell87|(X6=cell86|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X6=cell88|(X6=cell87|(X6=cell86|(X5=cell84|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))))&((((((X5=cell88|(X5=cell87|(X5=cell86|(X5=cell85|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X5=cell87|(X5=cell86|(X5=cell85|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X5=cell86|(X5=cell85|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X6=cell88|(X5=cell86|(X5=cell85|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))&((((X5=cell88|(X5=cell87|(X6=cell87|(X5=cell85|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X5=cell87|(X6=cell87|(X5=cell85|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X6=cell87|(X5=cell85|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X6=cell88|(X6=cell87|(X5=cell85|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))))&(((((X5=cell88|(X5=cell87|(X5=cell86|(X6=cell86|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X5=cell87|(X5=cell86|(X6=cell86|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X5=cell86|(X6=cell86|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X6=cell88|(X5=cell86|(X6=cell86|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))))&((((X5=cell88|(X5=cell87|(X6=cell87|(X6=cell86|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X5=cell87|(X6=cell87|(X6=cell86|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))&(((X5=cell88|(X6=cell88|(X6=cell87|(X6=cell86|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5))&((X6=cell89|(X6=cell88|(X6=cell87|(X6=cell86|(X6=cell85|epred33_2(X6,X5))))))|~epred34_2(X6,X5)))))))&((((((~epred33_2(X8,X7)|epred34_2(X8,X7))&((X7!=cell84|X8!=cell85)|epred34_2(X8,X7)))&((X7!=cell85|X8!=cell86)|epred34_2(X8,X7)))&((X7!=cell86|X8!=cell87)|epred34_2(X8,X7)))&((X7!=cell87|X8!=cell88)|epred34_2(X8,X7)))&((X7!=cell88|X8!=cell89)|epred34_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_222])])])])])).
cnf(c_0_245,plain,(epred33_2(X1,X2)|~epred32_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_223])).
cnf(c_0_246,plain,(epred32_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_224, c_0_225])).
cnf(c_0_247,plain,(epred31_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_203, c_0_226])).
cnf(c_0_248,plain,(epred30_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_182, c_0_227])).
cnf(c_0_249,plain,(epred29_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_162, c_0_228])).
cnf(c_0_250,plain,(epred28_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_142, c_0_229])).
cnf(c_0_251,plain,(epred27_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_122, c_0_230])).
cnf(c_0_252,plain,(epred26_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_102, c_0_231])).
cnf(c_0_253,plain,(epred25_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_84, c_0_232])).
cnf(c_0_254,plain,(epred24_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_68, c_0_233])).
fof(c_0_255, plain, (![X2]:![X1]:(epred16_2(X1,X2)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60))|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65))|(X2=cell56&X1=cell66))|(X2=cell57&X1=cell67))|(X2=cell58&X1=cell68))|(X2=cell59&X1=cell69))|(X2=cell60&X1=cell70))|(X2=cell61&X1=cell71))|(X2=cell62&X1=cell72))|(X2=cell63&X1=cell73))|(X2=cell64&X1=cell74))|(X2=cell65&X1=cell75))|(X2=cell66&X1=cell76))|(X2=cell67&X1=cell77))|(X2=cell68&X1=cell78))|(X2=cell69&X1=cell79))|(X2=cell70&X1=cell80))|(X2=cell71&X1=cell81))|(X2=cell72&X1=cell82))|(X2=cell73&X1=cell83))|(X2=cell74&X1=cell84))|(X2=cell75&X1=cell85))|(X2=cell76&X1=cell86))|(X2=cell77&X1=cell87))|(X2=cell78&X1=cell88))|(X2=cell79&X1=cell89))|(X2=cell80&X1=cell90)))), introduced(definition)).
fof(c_0_256, plain, (![X2]:![X1]:(epred15_2(X1,X2)<=>(((((epred14_2(X1,X2)|(X2=cell71&X1=cell81))|(X2=cell72&X1=cell82))|(X2=cell73&X1=cell83))|(X2=cell74&X1=cell84))|(X2=cell75&X1=cell85)))), inference(apply_def,[status(thm)],[c_0_234, c_0_213])).
fof(c_0_257, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell70|(X3=cell69|(X3=cell68|(X3=cell67|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X3=cell69|(X3=cell68|(X3=cell67|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))&(((X3=cell70|(X4=cell79|(X3=cell68|(X3=cell67|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X4=cell79|(X3=cell68|(X3=cell67|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))))&((((X3=cell70|(X3=cell69|(X4=cell78|(X3=cell67|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X3=cell69|(X4=cell78|(X3=cell67|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))&(((X3=cell70|(X4=cell79|(X4=cell78|(X3=cell67|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X4=cell79|(X4=cell78|(X3=cell67|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))))&(((((X3=cell70|(X3=cell69|(X3=cell68|(X4=cell77|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X3=cell69|(X3=cell68|(X4=cell77|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))&(((X3=cell70|(X4=cell79|(X3=cell68|(X4=cell77|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X4=cell79|(X3=cell68|(X4=cell77|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))))&((((X3=cell70|(X3=cell69|(X4=cell78|(X4=cell77|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X3=cell69|(X4=cell78|(X4=cell77|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))&(((X3=cell70|(X4=cell79|(X4=cell78|(X4=cell77|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X4=cell79|(X4=cell78|(X4=cell77|(X3=cell66|epred13_2(X4,X3))))))|~epred14_2(X4,X3))))))&((((((X3=cell70|(X3=cell69|(X3=cell68|(X3=cell67|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X3=cell69|(X3=cell68|(X3=cell67|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))&(((X3=cell70|(X4=cell79|(X3=cell68|(X3=cell67|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X4=cell79|(X3=cell68|(X3=cell67|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))))&((((X3=cell70|(X3=cell69|(X4=cell78|(X3=cell67|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X3=cell69|(X4=cell78|(X3=cell67|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))&(((X3=cell70|(X4=cell79|(X4=cell78|(X3=cell67|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X4=cell79|(X4=cell78|(X3=cell67|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))))&(((((X3=cell70|(X3=cell69|(X3=cell68|(X4=cell77|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X3=cell69|(X3=cell68|(X4=cell77|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))&(((X3=cell70|(X4=cell79|(X3=cell68|(X4=cell77|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X4=cell79|(X3=cell68|(X4=cell77|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))))&((((X3=cell70|(X3=cell69|(X4=cell78|(X4=cell77|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X3=cell69|(X4=cell78|(X4=cell77|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))&(((X3=cell70|(X4=cell79|(X4=cell78|(X4=cell77|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3))&((X4=cell80|(X4=cell79|(X4=cell78|(X4=cell77|(X4=cell76|epred13_2(X4,X3))))))|~epred14_2(X4,X3)))))))&((((((~epred13_2(X6,X5)|epred14_2(X6,X5))&((X5!=cell66|X6!=cell76)|epred14_2(X6,X5)))&((X5!=cell67|X6!=cell77)|epred14_2(X6,X5)))&((X5!=cell68|X6!=cell78)|epred14_2(X6,X5)))&((X5!=cell69|X6!=cell79)|epred14_2(X6,X5)))&((X5!=cell70|X6!=cell80)|epred14_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_235])])])])])).
cnf(c_0_258,plain,(epred13_2(X1,X2)|~epred12_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_236])).
cnf(c_0_259,plain,(epred12_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_237, c_0_238])).
cnf(c_0_260,plain,(epred11_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_216, c_0_239])).
cnf(c_0_261,plain,(epred10_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_195, c_0_240])).
cnf(c_0_262,plain,(epred9_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_175, c_0_241])).
cnf(c_0_263,plain,(epred10_2(X1,X2)|X1!=cell60|X2!=cell50), inference(split_conjunct,[status(thm)],[c_0_174])).
fof(c_0_264, plain, (![X4]:![X3]:(epred36_2(X3,X4)<=>(((((epred35_2(X3,X4)|(X4=cell95&X3=cell96))|(X4=cell96&X3=cell97))|(X4=cell97&X3=cell98))|(X4=cell98&X3=cell99))|(X4=cell99&X3=cell100)))), inference(apply_def,[status(thm)],[c_0_242, c_0_221])).
fof(c_0_265, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell94|(X5=cell93|(X5=cell92|(X5=cell91|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell93|(X5=cell92|(X5=cell91|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X5=cell92|(X5=cell91|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell92|(X5=cell91|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))&((((X5=cell94|(X5=cell93|(X6=cell93|(X5=cell91|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell93|(X6=cell93|(X5=cell91|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X6=cell93|(X5=cell91|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X5=cell91|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))))&(((((X5=cell94|(X5=cell93|(X5=cell92|(X6=cell92|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell93|(X5=cell92|(X6=cell92|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X5=cell92|(X6=cell92|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell92|(X6=cell92|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))&((((X5=cell94|(X5=cell93|(X6=cell93|(X6=cell92|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell93|(X6=cell93|(X6=cell92|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X6=cell93|(X6=cell92|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X6=cell92|(X5=cell89|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))))&((((((X5=cell94|(X5=cell93|(X5=cell92|(X5=cell91|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell93|(X5=cell92|(X5=cell91|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X5=cell92|(X5=cell91|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell92|(X5=cell91|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))&((((X5=cell94|(X5=cell93|(X6=cell93|(X5=cell91|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell93|(X6=cell93|(X5=cell91|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X6=cell93|(X5=cell91|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X5=cell91|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))))&(((((X5=cell94|(X5=cell93|(X5=cell92|(X6=cell92|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell93|(X5=cell92|(X6=cell92|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X5=cell92|(X6=cell92|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X5=cell92|(X6=cell92|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))))&((((X5=cell94|(X5=cell93|(X6=cell93|(X6=cell92|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X5=cell93|(X6=cell93|(X6=cell92|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))&(((X5=cell94|(X6=cell94|(X6=cell93|(X6=cell92|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5))&((X6=cell95|(X6=cell94|(X6=cell93|(X6=cell92|(X6=cell90|epred34_2(X6,X5))))))|~epred35_2(X6,X5)))))))&((((((~epred34_2(X8,X7)|epred35_2(X8,X7))&((X7!=cell89|X8!=cell90)|epred35_2(X8,X7)))&((X7!=cell91|X8!=cell92)|epred35_2(X8,X7)))&((X7!=cell92|X8!=cell93)|epred35_2(X8,X7)))&((X7!=cell93|X8!=cell94)|epred35_2(X8,X7)))&((X7!=cell94|X8!=cell95)|epred35_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_243])])])])])).
cnf(c_0_266,plain,(epred34_2(X1,X2)|~epred33_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_244])).
cnf(c_0_267,plain,(epred33_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_245, c_0_246])).
cnf(c_0_268,plain,(epred32_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_224, c_0_247])).
cnf(c_0_269,plain,(epred31_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_203, c_0_248])).
cnf(c_0_270,plain,(epred30_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_182, c_0_249])).
cnf(c_0_271,plain,(epred29_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_162, c_0_250])).
cnf(c_0_272,plain,(epred28_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_142, c_0_251])).
cnf(c_0_273,plain,(epred27_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_122, c_0_252])).
cnf(c_0_274,plain,(epred26_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_102, c_0_253])).
cnf(c_0_275,plain,(epred25_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_84, c_0_254])).
fof(c_0_276, plain, (![X2]:![X1]:(epred17_2(X1,X2)<=>(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60))|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65))|(X2=cell56&X1=cell66))|(X2=cell57&X1=cell67))|(X2=cell58&X1=cell68))|(X2=cell59&X1=cell69))|(X2=cell60&X1=cell70))|(X2=cell61&X1=cell71))|(X2=cell62&X1=cell72))|(X2=cell63&X1=cell73))|(X2=cell64&X1=cell74))|(X2=cell65&X1=cell75))|(X2=cell66&X1=cell76))|(X2=cell67&X1=cell77))|(X2=cell68&X1=cell78))|(X2=cell69&X1=cell79))|(X2=cell70&X1=cell80))|(X2=cell71&X1=cell81))|(X2=cell72&X1=cell82))|(X2=cell73&X1=cell83))|(X2=cell74&X1=cell84))|(X2=cell75&X1=cell85))|(X2=cell76&X1=cell86))|(X2=cell77&X1=cell87))|(X2=cell78&X1=cell88))|(X2=cell79&X1=cell89))|(X2=cell80&X1=cell90))|(X2=cell81&X1=cell91))|(X2=cell82&X1=cell92))|(X2=cell83&X1=cell93))|(X2=cell84&X1=cell94))|(X2=cell85&X1=cell95)))), introduced(definition)).
fof(c_0_277, plain, (![X2]:![X1]:(epred16_2(X1,X2)<=>(((((epred15_2(X1,X2)|(X2=cell76&X1=cell86))|(X2=cell77&X1=cell87))|(X2=cell78&X1=cell88))|(X2=cell79&X1=cell89))|(X2=cell80&X1=cell90)))), inference(apply_def,[status(thm)],[c_0_255, c_0_234])).
fof(c_0_278, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell75|(X3=cell74|(X3=cell73|(X3=cell72|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X3=cell74|(X3=cell73|(X3=cell72|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))&(((X3=cell75|(X4=cell84|(X3=cell73|(X3=cell72|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X4=cell84|(X3=cell73|(X3=cell72|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))))&((((X3=cell75|(X3=cell74|(X4=cell83|(X3=cell72|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X3=cell74|(X4=cell83|(X3=cell72|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))&(((X3=cell75|(X4=cell84|(X4=cell83|(X3=cell72|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X4=cell84|(X4=cell83|(X3=cell72|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))))&(((((X3=cell75|(X3=cell74|(X3=cell73|(X4=cell82|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X3=cell74|(X3=cell73|(X4=cell82|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))&(((X3=cell75|(X4=cell84|(X3=cell73|(X4=cell82|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X4=cell84|(X3=cell73|(X4=cell82|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))))&((((X3=cell75|(X3=cell74|(X4=cell83|(X4=cell82|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X3=cell74|(X4=cell83|(X4=cell82|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))&(((X3=cell75|(X4=cell84|(X4=cell83|(X4=cell82|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X4=cell84|(X4=cell83|(X4=cell82|(X3=cell71|epred14_2(X4,X3))))))|~epred15_2(X4,X3))))))&((((((X3=cell75|(X3=cell74|(X3=cell73|(X3=cell72|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X3=cell74|(X3=cell73|(X3=cell72|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))&(((X3=cell75|(X4=cell84|(X3=cell73|(X3=cell72|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X4=cell84|(X3=cell73|(X3=cell72|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))))&((((X3=cell75|(X3=cell74|(X4=cell83|(X3=cell72|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X3=cell74|(X4=cell83|(X3=cell72|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))&(((X3=cell75|(X4=cell84|(X4=cell83|(X3=cell72|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X4=cell84|(X4=cell83|(X3=cell72|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))))&(((((X3=cell75|(X3=cell74|(X3=cell73|(X4=cell82|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X3=cell74|(X3=cell73|(X4=cell82|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))&(((X3=cell75|(X4=cell84|(X3=cell73|(X4=cell82|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X4=cell84|(X3=cell73|(X4=cell82|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))))&((((X3=cell75|(X3=cell74|(X4=cell83|(X4=cell82|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X3=cell74|(X4=cell83|(X4=cell82|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))&(((X3=cell75|(X4=cell84|(X4=cell83|(X4=cell82|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3))&((X4=cell85|(X4=cell84|(X4=cell83|(X4=cell82|(X4=cell81|epred14_2(X4,X3))))))|~epred15_2(X4,X3)))))))&((((((~epred14_2(X6,X5)|epred15_2(X6,X5))&((X5!=cell71|X6!=cell81)|epred15_2(X6,X5)))&((X5!=cell72|X6!=cell82)|epred15_2(X6,X5)))&((X5!=cell73|X6!=cell83)|epred15_2(X6,X5)))&((X5!=cell74|X6!=cell84)|epred15_2(X6,X5)))&((X5!=cell75|X6!=cell85)|epred15_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_256])])])])])).
cnf(c_0_279,plain,(epred14_2(X1,X2)|~epred13_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_257])).
cnf(c_0_280,plain,(epred13_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_258, c_0_259])).
cnf(c_0_281,plain,(epred12_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_237, c_0_260])).
cnf(c_0_282,plain,(epred11_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_216, c_0_261])).
cnf(c_0_283,plain,(epred10_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_195, c_0_262])).
cnf(c_0_284,plain,(epred10_2(X1,cell50)|cell60!=X1), inference(er,[status(thm)],[c_0_263])).
fof(c_0_285, plain, (![X44]:![X43]:(epred42_2(X43,X44)<=>((((((((north(X44,X43)|?[X45]:(north(X44,X45)&north(X45,X43)))|?[X46]:?[X47]:((north(X44,X46)&north(X46,X47))&north(X47,X43)))|?[X48]:?[X49]:?[X50]:(((north(X44,X48)&north(X48,X49))&north(X49,X50))&north(X50,X43)))|?[X51]:?[X52]:?[X53]:?[X54]:((((north(X44,X51)&north(X51,X52))&north(X52,X53))&north(X53,X54))&north(X54,X43)))|?[X55]:?[X56]:?[X57]:?[X58]:?[X59]:(((((north(X44,X55)&north(X55,X56))&north(X56,X57))&north(X57,X58))&north(X58,X59))&north(X59,X43)))|?[X60]:?[X61]:?[X62]:?[X63]:?[X64]:?[X65]:((((((north(X44,X60)&north(X60,X61))&north(X61,X62))&north(X62,X63))&north(X63,X64))&north(X64,X65))&north(X65,X43)))|?[X66]:?[X67]:?[X68]:?[X69]:?[X70]:?[X71]:?[X72]:(((((((north(X44,X66)&north(X66,X67))&north(X67,X68))&north(X68,X69))&north(X69,X70))&north(X70,X71))&north(X71,X72))&north(X72,X43)))|?[X73]:?[X74]:?[X75]:?[X76]:?[X77]:?[X78]:?[X79]:?[X80]:((((((((north(X44,X73)&north(X73,X74))&north(X74,X75))&north(X75,X76))&north(X76,X77))&north(X77,X78))&north(X78,X79))&north(X79,X80))&north(X80,X43))))), introduced(definition)).
fof(c_0_286, plain, (![X44]:![X43]:(epred41_2(X43,X44)<=>((((((north(X44,X43)|?[X45]:(north(X44,X45)&north(X45,X43)))|?[X46]:?[X47]:((north(X44,X46)&north(X46,X47))&north(X47,X43)))|?[X48]:?[X49]:?[X50]:(((north(X44,X48)&north(X48,X49))&north(X49,X50))&north(X50,X43)))|?[X51]:?[X52]:?[X53]:?[X54]:((((north(X44,X51)&north(X51,X52))&north(X52,X53))&north(X53,X54))&north(X54,X43)))|?[X55]:?[X56]:?[X57]:?[X58]:?[X59]:(((((north(X44,X55)&north(X55,X56))&north(X56,X57))&north(X57,X58))&north(X58,X59))&north(X59,X43)))|?[X60]:?[X61]:?[X62]:?[X63]:?[X64]:?[X65]:((((((north(X44,X60)&north(X60,X61))&north(X61,X62))&north(X62,X63))&north(X63,X64))&north(X64,X65))&north(X65,X43))))), introduced(definition)).
fof(c_0_287, axiom, (![X3]:![X4]:(north(X4,X3)<=>epred36_2(X3,X4))), inference(apply_def,[status(thm)],[c_0_0, c_0_242])).
fof(c_0_288, plain, (![X5]:![X6]:![X7]:![X8]:((((((((X5=cell99|(X5=cell98|(X5=cell97|(X5=cell96|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell98|(X5=cell97|(X5=cell96|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X5=cell97|(X5=cell96|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell97|(X5=cell96|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))&((((X5=cell99|(X5=cell98|(X6=cell98|(X5=cell96|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell98|(X6=cell98|(X5=cell96|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X6=cell98|(X5=cell96|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X5=cell96|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))))&(((((X5=cell99|(X5=cell98|(X5=cell97|(X6=cell97|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell98|(X5=cell97|(X6=cell97|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X5=cell97|(X6=cell97|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell97|(X6=cell97|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))&((((X5=cell99|(X5=cell98|(X6=cell98|(X6=cell97|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell98|(X6=cell98|(X6=cell97|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X6=cell98|(X6=cell97|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X6=cell97|(X5=cell95|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))))&((((((X5=cell99|(X5=cell98|(X5=cell97|(X5=cell96|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell98|(X5=cell97|(X5=cell96|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X5=cell97|(X5=cell96|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell97|(X5=cell96|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))&((((X5=cell99|(X5=cell98|(X6=cell98|(X5=cell96|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell98|(X6=cell98|(X5=cell96|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X6=cell98|(X5=cell96|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X5=cell96|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))))&(((((X5=cell99|(X5=cell98|(X5=cell97|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell98|(X5=cell97|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X5=cell97|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X5=cell97|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))))&((((X5=cell99|(X5=cell98|(X6=cell98|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X5=cell98|(X6=cell98|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))&(((X5=cell99|(X6=cell99|(X6=cell98|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5))&((X6=cell100|(X6=cell99|(X6=cell98|(X6=cell97|(X6=cell96|epred35_2(X6,X5))))))|~epred36_2(X6,X5)))))))&((((((~epred35_2(X8,X7)|epred36_2(X8,X7))&((X7!=cell95|X8!=cell96)|epred36_2(X8,X7)))&((X7!=cell96|X8!=cell97)|epred36_2(X8,X7)))&((X7!=cell97|X8!=cell98)|epred36_2(X8,X7)))&((X7!=cell98|X8!=cell99)|epred36_2(X8,X7)))&((X7!=cell99|X8!=cell100)|epred36_2(X8,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_264])])])])])).
cnf(c_0_289,plain,(epred35_2(X1,X2)|~epred34_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_265])).
cnf(c_0_290,plain,(epred34_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_266, c_0_267])).
cnf(c_0_291,plain,(epred33_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_245, c_0_268])).
cnf(c_0_292,plain,(epred32_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_224, c_0_269])).
cnf(c_0_293,plain,(epred31_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_203, c_0_270])).
cnf(c_0_294,plain,(epred30_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_182, c_0_271])).
cnf(c_0_295,plain,(epred29_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_162, c_0_272])).
cnf(c_0_296,plain,(epred28_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_142, c_0_273])).
cnf(c_0_297,plain,(epred27_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_122, c_0_274])).
cnf(c_0_298,plain,(epred26_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_102, c_0_275])).
fof(c_0_299, plain, (![X2]:![X1]:(epred18_2(X1,X2)<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((X2=cell1&X1=cell11)|(X2=cell2&X1=cell12))|(X2=cell3&X1=cell13))|(X2=cell4&X1=cell14))|(X2=cell5&X1=cell15))|(X2=cell6&X1=cell16))|(X2=cell7&X1=cell17))|(X2=cell8&X1=cell18))|(X2=cell9&X1=cell19))|(X2=cell10&X1=cell20))|(X2=cell11&X1=cell21))|(X2=cell12&X1=cell22))|(X2=cell13&X1=cell23))|(X2=cell14&X1=cell24))|(X2=cell15&X1=cell25))|(X2=cell16&X1=cell26))|(X2=cell17&X1=cell27))|(X2=cell18&X1=cell28))|(X2=cell19&X1=cell29))|(X2=cell20&X1=cell30))|(X2=cell21&X1=cell31))|(X2=cell22&X1=cell32))|(X2=cell23&X1=cell33))|(X2=cell24&X1=cell34))|(X2=cell25&X1=cell35))|(X2=cell26&X1=cell36))|(X2=cell27&X1=cell37))|(X2=cell28&X1=cell38))|(X2=cell29&X1=cell39))|(X2=cell30&X1=cell40))|(X2=cell31&X1=cell41))|(X2=cell32&X1=cell42))|(X2=cell33&X1=cell43))|(X2=cell34&X1=cell44))|(X2=cell35&X1=cell45))|(X2=cell36&X1=cell46))|(X2=cell37&X1=cell47))|(X2=cell38&X1=cell48))|(X2=cell39&X1=cell49))|(X2=cell40&X1=cell50))|(X2=cell41&X1=cell51))|(X2=cell42&X1=cell52))|(X2=cell43&X1=cell53))|(X2=cell44&X1=cell54))|(X2=cell45&X1=cell55))|(X2=cell46&X1=cell56))|(X2=cell47&X1=cell57))|(X2=cell48&X1=cell58))|(X2=cell49&X1=cell59))|(X2=cell50&X1=cell60))|(X2=cell51&X1=cell61))|(X2=cell52&X1=cell62))|(X2=cell53&X1=cell63))|(X2=cell54&X1=cell64))|(X2=cell55&X1=cell65))|(X2=cell56&X1=cell66))|(X2=cell57&X1=cell67))|(X2=cell58&X1=cell68))|(X2=cell59&X1=cell69))|(X2=cell60&X1=cell70))|(X2=cell61&X1=cell71))|(X2=cell62&X1=cell72))|(X2=cell63&X1=cell73))|(X2=cell64&X1=cell74))|(X2=cell65&X1=cell75))|(X2=cell66&X1=cell76))|(X2=cell67&X1=cell77))|(X2=cell68&X1=cell78))|(X2=cell69&X1=cell79))|(X2=cell70&X1=cell80))|(X2=cell71&X1=cell81))|(X2=cell72&X1=cell82))|(X2=cell73&X1=cell83))|(X2=cell74&X1=cell84))|(X2=cell75&X1=cell85))|(X2=cell76&X1=cell86))|(X2=cell77&X1=cell87))|(X2=cell78&X1=cell88))|(X2=cell79&X1=cell89))|(X2=cell80&X1=cell90))|(X2=cell81&X1=cell91))|(X2=cell82&X1=cell92))|(X2=cell83&X1=cell93))|(X2=cell84&X1=cell94))|(X2=cell85&X1=cell95))|(X2=cell86&X1=cell96))|(X2=cell87&X1=cell97))|(X2=cell88&X1=cell98))|(X2=cell89&X1=cell99))|(X2=cell90&X1=cell100)))), introduced(definition)).
fof(c_0_300, plain, (![X2]:![X1]:(epred17_2(X1,X2)<=>(((((epred16_2(X1,X2)|(X2=cell81&X1=cell91))|(X2=cell82&X1=cell92))|(X2=cell83&X1=cell93))|(X2=cell84&X1=cell94))|(X2=cell85&X1=cell95)))), inference(apply_def,[status(thm)],[c_0_276, c_0_255])).
fof(c_0_301, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell80|(X3=cell79|(X3=cell78|(X3=cell77|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X3=cell79|(X3=cell78|(X3=cell77|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))&(((X3=cell80|(X4=cell89|(X3=cell78|(X3=cell77|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X4=cell89|(X3=cell78|(X3=cell77|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))))&((((X3=cell80|(X3=cell79|(X4=cell88|(X3=cell77|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X3=cell79|(X4=cell88|(X3=cell77|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))&(((X3=cell80|(X4=cell89|(X4=cell88|(X3=cell77|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X4=cell89|(X4=cell88|(X3=cell77|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))))&(((((X3=cell80|(X3=cell79|(X3=cell78|(X4=cell87|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X3=cell79|(X3=cell78|(X4=cell87|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))&(((X3=cell80|(X4=cell89|(X3=cell78|(X4=cell87|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X4=cell89|(X3=cell78|(X4=cell87|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))))&((((X3=cell80|(X3=cell79|(X4=cell88|(X4=cell87|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X3=cell79|(X4=cell88|(X4=cell87|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))&(((X3=cell80|(X4=cell89|(X4=cell88|(X4=cell87|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X4=cell89|(X4=cell88|(X4=cell87|(X3=cell76|epred15_2(X4,X3))))))|~epred16_2(X4,X3))))))&((((((X3=cell80|(X3=cell79|(X3=cell78|(X3=cell77|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X3=cell79|(X3=cell78|(X3=cell77|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))&(((X3=cell80|(X4=cell89|(X3=cell78|(X3=cell77|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X4=cell89|(X3=cell78|(X3=cell77|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))))&((((X3=cell80|(X3=cell79|(X4=cell88|(X3=cell77|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X3=cell79|(X4=cell88|(X3=cell77|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))&(((X3=cell80|(X4=cell89|(X4=cell88|(X3=cell77|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X4=cell89|(X4=cell88|(X3=cell77|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))))&(((((X3=cell80|(X3=cell79|(X3=cell78|(X4=cell87|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X3=cell79|(X3=cell78|(X4=cell87|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))&(((X3=cell80|(X4=cell89|(X3=cell78|(X4=cell87|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X4=cell89|(X3=cell78|(X4=cell87|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))))&((((X3=cell80|(X3=cell79|(X4=cell88|(X4=cell87|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X3=cell79|(X4=cell88|(X4=cell87|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))&(((X3=cell80|(X4=cell89|(X4=cell88|(X4=cell87|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3))&((X4=cell90|(X4=cell89|(X4=cell88|(X4=cell87|(X4=cell86|epred15_2(X4,X3))))))|~epred16_2(X4,X3)))))))&((((((~epred15_2(X6,X5)|epred16_2(X6,X5))&((X5!=cell76|X6!=cell86)|epred16_2(X6,X5)))&((X5!=cell77|X6!=cell87)|epred16_2(X6,X5)))&((X5!=cell78|X6!=cell88)|epred16_2(X6,X5)))&((X5!=cell79|X6!=cell89)|epred16_2(X6,X5)))&((X5!=cell80|X6!=cell90)|epred16_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_277])])])])])).
cnf(c_0_302,plain,(epred15_2(X1,X2)|~epred14_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_278])).
cnf(c_0_303,plain,(epred14_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_279, c_0_280])).
cnf(c_0_304,plain,(epred13_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_258, c_0_281])).
cnf(c_0_305,plain,(epred12_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_237, c_0_282])).
cnf(c_0_306,plain,(epred11_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_216, c_0_283])).
cnf(c_0_307,plain,(epred10_2(cell60,cell50)), inference(er,[status(thm)],[c_0_284])).
fof(c_0_308, plain, (![X44]:![X43]:(epred42_2(X43,X44)<=>((epred41_2(X43,X44)|?[X66]:?[X67]:?[X68]:?[X69]:?[X70]:?[X71]:?[X72]:(((((((north(X44,X66)&north(X66,X67))&north(X67,X68))&north(X68,X69))&north(X69,X70))&north(X70,X71))&north(X71,X72))&north(X72,X43)))|?[X73]:?[X74]:?[X75]:?[X76]:?[X77]:?[X78]:?[X79]:?[X80]:((((((((north(X44,X73)&north(X73,X74))&north(X74,X75))&north(X75,X76))&north(X76,X77))&north(X77,X78))&north(X78,X79))&north(X79,X80))&north(X80,X43))))), inference(apply_def,[status(thm)],[c_0_285, c_0_286])).
fof(c_0_309, plain, (![X5]:![X6]:![X7]:![X8]:((~north(X6,X5)|epred36_2(X5,X6))&(~epred36_2(X7,X8)|north(X8,X7)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_287])])])])).
cnf(c_0_310,plain,(epred36_2(X1,X2)|~epred35_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_288])).
cnf(c_0_311,plain,(epred35_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_289, c_0_290])).
cnf(c_0_312,plain,(epred34_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_266, c_0_291])).
cnf(c_0_313,plain,(epred33_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_245, c_0_292])).
cnf(c_0_314,plain,(epred32_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_224, c_0_293])).
cnf(c_0_315,plain,(epred31_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_203, c_0_294])).
cnf(c_0_316,plain,(epred30_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_182, c_0_295])).
cnf(c_0_317,plain,(epred29_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_162, c_0_296])).
cnf(c_0_318,plain,(epred28_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_142, c_0_297])).
cnf(c_0_319,plain,(epred27_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_122, c_0_298])).
fof(c_0_320, plain, (![X2]:![X1]:(epred18_2(X1,X2)<=>(((((epred17_2(X1,X2)|(X2=cell86&X1=cell96))|(X2=cell87&X1=cell97))|(X2=cell88&X1=cell98))|(X2=cell89&X1=cell99))|(X2=cell90&X1=cell100)))), inference(apply_def,[status(thm)],[c_0_299, c_0_276])).
fof(c_0_321, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell85|(X3=cell84|(X3=cell83|(X3=cell82|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X3=cell84|(X3=cell83|(X3=cell82|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))&(((X3=cell85|(X4=cell94|(X3=cell83|(X3=cell82|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X4=cell94|(X3=cell83|(X3=cell82|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))))&((((X3=cell85|(X3=cell84|(X4=cell93|(X3=cell82|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X3=cell84|(X4=cell93|(X3=cell82|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))&(((X3=cell85|(X4=cell94|(X4=cell93|(X3=cell82|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X4=cell94|(X4=cell93|(X3=cell82|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))))&(((((X3=cell85|(X3=cell84|(X3=cell83|(X4=cell92|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X3=cell84|(X3=cell83|(X4=cell92|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))&(((X3=cell85|(X4=cell94|(X3=cell83|(X4=cell92|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X4=cell94|(X3=cell83|(X4=cell92|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))))&((((X3=cell85|(X3=cell84|(X4=cell93|(X4=cell92|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X3=cell84|(X4=cell93|(X4=cell92|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))&(((X3=cell85|(X4=cell94|(X4=cell93|(X4=cell92|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X4=cell94|(X4=cell93|(X4=cell92|(X3=cell81|epred16_2(X4,X3))))))|~epred17_2(X4,X3))))))&((((((X3=cell85|(X3=cell84|(X3=cell83|(X3=cell82|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X3=cell84|(X3=cell83|(X3=cell82|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))&(((X3=cell85|(X4=cell94|(X3=cell83|(X3=cell82|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X4=cell94|(X3=cell83|(X3=cell82|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))))&((((X3=cell85|(X3=cell84|(X4=cell93|(X3=cell82|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X3=cell84|(X4=cell93|(X3=cell82|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))&(((X3=cell85|(X4=cell94|(X4=cell93|(X3=cell82|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X4=cell94|(X4=cell93|(X3=cell82|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))))&(((((X3=cell85|(X3=cell84|(X3=cell83|(X4=cell92|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X3=cell84|(X3=cell83|(X4=cell92|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))&(((X3=cell85|(X4=cell94|(X3=cell83|(X4=cell92|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X4=cell94|(X3=cell83|(X4=cell92|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))))&((((X3=cell85|(X3=cell84|(X4=cell93|(X4=cell92|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X3=cell84|(X4=cell93|(X4=cell92|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))&(((X3=cell85|(X4=cell94|(X4=cell93|(X4=cell92|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3))&((X4=cell95|(X4=cell94|(X4=cell93|(X4=cell92|(X4=cell91|epred16_2(X4,X3))))))|~epred17_2(X4,X3)))))))&((((((~epred16_2(X6,X5)|epred17_2(X6,X5))&((X5!=cell81|X6!=cell91)|epred17_2(X6,X5)))&((X5!=cell82|X6!=cell92)|epred17_2(X6,X5)))&((X5!=cell83|X6!=cell93)|epred17_2(X6,X5)))&((X5!=cell84|X6!=cell94)|epred17_2(X6,X5)))&((X5!=cell85|X6!=cell95)|epred17_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_300])])])])])).
cnf(c_0_322,plain,(epred16_2(X1,X2)|~epred15_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_301])).
cnf(c_0_323,plain,(epred15_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_302, c_0_303])).
cnf(c_0_324,plain,(epred14_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_279, c_0_304])).
cnf(c_0_325,plain,(epred13_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_258, c_0_305])).
cnf(c_0_326,plain,(epred12_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_237, c_0_306])).
cnf(c_0_327,plain,(epred11_2(cell60,cell50)), inference(spm,[status(thm)],[c_0_216, c_0_307])).
cnf(c_0_328,plain,(epred12_2(X1,X2)|X1!=cell70|X2!=cell60), inference(split_conjunct,[status(thm)],[c_0_215])).
fof(c_0_329, plain, (![X81]:![X82]:![X98]:![X99]:![X100]:![X101]:![X102]:![X103]:![X104]:![X105]:![X106]:![X107]:![X108]:![X109]:![X110]:![X111]:![X112]:![X113]:![X114]:((((((((((((((((((north(X81,esk66_2(X81,X82))|(north(X81,esk59_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))&((north(esk66_2(X81,X82),esk67_2(X81,X82))|(north(X81,esk59_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk67_2(X81,X82),esk68_2(X81,X82))|(north(X81,esk59_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk68_2(X81,X82),esk69_2(X81,X82))|(north(X81,esk59_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk69_2(X81,X82),esk70_2(X81,X82))|(north(X81,esk59_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk70_2(X81,X82),esk71_2(X81,X82))|(north(X81,esk59_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk71_2(X81,X82),esk72_2(X81,X82))|(north(X81,esk59_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk72_2(X81,X82),esk73_2(X81,X82))|(north(X81,esk59_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk73_2(X81,X82),X82)|(north(X81,esk59_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((((((((((north(X81,esk66_2(X81,X82))|(north(esk59_2(X81,X82),esk60_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))&((north(esk66_2(X81,X82),esk67_2(X81,X82))|(north(esk59_2(X81,X82),esk60_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk67_2(X81,X82),esk68_2(X81,X82))|(north(esk59_2(X81,X82),esk60_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk68_2(X81,X82),esk69_2(X81,X82))|(north(esk59_2(X81,X82),esk60_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk69_2(X81,X82),esk70_2(X81,X82))|(north(esk59_2(X81,X82),esk60_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk70_2(X81,X82),esk71_2(X81,X82))|(north(esk59_2(X81,X82),esk60_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk71_2(X81,X82),esk72_2(X81,X82))|(north(esk59_2(X81,X82),esk60_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk72_2(X81,X82),esk73_2(X81,X82))|(north(esk59_2(X81,X82),esk60_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk73_2(X81,X82),X82)|(north(esk59_2(X81,X82),esk60_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))))&((((((((((north(X81,esk66_2(X81,X82))|(north(esk60_2(X81,X82),esk61_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))&((north(esk66_2(X81,X82),esk67_2(X81,X82))|(north(esk60_2(X81,X82),esk61_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk67_2(X81,X82),esk68_2(X81,X82))|(north(esk60_2(X81,X82),esk61_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk68_2(X81,X82),esk69_2(X81,X82))|(north(esk60_2(X81,X82),esk61_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk69_2(X81,X82),esk70_2(X81,X82))|(north(esk60_2(X81,X82),esk61_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk70_2(X81,X82),esk71_2(X81,X82))|(north(esk60_2(X81,X82),esk61_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk71_2(X81,X82),esk72_2(X81,X82))|(north(esk60_2(X81,X82),esk61_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk72_2(X81,X82),esk73_2(X81,X82))|(north(esk60_2(X81,X82),esk61_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk73_2(X81,X82),X82)|(north(esk60_2(X81,X82),esk61_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))))&((((((((((north(X81,esk66_2(X81,X82))|(north(esk61_2(X81,X82),esk62_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))&((north(esk66_2(X81,X82),esk67_2(X81,X82))|(north(esk61_2(X81,X82),esk62_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk67_2(X81,X82),esk68_2(X81,X82))|(north(esk61_2(X81,X82),esk62_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk68_2(X81,X82),esk69_2(X81,X82))|(north(esk61_2(X81,X82),esk62_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk69_2(X81,X82),esk70_2(X81,X82))|(north(esk61_2(X81,X82),esk62_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk70_2(X81,X82),esk71_2(X81,X82))|(north(esk61_2(X81,X82),esk62_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk71_2(X81,X82),esk72_2(X81,X82))|(north(esk61_2(X81,X82),esk62_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk72_2(X81,X82),esk73_2(X81,X82))|(north(esk61_2(X81,X82),esk62_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk73_2(X81,X82),X82)|(north(esk61_2(X81,X82),esk62_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))))&((((((((((north(X81,esk66_2(X81,X82))|(north(esk62_2(X81,X82),esk63_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))&((north(esk66_2(X81,X82),esk67_2(X81,X82))|(north(esk62_2(X81,X82),esk63_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk67_2(X81,X82),esk68_2(X81,X82))|(north(esk62_2(X81,X82),esk63_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk68_2(X81,X82),esk69_2(X81,X82))|(north(esk62_2(X81,X82),esk63_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk69_2(X81,X82),esk70_2(X81,X82))|(north(esk62_2(X81,X82),esk63_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk70_2(X81,X82),esk71_2(X81,X82))|(north(esk62_2(X81,X82),esk63_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk71_2(X81,X82),esk72_2(X81,X82))|(north(esk62_2(X81,X82),esk63_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk72_2(X81,X82),esk73_2(X81,X82))|(north(esk62_2(X81,X82),esk63_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk73_2(X81,X82),X82)|(north(esk62_2(X81,X82),esk63_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))))&((((((((((north(X81,esk66_2(X81,X82))|(north(esk63_2(X81,X82),esk64_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))&((north(esk66_2(X81,X82),esk67_2(X81,X82))|(north(esk63_2(X81,X82),esk64_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk67_2(X81,X82),esk68_2(X81,X82))|(north(esk63_2(X81,X82),esk64_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk68_2(X81,X82),esk69_2(X81,X82))|(north(esk63_2(X81,X82),esk64_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk69_2(X81,X82),esk70_2(X81,X82))|(north(esk63_2(X81,X82),esk64_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk70_2(X81,X82),esk71_2(X81,X82))|(north(esk63_2(X81,X82),esk64_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk71_2(X81,X82),esk72_2(X81,X82))|(north(esk63_2(X81,X82),esk64_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk72_2(X81,X82),esk73_2(X81,X82))|(north(esk63_2(X81,X82),esk64_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk73_2(X81,X82),X82)|(north(esk63_2(X81,X82),esk64_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))))&((((((((((north(X81,esk66_2(X81,X82))|(north(esk64_2(X81,X82),esk65_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))&((north(esk66_2(X81,X82),esk67_2(X81,X82))|(north(esk64_2(X81,X82),esk65_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk67_2(X81,X82),esk68_2(X81,X82))|(north(esk64_2(X81,X82),esk65_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk68_2(X81,X82),esk69_2(X81,X82))|(north(esk64_2(X81,X82),esk65_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk69_2(X81,X82),esk70_2(X81,X82))|(north(esk64_2(X81,X82),esk65_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk70_2(X81,X82),esk71_2(X81,X82))|(north(esk64_2(X81,X82),esk65_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk71_2(X81,X82),esk72_2(X81,X82))|(north(esk64_2(X81,X82),esk65_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk72_2(X81,X82),esk73_2(X81,X82))|(north(esk64_2(X81,X82),esk65_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk73_2(X81,X82),X82)|(north(esk64_2(X81,X82),esk65_2(X81,X82))|epred41_2(X82,X81)))|~epred42_2(X82,X81))))&((((((((((north(X81,esk66_2(X81,X82))|(north(esk65_2(X81,X82),X82)|epred41_2(X82,X81)))|~epred42_2(X82,X81))&((north(esk66_2(X81,X82),esk67_2(X81,X82))|(north(esk65_2(X81,X82),X82)|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk67_2(X81,X82),esk68_2(X81,X82))|(north(esk65_2(X81,X82),X82)|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk68_2(X81,X82),esk69_2(X81,X82))|(north(esk65_2(X81,X82),X82)|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk69_2(X81,X82),esk70_2(X81,X82))|(north(esk65_2(X81,X82),X82)|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk70_2(X81,X82),esk71_2(X81,X82))|(north(esk65_2(X81,X82),X82)|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk71_2(X81,X82),esk72_2(X81,X82))|(north(esk65_2(X81,X82),X82)|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk72_2(X81,X82),esk73_2(X81,X82))|(north(esk65_2(X81,X82),X82)|epred41_2(X82,X81)))|~epred42_2(X82,X81)))&((north(esk73_2(X81,X82),X82)|(north(esk65_2(X81,X82),X82)|epred41_2(X82,X81)))|~epred42_2(X82,X81))))&(((~epred41_2(X99,X98)|epred42_2(X99,X98))&((((((((~north(X98,X100)|~north(X100,X101))|~north(X101,X102))|~north(X102,X103))|~north(X103,X104))|~north(X104,X105))|~north(X105,X106))|~north(X106,X99))|epred42_2(X99,X98)))&(((((((((~north(X98,X107)|~north(X107,X108))|~north(X108,X109))|~north(X109,X110))|~north(X110,X111))|~north(X111,X112))|~north(X112,X113))|~north(X113,X114))|~north(X114,X99))|epred42_2(X99,X98))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_308])])])])])])).
cnf(c_0_330,plain,(north(X1,X2)|~epred36_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_309])).
cnf(c_0_331,plain,(epred36_2(cell3,cell2)), inference(spm,[status(thm)],[c_0_310, c_0_311])).
cnf(c_0_332,plain,(epred35_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_289, c_0_312])).
cnf(c_0_333,plain,(epred34_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_266, c_0_313])).
cnf(c_0_334,plain,(epred33_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_245, c_0_314])).
cnf(c_0_335,plain,(epred32_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_224, c_0_315])).
cnf(c_0_336,plain,(epred31_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_203, c_0_316])).
cnf(c_0_337,plain,(epred30_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_182, c_0_317])).
cnf(c_0_338,plain,(epred29_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_162, c_0_318])).
cnf(c_0_339,plain,(epred28_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_142, c_0_319])).
fof(c_0_340, plain, (![X6]:![X5]:(epred39_2(X5,X6)<=>((((((((west(X6,X5)|?[X7]:(west(X6,X7)&west(X7,X5)))|?[X8]:?[X9]:((west(X6,X8)&west(X8,X9))&west(X9,X5)))|?[X10]:?[X11]:?[X12]:(((west(X6,X10)&west(X10,X11))&west(X11,X12))&west(X12,X5)))|?[X13]:?[X14]:?[X15]:?[X16]:((((west(X6,X13)&west(X13,X14))&west(X14,X15))&west(X15,X16))&west(X16,X5)))|?[X17]:?[X18]:?[X19]:?[X20]:?[X21]:(((((west(X6,X17)&west(X17,X18))&west(X18,X19))&west(X19,X20))&west(X20,X21))&west(X21,X5)))|?[X22]:?[X23]:?[X24]:?[X25]:?[X26]:?[X27]:((((((west(X6,X22)&west(X22,X23))&west(X23,X24))&west(X24,X25))&west(X25,X26))&west(X26,X27))&west(X27,X5)))|?[X28]:?[X29]:?[X30]:?[X31]:?[X32]:?[X33]:?[X34]:(((((((west(X6,X28)&west(X28,X29))&west(X29,X30))&west(X30,X31))&west(X31,X32))&west(X32,X33))&west(X33,X34))&west(X34,X5)))|?[X35]:?[X36]:?[X37]:?[X38]:?[X39]:?[X40]:?[X41]:?[X42]:((((((((west(X6,X35)&west(X35,X36))&west(X36,X37))&west(X37,X38))&west(X38,X39))&west(X39,X40))&west(X40,X41))&west(X41,X42))&west(X42,X5))))), introduced(definition)).
fof(c_0_341, plain, (![X6]:![X5]:(epred38_2(X5,X6)<=>((((((west(X6,X5)|?[X7]:(west(X6,X7)&west(X7,X5)))|?[X8]:?[X9]:((west(X6,X8)&west(X8,X9))&west(X9,X5)))|?[X10]:?[X11]:?[X12]:(((west(X6,X10)&west(X10,X11))&west(X11,X12))&west(X12,X5)))|?[X13]:?[X14]:?[X15]:?[X16]:((((west(X6,X13)&west(X13,X14))&west(X14,X15))&west(X15,X16))&west(X16,X5)))|?[X17]:?[X18]:?[X19]:?[X20]:?[X21]:(((((west(X6,X17)&west(X17,X18))&west(X18,X19))&west(X19,X20))&west(X20,X21))&west(X21,X5)))|?[X22]:?[X23]:?[X24]:?[X25]:?[X26]:?[X27]:((((((west(X6,X22)&west(X22,X23))&west(X23,X24))&west(X24,X25))&west(X25,X26))&west(X26,X27))&west(X27,X5))))), introduced(definition)).
fof(c_0_342, axiom, (![X1]:![X2]:(west(X2,X1)<=>epred18_2(X1,X2))), inference(apply_def,[status(thm)],[c_0_1, c_0_299])).
fof(c_0_343, plain, (![X3]:![X4]:![X5]:![X6]:((((((((X3=cell90|(X3=cell89|(X3=cell88|(X3=cell87|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X3=cell89|(X3=cell88|(X3=cell87|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))&(((X3=cell90|(X4=cell99|(X3=cell88|(X3=cell87|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X4=cell99|(X3=cell88|(X3=cell87|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))))&((((X3=cell90|(X3=cell89|(X4=cell98|(X3=cell87|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X3=cell89|(X4=cell98|(X3=cell87|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))&(((X3=cell90|(X4=cell99|(X4=cell98|(X3=cell87|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X4=cell99|(X4=cell98|(X3=cell87|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))))&(((((X3=cell90|(X3=cell89|(X3=cell88|(X4=cell97|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X3=cell89|(X3=cell88|(X4=cell97|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))&(((X3=cell90|(X4=cell99|(X3=cell88|(X4=cell97|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X4=cell99|(X3=cell88|(X4=cell97|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))))&((((X3=cell90|(X3=cell89|(X4=cell98|(X4=cell97|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X3=cell89|(X4=cell98|(X4=cell97|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))&(((X3=cell90|(X4=cell99|(X4=cell98|(X4=cell97|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X4=cell99|(X4=cell98|(X4=cell97|(X3=cell86|epred17_2(X4,X3))))))|~epred18_2(X4,X3))))))&((((((X3=cell90|(X3=cell89|(X3=cell88|(X3=cell87|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X3=cell89|(X3=cell88|(X3=cell87|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))&(((X3=cell90|(X4=cell99|(X3=cell88|(X3=cell87|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X4=cell99|(X3=cell88|(X3=cell87|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))))&((((X3=cell90|(X3=cell89|(X4=cell98|(X3=cell87|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X3=cell89|(X4=cell98|(X3=cell87|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))&(((X3=cell90|(X4=cell99|(X4=cell98|(X3=cell87|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X4=cell99|(X4=cell98|(X3=cell87|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))))&(((((X3=cell90|(X3=cell89|(X3=cell88|(X4=cell97|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X3=cell89|(X3=cell88|(X4=cell97|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))&(((X3=cell90|(X4=cell99|(X3=cell88|(X4=cell97|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X4=cell99|(X3=cell88|(X4=cell97|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))))&((((X3=cell90|(X3=cell89|(X4=cell98|(X4=cell97|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X3=cell89|(X4=cell98|(X4=cell97|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))&(((X3=cell90|(X4=cell99|(X4=cell98|(X4=cell97|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3))&((X4=cell100|(X4=cell99|(X4=cell98|(X4=cell97|(X4=cell96|epred17_2(X4,X3))))))|~epred18_2(X4,X3)))))))&((((((~epred17_2(X6,X5)|epred18_2(X6,X5))&((X5!=cell86|X6!=cell96)|epred18_2(X6,X5)))&((X5!=cell87|X6!=cell97)|epred18_2(X6,X5)))&((X5!=cell88|X6!=cell98)|epred18_2(X6,X5)))&((X5!=cell89|X6!=cell99)|epred18_2(X6,X5)))&((X5!=cell90|X6!=cell100)|epred18_2(X6,X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_320])])])])])).
cnf(c_0_344,plain,(epred17_2(X1,X2)|~epred16_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_321])).
cnf(c_0_345,plain,(epred16_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_322, c_0_323])).
cnf(c_0_346,plain,(epred15_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_302, c_0_324])).
cnf(c_0_347,plain,(epred14_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_279, c_0_325])).
cnf(c_0_348,plain,(epred13_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_258, c_0_326])).
cnf(c_0_349,plain,(epred12_2(cell60,cell50)), inference(spm,[status(thm)],[c_0_237, c_0_327])).
cnf(c_0_350,plain,(epred12_2(X1,cell60)|cell70!=X1), inference(er,[status(thm)],[c_0_328])).
cnf(c_0_351,plain,(epred42_2(X1,X2)|~north(X3,X1)|~north(X4,X3)|~north(X5,X4)|~north(X6,X5)|~north(X7,X6)|~north(X8,X7)|~north(X9,X8)|~north(X10,X9)|~north(X2,X10)), inference(split_conjunct,[status(thm)],[c_0_329])).
cnf(c_0_352,plain,(north(cell2,cell3)), inference(spm,[status(thm)],[c_0_330, c_0_331])).
cnf(c_0_353,plain,(epred36_2(cell4,cell3)), inference(spm,[status(thm)],[c_0_310, c_0_332])).
cnf(c_0_354,plain,(epred35_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_289, c_0_333])).
cnf(c_0_355,plain,(epred34_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_266, c_0_334])).
cnf(c_0_356,plain,(epred33_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_245, c_0_335])).
cnf(c_0_357,plain,(epred32_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_224, c_0_336])).
cnf(c_0_358,plain,(epred31_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_203, c_0_337])).
cnf(c_0_359,plain,(epred30_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_182, c_0_338])).
cnf(c_0_360,plain,(epred29_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_162, c_0_339])).
fof(c_0_361, plain, (![X6]:![X5]:(epred39_2(X5,X6)<=>((epred38_2(X5,X6)|?[X28]:?[X29]:?[X30]:?[X31]:?[X32]:?[X33]:?[X34]:(((((((west(X6,X28)&west(X28,X29))&west(X29,X30))&west(X30,X31))&west(X31,X32))&west(X32,X33))&west(X33,X34))&west(X34,X5)))|?[X35]:?[X36]:?[X37]:?[X38]:?[X39]:?[X40]:?[X41]:?[X42]:((((((((west(X6,X35)&west(X35,X36))&west(X36,X37))&west(X37,X38))&west(X38,X39))&west(X39,X40))&west(X40,X41))&west(X41,X42))&west(X42,X5))))), inference(apply_def,[status(thm)],[c_0_340, c_0_341])).
fof(c_0_362, plain, (![X3]:![X4]:![X5]:![X6]:((~west(X4,X3)|epred18_2(X3,X4))&(~epred18_2(X5,X6)|west(X6,X5)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_342])])])])).
cnf(c_0_363,plain,(epred18_2(X1,X2)|~epred17_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_343])).
cnf(c_0_364,plain,(epred17_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_344, c_0_345])).
cnf(c_0_365,plain,(epred16_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_322, c_0_346])).
cnf(c_0_366,plain,(epred15_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_302, c_0_347])).
cnf(c_0_367,plain,(epred14_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_279, c_0_348])).
cnf(c_0_368,plain,(epred13_2(cell60,cell50)), inference(spm,[status(thm)],[c_0_258, c_0_349])).
cnf(c_0_369,plain,(epred12_2(cell70,cell60)), inference(er,[status(thm)],[c_0_350])).
cnf(c_0_370,plain,(epred42_2(X1,X2)|~north(cell3,X3)|~north(X2,cell2)|~north(X3,X4)|~north(X4,X5)|~north(X5,X6)|~north(X6,X7)|~north(X7,X8)|~north(X8,X1)), inference(spm,[status(thm)],[c_0_351, c_0_352])).
cnf(c_0_371,plain,(north(cell3,cell4)), inference(spm,[status(thm)],[c_0_330, c_0_353])).
cnf(c_0_372,plain,(epred36_2(cell2,cell1)), inference(spm,[status(thm)],[c_0_310, c_0_354])).
cnf(c_0_373,plain,(epred35_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_289, c_0_355])).
cnf(c_0_374,plain,(epred34_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_266, c_0_356])).
cnf(c_0_375,plain,(epred33_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_245, c_0_357])).
cnf(c_0_376,plain,(epred32_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_224, c_0_358])).
cnf(c_0_377,plain,(epred31_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_203, c_0_359])).
cnf(c_0_378,plain,(epred30_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_182, c_0_360])).
fof(c_0_379, plain, (![X43]:![X44]:![X60]:![X61]:![X62]:![X63]:![X64]:![X65]:![X66]:![X67]:![X68]:![X69]:![X70]:![X71]:![X72]:![X73]:![X74]:![X75]:![X76]:((((((((((((((((((west(X43,esk30_2(X43,X44))|(west(X43,esk23_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))&((west(esk30_2(X43,X44),esk31_2(X43,X44))|(west(X43,esk23_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk31_2(X43,X44),esk32_2(X43,X44))|(west(X43,esk23_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk32_2(X43,X44),esk33_2(X43,X44))|(west(X43,esk23_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk33_2(X43,X44),esk34_2(X43,X44))|(west(X43,esk23_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk34_2(X43,X44),esk35_2(X43,X44))|(west(X43,esk23_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk35_2(X43,X44),esk36_2(X43,X44))|(west(X43,esk23_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk36_2(X43,X44),esk37_2(X43,X44))|(west(X43,esk23_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk37_2(X43,X44),X44)|(west(X43,esk23_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((((((((((west(X43,esk30_2(X43,X44))|(west(esk23_2(X43,X44),esk24_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))&((west(esk30_2(X43,X44),esk31_2(X43,X44))|(west(esk23_2(X43,X44),esk24_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk31_2(X43,X44),esk32_2(X43,X44))|(west(esk23_2(X43,X44),esk24_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk32_2(X43,X44),esk33_2(X43,X44))|(west(esk23_2(X43,X44),esk24_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk33_2(X43,X44),esk34_2(X43,X44))|(west(esk23_2(X43,X44),esk24_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk34_2(X43,X44),esk35_2(X43,X44))|(west(esk23_2(X43,X44),esk24_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk35_2(X43,X44),esk36_2(X43,X44))|(west(esk23_2(X43,X44),esk24_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk36_2(X43,X44),esk37_2(X43,X44))|(west(esk23_2(X43,X44),esk24_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk37_2(X43,X44),X44)|(west(esk23_2(X43,X44),esk24_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))))&((((((((((west(X43,esk30_2(X43,X44))|(west(esk24_2(X43,X44),esk25_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))&((west(esk30_2(X43,X44),esk31_2(X43,X44))|(west(esk24_2(X43,X44),esk25_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk31_2(X43,X44),esk32_2(X43,X44))|(west(esk24_2(X43,X44),esk25_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk32_2(X43,X44),esk33_2(X43,X44))|(west(esk24_2(X43,X44),esk25_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk33_2(X43,X44),esk34_2(X43,X44))|(west(esk24_2(X43,X44),esk25_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk34_2(X43,X44),esk35_2(X43,X44))|(west(esk24_2(X43,X44),esk25_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk35_2(X43,X44),esk36_2(X43,X44))|(west(esk24_2(X43,X44),esk25_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk36_2(X43,X44),esk37_2(X43,X44))|(west(esk24_2(X43,X44),esk25_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk37_2(X43,X44),X44)|(west(esk24_2(X43,X44),esk25_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))))&((((((((((west(X43,esk30_2(X43,X44))|(west(esk25_2(X43,X44),esk26_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))&((west(esk30_2(X43,X44),esk31_2(X43,X44))|(west(esk25_2(X43,X44),esk26_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk31_2(X43,X44),esk32_2(X43,X44))|(west(esk25_2(X43,X44),esk26_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk32_2(X43,X44),esk33_2(X43,X44))|(west(esk25_2(X43,X44),esk26_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk33_2(X43,X44),esk34_2(X43,X44))|(west(esk25_2(X43,X44),esk26_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk34_2(X43,X44),esk35_2(X43,X44))|(west(esk25_2(X43,X44),esk26_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk35_2(X43,X44),esk36_2(X43,X44))|(west(esk25_2(X43,X44),esk26_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk36_2(X43,X44),esk37_2(X43,X44))|(west(esk25_2(X43,X44),esk26_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk37_2(X43,X44),X44)|(west(esk25_2(X43,X44),esk26_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))))&((((((((((west(X43,esk30_2(X43,X44))|(west(esk26_2(X43,X44),esk27_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))&((west(esk30_2(X43,X44),esk31_2(X43,X44))|(west(esk26_2(X43,X44),esk27_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk31_2(X43,X44),esk32_2(X43,X44))|(west(esk26_2(X43,X44),esk27_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk32_2(X43,X44),esk33_2(X43,X44))|(west(esk26_2(X43,X44),esk27_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk33_2(X43,X44),esk34_2(X43,X44))|(west(esk26_2(X43,X44),esk27_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk34_2(X43,X44),esk35_2(X43,X44))|(west(esk26_2(X43,X44),esk27_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk35_2(X43,X44),esk36_2(X43,X44))|(west(esk26_2(X43,X44),esk27_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk36_2(X43,X44),esk37_2(X43,X44))|(west(esk26_2(X43,X44),esk27_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk37_2(X43,X44),X44)|(west(esk26_2(X43,X44),esk27_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))))&((((((((((west(X43,esk30_2(X43,X44))|(west(esk27_2(X43,X44),esk28_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))&((west(esk30_2(X43,X44),esk31_2(X43,X44))|(west(esk27_2(X43,X44),esk28_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk31_2(X43,X44),esk32_2(X43,X44))|(west(esk27_2(X43,X44),esk28_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk32_2(X43,X44),esk33_2(X43,X44))|(west(esk27_2(X43,X44),esk28_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk33_2(X43,X44),esk34_2(X43,X44))|(west(esk27_2(X43,X44),esk28_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk34_2(X43,X44),esk35_2(X43,X44))|(west(esk27_2(X43,X44),esk28_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk35_2(X43,X44),esk36_2(X43,X44))|(west(esk27_2(X43,X44),esk28_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk36_2(X43,X44),esk37_2(X43,X44))|(west(esk27_2(X43,X44),esk28_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk37_2(X43,X44),X44)|(west(esk27_2(X43,X44),esk28_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))))&((((((((((west(X43,esk30_2(X43,X44))|(west(esk28_2(X43,X44),esk29_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))&((west(esk30_2(X43,X44),esk31_2(X43,X44))|(west(esk28_2(X43,X44),esk29_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk31_2(X43,X44),esk32_2(X43,X44))|(west(esk28_2(X43,X44),esk29_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk32_2(X43,X44),esk33_2(X43,X44))|(west(esk28_2(X43,X44),esk29_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk33_2(X43,X44),esk34_2(X43,X44))|(west(esk28_2(X43,X44),esk29_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk34_2(X43,X44),esk35_2(X43,X44))|(west(esk28_2(X43,X44),esk29_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk35_2(X43,X44),esk36_2(X43,X44))|(west(esk28_2(X43,X44),esk29_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk36_2(X43,X44),esk37_2(X43,X44))|(west(esk28_2(X43,X44),esk29_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk37_2(X43,X44),X44)|(west(esk28_2(X43,X44),esk29_2(X43,X44))|epred38_2(X44,X43)))|~epred39_2(X44,X43))))&((((((((((west(X43,esk30_2(X43,X44))|(west(esk29_2(X43,X44),X44)|epred38_2(X44,X43)))|~epred39_2(X44,X43))&((west(esk30_2(X43,X44),esk31_2(X43,X44))|(west(esk29_2(X43,X44),X44)|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk31_2(X43,X44),esk32_2(X43,X44))|(west(esk29_2(X43,X44),X44)|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk32_2(X43,X44),esk33_2(X43,X44))|(west(esk29_2(X43,X44),X44)|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk33_2(X43,X44),esk34_2(X43,X44))|(west(esk29_2(X43,X44),X44)|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk34_2(X43,X44),esk35_2(X43,X44))|(west(esk29_2(X43,X44),X44)|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk35_2(X43,X44),esk36_2(X43,X44))|(west(esk29_2(X43,X44),X44)|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk36_2(X43,X44),esk37_2(X43,X44))|(west(esk29_2(X43,X44),X44)|epred38_2(X44,X43)))|~epred39_2(X44,X43)))&((west(esk37_2(X43,X44),X44)|(west(esk29_2(X43,X44),X44)|epred38_2(X44,X43)))|~epred39_2(X44,X43))))&(((~epred38_2(X61,X60)|epred39_2(X61,X60))&((((((((~west(X60,X62)|~west(X62,X63))|~west(X63,X64))|~west(X64,X65))|~west(X65,X66))|~west(X66,X67))|~west(X67,X68))|~west(X68,X61))|epred39_2(X61,X60)))&(((((((((~west(X60,X69)|~west(X69,X70))|~west(X70,X71))|~west(X71,X72))|~west(X72,X73))|~west(X73,X74))|~west(X74,X75))|~west(X75,X76))|~west(X76,X61))|epred39_2(X61,X60))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_361])])])])])])).
cnf(c_0_380,plain,(west(X1,X2)|~epred18_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_362])).
cnf(c_0_381,plain,(epred18_2(cell30,cell20)), inference(spm,[status(thm)],[c_0_363, c_0_364])).
cnf(c_0_382,plain,(epred17_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_344, c_0_365])).
cnf(c_0_383,plain,(epred16_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_322, c_0_366])).
cnf(c_0_384,plain,(epred15_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_302, c_0_367])).
cnf(c_0_385,plain,(epred14_2(cell60,cell50)), inference(spm,[status(thm)],[c_0_279, c_0_368])).
cnf(c_0_386,plain,(epred13_2(cell70,cell60)), inference(spm,[status(thm)],[c_0_258, c_0_369])).
cnf(c_0_387,plain,(epred14_2(X1,X2)|X1!=cell80|X2!=cell70), inference(split_conjunct,[status(thm)],[c_0_257])).
cnf(c_0_388,plain,(epred42_2(X1,X2)|~north(X2,cell2)|~north(cell4,X3)|~north(X3,X4)|~north(X4,X5)|~north(X5,X6)|~north(X6,X7)|~north(X7,X1)), inference(spm,[status(thm)],[c_0_370, c_0_371])).
cnf(c_0_389,plain,(north(cell1,cell2)), inference(spm,[status(thm)],[c_0_330, c_0_372])).
cnf(c_0_390,plain,(epred36_2(cell5,cell4)), inference(spm,[status(thm)],[c_0_310, c_0_373])).
cnf(c_0_391,plain,(epred35_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_289, c_0_374])).
cnf(c_0_392,plain,(epred34_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_266, c_0_375])).
cnf(c_0_393,plain,(epred33_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_245, c_0_376])).
cnf(c_0_394,plain,(epred32_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_224, c_0_377])).
cnf(c_0_395,plain,(epred31_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_203, c_0_378])).
cnf(c_0_396,plain,(epred39_2(X1,X2)|~west(X3,X1)|~west(X4,X3)|~west(X5,X4)|~west(X6,X5)|~west(X7,X6)|~west(X8,X7)|~west(X9,X8)|~west(X10,X9)|~west(X2,X10)), inference(split_conjunct,[status(thm)],[c_0_379])).
cnf(c_0_397,plain,(west(cell20,cell30)), inference(spm,[status(thm)],[c_0_380, c_0_381])).
cnf(c_0_398,plain,(epred18_2(cell40,cell30)), inference(spm,[status(thm)],[c_0_363, c_0_382])).
cnf(c_0_399,plain,(epred17_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_344, c_0_383])).
cnf(c_0_400,plain,(epred16_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_322, c_0_384])).
cnf(c_0_401,plain,(epred15_2(cell60,cell50)), inference(spm,[status(thm)],[c_0_302, c_0_385])).
cnf(c_0_402,plain,(epred14_2(cell70,cell60)), inference(spm,[status(thm)],[c_0_279, c_0_386])).
cnf(c_0_403,plain,(epred14_2(X1,cell70)|cell80!=X1), inference(er,[status(thm)],[c_0_387])).
cnf(c_0_404,plain,(epred42_2(X1,cell1)|~north(cell4,X2)|~north(X2,X3)|~north(X3,X4)|~north(X4,X5)|~north(X5,X6)|~north(X6,X1)), inference(spm,[status(thm)],[c_0_388, c_0_389])).
cnf(c_0_405,plain,(north(cell4,cell5)), inference(spm,[status(thm)],[c_0_330, c_0_390])).
cnf(c_0_406,plain,(epred36_2(cell6,cell5)), inference(spm,[status(thm)],[c_0_310, c_0_391])).
cnf(c_0_407,plain,(epred35_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_289, c_0_392])).
cnf(c_0_408,plain,(epred34_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_266, c_0_393])).
cnf(c_0_409,plain,(epred33_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_245, c_0_394])).
cnf(c_0_410,plain,(epred32_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_224, c_0_395])).
cnf(c_0_411,plain,(epred39_2(X1,X2)|~west(cell30,X3)|~west(X2,cell20)|~west(X3,X4)|~west(X4,X5)|~west(X5,X6)|~west(X6,X7)|~west(X7,X8)|~west(X8,X1)), inference(spm,[status(thm)],[c_0_396, c_0_397])).
cnf(c_0_412,plain,(west(cell30,cell40)), inference(spm,[status(thm)],[c_0_380, c_0_398])).
cnf(c_0_413,plain,(epred18_2(cell20,cell10)), inference(spm,[status(thm)],[c_0_363, c_0_399])).
cnf(c_0_414,plain,(epred17_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_344, c_0_400])).
cnf(c_0_415,plain,(epred16_2(cell60,cell50)), inference(spm,[status(thm)],[c_0_322, c_0_401])).
cnf(c_0_416,plain,(epred15_2(cell70,cell60)), inference(spm,[status(thm)],[c_0_302, c_0_402])).
cnf(c_0_417,plain,(epred14_2(cell80,cell70)), inference(er,[status(thm)],[c_0_403])).
cnf(c_0_418,plain,(epred42_2(X1,cell1)|~north(cell5,X2)|~north(X2,X3)|~north(X3,X4)|~north(X4,X5)|~north(X5,X1)), inference(spm,[status(thm)],[c_0_404, c_0_405])).
cnf(c_0_419,plain,(north(cell5,cell6)), inference(spm,[status(thm)],[c_0_330, c_0_406])).
cnf(c_0_420,plain,(epred36_2(cell7,cell6)), inference(spm,[status(thm)],[c_0_310, c_0_407])).
cnf(c_0_421,plain,(epred35_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_289, c_0_408])).
cnf(c_0_422,plain,(epred34_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_266, c_0_409])).
cnf(c_0_423,plain,(epred33_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_245, c_0_410])).
cnf(c_0_424,plain,(epred39_2(X1,X2)|~west(X2,cell20)|~west(cell40,X3)|~west(X3,X4)|~west(X4,X5)|~west(X5,X6)|~west(X6,X7)|~west(X7,X1)), inference(spm,[status(thm)],[c_0_411, c_0_412])).
cnf(c_0_425,plain,(west(cell10,cell20)), inference(spm,[status(thm)],[c_0_380, c_0_413])).
cnf(c_0_426,plain,(epred18_2(cell50,cell40)), inference(spm,[status(thm)],[c_0_363, c_0_414])).
cnf(c_0_427,plain,(epred17_2(cell60,cell50)), inference(spm,[status(thm)],[c_0_344, c_0_415])).
cnf(c_0_428,plain,(epred16_2(cell70,cell60)), inference(spm,[status(thm)],[c_0_322, c_0_416])).
cnf(c_0_429,plain,(epred15_2(cell80,cell70)), inference(spm,[status(thm)],[c_0_302, c_0_417])).
cnf(c_0_430,plain,(epred16_2(X1,X2)|X1!=cell90|X2!=cell80), inference(split_conjunct,[status(thm)],[c_0_301])).
cnf(c_0_431,plain,(epred42_2(X1,cell1)|~north(cell6,X2)|~north(X2,X3)|~north(X3,X4)|~north(X4,X1)), inference(spm,[status(thm)],[c_0_418, c_0_419])).
cnf(c_0_432,plain,(north(cell6,cell7)), inference(spm,[status(thm)],[c_0_330, c_0_420])).
cnf(c_0_433,plain,(epred36_2(cell8,cell7)), inference(spm,[status(thm)],[c_0_310, c_0_421])).
cnf(c_0_434,plain,(epred35_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_289, c_0_422])).
cnf(c_0_435,plain,(epred34_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_266, c_0_423])).
cnf(c_0_436,plain,(epred39_2(X1,cell10)|~west(cell40,X2)|~west(X2,X3)|~west(X3,X4)|~west(X4,X5)|~west(X5,X6)|~west(X6,X1)), inference(spm,[status(thm)],[c_0_424, c_0_425])).
cnf(c_0_437,plain,(west(cell40,cell50)), inference(spm,[status(thm)],[c_0_380, c_0_426])).
cnf(c_0_438,plain,(epred18_2(cell60,cell50)), inference(spm,[status(thm)],[c_0_363, c_0_427])).
cnf(c_0_439,plain,(epred17_2(cell70,cell60)), inference(spm,[status(thm)],[c_0_344, c_0_428])).
cnf(c_0_440,plain,(epred16_2(cell80,cell70)), inference(spm,[status(thm)],[c_0_322, c_0_429])).
cnf(c_0_441,plain,(epred16_2(X1,cell80)|cell90!=X1), inference(er,[status(thm)],[c_0_430])).
cnf(c_0_442,plain,(epred42_2(X1,cell1)|~north(cell7,X2)|~north(X2,X3)|~north(X3,X1)), inference(spm,[status(thm)],[c_0_431, c_0_432])).
cnf(c_0_443,plain,(north(cell7,cell8)), inference(spm,[status(thm)],[c_0_330, c_0_433])).
cnf(c_0_444,plain,(epred36_2(cell9,cell8)), inference(spm,[status(thm)],[c_0_310, c_0_434])).
cnf(c_0_445,plain,(epred35_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_289, c_0_435])).
cnf(c_0_446,plain,(epred39_2(X1,cell10)|~west(cell50,X2)|~west(X2,X3)|~west(X3,X4)|~west(X4,X5)|~west(X5,X1)), inference(spm,[status(thm)],[c_0_436, c_0_437])).
cnf(c_0_447,plain,(west(cell50,cell60)), inference(spm,[status(thm)],[c_0_380, c_0_438])).
cnf(c_0_448,plain,(epred18_2(cell70,cell60)), inference(spm,[status(thm)],[c_0_363, c_0_439])).
cnf(c_0_449,plain,(epred17_2(cell80,cell70)), inference(spm,[status(thm)],[c_0_344, c_0_440])).
cnf(c_0_450,plain,(epred16_2(cell90,cell80)), inference(er,[status(thm)],[c_0_441])).
fof(c_0_451, axiom, (![X43]:![X44]:(duenorth(X44,X43)<=>epred42_2(X43,X44))), inference(apply_def,[status(thm)],[c_0_2, c_0_285])).
cnf(c_0_452,plain,(epred42_2(X1,cell1)|~north(cell8,X2)|~north(X2,X1)), inference(spm,[status(thm)],[c_0_442, c_0_443])).
cnf(c_0_453,plain,(north(cell8,cell9)), inference(spm,[status(thm)],[c_0_330, c_0_444])).
cnf(c_0_454,plain,(epred36_2(cell10,cell9)), inference(spm,[status(thm)],[c_0_310, c_0_445])).
cnf(c_0_455,plain,(epred39_2(X1,cell10)|~west(cell60,X2)|~west(X2,X3)|~west(X3,X4)|~west(X4,X1)), inference(spm,[status(thm)],[c_0_446, c_0_447])).
cnf(c_0_456,plain,(west(cell60,cell70)), inference(spm,[status(thm)],[c_0_380, c_0_448])).
cnf(c_0_457,plain,(epred18_2(cell80,cell70)), inference(spm,[status(thm)],[c_0_363, c_0_449])).
cnf(c_0_458,plain,(epred17_2(cell90,cell80)), inference(spm,[status(thm)],[c_0_344, c_0_450])).
cnf(c_0_459,plain,(epred18_2(X1,X2)|X1!=cell100|X2!=cell90), inference(split_conjunct,[status(thm)],[c_0_343])).
fof(c_0_460, plain, (![X45]:![X46]:![X47]:![X48]:((~duenorth(X46,X45)|epred42_2(X45,X46))&(~epred42_2(X47,X48)|duenorth(X48,X47)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_451])])])])).
cnf(c_0_461,plain,(epred42_2(X1,cell1)|~north(cell9,X1)), inference(spm,[status(thm)],[c_0_452, c_0_453])).
cnf(c_0_462,plain,(north(cell9,cell10)), inference(spm,[status(thm)],[c_0_330, c_0_454])).
cnf(c_0_463,plain,(epred39_2(X1,cell10)|~west(cell70,X2)|~west(X2,X3)|~west(X3,X1)), inference(spm,[status(thm)],[c_0_455, c_0_456])).
cnf(c_0_464,plain,(west(cell70,cell80)), inference(spm,[status(thm)],[c_0_380, c_0_457])).
cnf(c_0_465,plain,(epred18_2(cell90,cell80)), inference(spm,[status(thm)],[c_0_363, c_0_458])).
cnf(c_0_466,plain,(epred18_2(X1,cell90)|cell100!=X1), inference(er,[status(thm)],[c_0_459])).
fof(c_0_467, plain, (![X83]:![X84]:![X85]:![X86]:((~vert(X84,X83)|(duenorth(X84,X83)|duenorth(X83,X84)))&((~duenorth(X86,X85)|vert(X86,X85))&(~duenorth(X85,X86)|vert(X86,X85))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_468,plain,(duenorth(X1,X2)|~epred42_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_460])).
cnf(c_0_469,plain,(epred42_2(cell10,cell1)), inference(spm,[status(thm)],[c_0_461, c_0_462])).
fof(c_0_470, axiom, (![X5]:![X6]:(duewest(X6,X5)<=>epred39_2(X5,X6))), inference(apply_def,[status(thm)],[c_0_4, c_0_340])).
cnf(c_0_471,plain,(epred39_2(X1,cell10)|~west(cell80,X2)|~west(X2,X1)), inference(spm,[status(thm)],[c_0_463, c_0_464])).
cnf(c_0_472,plain,(west(cell80,cell90)), inference(spm,[status(thm)],[c_0_380, c_0_465])).
cnf(c_0_473,plain,(epred18_2(cell100,cell90)), inference(er,[status(thm)],[c_0_466])).
fof(c_0_474, plain, (![X84]:![X85]:![X87]:![X88]:![X89]:((((vert(X85,esk1_2(X84,X85))|duewest(X85,X84))|~westof(X85,X84))&((duewest(esk1_2(X84,X85),X84)|duewest(X85,X84))|~westof(X85,X84)))&((~duewest(X88,X87)|westof(X88,X87))&((~vert(X88,X89)|~duewest(X89,X87))|westof(X88,X87))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).
cnf(c_0_475,plain,(vert(X1,X2)|~duenorth(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_467])).
cnf(c_0_476,plain,(duenorth(cell1,cell10)), inference(spm,[status(thm)],[c_0_468, c_0_469])).
fof(c_0_477, plain, (![X7]:![X8]:![X9]:![X10]:((~duewest(X8,X7)|epred39_2(X7,X8))&(~epred39_2(X9,X10)|duewest(X10,X9)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_470])])])])).
cnf(c_0_478,plain,(epred39_2(X1,cell10)|~west(cell90,X1)), inference(spm,[status(thm)],[c_0_471, c_0_472])).
cnf(c_0_479,plain,(west(cell90,cell100)), inference(spm,[status(thm)],[c_0_380, c_0_473])).
cnf(c_0_480,plain,(westof(X1,X2)|~duewest(X3,X2)|~vert(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_474])).
cnf(c_0_481,plain,(vert(cell1,cell10)), inference(spm,[status(thm)],[c_0_475, c_0_476])).
cnf(c_0_482,plain,(duewest(X1,X2)|~epred39_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_477])).
cnf(c_0_483,plain,(epred39_2(cell100,cell10)), inference(spm,[status(thm)],[c_0_478, c_0_479])).
fof(c_0_484, negated_conjecture, (~westof(cell1,cell100)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_6])])).
cnf(c_0_485,plain,(westof(cell1,X1)|~duewest(cell10,X1)), inference(spm,[status(thm)],[c_0_480, c_0_481])).
cnf(c_0_486,plain,(duewest(cell10,cell100)), inference(spm,[status(thm)],[c_0_482, c_0_483])).
cnf(c_0_487,negated_conjecture,(~westof(cell1,cell100)), inference(split_conjunct,[status(thm)],[c_0_484])).
cnf(c_0_488,plain,($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_486]), c_0_487]), ['proof']).
# SZS output end CNFRefutation.
EOF
