% Problem    : SWV478+1 : TPTP v6.1.0. Released v4.0.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n126.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 04:32:06 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2c
# and selection function SelectCQIArEqLast.
#
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation.
fof(c_0_0, conjecture, (![X12]:![X13]:![X5]:![X6]:(((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~(leq(host(X7),host(X14))))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~(setIn(X7,alive))|~(setIn(X14,alive)))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X14]:((((setIn(X7,alive)&setIn(X14,alive))&setIn(host(X14),index(down,host(X7))))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X14]:(((setIn(X14,alive)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&setIn(host(X14),index(down,host(X7)))))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&elem(m_Down(X15),queue(host(X7))))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&host(X15)=host(X14))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&elem(m_Down(X15),queue(host(X7)))))))&![X7]:![X14]:((((((setIn(X7,alive)&setIn(X14,alive))&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>X14=X7))&![X7]:![X15]:![X14]:((((~(leq(host(X15),host(X14)))&setIn(X14,alive))&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&elem(m_Down(X16),queue(host(X15)))))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&setIn(host(X16),index(down,host(X15)))))))&![X7]:![X15]:![X14]:((((((host(X14)!=host(X15)&setIn(X14,alive))&elem(m_Down(X7),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X15))=>~(((setIn(X15,alive)&index(status,host(X15))=norm)&index(ldr,host(X15))=host(X15)))))&![X7]:![X15]:![X14]:((((((host(X14)!=host(X15)&setIn(X14,alive))&elem(m_Down(X7),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=wait)&host(index(elid,host(X14)))=host(X15))=>~(((setIn(X15,alive)&index(status,host(X15))=norm)&index(ldr,host(X15))=host(X15)))))&![X7]:![X16]:![X15]:![X14]:(((((((![X17]:((~(leq(host(X14),X17))&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&setIn(X7,alive))&leq(host(X7),host(X14)))&elem(m_Down(X15),queue(host(X14))))&host(X16)=host(X14))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_1)=>~(elem(m_Ack(X7,X16),queue(host(X7))))))&![X7]:![X16]:![X15]:![X14]:(((((((![X17]:((~(leq(host(X14),X17))&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&setIn(X14,alive))&leq(nbr_proc,s(host(X14))))&elem(m_Down(X15),queue(host(X14))))&elem(m_Down(X16),queue(host(X14))))&host(X16)=s(host(X14)))&index(status,host(X14))=elec_1)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12))=>(setIn(X5,alive)=>(leq(host(X5),host(X6))=>((index(status,host(X5))=elec_2&host(X6)=index(pendack,host(X5)))=>(leq(nbr_proc,index(pendack,host(X5)))=>![X7]:(setIn(host(X7),index(acks,host(X5)))=>![X17]:(host(X5)=host(X17)=>![X18]:![X19]:![X20]:(host(X7)=host(X20)=>(host(X5)!=host(X20)=>(((((((![X21]:((~(leq(host(X20),X21))&leq(s(zero),X21))=>(setIn(X21,index(down,host(X20)))|X21=host(X19)))&setIn(X20,alive))&leq(nbr_proc,s(host(X20))))&elem(m_Down(X18),snoc(queue(host(X20)),m_Ldr(X5))))&elem(m_Down(X19),snoc(queue(host(X20)),m_Ldr(X5))))&host(X18)=s(host(X20)))&index(status,host(X20))=elec_1)=>~((setIn(X17,alive)&host(X5)=host(X17)))))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', conj)).
fof(c_0_1, axiom, (![X5]:![X6]:((leq(X5,X6)&leq(X6,X5))<=>X5=X6)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_61)).
fof(c_0_2, axiom, (![X3]:leq(host(X3),nbr_proc)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_04)).
fof(c_0_3, axiom, (![X5]:![X6]:![X4]:(elem(X5,cons(X6,X4))<=>(X5=X6|elem(X5,X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_46)).
fof(c_0_4, axiom, (![X5]:~(leq(s(X5),X5))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_58)).
fof(c_0_5, axiom, (![X5]:![X6]:(leq(X5,s(X6))<=>(X5=s(X6)|leq(X5,X6)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_64)).
fof(c_0_6, axiom, (![X3]:leq(s(zero),host(X3))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_02)).
fof(c_0_7, axiom, (![X5]:![X6]:![X4]:(elem(X5,snoc(X4,X6))<=>(X5=X6|elem(X5,X4)))), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_47)).
fof(c_0_8, axiom, (![X5]:![X6]:m_Down(X5)!=m_Ldr(X6)), file('/export/starexec/sandbox/benchmark/Axioms/SWV011+0.ax', axiom_18)).
fof(c_0_9, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)<=>((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X14]:((((setIn(X7,alive)&setIn(X14,alive))&setIn(host(X14),index(down,host(X7))))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X14]:(((setIn(X14,alive)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&setIn(host(X14),index(down,host(X7)))))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&elem(m_Down(X15),queue(host(X7))))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&host(X15)=host(X14))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&elem(m_Down(X15),queue(host(X7)))))))&![X7]:![X14]:((((((setIn(X7,alive)&setIn(X14,alive))&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>X14=X7))&![X7]:![X15]:![X14]:((((~leq(host(X15),host(X14))&setIn(X14,alive))&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&elem(m_Down(X16),queue(host(X15)))))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&setIn(host(X16),index(down,host(X15)))))))&![X7]:![X15]:![X14]:((((((host(X14)!=host(X15)&setIn(X14,alive))&elem(m_Down(X7),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X15))=>~(((setIn(X15,alive)&index(status,host(X15))=norm)&index(ldr,host(X15))=host(X15)))))&![X7]:![X15]:![X14]:((((((host(X14)!=host(X15)&setIn(X14,alive))&elem(m_Down(X7),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=wait)&host(index(elid,host(X14)))=host(X15))=>~(((setIn(X15,alive)&index(status,host(X15))=norm)&index(ldr,host(X15))=host(X15)))))&![X7]:![X16]:![X15]:![X14]:(((((((![X17]:((~leq(host(X14),X17)&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&setIn(X7,alive))&leq(host(X7),host(X14)))&elem(m_Down(X15),queue(host(X14))))&host(X16)=host(X14))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_1)=>~elem(m_Ack(X7,X16),queue(host(X7)))))&![X7]:![X16]:![X15]:![X14]:(((((((![X17]:((~leq(host(X14),X17)&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&setIn(X14,alive))&leq(nbr_proc,s(host(X14))))&elem(m_Down(X15),queue(host(X14))))&elem(m_Down(X16),queue(host(X14))))&host(X16)=s(host(X14)))&index(status,host(X14))=elec_1)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12)))), introduced(definition)).
fof(c_0_10, negated_conjecture, (~(![X12]:![X5]:![X6]:(epred1_3(X5,X6,X12)=>(setIn(X5,alive)=>(leq(host(X5),host(X6))=>((index(status,host(X5))=elec_2&host(X6)=index(pendack,host(X5)))=>(leq(nbr_proc,index(pendack,host(X5)))=>![X7]:(setIn(host(X7),index(acks,host(X5)))=>![X17]:(host(X5)=host(X17)=>![X18]:![X19]:![X20]:(host(X7)=host(X20)=>(host(X5)!=host(X20)=>(((((((![X21]:((~leq(host(X20),X21)&leq(s(zero),X21))=>(setIn(X21,index(down,host(X20)))|X21=host(X19)))&setIn(X20,alive))&leq(nbr_proc,s(host(X20))))&elem(m_Down(X18),snoc(queue(host(X20)),m_Ldr(X5))))&elem(m_Down(X19),snoc(queue(host(X20)),m_Ldr(X5))))&host(X18)=s(host(X20)))&index(status,host(X20))=elec_1)=>~((setIn(X17,alive)&host(X5)=host(X17))))))))))))))), inference(apply_def,[status(thm)],[inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])]), c_0_9])).
fof(c_0_11, negated_conjecture, (![X30]:(epred1_3(esk6_0,esk7_0,esk5_0)&(setIn(esk6_0,alive)&(leq(host(esk6_0),host(esk7_0))&((index(status,host(esk6_0))=elec_2&host(esk7_0)=index(pendack,host(esk6_0)))&(leq(nbr_proc,index(pendack,host(esk6_0)))&(setIn(host(esk8_0),index(acks,host(esk6_0)))&(host(esk6_0)=host(esk9_0)&(host(esk8_0)=host(esk12_0)&(host(esk6_0)!=host(esk12_0)&(((((((((leq(host(esk12_0),X30)|~leq(s(zero),X30))|(setIn(X30,index(down,host(esk12_0)))|X30=host(esk11_0)))&setIn(esk12_0,alive))&leq(nbr_proc,s(host(esk12_0))))&elem(m_Down(esk10_0),snoc(queue(host(esk12_0)),m_Ldr(esk6_0))))&elem(m_Down(esk11_0),snoc(queue(host(esk12_0)),m_Ldr(esk6_0))))&host(esk10_0)=s(host(esk12_0)))&index(status,host(esk12_0))=elec_1)&(setIn(esk9_0,alive)&host(esk6_0)=host(esk9_0))))))))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
fof(c_0_12, plain, (![X12]:![X6]:![X5]:(epred1_3(X5,X6,X12)=>((((((((((((((((![X7]:![X14]:(elem(m_Ldr(X14),queue(host(X7)))=>~leq(host(X7),host(X14)))&![X7]:![X14]:(elem(m_Down(X14),queue(host(X7)))=>host(X14)!=host(X7)))&![X7]:![X14]:((X14!=X7&host(X14)=host(X7))=>(~setIn(X7,alive)|~setIn(X14,alive))))&![X7]:(((index(status,host(X7))=elec_1|index(status,host(X7))=elec_2)&setIn(X7,alive))=>index(elid,host(X7))=X7))&![X7]:![X14]:((((setIn(X7,alive)&setIn(X14,alive))&setIn(host(X14),index(down,host(X7))))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X14]:(((setIn(X14,alive)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&setIn(host(X14),index(down,host(X7)))))))&![X7]:![X15]:![X14]:(((((setIn(X7,alive)&setIn(X14,alive))&elem(m_Down(X15),queue(host(X7))))&host(X14)=host(X15))&index(status,host(X14))=elec_2)=>leq(index(pendack,host(X14)),host(X7))))&![X7]:![X15]:![X14]:((((setIn(X14,alive)&host(X15)=host(X14))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>~((setIn(X7,alive)&elem(m_Down(X15),queue(host(X7)))))))&![X7]:![X14]:((((((setIn(X7,alive)&setIn(X14,alive))&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm)&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X14))=>X14=X7))&![X7]:![X15]:![X14]:((((~leq(host(X15),host(X14))&setIn(X14,alive))&elem(m_Down(X15),queue(host(X14))))&host(X15)=host(X7))=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&elem(m_Down(X16),queue(host(X15)))))))&![X7]:![X16]:![X15]:![X14]:(((((host(X15)!=host(X7)&setIn(X7,alive))&setIn(X15,alive))&host(X16)=host(X7))&host(X14)=host(X15))=>~((elem(m_Down(X14),queue(host(X7)))&setIn(host(X16),index(down,host(X15)))))))&![X7]:![X15]:![X14]:((((((host(X14)!=host(X15)&setIn(X14,alive))&elem(m_Down(X7),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=norm)&index(ldr,host(X14))=host(X15))=>~(((setIn(X15,alive)&index(status,host(X15))=norm)&index(ldr,host(X15))=host(X15)))))&![X7]:![X15]:![X14]:((((((host(X14)!=host(X15)&setIn(X14,alive))&elem(m_Down(X7),queue(host(X14))))&host(X15)=host(X7))&index(status,host(X14))=wait)&host(index(elid,host(X14)))=host(X15))=>~(((setIn(X15,alive)&index(status,host(X15))=norm)&index(ldr,host(X15))=host(X15)))))&![X7]:![X16]:![X15]:![X14]:(((((((![X17]:((~leq(host(X14),X17)&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&setIn(X7,alive))&leq(host(X7),host(X14)))&elem(m_Down(X15),queue(host(X14))))&host(X16)=host(X14))&index(status,host(X7))=elec_2)&index(status,host(X14))=elec_1)=>~elem(m_Ack(X7,X16),queue(host(X7)))))&![X7]:![X16]:![X15]:![X14]:(((((((![X17]:((~leq(host(X14),X17)&leq(s(zero),X17))=>(setIn(X17,index(down,host(X14)))|X17=host(X15)))&setIn(X14,alive))&leq(nbr_proc,s(host(X14))))&elem(m_Down(X15),queue(host(X14))))&elem(m_Down(X16),queue(host(X14))))&host(X16)=s(host(X14)))&index(status,host(X14))=elec_1)=>~(((setIn(X7,alive)&index(ldr,host(X7))=host(X7))&index(status,host(X7))=norm))))&queue(host(X5))=cons(m_Down(X6),X12)))), inference(split_equiv,[status(thm)],[c_0_9])).
fof(c_0_13, plain, (![X7]:![X8]:![X9]:![X10]:(((~leq(X7,X8)|~leq(X8,X7))|X7=X8)&((leq(X9,X10)|X9!=X10)&(leq(X10,X9)|X9!=X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
cnf(c_0_14,negated_conjecture,(leq(nbr_proc,s(host(esk12_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15,negated_conjecture,(host(esk10_0)=s(host(esk12_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, (![X4]:leq(host(X4),nbr_proc)), inference(variable_rename,[status(thm)],[c_0_2])).
fof(c_0_17, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((~elem(X7,cons(X8,X9))|(X7=X8|elem(X7,X9)))&((X10!=X11|elem(X10,cons(X11,X12)))&(~elem(X10,X12)|elem(X10,cons(X11,X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
fof(c_0_18, plain, (![X18]:![X19]:![X20]:![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:![X33]:![X34]:![X35]:![X36]:![X37]:![X38]:![X39]:![X40]:![X41]:![X42]:![X43]:![X44]:![X45]:![X46]:![X47]:![X48]:![X49]:![X50]:![X51]:![X52]:![X53]:![X54]:![X55]:![X56]:![X57]:![X58]:![X59]:![X60]:![X62]:![X63]:![X64]:![X66]:((((((((((((((((((~elem(m_Ldr(X22),queue(host(X21)))|~leq(host(X21),host(X22)))|~epred1_3(X20,X19,X18))&((~elem(m_Down(X24),queue(host(X23)))|host(X24)!=host(X23))|~epred1_3(X20,X19,X18)))&(((X26=X25|host(X26)!=host(X25))|(~setIn(X25,alive)|~setIn(X26,alive)))|~epred1_3(X20,X19,X18)))&((((index(status,host(X27))!=elec_1|~setIn(X27,alive))|index(elid,host(X27))=X27)|~epred1_3(X20,X19,X18))&(((index(status,host(X27))!=elec_2|~setIn(X27,alive))|index(elid,host(X27))=X27)|~epred1_3(X20,X19,X18))))&(((((~setIn(X28,alive)|~setIn(X29,alive))|~setIn(host(X29),index(down,host(X28))))|index(status,host(X29))!=elec_2)|leq(index(pendack,host(X29)),host(X28)))|~epred1_3(X20,X19,X18)))&((((~setIn(X31,alive)|index(status,host(X31))!=norm)|index(ldr,host(X31))!=host(X31))|(~setIn(X30,alive)|~setIn(host(X31),index(down,host(X30)))))|~epred1_3(X20,X19,X18)))&((((((~setIn(X32,alive)|~setIn(X34,alive))|~elem(m_Down(X33),queue(host(X32))))|host(X34)!=host(X33))|index(status,host(X34))!=elec_2)|leq(index(pendack,host(X34)),host(X32)))|~epred1_3(X20,X19,X18)))&(((((~setIn(X37,alive)|host(X36)!=host(X37))|index(status,host(X37))!=norm)|index(ldr,host(X37))!=host(X37))|(~setIn(X35,alive)|~elem(m_Down(X36),queue(host(X35)))))|~epred1_3(X20,X19,X18)))&(((((((~setIn(X38,alive)|~setIn(X39,alive))|index(ldr,host(X38))!=host(X38))|index(status,host(X38))!=norm)|index(status,host(X39))!=norm)|index(ldr,host(X39))!=host(X39))|X39=X38)|~epred1_3(X20,X19,X18)))&(((((leq(host(X41),host(X42))|~setIn(X42,alive))|~elem(m_Down(X41),queue(host(X42))))|host(X41)!=host(X40))|((~setIn(X40,alive)|index(ldr,host(X40))!=host(X40))|index(status,host(X40))!=norm))|~epred1_3(X20,X19,X18)))&((((((host(X45)=host(X43)|~setIn(X43,alive))|~setIn(X45,alive))|host(X44)!=host(X43))|host(X46)!=host(X45))|(~elem(m_Down(X46),queue(host(X43)))|~elem(m_Down(X44),queue(host(X45)))))|~epred1_3(X20,X19,X18)))&((((((host(X49)=host(X47)|~setIn(X47,alive))|~setIn(X49,alive))|host(X48)!=host(X47))|host(X50)!=host(X49))|(~elem(m_Down(X50),queue(host(X47)))|~setIn(host(X48),index(down,host(X49)))))|~epred1_3(X20,X19,X18)))&(((((((host(X53)=host(X52)|~setIn(X53,alive))|~elem(m_Down(X51),queue(host(X53))))|host(X52)!=host(X51))|index(status,host(X53))!=norm)|index(ldr,host(X53))!=host(X52))|((~setIn(X52,alive)|index(status,host(X52))!=norm)|index(ldr,host(X52))!=host(X52)))|~epred1_3(X20,X19,X18)))&(((((((host(X56)=host(X55)|~setIn(X56,alive))|~elem(m_Down(X54),queue(host(X56))))|host(X55)!=host(X54))|index(status,host(X56))!=wait)|host(index(elid,host(X56)))!=host(X55))|((~setIn(X55,alive)|index(status,host(X55))!=norm)|index(ldr,host(X55))!=host(X55)))|~epred1_3(X20,X19,X18)))&((((((((((~leq(host(X60),esk13_7(X18,X19,X20,X57,X58,X59,X60))|~setIn(X57,alive))|~leq(host(X57),host(X60)))|~elem(m_Down(X59),queue(host(X60))))|host(X58)!=host(X60))|index(status,host(X57))!=elec_2)|index(status,host(X60))!=elec_1)|~elem(m_Ack(X57,X58),queue(host(X57))))|~epred1_3(X20,X19,X18))&((((((((leq(s(zero),esk13_7(X18,X19,X20,X57,X58,X59,X60))|~setIn(X57,alive))|~leq(host(X57),host(X60)))|~elem(m_Down(X59),queue(host(X60))))|host(X58)!=host(X60))|index(status,host(X57))!=elec_2)|index(status,host(X60))!=elec_1)|~elem(m_Ack(X57,X58),queue(host(X57))))|~epred1_3(X20,X19,X18)))&(((((((((~setIn(esk13_7(X18,X19,X20,X57,X58,X59,X60),index(down,host(X60)))|~setIn(X57,alive))|~leq(host(X57),host(X60)))|~elem(m_Down(X59),queue(host(X60))))|host(X58)!=host(X60))|index(status,host(X57))!=elec_2)|index(status,host(X60))!=elec_1)|~elem(m_Ack(X57,X58),queue(host(X57))))|~epred1_3(X20,X19,X18))&((((((((esk13_7(X18,X19,X20,X57,X58,X59,X60)!=host(X59)|~setIn(X57,alive))|~leq(host(X57),host(X60)))|~elem(m_Down(X59),queue(host(X60))))|host(X58)!=host(X60))|index(status,host(X57))!=elec_2)|index(status,host(X60))!=elec_1)|~elem(m_Ack(X57,X58),queue(host(X57))))|~epred1_3(X20,X19,X18)))))&((((((((((~leq(host(X64),esk14_6(X18,X19,X20,X62,X63,X64))|~setIn(X64,alive))|~leq(nbr_proc,s(host(X64))))|~elem(m_Down(X63),queue(host(X64))))|~elem(m_Down(X62),queue(host(X64))))|host(X62)!=s(host(X64)))|index(status,host(X64))!=elec_1)|((~setIn(X66,alive)|index(ldr,host(X66))!=host(X66))|index(status,host(X66))!=norm))|~epred1_3(X20,X19,X18))&((((((((leq(s(zero),esk14_6(X18,X19,X20,X62,X63,X64))|~setIn(X64,alive))|~leq(nbr_proc,s(host(X64))))|~elem(m_Down(X63),queue(host(X64))))|~elem(m_Down(X62),queue(host(X64))))|host(X62)!=s(host(X64)))|index(status,host(X64))!=elec_1)|((~setIn(X66,alive)|index(ldr,host(X66))!=host(X66))|index(status,host(X66))!=norm))|~epred1_3(X20,X19,X18)))&(((((((((~setIn(esk14_6(X18,X19,X20,X62,X63,X64),index(down,host(X64)))|~setIn(X64,alive))|~leq(nbr_proc,s(host(X64))))|~elem(m_Down(X63),queue(host(X64))))|~elem(m_Down(X62),queue(host(X64))))|host(X62)!=s(host(X64)))|index(status,host(X64))!=elec_1)|((~setIn(X66,alive)|index(ldr,host(X66))!=host(X66))|index(status,host(X66))!=norm))|~epred1_3(X20,X19,X18))&((((((((esk14_6(X18,X19,X20,X62,X63,X64)!=host(X63)|~setIn(X64,alive))|~leq(nbr_proc,s(host(X64))))|~elem(m_Down(X63),queue(host(X64))))|~elem(m_Down(X62),queue(host(X64))))|host(X62)!=s(host(X64)))|index(status,host(X64))!=elec_1)|((~setIn(X66,alive)|index(ldr,host(X66))!=host(X66))|index(status,host(X66))!=norm))|~epred1_3(X20,X19,X18)))))&(queue(host(X20))=cons(m_Down(X19),X18)|~epred1_3(X20,X19,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])])).
cnf(c_0_19,plain,(X1=X2|~leq(X2,X1)|~leq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20,negated_conjecture,(leq(nbr_proc,host(esk10_0))), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21,plain,(leq(host(X1),nbr_proc)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22,plain,(elem(X1,cons(X2,X3))|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23,plain,(queue(host(X1))=cons(m_Down(X2),X3)|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24,negated_conjecture,(epred1_3(esk6_0,esk7_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25,negated_conjecture,(leq(nbr_proc,index(pendack,host(esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_26,negated_conjecture,(host(esk7_0)=index(pendack,host(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_27, plain, (![X5]:~leq(s(X5),X5)), inference(fof_simplification,[status(thm)],[c_0_4])).
fof(c_0_28, plain, (![X7]:![X8]:![X9]:![X10]:((~leq(X7,s(X8))|(X7=s(X8)|leq(X7,X8)))&((X9!=s(X10)|leq(X9,s(X10)))&(~leq(X9,X10)|leq(X9,s(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
cnf(c_0_29,negated_conjecture,(host(esk10_0)=nbr_proc), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
cnf(c_0_30,plain,(elem(X1,cons(X1,X2))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_31,negated_conjecture,(cons(m_Down(esk7_0),esk5_0)=queue(host(esk6_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32,negated_conjecture,(leq(nbr_proc,host(esk7_0))), inference(rw,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_33, plain, (![X6]:~leq(s(X6),X6)), inference(variable_rename,[status(thm)],[c_0_27])).
cnf(c_0_34,plain,(leq(X1,X2)|X1=s(X2)|~leq(X1,s(X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35,negated_conjecture,(s(host(esk12_0))=nbr_proc), inference(rw,[status(thm)],[c_0_15, c_0_29])).
cnf(c_0_36,plain,(~epred1_3(X1,X2,X3)|host(X4)!=host(X5)|~elem(m_Down(X4),queue(host(X5)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_37,negated_conjecture,(elem(m_Down(esk7_0),queue(host(esk6_0)))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38,negated_conjecture,(host(esk7_0)=nbr_proc), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_32]), c_0_21])])).
cnf(c_0_39,plain,(~leq(s(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40,negated_conjecture,(nbr_proc=X1|leq(X1,host(esk12_0))|~leq(X1,nbr_proc)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41,plain,(host(esk6_0)!=nbr_proc|~epred1_3(X1,X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_42,plain,(leq(index(pendack,host(X4)),host(X5))|~epred1_3(X1,X2,X3)|index(status,host(X4))!=elec_2|~setIn(host(X4),index(down,host(X5)))|~setIn(X4,alive)|~setIn(X5,alive)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_43,negated_conjecture,(setIn(esk12_0,alive)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_44,negated_conjecture,(~leq(host(esk10_0),host(esk12_0))), inference(spm,[status(thm)],[c_0_39, c_0_15])).
fof(c_0_45, plain, (![X4]:leq(s(zero),host(X4))), inference(variable_rename,[status(thm)],[c_0_6])).
cnf(c_0_46,negated_conjecture,(host(X1)=nbr_proc|leq(host(X1),host(esk12_0))), inference(spm,[status(thm)],[c_0_40, c_0_21])).
cnf(c_0_47,negated_conjecture,(host(esk6_0)=host(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_48,negated_conjecture,(host(esk6_0)!=nbr_proc), inference(spm,[status(thm)],[c_0_41, c_0_24])).
cnf(c_0_49,negated_conjecture,(leq(index(pendack,host(X1)),host(esk12_0))|index(status,host(X1))!=elec_2|~epred1_3(X2,X3,X4)|~setIn(host(X1),index(down,host(esk12_0)))|~setIn(X1,alive)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50,negated_conjecture,(setIn(esk6_0,alive)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_51,negated_conjecture,(index(status,host(esk6_0))=elec_2), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_52,negated_conjecture,(~leq(nbr_proc,host(esk12_0))), inference(rw,[status(thm)],[c_0_44, c_0_29])).
cnf(c_0_53,negated_conjecture,(X1=host(esk11_0)|setIn(X1,index(down,host(esk12_0)))|leq(host(esk12_0),X1)|~leq(s(zero),X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_54,plain,(leq(s(zero),host(X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_55,negated_conjecture,(leq(host(esk6_0),host(esk12_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_56,negated_conjecture,(host(esk6_0)!=host(esk12_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_57, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((~elem(X7,snoc(X9,X8))|(X7=X8|elem(X7,X9)))&((X10!=X11|elem(X10,snoc(X12,X11)))&(~elem(X10,X12)|elem(X10,snoc(X12,X11)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
fof(c_0_58, plain, (![X7]:![X8]:m_Down(X7)!=m_Ldr(X8)), inference(variable_rename,[status(thm)],[c_0_8])).
cnf(c_0_59,negated_conjecture,(~epred1_3(X1,X2,X3)|~setIn(host(esk6_0),index(down,host(esk12_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_26]), c_0_38]), c_0_51])]), c_0_52])).
cnf(c_0_60,negated_conjecture,(host(esk11_0)=host(X1)|leq(host(esk12_0),host(X1))|setIn(host(X1),index(down,host(esk12_0)))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61,negated_conjecture,(~leq(host(esk12_0),host(esk6_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_55]), c_0_56])).
cnf(c_0_62,plain,(elem(X1,X2)|X1=X3|~elem(X1,snoc(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_63,negated_conjecture,(elem(m_Down(esk11_0),snoc(queue(host(esk12_0)),m_Ldr(esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_64,plain,(m_Down(X1)!=m_Ldr(X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65,negated_conjecture,(host(esk11_0)=host(esk6_0)|~epred1_3(X1,X2,X3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])).
cnf(c_0_66,plain,(leq(index(pendack,host(X4)),host(X5))|~epred1_3(X1,X2,X3)|index(status,host(X4))!=elec_2|host(X4)!=host(X6)|~elem(m_Down(X6),queue(host(X5)))|~setIn(X4,alive)|~setIn(X5,alive)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_67,negated_conjecture,(elem(m_Down(esk11_0),queue(host(esk12_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])).
cnf(c_0_68,negated_conjecture,(host(esk11_0)=host(esk6_0)), inference(spm,[status(thm)],[c_0_65, c_0_24])).
cnf(c_0_69,negated_conjecture,(leq(index(pendack,host(X1)),host(esk12_0))|index(status,host(X1))!=elec_2|host(X1)!=host(esk6_0)|~epred1_3(X2,X3,X4)|~setIn(X1,alive)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_43])]), c_0_68])).
cnf(c_0_70,negated_conjecture,(~epred1_3(X1,X2,X3)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_50]), c_0_26]), c_0_38]), c_0_51])]), c_0_52])).
cnf(c_0_71,negated_conjecture,($false), inference(sr,[status(thm)],[c_0_24, c_0_70]), ['proof']).
# SZS output end CNFRefutation.
EOF
