% Problem    : SYN545-1 : TPTP v6.1.0. Released v2.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n140.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:29:41 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic H_____011_C07_F1_PI_AE_SP_S0V
# and selection function PSelectComplexExceptRRHorn.
#
# No success with AutoSched0
# Trying AutoSched1 for 73 seconds
# AutoSched1-Mode selected heuristic G_E___207_C18_F1_AE_Q7_CS_SP_PI_CO_PS_S00
# and selection function NoSelection.
#
# Presaturation interreduction done
# No success with AutoSched1
# Trying AutoSched2 for 23 seconds
# AutoSched2-Mode selected heuristic H_____011_C18_F1_PI_SE_SP_S2S
# and selection function SelectNewComplexAHP.
#

# No proof found!
# SZS status Satisfiable
# SZS output start Saturation.
cnf(c_0_0,negated_conjecture,(ndr1_0|~ssSkC17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause2)).
cnf(c_0_1,negated_conjecture,(ndr1_0|ssSkC17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause13)).
cnf(c_0_2,negated_conjecture,(c4_2(X1,a293)|c4_1(X1)|~ssSkP10(X1)|~ndr1_0|~c2_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause343)).
cnf(c_0_3,negated_conjecture,(c2_0|c8_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause119)).
cnf(c_0_4,negated_conjecture,(c1_0|c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause4)).
cnf(c_0_5,negated_conjecture,(c8_1(X2)|c6_1(X2)|c2_1(X2)|ndr1_1(X1)|c5_1(X1)|~ndr1_0|~c3_1(X1)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause399)).
cnf(c_0_6,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c1_1(X3)|~c6_1(a160)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause455)).
cnf(c_0_7,negated_conjecture,(ndr1_1(X1)|c4_1(X1)|~ssSkP10(X1)|~ndr1_0|~c2_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause325)).
cnf(c_0_8,negated_conjecture,(~c6_2(X1,X2)|~c3_2(X1,X2)|~c1_2(X1,X2)|~ndr1_1(X1)|~c4_2(X1,X3)|~c1_2(X1,X3)|~c8_1(X1)|~ndr1_0|~c4_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause496)).
cnf(c_0_9,negated_conjecture,(c4_2(X1,a185)|c3_1(X1)|c1_0|~c5_1(X1)|~ndr1_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause355)).
cnf(c_0_10,negated_conjecture,(c5_2(X1,a278)|c1_1(X1)|c8_0|c7_0|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause326)).
cnf(c_0_11,negated_conjecture,(c7_2(X1,a292)|ssSkP10(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause120)).
cnf(c_0_12,negated_conjecture,(c1_1(X2)|c6_1(X2)|ndr1_1(X1)|c6_0|~ssSkP7(X1)|~ssSkP8(X1)|~c8_1(X2)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause381)).
cnf(c_0_13,negated_conjecture,(c7_2(a294,a295)|c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause85)).
cnf(c_0_14,negated_conjecture,(c3_0|~c6_1(a168)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause279)).
cnf(c_0_15,negated_conjecture,(c7_2(a204,a205)|c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause101)).
cnf(c_0_16,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|c8_2(a160,a161)|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c1_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause453)).
cnf(c_0_17,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|c2_2(a160,a161)|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c1_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause454)).
cnf(c_0_18,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|ndr1_1(a160)|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c1_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause446)).
cnf(c_0_19,negated_conjecture,(c3_0|c6_0|c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause45)).
cnf(c_0_20,negated_conjecture,(c1_0|c2_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause46)).
cnf(c_0_21,negated_conjecture,(~c1_1(a217)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause187)).
cnf(c_0_22,negated_conjecture,(ndr1_1(X1)|ssSkP7(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause51)).
cnf(c_0_23,negated_conjecture,(c7_2(X1,a248)|c2_1(X1)|ssSkC20|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause320)).
cnf(c_0_24,negated_conjecture,(ndr1_1(a294)|c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause30)).
cnf(c_0_25,negated_conjecture,(c4_1(X2)|c2_0|~c2_1(X1)|~c7_1(X1)|~ndr1_0|~c5_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause363)).
cnf(c_0_26,negated_conjecture,(c7_1(a217)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause98)).
cnf(c_0_27,negated_conjecture,(c8_1(a217)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause97)).
cnf(c_0_28,negated_conjecture,(c3_2(X1,a248)|c2_1(X1)|ssSkC20|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause319)).
cnf(c_0_29,negated_conjecture,(ndr1_1(X1)|c2_1(X1)|ssSkC20|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause304)).
cnf(c_0_30,negated_conjecture,(ndr1_1(a204)|c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause38)).
cnf(c_0_31,negated_conjecture,(ssSkC4|~c3_1(a174)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause168)).
cnf(c_0_32,negated_conjecture,(c1_1(X1)|c3_1(X1)|c1_1(a256)|c6_0|~c5_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause328)).
cnf(c_0_33,negated_conjecture,(ndr1_1(X1)|c3_1(X1)|c1_0|~c5_1(X1)|~ndr1_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause333)).
cnf(c_0_34,negated_conjecture,(c1_1(X1)|c3_1(X1)|ndr1_1(a256)|c6_0|~c5_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause327)).
cnf(c_0_35,negated_conjecture,(c3_2(X1,X2)|c2_2(X1,X2)|c1_1(X1)|c2_2(X3,X4)|c1_2(X3,X4)|c4_1(X3)|~c5_2(X1,X2)|~ndr1_1(X1)|~c7_1(X1)|~ndr1_0|~c7_2(X3,X4)|~ndr1_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause498)).
cnf(c_0_36,negated_conjecture,(ndr1_1(X1)|c1_1(X1)|c8_0|c7_0|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause307)).
cnf(c_0_37,negated_conjecture,(ndr1_1(X1)|ssSkP10(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause47)).
cnf(c_0_38,negated_conjecture,(ndr1_1(X1)|ssSkP8(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause50)).
cnf(c_0_39,negated_conjecture,(c6_2(a233,a234)|ssSkC17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause150)).
cnf(c_0_40,negated_conjecture,(c8_2(X1,a292)|ssSkP10(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause121)).
cnf(c_0_41,negated_conjecture,(c3_0|~c2_1(a168)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause278)).
cnf(c_0_42,negated_conjecture,(c3_1(X1)|c4_1(X1)|c5_1(X1)|ndr1_1(a296)|c6_0|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause306)).
cnf(c_0_43,negated_conjecture,(c1_2(X1,X2)|c2_1(X1)|c3_2(X3,X4)|c6_2(a231,a232)|~c7_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c3_1(X3)|~c5_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause512)).
cnf(c_0_44,negated_conjecture,(c5_1(a174)|ssSkC4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause78)).
cnf(c_0_45,negated_conjecture,(c7_2(X1,a185)|c3_1(X1)|c1_0|~c5_1(X1)|~ndr1_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause354)).
cnf(c_0_46,negated_conjecture,(ssSkC15|~c4_1(a223)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause155)).
cnf(c_0_47,negated_conjecture,(c4_2(X1,X2)|c3_2(X1,X3)|c4_2(X1,X3)|c5_2(X1,a147)|~c5_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause473)).
cnf(c_0_48,negated_conjecture,(c8_2(X1,X2)|c2_2(X1,X2)|c2_1(X1)|~c4_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c4_1(a159)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause468)).
cnf(c_0_49,negated_conjecture,(c4_2(X1,a258)|c1_1(X1)|ssSkC22|~c7_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause318)).
cnf(c_0_50,negated_conjecture,(c7_1(a202)|ssSkC11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause72)).
cnf(c_0_51,negated_conjecture,(ndr1_1(X1)|c3_1(X1)|c2_1(X1)|ssSkC7|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause297)).
cnf(c_0_52,negated_conjecture,(c3_1(X1)|c4_1(X1)|c5_1(X1)|c5_2(a296,a297)|c6_0|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause323)).
cnf(c_0_53,negated_conjecture,(c5_2(a275,X1)|ssSkC25|~c7_2(a275,X1)|~c2_2(a275,X1)|~ndr1_1(a275)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause400)).
cnf(c_0_54,negated_conjecture,(ndr1_1(a275)|ssSkC25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause54)).
cnf(c_0_55,negated_conjecture,(ndr1_1(X1)|c1_1(X1)|ssSkC22|~c7_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause303)).
cnf(c_0_56,negated_conjecture,(c4_2(a265,a266)|ssSkC23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause134)).
cnf(c_0_57,negated_conjecture,(c6_2(a223,X1)|ssSkC15|~c4_2(a223,X1)|~ndr1_1(a223)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause336)).
cnf(c_0_58,negated_conjecture,(c5_2(X1,a185)|c3_1(X1)|c1_0|~c5_1(X1)|~ndr1_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause356)).
cnf(c_0_59,negated_conjecture,(c8_2(X1,X2)|c2_2(X1,X2)|c2_1(X1)|c2_1(a159)|~c4_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause456)).
cnf(c_0_60,negated_conjecture,(c2_2(a275,a276)|ssSkC25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause128)).
cnf(c_0_61,negated_conjecture,(c8_2(X1,a207)|c3_0|c4_0|~ssSkP0(X1)|~c5_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause426)).
cnf(c_0_62,negated_conjecture,(c1_2(X1,X2)|c2_1(X1)|c3_2(X3,X4)|~c7_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c3_1(X3)|~c5_1(X3)|~c8_1(a231)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause514)).
cnf(c_0_63,negated_conjecture,(ndr1_1(a233)|ssSkC17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause66)).
cnf(c_0_64,negated_conjecture,(c7_2(a265,X1)|ssSkC23|~c4_2(a265,X1)|~ndr1_1(a265)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause334)).
cnf(c_0_65,negated_conjecture,(ndr1_1(a265)|ssSkC23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause57)).
cnf(c_0_66,negated_conjecture,(c7_2(a275,a276)|ssSkC25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause129)).
cnf(c_0_67,negated_conjecture,(c5_2(X1,X2)|c3_2(X1,X2)|c4_1(X1)|c3_1(X1)|c2_1(X3)|c3_1(X3)|~c8_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c3_2(X3,X4)|~c6_2(X3,X4)|~c8_2(X3,X4)|~ndr1_1(X3)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause510)).
cnf(c_0_68,negated_conjecture,(c4_2(X1,X2)|c5_2(X1,X2)|c7_2(X1,X3)|c6_2(X1,X3)|c1_0|c5_0|~c1_2(X1,X2)|~ndr1_1(X1)|~c8_2(X1,X3)|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause486)).
cnf(c_0_69,negated_conjecture,(c2_1(a181)|c5_0|~c5_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~c8_2(X1,X3)|~c1_2(X1,X3)|~c3_2(X1,X3)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause491)).
cnf(c_0_70,negated_conjecture,(ssSkC21|~c8_1(a253)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause139)).
cnf(c_0_71,negated_conjecture,(c1_0|c5_0|~c4_1(a158)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause207)).
cnf(c_0_72,negated_conjecture,(c4_2(X1,X2)|c7_2(X1,X3)|c4_2(X1,X3)|c7_1(X1)|ndr1_1(a208)|~c1_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,X3)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause488)).
cnf(c_0_73,negated_conjecture,(c8_0|c7_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause116)).
cnf(c_0_74,negated_conjecture,(c4_2(X1,X2)|c7_2(X1,X3)|c4_2(X1,X3)|c7_1(X1)|c8_1(a208)|~c1_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,X3)|~ndr1_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause489)).
cnf(c_0_75,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|c7_2(a193,a194)|c1_0|~c6_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause464)).
cnf(c_0_76,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|ndr1_1(a193)|c1_0|~c6_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause452)).
cnf(c_0_77,negated_conjecture,(c5_2(X1,X3)|c6_1(X1)|c4_0|c2_0|~c5_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,X3)|~c8_2(X1,X3)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause482)).
cnf(c_0_78,negated_conjecture,(c1_2(X1,a190)|c3_1(X1)|c2_1(X1)|ssSkC7|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause305)).
cnf(c_0_79,negated_conjecture,(c5_2(a241,a243)|ssSkC19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause145)).
cnf(c_0_80,negated_conjecture,(c3_2(X1,X2)|c8_2(X1,X2)|c2_1(X1)|c8_0|~c2_2(X1,X2)|~ndr1_1(X1)|~c5_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause429)).
cnf(c_0_81,negated_conjecture,(c1_2(a233,a234)|ssSkC17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause152)).
cnf(c_0_82,negated_conjecture,(c4_2(a197,a198)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause103)).
cnf(c_0_83,negated_conjecture,(c4_2(a241,a243)|ssSkC19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause144)).
cnf(c_0_84,negated_conjecture,(c5_0|~c8_1(a197)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause104)).
cnf(c_0_85,negated_conjecture,(c5_0|~c8_1(a162)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause112)).
cnf(c_0_86,negated_conjecture,(ssSkC27|~c8_1(a288)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause122)).
cnf(c_0_87,negated_conjecture,(c3_2(X1,a258)|c1_1(X1)|ssSkC22|~c7_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause317)).
cnf(c_0_88,negated_conjecture,(c1_2(X1,X2)|c2_1(X1)|c3_2(X3,X4)|c4_2(a231,a232)|~c7_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c3_1(X3)|~c5_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause511)).
cnf(c_0_89,negated_conjecture,(c1_2(X1,X2)|c2_1(X1)|c3_2(X3,X4)|c5_2(a231,a232)|~c7_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c3_1(X3)|~c5_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause513)).
cnf(c_0_90,negated_conjecture,(c1_2(X1,X2)|c2_1(X1)|c3_2(X3,X4)|ndr1_1(a231)|~c7_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c3_1(X3)|~c5_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause508)).
cnf(c_0_91,negated_conjecture,(c1_2(X1,X2)|c2_1(X1)|c3_2(X3,X4)|c4_1(a231)|~c7_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0|~c6_2(X3,X4)|~ndr1_1(X3)|~c3_1(X3)|~c5_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause509)).
cnf(c_0_92,negated_conjecture,(c3_1(a287)|c3_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause87)).
cnf(c_0_93,negated_conjecture,(c1_2(a271,a273)|ssSkC24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause131)).
cnf(c_0_94,negated_conjecture,(c1_2(a271,a272)|ssSkC24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause133)).
cnf(c_0_95,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|c1_0|~c6_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c1_1(a193)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause465)).
cnf(c_0_96,negated_conjecture,(c3_2(X1,X2)|c4_2(X1,X2)|c3_0|c8_0|~c5_2(X1,X2)|~ndr1_1(X1)|~c6_1(X1)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause447)).
cnf(c_0_97,negated_conjecture,(ssSkC10|~c1_1(a200)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause165)).
cnf(c_0_98,negated_conjecture,(c6_2(a268,a270)|c4_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause178)).
cnf(c_0_99,negated_conjecture,(c8_2(X1,X2)|c4_1(X1)|c4_0|~c7_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~c8_1(X1)|~ndr1_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause460)).
cnf(c_0_100,negated_conjecture,(c6_2(X1,a206)|ssSkP0(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause163)).
cnf(c_0_101,negated_conjecture,(c4_2(X1,a206)|ssSkP0(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause164)).
cnf(c_0_102,negated_conjecture,(c2_1(X1)|ndr1_1(X2)|c1_1(X2)|c8_0|~c7_1(X1)|~c5_1(X1)|~ndr1_0|~c7_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause411)).
cnf(c_0_103,negated_conjecture,(c8_1(X2)|c6_1(X2)|c2_1(X2)|c5_1(X1)|~ndr1_0|~c8_2(X1,a151)|~c3_1(X1)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause423)).
cnf(c_0_104,negated_conjecture,(c3_2(X1,a206)|ssSkP0(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause162)).
cnf(c_0_105,negated_conjecture,(c5_2(a268,a269)|c4_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause182)).
cnf(c_0_106,negated_conjecture,(c2_2(a271,a272)|ssSkC24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause132)).
cnf(c_0_107,negated_conjecture,(c4_1(X1)|~ssSkP10(X1)|~c3_2(X1,a293)|~ndr1_0|~c2_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause365)).
cnf(c_0_108,negated_conjecture,(c4_1(X1)|~ssSkP10(X1)|~c1_2(X1,a293)|~ndr1_0|~c2_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause366)).
cnf(c_0_109,negated_conjecture,(c1_1(a252)|~c8_1(X1)|~ndr1_0|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause298)).
cnf(c_0_110,negated_conjecture,(c3_2(a233,a234)|ssSkC17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause151)).
cnf(c_0_111,negated_conjecture,(c7_0|c8_0|~c2_1(a215)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause189)).
cnf(c_0_112,negated_conjecture,(c7_1(a215)|c7_0|c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause100)).
cnf(c_0_113,negated_conjecture,(c3_0|c5_0|~c5_1(a287)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause176)).
cnf(c_0_114,negated_conjecture,(c1_1(a158)|c1_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause113)).
cnf(c_0_115,negated_conjecture,(ssSkC25|~c1_1(a275)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause130)).
cnf(c_0_116,negated_conjecture,(ssSkC13|~c1_1(a218)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause160)).
cnf(c_0_117,negated_conjecture,(c1_0|~c1_1(a182)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause108)).
cnf(c_0_118,negated_conjecture,(ssSkC14|~c1_1(a221)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause156)).
cnf(c_0_119,negated_conjecture,(c4_2(X1,X2)|c7_2(X1,X3)|c4_2(X1,X3)|c7_1(X1)|~c1_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,X3)|~ndr1_0|~c3_2(a208,a209)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause501)).
cnf(c_0_120,negated_conjecture,(c4_2(X1,X2)|c7_2(X1,X3)|c4_2(X1,X3)|c7_1(X1)|~c1_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,X3)|~ndr1_0|~c2_2(a208,a209)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause500)).
cnf(c_0_121,negated_conjecture,(c8_2(X1,X2)|c2_2(X1,X2)|c2_1(X1)|c2_2(a159,X3)|~c4_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c4_2(a159,X3)|~c8_2(a159,X3)|~ndr1_1(a159)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause507)).
cnf(c_0_122,negated_conjecture,(ndr1_1(a241)|ssSkC19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause62)).
cnf(c_0_123,negated_conjecture,(c4_2(X1,X2)|c7_2(X1,X3)|c4_2(X1,X3)|c7_1(X1)|~c1_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,X3)|~ndr1_0|~c6_2(a208,a209)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause499)).
cnf(c_0_124,negated_conjecture,(ndr1_1(a197)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause40)).
cnf(c_0_125,negated_conjecture,(c4_2(a218,a219)|ssSkC13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause158)).
cnf(c_0_126,negated_conjecture,(c4_2(X1,X2)|c7_2(X1,X3)|c4_2(X1,X3)|c7_1(X1)|~c1_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,X3)|~ndr1_0|~c1_1(a208)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause495)).
cnf(c_0_127,negated_conjecture,(ndr1_1(a268)|c4_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause90)).
cnf(c_0_128,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,X2)|c1_2(X1,a261)|c6_0|c5_0|~ssSkP5(X1)|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause435)).
cnf(c_0_129,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,X2)|c2_2(X1,a261)|c6_0|c5_0|~ssSkP5(X1)|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause434)).
cnf(c_0_130,negated_conjecture,(c4_2(X1,X2)|c3_2(X1,X3)|c4_2(X1,X3)|~c5_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c7_2(X1,a147)|~ndr1_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause480)).
cnf(c_0_131,negated_conjecture,(c4_2(X1,X2)|c3_2(X1,X3)|c4_2(X1,X3)|~c5_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c3_2(X1,a147)|~ndr1_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause481)).
cnf(c_0_132,negated_conjecture,(c7_2(a223,X1)|c2_2(a223,X1)|ssSkC15|~ndr1_1(a223)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause313)).
cnf(c_0_133,negated_conjecture,(ssSkC14|~c2_1(a221)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause157)).
cnf(c_0_134,negated_conjecture,(ssSkC23|~c6_1(a265)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause136)).
cnf(c_0_135,negated_conjecture,(ssSkC3|~c6_1(a173)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause169)).
cnf(c_0_136,negated_conjecture,(c7_2(a148,a149)|c7_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause281)).
cnf(c_0_137,negated_conjecture,(c5_2(a252,X2)|~c8_1(X1)|~ndr1_0|~c2_2(a252,X2)|~c4_2(a252,X2)|~ndr1_1(a252)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause449)).
cnf(c_0_138,negated_conjecture,(c6_2(a174,X1)|ssSkC4|~c8_2(a174,X1)|~c5_2(a174,X1)|~ndr1_1(a174)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause403)).
cnf(c_0_139,negated_conjecture,(c8_2(a204,X1)|c2_2(a204,X1)|c1_0|~ndr1_1(a204)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause302)).
cnf(c_0_140,negated_conjecture,(c8_1(X2)|c6_1(X2)|c2_1(X2)|c5_1(X1)|~ndr1_0|~c1_2(X1,a151)|~c3_1(X1)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause424)).
cnf(c_0_141,negated_conjecture,(c8_1(X2)|c6_1(X2)|c2_1(X2)|c5_1(X1)|~ndr1_0|~c5_2(X1,a151)|~c3_1(X1)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause425)).
cnf(c_0_142,negated_conjecture,(c5_2(a148,X1)|c7_0|~c8_2(a148,X1)|~c4_2(a148,X1)|~ndr1_1(a148)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause421)).
cnf(c_0_143,negated_conjecture,(ndr1_1(a148)|c7_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause209)).
cnf(c_0_144,negated_conjecture,(c5_1(a271)|ssSkC24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause56)).
cnf(c_0_145,negated_conjecture,(c5_0|~c6_1(a197)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause105)).
cnf(c_0_146,negated_conjecture,(c7_2(X1,a264)|~ssSkP6(X1)|~c6_1(X1)|~ndr1_0|~c5_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause369)).
cnf(c_0_147,negated_conjecture,(c1_2(X1,a264)|~ssSkP6(X1)|~c6_1(X1)|~ndr1_0|~c5_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause368)).
cnf(c_0_148,negated_conjecture,(c5_2(X1,a264)|~ssSkP6(X1)|~c6_1(X1)|~ndr1_0|~c5_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause367)).
cnf(c_0_149,negated_conjecture,(c3_2(a155,X1)|ssSkC1|~c4_2(a155,X1)|~c2_2(a155,X1)|~ndr1_1(a155)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause404)).
cnf(c_0_150,negated_conjecture,(ndr1_1(a155)|ssSkC1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause81)).
cnf(c_0_151,negated_conjecture,(ndr1_1(X1)|~ssSkP6(X1)|~c6_1(X1)|~ndr1_0|~c5_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause345)).
cnf(c_0_152,negated_conjecture,(~c8_1(X1)|~ndr1_0|~c7_1(a252)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause308)).
cnf(c_0_153,negated_conjecture,(ndr1_1(a271)|ssSkC24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause55)).
cnf(c_0_154,negated_conjecture,(ndr1_1(a218)|ssSkC13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause68)).
cnf(c_0_155,negated_conjecture,(c6_1(a200)|ssSkC10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause73)).
cnf(c_0_156,negated_conjecture,(c8_2(a268,a270)|c4_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause179)).
cnf(c_0_157,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c1_1(X3)|~c5_2(a160,a161)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause467)).
cnf(c_0_158,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|c1_0|~c6_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c4_2(a193,a194)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause475)).
cnf(c_0_159,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|c1_0|~c6_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c6_2(a193,a194)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause476)).
cnf(c_0_160,negated_conjecture,(c5_2(X1,X2)|~ssSkC24|~c6_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0|~c1_2(a274,X3)|~c5_2(a274,X3)|~c6_2(a274,X3)|~ndr1_1(a274)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause497)).
cnf(c_0_161,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|~c7_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c1_1(X3)|~c6_2(a160,X4)|~ndr1_1(a160)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause485)).
cnf(c_0_162,negated_conjecture,(c5_2(X1,X2)|c5_2(a274,X3)|~ssSkC24|~c6_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0|~c7_2(a274,X3)|~c4_2(a274,X3)|~ndr1_1(a274)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause492)).
cnf(c_0_163,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X3)|c4_2(X1,X3)|c1_2(X1,X3)|~ssSkC6|~c3_2(X1,X2)|~c8_2(X1,X2)|~ndr1_1(X1)|~c7_1(X1)|~ndr1_0|~c6_1(a189)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause490)).
cnf(c_0_164,negated_conjecture,(c1_1(X1)|c3_1(X1)|c6_0|~c5_1(X1)|~ndr1_0|~c3_2(a256,X2)|~c5_2(a256,X2)|~c4_2(a256,X2)|~ndr1_1(a256)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause477)).
cnf(c_0_165,negated_conjecture,(c2_2(X1,X2)|c4_2(X1,X2)|~ssSkC11|~c5_2(X1,X2)|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0|~c1_2(a203,X3)|~c5_2(a203,X3)|~c8_2(a203,X3)|~ndr1_1(a203)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause502)).
cnf(c_0_166,negated_conjecture,(c3_2(X1,X2)|c6_2(X1,X3)|c7_2(X1,X3)|c8_1(X1)|~ssSkC7|~c2_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c8_2(a191,a192)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause484)).
cnf(c_0_167,negated_conjecture,(c3_1(X2)|c4_1(X2)|c5_1(X2)|c2_2(a296,X1)|c6_0|~ndr1_0|~c6_2(a296,X1)|~c7_2(a296,X1)|~ndr1_1(a296)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause461)).
cnf(c_0_168,negated_conjecture,(c8_2(X1,X2)|c1_2(X1,X2)|c6_0|c5_0|~ssSkP5(X1)|~c7_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,a261)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause444)).
cnf(c_0_169,negated_conjecture,(c3_2(X1,X2)|c6_2(X1,X3)|c7_2(X1,X3)|c8_1(X1)|c7_2(a191,a192)|~ssSkC7|~c2_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause478)).
cnf(c_0_170,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c1_2(X1,a220)|c4_0|~ssSkC13|~c8_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,X3)|~c7_2(X1,X3)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause487)).
cnf(c_0_171,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|c1_0|~c6_2(X1,X2)|~c7_2(X1,X2)|~ndr1_1(X1)|~c3_1(X1)|~ndr1_0|~c7_1(a193)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause466)).
cnf(c_0_172,negated_conjecture,(c3_2(X1,X2)|c6_2(X1,X3)|c7_2(X1,X3)|c8_1(X1)|~ssSkC7|~c2_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c5_1(a191)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause479)).
cnf(c_0_173,negated_conjecture,(c7_2(a155,a156)|ssSkC1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause172)).
cnf(c_0_174,negated_conjecture,(c3_0|c4_0|~ssSkP0(X1)|~c5_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c3_2(X1,a207)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause437)).
cnf(c_0_175,negated_conjecture,(c3_0|c4_0|~ssSkP0(X1)|~c5_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c5_2(X1,a207)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause438)).
cnf(c_0_176,negated_conjecture,(c3_2(X1,X2)|c6_2(X1,X3)|c7_2(X1,X3)|c8_1(X1)|ndr1_1(a191)|~ssSkC7|~c2_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause471)).
cnf(c_0_177,negated_conjecture,(c3_2(X1,X2)|c6_2(X1,X3)|c7_2(X1,X3)|c8_1(X1)|c1_1(a191)|~ssSkC7|~c2_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause472)).
cnf(c_0_178,negated_conjecture,(c2_2(a268,a269)|c4_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause181)).
cnf(c_0_179,negated_conjecture,(c7_2(X2,X3)|c5_2(X2,X3)|c6_1(X2)|c7_2(X1,X4)|c5_2(X1,X4)|c3_2(X1,X4)|c7_2(X1,a291)|~ssSkP9(X1)|~ssSkC27|~c3_2(X2,X3)|~ndr1_1(X2)|~c2_1(X2)|~ndr1_0|~ndr1_1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause503)).
cnf(c_0_180,negated_conjecture,(c2_1(X1)|c1_1(X2)|c8_0|~c7_1(X1)|~c5_1(X1)|~ndr1_0|~c1_2(X2,a247)|~c7_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause431)).
cnf(c_0_181,negated_conjecture,(c2_1(X1)|c1_1(X2)|c8_0|~c7_1(X1)|~c5_1(X1)|~ndr1_0|~c5_2(X2,a247)|~c7_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause430)).
cnf(c_0_182,negated_conjecture,(c7_2(X2,X3)|c5_2(X2,X3)|c6_1(X2)|c7_2(X1,X4)|c5_2(X1,X4)|c3_2(X1,X4)|c2_2(X1,a291)|~ssSkP9(X1)|~ssSkC27|~c3_2(X2,X3)|~ndr1_1(X2)|~c2_1(X2)|~ndr1_0|~ndr1_1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause504)).
cnf(c_0_183,negated_conjecture,(c1_1(X2)|c6_1(X2)|c6_0|~ssSkP7(X1)|~ssSkP8(X1)|~c8_1(X2)|~ndr1_0|~c7_2(X1,a286)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause409)).
cnf(c_0_184,negated_conjecture,(c1_1(X2)|c6_1(X2)|c6_0|~ssSkP7(X1)|~ssSkP8(X1)|~c8_1(X2)|~ndr1_0|~c1_2(X1,a286)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause410)).
cnf(c_0_185,negated_conjecture,(c2_1(a196)|~ssSkC8|~c1_2(X1,X2)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c4_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause474)).
cnf(c_0_186,negated_conjecture,(c1_0|~c7_2(a182,X1)|~c6_2(a182,X1)|~c1_2(a182,X1)|~ndr1_1(a182)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause407)).
cnf(c_0_187,negated_conjecture,(c7_2(X1,X2)|c7_1(X1)|c8_0|~c4_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~c8_2(X1,a230)|~ndr1_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause462)).
cnf(c_0_188,negated_conjecture,(c7_2(X1,X2)|c7_1(X1)|c8_0|~c4_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~c6_2(X1,a230)|~ndr1_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause463)).
cnf(c_0_189,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c6_0|~ssSkP2(X1)|~ssSkC18|~c5_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,a240)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause445)).
cnf(c_0_190,negated_conjecture,(c2_2(a265,a266)|ssSkC23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause135)).
cnf(c_0_191,negated_conjecture,(c2_2(a237,a238)|ssSkC18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause148)).
cnf(c_0_192,negated_conjecture,(c2_2(a155,a156)|ssSkC1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause171)).
cnf(c_0_193,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c4_0|~ssSkC13|~c8_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,X3)|~c7_2(X1,X3)|~c3_2(X1,a220)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause493)).
cnf(c_0_194,negated_conjecture,(c5_2(X1,X2)|c1_2(X1,X2)|c2_1(X1)|c1_0|~c3_2(X1,X2)|~ndr1_1(X1)|~c2_2(X1,a262)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause450)).
cnf(c_0_195,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c4_0|~ssSkC13|~c8_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_2(X1,X3)|~c7_2(X1,X3)|~c5_2(X1,a220)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause494)).
cnf(c_0_196,negated_conjecture,(c7_2(X1,X2)|c7_2(X1,a230)|c7_1(X1)|c8_0|~c4_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause451)).
cnf(c_0_197,negated_conjecture,(c7_2(X1,X2)|~c3_2(X1,X2)|~c5_2(X1,X2)|~ndr1_1(X1)|~c3_2(X1,a152)|~c5_1(X1)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause470)).
cnf(c_0_198,negated_conjecture,(c3_2(a229,X1)|c1_2(a229,X1)|~c4_2(a229,X1)|~ndr1_1(a229)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause395)).
cnf(c_0_199,negated_conjecture,(c2_2(a221,X1)|c4_2(a221,X1)|ssSkC14|~c7_2(a221,X1)|~ndr1_1(a221)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause385)).
cnf(c_0_200,negated_conjecture,(c7_2(X2,X3)|c5_2(X2,X3)|c6_1(X2)|c7_2(X1,X4)|c5_2(X1,X4)|c3_2(X1,X4)|~ssSkP9(X1)|~ssSkC27|~c3_2(X2,X3)|~ndr1_1(X2)|~c2_1(X2)|~ndr1_0|~ndr1_1(X1)|~c4_2(X1,a291)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause506)).
cnf(c_0_201,negated_conjecture,(c7_2(X1,X2)|~c3_2(X1,X2)|~c5_2(X1,X2)|~ndr1_1(X1)|~c4_2(X1,a152)|~c5_1(X1)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause469)).
cnf(c_0_202,negated_conjecture,(c4_2(X1,X2)|c3_2(X1,X2)|c6_2(X1,X2)|c3_1(X1)|~ssSkC3|~ssSkC4|~ndr1_1(X1)|~c2_2(X1,a175)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause439)).
cnf(c_0_203,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c7_2(X1,a240)|c6_0|~ssSkP2(X1)|~ssSkC18|~c5_2(X1,X2)|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause436)).
cnf(c_0_204,negated_conjecture,(c5_2(X1,X2)|c1_2(X1,X2)|c4_2(X1,a262)|c2_1(X1)|c1_0|~c3_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause442)).
cnf(c_0_205,negated_conjecture,(c5_2(X1,X2)|c1_2(X1,X2)|c1_2(X1,a262)|c2_1(X1)|c1_0|~c3_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause443)).
cnf(c_0_206,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|c5_0|~c8_2(X1,X2)|~c5_2(X1,X2)|~ndr1_1(X1)|~c7_1(X1)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause459)).
cnf(c_0_207,negated_conjecture,(c7_2(X1,X2)|c5_1(X1)|c4_0|~c5_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~c7_1(X1)|~ndr1_0|~c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause458)).
cnf(c_0_208,negated_conjecture,(c7_2(X1,X2)|c7_2(X1,a152)|~c3_2(X1,X2)|~c5_2(X1,X2)|~ndr1_1(X1)|~c5_1(X1)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause457)).
cnf(c_0_209,negated_conjecture,(c4_2(X1,X2)|c3_2(X1,X2)|c6_2(X1,X2)|c3_2(X1,a175)|c3_1(X1)|~ssSkC3|~ssSkC4|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause428)).
cnf(c_0_210,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X4)|c4_2(X1,X4)|c1_2(X1,X4)|c7_2(a189,X3)|~ssSkC6|~c3_2(X1,X2)|~c8_2(X1,X2)|~ndr1_1(X1)|~c7_1(X1)|~ndr1_0|~c2_2(a189,X3)|~c8_2(a189,X3)|~ndr1_1(a189)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause515)).
cnf(c_0_211,negated_conjecture,(c4_2(X1,X2)|c3_2(X1,X2)|c6_2(X1,X2)|c5_2(X1,a175)|c3_1(X1)|~ssSkC3|~ssSkC4|~ndr1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause427)).
cnf(c_0_212,negated_conjecture,(c4_2(X1,X2)|c6_2(X1,X2)|c8_1(X1)|c4_0|~c1_2(X1,X2)|~ndr1_1(X1)|~c6_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause448)).
cnf(c_0_213,negated_conjecture,(c7_2(X1,X2)|c5_2(X1,X2)|c8_1(X1)|c1_1(X1)|c6_0|~c1_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause440)).
cnf(c_0_214,negated_conjecture,(c1_2(a182,X1)|c1_0|~c2_2(a182,X1)|~c4_2(a182,X1)|~ndr1_1(a182)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause396)).
cnf(c_0_215,negated_conjecture,(c2_2(a229,X1)|c1_2(a229,X1)|c8_2(a229,X1)|~ndr1_1(a229)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause375)).
cnf(c_0_216,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X4)|c4_2(X1,X4)|c1_2(X1,X4)|c6_2(a189,X3)|~ssSkC6|~c3_2(X1,X2)|~c8_2(X1,X2)|~ndr1_1(X1)|~c7_1(X1)|~ndr1_0|~c8_2(a189,X3)|~c1_2(a189,X3)|~ndr1_1(a189)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause516)).
cnf(c_0_217,negated_conjecture,(c2_2(X1,X2)|c4_2(X1,X2)|c8_1(a203)|~ssSkC11|~c5_2(X1,X2)|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause433)).
cnf(c_0_218,negated_conjecture,(c2_2(X1,X2)|c4_2(X1,X2)|c4_1(a203)|~ssSkC11|~c5_2(X1,X2)|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause432)).
cnf(c_0_219,negated_conjecture,(c3_1(X1)|c1_0|~c3_2(X1,X2)|~c2_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause441)).
cnf(c_0_220,negated_conjecture,(c2_2(X1,X2)|c7_2(X1,X2)|c7_2(X1,a201)|~ssSkC9|~ssSkC10|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause417)).
cnf(c_0_221,negated_conjecture,(c2_2(X1,X2)|c7_2(X1,X2)|c8_2(X1,a201)|~ssSkC9|~ssSkC10|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause416)).
cnf(c_0_222,negated_conjecture,(c2_2(X1,X2)|c7_2(X1,X2)|c1_2(X1,a201)|~ssSkC9|~ssSkC10|~ndr1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause415)).
cnf(c_0_223,negated_conjecture,(c4_2(a181,X4)|c5_0|~c5_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~c8_2(X1,X3)|~c1_2(X1,X3)|~c3_2(X1,X3)|~c3_1(X1)|~ndr1_0|~c2_2(a181,X4)|~c3_2(a181,X4)|~ndr1_1(a181)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause517)).
cnf(c_0_224,negated_conjecture,(c5_0|~c5_2(X1,X2)|~c6_2(X1,X2)|~ndr1_1(X1)|~c8_2(X1,X3)|~c1_2(X1,X3)|~c3_2(X1,X3)|~c3_1(X1)|~ndr1_0|~c2_2(a181,X4)|~c6_2(a181,X4)|~c4_2(a181,X4)|~ndr1_1(a181)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause518)).
cnf(c_0_225,negated_conjecture,(c4_2(a158,X1)|c8_2(a158,X1)|c1_0|c5_0|~c1_2(a158,X1)|~ndr1_1(a158)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause398)).
cnf(c_0_226,negated_conjecture,(c1_1(X1)|c3_1(X1)|c6_0|~c5_1(X1)|~ndr1_0|~c7_2(a256,a257)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause372)).
cnf(c_0_227,negated_conjecture,(c1_1(X1)|c3_1(X1)|c6_0|~c5_1(X1)|~ndr1_0|~c1_2(a256,a257)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause371)).
cnf(c_0_228,negated_conjecture,(c2_2(a168,X1)|c3_0|~c5_2(a168,X1)|~c8_2(a168,X1)|~ndr1_1(a168)|~c4_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause419)).
cnf(c_0_229,negated_conjecture,(c5_2(X1,X2)|~ssSkC24|~c6_2(X1,X2)|~ndr1_1(X1)|~c1_1(X1)|~ndr1_0|~c7_1(a274)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause422)).
cnf(c_0_230,negated_conjecture,(c1_1(X2)|c5_1(X1)|~ssSkP4(X1)|~ssSkC21|~c3_1(X2)|~ndr1_0|~c3_2(X1,a255)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause391)).
cnf(c_0_231,negated_conjecture,(c1_1(X2)|c5_1(X1)|~ssSkP4(X1)|~ssSkC21|~c3_1(X2)|~ndr1_0|~c8_2(X1,a255)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause392)).
cnf(c_0_232,negated_conjecture,(c3_1(X1)|c4_1(X1)|c5_1(X1)|c6_0|~ndr1_0|~c7_2(a296,a297)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause341)).
cnf(c_0_233,negated_conjecture,(c3_1(X1)|c4_1(X1)|c5_1(X1)|c6_0|~ndr1_0|~c2_2(a296,a297)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause342)).
cnf(c_0_234,negated_conjecture,(c1_1(X1)|c8_0|c7_0|~c3_2(X1,a278)|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause344)).
cnf(c_0_235,negated_conjecture,(ssSkC12|~c3_1(a211)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause161)).
cnf(c_0_236,negated_conjecture,(c5_2(a186,X1)|c2_0|~c8_2(a186,X1)|~c1_2(a186,X1)|~ndr1_1(a186)|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause418)).
cnf(c_0_237,negated_conjecture,(ssSkC6|~c3_1(a187)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause166)).
cnf(c_0_238,negated_conjecture,(c5_2(a222,X1)|c2_2(a222,X1)|~ssSkC14|~ndr1_1(a222)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause337)).
cnf(c_0_239,negated_conjecture,(c1_1(X2)|c7_2(X1,a255)|c5_1(X1)|~ssSkP4(X1)|~ssSkC21|~c3_1(X2)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause373)).
cnf(c_0_240,negated_conjecture,(c3_1(X1)|c5_1(X1)|~c4_1(X1)|~ndr1_0|~c1_2(a171,a172)|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause397)).
cnf(c_0_241,negated_conjecture,(c1_0|~c3_1(a294)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause86)).
cnf(c_0_242,negated_conjecture,(ssSkC27|~c8_2(a288,X1)|~c5_2(a288,X1)|~c4_2(a288,X1)|~ndr1_1(a288)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause412)).
cnf(c_0_243,negated_conjecture,(ndr1_1(a288)|ssSkC27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause49)).
cnf(c_0_244,negated_conjecture,(c1_1(X2)|ndr1_1(X1)|c5_1(X1)|~ssSkP4(X1)|~ssSkC21|~c3_1(X2)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause346)).
cnf(c_0_245,negated_conjecture,(c4_1(a211)|ssSkC12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause70)).
cnf(c_0_246,negated_conjecture,(c7_2(a176,X1)|ssSkC5|~c2_2(a176,X1)|~c3_2(a176,X1)|~ndr1_1(a176)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause402)).
cnf(c_0_247,negated_conjecture,(ndr1_1(a176)|ssSkC5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause77)).
cnf(c_0_248,negated_conjecture,(c4_1(a187)|ssSkC6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause76)).
cnf(c_0_249,negated_conjecture,(c3_1(X1)|c5_1(X1)|c5_2(a171,a172)|~c4_1(X1)|~ndr1_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause377)).
cnf(c_0_250,negated_conjecture,(c4_1(a154)|ssSkC0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause82)).
cnf(c_0_251,negated_conjecture,(c1_1(X1)|ssSkC22|~c7_2(X1,a258)|~c7_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause339)).
cnf(c_0_252,negated_conjecture,(c2_1(X1)|ssSkC20|~c6_2(X1,a248)|~c8_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause340)).
cnf(c_0_253,negated_conjecture,(c3_1(X1)|c4_1(X1)|c5_1(X1)|c6_0|~ndr1_0|~c7_1(a296)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause324)).
cnf(c_0_254,negated_conjecture,(c3_1(X1)|c5_1(X1)|c2_2(a171,a172)|~c4_1(X1)|~ndr1_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause378)).
cnf(c_0_255,negated_conjecture,(c1_0|~c4_2(X1,a236)|~c6_1(X1)|~c8_1(X1)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause394)).
cnf(c_0_256,negated_conjecture,(c1_0|~c6_2(X1,a236)|~c6_1(X1)|~c8_1(X1)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause393)).
cnf(c_0_257,negated_conjecture,(c1_1(X1)|c1_0|~ssSkP1(X1)|~ssSkC15|~c4_2(X1,a225)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause329)).
cnf(c_0_258,negated_conjecture,(c2_1(X1)|c4_0|~c2_2(X1,a214)|~c1_1(X1)|~ndr1_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause376)).
cnf(c_0_259,negated_conjecture,(c5_1(X1)|c7_0|~ssSkC12|~c3_2(X1,a213)|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause353)).
cnf(c_0_260,negated_conjecture,(c3_1(X1)|c8_0|~ssSkP3(X1)|~c3_2(X1,a251)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause349)).
cnf(c_0_261,negated_conjecture,(c3_1(X1)|c8_0|~ssSkP3(X1)|~c8_2(X1,a251)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause347)).
cnf(c_0_262,negated_conjecture,(c3_1(X1)|c8_0|~ssSkP3(X1)|~c2_2(X1,a251)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause348)).
cnf(c_0_263,negated_conjecture,(c3_1(X1)|c5_1(X1)|~c4_1(X1)|~ndr1_0|~c2_1(a171)|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause379)).
cnf(c_0_264,negated_conjecture,(c6_2(a241,X1)|c5_2(a241,X1)|ssSkC19|~c4_2(a241,X1)|~ndr1_1(a241)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause383)).
cnf(c_0_265,negated_conjecture,(c5_2(X1,a236)|c1_0|~c6_1(X1)|~c8_1(X1)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause374)).
cnf(c_0_266,negated_conjecture,(c3_1(X1)|c2_1(X1)|ssSkC7|~c3_2(X1,a190)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause321)).
cnf(c_0_267,negated_conjecture,(c3_1(X1)|c2_1(X1)|ssSkC7|~c6_2(X1,a190)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause322)).
cnf(c_0_268,negated_conjecture,(ssSkP5(X1)|~c2_2(X1,a260)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause232)).
cnf(c_0_269,negated_conjecture,(c6_2(a279,X1)|c3_2(a279,X1)|ssSkC26|~c5_2(a279,X1)|~ndr1_1(a279)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause382)).
cnf(c_0_270,negated_conjecture,(ndr1_1(a279)|ssSkC26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause52)).
cnf(c_0_271,negated_conjecture,(c8_1(a287)|c3_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause88)).
cnf(c_0_272,negated_conjecture,(c8_2(a162,X1)|c3_2(a162,X1)|c5_0|~c7_2(a162,X1)|~ndr1_1(a162)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause380)).
cnf(c_0_273,negated_conjecture,(ndr1_1(a162)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause43)).
cnf(c_0_274,negated_conjecture,(c8_1(a204)|c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause39)).
cnf(c_0_275,negated_conjecture,(c3_1(X1)|c5_1(X1)|ndr1_1(a171)|~c4_1(X1)|~ndr1_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause357)).
cnf(c_0_276,negated_conjecture,(c3_1(X1)|c5_1(X1)|c4_1(a171)|~c4_1(X1)|~ndr1_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause358)).
cnf(c_0_277,negated_conjecture,(c8_2(X1,a225)|c1_1(X1)|c1_0|~ssSkP1(X1)|~ssSkC15|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause310)).
cnf(c_0_278,negated_conjecture,(c2_2(X1,a225)|c1_1(X1)|c1_0|~ssSkP1(X1)|~ssSkC15|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause311)).
cnf(c_0_279,negated_conjecture,(ssSkC2|~c8_2(a164,X1)|~c2_2(a164,X1)|~ndr1_1(a164)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause361)).
cnf(c_0_280,negated_conjecture,(ndr1_1(a164)|ssSkC2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause79)).
cnf(c_0_281,negated_conjecture,(c7_2(X1,a214)|c2_1(X1)|c4_0|~c1_1(X1)|~ndr1_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause351)).
cnf(c_0_282,negated_conjecture,(c4_2(X1,a214)|c2_1(X1)|c4_0|~c1_1(X1)|~ndr1_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause352)).
cnf(c_0_283,negated_conjecture,(c7_2(X1,a213)|c5_1(X1)|c7_0|~ssSkC12|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause332)).
cnf(c_0_284,negated_conjecture,(c2_2(X1,a213)|c5_1(X1)|c7_0|~ssSkC12|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause331)).
cnf(c_0_285,negated_conjecture,(ssSkC17|~c4_2(a233,X1)|~c3_2(a233,X1)|~ndr1_1(a233)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause359)).
cnf(c_0_286,negated_conjecture,(c6_2(a233,X1)|ssSkC17|~c8_2(a233,X1)|~ndr1_1(a233)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause335)).
cnf(c_0_287,negated_conjecture,(ndr1_1(X1)|c1_0|~c6_1(X1)|~c8_1(X1)|~ndr1_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause350)).
cnf(c_0_288,negated_conjecture,(ndr1_1(X1)|c1_1(X1)|c1_0|~ssSkP1(X1)|~ssSkC15|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause299)).
cnf(c_0_289,negated_conjecture,(ssSkP8(X1)|~c8_2(X1,a285)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause220)).
cnf(c_0_290,negated_conjecture,(ssSkP9(X1)|~c8_2(X1,a290)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause214)).
cnf(c_0_291,negated_conjecture,(c7_0|~c1_2(a148,a150)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause293)).
cnf(c_0_292,negated_conjecture,(c7_0|~c1_2(a148,a149)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause295)).
cnf(c_0_293,negated_conjecture,(ndr1_1(X1)|c2_1(X1)|c4_0|~c1_1(X1)|~ndr1_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause330)).
cnf(c_0_294,negated_conjecture,(ndr1_1(X1)|c5_1(X1)|c7_0|~ssSkC12|~c4_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause312)).
cnf(c_0_295,negated_conjecture,(ndr1_1(X1)|c3_1(X1)|c8_0|~ssSkP3(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause309)).
cnf(c_0_296,negated_conjecture,(c7_0|~c8_2(a148,a150)|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause294)).
cnf(c_0_297,negated_conjecture,(~ssSkC14|~c6_1(a222)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause269)).
cnf(c_0_298,negated_conjecture,(~ssSkC14|~c1_1(a222)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause268)).
cnf(c_0_299,negated_conjecture,(c2_0|~c7_1(a186)|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause272)).
cnf(c_0_300,negated_conjecture,(c2_2(a148,a149)|c7_0|~c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause282)).
cnf(c_0_301,negated_conjecture,(c1_1(a186)|c2_0|~c3_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause192)).
cnf(c_0_302,negated_conjecture,(c4_2(X1,a210)|~c4_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause300)).
cnf(c_0_303,negated_conjecture,(c8_2(X1,a210)|~c4_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause301)).
cnf(c_0_304,negated_conjecture,(c1_1(a229)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause95)).
cnf(c_0_305,negated_conjecture,(ndr1_1(X1)|~c4_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause296)).
cnf(c_0_306,negated_conjecture,(~c8_0|~c1_0|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause283)).
cnf(c_0_307,negated_conjecture,(~ssSkC8|~c1_2(X1,X2)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c4_1(X1)|~c2_1(X1)|~ndr1_0|~c1_2(a196,X3)|~c4_2(a196,X3)|~ndr1_1(a196)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause505)).
cnf(c_0_308,negated_conjecture,(~ssSkC8|~c1_2(X1,X2)|~c6_2(X1,X2)|~c3_2(X1,X2)|~ndr1_1(X1)|~c4_1(X1)|~c2_1(X1)|~ndr1_0|~c4_1(a196)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause483)).
cnf(c_0_309,negated_conjecture,(~c5_2(X1,a210)|~c4_1(X1)|~ndr1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause316)).
cnf(c_0_310,negated_conjecture,(c5_2(a249,X1)|~ssSkC20|~c4_2(a249,X1)|~c2_2(a249,X1)|~ndr1_1(a249)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause413)).
cnf(c_0_311,negated_conjecture,(c7_2(a277,X1)|c4_0|~ssSkC25|~c1_2(a277,X1)|~c2_2(a277,X1)|~ndr1_1(a277)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause405)).
cnf(c_0_312,negated_conjecture,(c3_2(a170,X1)|c7_0|c6_0|~c2_2(a170,X1)|~c4_2(a170,X1)|~ndr1_1(a170)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause408)).
cnf(c_0_313,negated_conjecture,(c2_2(a227,X1)|~ssSkC16|~c7_2(a227,X1)|~c4_2(a227,X1)|~ndr1_1(a227)|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause414)).
cnf(c_0_314,negated_conjecture,(c5_2(a227,X1)|c2_2(a227,X1)|~ssSkC16|~c4_2(a227,X1)|~ndr1_1(a227)|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause406)).
cnf(c_0_315,negated_conjecture,(c8_2(a153,X1)|c5_2(a153,X1)|~c3_2(a153,X1)|~ndr1_1(a153)|~c4_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause420)).
cnf(c_0_316,negated_conjecture,(c6_2(a277,X1)|c4_2(a277,X1)|c4_0|~ssSkC25|~c8_2(a277,X1)|~ndr1_1(a277)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause390)).
cnf(c_0_317,negated_conjecture,(c7_2(a226,X1)|ssSkC16|~c6_2(a226,X1)|~c5_2(a226,X1)|~ndr1_1(a226)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause401)).
cnf(c_0_318,negated_conjecture,(c6_2(a202,X1)|c7_2(a202,X1)|ssSkC11|~c3_2(a202,X1)|~ndr1_1(a202)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause387)).
cnf(c_0_319,negated_conjecture,(c5_2(a226,X1)|c6_2(a226,X1)|ssSkC16|~c8_2(a226,X1)|~ndr1_1(a226)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause384)).
cnf(c_0_320,negated_conjecture,(c1_2(a259,X1)|c8_2(a259,X1)|c3_2(a259,X1)|c7_0|~ssSkC22|~ndr1_1(a259)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause370)).
cnf(c_0_321,negated_conjecture,(c8_2(a173,X1)|c5_2(a173,X1)|ssSkC3|~c7_2(a173,X1)|~ndr1_1(a173)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause389)).
cnf(c_0_322,negated_conjecture,(c4_2(a195,X1)|c5_2(a195,X1)|ssSkC8|~c6_2(a195,X1)|~ndr1_1(a195)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause388)).
cnf(c_0_323,negated_conjecture,(c3_2(a202,X1)|c8_2(a202,X1)|ssSkC11|~c6_2(a202,X1)|~ndr1_1(a202)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause386)).
cnf(c_0_324,negated_conjecture,(c2_2(a253,X1)|c3_2(a253,X1)|c5_2(a253,X1)|ssSkC21|~ndr1_1(a253)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause364)).
cnf(c_0_325,negated_conjecture,(c4_0|~ssSkC25|~c6_2(a277,X1)|~c7_2(a277,X1)|~ndr1_1(a277)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause362)).
cnf(c_0_326,negated_conjecture,(c2_2(a179,X1)|c7_0|~ssSkC5|~c3_2(a179,X1)|~ndr1_1(a179)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause338)).
cnf(c_0_327,negated_conjecture,(ssSkC3|~c8_2(a173,X1)|~c1_2(a173,X1)|~ndr1_1(a173)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause360)).
cnf(c_0_328,negated_conjecture,(c8_2(a235,X1)|~ssSkC17|~c4_2(a235,X1)|~ndr1_1(a235)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause315)).
cnf(c_0_329,negated_conjecture,(c6_2(a267,X1)|c3_2(a267,X1)|c7_0|~ssSkC23|~ndr1_1(a267)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause314)).
cnf(c_0_330,negated_conjecture,(c7_0|c8_0|~c1_2(a215,a216)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause270)).
cnf(c_0_331,negated_conjecture,(c7_0|c8_0|~c4_2(a215,a216)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause271)).
cnf(c_0_332,negated_conjecture,(c4_0|c5_0|~c6_2(a268,a269)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause261)).
cnf(c_0_333,negated_conjecture,(ssSkC1|~c5_2(a155,a156)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause256)).
cnf(c_0_334,negated_conjecture,(ssSkC2|~c2_2(a164,a165)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause254)).
cnf(c_0_335,negated_conjecture,(ssSkC2|~c5_2(a164,a165)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause255)).
cnf(c_0_336,negated_conjecture,(ssSkC5|~c6_2(a176,a177)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause252)).
cnf(c_0_337,negated_conjecture,(ssSkC5|~c2_2(a176,a177)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause251)).
cnf(c_0_338,negated_conjecture,(c6_2(a227,a228)|~ssSkC16|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause267)).
cnf(c_0_339,negated_conjecture,(c2_2(a281,a282)|~ssSkC26|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause260)).
cnf(c_0_340,negated_conjecture,(c8_2(a244,a246)|~ssSkC19|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause264)).
cnf(c_0_341,negated_conjecture,(c8_2(a281,a282)|~ssSkC26|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause259)).
cnf(c_0_342,negated_conjecture,(c8_2(a281,a283)|~ssSkC26|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause257)).
cnf(c_0_343,negated_conjecture,(ssSkC5|~c4_2(a176,a178)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause249)).
cnf(c_0_344,negated_conjecture,(ssSkC5|~c3_2(a176,a177)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause253)).
cnf(c_0_345,negated_conjecture,(ssSkC5|~c7_2(a176,a178)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause250)).
cnf(c_0_346,negated_conjecture,(ssSkC6|~c6_2(a187,a188)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause248)).
cnf(c_0_347,negated_conjecture,(ssSkC6|~c4_2(a187,a188)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause247)).
cnf(c_0_348,negated_conjecture,(ssSkC12|~c1_2(a211,a212)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause245)).
cnf(c_0_349,negated_conjecture,(ssSkC12|~c8_2(a211,a212)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause246)).
cnf(c_0_350,negated_conjecture,(ssSkC13|~c5_2(a218,a219)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause244)).
cnf(c_0_351,negated_conjecture,(c4_2(a244,a245)|~ssSkC19|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause265)).
cnf(c_0_352,negated_conjecture,(ssSkC18|~c1_2(a237,a238)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause242)).
cnf(c_0_353,negated_conjecture,(c7_2(a166,a167)|~ssSkC2|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause280)).
cnf(c_0_354,negated_conjecture,(c7_2(a227,a228)|~ssSkC16|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause266)).
cnf(c_0_355,negated_conjecture,(c7_2(a281,a282)|~ssSkC26|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause258)).
cnf(c_0_356,negated_conjecture,(c6_2(a179,a180)|c7_0|~ssSkC5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause198)).
cnf(c_0_357,negated_conjecture,(ssSkC18|~c7_2(a237,a238)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause241)).
cnf(c_0_358,negated_conjecture,(ssSkC19|~c6_2(a241,a243)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause236)).
cnf(c_0_359,negated_conjecture,(ssSkC19|~c1_2(a241,a242)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause237)).
cnf(c_0_360,negated_conjecture,(ssSkC19|~c4_2(a241,a242)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause239)).
cnf(c_0_361,negated_conjecture,(ssSkC19|~c7_2(a241,a242)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause238)).
cnf(c_0_362,negated_conjecture,(ssSkC23|~c8_2(a265,a266)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause228)).
cnf(c_0_363,negated_conjecture,(ssSkC24|~c4_2(a271,a273)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause226)).
cnf(c_0_364,negated_conjecture,(ssSkC24|~c3_2(a271,a273)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause225)).
cnf(c_0_365,negated_conjecture,(ssSkC24|~c7_2(a271,a272)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause227)).
cnf(c_0_366,negated_conjecture,(c2_2(a179,a180)|c7_0|~ssSkC5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause196)).
cnf(c_0_367,negated_conjecture,(c5_2(a183,a184)|c7_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause274)).
cnf(c_0_368,negated_conjecture,(c3_2(a183,a184)|c7_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause275)).
cnf(c_0_369,negated_conjecture,(c3_2(a179,a180)|c7_0|~ssSkC5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause197)).
cnf(c_0_370,negated_conjecture,(c7_2(a183,a184)|c7_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause273)).
cnf(c_0_371,negated_conjecture,(c3_2(a215,a216)|c7_0|c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause188)).
cnf(c_0_372,negated_conjecture,(ssSkC25|~c3_2(a275,a276)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause224)).
cnf(c_0_373,negated_conjecture,(ssSkC26|~c1_2(a279,a280)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause223)).
cnf(c_0_374,negated_conjecture,(ssSkC26|~c4_2(a279,a280)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause222)).
cnf(c_0_375,negated_conjecture,(ssSkC27|~c5_2(a288,a289)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause219)).
cnf(c_0_376,negated_conjecture,(ssSkC27|~c8_2(a288,a289)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause217)).
cnf(c_0_377,negated_conjecture,(ssSkC27|~c4_2(a288,a289)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause218)).
cnf(c_0_378,negated_conjecture,(c5_0|~c6_2(a162,a163)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause206)).
cnf(c_0_379,negated_conjecture,(c5_0|~c6_2(a197,a198)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause190)).
cnf(c_0_380,negated_conjecture,(c1_2(a268,a270)|c4_0|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause180)).
cnf(c_0_381,negated_conjecture,(c5_0|~c5_2(a162,a163)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause205)).
cnf(c_0_382,negated_conjecture,(c5_0|~c1_2(a197,a198)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause191)).
cnf(c_0_383,negated_conjecture,(c1_0|~c2_2(a294,a295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause174)).
cnf(c_0_384,negated_conjecture,(c1_0|~c8_2(a294,a295)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause175)).
cnf(c_0_385,negated_conjecture,(c7_0|~ssSkC22|~c7_1(a259)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause184)).
cnf(c_0_386,negated_conjecture,(c7_0|~ssSkC5|~c2_1(a179)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause199)).
cnf(c_0_387,negated_conjecture,(c7_0|c6_0|~c7_1(a170)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause201)).
cnf(c_0_388,negated_conjecture,(c7_0|c6_0|~c5_1(a170)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause200)).
cnf(c_0_389,negated_conjecture,(c4_0|c5_0|~c7_1(a268)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause183)).
cnf(c_0_390,negated_conjecture,(c7_1(a166)|~ssSkC2|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause204)).
cnf(c_0_391,negated_conjecture,(c1_1(a166)|~ssSkC2|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause203)).
cnf(c_0_392,negated_conjecture,(c5_1(a157)|~ssSkC0|~ssSkC1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause114)).
cnf(c_0_393,negated_conjecture,(c3_2(a164,a165)|ssSkC2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause170)).
cnf(c_0_394,negated_conjecture,(c8_2(a176,a178)|ssSkC5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause167)).
cnf(c_0_395,negated_conjecture,(c8_1(a169)|~c4_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause277)).
cnf(c_0_396,negated_conjecture,(c2_1(a169)|~c4_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause276)).
cnf(c_0_397,negated_conjecture,(ndr1_1(a166)|~ssSkC2|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause202)).
cnf(c_0_398,negated_conjecture,(ndr1_1(a227)|~ssSkC16|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause186)).
cnf(c_0_399,negated_conjecture,(ndr1_1(a244)|~ssSkC19|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause185)).
cnf(c_0_400,negated_conjecture,(ndr1_1(a281)|~ssSkC26|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause177)).
cnf(c_0_401,negated_conjecture,(c6_1(a259)|c7_0|~ssSkC22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause91)).
cnf(c_0_402,negated_conjecture,(c4_1(a183)|c7_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause195)).
cnf(c_0_403,negated_conjecture,(c1_1(a183)|c7_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause194)).
cnf(c_0_404,negated_conjecture,(c8_2(a279,a280)|ssSkC26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause127)).
cnf(c_0_405,negated_conjecture,(ndr1_1(a183)|c7_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause193)).
cnf(c_0_406,negated_conjecture,(ndr1_1(a179)|c7_0|~ssSkC5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause109)).
cnf(c_0_407,negated_conjecture,(ssSkP1(X1)|~c5_2(X1,a224)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause243)).
cnf(c_0_408,negated_conjecture,(ssSkP2(X1)|~c5_2(X1,a239)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause240)).
cnf(c_0_409,negated_conjecture,(ssSkP4(X1)|~c6_2(X1,a254)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause235)).
cnf(c_0_410,negated_conjecture,(ssSkP4(X1)|~c5_2(X1,a254)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause234)).
cnf(c_0_411,negated_conjecture,(c8_2(a162,a163)|c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause111)).
cnf(c_0_412,negated_conjecture,(c3_2(a204,a205)|c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause102)).
cnf(c_0_413,negated_conjecture,(ndr1_1(a215)|c7_0|c8_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause99)).
cnf(c_0_414,negated_conjecture,(ssSkP5(X1)|~c1_2(X1,a260)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause233)).
cnf(c_0_415,negated_conjecture,(ssSkP5(X1)|~c7_2(X1,a260)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause231)).
cnf(c_0_416,negated_conjecture,(ssSkP6(X1)|~c5_2(X1,a263)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause229)).
cnf(c_0_417,negated_conjecture,(ssSkP6(X1)|~c7_2(X1,a263)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause230)).
cnf(c_0_418,negated_conjecture,(ssSkP8(X1)|~c7_2(X1,a285)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause221)).
cnf(c_0_419,negated_conjecture,(ssSkP9(X1)|~c1_2(X1,a290)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause216)).
cnf(c_0_420,negated_conjecture,(ssSkC1|~c7_1(a155)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause173)).
cnf(c_0_421,negated_conjecture,(ssSkP9(X1)|~c3_2(X1,a290)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause215)).
cnf(c_0_422,negated_conjecture,(ssSkP10(X1)|~c5_2(X1,a292)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause213)).
cnf(c_0_423,negated_conjecture,(c3_2(X1,a224)|ssSkP1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause154)).
cnf(c_0_424,negated_conjecture,(c8_2(X1,a239)|ssSkP2(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause147)).
cnf(c_0_425,negated_conjecture,(c7_2(X1,a239)|ssSkP2(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause146)).
cnf(c_0_426,negated_conjecture,(ssSkC13|~c5_1(a218)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause159)).
cnf(c_0_427,negated_conjecture,(ssSkC16|~c5_1(a226)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause153)).
cnf(c_0_428,negated_conjecture,(ssSkC18|~c7_1(a237)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause149)).
cnf(c_0_429,negated_conjecture,(ssSkC21|~c7_1(a253)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause140)).
cnf(c_0_430,negated_conjecture,(c5_2(X1,a250)|ssSkP3(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause143)).
cnf(c_0_431,negated_conjecture,(c8_2(X1,a250)|ssSkP3(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause142)).
cnf(c_0_432,negated_conjecture,(c7_2(X1,a250)|ssSkP3(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause141)).
cnf(c_0_433,negated_conjecture,(c1_2(X1,a254)|ssSkP4(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause138)).
cnf(c_0_434,negated_conjecture,(c4_2(X1,a263)|ssSkP6(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause137)).
cnf(c_0_435,negated_conjecture,(c5_2(X1,a284)|ssSkP7(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause126)).
cnf(c_0_436,negated_conjecture,(c1_2(X1,a284)|ssSkP7(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause125)).
cnf(c_0_437,negated_conjecture,(c3_2(X1,a284)|ssSkP7(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause124)).
cnf(c_0_438,negated_conjecture,(c3_2(X1,a285)|ssSkP8(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause123)).
cnf(c_0_439,negated_conjecture,(c6_1(a154)|ssSkC0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause83)).
cnf(c_0_440,negated_conjecture,(ndr1_1(X1)|ssSkP0(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause71)).
cnf(c_0_441,negated_conjecture,(ndr1_1(X1)|ssSkP1(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause67)).
cnf(c_0_442,negated_conjecture,(ndr1_1(X1)|ssSkP2(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause63)).
cnf(c_0_443,negated_conjecture,(ndr1_1(X1)|ssSkP3(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause61)).
cnf(c_0_444,negated_conjecture,(c2_1(a154)|ssSkC0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause84)).
cnf(c_0_445,negated_conjecture,(c6_1(a164)|ssSkC2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause80)).
cnf(c_0_446,negated_conjecture,(ndr1_1(a187)|ssSkC6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause75)).
cnf(c_0_447,negated_conjecture,(c1_1(a199)|ssSkC9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause74)).
cnf(c_0_448,negated_conjecture,(ndr1_1(X1)|ssSkP4(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause60)).
cnf(c_0_449,negated_conjecture,(ndr1_1(X1)|ssSkP5(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause59)).
cnf(c_0_450,negated_conjecture,(ndr1_1(X1)|ssSkP6(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause58)).
cnf(c_0_451,negated_conjecture,(ndr1_1(X1)|ssSkP9(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause48)).
cnf(c_0_452,negated_conjecture,(c4_0|~c1_0|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause212)).
cnf(c_0_453,negated_conjecture,(ndr1_1(a211)|ssSkC12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause69)).
cnf(c_0_454,negated_conjecture,(c1_1(a237)|ssSkC18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause65)).
cnf(c_0_455,negated_conjecture,(ndr1_1(a237)|ssSkC18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause64)).
cnf(c_0_456,negated_conjecture,(c2_1(a279)|ssSkC26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause53)).
cnf(c_0_457,negated_conjecture,(c3_0|~c5_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause210)).
cnf(c_0_458,negated_conjecture,(c5_0|~c4_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause211)).
cnf(c_0_459,negated_conjecture,(c7_0|c8_0|~c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause117)).
cnf(c_0_460,negated_conjecture,(c4_0|c3_0|~c5_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause118)).
cnf(c_0_461,negated_conjecture,(c2_1(a294)|c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause31)).
cnf(c_0_462,negated_conjecture,(~ssSkC2|~c8_2(a166,a167)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause289)).
cnf(c_0_463,negated_conjecture,(~ssSkC2|~c3_2(a166,a167)|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause290)).
cnf(c_0_464,negated_conjecture,(~ssSkC19|~c2_2(a244,a245)|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause287)).
cnf(c_0_465,negated_conjecture,(~ssSkC19|~c3_2(a244,a246)|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause285)).
cnf(c_0_466,negated_conjecture,(~ssSkC19|~c7_2(a244,a246)|~c6_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause286)).
cnf(c_0_467,negated_conjecture,(~ssSkC26|~c1_2(a281,a283)|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause284)).
cnf(c_0_468,negated_conjecture,(~ssSkC20|~c7_1(a249)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause263)).
cnf(c_0_469,negated_conjecture,(~ssSkC20|~c1_1(a249)|~c2_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause262)).
cnf(c_0_470,negated_conjecture,(~ssSkC0|~ssSkC1|~c4_1(a157)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause208)).
cnf(c_0_471,negated_conjecture,(~c6_1(a169)|~c4_0|~c7_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause288)).
cnf(c_0_472,negated_conjecture,(~c4_1(a153)|~c4_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause291)).
cnf(c_0_473,negated_conjecture,(~c1_1(a153)|~c4_0|~c1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause292)).
cnf(c_0_474,negated_conjecture,(~ssSkC17|~c3_1(a235)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause94)).
cnf(c_0_475,negated_conjecture,(~ssSkC17|~c1_1(a235)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause93)).
cnf(c_0_476,negated_conjecture,(ndr1_0), inference(spm,[status(thm)],[c_0_0, c_0_1]), ['final']).
cnf(c_0_477,negated_conjecture,(c4_2(X1,a293)|c4_1(X1)|~ssSkP10(X1)|~c2_0|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_476])]), ['final']).
cnf(c_0_478,negated_conjecture,(c8_0|c2_0|c1_0), inference(spm,[status(thm)],[c_0_3, c_0_4]), ['final']).
cnf(c_0_479,negated_conjecture,(c6_1(X1)|c5_1(X2)|c8_1(X1)|c2_1(X1)|ndr1_1(X2)|~c3_1(X2)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_476])]), ['final']).
cnf(c_0_480,negated_conjecture,(c4_2(X1,a293)|c4_1(X1)|c8_0|c1_0|~ssSkP10(X1)|~c4_0), inference(spm,[status(thm)],[c_0_477, c_0_478])).
cnf(c_0_481,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|~c7_2(X1,X2)|~c6_1(a160)|~c1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_476])]), ['final']).
cnf(c_0_482,negated_conjecture,(c6_1(X1)|c5_1(X2)|c8_1(X1)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)), inference(spm,[status(thm)],[c_0_479, c_0_4]), ['final']).
cnf(c_0_483,negated_conjecture,(c4_2(X1,a293)|c4_1(X1)|c8_0|c1_0|~ssSkP10(X1)), inference(spm,[status(thm)],[c_0_480, c_0_4]), ['final']).
cnf(c_0_484,negated_conjecture,(c4_1(X1)|ndr1_1(X1)|~ssSkP10(X1)|~c2_0|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_476])]), ['final']).
cnf(c_0_485,negated_conjecture,(~c6_2(X1,X2)|~c1_2(X1,X3)|~c1_2(X1,X2)|~c4_2(X1,X3)|~c3_2(X1,X2)|~c8_1(X1)|~c2_0|~ndr1_1(X1)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_476])]), ['final']).
cnf(c_0_486,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c5_1(X3)|c8_1(a160)|c8_1(X4)|c2_1(a160)|ndr1_1(X3)|c1_0|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X4)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_481, c_0_482]), ['final']).
cnf(c_0_487,negated_conjecture,(c4_2(X1,a185)|c3_1(X1)|c1_0|~c5_1(X1)|~c3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_476])]), ['final']).
cnf(c_0_488,negated_conjecture,(c5_2(X1,a278)|c1_1(X1)|c8_0|c7_0|~c4_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_476])]), ['final']).
cnf(c_0_489,negated_conjecture,(c4_2(X1,a293)|c7_2(X1,a292)|c4_1(X1)|c8_0|c1_0), inference(spm,[status(thm)],[c_0_483, c_0_11]), ['final']).
cnf(c_0_490,negated_conjecture,(c4_1(X1)|c8_0|ndr1_1(X1)|c1_0|~ssSkP10(X1)|~c4_0), inference(spm,[status(thm)],[c_0_484, c_0_478])).
cnf(c_0_491,negated_conjecture,(c6_1(X1)|c1_1(X1)|c6_0|ndr1_1(X2)|~ssSkP7(X2)|~ssSkP8(X2)|~c8_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_476])])).
cnf(c_0_492,negated_conjecture,(c8_0|c1_0|~c6_2(X1,X2)|~c1_2(X1,X3)|~c1_2(X1,X2)|~c4_2(X1,X3)|~c3_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)|~c4_0), inference(spm,[status(thm)],[c_0_485, c_0_478])).
cnf(c_0_493,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_1(X2)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)|~ndr1_1(a294)), inference(spm,[status(thm)],[c_0_486, c_0_13])).
cnf(c_0_494,negated_conjecture,(c3_0|c1_0|~c6_1(a168)), inference(spm,[status(thm)],[c_0_14, c_0_4]), ['final']).
cnf(c_0_495,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_1(X2)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)|~ndr1_1(a204)), inference(spm,[status(thm)],[c_0_486, c_0_15])).
cnf(c_0_496,negated_conjecture,(c2_2(X1,X2)|c8_2(a160,a161)|c4_1(X1)|c5_1(X1)|c8_1(X3)|~c7_2(X1,X2)|~c1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_476])]), ['final']).
cnf(c_0_497,negated_conjecture,(c2_2(a160,a161)|c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|~c7_2(X1,X2)|~c1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_476])]), ['final']).
cnf(c_0_498,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|ndr1_1(a160)|~c7_2(X1,X2)|~c1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_476])]), ['final']).
cnf(c_0_499,negated_conjecture,(c4_2(X1,a185)|c3_1(X1)|c6_0|c8_0|c1_0|~c5_1(X1)), inference(spm,[status(thm)],[c_0_487, c_0_19]), ['final']).
cnf(c_0_500,negated_conjecture,(c4_1(X1)|ndr1_1(X1)|c5_0|c1_0|~ssSkP10(X1)|~c4_0), inference(spm,[status(thm)],[c_0_484, c_0_20])).
cnf(c_0_501,negated_conjecture,(c1_0|~c1_1(a217)), inference(spm,[status(thm)],[c_0_21, c_0_4]), ['final']).
cnf(c_0_502,negated_conjecture,(c5_2(X1,a278)|c4_2(X1,a293)|c7_2(X1,a292)|c1_1(X1)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_489]), ['final']).
cnf(c_0_503,negated_conjecture,(c4_1(X1)|c8_0|ndr1_1(X1)|c1_0|~ssSkP10(X1)), inference(spm,[status(thm)],[c_0_490, c_0_4])).
cnf(c_0_504,negated_conjecture,(c6_1(X1)|c1_1(X1)|c6_0|ndr1_1(X2)|~ssSkP8(X2)|~c8_1(X1)), inference(spm,[status(thm)],[c_0_491, c_0_22])).
cnf(c_0_505,negated_conjecture,(c8_0|c1_0|~c6_2(X1,X2)|~c1_2(X1,X3)|~c1_2(X1,X2)|~c4_2(X1,X3)|~c3_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_492, c_0_4]), ['final']).
cnf(c_0_506,negated_conjecture,(ssSkC20|c7_2(X1,a248)|c2_1(X1)|~c8_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_476])]), ['final']).
cnf(c_0_507,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_1(X2)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_493, c_0_24]), ['final']).
cnf(c_0_508,negated_conjecture,(c4_1(X1)|c2_0|~c7_1(X2)|~c5_1(X1)|~c2_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_476])]), ['final']).
cnf(c_0_509,negated_conjecture,(c7_1(a217)|c1_0), inference(spm,[status(thm)],[c_0_26, c_0_4]), ['final']).
cnf(c_0_510,negated_conjecture,(c8_1(a217)|c1_0), inference(spm,[status(thm)],[c_0_27, c_0_4]), ['final']).
cnf(c_0_511,negated_conjecture,(c5_1(X1)|c3_0|c8_1(a168)|c2_1(a168)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_494, c_0_482]), ['final']).
cnf(c_0_512,negated_conjecture,(ssSkC20|c3_2(X1,a248)|c2_1(X1)|~c8_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_476])]), ['final']).
cnf(c_0_513,negated_conjecture,(ssSkC20|c2_1(X1)|ndr1_1(X1)|~c8_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_476])]), ['final']).
cnf(c_0_514,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_1(X2)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_495, c_0_30]), ['final']).
cnf(c_0_515,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|c8_1(X1)|c1_0|~c1_1(X1)|~ndr1_1(a294)), inference(spm,[status(thm)],[c_0_496, c_0_13])).
cnf(c_0_516,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c8_1(X1)|c1_0|~c1_1(X1)|~ndr1_1(a294)), inference(spm,[status(thm)],[c_0_497, c_0_13])).
cnf(c_0_517,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|c8_1(X1)|c1_0|~c1_1(X1)|~ndr1_1(a204)), inference(spm,[status(thm)],[c_0_496, c_0_15])).
cnf(c_0_518,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c8_1(X1)|c1_0|~c1_1(X1)|~ndr1_1(a204)), inference(spm,[status(thm)],[c_0_497, c_0_15])).
cnf(c_0_519,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c8_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)|~ndr1_1(a204)), inference(spm,[status(thm)],[c_0_498, c_0_15])).
cnf(c_0_520,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c8_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)|~ndr1_1(a294)), inference(spm,[status(thm)],[c_0_498, c_0_13])).
cnf(c_0_521,negated_conjecture,(c4_2(a174,a185)|c6_0|c8_0|ssSkC4|c1_0|~c5_1(a174)), inference(spm,[status(thm)],[c_0_31, c_0_499])).
cnf(c_0_522,negated_conjecture,(c3_1(X1)|c1_1(a256)|c1_1(X1)|c6_0|~c5_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_476])]), ['final']).
cnf(c_0_523,negated_conjecture,(c3_1(X1)|ndr1_1(X1)|c1_0|~c5_1(X1)|~c3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_476])]), ['final']).
cnf(c_0_524,negated_conjecture,(c3_1(X1)|c1_1(X1)|c6_0|ndr1_1(a256)|~c5_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_476])]), ['final']).
cnf(c_0_525,negated_conjecture,(c4_1(X1)|ndr1_1(X1)|c5_0|c1_0|~ssSkP10(X1)), inference(spm,[status(thm)],[c_0_500, c_0_4])).
cnf(c_0_526,negated_conjecture,(c4_2(X1,a293)|c4_1(X1)|c5_0|c1_0|~ssSkP10(X1)|~c4_0), inference(spm,[status(thm)],[c_0_477, c_0_20])).
cnf(c_0_527,negated_conjecture,(c2_2(X1,X2)|c2_2(X3,X4)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(X1)|c1_1(X3)|~c5_2(X3,X4)|~c7_2(X1,X2)|~c7_1(X3)|~ndr1_1(X1)|~ndr1_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_476])]), ['final']).
cnf(c_0_528,negated_conjecture,(c5_2(a217,a278)|c4_2(a217,a293)|c7_2(a217,a292)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_501, c_0_502]), ['final']).
cnf(c_0_529,negated_conjecture,(c1_1(X1)|c8_0|c7_0|ndr1_1(X1)|~c4_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_476])]), ['final']).
cnf(c_0_530,negated_conjecture,(c4_1(X1)|c8_0|ndr1_1(X1)|c1_0), inference(spm,[status(thm)],[c_0_503, c_0_37]), ['final']).
cnf(c_0_531,negated_conjecture,(c6_1(X1)|c1_1(X1)|c6_0|ndr1_1(X2)|~c8_1(X1)), inference(spm,[status(thm)],[c_0_504, c_0_38]), ['final']).
cnf(c_0_532,negated_conjecture,(c8_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c8_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_505, c_0_39]), ['final']).
cnf(c_0_533,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(X2)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_506, c_0_507]), ['final']).
cnf(c_0_534,negated_conjecture,(c8_2(X1,a292)|c4_2(X1,a293)|c4_1(X1)|c8_0|c1_0), inference(spm,[status(thm)],[c_0_483, c_0_40]), ['final']).
cnf(c_0_535,negated_conjecture,(c4_1(X1)|c2_0|c1_0|~c5_1(X1)|~c2_1(a217)), inference(spm,[status(thm)],[c_0_508, c_0_509]), ['final']).
cnf(c_0_536,negated_conjecture,(ssSkC20|c7_2(a217,a248)|c2_1(a217)|c1_0), inference(spm,[status(thm)],[c_0_506, c_0_510]), ['final']).
cnf(c_0_537,negated_conjecture,(c3_0|c1_0|~c2_1(a168)), inference(spm,[status(thm)],[c_0_41, c_0_4]), ['final']).
cnf(c_0_538,negated_conjecture,(ssSkC20|c7_2(a168,a248)|c5_1(X1)|c3_0|c2_1(a168)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_511])).
cnf(c_0_539,negated_conjecture,(c3_1(X1)|c4_1(X1)|c5_1(X1)|c6_0|ndr1_1(a296)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_476])]), ['final']).
cnf(c_0_540,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(X2)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_512, c_0_507]), ['final']).
cnf(c_0_541,negated_conjecture,(ssSkC20|c5_1(X1)|c3_0|c2_1(a168)|ndr1_1(a168)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_511])).
cnf(c_0_542,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(X2)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_506, c_0_514]), ['final']).
cnf(c_0_543,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(X2)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_512, c_0_514]), ['final']).
cnf(c_0_544,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|c8_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_515, c_0_24]), ['final']).
cnf(c_0_545,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c8_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_516, c_0_24]), ['final']).
cnf(c_0_546,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|c8_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_517, c_0_30]), ['final']).
cnf(c_0_547,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c8_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_518, c_0_30]), ['final']).
cnf(c_0_548,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c8_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_519, c_0_30]), ['final']).
cnf(c_0_549,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c8_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_520, c_0_24]), ['final']).
cnf(c_0_550,negated_conjecture,(c6_2(a231,a232)|c1_2(X1,X2)|c3_2(X3,X4)|c2_1(X1)|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_476])]), ['final']).
cnf(c_0_551,negated_conjecture,(c4_2(a174,a185)|c6_0|c8_0|ssSkC4|c1_0), inference(spm,[status(thm)],[c_0_521, c_0_44]), ['final']).
cnf(c_0_552,negated_conjecture,(c1_1(a256)|c1_1(a174)|c6_0|ssSkC4|~c5_1(a174)), inference(spm,[status(thm)],[c_0_31, c_0_522])).
cnf(c_0_553,negated_conjecture,(c3_1(X1)|c6_0|c8_0|ndr1_1(X1)|c1_0|~c5_1(X1)), inference(spm,[status(thm)],[c_0_523, c_0_19]), ['final']).
cnf(c_0_554,negated_conjecture,(c3_1(X1)|c7_2(X1,a185)|c1_0|~c5_1(X1)|~c3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_476])]), ['final']).
cnf(c_0_555,negated_conjecture,(c1_1(a174)|c6_0|ndr1_1(a256)|ssSkC4|~c5_1(a174)), inference(spm,[status(thm)],[c_0_31, c_0_524])).
cnf(c_0_556,negated_conjecture,(c4_1(X1)|ndr1_1(X1)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_525, c_0_37]), ['final']).
cnf(c_0_557,negated_conjecture,(c4_2(X1,a293)|c4_1(X1)|c5_0|c1_0|~ssSkP10(X1)), inference(spm,[status(thm)],[c_0_526, c_0_4]), ['final']).
cnf(c_0_558,negated_conjecture,(c2_2(a217,a278)|c2_2(X1,X2)|c1_2(X1,X2)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c4_1(X1)|c1_1(a217)|c8_0|c7_0|c1_0|~c7_2(X1,X2)|~c7_1(a217)|~ndr1_1(a217)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_528])).
cnf(c_0_559,negated_conjecture,(c1_1(X1)|c8_0|c7_0|ndr1_1(X1)|c1_0), inference(spm,[status(thm)],[c_0_529, c_0_530]), ['final']).
cnf(c_0_560,negated_conjecture,(c1_1(a168)|c3_0|c6_0|ndr1_1(X1)|c1_0|~c8_1(a168)), inference(spm,[status(thm)],[c_0_494, c_0_531]), ['final']).
cnf(c_0_561,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_533]), ['final']).
cnf(c_0_562,negated_conjecture,(c5_2(X1,a278)|c8_2(X1,a292)|c4_2(X1,a293)|c1_1(X1)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_534]), ['final']).
cnf(c_0_563,negated_conjecture,(ssSkC20|c7_2(a217,a248)|c4_1(X1)|c2_0|c1_0|~c5_1(X1)), inference(spm,[status(thm)],[c_0_535, c_0_536]), ['final']).
cnf(c_0_564,negated_conjecture,(ssSkC20|c7_2(a168,a248)|c5_1(X1)|c3_0|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_537, c_0_538]), ['final']).
cnf(c_0_565,negated_conjecture,(c3_1(a223)|c5_1(a223)|c6_0|ndr1_1(a296)|ssSkC15), inference(spm,[status(thm)],[c_0_46, c_0_539]), ['final']).
cnf(c_0_566,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_540]), ['final']).
cnf(c_0_567,negated_conjecture,(ssSkC20|c5_1(X1)|c3_0|ndr1_1(a168)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_537, c_0_541]), ['final']).
cnf(c_0_568,negated_conjecture,(ssSkC20|c2_1(a217)|ndr1_1(a217)|c1_0), inference(spm,[status(thm)],[c_0_513, c_0_510]), ['final']).
cnf(c_0_569,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c1_1(a160)|c5_1(X1)|c6_0|c8_1(X3)|ndr1_1(X4)|~c7_2(X1,X2)|~c1_1(X3)|~c8_1(a160)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_481, c_0_531]), ['final']).
cnf(c_0_570,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_542]), ['final']).
cnf(c_0_571,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_543]), ['final']).
cnf(c_0_572,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|c8_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_544]), ['final']).
cnf(c_0_573,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c8_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_545]), ['final']).
cnf(c_0_574,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|c8_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_546]), ['final']).
cnf(c_0_575,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c8_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_547]), ['final']).
cnf(c_0_576,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c8_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_548]), ['final']).
cnf(c_0_577,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c8_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_549]), ['final']).
cnf(c_0_578,negated_conjecture,(c5_2(X1,a147)|c4_2(X1,X2)|c4_2(X1,X3)|c3_2(X1,X2)|~c5_2(X1,X3)|~c7_2(X1,X3)|~c3_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_476])]), ['final']).
cnf(c_0_579,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(X1,X2)|c6_0|c8_0|c2_1(a174)|ssSkC4|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a174,a185)|~c1_1(a174)|~c5_1(X1)|~ndr1_1(a174)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_550, c_0_551]), ['final']).
cnf(c_0_580,negated_conjecture,(c1_1(a256)|c1_1(a174)|c6_0|ssSkC4), inference(spm,[status(thm)],[c_0_552, c_0_44]), ['final']).
cnf(c_0_581,negated_conjecture,(c6_0|c8_0|ndr1_1(a174)|ssSkC4|c1_0|~c5_1(a174)), inference(spm,[status(thm)],[c_0_31, c_0_553])).
cnf(c_0_582,negated_conjecture,(c3_1(X1)|c7_2(X1,a185)|c6_0|c8_0|c1_0|~c5_1(X1)), inference(spm,[status(thm)],[c_0_554, c_0_19]), ['final']).
cnf(c_0_583,negated_conjecture,(c1_1(a174)|c6_0|ndr1_1(a256)|ssSkC4), inference(spm,[status(thm)],[c_0_555, c_0_44]), ['final']).
cnf(c_0_584,negated_conjecture,(ndr1_1(a223)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_46, c_0_556]), ['final']).
cnf(c_0_585,negated_conjecture,(c4_2(X1,a293)|c7_2(X1,a292)|c4_1(X1)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_557, c_0_11]), ['final']).
cnf(c_0_586,negated_conjecture,(c2_2(a217,a278)|c2_2(X1,X2)|c1_2(X1,X2)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c4_1(X1)|c1_1(a217)|c8_0|c7_0|c1_0|~c7_2(X1,X2)|~ndr1_1(a217)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_558, c_0_509])).
cnf(c_0_587,negated_conjecture,(c8_0|c7_0|ndr1_1(a217)|c1_0), inference(spm,[status(thm)],[c_0_501, c_0_559]), ['final']).
cnf(c_0_588,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c2_1(X1)|~c4_2(X1,X2)|~c3_1(X1)|~c4_1(a159)|~c2_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_476])]), ['final']).
cnf(c_0_589,negated_conjecture,(c4_2(X1,a258)|c1_1(X1)|ssSkC22|~c7_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_476])]), ['final']).
cnf(c_0_590,negated_conjecture,(c8_0|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_46, c_0_530]), ['final']).
cnf(c_0_591,negated_conjecture,(c4_1(X1)|c2_0|ssSkC11|~c5_1(X1)|~c2_1(a202)), inference(spm,[status(thm)],[c_0_508, c_0_50]), ['final']).
cnf(c_0_592,negated_conjecture,(ssSkC7|c3_1(X1)|c2_1(X1)|ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_476])]), ['final']).
cnf(c_0_593,negated_conjecture,(c1_1(a168)|c5_1(X1)|c3_0|c6_0|c2_1(a168)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_560, c_0_511])).
cnf(c_0_594,negated_conjecture,(c5_2(a296,a297)|c3_1(X1)|c4_1(X1)|c5_1(X1)|c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_476])]), ['final']).
cnf(c_0_595,negated_conjecture,(c5_2(a275,X1)|ssSkC25|~c2_2(a275,X1)|~c7_2(a275,X1)), inference(spm,[status(thm)],[c_0_53, c_0_54]), ['final']).
cnf(c_0_596,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_561, c_0_502])).
cnf(c_0_597,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_561, c_0_562])).
cnf(c_0_598,negated_conjecture,(ssSkC20|c7_2(a217,a248)|c2_0|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_46, c_0_563]), ['final']).
cnf(c_0_599,negated_conjecture,(ssSkC20|c7_2(a168,a248)|c5_1(a223)|c3_0|c6_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_564, c_0_565]), ['final']).
cnf(c_0_600,negated_conjecture,(c1_1(X1)|ssSkC22|ndr1_1(X1)|~c7_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_476])]), ['final']).
cnf(c_0_601,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_566, c_0_502])).
cnf(c_0_602,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_566, c_0_562])).
cnf(c_0_603,negated_conjecture,(ssSkC20|c5_1(a223)|c3_0|c6_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_567, c_0_565]), ['final']).
cnf(c_0_604,negated_conjecture,(ssSkC20|c4_1(X1)|c2_0|ndr1_1(a217)|c1_0|~c5_1(X1)), inference(spm,[status(thm)],[c_0_535, c_0_568]), ['final']).
cnf(c_0_605,negated_conjecture,(c2_2(a294,a295)|c2_2(X1,X2)|c4_1(a294)|c4_1(X1)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c5_1(X3)|c6_0|c8_1(X4)|c8_1(X5)|c2_1(a160)|ndr1_1(X6)|ndr1_1(X3)|c1_0|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X4)|~c1_1(X5)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_569, c_0_507])).
cnf(c_0_606,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_570, c_0_502])).
cnf(c_0_607,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_570, c_0_562])).
cnf(c_0_608,negated_conjecture,(c2_2(a204,a205)|c2_2(X1,X2)|c4_1(a204)|c4_1(X1)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c5_1(X3)|c6_0|c8_1(X4)|c8_1(X5)|c2_1(a160)|ndr1_1(X6)|ndr1_1(X3)|c1_0|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X4)|~c1_1(X5)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_569, c_0_514])).
cnf(c_0_609,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_571, c_0_502])).
cnf(c_0_610,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_571, c_0_562])).
cnf(c_0_611,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_572, c_0_502])).
cnf(c_0_612,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_572, c_0_562])).
cnf(c_0_613,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_573, c_0_502])).
cnf(c_0_614,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_573, c_0_562])).
cnf(c_0_615,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_574, c_0_502])).
cnf(c_0_616,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_574, c_0_562])).
cnf(c_0_617,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_575, c_0_502])).
cnf(c_0_618,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_575, c_0_562])).
cnf(c_0_619,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_576, c_0_502])).
cnf(c_0_620,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_576, c_0_562])).
cnf(c_0_621,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_577, c_0_502])).
cnf(c_0_622,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_577, c_0_562])).
cnf(c_0_623,negated_conjecture,(c6_2(a231,a232)|c1_2(a265,a266)|c3_2(X1,X2)|c2_1(a265)|ssSkC23|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a265,a266)|~c1_1(a265)|~c5_1(X1)|~ndr1_1(a265)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_550, c_0_56]), ['final']).
cnf(c_0_624,negated_conjecture,(c5_2(X1,a147)|c4_2(X1,X2)|c4_2(X1,X3)|c3_2(X1,X2)|c6_0|c8_0|~c5_2(X1,X3)|~c7_2(X1,X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_578, c_0_19]), ['final']).
cnf(c_0_625,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(X1,X2)|c1_1(a256)|c6_0|c8_0|c2_1(a174)|ssSkC4|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a174,a185)|~c5_1(X1)|~ndr1_1(a174)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_579, c_0_580])).
cnf(c_0_626,negated_conjecture,(c6_0|c8_0|ndr1_1(a174)|ssSkC4|c1_0), inference(spm,[status(thm)],[c_0_581, c_0_44]), ['final']).
cnf(c_0_627,negated_conjecture,(c7_2(a174,a185)|c6_0|c8_0|ssSkC4|c1_0|~c5_1(a174)), inference(spm,[status(thm)],[c_0_31, c_0_582])).
cnf(c_0_628,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(X1,X2)|c6_0|c8_0|c2_1(a174)|ndr1_1(a256)|ssSkC4|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a174,a185)|~c5_1(X1)|~ndr1_1(a174)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_579, c_0_583])).
cnf(c_0_629,negated_conjecture,(c6_2(a223,X1)|ssSkC15|c5_0|c1_0|~c4_2(a223,X1)), inference(spm,[status(thm)],[c_0_57, c_0_584]), ['final']).
cnf(c_0_630,negated_conjecture,(c4_2(a223,a293)|c7_2(a223,a292)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_46, c_0_585]), ['final']).
cnf(c_0_631,negated_conjecture,(c2_2(a217,a278)|c2_2(X1,X2)|c1_2(X1,X2)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c4_1(X1)|c1_1(a217)|c8_0|c7_0|c1_0|~c7_2(X1,X2)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_586, c_0_587]), ['final']).
cnf(c_0_632,negated_conjecture,(ssSkC20|c3_2(a217,a248)|c2_1(a217)|c1_0), inference(spm,[status(thm)],[c_0_512, c_0_510]), ['final']).
cnf(c_0_633,negated_conjecture,(ssSkC20|c3_2(a168,a248)|c5_1(X1)|c3_0|c2_1(a168)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_511])).
cnf(c_0_634,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c8_0|c2_1(X1)|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~c4_1(a159)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_588, c_0_478]), ['final']).
cnf(c_0_635,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c2_1(X1)|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~c4_1(a159)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_588, c_0_20]), ['final']).
cnf(c_0_636,negated_conjecture,(c8_2(X1,a292)|c4_2(X1,a293)|c4_1(X1)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_557, c_0_40]), ['final']).
cnf(c_0_637,negated_conjecture,(c5_2(X1,a185)|c3_1(X1)|c1_0|~c5_1(X1)|~c3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_476])]), ['final']).
cnf(c_0_638,negated_conjecture,(c4_2(a217,a258)|c1_1(a217)|ssSkC22|c1_0), inference(spm,[status(thm)],[c_0_589, c_0_509])).
cnf(c_0_639,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c2_1(a159)|c2_1(X1)|~c4_2(X1,X2)|~c3_1(X1)|~c2_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_476])]), ['final']).
cnf(c_0_640,negated_conjecture,(c6_2(a223,X1)|c8_0|ssSkC15|c1_0|~c4_2(a223,X1)), inference(spm,[status(thm)],[c_0_57, c_0_590]), ['final']).
cnf(c_0_641,negated_conjecture,(c4_2(a223,a293)|c7_2(a223,a292)|c8_0|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_46, c_0_489]), ['final']).
cnf(c_0_642,negated_conjecture,(ssSkC7|c3_1(a202)|c4_1(X1)|c2_0|ndr1_1(a202)|ssSkC11|~c5_1(X1)), inference(spm,[status(thm)],[c_0_591, c_0_592]), ['final']).
cnf(c_0_643,negated_conjecture,(c1_1(a168)|c5_1(X1)|c3_0|c6_0|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_537, c_0_593]), ['final']).
cnf(c_0_644,negated_conjecture,(c5_2(a296,a297)|c3_1(a223)|c5_1(a223)|c6_0|ssSkC15), inference(spm,[status(thm)],[c_0_46, c_0_594]), ['final']).
cnf(c_0_645,negated_conjecture,(ssSkC7|c3_1(a217)|c4_1(X1)|c2_0|ndr1_1(a217)|c1_0|~c5_1(X1)), inference(spm,[status(thm)],[c_0_535, c_0_592]), ['final']).
cnf(c_0_646,negated_conjecture,(c5_2(a275,a276)|ssSkC25|~c7_2(a275,a276)), inference(spm,[status(thm)],[c_0_595, c_0_60])).
cnf(c_0_647,negated_conjecture,(c8_2(X1,a207)|c3_0|c4_0|~c5_2(X1,X2)|~c4_2(X1,X2)|~ssSkP0(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_476])]), ['final']).
cnf(c_0_648,negated_conjecture,(c1_2(X1,X2)|c3_2(X3,X4)|c2_1(X1)|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~c8_1(a231)|~ndr1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_476])]), ['final']).
cnf(c_0_649,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_596, c_0_63])).
cnf(c_0_650,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_597, c_0_63])).
cnf(c_0_651,negated_conjecture,(ssSkC20|c7_2(a217,a248)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_598, c_0_599]), ['final']).
cnf(c_0_652,negated_conjecture,(c1_1(a217)|ssSkC22|ndr1_1(a217)|c1_0), inference(spm,[status(thm)],[c_0_600, c_0_509])).
cnf(c_0_653,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_601, c_0_63])).
cnf(c_0_654,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_602, c_0_63])).
cnf(c_0_655,negated_conjecture,(ssSkC20|c7_2(a217,a248)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_598, c_0_603]), ['final']).
cnf(c_0_656,negated_conjecture,(ssSkC20|c2_0|ndr1_1(a217)|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_46, c_0_604]), ['final']).
cnf(c_0_657,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c8_1(X2)|c8_1(X3)|c2_1(a160)|ndr1_1(X4)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)|~c1_1(X3)|~ndr1_1(a294)), inference(spm,[status(thm)],[c_0_605, c_0_13])).
cnf(c_0_658,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_606, c_0_63])).
cnf(c_0_659,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_607, c_0_63])).
cnf(c_0_660,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c8_1(X2)|c8_1(X3)|c2_1(a160)|ndr1_1(X4)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)|~c1_1(X3)|~ndr1_1(a204)), inference(spm,[status(thm)],[c_0_608, c_0_15])).
cnf(c_0_661,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_609, c_0_63])).
cnf(c_0_662,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_610, c_0_63])).
cnf(c_0_663,negated_conjecture,(c5_0|c1_0|~c6_2(X1,X2)|~c1_2(X1,X3)|~c1_2(X1,X2)|~c4_2(X1,X3)|~c3_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)|~c4_0), inference(spm,[status(thm)],[c_0_485, c_0_20])).
cnf(c_0_664,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_611, c_0_63])).
cnf(c_0_665,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_612, c_0_63])).
cnf(c_0_666,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_613, c_0_63])).
cnf(c_0_667,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_614, c_0_63])).
cnf(c_0_668,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_615, c_0_63])).
cnf(c_0_669,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_616, c_0_63])).
cnf(c_0_670,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_617, c_0_63])).
cnf(c_0_671,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_618, c_0_63])).
cnf(c_0_672,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_619, c_0_63])).
cnf(c_0_673,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_620, c_0_63])).
cnf(c_0_674,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_621, c_0_63])).
cnf(c_0_675,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_622, c_0_63])).
cnf(c_0_676,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c7_2(X1,a248)|c4_1(a294)|c5_1(a294)|c5_1(X2)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_533]), ['final']).
cnf(c_0_677,negated_conjecture,(c6_2(a231,a232)|c5_2(a265,a278)|c1_2(a265,a266)|c4_2(a265,a293)|c3_2(X1,X2)|c7_2(a265,a292)|c8_0|c7_0|c2_1(a265)|ssSkC23|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a265,a266)|~c5_1(X1)|~ndr1_1(a265)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_623, c_0_502])).
cnf(c_0_678,negated_conjecture,(c7_2(a265,X1)|ssSkC23|~c4_2(a265,X1)), inference(spm,[status(thm)],[c_0_64, c_0_65]), ['final']).
cnf(c_0_679,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(X1,a248)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X2)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_533]), ['final']).
cnf(c_0_680,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c7_2(X1,a248)|c4_1(a294)|c5_1(a294)|c5_1(X2)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_540]), ['final']).
cnf(c_0_681,negated_conjecture,(c6_2(a231,a232)|c5_2(a265,a278)|c1_2(a265,a266)|c8_2(a265,a292)|c4_2(a265,a293)|c3_2(X1,X2)|c8_0|c7_0|c2_1(a265)|ssSkC23|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a265,a266)|~c5_1(X1)|~ndr1_1(a265)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_623, c_0_562])).
cnf(c_0_682,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c3_2(X1,a248)|c4_1(a294)|c5_1(a294)|c5_1(X2)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_540]), ['final']).
cnf(c_0_683,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c7_2(X1,a248)|c4_1(a204)|c5_1(a204)|c5_1(X2)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_542]), ['final']).
cnf(c_0_684,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(X1,a248)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X2)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_542]), ['final']).
cnf(c_0_685,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c7_2(X1,a248)|c4_1(a204)|c5_1(a204)|c5_1(X2)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_543]), ['final']).
cnf(c_0_686,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c3_2(X1,a248)|c4_1(a204)|c5_1(a204)|c5_1(X2)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_543]), ['final']).
cnf(c_0_687,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a160)), inference(ef,[status(thm)],[c_0_507]), ['final']).
cnf(c_0_688,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a160)), inference(ef,[status(thm)],[c_0_514]), ['final']).
cnf(c_0_689,negated_conjecture,(c5_2(a275,a147)|c4_2(a275,a276)|c4_2(a275,X1)|c3_2(a275,X1)|c6_0|c8_0|ssSkC25|~c5_2(a275,a276)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_624, c_0_66])).
cnf(c_0_690,negated_conjecture,(c5_2(X1,X2)|c3_2(X1,X2)|c3_1(X3)|c3_1(X1)|c4_1(X1)|c2_1(X3)|~c6_2(X3,X4)|~c8_2(X3,X4)|~c8_2(X1,X2)|~c3_2(X3,X4)|~ndr1_1(X3)|~ndr1_1(X1)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_476])]), ['final']).
cnf(c_0_691,negated_conjecture,(c6_2(X1,X2)|c5_2(X1,X3)|c4_2(X1,X3)|c7_2(X1,X2)|c5_0|c1_0|~c1_2(X1,X3)|~c8_2(X1,X2)|~c1_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_476])]), ['final']).
cnf(c_0_692,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(X1,X2)|c1_1(a256)|c6_0|c8_0|c2_1(a174)|ssSkC4|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a174,a185)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_625, c_0_626])).
cnf(c_0_693,negated_conjecture,(c7_2(a174,a185)|c6_0|c8_0|ssSkC4|c1_0), inference(spm,[status(thm)],[c_0_627, c_0_44]), ['final']).
cnf(c_0_694,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(X1,X2)|c6_0|c8_0|c2_1(a174)|ndr1_1(a256)|ssSkC4|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a174,a185)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_628, c_0_626])).
cnf(c_0_695,negated_conjecture,(c2_1(a181)|c5_0|~c6_2(X1,X2)|~c5_2(X1,X2)|~c1_2(X1,X3)|~c8_2(X1,X3)|~c3_2(X1,X3)|~c3_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_476])]), ['final']).
cnf(c_0_696,negated_conjecture,(c6_2(a223,a293)|c7_2(a223,a292)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_629, c_0_630]), ['final']).
cnf(c_0_697,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_507]), ['final']).
cnf(c_0_698,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_514]), ['final']).
cnf(c_0_699,negated_conjecture,(c4_2(a158,a293)|c7_2(a158,a292)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_71, c_0_585]), ['final']).
cnf(c_0_700,negated_conjecture,(c2_2(a217,a278)|c2_2(a294,a295)|c1_2(a294,a295)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c4_1(a294)|c1_1(a217)|c8_0|c7_0|c1_0|~ndr1_1(a294)), inference(spm,[status(thm)],[c_0_631, c_0_13])).
cnf(c_0_701,negated_conjecture,(c2_2(a217,a278)|c2_2(a204,a205)|c1_2(a204,a205)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c4_1(a204)|c1_1(a217)|c8_0|c7_0|c1_0|~ndr1_1(a204)), inference(spm,[status(thm)],[c_0_631, c_0_15])).
cnf(c_0_702,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c8_2(a160,a161)|c7_2(X1,a248)|c4_1(a294)|c5_1(a294)|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_544]), ['final']).
cnf(c_0_703,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c8_2(a160,a161)|c3_2(X1,a248)|c4_1(a294)|c5_1(a294)|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_544]), ['final']).
cnf(c_0_704,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c8_2(a160,a161)|c7_2(X1,a248)|c4_1(a204)|c5_1(a204)|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_546]), ['final']).
cnf(c_0_705,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c8_2(a160,a161)|c3_2(X1,a248)|c4_1(a204)|c5_1(a204)|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_546]), ['final']).
cnf(c_0_706,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c7_2(X1,a248)|c4_1(a294)|c5_1(a294)|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_545]), ['final']).
cnf(c_0_707,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c3_2(X1,a248)|c4_1(a294)|c5_1(a294)|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_545]), ['final']).
cnf(c_0_708,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c7_2(X1,a248)|c4_1(a204)|c5_1(a204)|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_547]), ['final']).
cnf(c_0_709,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c3_2(X1,a248)|c4_1(a204)|c5_1(a204)|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_547]), ['final']).
cnf(c_0_710,negated_conjecture,(ssSkC20|c3_2(a217,a248)|c4_1(X1)|c2_0|c1_0|~c5_1(X1)), inference(spm,[status(thm)],[c_0_535, c_0_632]), ['final']).
cnf(c_0_711,negated_conjecture,(ssSkC20|c3_2(a168,a248)|c5_1(X1)|c3_0|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_537, c_0_633]), ['final']).
cnf(c_0_712,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(X1,a248)|c4_1(a294)|c5_1(a294)|c2_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_549]), ['final']).
cnf(c_0_713,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(X1,a248)|c4_1(a294)|c5_1(a294)|c2_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_549]), ['final']).
cnf(c_0_714,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(X1,a248)|c4_1(a204)|c5_1(a204)|c2_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_548]), ['final']).
cnf(c_0_715,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(X1,a248)|c4_1(a204)|c5_1(a204)|c2_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_548]), ['final']).
cnf(c_0_716,negated_conjecture,(c2_2(X1,X2)|c5_2(a296,a297)|c8_2(X1,X2)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(X1)|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_634, c_0_594]), ['final']).
cnf(c_0_717,negated_conjecture,(c2_2(X1,X2)|c5_2(a296,a297)|c8_2(X1,X2)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(X1)|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_635, c_0_594]), ['final']).
cnf(c_0_718,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c4_2(a159,a293)|c7_2(a159,a292)|c8_0|c2_1(X1)|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_634, c_0_489]), ['final']).
cnf(c_0_719,negated_conjecture,(c2_2(X1,X2)|c8_2(a159,a292)|c8_2(X1,X2)|c4_2(a159,a293)|c8_0|c2_1(X1)|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_634, c_0_534]), ['final']).
cnf(c_0_720,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c4_2(a159,a293)|c7_2(a159,a292)|c2_1(X1)|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_635, c_0_585]), ['final']).
cnf(c_0_721,negated_conjecture,(c2_2(X1,X2)|c8_2(a159,a292)|c8_2(X1,X2)|c4_2(a159,a293)|c2_1(X1)|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_635, c_0_636]), ['final']).
cnf(c_0_722,negated_conjecture,(c5_2(X1,a185)|c3_1(X1)|c6_0|c8_0|c1_0|~c5_1(X1)), inference(spm,[status(thm)],[c_0_637, c_0_19]), ['final']).
cnf(c_0_723,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|ndr1_1(a208)|~c6_2(X1,X2)|~c1_2(X1,X3)|~c7_2(X1,X3)|~c8_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_476])]), ['final']).
cnf(c_0_724,negated_conjecture,(c6_0|c8_0|c7_0), inference(spm,[status(thm)],[c_0_73, c_0_19]), ['final']).
cnf(c_0_725,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|c8_1(a208)|~c6_2(X1,X2)|~c1_2(X1,X3)|~c7_2(X1,X3)|~c8_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_476])]), ['final']).
cnf(c_0_726,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(X1)|ndr1_1(a296)|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_634, c_0_539]), ['final']).
cnf(c_0_727,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(X1)|ndr1_1(a296)|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_635, c_0_539]), ['final']).
cnf(c_0_728,negated_conjecture,(c4_2(a217,a258)|ssSkC22|c1_0), inference(spm,[status(thm)],[c_0_501, c_0_638]), ['final']).
cnf(c_0_729,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c8_0|c2_1(X1)|ndr1_1(a159)|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_634, c_0_530]), ['final']).
cnf(c_0_730,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c2_1(X1)|ndr1_1(a159)|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_635, c_0_556]), ['final']).
cnf(c_0_731,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c8_0|c2_1(a159)|c2_1(X1)|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_639, c_0_478]), ['final']).
cnf(c_0_732,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,X2)|c2_1(a159)|c2_1(X1)|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_639, c_0_20]), ['final']).
cnf(c_0_733,negated_conjecture,(c2_2(X1,X2)|c7_2(a193,a194)|c5_1(X1)|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_476])]), ['final']).
cnf(c_0_734,negated_conjecture,(c6_2(a223,a293)|c7_2(a223,a292)|c8_0|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_640, c_0_641]), ['final']).
cnf(c_0_735,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|ndr1_1(a193)|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_476])]), ['final']).
cnf(c_0_736,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|c2_1(X1)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_544]), ['final']).
cnf(c_0_737,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|c2_1(X1)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_546]), ['final']).
cnf(c_0_738,negated_conjecture,(ssSkC7|c3_1(a202)|c2_0|ndr1_1(a202)|ssSkC11|ssSkC15|~c5_1(a223)), inference(spm,[status(thm)],[c_0_46, c_0_642]), ['final']).
cnf(c_0_739,negated_conjecture,(c5_2(a296,a297)|c1_1(a168)|c5_1(a223)|c3_0|c6_0|ndr1_1(a223)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_643, c_0_644]), ['final']).
cnf(c_0_740,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c2_1(X1)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_545]), ['final']).
cnf(c_0_741,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c2_1(X1)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_547]), ['final']).
cnf(c_0_742,negated_conjecture,(ssSkC7|c3_1(a217)|c2_0|ndr1_1(a217)|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_46, c_0_645]), ['final']).
cnf(c_0_743,negated_conjecture,(c1_1(a168)|c5_1(a223)|c3_0|c6_0|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_643, c_0_565]), ['final']).
cnf(c_0_744,negated_conjecture,(c5_2(X1,X2)|c6_1(X1)|c2_0|c4_0|~c6_2(X1,X2)|~c5_2(X1,X3)|~c8_2(X1,X2)|~c7_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_476])]), ['final']).
cnf(c_0_745,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c2_1(X1)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_549]), ['final']).
cnf(c_0_746,negated_conjecture,(ssSkC7|c1_2(X1,a190)|c3_1(X1)|c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_476])]), ['final']).
cnf(c_0_747,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c2_1(X1)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_548]), ['final']).
cnf(c_0_748,negated_conjecture,(c5_2(a275,a276)|ssSkC25), inference(spm,[status(thm)],[c_0_646, c_0_66]), ['final']).
cnf(c_0_749,negated_conjecture,(c8_2(a241,a207)|c3_0|ssSkC19|c4_0|~c4_2(a241,a243)|~ssSkP0(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_647, c_0_79]), ['final']).
cnf(c_0_750,negated_conjecture,(c8_2(X1,X2)|c3_2(X1,X2)|c8_0|c2_1(X1)|~c2_2(X1,X2)|~c5_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_476])]), ['final']).
cnf(c_0_751,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c1_2(X1,X2)|c3_2(X3,X4)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X5)|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_533]), ['final']).
cnf(c_0_752,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c1_2(X1,X2)|c3_2(a160,a248)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c5_1(X5)|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_540]), ['final']).
cnf(c_0_753,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c1_2(X1,X2)|c3_2(X3,X4)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X5)|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_542]), ['final']).
cnf(c_0_754,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c1_2(X1,X2)|c3_2(a160,a248)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c5_1(X5)|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_543]), ['final']).
cnf(c_0_755,negated_conjecture,(c2_2(a294,a295)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c5_1(X5)|c8_1(a160)|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_507]), ['final']).
cnf(c_0_756,negated_conjecture,(c2_2(a204,a205)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c5_1(X5)|c8_1(a160)|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_514]), ['final']).
cnf(c_0_757,negated_conjecture,(c2_2(a294,a295)|c1_2(X1,X2)|c8_2(a160,a161)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_544]), ['final']).
cnf(c_0_758,negated_conjecture,(c2_2(a204,a205)|c1_2(X1,X2)|c8_2(a160,a161)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_546]), ['final']).
cnf(c_0_759,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_545]), ['final']).
cnf(c_0_760,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_547]), ['final']).
cnf(c_0_761,negated_conjecture,(c2_2(a294,a295)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c2_1(X1)|ndr1_1(a160)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_549]), ['final']).
cnf(c_0_762,negated_conjecture,(c2_2(a204,a205)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c2_1(X1)|ndr1_1(a160)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_548]), ['final']).
cnf(c_0_763,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_649, c_0_81])).
cnf(c_0_764,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_650, c_0_81])).
cnf(c_0_765,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_496, c_0_651]), ['final']).
cnf(c_0_766,negated_conjecture,(ssSkC22|ndr1_1(a217)|c1_0), inference(spm,[status(thm)],[c_0_501, c_0_652]), ['final']).
cnf(c_0_767,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_653, c_0_81])).
cnf(c_0_768,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_654, c_0_81])).
cnf(c_0_769,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_496, c_0_655])).
cnf(c_0_770,negated_conjecture,(ssSkC20|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_656, c_0_603]), ['final']).
cnf(c_0_771,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c8_1(X2)|c8_1(X3)|c2_1(a160)|ndr1_1(X4)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)|~c1_1(X3)), inference(spm,[status(thm)],[c_0_657, c_0_24])).
cnf(c_0_772,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_658, c_0_81])).
cnf(c_0_773,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_659, c_0_81])).
cnf(c_0_774,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_497, c_0_651]), ['final']).
cnf(c_0_775,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c8_1(X2)|c8_1(X3)|c2_1(a160)|ndr1_1(X4)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)|~c1_1(X3)), inference(spm,[status(thm)],[c_0_660, c_0_30])).
cnf(c_0_776,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_661, c_0_81])).
cnf(c_0_777,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_662, c_0_81])).
cnf(c_0_778,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_497, c_0_655])).
cnf(c_0_779,negated_conjecture,(c5_0|c1_0|~c6_2(X1,X2)|~c1_2(X1,X3)|~c1_2(X1,X2)|~c4_2(X1,X3)|~c3_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_663, c_0_4]), ['final']).
cnf(c_0_780,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a160)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_498, c_0_651]), ['final']).
cnf(c_0_781,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a160)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_498, c_0_655])).
cnf(c_0_782,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_664, c_0_81])).
cnf(c_0_783,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_665, c_0_81])).
cnf(c_0_784,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_666, c_0_81])).
cnf(c_0_785,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_667, c_0_81])).
cnf(c_0_786,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_668, c_0_81])).
cnf(c_0_787,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_669, c_0_81])).
cnf(c_0_788,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_670, c_0_81])).
cnf(c_0_789,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_671, c_0_81])).
cnf(c_0_790,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_672, c_0_81])).
cnf(c_0_791,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_673, c_0_81])).
cnf(c_0_792,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_674, c_0_81])).
cnf(c_0_793,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)), inference(spm,[status(thm)],[c_0_675, c_0_81])).
cnf(c_0_794,negated_conjecture,(c6_2(a231,a232)|c1_2(a197,a198)|c3_2(X1,X2)|c2_1(a197)|c5_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a197,a198)|~c1_1(a197)|~c5_1(X1)|~ndr1_1(a197)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_550, c_0_82]), ['final']).
cnf(c_0_795,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(a160,a248)|c7_2(X1,a248)|c1_1(a294)|c5_1(a294)|c5_1(X2)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_676]), ['final']).
cnf(c_0_796,negated_conjecture,(c6_2(a231,a232)|c5_2(a265,a278)|c1_2(a265,a266)|c4_2(a265,a293)|c3_2(X1,X2)|c7_2(a265,a292)|c8_0|c7_0|c2_1(a265)|ssSkC23|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a265,a266)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_677, c_0_65])).
cnf(c_0_797,negated_conjecture,(c7_2(a265,a266)|ssSkC23), inference(spm,[status(thm)],[c_0_678, c_0_56]), ['final']).
cnf(c_0_798,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(X1,a248)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X2)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_679]), ['final']).
cnf(c_0_799,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a160,a248)|c7_2(X1,a248)|c1_1(a294)|c5_1(a294)|c5_1(X2)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_680]), ['final']).
cnf(c_0_800,negated_conjecture,(c6_2(a231,a232)|c5_2(a265,a278)|c1_2(a265,a266)|c8_2(a265,a292)|c4_2(a265,a293)|c3_2(X1,X2)|c8_0|c7_0|c2_1(a265)|ssSkC23|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a265,a266)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_681, c_0_65])).
cnf(c_0_801,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a160,a248)|c3_2(X1,a248)|c1_1(a294)|c5_1(a294)|c5_1(X2)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_682]), ['final']).
cnf(c_0_802,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(a160,a248)|c7_2(X1,a248)|c1_1(a204)|c5_1(a204)|c5_1(X2)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_683]), ['final']).
cnf(c_0_803,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(X1,a248)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X2)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_684]), ['final']).
cnf(c_0_804,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a160,a248)|c7_2(X1,a248)|c1_1(a204)|c5_1(a204)|c5_1(X2)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_685]), ['final']).
cnf(c_0_805,negated_conjecture,(c6_2(a231,a232)|c1_2(a241,a243)|c3_2(X1,X2)|c2_1(a241)|ssSkC19|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a241,a243)|~c1_1(a241)|~c5_1(X1)|~ndr1_1(a241)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_550, c_0_83]), ['final']).
cnf(c_0_806,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a160,a248)|c3_2(X1,a248)|c1_1(a204)|c5_1(a204)|c5_1(X2)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_686]), ['final']).
cnf(c_0_807,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_533]), ['final']).
cnf(c_0_808,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_533]), ['final']).
cnf(c_0_809,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_533]), ['final']).
cnf(c_0_810,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_533]), ['final']).
cnf(c_0_811,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_540]), ['final']).
cnf(c_0_812,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_540]), ['final']).
cnf(c_0_813,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_540]), ['final']).
cnf(c_0_814,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_540]), ['final']).
cnf(c_0_815,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_542]), ['final']).
cnf(c_0_816,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_542]), ['final']).
cnf(c_0_817,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_542]), ['final']).
cnf(c_0_818,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_542]), ['final']).
cnf(c_0_819,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_543]), ['final']).
cnf(c_0_820,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_543]), ['final']).
cnf(c_0_821,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_543]), ['final']).
cnf(c_0_822,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_543]), ['final']).
cnf(c_0_823,negated_conjecture,(c2_2(a294,a295)|c5_2(a160,a278)|c4_2(a160,a293)|c7_2(a160,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_687, c_0_502]), ['final']).
cnf(c_0_824,negated_conjecture,(c2_2(a294,a295)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_687, c_0_562]), ['final']).
cnf(c_0_825,negated_conjecture,(c2_2(a204,a205)|c5_2(a160,a278)|c4_2(a160,a293)|c7_2(a160,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_688, c_0_502]), ['final']).
cnf(c_0_826,negated_conjecture,(c2_2(a204,a205)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_688, c_0_562]), ['final']).
cnf(c_0_827,negated_conjecture,(c3_2(X1,a258)|c1_1(X1)|ssSkC22|~c7_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_476])]), ['final']).
cnf(c_0_828,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(X3,X4)|c2_1(X1)|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_476])]), ['final']).
cnf(c_0_829,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(X3,X4)|c2_1(X1)|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_476])]), ['final']).
cnf(c_0_830,negated_conjecture,(c1_2(X1,X2)|c3_2(X3,X4)|c2_1(X1)|ndr1_1(a231)|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_476])]), ['final']).
cnf(c_0_831,negated_conjecture,(c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a231)|c2_1(X1)|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_476])]), ['final']).
cnf(c_0_832,negated_conjecture,(c5_2(a275,a147)|c4_2(a275,a276)|c4_2(a275,X1)|c3_2(a275,X1)|c6_0|c8_0|ssSkC25|~c5_2(a275,a276)), inference(spm,[status(thm)],[c_0_689, c_0_54])).
cnf(c_0_833,negated_conjecture,(c5_2(X1,X2)|c3_2(X1,X2)|c3_1(X3)|c3_1(X1)|c4_1(X1)|c2_1(X3)|c1_0|~c6_2(X3,X4)|~c8_2(X3,X4)|~c8_2(X1,X2)|~c3_2(X3,X4)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_690, c_0_4]), ['final']).
cnf(c_0_834,negated_conjecture,(ssSkC20|c7_2(a168,a248)|c5_1(a287)|c3_0|ndr1_1(a287)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_564, c_0_92])).
cnf(c_0_835,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a273)|c4_2(a271,a273)|c7_2(a271,X1)|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)|~c1_1(a271)|~ndr1_1(a271)), inference(spm,[status(thm)],[c_0_691, c_0_93]), ['final']).
cnf(c_0_836,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a272)|c4_2(a271,a272)|c7_2(a271,X1)|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)|~c1_1(a271)|~ndr1_1(a271)), inference(spm,[status(thm)],[c_0_691, c_0_94]), ['final']).
cnf(c_0_837,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(X1,X2)|c1_1(a256)|c6_0|c8_0|c2_1(a174)|ssSkC4|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_692, c_0_693]), ['final']).
cnf(c_0_838,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(X1,X2)|c6_0|c8_0|c2_1(a174)|ndr1_1(a256)|ssSkC4|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_694, c_0_693]), ['final']).
cnf(c_0_839,negated_conjecture,(c7_2(a223,a292)|c2_1(a181)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_695, c_0_696]), ['final']).
cnf(c_0_840,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_697, c_0_559]), ['final']).
cnf(c_0_841,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_698, c_0_559]), ['final']).
cnf(c_0_842,negated_conjecture,(c6_2(a231,a232)|c1_2(a158,a293)|c3_2(X1,X2)|c7_2(a158,a292)|c2_1(a158)|c5_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a158,a293)|~c1_1(a158)|~c5_1(X1)|~ndr1_1(a158)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_550, c_0_699])).
cnf(c_0_843,negated_conjecture,(c2_2(a217,a278)|c2_2(a294,a295)|c1_2(a294,a295)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c4_1(a294)|c1_1(a217)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_700, c_0_24]), ['final']).
cnf(c_0_844,negated_conjecture,(c2_2(a217,a278)|c2_2(a204,a205)|c1_2(a204,a205)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c4_1(a204)|c1_1(a217)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_701, c_0_30]), ['final']).
cnf(c_0_845,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c7_2(X1,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_702]), ['final']).
cnf(c_0_846,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c3_2(X1,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_703]), ['final']).
cnf(c_0_847,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c7_2(X1,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_704]), ['final']).
cnf(c_0_848,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c3_2(X1,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_705]), ['final']).
cnf(c_0_849,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(X1,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_706]), ['final']).
cnf(c_0_850,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(X1,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_707]), ['final']).
cnf(c_0_851,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(X1,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_708]), ['final']).
cnf(c_0_852,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(X1,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_709]), ['final']).
cnf(c_0_853,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(X1,X2)|~c1_1(a193)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_476])]), ['final']).
cnf(c_0_854,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c7_2(a168,a248)|c5_1(a223)|c3_0|c6_0|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_564, c_0_644]), ['final']).
cnf(c_0_855,negated_conjecture,(ssSkC20|c3_2(a217,a248)|c2_0|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_46, c_0_710]), ['final']).
cnf(c_0_856,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c3_2(a168,a248)|c5_1(a223)|c3_0|c6_0|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_711, c_0_644]), ['final']).
cnf(c_0_857,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(X1,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_712]), ['final']).
cnf(c_0_858,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(X1,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_713]), ['final']).
cnf(c_0_859,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(X1,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_714]), ['final']).
cnf(c_0_860,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(X1,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_715]), ['final']).
cnf(c_0_861,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_716, c_0_641])).
cnf(c_0_862,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c5_1(a223)|c3_0|c6_0|ndr1_1(a168)|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_567, c_0_644]), ['final']).
cnf(c_0_863,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_717, c_0_630])).
cnf(c_0_864,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c8_0|c2_1(a223)|ssSkC15|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_718, c_0_641]), ['final']).
cnf(c_0_865,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c8_0|c2_1(a223)|ssSkC15|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_719, c_0_641]), ['final']).
cnf(c_0_866,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c2_1(a223)|ssSkC15|c5_0|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_720, c_0_630]), ['final']).
cnf(c_0_867,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c2_1(a223)|ssSkC15|c5_0|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_721, c_0_630]), ['final']).
cnf(c_0_868,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c6_0|c2_1(a158)|c5_0|c1_0|~c3_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_717, c_0_699])).
cnf(c_0_869,negated_conjecture,(ndr1_1(a158)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_71, c_0_556]), ['final']).
cnf(c_0_870,negated_conjecture,(c5_2(a174,a185)|c6_0|c8_0|ssSkC4|c1_0|~c5_1(a174)), inference(spm,[status(thm)],[c_0_31, c_0_722])).
cnf(c_0_871,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a158,a292)|c2_1(a158)|c5_0|c1_0|~c3_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_720, c_0_699]), ['final']).
cnf(c_0_872,negated_conjecture,(c5_2(a296,a297)|c3_1(a158)|c5_1(a158)|c6_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_71, c_0_594]), ['final']).
cnf(c_0_873,negated_conjecture,(c2_2(a158,a293)|c8_2(a159,a292)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a158,a292)|c2_1(a158)|c5_0|c1_0|~c3_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_721, c_0_699]), ['final']).
cnf(c_0_874,negated_conjecture,(c3_1(a158)|c5_1(a158)|c6_0|ndr1_1(a296)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_71, c_0_539]), ['final']).
cnf(c_0_875,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|c6_0|c7_0|ndr1_1(a208)|~c6_2(X1,X2)|~c1_2(X1,X3)|~c7_2(X1,X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_723, c_0_724]), ['final']).
cnf(c_0_876,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|c6_0|c8_1(a208)|c7_0|~c6_2(X1,X2)|~c1_2(X1,X3)|~c7_2(X1,X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_725, c_0_724]), ['final']).
cnf(c_0_877,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_687, c_0_559]), ['final']).
cnf(c_0_878,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_726, c_0_641]), ['final']).
cnf(c_0_879,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_727, c_0_630]), ['final']).
cnf(c_0_880,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|c5_0|c1_0|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_544]), ['final']).
cnf(c_0_881,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|c5_0|c1_0|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_544]), ['final']).
cnf(c_0_882,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_544]), ['final']).
cnf(c_0_883,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_544]), ['final']).
cnf(c_0_884,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_688, c_0_559]), ['final']).
cnf(c_0_885,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|c5_0|c1_0|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_546]), ['final']).
cnf(c_0_886,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|c5_0|c1_0|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_546]), ['final']).
cnf(c_0_887,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_546]), ['final']).
cnf(c_0_888,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_546]), ['final']).
cnf(c_0_889,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_0|c1_0|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_545]), ['final']).
cnf(c_0_890,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_0|c1_0|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_545]), ['final']).
cnf(c_0_891,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_545]), ['final']).
cnf(c_0_892,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_545]), ['final']).
cnf(c_0_893,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_0|c1_0|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_547]), ['final']).
cnf(c_0_894,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_0|c1_0|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_547]), ['final']).
cnf(c_0_895,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_547]), ['final']).
cnf(c_0_896,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_547]), ['final']).
cnf(c_0_897,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_716, c_0_83])).
cnf(c_0_898,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_716, c_0_728])).
cnf(c_0_899,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(a241)|ssSkC19|c5_0|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_717, c_0_83])).
cnf(c_0_900,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c3_1(a159)|c5_1(a159)|c6_0|ssSkC22|c2_1(a217)|c5_0|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_717, c_0_728])).
cnf(c_0_901,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c6_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0|~c3_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_727, c_0_699]), ['final']).
cnf(c_0_902,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|ndr1_1(a160)|c5_0|c1_0|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_549]), ['final']).
cnf(c_0_903,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|ndr1_1(a160)|c5_0|c1_0|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_549]), ['final']).
cnf(c_0_904,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|ndr1_1(a160)|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_549]), ['final']).
cnf(c_0_905,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|ndr1_1(a160)|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_549]), ['final']).
cnf(c_0_906,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|ndr1_1(a160)|c5_0|c1_0|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_548]), ['final']).
cnf(c_0_907,negated_conjecture,(c2_2(a197,a198)|c5_2(a296,a297)|c8_2(a197,a198)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(a197)|c5_0|c1_0|~c3_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_717, c_0_82])).
cnf(c_0_908,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|ndr1_1(a160)|c5_0|c1_0|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_548]), ['final']).
cnf(c_0_909,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|ndr1_1(a160)|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_548]), ['final']).
cnf(c_0_910,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|ndr1_1(a160)|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_548]), ['final']).
cnf(c_0_911,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c8_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_729, c_0_641]), ['final']).
cnf(c_0_912,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c2_1(a223)|ndr1_1(a159)|ssSkC15|c5_0|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_730, c_0_630]), ['final']).
cnf(c_0_913,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c8_0|c2_1(a159)|c2_1(a223)|ssSkC15|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_731, c_0_641]), ['final']).
cnf(c_0_914,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c2_1(a159)|c2_1(a223)|ssSkC15|c5_0|c1_0|~c3_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_732, c_0_630]), ['final']).
cnf(c_0_915,negated_conjecture,(c2_2(a223,a293)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|ssSkC15|c5_0|c1_0|~c3_1(a223)|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_733, c_0_696]), ['final']).
cnf(c_0_916,negated_conjecture,(c2_2(a223,a293)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|c8_0|ssSkC15|c1_0|~c3_1(a223)|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_733, c_0_734]), ['final']).
cnf(c_0_917,negated_conjecture,(c4_2(X1,X2)|c3_2(X1,X2)|c3_0|c8_0|~c5_2(X1,X2)|~c6_1(X1)|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_476])]), ['final']).
cnf(c_0_918,negated_conjecture,(c5_2(a200,a278)|c4_2(a200,a293)|c7_2(a200,a292)|c8_0|c7_0|ssSkC10|c1_0), inference(spm,[status(thm)],[c_0_97, c_0_502]), ['final']).
cnf(c_0_919,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c7_2(a158,a292)|c2_1(a158)|ndr1_1(a159)|c5_0|c1_0|~c3_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_730, c_0_699]), ['final']).
cnf(c_0_920,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c7_2(a158,a292)|c2_1(a159)|c2_1(a158)|c5_0|c1_0|~c3_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_732, c_0_699]), ['final']).
cnf(c_0_921,negated_conjecture,(c2_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|ndr1_1(a193)|ssSkC15|c5_0|c1_0|~c3_1(a223)|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_735, c_0_696]), ['final']).
cnf(c_0_922,negated_conjecture,(c2_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c8_0|ndr1_1(a193)|ssSkC15|c1_0|~c3_1(a223)|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_735, c_0_734]), ['final']).
cnf(c_0_923,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_736]), ['final']).
cnf(c_0_924,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_737]), ['final']).
cnf(c_0_925,negated_conjecture,(ssSkC7|c5_2(a296,a297)|c3_1(a202)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(X1)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_738, c_0_739]), ['final']).
cnf(c_0_926,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_740]), ['final']).
cnf(c_0_927,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_741]), ['final']).
cnf(c_0_928,negated_conjecture,(ssSkC7|c5_2(a296,a297)|c3_1(a217)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_742, c_0_739]), ['final']).
cnf(c_0_929,negated_conjecture,(ssSkC7|c3_1(a202)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_738, c_0_743]), ['final']).
cnf(c_0_930,negated_conjecture,(c5_2(a268,a270)|c6_1(a268)|c2_0|c4_0|c5_0|~c5_2(a268,X1)|~c8_2(a268,a270)|~c7_2(a268,X1)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_744, c_0_98])).
cnf(c_0_931,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_745]), ['final']).
cnf(c_0_932,negated_conjecture,(ssSkC7|c1_2(a202,a190)|c3_1(a202)|c4_1(X1)|c2_0|ssSkC11|~c5_1(X1)), inference(spm,[status(thm)],[c_0_591, c_0_746]), ['final']).
cnf(c_0_933,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_747]), ['final']).
cnf(c_0_934,negated_conjecture,(ssSkC7|c3_1(a217)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_742, c_0_743]), ['final']).
cnf(c_0_935,negated_conjecture,(ssSkC7|c1_2(a217,a190)|c3_1(a217)|c4_1(X1)|c2_0|c1_0|~c5_1(X1)), inference(spm,[status(thm)],[c_0_535, c_0_746]), ['final']).
cnf(c_0_936,negated_conjecture,(ssSkC20|c3_2(a168,a248)|c5_1(a223)|c3_0|c6_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_711, c_0_565]), ['final']).
cnf(c_0_937,negated_conjecture,(c8_2(X1,X2)|c4_1(X1)|c4_0|~c6_2(X1,X2)|~c7_2(X1,X2)|~c3_0|~c8_1(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_476])]), ['final']).
cnf(c_0_938,negated_conjecture,(c8_2(a275,a207)|c3_0|ssSkC25|c4_0|~c4_2(a275,a276)|~ssSkP0(a275)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_647, c_0_748]), ['final']).
cnf(c_0_939,negated_conjecture,(c6_2(a241,a206)|c8_2(a241,a207)|c3_0|ssSkC19|c4_0|~c4_2(a241,a243)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_749, c_0_100])).
cnf(c_0_940,negated_conjecture,(c8_2(a241,a207)|c4_2(a241,a206)|c3_0|ssSkC19|c4_0|~c4_2(a241,a243)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_749, c_0_101])).
cnf(c_0_941,negated_conjecture,(c1_1(X1)|c8_0|c2_1(X2)|ndr1_1(X1)|~c7_1(X1)|~c7_1(X2)|~c5_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_476])]), ['final']).
cnf(c_0_942,negated_conjecture,(c6_1(X1)|c5_1(X2)|c8_1(X1)|c2_1(X1)|~c8_2(X2,a151)|~c3_1(X2)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_476])]), ['final']).
cnf(c_0_943,negated_conjecture,(c8_2(a241,a207)|c3_2(a241,a206)|c3_0|ssSkC19|c4_0|~c4_2(a241,a243)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_749, c_0_104])).
cnf(c_0_944,negated_conjecture,(c8_2(a268,a207)|c3_0|c4_0|c5_0|~c4_2(a268,a269)|~ssSkP0(a268)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_647, c_0_105]), ['final']).
cnf(c_0_945,negated_conjecture,(c8_2(a271,a272)|c3_2(a271,a272)|c8_0|c2_1(a271)|ssSkC24|~c5_1(a271)|~ndr1_1(a271)), inference(spm,[status(thm)],[c_0_750, c_0_106])).
cnf(c_0_946,negated_conjecture,(c4_1(X1)|~c3_2(X1,a293)|~ssSkP10(X1)|~c2_0|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_476])]), ['final']).
cnf(c_0_947,negated_conjecture,(c4_1(X1)|~c1_2(X1,a293)|~ssSkP10(X1)|~c2_0|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_476])]), ['final']).
cnf(c_0_948,negated_conjecture,(c1_1(a252)|~c8_1(X1)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_476])]), ['final']).
cnf(c_0_949,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a160,a248)|c7_2(a231,a292)|c4_1(a294)|c5_1(a294)|c5_1(X5)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_751, c_0_502]), ['final']).
cnf(c_0_950,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X5)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_751, c_0_562]), ['final']).
cnf(c_0_951,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a160,a248)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a294)|c5_1(a294)|c5_1(X5)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_752, c_0_502]), ['final']).
cnf(c_0_952,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a160,a248)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c5_1(X5)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_752, c_0_562]), ['final']).
cnf(c_0_953,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a160,a248)|c7_2(a231,a292)|c4_1(a204)|c5_1(a204)|c5_1(X5)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_753, c_0_502]), ['final']).
cnf(c_0_954,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X5)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_753, c_0_562]), ['final']).
cnf(c_0_955,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a160,a248)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a204)|c5_1(a204)|c5_1(X5)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_754, c_0_502]), ['final']).
cnf(c_0_956,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a160,a248)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c5_1(X5)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_754, c_0_562]), ['final']).
cnf(c_0_957,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a294)|c5_1(a294)|c5_1(X5)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_755, c_0_502]), ['final']).
cnf(c_0_958,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c5_1(X5)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_755, c_0_562]), ['final']).
cnf(c_0_959,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a204)|c5_1(a204)|c5_1(X5)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_756, c_0_502]), ['final']).
cnf(c_0_960,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c5_1(X5)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_756, c_0_562]), ['final']).
cnf(c_0_961,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a160,a161)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_757, c_0_502]), ['final']).
cnf(c_0_962,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a160,a161)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_757, c_0_562]), ['final']).
cnf(c_0_963,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a160,a161)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_758, c_0_502]), ['final']).
cnf(c_0_964,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a160,a161)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_758, c_0_562]), ['final']).
cnf(c_0_965,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_759, c_0_502]), ['final']).
cnf(c_0_966,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_759, c_0_562]), ['final']).
cnf(c_0_967,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_760, c_0_502]), ['final']).
cnf(c_0_968,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_760, c_0_562]), ['final']).
cnf(c_0_969,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_761, c_0_502]), ['final']).
cnf(c_0_970,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_761, c_0_562]), ['final']).
cnf(c_0_971,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(X3,X4)|c7_2(a231,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_762, c_0_502]), ['final']).
cnf(c_0_972,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(X3,X4)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_762, c_0_562]), ['final']).
cnf(c_0_973,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_763, c_0_110]), ['final']).
cnf(c_0_974,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_764, c_0_110]), ['final']).
cnf(c_0_975,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ssSkC22|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_765, c_0_766]), ['final']).
cnf(c_0_976,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_767, c_0_110]), ['final']).
cnf(c_0_977,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_768, c_0_110]), ['final']).
cnf(c_0_978,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_769, c_0_770]), ['final']).
cnf(c_0_979,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c8_1(X2)|c2_1(a160)|ndr1_1(X3)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(ef,[status(thm)],[c_0_771]), ['final']).
cnf(c_0_980,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_772, c_0_110]), ['final']).
cnf(c_0_981,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_773, c_0_110]), ['final']).
cnf(c_0_982,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ssSkC22|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_774, c_0_766]), ['final']).
cnf(c_0_983,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c8_1(X2)|c2_1(a160)|ndr1_1(X3)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(ef,[status(thm)],[c_0_775]), ['final']).
cnf(c_0_984,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_776, c_0_110]), ['final']).
cnf(c_0_985,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_777, c_0_110]), ['final']).
cnf(c_0_986,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_778, c_0_770]), ['final']).
cnf(c_0_987,negated_conjecture,(c7_2(a223,a292)|c8_0|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c8_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_505, c_0_734]), ['final']).
cnf(c_0_988,negated_conjecture,(c7_2(a223,a292)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c8_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_779, c_0_696]), ['final']).
cnf(c_0_989,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ssSkC22|ndr1_1(a160)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_780, c_0_766]), ['final']).
cnf(c_0_990,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a160)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_781, c_0_770]), ['final']).
cnf(c_0_991,negated_conjecture,(c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c8_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_779, c_0_39]), ['final']).
cnf(c_0_992,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_782, c_0_110]), ['final']).
cnf(c_0_993,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_783, c_0_110]), ['final']).
cnf(c_0_994,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_784, c_0_110]), ['final']).
cnf(c_0_995,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_785, c_0_110]), ['final']).
cnf(c_0_996,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_786, c_0_110]), ['final']).
cnf(c_0_997,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_787, c_0_110]), ['final']).
cnf(c_0_998,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_788, c_0_110]), ['final']).
cnf(c_0_999,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_789, c_0_110]), ['final']).
cnf(c_0_1000,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_790, c_0_110]), ['final']).
cnf(c_0_1001,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_791, c_0_110]), ['final']).
cnf(c_0_1002,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_792, c_0_110]), ['final']).
cnf(c_0_1003,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c1_2(a233,X1)|~c4_2(a233,X1)), inference(spm,[status(thm)],[c_0_793, c_0_110]), ['final']).
cnf(c_0_1004,negated_conjecture,(c6_2(a231,a232)|c5_2(a197,a278)|c1_2(a197,a198)|c4_2(a197,a293)|c3_2(X1,X2)|c7_2(a197,a292)|c8_0|c7_0|c2_1(a197)|c5_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a197,a198)|~c5_1(X1)|~ndr1_1(a197)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_794, c_0_502])).
cnf(c_0_1005,negated_conjecture,(c6_2(a231,a232)|c5_2(a197,a278)|c1_2(a197,a198)|c8_2(a197,a292)|c4_2(a197,a293)|c3_2(X1,X2)|c8_0|c7_0|c2_1(a197)|c5_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a197,a198)|~c5_1(X1)|~ndr1_1(a197)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_794, c_0_562])).
cnf(c_0_1006,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(a160,a248)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_795]), ['final']).
cnf(c_0_1007,negated_conjecture,(c6_2(a231,a232)|c5_2(a265,a278)|c1_2(a265,a266)|c4_2(a265,a293)|c3_2(X1,X2)|c7_2(a265,a292)|c8_0|c7_0|c2_1(a265)|ssSkC23|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_796, c_0_797]), ['final']).
cnf(c_0_1008,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a215,a248)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_798]), ['final']).
cnf(c_0_1009,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a160,a248)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_799]), ['final']).
cnf(c_0_1010,negated_conjecture,(c6_2(a231,a232)|c5_2(a265,a278)|c1_2(a265,a266)|c8_2(a265,a292)|c4_2(a265,a293)|c3_2(X1,X2)|c8_0|c7_0|c2_1(a265)|ssSkC23|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_800, c_0_797]), ['final']).
cnf(c_0_1011,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a160,a248)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_801]), ['final']).
cnf(c_0_1012,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(a160,a248)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_802]), ['final']).
cnf(c_0_1013,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a215,a248)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_803]), ['final']).
cnf(c_0_1014,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a160,a248)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_804]), ['final']).
cnf(c_0_1015,negated_conjecture,(c6_2(a231,a232)|c5_2(a241,a278)|c1_2(a241,a243)|c4_2(a241,a293)|c3_2(X1,X2)|c7_2(a241,a292)|c8_0|c7_0|c2_1(a241)|ssSkC19|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a241,a243)|~c5_1(X1)|~ndr1_1(a241)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_805, c_0_502])).
cnf(c_0_1016,negated_conjecture,(c6_2(a231,a232)|c5_2(a241,a278)|c1_2(a241,a243)|c8_2(a241,a292)|c4_2(a241,a293)|c3_2(X1,X2)|c8_0|c7_0|c2_1(a241)|ssSkC19|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a241,a243)|~c5_1(X1)|~ndr1_1(a241)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_805, c_0_562])).
cnf(c_0_1017,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a160,a248)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_806]), ['final']).
cnf(c_0_1018,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a160,a248)|c7_2(a197,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_807, c_0_502]), ['final']).
cnf(c_0_1019,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_807, c_0_562]), ['final']).
cnf(c_0_1020,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a160,a248)|c7_2(a162,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_808, c_0_502]), ['final']).
cnf(c_0_1021,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_808, c_0_562]), ['final']).
cnf(c_0_1022,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a160,a248)|c7_2(a288,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_809, c_0_502]), ['final']).
cnf(c_0_1023,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_809, c_0_562]), ['final']).
cnf(c_0_1024,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a160,a248)|c7_2(a253,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_810, c_0_502]), ['final']).
cnf(c_0_1025,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_810, c_0_562]), ['final']).
cnf(c_0_1026,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a197,a278)|c4_2(a197,a293)|c3_2(a160,a248)|c7_2(a197,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_811, c_0_502]), ['final']).
cnf(c_0_1027,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_811, c_0_562]), ['final']).
cnf(c_0_1028,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a162,a278)|c4_2(a162,a293)|c3_2(a160,a248)|c7_2(a162,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_812, c_0_502]), ['final']).
cnf(c_0_1029,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_812, c_0_562]), ['final']).
cnf(c_0_1030,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a288,a278)|c4_2(a288,a293)|c3_2(a160,a248)|c7_2(a288,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_813, c_0_502]), ['final']).
cnf(c_0_1031,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_813, c_0_562]), ['final']).
cnf(c_0_1032,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a253,a278)|c4_2(a253,a293)|c3_2(a160,a248)|c7_2(a253,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_814, c_0_502]), ['final']).
cnf(c_0_1033,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_814, c_0_562]), ['final']).
cnf(c_0_1034,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a160,a248)|c7_2(a197,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_815, c_0_502]), ['final']).
cnf(c_0_1035,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_815, c_0_562]), ['final']).
cnf(c_0_1036,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a160,a248)|c7_2(a162,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_816, c_0_502]), ['final']).
cnf(c_0_1037,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_816, c_0_562]), ['final']).
cnf(c_0_1038,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a160,a248)|c7_2(a288,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_817, c_0_502]), ['final']).
cnf(c_0_1039,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_817, c_0_562]), ['final']).
cnf(c_0_1040,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a160,a248)|c7_2(a253,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_818, c_0_502]), ['final']).
cnf(c_0_1041,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_818, c_0_562]), ['final']).
cnf(c_0_1042,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a197,a278)|c4_2(a197,a293)|c3_2(a160,a248)|c7_2(a197,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_819, c_0_502]), ['final']).
cnf(c_0_1043,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_819, c_0_562]), ['final']).
cnf(c_0_1044,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a162,a278)|c4_2(a162,a293)|c3_2(a160,a248)|c7_2(a162,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_820, c_0_502]), ['final']).
cnf(c_0_1045,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_820, c_0_562]), ['final']).
cnf(c_0_1046,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a288,a278)|c4_2(a288,a293)|c3_2(a160,a248)|c7_2(a288,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_821, c_0_502]), ['final']).
cnf(c_0_1047,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_821, c_0_562]), ['final']).
cnf(c_0_1048,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a253,a278)|c4_2(a253,a293)|c3_2(a160,a248)|c7_2(a253,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_822, c_0_502]), ['final']).
cnf(c_0_1049,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_822, c_0_562]), ['final']).
cnf(c_0_1050,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a160,a278)|c4_2(a160,a293)|c7_2(a160,a248)|c7_2(a160,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_823]), ['final']).
cnf(c_0_1051,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a160,a278)|c4_2(a160,a293)|c3_2(a160,a248)|c7_2(a160,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_823]), ['final']).
cnf(c_0_1052,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_824]), ['final']).
cnf(c_0_1053,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_824]), ['final']).
cnf(c_0_1054,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a160,a278)|c4_2(a160,a293)|c7_2(a160,a248)|c7_2(a160,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_825]), ['final']).
cnf(c_0_1055,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a160,a278)|c4_2(a160,a293)|c3_2(a160,a248)|c7_2(a160,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_825]), ['final']).
cnf(c_0_1056,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_826]), ['final']).
cnf(c_0_1057,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_826]), ['final']).
cnf(c_0_1058,negated_conjecture,(c4_2(a215,a258)|c1_1(a215)|c8_0|ssSkC22|c7_0), inference(spm,[status(thm)],[c_0_589, c_0_112]), ['final']).
cnf(c_0_1059,negated_conjecture,(c3_2(a215,a258)|c1_1(a215)|c8_0|ssSkC22|c7_0), inference(spm,[status(thm)],[c_0_827, c_0_112]), ['final']).
cnf(c_0_1060,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a223,a293)|c7_2(a223,a292)|c2_1(X1)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_828, c_0_696]), ['final']).
cnf(c_0_1061,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a223,a293)|c7_2(a223,a292)|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_828, c_0_734]), ['final']).
cnf(c_0_1062,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c2_1(X1)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_829, c_0_696]), ['final']).
cnf(c_0_1063,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_829, c_0_734]), ['final']).
cnf(c_0_1064,negated_conjecture,(ssSkC20|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_656, c_0_599]), ['final']).
cnf(c_0_1065,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c2_1(X1)|ndr1_1(a231)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_830, c_0_696]), ['final']).
cnf(c_0_1066,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c8_0|c2_1(X1)|ndr1_1(a231)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_830, c_0_734]), ['final']).
cnf(c_0_1067,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a231)|c2_1(X1)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_831, c_0_696]), ['final']).
cnf(c_0_1068,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a231)|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_831, c_0_734]), ['final']).
cnf(c_0_1069,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_513, c_0_533]), ['final']).
cnf(c_0_1070,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_513, c_0_540]), ['final']).
cnf(c_0_1071,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_513, c_0_542]), ['final']).
cnf(c_0_1072,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_513, c_0_543]), ['final']).
cnf(c_0_1073,negated_conjecture,(c5_2(a275,a147)|c4_2(a275,a276)|c4_2(a275,X1)|c3_2(a275,X1)|c6_0|c8_0|ssSkC25), inference(spm,[status(thm)],[c_0_832, c_0_748]), ['final']).
cnf(c_0_1074,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_507]), ['final']).
cnf(c_0_1075,negated_conjecture,(c5_2(X1,X2)|c3_2(X1,X2)|c3_1(a223)|c3_1(X1)|c7_2(a223,a292)|c4_1(X1)|c2_1(a223)|ssSkC15|c5_0|c1_0|~c8_2(a223,a293)|~c8_2(X1,X2)|~c3_2(a223,a293)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_833, c_0_696])).
cnf(c_0_1076,negated_conjecture,(c5_2(X1,X2)|c3_2(X1,X2)|c3_1(a223)|c3_1(X1)|c7_2(a223,a292)|c4_1(X1)|c8_0|c2_1(a223)|ssSkC15|c1_0|~c8_2(a223,a293)|~c8_2(X1,X2)|~c3_2(a223,a293)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_833, c_0_734])).
cnf(c_0_1077,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_514]), ['final']).
cnf(c_0_1078,negated_conjecture,(ssSkC20|c7_2(a168,a248)|c3_0|ndr1_1(a287)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_113, c_0_834]), ['final']).
cnf(c_0_1079,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a273)|c5_2(a271,a278)|c4_2(a271,a273)|c4_2(a271,a293)|c7_2(a271,a292)|c7_2(a271,X1)|c8_0|c7_0|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)|~ndr1_1(a271)), inference(spm,[status(thm)],[c_0_835, c_0_502])).
cnf(c_0_1080,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a273)|c5_2(a271,a278)|c8_2(a271,a292)|c4_2(a271,a273)|c4_2(a271,a293)|c7_2(a271,X1)|c8_0|c7_0|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)|~ndr1_1(a271)), inference(spm,[status(thm)],[c_0_835, c_0_562])).
cnf(c_0_1081,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a272)|c5_2(a271,a278)|c4_2(a271,a272)|c4_2(a271,a293)|c7_2(a271,a292)|c7_2(a271,X1)|c8_0|c7_0|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)|~ndr1_1(a271)), inference(spm,[status(thm)],[c_0_836, c_0_502])).
cnf(c_0_1082,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a272)|c5_2(a271,a278)|c8_2(a271,a292)|c4_2(a271,a272)|c4_2(a271,a293)|c7_2(a271,X1)|c8_0|c7_0|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)|~ndr1_1(a271)), inference(spm,[status(thm)],[c_0_836, c_0_562])).
cnf(c_0_1083,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c6_0|c8_0|c2_1(a174)|ssSkC4|ssSkC15|c1_0|~c3_1(a223)|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_837, c_0_734]), ['final']).
cnf(c_0_1084,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a174)|ndr1_1(a256)|ssSkC4|ssSkC15|c1_0|~c3_1(a223)|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_838, c_0_734]), ['final']).
cnf(c_0_1085,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a268,a270)|c2_1(X1)|c4_0|c5_0|~c4_2(X1,X2)|~c3_1(a268)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_828, c_0_98]), ['final']).
cnf(c_0_1086,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a268,a270)|c2_1(X1)|c4_0|c5_0|~c4_2(X1,X2)|~c3_1(a268)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_829, c_0_98]), ['final']).
cnf(c_0_1087,negated_conjecture,(c5_2(a296,a297)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a181)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_839, c_0_644])).
cnf(c_0_1088,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_840]), ['final']).
cnf(c_0_1089,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_840]), ['final']).
cnf(c_0_1090,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_841]), ['final']).
cnf(c_0_1091,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_841]), ['final']).
cnf(c_0_1092,negated_conjecture,(c6_2(a231,a232)|c1_2(a158,a293)|c3_2(X1,X2)|c7_2(a158,a292)|c2_1(a158)|c5_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a158,a293)|~c5_1(X1)|~ndr1_1(a158)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_842, c_0_114])).
cnf(c_0_1093,negated_conjecture,(c2_2(a217,a278)|c2_2(a294,a295)|c5_2(a294,a278)|c1_2(a294,a295)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c1_1(a294)|c1_1(a217)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_843])).
cnf(c_0_1094,negated_conjecture,(c2_2(a217,a278)|c2_2(a204,a205)|c5_2(a204,a278)|c1_2(a204,a205)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c1_1(a204)|c1_1(a217)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_844])).
cnf(c_0_1095,negated_conjecture,(c1_2(X1,X2)|c3_2(a268,a270)|c2_1(X1)|ndr1_1(a231)|c4_0|c5_0|~c4_2(X1,X2)|~c3_1(a268)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_830, c_0_98]), ['final']).
cnf(c_0_1096,negated_conjecture,(c1_2(X1,X2)|c3_2(a268,a270)|c4_1(a231)|c2_1(X1)|c4_0|c5_0|~c4_2(X1,X2)|~c3_1(a268)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_831, c_0_98]), ['final']).
cnf(c_0_1097,negated_conjecture,(c5_2(a275,a278)|c4_2(a275,a293)|c7_2(a275,a292)|c8_0|c7_0|ssSkC25|c1_0), inference(spm,[status(thm)],[c_0_115, c_0_502]), ['final']).
cnf(c_0_1098,negated_conjecture,(c5_2(a218,a278)|c4_2(a218,a293)|c7_2(a218,a292)|c8_0|c7_0|ssSkC13|c1_0), inference(spm,[status(thm)],[c_0_116, c_0_502]), ['final']).
cnf(c_0_1099,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c7_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_845]), ['final']).
cnf(c_0_1100,negated_conjecture,(c5_2(a182,a278)|c4_2(a182,a293)|c7_2(a182,a292)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_117, c_0_502]), ['final']).
cnf(c_0_1101,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c3_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_846]), ['final']).
cnf(c_0_1102,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c7_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_847]), ['final']).
cnf(c_0_1103,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c3_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_848]), ['final']).
cnf(c_0_1104,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_849]), ['final']).
cnf(c_0_1105,negated_conjecture,(c5_2(a221,a278)|c4_2(a221,a293)|c7_2(a221,a292)|c8_0|c7_0|ssSkC14|c1_0), inference(spm,[status(thm)],[c_0_118, c_0_502]), ['final']).
cnf(c_0_1106,negated_conjecture,(c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a181)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_839, c_0_565])).
cnf(c_0_1107,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_850]), ['final']).
cnf(c_0_1108,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_851]), ['final']).
cnf(c_0_1109,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_852]), ['final']).
cnf(c_0_1110,negated_conjecture,(c2_2(X1,X2)|c5_2(a193,a278)|c4_2(a193,a293)|c7_2(a193,a292)|c5_1(X1)|c8_0|c7_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_853, c_0_502]), ['final']).
cnf(c_0_1111,negated_conjecture,(c2_2(X1,X2)|c5_2(a193,a278)|c8_2(a193,a292)|c4_2(a193,a293)|c5_1(X1)|c8_0|c7_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_853, c_0_562]), ['final']).
cnf(c_0_1112,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c7_2(a217,a248)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_598, c_0_854]), ['final']).
cnf(c_0_1113,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_656, c_0_854]), ['final']).
cnf(c_0_1114,negated_conjecture,(c2_1(a181)|c5_0|ssSkC17|~c5_2(a233,a234)|~c1_2(a233,X1)|~c8_2(a233,X1)|~c3_2(a233,X1)|~c3_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_695, c_0_39]), ['final']).
cnf(c_0_1115,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c3_2(a217,a248)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_855, c_0_854]), ['final']).
cnf(c_0_1116,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c3_2(a168,a248)|c7_2(a217,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_598, c_0_856]), ['final']).
cnf(c_0_1117,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_845]), ['final']).
cnf(c_0_1118,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_846]), ['final']).
cnf(c_0_1119,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_847]), ['final']).
cnf(c_0_1120,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_848]), ['final']).
cnf(c_0_1121,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_849]), ['final']).
cnf(c_0_1122,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_850]), ['final']).
cnf(c_0_1123,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_851]), ['final']).
cnf(c_0_1124,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c3_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_656, c_0_856]), ['final']).
cnf(c_0_1125,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_852]), ['final']).
cnf(c_0_1126,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_857]), ['final']).
cnf(c_0_1127,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_858]), ['final']).
cnf(c_0_1128,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_859]), ['final']).
cnf(c_0_1129,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_860]), ['final']).
cnf(c_0_1130,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c3_1(a223)), inference(spm,[status(thm)],[c_0_861, c_0_590]), ['final']).
cnf(c_0_1131,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c3_2(a217,a248)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_855, c_0_862]), ['final']).
cnf(c_0_1132,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0|~c3_1(a223)), inference(spm,[status(thm)],[c_0_863, c_0_584]), ['final']).
cnf(c_0_1133,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c7_2(a217,a248)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_598, c_0_862]), ['final']).
cnf(c_0_1134,negated_conjecture,(c5_2(X1,X2)|c3_2(X1,X2)|c3_1(a233)|c3_1(X1)|c4_1(X1)|c2_1(a233)|ssSkC17|c1_0|~c8_2(a233,a234)|~c8_2(X1,X2)|~c3_2(a233,a234)|~ndr1_1(a233)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_833, c_0_39])).
cnf(c_0_1135,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_864, c_0_644])).
cnf(c_0_1136,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_865, c_0_644])).
cnf(c_0_1137,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_866, c_0_644])).
cnf(c_0_1138,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_867, c_0_644])).
cnf(c_0_1139,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c3_2(a217,a248)|c3_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_855, c_0_856]), ['final']).
cnf(c_0_1140,negated_conjecture,(c2_1(a181)|c4_0|c5_0|~c5_2(a268,a270)|~c1_2(a268,X1)|~c8_2(a268,X1)|~c3_2(a268,X1)|~c3_1(a268)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_695, c_0_98]), ['final']).
cnf(c_0_1141,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a168)|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_656, c_0_862]), ['final']).
cnf(c_0_1142,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_857]), ['final']).
cnf(c_0_1143,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_858]), ['final']).
cnf(c_0_1144,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_859]), ['final']).
cnf(c_0_1145,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c1_0|~c1_1(a215)), inference(spm,[status(thm)],[c_0_111, c_0_860]), ['final']).
cnf(c_0_1146,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c6_0|c2_1(a158)|c5_0|c1_0|~c3_1(a158)), inference(spm,[status(thm)],[c_0_868, c_0_869]), ['final']).
cnf(c_0_1147,negated_conjecture,(c5_2(a174,a185)|c6_0|c8_0|ssSkC4|c1_0), inference(spm,[status(thm)],[c_0_870, c_0_44]), ['final']).
cnf(c_0_1148,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_864, c_0_565])).
cnf(c_0_1149,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_865, c_0_565])).
cnf(c_0_1150,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_866, c_0_565])).
cnf(c_0_1151,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_867, c_0_565])).
cnf(c_0_1152,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|c5_0|c1_0|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_871, c_0_872])).
cnf(c_0_1153,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a159,a292)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|c5_0|c1_0|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_873, c_0_872])).
cnf(c_0_1154,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ndr1_1(a296)|ssSkC19|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_726, c_0_83]), ['final']).
cnf(c_0_1155,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a296)|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_726, c_0_728]), ['final']).
cnf(c_0_1156,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(a241)|ndr1_1(a296)|ssSkC19|c5_0|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_727, c_0_83]), ['final']).
cnf(c_0_1157,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c3_1(a159)|c5_1(a159)|c6_0|ssSkC22|c2_1(a217)|ndr1_1(a296)|c5_0|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_727, c_0_728]), ['final']).
cnf(c_0_1158,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a159,a293)|c7_2(a159,a292)|c8_0|c2_1(a241)|ssSkC19|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_718, c_0_83]), ['final']).
cnf(c_0_1159,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a159,a293)|c7_2(a159,a292)|c8_0|ssSkC22|c2_1(a217)|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_718, c_0_728]), ['final']).
cnf(c_0_1160,negated_conjecture,(c2_2(a241,a243)|c8_2(a159,a292)|c8_2(a241,a243)|c4_2(a159,a293)|c8_0|c2_1(a241)|ssSkC19|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_719, c_0_83]), ['final']).
cnf(c_0_1161,negated_conjecture,(c2_2(a217,a258)|c8_2(a159,a292)|c8_2(a217,a258)|c4_2(a159,a293)|c8_0|ssSkC22|c2_1(a217)|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_719, c_0_728]), ['final']).
cnf(c_0_1162,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a159,a293)|c7_2(a159,a292)|c2_1(a241)|ssSkC19|c5_0|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_720, c_0_83]), ['final']).
cnf(c_0_1163,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a159,a293)|c7_2(a159,a292)|ssSkC22|c2_1(a217)|c5_0|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_720, c_0_728]), ['final']).
cnf(c_0_1164,negated_conjecture,(c2_2(a241,a243)|c8_2(a159,a292)|c8_2(a241,a243)|c4_2(a159,a293)|c2_1(a241)|ssSkC19|c5_0|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_721, c_0_83]), ['final']).
cnf(c_0_1165,negated_conjecture,(c2_2(a217,a258)|c8_2(a159,a292)|c8_2(a217,a258)|c4_2(a159,a293)|ssSkC22|c2_1(a217)|c5_0|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_721, c_0_728]), ['final']).
cnf(c_0_1166,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_871, c_0_874])).
cnf(c_0_1167,negated_conjecture,(c4_2(a268,a270)|c4_2(a268,X1)|c7_2(a268,a270)|c7_1(a268)|c6_0|c7_0|ndr1_1(a208)|c4_0|c5_0|~c1_2(a268,X1)|~c7_2(a268,X1)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_875, c_0_98])).
cnf(c_0_1168,negated_conjecture,(c4_2(a268,a270)|c4_2(a268,X1)|c7_2(a268,a270)|c7_1(a268)|c6_0|c8_1(a208)|c7_0|c4_0|c5_0|~c1_2(a268,X1)|~c7_2(a268,X1)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_876, c_0_98])).
cnf(c_0_1169,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|~c6_2(X1,X2)|~c1_2(X1,X3)|~c3_2(a208,a209)|~c7_2(X1,X3)|~c8_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_476])]), ['final']).
cnf(c_0_1170,negated_conjecture,(c2_2(a158,a293)|c8_2(a159,a292)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_873, c_0_874])).
cnf(c_0_1171,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|~c6_2(X1,X2)|~c2_2(a208,a209)|~c1_2(X1,X3)|~c7_2(X1,X3)|~c8_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_476])]), ['final']).
cnf(c_0_1172,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(a197)|ndr1_1(a296)|c5_0|c1_0|~c3_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_727, c_0_82]), ['final']).
cnf(c_0_1173,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_877]), ['final']).
cnf(c_0_1174,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_878, c_0_565])).
cnf(c_0_1175,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_879, c_0_565])).
cnf(c_0_1176,negated_conjecture,(c2_2(a294,a295)|c5_2(a197,a278)|c8_2(a160,a161)|c4_2(a197,a293)|c7_2(a197,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_880, c_0_502]), ['final']).
cnf(c_0_1177,negated_conjecture,(c2_2(a294,a295)|c5_2(a197,a278)|c8_2(a160,a161)|c8_2(a197,a292)|c4_2(a197,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_880, c_0_562]), ['final']).
cnf(c_0_1178,negated_conjecture,(c2_2(a294,a295)|c5_2(a162,a278)|c8_2(a160,a161)|c4_2(a162,a293)|c7_2(a162,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_881, c_0_502]), ['final']).
cnf(c_0_1179,negated_conjecture,(c2_2(a294,a295)|c5_2(a162,a278)|c8_2(a160,a161)|c8_2(a162,a292)|c4_2(a162,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_881, c_0_562]), ['final']).
cnf(c_0_1180,negated_conjecture,(c2_2(a294,a295)|c5_2(a288,a278)|c8_2(a160,a161)|c4_2(a288,a293)|c7_2(a288,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_882, c_0_502]), ['final']).
cnf(c_0_1181,negated_conjecture,(c2_2(a294,a295)|c5_2(a288,a278)|c8_2(a160,a161)|c8_2(a288,a292)|c4_2(a288,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_882, c_0_562]), ['final']).
cnf(c_0_1182,negated_conjecture,(c2_2(a294,a295)|c5_2(a253,a278)|c8_2(a160,a161)|c4_2(a253,a293)|c7_2(a253,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_883, c_0_502]), ['final']).
cnf(c_0_1183,negated_conjecture,(c2_2(a294,a295)|c5_2(a253,a278)|c8_2(a160,a161)|c8_2(a253,a292)|c4_2(a253,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_883, c_0_562]), ['final']).
cnf(c_0_1184,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c4_2(a159,a293)|c7_2(a159,a292)|c2_1(a197)|c5_0|c1_0|~c3_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_720, c_0_82]), ['final']).
cnf(c_0_1185,negated_conjecture,(c2_2(a197,a198)|c8_2(a159,a292)|c8_2(a197,a198)|c4_2(a159,a293)|c2_1(a197)|c5_0|c1_0|~c3_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_721, c_0_82]), ['final']).
cnf(c_0_1186,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_884]), ['final']).
cnf(c_0_1187,negated_conjecture,(c2_2(a204,a205)|c5_2(a197,a278)|c8_2(a160,a161)|c4_2(a197,a293)|c7_2(a197,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_885, c_0_502]), ['final']).
cnf(c_0_1188,negated_conjecture,(c2_2(a204,a205)|c5_2(a197,a278)|c8_2(a160,a161)|c8_2(a197,a292)|c4_2(a197,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_885, c_0_562]), ['final']).
cnf(c_0_1189,negated_conjecture,(c2_2(a204,a205)|c5_2(a162,a278)|c8_2(a160,a161)|c4_2(a162,a293)|c7_2(a162,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_886, c_0_502]), ['final']).
cnf(c_0_1190,negated_conjecture,(c2_2(a204,a205)|c5_2(a162,a278)|c8_2(a160,a161)|c8_2(a162,a292)|c4_2(a162,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_886, c_0_562]), ['final']).
cnf(c_0_1191,negated_conjecture,(c2_2(a204,a205)|c5_2(a288,a278)|c8_2(a160,a161)|c4_2(a288,a293)|c7_2(a288,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_887, c_0_502]), ['final']).
cnf(c_0_1192,negated_conjecture,(c2_2(a204,a205)|c5_2(a288,a278)|c8_2(a160,a161)|c8_2(a288,a292)|c4_2(a288,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_887, c_0_562]), ['final']).
cnf(c_0_1193,negated_conjecture,(c2_2(a204,a205)|c5_2(a253,a278)|c8_2(a160,a161)|c4_2(a253,a293)|c7_2(a253,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_888, c_0_502]), ['final']).
cnf(c_0_1194,negated_conjecture,(c2_2(a204,a205)|c5_2(a253,a278)|c8_2(a160,a161)|c8_2(a253,a292)|c4_2(a253,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_888, c_0_562]), ['final']).
cnf(c_0_1195,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_889, c_0_502]), ['final']).
cnf(c_0_1196,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_889, c_0_562]), ['final']).
cnf(c_0_1197,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a162,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_890, c_0_502]), ['final']).
cnf(c_0_1198,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_890, c_0_562]), ['final']).
cnf(c_0_1199,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a288,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_891, c_0_502]), ['final']).
cnf(c_0_1200,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_891, c_0_562]), ['final']).
cnf(c_0_1201,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_892, c_0_502]), ['final']).
cnf(c_0_1202,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_892, c_0_562]), ['final']).
cnf(c_0_1203,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_893, c_0_502]), ['final']).
cnf(c_0_1204,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_893, c_0_562]), ['final']).
cnf(c_0_1205,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a162,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_894, c_0_502]), ['final']).
cnf(c_0_1206,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_894, c_0_562]), ['final']).
cnf(c_0_1207,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a288,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_895, c_0_502]), ['final']).
cnf(c_0_1208,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_895, c_0_562]), ['final']).
cnf(c_0_1209,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_896, c_0_502]), ['final']).
cnf(c_0_1210,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_896, c_0_562]), ['final']).
cnf(c_0_1211,negated_conjecture,(c2_2(a159,X1)|c2_2(X2,X3)|c8_2(X2,X3)|c2_1(X2)|~c8_2(a159,X1)|~c4_2(a159,X1)|~c4_2(X2,X3)|~c3_1(X2)|~c2_0|~ndr1_1(a159)|~ndr1_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_476])]), ['final']).
cnf(c_0_1212,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c3_1(a241)), inference(spm,[status(thm)],[c_0_897, c_0_122]), ['final']).
cnf(c_0_1213,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|~c6_2(a208,a209)|~c6_2(X1,X2)|~c1_2(X1,X3)|~c7_2(X1,X3)|~c8_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_476])]), ['final']).
cnf(c_0_1214,negated_conjecture,(c4_2(a233,a234)|c4_2(a233,X1)|c7_2(a233,a234)|c7_1(a233)|c6_0|c7_0|ndr1_1(a208)|ssSkC17|~c1_2(a233,X1)|~c7_2(a233,X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_875, c_0_39])).
cnf(c_0_1215,negated_conjecture,(c4_2(a233,a234)|c4_2(a233,X1)|c7_2(a233,a234)|c7_1(a233)|c6_0|c8_1(a208)|c7_0|ssSkC17|~c1_2(a233,X1)|~c7_2(a233,X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_876, c_0_39])).
cnf(c_0_1216,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c3_1(a217)), inference(spm,[status(thm)],[c_0_898, c_0_766]), ['final']).
cnf(c_0_1217,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(a241)|ssSkC19|c5_0|c1_0|~c3_1(a241)), inference(spm,[status(thm)],[c_0_899, c_0_122]), ['final']).
cnf(c_0_1218,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c3_1(a159)|c5_1(a159)|c6_0|ssSkC22|c2_1(a217)|c5_0|c1_0|~c3_1(a217)), inference(spm,[status(thm)],[c_0_900, c_0_766]), ['final']).
cnf(c_0_1219,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_901, c_0_874])).
cnf(c_0_1220,negated_conjecture,(c2_2(a294,a295)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_902, c_0_502]), ['final']).
cnf(c_0_1221,negated_conjecture,(c2_2(a294,a295)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_902, c_0_562]), ['final']).
cnf(c_0_1222,negated_conjecture,(c2_2(a294,a295)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a162,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_903, c_0_502]), ['final']).
cnf(c_0_1223,negated_conjecture,(c2_2(a294,a295)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_903, c_0_562]), ['final']).
cnf(c_0_1224,negated_conjecture,(c2_2(a294,a295)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a288,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_904, c_0_502]), ['final']).
cnf(c_0_1225,negated_conjecture,(c2_2(a294,a295)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_904, c_0_562]), ['final']).
cnf(c_0_1226,negated_conjecture,(c2_2(a294,a295)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_905, c_0_502]), ['final']).
cnf(c_0_1227,negated_conjecture,(c2_2(a294,a295)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_905, c_0_562]), ['final']).
cnf(c_0_1228,negated_conjecture,(c2_2(a204,a205)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_906, c_0_502]), ['final']).
cnf(c_0_1229,negated_conjecture,(c2_2(a197,a198)|c5_2(a296,a297)|c8_2(a197,a198)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(a197)|c5_0|c1_0|~c3_1(a197)), inference(spm,[status(thm)],[c_0_907, c_0_124]), ['final']).
cnf(c_0_1230,negated_conjecture,(c2_2(a204,a205)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_906, c_0_562]), ['final']).
cnf(c_0_1231,negated_conjecture,(c2_2(a204,a205)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a162,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_908, c_0_502]), ['final']).
cnf(c_0_1232,negated_conjecture,(c2_2(a204,a205)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_908, c_0_562]), ['final']).
cnf(c_0_1233,negated_conjecture,(c2_2(a204,a205)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a288,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_909, c_0_502]), ['final']).
cnf(c_0_1234,negated_conjecture,(c2_2(a204,a205)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_909, c_0_562]), ['final']).
cnf(c_0_1235,negated_conjecture,(c2_2(a204,a205)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_910, c_0_502]), ['final']).
cnf(c_0_1236,negated_conjecture,(c2_2(a204,a205)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_910, c_0_562]), ['final']).
cnf(c_0_1237,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_911, c_0_644])).
cnf(c_0_1238,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c5_0|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_912, c_0_644])).
cnf(c_0_1239,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a159)|c2_1(a223)|ssSkC15|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_913, c_0_644])).
cnf(c_0_1240,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a159)|c2_1(a223)|ssSkC15|c5_0|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_914, c_0_644])).
cnf(c_0_1241,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|c6_0|ssSkC15|c5_0|c1_0|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_915, c_0_644])).
cnf(c_0_1242,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|ssSkC15|c1_0|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_916, c_0_644])).
cnf(c_0_1243,negated_conjecture,(c2_2(a218,a219)|c5_2(a296,a297)|c8_2(a218,a219)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a218)|ssSkC13|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_716, c_0_125])).
cnf(c_0_1244,negated_conjecture,(c4_2(a200,a278)|c4_2(a200,a293)|c3_2(a200,a278)|c7_2(a200,a292)|c3_0|c8_0|c7_0|ssSkC10|c1_0|~c6_1(a200)|~c8_1(a200)|~ndr1_1(a200)), inference(spm,[status(thm)],[c_0_917, c_0_918])).
cnf(c_0_1245,negated_conjecture,(c2_2(a218,a219)|c5_2(a296,a297)|c8_2(a218,a219)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(a218)|ssSkC13|c5_0|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_717, c_0_125])).
cnf(c_0_1246,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a159)|ndr1_1(a296)|ssSkC15|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_911, c_0_565])).
cnf(c_0_1247,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a159)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_912, c_0_565])).
cnf(c_0_1248,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a159)|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_913, c_0_565])).
cnf(c_0_1249,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a159)|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_914, c_0_565])).
cnf(c_0_1250,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a159)|c5_0|c1_0|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_919, c_0_872])).
cnf(c_0_1251,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a159)|c2_1(a158)|c5_0|c1_0|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_920, c_0_872])).
cnf(c_0_1252,negated_conjecture,(c2_2(a223,a293)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|c6_0|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_915, c_0_565])).
cnf(c_0_1253,negated_conjecture,(c2_2(a223,a293)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|ndr1_1(a296)|ssSkC15|c1_0|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_916, c_0_565])).
cnf(c_0_1254,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c7_2(a223,a292)|c5_1(a223)|c6_0|ndr1_1(a193)|ssSkC15|c5_0|c1_0|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_921, c_0_644])).
cnf(c_0_1255,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|ndr1_1(a193)|ssSkC15|c1_0|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_922, c_0_644])).
cnf(c_0_1256,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c8_0|c2_1(a241)|ndr1_1(a159)|ssSkC19|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_729, c_0_83]), ['final']).
cnf(c_0_1257,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a159)|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_729, c_0_728]), ['final']).
cnf(c_0_1258,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c2_1(a241)|ndr1_1(a159)|ssSkC19|c5_0|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_730, c_0_83]), ['final']).
cnf(c_0_1259,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|ssSkC22|c2_1(a217)|ndr1_1(a159)|c5_0|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_730, c_0_728]), ['final']).
cnf(c_0_1260,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c8_0|c2_1(a159)|c2_1(a241)|ssSkC19|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_731, c_0_83]), ['final']).
cnf(c_0_1261,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c8_0|ssSkC22|c2_1(a159)|c2_1(a217)|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_731, c_0_728]), ['final']).
cnf(c_0_1262,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c2_1(a159)|c2_1(a241)|ssSkC19|c5_0|c1_0|~c3_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_732, c_0_83]), ['final']).
cnf(c_0_1263,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|ssSkC22|c2_1(a159)|c2_1(a217)|c5_0|c1_0|~c3_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_732, c_0_728]), ['final']).
cnf(c_0_1264,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a168)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_923])).
cnf(c_0_1265,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a168)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_924])).
cnf(c_0_1266,negated_conjecture,(ssSkC7|c5_2(a296,a297)|c1_1(a168)|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(X1)|ndr1_1(X2)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_643, c_0_925])).
cnf(c_0_1267,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a168)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_926])).
cnf(c_0_1268,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a168)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_927])).
cnf(c_0_1269,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a159)|ndr1_1(a296)|c5_0|c1_0|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_919, c_0_874])).
cnf(c_0_1270,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a159)|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_920, c_0_874])).
cnf(c_0_1271,negated_conjecture,(c2_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|ndr1_1(a193)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_921, c_0_565])).
cnf(c_0_1272,negated_conjecture,(c2_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|ndr1_1(a193)|ndr1_1(a296)|ssSkC15|c1_0|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_922, c_0_565])).
cnf(c_0_1273,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|~c6_2(X1,X2)|~c1_2(X1,X3)|~c7_2(X1,X3)|~c1_1(a208)|~c8_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_476])]), ['final']).
cnf(c_0_1274,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c2_1(a197)|ndr1_1(a159)|c5_0|c1_0|~c3_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_730, c_0_82]), ['final']).
cnf(c_0_1275,negated_conjecture,(ssSkC7|c5_2(a296,a297)|c1_1(a168)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(X1)|ndr1_1(X2)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_643, c_0_928])).
cnf(c_0_1276,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c2_1(a159)|c2_1(a197)|c5_0|c1_0|~c3_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_732, c_0_82]), ['final']).
cnf(c_0_1277,negated_conjecture,(ssSkC7|c1_1(a168)|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ndr1_1(X2)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_643, c_0_929])).
cnf(c_0_1278,negated_conjecture,(c5_2(a233,a234)|c6_1(a233)|c2_0|c4_0|ssSkC17|~c5_2(a233,X1)|~c8_2(a233,a234)|~c7_2(a233,X1)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_744, c_0_39])).
cnf(c_0_1279,negated_conjecture,(c5_2(a268,a270)|c6_1(a268)|c2_0|c4_0|c5_0|~c5_2(a268,X1)|~c8_2(a268,a270)|~c7_2(a268,X1)), inference(spm,[status(thm)],[c_0_930, c_0_127])).
cnf(c_0_1280,negated_conjecture,(c5_2(a294,a147)|c4_2(a294,a295)|c4_2(a294,X1)|c3_2(a294,X1)|c6_0|c8_0|c1_0|~c5_2(a294,a295)|~ndr1_1(a294)), inference(spm,[status(thm)],[c_0_624, c_0_13])).
cnf(c_0_1281,negated_conjecture,(c5_2(a204,a147)|c4_2(a204,a205)|c4_2(a204,X1)|c3_2(a204,X1)|c6_0|c8_0|c1_0|~c5_2(a204,a205)|~ndr1_1(a204)), inference(spm,[status(thm)],[c_0_624, c_0_15])).
cnf(c_0_1282,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a160)|ndr1_1(a168)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_931])).
cnf(c_0_1283,negated_conjecture,(ssSkC7|c1_2(a202,a190)|c3_1(a202)|c2_0|ssSkC11|ssSkC15|~c5_1(a223)), inference(spm,[status(thm)],[c_0_46, c_0_932]), ['final']).
cnf(c_0_1284,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a160)|ndr1_1(a168)|c1_0|~c1_1(a168)), inference(spm,[status(thm)],[c_0_537, c_0_933])).
cnf(c_0_1285,negated_conjecture,(ssSkC7|c1_1(a168)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ndr1_1(X2)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_643, c_0_934])).
cnf(c_0_1286,negated_conjecture,(ssSkC7|c1_2(a217,a190)|c3_1(a217)|c2_0|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_46, c_0_935]), ['final']).
cnf(c_0_1287,negated_conjecture,(c1_2(X1,a261)|c1_2(X1,X2)|c8_2(X1,X2)|c6_0|c5_0|~c7_2(X1,X2)|~ssSkP5(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_476])]), ['final']).
cnf(c_0_1288,negated_conjecture,(c2_2(X1,a261)|c1_2(X1,X2)|c8_2(X1,X2)|c6_0|c5_0|~c7_2(X1,X2)|~ssSkP5(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_476])]), ['final']).
cnf(c_0_1289,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c3_2(X1,X2)|~c5_2(X1,X3)|~c7_2(X1,a147)|~c7_2(X1,X3)|~c3_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_476])]), ['final']).
cnf(c_0_1290,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c3_2(X1,X2)|~c5_2(X1,X3)|~c3_2(X1,a147)|~c7_2(X1,X3)|~c3_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_476])]), ['final']).
cnf(c_0_1291,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c3_1(a202)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_738, c_0_854]), ['final']).
cnf(c_0_1292,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c3_2(a168,a248)|c3_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_738, c_0_856]), ['final']).
cnf(c_0_1293,negated_conjecture,(ssSkC7|ssSkC20|c3_1(a202)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(a296)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_738, c_0_599]), ['final']).
cnf(c_0_1294,negated_conjecture,(ssSkC7|ssSkC20|c3_2(a168,a248)|c3_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(a296)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_738, c_0_936]), ['final']).
cnf(c_0_1295,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c3_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a168)|ndr1_1(a223)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_738, c_0_862]), ['final']).
cnf(c_0_1296,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_883, c_0_559]), ['final']).
cnf(c_0_1297,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_888, c_0_559]), ['final']).
cnf(c_0_1298,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_892, c_0_559]), ['final']).
cnf(c_0_1299,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_896, c_0_559]), ['final']).
cnf(c_0_1300,negated_conjecture,(c2_2(a223,X1)|c7_2(a223,X1)|c8_0|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_132, c_0_590]), ['final']).
cnf(c_0_1301,negated_conjecture,(ssSkC7|c3_1(a221)|ndr1_1(a221)|ssSkC14), inference(spm,[status(thm)],[c_0_133, c_0_592]), ['final']).
cnf(c_0_1302,negated_conjecture,(c1_1(a168)|c5_1(a287)|c3_0|c6_0|ndr1_1(a287)|ndr1_1(X1)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_643, c_0_92])).
cnf(c_0_1303,negated_conjecture,(ssSkC7|c3_1(a168)|c3_0|ndr1_1(a168)|c1_0), inference(spm,[status(thm)],[c_0_537, c_0_592]), ['final']).
cnf(c_0_1304,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_905, c_0_559]), ['final']).
cnf(c_0_1305,negated_conjecture,(c1_1(a265)|c6_0|ndr1_1(X1)|ssSkC23|~c8_1(a265)), inference(spm,[status(thm)],[c_0_134, c_0_531]), ['final']).
cnf(c_0_1306,negated_conjecture,(c5_1(X1)|c8_1(a265)|c2_1(a265)|ndr1_1(X1)|ssSkC23|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_134, c_0_482]), ['final']).
cnf(c_0_1307,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_910, c_0_559]), ['final']).
cnf(c_0_1308,negated_conjecture,(ssSkC7|ssSkC20|c3_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_738, c_0_603]), ['final']).
cnf(c_0_1309,negated_conjecture,(c1_1(a173)|c6_0|ndr1_1(X1)|ssSkC3|~c8_1(a173)), inference(spm,[status(thm)],[c_0_135, c_0_531]), ['final']).
cnf(c_0_1310,negated_conjecture,(c5_1(X1)|c8_1(a173)|c2_1(a173)|ndr1_1(X1)|ssSkC3|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_135, c_0_482]), ['final']).
cnf(c_0_1311,negated_conjecture,(c8_2(X1,X2)|c4_1(X1)|c6_0|c8_0|c4_0|~c6_2(X1,X2)|~c7_2(X1,X2)|~c8_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_937, c_0_19]), ['final']).
cnf(c_0_1312,negated_conjecture,(c7_2(a148,a149)|c6_0|c7_0), inference(spm,[status(thm)],[c_0_136, c_0_724]), ['final']).
cnf(c_0_1313,negated_conjecture,(c5_2(X1,a278)|c1_1(X1)|c8_0|c2_0|c7_0|~c7_1(X2)|~c5_1(X1)|~c2_1(X2)), inference(spm,[status(thm)],[c_0_488, c_0_508]), ['final']).
cnf(c_0_1314,negated_conjecture,(c5_2(a252,X1)|~c2_2(a252,X1)|~c4_2(a252,X1)|~c8_1(X2)|~ndr1_1(a252)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_137, c_0_476])]), ['final']).
cnf(c_0_1315,negated_conjecture,(c1_1(X1)|c8_0|c2_0|c7_0|ndr1_1(X1)|~c7_1(X2)|~c5_1(X1)|~c2_1(X2)), inference(spm,[status(thm)],[c_0_529, c_0_508]), ['final']).
cnf(c_0_1316,negated_conjecture,(c6_2(a174,X1)|c6_0|c8_0|ssSkC4|c1_0|~c5_2(a174,X1)|~c8_2(a174,X1)), inference(spm,[status(thm)],[c_0_138, c_0_626]), ['final']).
cnf(c_0_1317,negated_conjecture,(c6_2(a275,a206)|c8_2(a275,a207)|c3_0|ssSkC25|c4_0|~c4_2(a275,a276)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_938, c_0_100])).
cnf(c_0_1318,negated_conjecture,(c8_2(a275,a207)|c4_2(a275,a206)|c3_0|ssSkC25|c4_0|~c4_2(a275,a276)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_938, c_0_101])).
cnf(c_0_1319,negated_conjecture,(c8_2(a275,a207)|c3_2(a275,a206)|c3_0|ssSkC25|c4_0|~c4_2(a275,a276)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_938, c_0_104])).
cnf(c_0_1320,negated_conjecture,(c8_0|c7_0|ndr1_1(a182)|c1_0), inference(spm,[status(thm)],[c_0_117, c_0_559]), ['final']).
cnf(c_0_1321,negated_conjecture,(c6_2(a241,a206)|c8_2(a241,a207)|c3_0|ssSkC19|c4_0|~c4_2(a241,a243)), inference(spm,[status(thm)],[c_0_939, c_0_122])).
cnf(c_0_1322,negated_conjecture,(c8_2(a241,a207)|c4_2(a241,a206)|c3_0|ssSkC19|c4_0|~c4_2(a241,a243)), inference(spm,[status(thm)],[c_0_940, c_0_122])).
cnf(c_0_1323,negated_conjecture,(c1_1(a202)|c8_0|c2_1(X1)|ndr1_1(a202)|ssSkC11|~c7_1(X1)|~c5_1(X1)), inference(spm,[status(thm)],[c_0_941, c_0_50]), ['final']).
cnf(c_0_1324,negated_conjecture,(c1_1(a217)|c8_0|c2_1(X1)|ndr1_1(a217)|c1_0|~c7_1(X1)|~c5_1(X1)), inference(spm,[status(thm)],[c_0_941, c_0_509]), ['final']).
cnf(c_0_1325,negated_conjecture,(c6_1(X1)|c5_1(X2)|c8_1(X1)|c2_1(X1)|c1_0|~c8_2(X2,a151)|~c3_1(X2)), inference(spm,[status(thm)],[c_0_942, c_0_4]), ['final']).
cnf(c_0_1326,negated_conjecture,(c2_2(a204,X1)|c8_2(a204,X1)|c1_0), inference(spm,[status(thm)],[c_0_139, c_0_30]), ['final']).
cnf(c_0_1327,negated_conjecture,(c6_1(X1)|c5_1(X2)|c8_1(X1)|c2_1(X1)|~c1_2(X2,a151)|~c3_1(X2)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_140, c_0_476])]), ['final']).
cnf(c_0_1328,negated_conjecture,(c8_2(a241,a207)|c3_2(a241,a206)|c3_0|ssSkC19|c4_0|~c4_2(a241,a243)), inference(spm,[status(thm)],[c_0_943, c_0_122])).
cnf(c_0_1329,negated_conjecture,(c6_2(a268,a206)|c8_2(a268,a207)|c3_0|c4_0|c5_0|~c4_2(a268,a269)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_944, c_0_100])).
cnf(c_0_1330,negated_conjecture,(c8_2(a268,a207)|c4_2(a268,a206)|c3_0|c4_0|c5_0|~c4_2(a268,a269)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_944, c_0_101])).
cnf(c_0_1331,negated_conjecture,(c8_2(a268,a207)|c3_2(a268,a206)|c3_0|c4_0|c5_0|~c4_2(a268,a269)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_944, c_0_104])).
cnf(c_0_1332,negated_conjecture,(c6_1(X1)|c5_1(X2)|c8_1(X1)|c2_1(X1)|~c5_2(X2,a151)|~c3_1(X2)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_141, c_0_476])]), ['final']).
cnf(c_0_1333,negated_conjecture,(c5_2(a148,X1)|c6_0|c7_0|~c8_2(a148,X1)|~c4_2(a148,X1)|~ndr1_1(a148)), inference(spm,[status(thm)],[c_0_142, c_0_724])).
cnf(c_0_1334,negated_conjecture,(c6_0|c7_0|ndr1_1(a148)), inference(spm,[status(thm)],[c_0_143, c_0_724]), ['final']).
cnf(c_0_1335,negated_conjecture,(c8_2(a271,a272)|c3_2(a271,a272)|c8_0|c2_1(a271)|ssSkC24|~ndr1_1(a271)), inference(spm,[status(thm)],[c_0_945, c_0_144])).
cnf(c_0_1336,negated_conjecture,(ssSkC20|c3_2(a168,a248)|c5_1(a287)|c3_0|ndr1_1(a287)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_711, c_0_92])).
cnf(c_0_1337,negated_conjecture,(c8_0|c7_0|ndr1_1(a221)|ssSkC14|c1_0), inference(spm,[status(thm)],[c_0_118, c_0_559]), ['final']).
cnf(c_0_1338,negated_conjecture,(ssSkC20|c5_1(a287)|c3_0|ndr1_1(a168)|ndr1_1(a287)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_567, c_0_92])).
cnf(c_0_1339,negated_conjecture,(c5_1(X1)|c8_1(a197)|c2_1(a197)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_145, c_0_482])).
cnf(c_0_1340,negated_conjecture,(c4_1(X1)|c8_0|c1_0|~c3_2(X1,a293)|~ssSkP10(X1)|~c4_0), inference(spm,[status(thm)],[c_0_946, c_0_478])).
cnf(c_0_1341,negated_conjecture,(c4_1(X1)|c5_0|c1_0|~c3_2(X1,a293)|~ssSkP10(X1)|~c4_0), inference(spm,[status(thm)],[c_0_946, c_0_20])).
cnf(c_0_1342,negated_conjecture,(c4_1(X1)|c8_0|c1_0|~c1_2(X1,a293)|~ssSkP10(X1)|~c4_0), inference(spm,[status(thm)],[c_0_947, c_0_478])).
cnf(c_0_1343,negated_conjecture,(c4_1(X1)|c5_0|c1_0|~c1_2(X1,a293)|~ssSkP10(X1)|~c4_0), inference(spm,[status(thm)],[c_0_947, c_0_20])).
cnf(c_0_1344,negated_conjecture,(c7_2(X1,a264)|~c6_1(X1)|~ssSkP6(X1)|~c3_0|~c5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_146, c_0_476])]), ['final']).
cnf(c_0_1345,negated_conjecture,(c1_2(X1,a264)|~c6_1(X1)|~ssSkP6(X1)|~c3_0|~c5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_147, c_0_476])]), ['final']).
cnf(c_0_1346,negated_conjecture,(c5_2(X1,a264)|~c6_1(X1)|~ssSkP6(X1)|~c3_0|~c5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_148, c_0_476])]), ['final']).
cnf(c_0_1347,negated_conjecture,(c3_2(a155,X1)|ssSkC1|~c2_2(a155,X1)|~c4_2(a155,X1)), inference(spm,[status(thm)],[c_0_149, c_0_150]), ['final']).
cnf(c_0_1348,negated_conjecture,(ndr1_1(X1)|~c6_1(X1)|~ssSkP6(X1)|~c3_0|~c5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_151, c_0_476])]), ['final']).
cnf(c_0_1349,negated_conjecture,(c2_2(a223,X1)|c7_2(a223,X1)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_132, c_0_584]), ['final']).
cnf(c_0_1350,negated_conjecture,(c3_2(a217,a258)|c1_1(a217)|ssSkC22|c1_0), inference(spm,[status(thm)],[c_0_827, c_0_509])).
cnf(c_0_1351,negated_conjecture,(~c7_1(a252)|~c8_1(X1)|~c4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_476])]), ['final']).
cnf(c_0_1352,negated_conjecture,(c1_1(a252)|c1_0|~c8_1(X1)), inference(spm,[status(thm)],[c_0_948, c_0_4])).
cnf(c_0_1353,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a160,a248)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c5_1(X3)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_949, c_0_734]), ['final']).
cnf(c_0_1354,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a160,a248)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c5_1(X3)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_950, c_0_734]), ['final']).
cnf(c_0_1355,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a160,a248)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c5_1(X3)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_951, c_0_734]), ['final']).
cnf(c_0_1356,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a160,a248)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c5_1(X3)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_952, c_0_734]), ['final']).
cnf(c_0_1357,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a160,a248)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c5_1(X3)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_953, c_0_734]), ['final']).
cnf(c_0_1358,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a160,a248)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c5_1(X3)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_954, c_0_734]), ['final']).
cnf(c_0_1359,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a160,a248)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c5_1(X3)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_955, c_0_734]), ['final']).
cnf(c_0_1360,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a160,a248)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c5_1(X3)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_956, c_0_734]), ['final']).
cnf(c_0_1361,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c5_1(X3)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_957, c_0_734]), ['final']).
cnf(c_0_1362,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c5_1(X3)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_958, c_0_734]), ['final']).
cnf(c_0_1363,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c5_1(X3)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_959, c_0_734]), ['final']).
cnf(c_0_1364,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c5_1(X3)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_960, c_0_734]), ['final']).
cnf(c_0_1365,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a160,a161)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_961, c_0_734]), ['final']).
cnf(c_0_1366,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a160,a161)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_962, c_0_734]), ['final']).
cnf(c_0_1367,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a160,a161)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_963, c_0_734]), ['final']).
cnf(c_0_1368,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a160,a161)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_964, c_0_734]), ['final']).
cnf(c_0_1369,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_965, c_0_734]), ['final']).
cnf(c_0_1370,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_966, c_0_734]), ['final']).
cnf(c_0_1371,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_967, c_0_734]), ['final']).
cnf(c_0_1372,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_968, c_0_734]), ['final']).
cnf(c_0_1373,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_969, c_0_734]), ['final']).
cnf(c_0_1374,negated_conjecture,(c2_2(a294,a295)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_970, c_0_734]), ['final']).
cnf(c_0_1375,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a231,a292)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_971, c_0_734]), ['final']).
cnf(c_0_1376,negated_conjecture,(c2_2(a204,a205)|c5_2(a231,a278)|c1_2(X1,X2)|c8_2(a231,a292)|c4_2(a231,a293)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|c2_1(X1)|ndr1_1(a160)|ssSkC15|c1_0|~c4_2(X1,X2)|~c3_1(a223)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_972, c_0_734]), ['final']).
cnf(c_0_1377,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c4_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_973, c_0_81]), ['final']).
cnf(c_0_1378,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c4_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_974, c_0_81]), ['final']).
cnf(c_0_1379,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c1_2(X1,X2)|c8_2(a160,a161)|c3_2(X3,X4)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_975]), ['final']).
cnf(c_0_1380,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c4_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_976, c_0_81]), ['final']).
cnf(c_0_1381,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c4_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_977, c_0_81]), ['final']).
cnf(c_0_1382,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c1_2(X1,X2)|c8_2(a160,a161)|c3_2(X3,X4)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_978]), ['final']).
cnf(c_0_1383,negated_conjecture,(c2_2(a294,a295)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X5)|c6_0|c2_1(a160)|c2_1(X1)|ndr1_1(X6)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_979]), ['final']).
cnf(c_0_1384,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c4_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_980, c_0_81]), ['final']).
cnf(c_0_1385,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c4_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_981, c_0_81]), ['final']).
cnf(c_0_1386,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c1_2(X1,X2)|c3_2(X3,X4)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_982]), ['final']).
cnf(c_0_1387,negated_conjecture,(c2_2(a204,a205)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X5)|c6_0|c2_1(a160)|c2_1(X1)|ndr1_1(X6)|ndr1_1(X5)|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c3_1(X5)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_983]), ['final']).
cnf(c_0_1388,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c3_2(a160,a248)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c4_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_984, c_0_81]), ['final']).
cnf(c_0_1389,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC17|c1_0|~c4_2(a233,a234)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_985, c_0_81]), ['final']).
cnf(c_0_1390,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_986]), ['final']).
cnf(c_0_1391,negated_conjecture,(c2_2(a294,a295)|c7_2(a223,a292)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c8_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X3)|~c4_2(a223,X3)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_979]), ['final']).
cnf(c_0_1392,negated_conjecture,(c2_2(a294,a295)|c7_2(a223,a292)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X3)|~c4_2(a223,X3)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_979]), ['final']).
cnf(c_0_1393,negated_conjecture,(c2_2(a204,a205)|c7_2(a223,a292)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c8_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X3)|~c4_2(a223,X3)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_983]), ['final']).
cnf(c_0_1394,negated_conjecture,(c2_2(a204,a205)|c7_2(a223,a292)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X3)|~c4_2(a223,X3)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_983]), ['final']).
cnf(c_0_1395,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c1_2(X1,X2)|c3_2(X3,X4)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a160)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_989]), ['final']).
cnf(c_0_1396,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a223,a292)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c2_1(a160)|ndr1_1(X1)|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X2)|~c4_2(a223,X2)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_533]), ['final']).
cnf(c_0_1397,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c2_1(a160)|ndr1_1(X1)|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X2)|~c4_2(a223,X2)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_540]), ['final']).
cnf(c_0_1398,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a223,a292)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c2_1(a160)|ndr1_1(X1)|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X2)|~c4_2(a223,X2)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_542]), ['final']).
cnf(c_0_1399,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c2_1(a160)|ndr1_1(X1)|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X2)|~c4_2(a223,X2)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_543]), ['final']).
cnf(c_0_1400,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c1_2(X1,X2)|c3_2(X3,X4)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a160)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c6_2(X3,X4)|~c4_2(X1,X2)|~c3_1(X3)|~c7_2(X1,X2)|~c1_1(a231)|~c1_1(X1)|~c5_1(X3)|~ndr1_1(X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_648, c_0_990]), ['final']).
cnf(c_0_1401,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a223,a292)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X2)|~c4_2(a223,X2)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_533]), ['final']).
cnf(c_0_1402,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X2)|~c4_2(a223,X2)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_540]), ['final']).
cnf(c_0_1403,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a223,a292)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X2)|~c4_2(a223,X2)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_542]), ['final']).
cnf(c_0_1404,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X2)|~c4_2(a223,X2)|~c3_2(a223,a293)|~c3_1(X1)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_543]), ['final']).
cnf(c_0_1405,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c8_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X3)|~c4_2(a233,X3)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_979]), ['final']).
cnf(c_0_1406,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X3)|~c4_2(a233,X3)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_979]), ['final']).
cnf(c_0_1407,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c8_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X3)|~c4_2(a233,X3)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_532, c_0_983]), ['final']).
cnf(c_0_1408,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X3)|~c4_2(a233,X3)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_983]), ['final']).
cnf(c_0_1409,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_533]), ['final']).
cnf(c_0_1410,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(a160,a248)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_540]), ['final']).
cnf(c_0_1411,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_542]), ['final']).
cnf(c_0_1412,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_992, c_0_81]), ['final']).
cnf(c_0_1413,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_993, c_0_81]), ['final']).
cnf(c_0_1414,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_994, c_0_81]), ['final']).
cnf(c_0_1415,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_995, c_0_81]), ['final']).
cnf(c_0_1416,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_996, c_0_81]), ['final']).
cnf(c_0_1417,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a160,a161)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_997, c_0_81]), ['final']).
cnf(c_0_1418,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_998, c_0_81]), ['final']).
cnf(c_0_1419,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_999, c_0_81]), ['final']).
cnf(c_0_1420,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(a160,a248)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c2_1(a160)|ndr1_1(X1)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X2)|~c4_2(a233,X2)|~c3_2(a233,a234)|~c3_1(X1)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_543]), ['final']).
cnf(c_0_1421,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_1000, c_0_81]), ['final']).
cnf(c_0_1422,negated_conjecture,(c2_2(a204,a205)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_1001, c_0_81]), ['final']).
cnf(c_0_1423,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c4_2(a233,a293)|c7_2(a233,a292)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_1002, c_0_81]), ['final']).
cnf(c_0_1424,negated_conjecture,(c2_2(a294,a295)|c5_2(a233,a278)|c8_2(a233,a292)|c4_2(a233,a293)|c4_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC17|c1_0|~c4_2(a233,a234)), inference(spm,[status(thm)],[c_0_1003, c_0_81]), ['final']).
cnf(c_0_1425,negated_conjecture,(c6_2(a231,a232)|c5_2(a197,a278)|c1_2(a197,a198)|c4_2(a197,a293)|c3_2(X1,X2)|c7_2(a197,a292)|c8_0|c7_0|c2_1(a197)|c5_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a197,a198)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1004, c_0_124]), ['final']).
cnf(c_0_1426,negated_conjecture,(c6_2(a231,a232)|c5_2(a197,a278)|c1_2(a197,a198)|c8_2(a197,a292)|c4_2(a197,a293)|c3_2(X1,X2)|c8_0|c7_0|c2_1(a197)|c5_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a197,a198)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1005, c_0_124]), ['final']).
cnf(c_0_1427,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c7_2(a160,a248)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1006, c_0_502]), ['final']).
cnf(c_0_1428,negated_conjecture,(c6_2(a231,a232)|c5_2(a265,a278)|c1_2(a265,a266)|c4_2(a265,a293)|c3_2(a223,a293)|c7_2(a265,a292)|c7_2(a223,a292)|c8_0|c7_0|c2_1(a265)|ssSkC15|ssSkC23|c1_0|~c3_1(a223)|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1007, c_0_734]), ['final']).
cnf(c_0_1429,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c7_2(a160,a248)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1006, c_0_562]), ['final']).
cnf(c_0_1430,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a160,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1008, c_0_502]), ['final']).
cnf(c_0_1431,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1008, c_0_562]), ['final']).
cnf(c_0_1432,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a160,a248)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1009, c_0_502]), ['final']).
cnf(c_0_1433,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a160,a248)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1009, c_0_562]), ['final']).
cnf(c_0_1434,negated_conjecture,(c6_2(a231,a232)|c5_2(a265,a278)|c1_2(a265,a266)|c8_2(a265,a292)|c4_2(a265,a293)|c3_2(a223,a293)|c7_2(a223,a292)|c8_0|c7_0|c2_1(a265)|ssSkC15|ssSkC23|c1_0|~c3_1(a223)|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1010, c_0_734]), ['final']).
cnf(c_0_1435,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a160,a248)|c3_2(a215,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1011, c_0_502]), ['final']).
cnf(c_0_1436,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a160,a248)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1011, c_0_562]), ['final']).
cnf(c_0_1437,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c7_2(a160,a248)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1012, c_0_502]), ['final']).
cnf(c_0_1438,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c7_2(a160,a248)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1012, c_0_562]), ['final']).
cnf(c_0_1439,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a160,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1013, c_0_502]), ['final']).
cnf(c_0_1440,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1013, c_0_562]), ['final']).
cnf(c_0_1441,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a160,a248)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1014, c_0_502]), ['final']).
cnf(c_0_1442,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a160,a248)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1014, c_0_562]), ['final']).
cnf(c_0_1443,negated_conjecture,(c6_2(a231,a232)|c5_2(a241,a278)|c1_2(a241,a243)|c4_2(a241,a293)|c3_2(X1,X2)|c7_2(a241,a292)|c8_0|c7_0|c2_1(a241)|ssSkC19|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a241,a243)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1015, c_0_122]), ['final']).
cnf(c_0_1444,negated_conjecture,(c6_2(a231,a232)|c5_2(a241,a278)|c1_2(a241,a243)|c8_2(a241,a292)|c4_2(a241,a293)|c3_2(X1,X2)|c8_0|c7_0|c2_1(a241)|ssSkC19|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a241,a243)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1016, c_0_122]), ['final']).
cnf(c_0_1445,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a160,a248)|c3_2(a215,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1017, c_0_502]), ['final']).
cnf(c_0_1446,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a160,a248)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1017, c_0_562]), ['final']).
cnf(c_0_1447,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a160,a248)|c7_2(a197,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1018]), ['final']).
cnf(c_0_1448,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1019]), ['final']).
cnf(c_0_1449,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a160,a248)|c7_2(a162,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1020]), ['final']).
cnf(c_0_1450,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1021]), ['final']).
cnf(c_0_1451,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a160,a248)|c7_2(a288,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1022]), ['final']).
cnf(c_0_1452,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1023]), ['final']).
cnf(c_0_1453,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a160,a248)|c7_2(a253,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1024]), ['final']).
cnf(c_0_1454,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1025]), ['final']).
cnf(c_0_1455,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c4_2(a197,a293)|c3_2(a160,a248)|c7_2(a197,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1026]), ['final']).
cnf(c_0_1456,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c3_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1027]), ['final']).
cnf(c_0_1457,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c4_2(a162,a293)|c3_2(a160,a248)|c7_2(a162,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1028]), ['final']).
cnf(c_0_1458,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c3_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1029]), ['final']).
cnf(c_0_1459,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c4_2(a288,a293)|c3_2(a160,a248)|c7_2(a288,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1030]), ['final']).
cnf(c_0_1460,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c3_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1031]), ['final']).
cnf(c_0_1461,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c4_2(a253,a293)|c3_2(a160,a248)|c7_2(a253,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1032]), ['final']).
cnf(c_0_1462,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c3_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1033]), ['final']).
cnf(c_0_1463,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a160,a248)|c7_2(a197,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1034]), ['final']).
cnf(c_0_1464,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1035]), ['final']).
cnf(c_0_1465,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a160,a248)|c7_2(a162,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1036]), ['final']).
cnf(c_0_1466,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1037]), ['final']).
cnf(c_0_1467,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a160,a248)|c7_2(a288,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1038]), ['final']).
cnf(c_0_1468,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1039]), ['final']).
cnf(c_0_1469,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a160,a248)|c7_2(a253,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1040]), ['final']).
cnf(c_0_1470,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1041]), ['final']).
cnf(c_0_1471,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c4_2(a197,a293)|c3_2(a160,a248)|c7_2(a197,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1042]), ['final']).
cnf(c_0_1472,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c3_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1043]), ['final']).
cnf(c_0_1473,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c4_2(a162,a293)|c3_2(a160,a248)|c7_2(a162,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1044]), ['final']).
cnf(c_0_1474,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c3_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1045]), ['final']).
cnf(c_0_1475,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c4_2(a288,a293)|c3_2(a160,a248)|c7_2(a288,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1046]), ['final']).
cnf(c_0_1476,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c3_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1047]), ['final']).
cnf(c_0_1477,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c4_2(a253,a293)|c3_2(a160,a248)|c7_2(a253,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1048]), ['final']).
cnf(c_0_1478,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c3_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1049]), ['final']).
cnf(c_0_1479,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a160,a278)|c4_2(a160,a293)|c7_2(a160,a248)|c7_2(a160,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1050]), ['final']).
cnf(c_0_1480,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a160,a278)|c4_2(a160,a293)|c3_2(a160,a248)|c7_2(a160,a292)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1051]), ['final']).
cnf(c_0_1481,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1052]), ['final']).
cnf(c_0_1482,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c3_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1053]), ['final']).
cnf(c_0_1483,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a160,a278)|c4_2(a160,a293)|c7_2(a160,a248)|c7_2(a160,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1054]), ['final']).
cnf(c_0_1484,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a160,a278)|c4_2(a160,a293)|c3_2(a160,a248)|c7_2(a160,a292)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1055]), ['final']).
cnf(c_0_1485,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1056]), ['final']).
cnf(c_0_1486,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a160,a278)|c8_2(a160,a292)|c4_2(a160,a293)|c3_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1057]), ['final']).
cnf(c_0_1487,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c4_2(a215,a258)|c7_2(a160,a248)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1006, c_0_1058]), ['final']).
cnf(c_0_1488,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a215,a258)|c7_2(a160,a248)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1006, c_0_1059]), ['final']).
cnf(c_0_1489,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a160)), inference(ef,[status(thm)],[c_0_795]), ['final']).
cnf(c_0_1490,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(X1,a248)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X2)|c6_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_979]), ['final']).
cnf(c_0_1491,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(X1,a248)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X2)|c6_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_979]), ['final']).
cnf(c_0_1492,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X3)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_513, c_0_979]), ['final']).
cnf(c_0_1493,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_979]), ['final']).
cnf(c_0_1494,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_979]), ['final']).
cnf(c_0_1495,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_979]), ['final']).
cnf(c_0_1496,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c1_1(a160)|c5_1(a294)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_979]), ['final']).
cnf(c_0_1497,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(X1,a248)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X2)|c6_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_983]), ['final']).
cnf(c_0_1498,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(X1,a248)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X2)|c6_0|c2_1(a160)|c2_1(X1)|ndr1_1(X3)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_983]), ['final']).
cnf(c_0_1499,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X3)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_513, c_0_983]), ['final']).
cnf(c_0_1500,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a197)), inference(spm,[status(thm)],[c_0_84, c_0_983]), ['final']).
cnf(c_0_1501,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(a162)), inference(spm,[status(thm)],[c_0_85, c_0_983]), ['final']).
cnf(c_0_1502,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC27|c1_0|~c3_1(X1)|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_983]), ['final']).
cnf(c_0_1503,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c1_1(a160)|c5_1(a204)|c5_1(X1)|c6_0|c2_1(a160)|ndr1_1(X2)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_983]), ['final']).
cnf(c_0_1504,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c4_2(a215,a258)|c3_2(a215,a248)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1008, c_0_1058]), ['final']).
cnf(c_0_1505,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a215,a248)|c3_2(a215,a258)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1008, c_0_1059]), ['final']).
cnf(c_0_1506,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c4_2(a215,a258)|c3_2(a160,a248)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1009, c_0_1058]), ['final']).
cnf(c_0_1507,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a160,a248)|c3_2(a215,a258)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1009, c_0_1059]), ['final']).
cnf(c_0_1508,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c8_0|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_544]), ['final']).
cnf(c_0_1509,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c8_0|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_545]), ['final']).
cnf(c_0_1510,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c8_0|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_546]), ['final']).
cnf(c_0_1511,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c4_2(a215,a258)|c3_2(a160,a248)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1011, c_0_1058]), ['final']).
cnf(c_0_1512,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a160,a248)|c3_2(a215,a248)|c3_2(a215,a258)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1011, c_0_1059]), ['final']).
cnf(c_0_1513,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a160)), inference(ef,[status(thm)],[c_0_801]), ['final']).
cnf(c_0_1514,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c8_0|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_547]), ['final']).
cnf(c_0_1515,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_544]), ['final']).
cnf(c_0_1516,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_545]), ['final']).
cnf(c_0_1517,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_546]), ['final']).
cnf(c_0_1518,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_547]), ['final']).
cnf(c_0_1519,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c4_2(a215,a258)|c7_2(a160,a248)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1012, c_0_1058]), ['final']).
cnf(c_0_1520,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a215,a258)|c7_2(a160,a248)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1012, c_0_1059]), ['final']).
cnf(c_0_1521,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a160)), inference(ef,[status(thm)],[c_0_802]), ['final']).
cnf(c_0_1522,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c4_2(a215,a258)|c3_2(a215,a248)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1013, c_0_1058]), ['final']).
cnf(c_0_1523,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a215,a248)|c3_2(a215,a258)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1013, c_0_1059]), ['final']).
cnf(c_0_1524,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c2_1(X1)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1060, c_0_522]), ['final']).
cnf(c_0_1525,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c2_1(X1)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1060, c_0_524]), ['final']).
cnf(c_0_1526,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1061, c_0_522]), ['final']).
cnf(c_0_1527,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c4_2(a215,a258)|c3_2(a160,a248)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1014, c_0_1058]), ['final']).
cnf(c_0_1528,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a160,a248)|c3_2(a215,a258)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1014, c_0_1059]), ['final']).
cnf(c_0_1529,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c8_0|c2_1(X1)|ndr1_1(a256)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1061, c_0_524]), ['final']).
cnf(c_0_1530,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c2_1(X1)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1062, c_0_522]), ['final']).
cnf(c_0_1531,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c2_1(X1)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1062, c_0_524]), ['final']).
cnf(c_0_1532,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c4_2(a215,a258)|c3_2(a160,a248)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1017, c_0_1058]), ['final']).
cnf(c_0_1533,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a160,a248)|c3_2(a215,a248)|c3_2(a215,a258)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|ssSkC22|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1017, c_0_1059]), ['final']).
cnf(c_0_1534,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(a160)), inference(ef,[status(thm)],[c_0_806]), ['final']).
cnf(c_0_1535,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1063, c_0_522]), ['final']).
cnf(c_0_1536,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c8_0|c2_1(X1)|ndr1_1(a256)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1063, c_0_524]), ['final']).
cnf(c_0_1537,negated_conjecture,(c2_2(a204,a205)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|c8_0|ndr1_1(a160)|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_548]), ['final']).
cnf(c_0_1538,negated_conjecture,(c2_2(a294,a295)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|c8_0|ndr1_1(a160)|ssSkC15|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_987, c_0_549]), ['final']).
cnf(c_0_1539,negated_conjecture,(c2_2(a204,a205)|c7_2(a223,a292)|c4_1(a204)|c5_1(a204)|ndr1_1(a160)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_548]), ['final']).
cnf(c_0_1540,negated_conjecture,(c2_2(a294,a295)|c7_2(a223,a292)|c4_1(a294)|c5_1(a294)|ndr1_1(a160)|ssSkC15|c5_0|c1_0|~c1_2(a223,a293)|~c1_2(a223,X1)|~c4_2(a223,X1)|~c3_2(a223,a293)|~c1_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_988, c_0_549]), ['final']).
cnf(c_0_1541,negated_conjecture,(ssSkC20|c2_2(a168,a248)|c4_1(a168)|c5_1(a168)|c5_1(X1)|c3_0|c6_0|c8_1(a160)|c8_1(X2)|c2_0|c2_1(a160)|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0|~c3_1(X1)|~c1_1(X2)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_486, c_0_1064]), ['final']).
cnf(c_0_1542,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c2_1(X1)|ndr1_1(a231)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1065, c_0_522]), ['final']).
cnf(c_0_1543,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c2_1(X1)|ndr1_1(a231)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1065, c_0_524]), ['final']).
cnf(c_0_1544,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c8_0|c2_1(X1)|ndr1_1(a231)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1066, c_0_522]), ['final']).
cnf(c_0_1545,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c8_0|c2_1(X1)|ndr1_1(a231)|ndr1_1(a256)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1066, c_0_524]), ['final']).
cnf(c_0_1546,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a231)|c1_1(a256)|c1_1(a223)|c6_0|c2_1(X1)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1067, c_0_522]), ['final']).
cnf(c_0_1547,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a231)|c1_1(a223)|c6_0|c2_1(X1)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1067, c_0_524]), ['final']).
cnf(c_0_1548,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a231)|c1_1(a256)|c1_1(a223)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1068, c_0_522]), ['final']).
cnf(c_0_1549,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a231)|c1_1(a223)|c6_0|c8_0|c2_1(X1)|ndr1_1(a256)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1068, c_0_524]), ['final']).
cnf(c_0_1550,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c4_2(a223,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1061, c_0_499]), ['final']).
cnf(c_0_1551,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a223,a293)|c7_2(a223,a292)|c7_2(a223,a185)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1061, c_0_582]), ['final']).
cnf(c_0_1552,negated_conjecture,(c5_2(a223,a185)|c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1061, c_0_722]), ['final']).
cnf(c_0_1553,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c4_2(a223,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1063, c_0_499]), ['final']).
cnf(c_0_1554,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c7_2(a223,a185)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1063, c_0_582]), ['final']).
cnf(c_0_1555,negated_conjecture,(c5_2(a231,a232)|c5_2(a223,a185)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1063, c_0_722]), ['final']).
cnf(c_0_1556,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_488, c_0_1069]), ['final']).
cnf(c_0_1557,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_488, c_0_1070]), ['final']).
cnf(c_0_1558,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_488, c_0_1071]), ['final']).
cnf(c_0_1559,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_488, c_0_1072]), ['final']).
cnf(c_0_1560,negated_conjecture,(c1_2(X1,X2)|c4_2(a223,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(X1)|ndr1_1(a231)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1066, c_0_499]), ['final']).
cnf(c_0_1561,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c7_2(a223,a185)|c6_0|c8_0|c2_1(X1)|ndr1_1(a231)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1066, c_0_582]), ['final']).
cnf(c_0_1562,negated_conjecture,(c2_2(a275,a147)|c2_2(X1,X2)|c1_2(X1,X2)|c4_2(a275,a276)|c4_2(a275,X3)|c3_2(a275,a147)|c3_2(a275,X3)|c4_1(X1)|c1_1(a275)|c6_0|c8_0|ssSkC25|~c7_2(X1,X2)|~c7_1(a275)|~ndr1_1(a275)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1073]), ['final']).
cnf(c_0_1563,negated_conjecture,(c5_2(a223,a185)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(X1)|ndr1_1(a231)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1066, c_0_722]), ['final']).
cnf(c_0_1564,negated_conjecture,(c1_2(X1,X2)|c4_2(a223,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a231)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1068, c_0_499]), ['final']).
cnf(c_0_1565,negated_conjecture,(c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c7_2(a223,a185)|c4_1(a231)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1068, c_0_582]), ['final']).
cnf(c_0_1566,negated_conjecture,(c5_2(a223,a185)|c1_2(X1,X2)|c3_2(a223,a293)|c7_2(a223,a292)|c4_1(a231)|c6_0|c8_0|c2_1(X1)|ssSkC15|c1_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a223)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1068, c_0_722]), ['final']).
cnf(c_0_1567,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c7_2(a168,a248)|c7_2(X1,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_975]), ['final']).
cnf(c_0_1568,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c3_2(X1,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_975]), ['final']).
cnf(c_0_1569,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c7_2(a168,a248)|c7_2(X1,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_982]), ['final']).
cnf(c_0_1570,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c3_2(X1,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_982]), ['final']).
cnf(c_0_1571,negated_conjecture,(c2_2(a294,a295)|c8_2(a160,a161)|c4_1(a294)|c5_1(a294)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_544]), ['final']).
cnf(c_0_1572,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_545]), ['final']).
cnf(c_0_1573,negated_conjecture,(c2_2(a204,a205)|c8_2(a160,a161)|c4_1(a204)|c5_1(a204)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_546]), ['final']).
cnf(c_0_1574,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_547]), ['final']).
cnf(c_0_1575,negated_conjecture,(c2_2(a294,a295)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1074, c_0_502]), ['final']).
cnf(c_0_1576,negated_conjecture,(c2_2(a294,a295)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1074, c_0_562]), ['final']).
cnf(c_0_1577,negated_conjecture,(c2_2(a294,a295)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_697, c_0_502]), ['final']).
cnf(c_0_1578,negated_conjecture,(c2_2(a294,a295)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_697, c_0_562]), ['final']).
cnf(c_0_1579,negated_conjecture,(c5_2(X1,X2)|c3_2(X1,X2)|c3_1(a223)|c3_1(X1)|c7_2(a223,a292)|c4_1(X1)|c2_1(a223)|ssSkC15|c5_0|c1_0|~c8_2(a223,a293)|~c8_2(X1,X2)|~c3_2(a223,a293)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1075, c_0_584]), ['final']).
cnf(c_0_1580,negated_conjecture,(c5_2(X1,X2)|c3_2(X1,X2)|c3_1(a223)|c3_1(X1)|c7_2(a223,a292)|c4_1(X1)|c8_0|c2_1(a223)|ssSkC15|c1_0|~c8_2(a223,a293)|~c8_2(X1,X2)|~c3_2(a223,a293)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1076, c_0_590]), ['final']).
cnf(c_0_1581,negated_conjecture,(c2_2(a204,a205)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1077, c_0_502]), ['final']).
cnf(c_0_1582,negated_conjecture,(c4_2(a223,a185)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a181)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_839, c_0_499]), ['final']).
cnf(c_0_1583,negated_conjecture,(c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c2_1(a181)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_839, c_0_522]), ['final']).
cnf(c_0_1584,negated_conjecture,(c7_2(a223,a292)|c1_1(a223)|c6_0|c2_1(a181)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_839, c_0_524]), ['final']).
cnf(c_0_1585,negated_conjecture,(c7_2(a223,a292)|c7_2(a223,a185)|c6_0|c8_0|c2_1(a181)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_839, c_0_582]), ['final']).
cnf(c_0_1586,negated_conjecture,(c5_2(a223,a185)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a181)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_839, c_0_722]), ['final']).
cnf(c_0_1587,negated_conjecture,(ssSkC20|c2_2(a217,a278)|c2_2(a168,a248)|c1_2(a168,a248)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c4_1(a168)|c1_1(a217)|c3_0|c8_0|c7_0|ndr1_1(a287)|c5_0|c1_0|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_631, c_0_1078]), ['final']).
cnf(c_0_1588,negated_conjecture,(c2_2(a204,a205)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1077, c_0_562]), ['final']).
cnf(c_0_1589,negated_conjecture,(c2_2(a204,a205)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_698, c_0_502]), ['final']).
cnf(c_0_1590,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_975]), ['final']).
cnf(c_0_1591,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_982]), ['final']).
cnf(c_0_1592,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c7_2(a168,a248)|c7_2(X1,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a160)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_989]), ['final']).
cnf(c_0_1593,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c3_2(X1,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a160)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_989]), ['final']).
cnf(c_0_1594,negated_conjecture,(c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|ndr1_1(a160)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_548]), ['final']).
cnf(c_0_1595,negated_conjecture,(c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|ndr1_1(a160)|c5_0|ssSkC17|c1_0|~c1_2(a233,a234)|~c1_2(a233,X1)|~c4_2(a233,X1)|~c3_2(a233,a234)|~c1_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_991, c_0_549]), ['final']).
cnf(c_0_1596,negated_conjecture,(c2_2(a204,a205)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c8_0|c7_0|c2_1(a160)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_698, c_0_562]), ['final']).
cnf(c_0_1597,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a273)|c5_2(a271,a278)|c4_2(a271,a273)|c4_2(a271,a293)|c7_2(a271,a292)|c7_2(a271,X1)|c8_0|c7_0|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)), inference(spm,[status(thm)],[c_0_1079, c_0_153]), ['final']).
cnf(c_0_1598,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a273)|c5_2(a271,a278)|c8_2(a271,a292)|c4_2(a271,a273)|c4_2(a271,a293)|c7_2(a271,X1)|c8_0|c7_0|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)), inference(spm,[status(thm)],[c_0_1080, c_0_153]), ['final']).
cnf(c_0_1599,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a272)|c5_2(a271,a278)|c4_2(a271,a272)|c4_2(a271,a293)|c7_2(a271,a292)|c7_2(a271,X1)|c8_0|c7_0|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)), inference(spm,[status(thm)],[c_0_1081, c_0_153]), ['final']).
cnf(c_0_1600,negated_conjecture,(c6_2(a271,X1)|c5_2(a271,a272)|c5_2(a271,a278)|c8_2(a271,a292)|c4_2(a271,a272)|c4_2(a271,a293)|c7_2(a271,X1)|c8_0|c7_0|ssSkC24|c5_0|c1_0|~c8_2(a271,X1)), inference(spm,[status(thm)],[c_0_1082, c_0_153]), ['final']).
cnf(c_0_1601,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c4_2(a223,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c6_0|c8_0|c2_1(a174)|ssSkC4|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1083, c_0_499]), ['final']).
cnf(c_0_1602,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c7_2(a223,a185)|c1_1(a256)|c6_0|c8_0|c2_1(a174)|ssSkC4|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1083, c_0_582]), ['final']).
cnf(c_0_1603,negated_conjecture,(c6_2(a231,a232)|c5_2(a223,a185)|c1_2(a174,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c6_0|c8_0|c2_1(a174)|ssSkC4|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1083, c_0_722]), ['final']).
cnf(c_0_1604,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c8_0|c2_1(a174)|ssSkC4|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1083, c_0_522]), ['final']).
cnf(c_0_1605,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c4_2(a223,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a174)|ndr1_1(a256)|ssSkC4|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1084, c_0_499]), ['final']).
cnf(c_0_1606,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c7_2(a223,a185)|c6_0|c8_0|c2_1(a174)|ndr1_1(a256)|ssSkC4|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1084, c_0_582]), ['final']).
cnf(c_0_1607,negated_conjecture,(c6_2(a231,a232)|c5_2(a223,a185)|c1_2(a174,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a174)|ndr1_1(a256)|ssSkC4|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1084, c_0_722]), ['final']).
cnf(c_0_1608,negated_conjecture,(c6_2(a231,a232)|c1_2(a174,a185)|c3_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c8_0|c2_1(a174)|ndr1_1(a256)|ssSkC4|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1084, c_0_524]), ['final']).
cnf(c_0_1609,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a268,a270)|c1_1(a256)|c1_1(a268)|c6_0|c2_1(X1)|c4_0|c5_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1085, c_0_522]), ['final']).
cnf(c_0_1610,negated_conjecture,(c1_2(X1,X2)|c4_2(a231,a232)|c3_2(a268,a270)|c1_1(a268)|c6_0|c2_1(X1)|ndr1_1(a256)|c4_0|c5_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1085, c_0_524]), ['final']).
cnf(c_0_1611,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a268,a270)|c1_1(a256)|c1_1(a268)|c6_0|c2_1(X1)|c4_0|c5_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1086, c_0_522]), ['final']).
cnf(c_0_1612,negated_conjecture,(c5_2(a231,a232)|c1_2(X1,X2)|c3_2(a268,a270)|c1_1(a268)|c6_0|c2_1(X1)|ndr1_1(a256)|c4_0|c5_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1086, c_0_524]), ['final']).
cnf(c_0_1613,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c7_2(X1,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_978]), ['final']).
cnf(c_0_1614,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c3_2(X1,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_978]), ['final']).
cnf(c_0_1615,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c7_2(X1,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_986]), ['final']).
cnf(c_0_1616,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c3_2(X1,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_986]), ['final']).
cnf(c_0_1617,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|c2_1(X1)|ndr1_1(a160)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_989]), ['final']).
cnf(c_0_1618,negated_conjecture,(c5_2(a296,a297)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a181)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)), inference(spm,[status(thm)],[c_0_1087, c_0_584]), ['final']).
cnf(c_0_1619,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c7_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1088]), ['final']).
cnf(c_0_1620,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a160,a248)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1089]), ['final']).
cnf(c_0_1621,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c7_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1090]), ['final']).
cnf(c_0_1622,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a160,a248)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a253)|ndr1_1(X1)|ssSkC21|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1091]), ['final']).
cnf(c_0_1623,negated_conjecture,(c6_2(a231,a232)|c1_2(a158,a293)|c3_2(X1,X2)|c7_2(a158,a292)|c2_1(a158)|c5_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a158,a293)|~c5_1(X1)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1092, c_0_869]), ['final']).
cnf(c_0_1624,negated_conjecture,(c2_2(a217,a278)|c2_2(a294,a295)|c5_2(a294,a278)|c1_2(a294,a295)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c1_1(a294)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_501, c_0_1093]), ['final']).
cnf(c_0_1625,negated_conjecture,(c2_2(a217,a278)|c2_2(a204,a205)|c5_2(a204,a278)|c1_2(a204,a205)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c1_1(a204)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_501, c_0_1094]), ['final']).
cnf(c_0_1626,negated_conjecture,(c1_2(X1,X2)|c3_2(a268,a270)|c1_1(a256)|c1_1(a268)|c6_0|c2_1(X1)|ndr1_1(a231)|c4_0|c5_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1095, c_0_522]), ['final']).
cnf(c_0_1627,negated_conjecture,(c1_2(X1,X2)|c3_2(a268,a270)|c1_1(a268)|c6_0|c2_1(X1)|ndr1_1(a231)|ndr1_1(a256)|c4_0|c5_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1095, c_0_524]), ['final']).
cnf(c_0_1628,negated_conjecture,(c1_2(X1,X2)|c3_2(a268,a270)|c4_1(a231)|c1_1(a256)|c1_1(a268)|c6_0|c2_1(X1)|c4_0|c5_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1096, c_0_522]), ['final']).
cnf(c_0_1629,negated_conjecture,(c1_2(X1,X2)|c3_2(a268,a270)|c4_1(a231)|c1_1(a268)|c6_0|c2_1(X1)|ndr1_1(a256)|c4_0|c5_0|~c4_2(X1,X2)|~c7_2(X1,X2)|~c1_1(X1)|~c5_1(a268)|~ndr1_1(a268)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1096, c_0_524]), ['final']).
cnf(c_0_1630,negated_conjecture,(c2_2(a275,a278)|c2_2(X1,X2)|c1_2(X1,X2)|c4_2(a275,a293)|c3_2(a275,a278)|c7_2(a275,a292)|c4_1(X1)|c1_1(a275)|c8_0|c7_0|ssSkC25|c1_0|~c7_2(X1,X2)|~c7_1(a275)|~ndr1_1(a275)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1097]), ['final']).
cnf(c_0_1631,negated_conjecture,(c2_2(a218,a278)|c2_2(X1,X2)|c1_2(X1,X2)|c4_2(a218,a293)|c3_2(a218,a278)|c7_2(a218,a292)|c4_1(X1)|c1_1(a218)|c8_0|c7_0|ssSkC13|c1_0|~c7_2(X1,X2)|~c7_1(a218)|~ndr1_1(a218)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1098]), ['final']).
cnf(c_0_1632,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_878, c_0_522]), ['final']).
cnf(c_0_1633,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c1_1(a223)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ndr1_1(a256)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_878, c_0_524]), ['final']).
cnf(c_0_1634,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c5_1(a159)|c6_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_879, c_0_522]), ['final']).
cnf(c_0_1635,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c1_1(a223)|c5_1(a159)|c6_0|c2_1(a223)|ndr1_1(a296)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_879, c_0_524]), ['final']).
cnf(c_0_1636,negated_conjecture,(c2_2(a200,a278)|c2_2(X1,X2)|c1_2(X1,X2)|c4_2(a200,a293)|c3_2(a200,a278)|c7_2(a200,a292)|c4_1(X1)|c1_1(a200)|c8_0|c7_0|ssSkC10|c1_0|~c7_2(X1,X2)|~c7_1(a200)|~ndr1_1(a200)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_918]), ['final']).
cnf(c_0_1637,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_864, c_0_522]), ['final']).
cnf(c_0_1638,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c1_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a256)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_864, c_0_524]), ['final']).
cnf(c_0_1639,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_865, c_0_522]), ['final']).
cnf(c_0_1640,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a256)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_865, c_0_524]), ['final']).
cnf(c_0_1641,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_866, c_0_522]), ['final']).
cnf(c_0_1642,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c1_1(a223)|c6_0|c2_1(a223)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_866, c_0_524]), ['final']).
cnf(c_0_1643,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_867, c_0_522]), ['final']).
cnf(c_0_1644,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c2_1(a223)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_867, c_0_524]), ['final']).
cnf(c_0_1645,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c8_2(a160,a161)|c4_2(a168,a293)|c7_2(a168,a248)|c7_2(a168,a292)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1099, c_0_502]), ['final']).
cnf(c_0_1646,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c8_2(a160,a161)|c8_2(a168,a292)|c4_2(a168,a293)|c7_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1099, c_0_562]), ['final']).
cnf(c_0_1647,negated_conjecture,(c2_2(a182,a278)|c2_2(X1,X2)|c1_2(X1,X2)|c4_2(a182,a293)|c3_2(a182,a278)|c7_2(a182,a292)|c4_1(X1)|c1_1(a182)|c8_0|c7_0|c1_0|~c7_2(X1,X2)|~c7_1(a182)|~ndr1_1(a182)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1100]), ['final']).
cnf(c_0_1648,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_978]), ['final']).
cnf(c_0_1649,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_986]), ['final']).
cnf(c_0_1650,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c7_2(X1,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a160)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_990]), ['final']).
cnf(c_0_1651,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c3_2(X1,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a160)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_990]), ['final']).
cnf(c_0_1652,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c8_2(a160,a161)|c4_2(a168,a293)|c3_2(a168,a248)|c7_2(a168,a292)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1101, c_0_502]), ['final']).
cnf(c_0_1653,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c8_2(a160,a161)|c8_2(a168,a292)|c4_2(a168,a293)|c3_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1101, c_0_562]), ['final']).
cnf(c_0_1654,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c8_2(a160,a161)|c4_2(a168,a293)|c7_2(a168,a248)|c7_2(a168,a292)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1102, c_0_502]), ['final']).
cnf(c_0_1655,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c8_2(a160,a161)|c8_2(a168,a292)|c4_2(a168,a293)|c7_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1102, c_0_562]), ['final']).
cnf(c_0_1656,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c8_2(a160,a161)|c4_2(a168,a293)|c3_2(a168,a248)|c7_2(a168,a292)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1103, c_0_502]), ['final']).
cnf(c_0_1657,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c8_2(a160,a161)|c8_2(a168,a292)|c4_2(a168,a293)|c3_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1103, c_0_562]), ['final']).
cnf(c_0_1658,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c4_2(a168,a293)|c7_2(a168,a248)|c7_2(a168,a292)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1104, c_0_502]), ['final']).
cnf(c_0_1659,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c8_2(a168,a292)|c4_2(a168,a293)|c7_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1104, c_0_562]), ['final']).
cnf(c_0_1660,negated_conjecture,(c2_2(a221,a278)|c2_2(X1,X2)|c1_2(X1,X2)|c4_2(a221,a293)|c3_2(a221,a278)|c7_2(a221,a292)|c4_1(X1)|c1_1(a221)|c8_0|c7_0|ssSkC14|c1_0|~c7_2(X1,X2)|~c7_1(a221)|~ndr1_1(a221)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1105]), ['final']).
cnf(c_0_1661,negated_conjecture,(c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a181)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~c5_2(a223,a293)|~c1_2(a223,X1)|~c8_2(a223,X1)|~c3_2(a223,X1)), inference(spm,[status(thm)],[c_0_1106, c_0_584]), ['final']).
cnf(c_0_1662,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c4_2(a168,a293)|c3_2(a168,a248)|c7_2(a168,a292)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1107, c_0_502]), ['final']).
cnf(c_0_1663,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c8_2(a168,a292)|c4_2(a168,a293)|c3_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1107, c_0_562]), ['final']).
cnf(c_0_1664,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c4_2(a168,a293)|c7_2(a168,a248)|c7_2(a168,a292)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1108, c_0_502]), ['final']).
cnf(c_0_1665,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c8_2(a168,a292)|c4_2(a168,a293)|c7_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1108, c_0_562]), ['final']).
cnf(c_0_1666,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c4_2(a168,a293)|c3_2(a168,a248)|c7_2(a168,a292)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1109, c_0_502]), ['final']).
cnf(c_0_1667,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c8_2(a168,a292)|c4_2(a168,a293)|c3_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1109, c_0_562]), ['final']).
cnf(c_0_1668,negated_conjecture,(c2_2(a223,a293)|c5_2(a193,a278)|c4_2(a193,a293)|c7_2(a193,a292)|c7_2(a223,a292)|c5_1(a223)|c8_0|c7_0|ssSkC15|c1_0|~c3_1(a223)|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1110, c_0_734]), ['final']).
cnf(c_0_1669,negated_conjecture,(c2_2(a223,a293)|c5_2(a193,a278)|c8_2(a193,a292)|c4_2(a193,a293)|c7_2(a223,a292)|c5_1(a223)|c8_0|c7_0|ssSkC15|c1_0|~c3_1(a223)|~c7_2(a223,a293)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_1111, c_0_734]), ['final']).
cnf(c_0_1670,negated_conjecture,(ssSkC20|c2_2(a296,a297)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a296,a297)|c7_2(a217,a248)|c7_2(a168,a248)|c4_1(X1)|c1_1(a296)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0|~c7_2(X1,X2)|~c7_1(a296)|~ndr1_1(a296)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1112]), ['final']).
cnf(c_0_1671,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a223,a185)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_878, c_0_499]), ['final']).
cnf(c_0_1672,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c7_2(a223,a185)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_878, c_0_582]), ['final']).
cnf(c_0_1673,negated_conjecture,(c2_2(a223,a293)|c5_2(a223,a185)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_878, c_0_722]), ['final']).
cnf(c_0_1674,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c4_2(a158,a185)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_901, c_0_499]), ['final']).
cnf(c_0_1675,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c7_2(a158,a185)|c5_1(a159)|c6_0|c8_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_901, c_0_582]), ['final']).
cnf(c_0_1676,negated_conjecture,(c2_2(a158,a293)|c5_2(a158,a185)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_901, c_0_722]), ['final']).
cnf(c_0_1677,negated_conjecture,(ssSkC20|c2_2(a296,a297)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a296,a297)|c7_2(a168,a248)|c4_1(X1)|c1_1(a296)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ssSkC15|c1_0|~c7_2(X1,X2)|~c7_1(a296)|~ndr1_1(a296)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1113]), ['final']).
cnf(c_0_1678,negated_conjecture,(c4_2(a233,a185)|c6_0|c8_0|c2_1(a181)|c5_0|ssSkC17|c1_0|~c5_2(a233,a234)|~c1_2(a233,X1)|~c8_2(a233,X1)|~c3_2(a233,X1)|~c5_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_1114, c_0_499]), ['final']).
cnf(c_0_1679,negated_conjecture,(c7_2(a233,a185)|c6_0|c8_0|c2_1(a181)|c5_0|ssSkC17|c1_0|~c5_2(a233,a234)|~c1_2(a233,X1)|~c8_2(a233,X1)|~c3_2(a233,X1)|~c5_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_1114, c_0_582]), ['final']).
cnf(c_0_1680,negated_conjecture,(c5_2(a233,a185)|c6_0|c8_0|c2_1(a181)|c5_0|ssSkC17|c1_0|~c5_2(a233,a234)|~c1_2(a233,X1)|~c8_2(a233,X1)|~c3_2(a233,X1)|~c5_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_1114, c_0_722]), ['final']).
cnf(c_0_1681,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c4_2(a223,a185)|c7_2(a159,a292)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_864, c_0_499]), ['final']).
cnf(c_0_1682,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c7_2(a223,a185)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_864, c_0_582]), ['final']).
cnf(c_0_1683,negated_conjecture,(c2_2(a223,a293)|c5_2(a223,a185)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_864, c_0_722]), ['final']).
cnf(c_0_1684,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c4_2(a223,a185)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_865, c_0_499]), ['final']).
cnf(c_0_1685,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c7_2(a223,a185)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_865, c_0_582]), ['final']).
cnf(c_0_1686,negated_conjecture,(c2_2(a223,a293)|c5_2(a223,a185)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_865, c_0_722]), ['final']).
cnf(c_0_1687,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c4_2(a159,a293)|c4_2(a158,a185)|c7_2(a159,a292)|c7_2(a158,a292)|c6_0|c8_0|c2_1(a158)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_871, c_0_499]), ['final']).
cnf(c_0_1688,negated_conjecture,(ssSkC20|c2_2(a296,a297)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a296,a297)|c3_2(a217,a248)|c7_2(a168,a248)|c4_1(X1)|c1_1(a296)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0|~c7_2(X1,X2)|~c7_1(a296)|~ndr1_1(a296)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1115]), ['final']).
cnf(c_0_1689,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a158,a292)|c7_2(a158,a185)|c6_0|c8_0|c2_1(a158)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_871, c_0_582]), ['final']).
cnf(c_0_1690,negated_conjecture,(c2_2(a158,a293)|c5_2(a158,a185)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a158,a292)|c6_0|c8_0|c2_1(a158)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_871, c_0_722]), ['final']).
cnf(c_0_1691,negated_conjecture,(c2_2(a158,a293)|c8_2(a159,a292)|c8_2(a158,a293)|c4_2(a159,a293)|c4_2(a158,a185)|c7_2(a158,a292)|c6_0|c8_0|c2_1(a158)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_873, c_0_499]), ['final']).
cnf(c_0_1692,negated_conjecture,(c2_2(a158,a293)|c8_2(a159,a292)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a158,a292)|c7_2(a158,a185)|c6_0|c8_0|c2_1(a158)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_873, c_0_582]), ['final']).
cnf(c_0_1693,negated_conjecture,(c2_2(a158,a293)|c5_2(a158,a185)|c8_2(a159,a292)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a158,a292)|c6_0|c8_0|c2_1(a158)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_873, c_0_722]), ['final']).
cnf(c_0_1694,negated_conjecture,(ssSkC20|c2_2(a296,a297)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a296,a297)|c3_2(a168,a248)|c7_2(a217,a248)|c4_1(X1)|c1_1(a296)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0|~c7_2(X1,X2)|~c7_1(a296)|~ndr1_1(a296)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1116]), ['final']).
cnf(c_0_1695,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a160,a161)|c4_2(a215,a293)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1117, c_0_502]), ['final']).
cnf(c_0_1696,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a160,a161)|c8_2(a215,a292)|c4_2(a215,a293)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1117, c_0_562]), ['final']).
cnf(c_0_1697,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a160,a161)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1118, c_0_502]), ['final']).
cnf(c_0_1698,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a160,a161)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1118, c_0_562]), ['final']).
cnf(c_0_1699,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a160,a161)|c4_2(a215,a293)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1119, c_0_502]), ['final']).
cnf(c_0_1700,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a160,a161)|c8_2(a215,a292)|c4_2(a215,a293)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1119, c_0_562]), ['final']).
cnf(c_0_1701,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a160,a161)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1120, c_0_502]), ['final']).
cnf(c_0_1702,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a160,a161)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1120, c_0_562]), ['final']).
cnf(c_0_1703,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1121, c_0_502]), ['final']).
cnf(c_0_1704,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1121, c_0_562]), ['final']).
cnf(c_0_1705,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1122, c_0_502]), ['final']).
cnf(c_0_1706,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1122, c_0_562]), ['final']).
cnf(c_0_1707,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1123, c_0_502]), ['final']).
cnf(c_0_1708,negated_conjecture,(ssSkC20|c2_2(a296,a297)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a296,a297)|c3_2(a168,a248)|c4_1(X1)|c1_1(a296)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ssSkC15|c1_0|~c7_2(X1,X2)|~c7_1(a296)|~ndr1_1(a296)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1124]), ['final']).
cnf(c_0_1709,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1123, c_0_562]), ['final']).
cnf(c_0_1710,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1125, c_0_502]), ['final']).
cnf(c_0_1711,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1125, c_0_562]), ['final']).
cnf(c_0_1712,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c4_2(a168,a293)|c7_2(a168,a248)|c7_2(a168,a292)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1126, c_0_502]), ['final']).
cnf(c_0_1713,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c8_2(a168,a292)|c4_2(a168,a293)|c7_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1126, c_0_562]), ['final']).
cnf(c_0_1714,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c4_2(a168,a293)|c3_2(a168,a248)|c7_2(a168,a292)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1127, c_0_502]), ['final']).
cnf(c_0_1715,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a168,a278)|c8_2(a168,a292)|c4_2(a168,a293)|c3_2(a168,a248)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1127, c_0_562]), ['final']).
cnf(c_0_1716,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c4_2(a168,a293)|c7_2(a168,a248)|c7_2(a168,a292)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1128, c_0_502]), ['final']).
cnf(c_0_1717,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c8_2(a168,a292)|c4_2(a168,a293)|c7_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1128, c_0_562]), ['final']).
cnf(c_0_1718,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c4_2(a168,a293)|c3_2(a168,a248)|c7_2(a168,a292)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1129, c_0_502]), ['final']).
cnf(c_0_1719,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a168,a278)|c8_2(a168,a292)|c4_2(a168,a293)|c3_2(a168,a248)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1129, c_0_562]), ['final']).
cnf(c_0_1720,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_1130, c_0_522]), ['final']).
cnf(c_0_1721,negated_conjecture,(ssSkC20|c2_2(a296,a297)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a296,a297)|c3_2(a217,a248)|c4_1(X1)|c1_1(a296)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ssSkC15|c1_0|~c7_2(X1,X2)|~c7_1(a296)|~ndr1_1(a296)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1131]), ['final']).
cnf(c_0_1722,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|c2_1(X1)|ndr1_1(a160)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0|~c1_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_990]), ['final']).
cnf(c_0_1723,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c1_1(a223)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ndr1_1(a256)|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_1130, c_0_524]), ['final']).
cnf(c_0_1724,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c5_1(a159)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_1132, c_0_522]), ['final']).
cnf(c_0_1725,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c1_1(a223)|c5_1(a159)|c6_0|c2_1(a223)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_1132, c_0_524]), ['final']).
cnf(c_0_1726,negated_conjecture,(ssSkC20|c2_2(a296,a297)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a296,a297)|c7_2(a217,a248)|c4_1(X1)|c1_1(a296)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ssSkC15|c1_0|~c7_2(X1,X2)|~c7_1(a296)|~ndr1_1(a296)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1133]), ['final']).
cnf(c_0_1727,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c7_2(X1,a248)|c4_1(a294)|c5_1(a294)|c5_1(X2)|c8_1(a160)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_507]), ['final']).
cnf(c_0_1728,negated_conjecture,(ssSkC20|c2_2(a168,a248)|c4_1(a168)|c5_1(a168)|c5_1(X1)|c3_0|c8_1(a160)|c8_1(X2)|c2_1(a160)|ndr1_1(a287)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)|~c1_1(X2)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_486, c_0_1078]), ['final']).
cnf(c_0_1729,negated_conjecture,(c5_2(X1,X2)|c3_2(X1,X2)|c3_1(a233)|c3_1(X1)|c4_1(X1)|c2_1(a233)|ssSkC17|c1_0|~c8_2(a233,a234)|~c8_2(X1,X2)|~c3_2(a233,a234)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1134, c_0_63]), ['final']).
cnf(c_0_1730,negated_conjecture,(c1_1(a256)|c1_1(a233)|c6_0|c2_1(a181)|c5_0|ssSkC17|~c5_2(a233,a234)|~c1_2(a233,X1)|~c8_2(a233,X1)|~c3_2(a233,X1)|~c5_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_1114, c_0_522]), ['final']).
cnf(c_0_1731,negated_conjecture,(c1_1(a233)|c6_0|c2_1(a181)|ndr1_1(a256)|c5_0|ssSkC17|~c5_2(a233,a234)|~c1_2(a233,X1)|~c8_2(a233,X1)|~c3_2(a233,X1)|~c5_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_1114, c_0_524]), ['final']).
cnf(c_0_1732,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1135, c_0_590]), ['final']).
cnf(c_0_1733,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1136, c_0_590]), ['final']).
cnf(c_0_1734,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1137, c_0_584]), ['final']).
cnf(c_0_1735,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1138, c_0_584]), ['final']).
cnf(c_0_1736,negated_conjecture,(ssSkC20|c2_2(a296,a297)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a296,a297)|c3_2(a217,a248)|c3_2(a168,a248)|c4_1(X1)|c1_1(a296)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0|~c7_2(X1,X2)|~c7_1(a296)|~ndr1_1(a296)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1139]), ['final']).
cnf(c_0_1737,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c3_2(X1,a248)|c4_1(a294)|c5_1(a294)|c5_1(X2)|c8_1(a160)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_507]), ['final']).
cnf(c_0_1738,negated_conjecture,(c1_1(a256)|c1_1(a268)|c6_0|c2_1(a181)|c4_0|c5_0|~c5_2(a268,a270)|~c1_2(a268,X1)|~c8_2(a268,X1)|~c3_2(a268,X1)|~c5_1(a268)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_1140, c_0_522]), ['final']).
cnf(c_0_1739,negated_conjecture,(c1_1(a268)|c6_0|c2_1(a181)|ndr1_1(a256)|c4_0|c5_0|~c5_2(a268,a270)|~c1_2(a268,X1)|~c8_2(a268,X1)|~c3_2(a268,X1)|~c5_1(a268)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_1140, c_0_524]), ['final']).
cnf(c_0_1740,negated_conjecture,(ssSkC20|c2_2(a296,a297)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a296,a297)|c4_1(X1)|c1_1(a296)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a168)|ndr1_1(a223)|ssSkC15|c1_0|~c7_2(X1,X2)|~c7_1(a296)|~ndr1_1(a296)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1141]), ['final']).
cnf(c_0_1741,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c7_2(X1,a248)|c4_1(a204)|c5_1(a204)|c5_1(X2)|c8_1(a160)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_514]), ['final']).
cnf(c_0_1742,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c3_2(X1,a248)|c4_1(a204)|c5_1(a204)|c5_1(X2)|c8_1(a160)|c2_1(a160)|c2_1(X1)|ndr1_1(X2)|c1_0|~c3_1(X2)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_514]), ['final']).
cnf(c_0_1743,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1142, c_0_502]), ['final']).
cnf(c_0_1744,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1142, c_0_562]), ['final']).
cnf(c_0_1745,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a215,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1143, c_0_502]), ['final']).
cnf(c_0_1746,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1143, c_0_562]), ['final']).
cnf(c_0_1747,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c7_2(a215,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1144, c_0_502]), ['final']).
cnf(c_0_1748,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1144, c_0_562]), ['final']).
cnf(c_0_1749,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c4_2(a215,a293)|c3_2(a215,a248)|c7_2(a215,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1145, c_0_502]), ['final']).
cnf(c_0_1750,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a215,a278)|c8_2(a215,a292)|c4_2(a215,a293)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1145, c_0_562]), ['final']).
cnf(c_0_1751,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_975]), ['final']).
cnf(c_0_1752,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_975]), ['final']).
cnf(c_0_1753,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c4_2(a223,a185)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_1130, c_0_499]), ['final']).
cnf(c_0_1754,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c7_2(a223,a185)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_1130, c_0_582]), ['final']).
cnf(c_0_1755,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c5_2(a223,a185)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)), inference(spm,[status(thm)],[c_0_1130, c_0_722]), ['final']).
cnf(c_0_1756,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c4_1(a294)|c5_1(a294)|c5_1(X1)|c8_1(a160)|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_513, c_0_507]), ['final']).
cnf(c_0_1757,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_982]), ['final']).
cnf(c_0_1758,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_982]), ['final']).
cnf(c_0_1759,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c4_2(a158,a185)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a158)|c5_0|c1_0|~c5_1(a158)), inference(spm,[status(thm)],[c_0_1146, c_0_499]), ['final']).
cnf(c_0_1760,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c7_2(a158,a185)|c5_1(a159)|c6_0|c8_0|c2_1(a158)|c5_0|c1_0|~c5_1(a158)), inference(spm,[status(thm)],[c_0_1146, c_0_582]), ['final']).
cnf(c_0_1761,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c5_2(a158,a185)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c6_0|c8_0|c2_1(a158)|c5_0|c1_0|~c5_1(a158)), inference(spm,[status(thm)],[c_0_1146, c_0_722]), ['final']).
cnf(c_0_1762,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c4_1(a204)|c5_1(a204)|c5_1(X1)|c8_1(a160)|c2_1(a160)|c2_1(X2)|ndr1_1(X2)|ndr1_1(X1)|c1_0|~c3_1(X1)|~c1_1(X2)), inference(spm,[status(thm)],[c_0_513, c_0_514]), ['final']).
cnf(c_0_1763,negated_conjecture,(c2_2(a174,a185)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a174,a185)|c4_1(X1)|c1_1(a174)|c6_0|c8_0|ssSkC4|c1_0|~c7_2(X1,X2)|~c7_1(a174)|~ndr1_1(a174)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_1147]), ['final']).
cnf(c_0_1764,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1148, c_0_590]), ['final']).
cnf(c_0_1765,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1149, c_0_590]), ['final']).
cnf(c_0_1766,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1150, c_0_584]), ['final']).
cnf(c_0_1767,negated_conjecture,(c2_2(a223,a293)|c8_2(a159,a292)|c8_2(a223,a293)|c4_2(a159,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1151, c_0_584]), ['final']).
cnf(c_0_1768,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1152, c_0_869]), ['final']).
cnf(c_0_1769,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a159,a292)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1153, c_0_869]), ['final']).
cnf(c_0_1770,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c3_1(a159)|c1_1(a256)|c1_1(a241)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ndr1_1(a296)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1154, c_0_522]), ['final']).
cnf(c_0_1771,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c3_1(a159)|c1_1(a241)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ndr1_1(a296)|ndr1_1(a256)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1154, c_0_524]), ['final']).
cnf(c_0_1772,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c3_1(a159)|c1_1(a256)|c1_1(a217)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a296)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1155, c_0_522]), ['final']).
cnf(c_0_1773,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c3_1(a159)|c1_1(a217)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a296)|ndr1_1(a256)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1155, c_0_524]), ['final']).
cnf(c_0_1774,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c3_1(a159)|c1_1(a256)|c1_1(a241)|c5_1(a159)|c6_0|c2_1(a241)|ndr1_1(a296)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1156, c_0_522]), ['final']).
cnf(c_0_1775,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c3_1(a159)|c1_1(a241)|c5_1(a159)|c6_0|c2_1(a241)|ndr1_1(a296)|ndr1_1(a256)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1156, c_0_524]), ['final']).
cnf(c_0_1776,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c3_1(a159)|c1_1(a256)|c1_1(a217)|c5_1(a159)|c6_0|ssSkC22|c2_1(a217)|ndr1_1(a296)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1157, c_0_522]), ['final']).
cnf(c_0_1777,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c3_1(a159)|c1_1(a217)|c5_1(a159)|c6_0|ssSkC22|c2_1(a217)|ndr1_1(a296)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1157, c_0_524]), ['final']).
cnf(c_0_1778,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a256)|c1_1(a241)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1158, c_0_522]), ['final']).
cnf(c_0_1779,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a241)|c6_0|c8_0|c2_1(a241)|ndr1_1(a256)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1158, c_0_524]), ['final']).
cnf(c_0_1780,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a256)|c1_1(a217)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1159, c_0_522]), ['final']).
cnf(c_0_1781,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a217)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a256)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1159, c_0_524]), ['final']).
cnf(c_0_1782,negated_conjecture,(c2_2(a241,a243)|c8_2(a159,a292)|c8_2(a241,a243)|c4_2(a159,a293)|c1_1(a256)|c1_1(a241)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1160, c_0_522]), ['final']).
cnf(c_0_1783,negated_conjecture,(c2_2(a241,a243)|c8_2(a159,a292)|c8_2(a241,a243)|c4_2(a159,a293)|c1_1(a241)|c6_0|c8_0|c2_1(a241)|ndr1_1(a256)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1160, c_0_524]), ['final']).
cnf(c_0_1784,negated_conjecture,(c2_2(a217,a258)|c8_2(a159,a292)|c8_2(a217,a258)|c4_2(a159,a293)|c1_1(a256)|c1_1(a217)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1161, c_0_522]), ['final']).
cnf(c_0_1785,negated_conjecture,(c2_2(a217,a258)|c8_2(a159,a292)|c8_2(a217,a258)|c4_2(a159,a293)|c1_1(a217)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a256)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1161, c_0_524]), ['final']).
cnf(c_0_1786,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a256)|c1_1(a241)|c6_0|c2_1(a241)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1162, c_0_522]), ['final']).
cnf(c_0_1787,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a241)|c6_0|c2_1(a241)|ndr1_1(a256)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1162, c_0_524]), ['final']).
cnf(c_0_1788,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a256)|c1_1(a217)|c6_0|ssSkC22|c2_1(a217)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1163, c_0_522]), ['final']).
cnf(c_0_1789,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a217)|c6_0|ssSkC22|c2_1(a217)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1163, c_0_524]), ['final']).
cnf(c_0_1790,negated_conjecture,(c2_2(a241,a243)|c8_2(a159,a292)|c8_2(a241,a243)|c4_2(a159,a293)|c1_1(a256)|c1_1(a241)|c6_0|c2_1(a241)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1164, c_0_522]), ['final']).
cnf(c_0_1791,negated_conjecture,(c2_2(a241,a243)|c8_2(a159,a292)|c8_2(a241,a243)|c4_2(a159,a293)|c1_1(a241)|c6_0|c2_1(a241)|ndr1_1(a256)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1164, c_0_524]), ['final']).
cnf(c_0_1792,negated_conjecture,(c2_2(a217,a258)|c8_2(a159,a292)|c8_2(a217,a258)|c4_2(a159,a293)|c1_1(a256)|c1_1(a217)|c6_0|ssSkC22|c2_1(a217)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1165, c_0_522]), ['final']).
cnf(c_0_1793,negated_conjecture,(c2_2(a217,a258)|c8_2(a159,a292)|c8_2(a217,a258)|c4_2(a159,a293)|c1_1(a217)|c6_0|ssSkC22|c2_1(a217)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1165, c_0_524]), ['final']).
cnf(c_0_1794,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_978]), ['final']).
cnf(c_0_1795,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c8_2(a160,a161)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_978]), ['final']).
cnf(c_0_1796,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1130, c_0_644]), ['final']).
cnf(c_0_1797,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c5_1(a223)|c6_0|c2_1(a223)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1132, c_0_644]), ['final']).
cnf(c_0_1798,negated_conjecture,(ssSkC20|c2_2(a168,a248)|c8_2(a160,a161)|c4_1(a168)|c5_1(a168)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_496, c_0_1064]), ['final']).
cnf(c_0_1799,negated_conjecture,(c2_2(a233,a234)|c5_2(a193,a278)|c4_2(a193,a293)|c7_2(a193,a292)|c5_1(a233)|c8_0|c7_0|ssSkC17|c1_0|~c3_1(a233)|~c7_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_1110, c_0_39]), ['final']).
cnf(c_0_1800,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_986]), ['final']).
cnf(c_0_1801,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_986]), ['final']).
cnf(c_0_1802,negated_conjecture,(c2_2(a233,a234)|c5_2(a193,a278)|c8_2(a193,a292)|c4_2(a193,a293)|c5_1(a233)|c8_0|c7_0|ssSkC17|c1_0|~c3_1(a233)|~c7_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_1111, c_0_39]), ['final']).
cnf(c_0_1803,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|ndr1_1(a160)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_989]), ['final']).
cnf(c_0_1804,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c7_2(a168,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ssSkC22|ndr1_1(a160)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_989]), ['final']).
cnf(c_0_1805,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1166, c_0_869]), ['final']).
cnf(c_0_1806,negated_conjecture,(c4_2(a268,a270)|c4_2(a268,X1)|c7_2(a268,a270)|c7_1(a268)|c6_0|c7_0|ndr1_1(a208)|c4_0|c5_0|~c1_2(a268,X1)|~c7_2(a268,X1)), inference(spm,[status(thm)],[c_0_1167, c_0_127]), ['final']).
cnf(c_0_1807,negated_conjecture,(c4_2(a268,a270)|c4_2(a268,X1)|c7_2(a268,a270)|c7_1(a268)|c6_0|c8_1(a208)|c7_0|c4_0|c5_0|~c1_2(a268,X1)|~c7_2(a268,X1)), inference(spm,[status(thm)],[c_0_1168, c_0_127]), ['final']).
cnf(c_0_1808,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|c6_0|c7_0|~c6_2(X1,X2)|~c1_2(X1,X3)|~c3_2(a208,a209)|~c7_2(X1,X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1169, c_0_724]), ['final']).
cnf(c_0_1809,negated_conjecture,(c2_2(a158,a293)|c8_2(a159,a292)|c8_2(a158,a293)|c4_2(a159,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1170, c_0_869]), ['final']).
cnf(c_0_1810,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a241,a185)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ndr1_1(a296)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1154, c_0_499]), ['final']).
cnf(c_0_1811,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a168,a248)|c4_1(a168)|c5_1(a168)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_497, c_0_1064]), ['final']).
cnf(c_0_1812,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c3_1(a159)|c7_2(a241,a185)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ndr1_1(a296)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1154, c_0_582]), ['final']).
cnf(c_0_1813,negated_conjecture,(c2_2(a241,a243)|c5_2(a241,a185)|c8_2(a241,a243)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ndr1_1(a296)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1154, c_0_722]), ['final']).
cnf(c_0_1814,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a217,a185)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a296)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1155, c_0_499]), ['final']).
cnf(c_0_1815,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c3_1(a159)|c7_2(a217,a185)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a296)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1155, c_0_582]), ['final']).
cnf(c_0_1816,negated_conjecture,(c2_2(a217,a258)|c5_2(a217,a185)|c8_2(a217,a258)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a296)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1155, c_0_722]), ['final']).
cnf(c_0_1817,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|c6_0|c7_0|~c6_2(X1,X2)|~c2_2(a208,a209)|~c1_2(X1,X3)|~c7_2(X1,X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1171, c_0_724]), ['final']).
cnf(c_0_1818,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c4_2(a197,a185)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a197)|ndr1_1(a296)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1172, c_0_499]), ['final']).
cnf(c_0_1819,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c3_1(a159)|c1_1(a256)|c1_1(a197)|c5_1(a159)|c6_0|c2_1(a197)|ndr1_1(a296)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1172, c_0_522]), ['final']).
cnf(c_0_1820,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c3_1(a159)|c1_1(a197)|c5_1(a159)|c6_0|c2_1(a197)|ndr1_1(a296)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1172, c_0_524]), ['final']).
cnf(c_0_1821,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c3_1(a159)|c7_2(a197,a185)|c5_1(a159)|c6_0|c8_0|c2_1(a197)|ndr1_1(a296)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1172, c_0_582]), ['final']).
cnf(c_0_1822,negated_conjecture,(c2_2(a197,a198)|c5_2(a197,a185)|c8_2(a197,a198)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a197)|ndr1_1(a296)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1172, c_0_722]), ['final']).
cnf(c_0_1823,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a159,a293)|c4_2(a241,a185)|c7_2(a159,a292)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1158, c_0_499]), ['final']).
cnf(c_0_1824,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a241,a185)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1158, c_0_582]), ['final']).
cnf(c_0_1825,negated_conjecture,(c2_2(a241,a243)|c5_2(a241,a185)|c8_2(a241,a243)|c4_2(a159,a293)|c7_2(a159,a292)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1158, c_0_722]), ['final']).
cnf(c_0_1826,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a159,a293)|c4_2(a217,a185)|c7_2(a159,a292)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1159, c_0_499]), ['final']).
cnf(c_0_1827,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a217,a185)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1159, c_0_582]), ['final']).
cnf(c_0_1828,negated_conjecture,(c2_2(a217,a258)|c5_2(a217,a185)|c8_2(a217,a258)|c4_2(a159,a293)|c7_2(a159,a292)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1159, c_0_722]), ['final']).
cnf(c_0_1829,negated_conjecture,(c2_2(a241,a243)|c8_2(a159,a292)|c8_2(a241,a243)|c4_2(a159,a293)|c4_2(a241,a185)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1160, c_0_499]), ['final']).
cnf(c_0_1830,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c1_1(a294)|c5_1(a294)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1173]), ['final']).
cnf(c_0_1831,negated_conjecture,(c2_2(a241,a243)|c8_2(a159,a292)|c8_2(a241,a243)|c4_2(a159,a293)|c7_2(a241,a185)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1160, c_0_582]), ['final']).
cnf(c_0_1832,negated_conjecture,(c2_2(a241,a243)|c5_2(a241,a185)|c8_2(a159,a292)|c8_2(a241,a243)|c4_2(a159,a293)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1160, c_0_722]), ['final']).
cnf(c_0_1833,negated_conjecture,(c2_2(a217,a258)|c8_2(a159,a292)|c8_2(a217,a258)|c4_2(a159,a293)|c4_2(a217,a185)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1161, c_0_499]), ['final']).
cnf(c_0_1834,negated_conjecture,(c2_2(a217,a258)|c8_2(a159,a292)|c8_2(a217,a258)|c4_2(a159,a293)|c7_2(a217,a185)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1161, c_0_582]), ['final']).
cnf(c_0_1835,negated_conjecture,(c2_2(a217,a258)|c5_2(a217,a185)|c8_2(a159,a292)|c8_2(a217,a258)|c4_2(a159,a293)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1161, c_0_722]), ['final']).
cnf(c_0_1836,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c5_1(a158)|c6_0|c2_1(a158)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1146, c_0_872]), ['final']).
cnf(c_0_1837,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1174, c_0_590]), ['final']).
cnf(c_0_1838,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c3_1(a159)|c7_2(a223,a292)|c5_1(a159)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1175, c_0_584]), ['final']).
cnf(c_0_1839,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c8_2(a160,a161)|c4_2(a197,a293)|c7_2(a197,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1176]), ['final']).
cnf(c_0_1840,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c8_2(a160,a161)|c8_2(a197,a292)|c4_2(a197,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1177]), ['final']).
cnf(c_0_1841,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c8_2(a160,a161)|c4_2(a162,a293)|c7_2(a162,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1178]), ['final']).
cnf(c_0_1842,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c8_2(a160,a161)|c8_2(a162,a292)|c4_2(a162,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1179]), ['final']).
cnf(c_0_1843,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c8_2(a160,a161)|c4_2(a288,a293)|c7_2(a288,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1180]), ['final']).
cnf(c_0_1844,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c8_2(a160,a161)|c8_2(a288,a292)|c4_2(a288,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1181]), ['final']).
cnf(c_0_1845,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c8_2(a160,a161)|c4_2(a253,a293)|c7_2(a253,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1182]), ['final']).
cnf(c_0_1846,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c8_2(a160,a161)|c8_2(a253,a292)|c4_2(a253,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1183]), ['final']).
cnf(c_0_1847,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c4_2(a159,a293)|c4_2(a197,a185)|c7_2(a159,a292)|c6_0|c8_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1184, c_0_499]), ['final']).
cnf(c_0_1848,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a256)|c1_1(a197)|c6_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1184, c_0_522]), ['final']).
cnf(c_0_1849,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c4_2(a159,a293)|c7_2(a159,a292)|c1_1(a197)|c6_0|c2_1(a197)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1184, c_0_524]), ['final']).
cnf(c_0_1850,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c4_2(a159,a293)|c7_2(a159,a292)|c7_2(a197,a185)|c6_0|c8_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1184, c_0_582]), ['final']).
cnf(c_0_1851,negated_conjecture,(c2_2(a197,a198)|c5_2(a197,a185)|c8_2(a197,a198)|c4_2(a159,a293)|c7_2(a159,a292)|c6_0|c8_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1184, c_0_722]), ['final']).
cnf(c_0_1852,negated_conjecture,(c2_2(a197,a198)|c8_2(a159,a292)|c8_2(a197,a198)|c4_2(a159,a293)|c4_2(a197,a185)|c6_0|c8_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1185, c_0_499]), ['final']).
cnf(c_0_1853,negated_conjecture,(c2_2(a197,a198)|c8_2(a159,a292)|c8_2(a197,a198)|c4_2(a159,a293)|c1_1(a256)|c1_1(a197)|c6_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1185, c_0_522]), ['final']).
cnf(c_0_1854,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c1_1(a204)|c5_1(a204)|c5_1(X1)|c8_0|c7_0|c2_1(a160)|ndr1_1(a160)|ndr1_1(X1)|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_1186]), ['final']).
cnf(c_0_1855,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c8_2(a160,a161)|c4_2(a197,a293)|c7_2(a197,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1187]), ['final']).
cnf(c_0_1856,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c8_2(a160,a161)|c8_2(a197,a292)|c4_2(a197,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1188]), ['final']).
cnf(c_0_1857,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c8_2(a160,a161)|c4_2(a162,a293)|c7_2(a162,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1189]), ['final']).
cnf(c_0_1858,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c8_2(a160,a161)|c8_2(a162,a292)|c4_2(a162,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1190]), ['final']).
cnf(c_0_1859,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c8_2(a160,a161)|c4_2(a288,a293)|c7_2(a288,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1191]), ['final']).
cnf(c_0_1860,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c8_2(a160,a161)|c8_2(a288,a292)|c4_2(a288,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1192]), ['final']).
cnf(c_0_1861,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c8_2(a160,a161)|c4_2(a253,a293)|c7_2(a253,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1193]), ['final']).
cnf(c_0_1862,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c8_2(a160,a161)|c8_2(a253,a292)|c4_2(a253,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1194]), ['final']).
cnf(c_0_1863,negated_conjecture,(c2_2(a197,a198)|c8_2(a159,a292)|c8_2(a197,a198)|c4_2(a159,a293)|c1_1(a197)|c6_0|c2_1(a197)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1185, c_0_524]), ['final']).
cnf(c_0_1864,negated_conjecture,(c2_2(a197,a198)|c8_2(a159,a292)|c8_2(a197,a198)|c4_2(a159,a293)|c7_2(a197,a185)|c6_0|c8_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1185, c_0_582]), ['final']).
cnf(c_0_1865,negated_conjecture,(c2_2(a197,a198)|c5_2(a197,a185)|c8_2(a159,a292)|c8_2(a197,a198)|c4_2(a159,a293)|c6_0|c8_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1185, c_0_722]), ['final']).
cnf(c_0_1866,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_911, c_0_522]), ['final']).
cnf(c_0_1867,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a159)|ndr1_1(a256)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_911, c_0_524]), ['final']).
cnf(c_0_1868,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1195]), ['final']).
cnf(c_0_1869,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1196]), ['final']).
cnf(c_0_1870,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a162,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1197]), ['final']).
cnf(c_0_1871,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1198]), ['final']).
cnf(c_0_1872,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a288,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1199]), ['final']).
cnf(c_0_1873,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1200]), ['final']).
cnf(c_0_1874,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1201]), ['final']).
cnf(c_0_1875,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1202]), ['final']).
cnf(c_0_1876,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_912, c_0_522]), ['final']).
cnf(c_0_1877,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c2_1(a223)|ndr1_1(a159)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_912, c_0_524]), ['final']).
cnf(c_0_1878,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c8_0|c2_1(a159)|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_913, c_0_522]), ['final']).
cnf(c_0_1879,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c8_0|c2_1(a159)|c2_1(a223)|ndr1_1(a256)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_913, c_0_524]), ['final']).
cnf(c_0_1880,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1203]), ['final']).
cnf(c_0_1881,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1204]), ['final']).
cnf(c_0_1882,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a162,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1205]), ['final']).
cnf(c_0_1883,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1206]), ['final']).
cnf(c_0_1884,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a288,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1207]), ['final']).
cnf(c_0_1885,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1208]), ['final']).
cnf(c_0_1886,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1209]), ['final']).
cnf(c_0_1887,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1210]), ['final']).
cnf(c_0_1888,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c1_1(a256)|c1_1(a223)|c6_0|c2_1(a159)|c2_1(a223)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_914, c_0_522]), ['final']).
cnf(c_0_1889,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c1_1(a223)|c6_0|c2_1(a159)|c2_1(a223)|ndr1_1(a256)|ssSkC15|c5_0|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_914, c_0_524]), ['final']).
cnf(c_0_1890,negated_conjecture,(c2_2(a159,X1)|c2_2(X2,X3)|c8_2(X2,X3)|c8_0|c2_1(X2)|c1_0|~c8_2(a159,X1)|~c4_2(a159,X1)|~c4_2(X2,X3)|~c3_1(X2)|~ndr1_1(a159)|~ndr1_1(X2)), inference(spm,[status(thm)],[c_0_1211, c_0_478]), ['final']).
cnf(c_0_1891,negated_conjecture,(c2_2(a159,X1)|c2_2(X2,X3)|c8_2(X2,X3)|c2_1(X2)|c5_0|c1_0|~c8_2(a159,X1)|~c4_2(a159,X1)|~c4_2(X2,X3)|~c3_1(X2)|~ndr1_1(a159)|~ndr1_1(X2)), inference(spm,[status(thm)],[c_0_1211, c_0_20]), ['final']).
cnf(c_0_1892,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c3_1(a159)|c1_1(a256)|c1_1(a241)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)), inference(spm,[status(thm)],[c_0_1212, c_0_522]), ['final']).
cnf(c_0_1893,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|c6_0|c7_0|~c6_2(a208,a209)|~c6_2(X1,X2)|~c1_2(X1,X3)|~c7_2(X1,X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1213, c_0_724]), ['final']).
cnf(c_0_1894,negated_conjecture,(c4_2(a233,a234)|c4_2(a233,X1)|c7_2(a233,a234)|c7_1(a233)|c6_0|c7_0|ndr1_1(a208)|ssSkC17|~c1_2(a233,X1)|~c7_2(a233,X1)), inference(spm,[status(thm)],[c_0_1214, c_0_63]), ['final']).
cnf(c_0_1895,negated_conjecture,(c4_2(a233,a234)|c4_2(a233,X1)|c7_2(a233,a234)|c7_1(a233)|c6_0|c8_1(a208)|c7_0|ssSkC17|~c1_2(a233,X1)|~c7_2(a233,X1)), inference(spm,[status(thm)],[c_0_1215, c_0_63]), ['final']).
cnf(c_0_1896,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c3_1(a159)|c1_1(a241)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ndr1_1(a256)|ssSkC19|c1_0|~c5_1(a241)), inference(spm,[status(thm)],[c_0_1212, c_0_524]), ['final']).
cnf(c_0_1897,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c3_1(a159)|c1_1(a256)|c1_1(a217)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)), inference(spm,[status(thm)],[c_0_1216, c_0_522]), ['final']).
cnf(c_0_1898,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a160)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC27|c1_0|~c1_1(a288)), inference(spm,[status(thm)],[c_0_86, c_0_990]), ['final']).
cnf(c_0_1899,negated_conjecture,(ssSkC20|c2_2(a217,a248)|c4_1(a217)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a160)|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|ssSkC21|c1_0|~c1_1(a253)), inference(spm,[status(thm)],[c_0_70, c_0_990]), ['final']).
cnf(c_0_1900,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c3_1(a159)|c1_1(a217)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a256)|c1_0|~c5_1(a217)), inference(spm,[status(thm)],[c_0_1216, c_0_524]), ['final']).
cnf(c_0_1901,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c3_1(a159)|c1_1(a256)|c1_1(a241)|c5_1(a159)|c6_0|c2_1(a241)|ssSkC19|c5_0|c1_0|~c5_1(a241)), inference(spm,[status(thm)],[c_0_1217, c_0_522]), ['final']).
cnf(c_0_1902,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c3_1(a159)|c1_1(a241)|c5_1(a159)|c6_0|c2_1(a241)|ndr1_1(a256)|ssSkC19|c5_0|c1_0|~c5_1(a241)), inference(spm,[status(thm)],[c_0_1217, c_0_524]), ['final']).
cnf(c_0_1903,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c3_1(a159)|c1_1(a256)|c1_1(a217)|c5_1(a159)|c6_0|ssSkC22|c2_1(a217)|c5_0|c1_0|~c5_1(a217)), inference(spm,[status(thm)],[c_0_1218, c_0_522]), ['final']).
cnf(c_0_1904,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c3_1(a159)|c1_1(a217)|c5_1(a159)|c6_0|ssSkC22|c2_1(a217)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a217)), inference(spm,[status(thm)],[c_0_1218, c_0_524]), ['final']).
cnf(c_0_1905,negated_conjecture,(ssSkC20|c2_2(a168,a248)|c4_1(a168)|c5_1(a168)|c3_0|c6_0|c8_1(X1)|c2_0|ndr1_1(a160)|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0|~c1_1(X1)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_498, c_0_1064]), ['final']).
cnf(c_0_1906,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c3_1(a159)|c7_2(a158,a292)|c5_1(a159)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1219, c_0_869]), ['final']).
cnf(c_0_1907,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a223,a185)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_911, c_0_499]), ['final']).
cnf(c_0_1908,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c7_2(a223,a185)|c6_0|c8_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_911, c_0_582]), ['final']).
cnf(c_0_1909,negated_conjecture,(c2_2(a223,a293)|c5_2(a223,a185)|c8_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_911, c_0_722]), ['final']).
cnf(c_0_1910,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c4_2(a223,a185)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a159)|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_913, c_0_499]), ['final']).
cnf(c_0_1911,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c7_2(a223,a185)|c6_0|c8_0|c2_1(a159)|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_913, c_0_582]), ['final']).
cnf(c_0_1912,negated_conjecture,(c2_2(a223,a293)|c5_2(a223,a185)|c8_2(a223,a293)|c7_2(a223,a292)|c6_0|c8_0|c2_1(a159)|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_913, c_0_722]), ['final']).
cnf(c_0_1913,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c4_2(a158,a185)|c7_2(a158,a292)|c6_0|c8_0|c2_1(a158)|ndr1_1(a159)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_919, c_0_499]), ['final']).
cnf(c_0_1914,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c7_2(a158,a292)|c7_2(a158,a185)|c6_0|c8_0|c2_1(a158)|ndr1_1(a159)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_919, c_0_582]), ['final']).
cnf(c_0_1915,negated_conjecture,(c2_2(a158,a293)|c5_2(a158,a185)|c8_2(a158,a293)|c7_2(a158,a292)|c6_0|c8_0|c2_1(a158)|ndr1_1(a159)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_919, c_0_722]), ['final']).
cnf(c_0_1916,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c4_2(a158,a185)|c7_2(a158,a292)|c6_0|c8_0|c2_1(a159)|c2_1(a158)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_920, c_0_499]), ['final']).
cnf(c_0_1917,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c7_2(a158,a292)|c7_2(a158,a185)|c6_0|c8_0|c2_1(a159)|c2_1(a158)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_920, c_0_582]), ['final']).
cnf(c_0_1918,negated_conjecture,(c2_2(a158,a293)|c5_2(a158,a185)|c8_2(a158,a293)|c7_2(a158,a292)|c6_0|c8_0|c2_1(a159)|c2_1(a158)|c5_0|c1_0|~c5_1(a158)|~ndr1_1(a158)), inference(spm,[status(thm)],[c_0_920, c_0_722]), ['final']).
cnf(c_0_1919,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c4_2(a241,a185)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)), inference(spm,[status(thm)],[c_0_1212, c_0_499]), ['final']).
cnf(c_0_1920,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1220]), ['final']).
cnf(c_0_1921,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c8_2(a241,a243)|c3_1(a159)|c7_2(a241,a185)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)), inference(spm,[status(thm)],[c_0_1212, c_0_582]), ['final']).
cnf(c_0_1922,negated_conjecture,(c2_2(a241,a243)|c5_2(a296,a297)|c5_2(a241,a185)|c8_2(a241,a243)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)), inference(spm,[status(thm)],[c_0_1212, c_0_722]), ['final']).
cnf(c_0_1923,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c4_2(a217,a185)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)), inference(spm,[status(thm)],[c_0_1216, c_0_499]), ['final']).
cnf(c_0_1924,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c8_2(a217,a258)|c3_1(a159)|c7_2(a217,a185)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)), inference(spm,[status(thm)],[c_0_1216, c_0_582]), ['final']).
cnf(c_0_1925,negated_conjecture,(c2_2(a217,a258)|c5_2(a296,a297)|c5_2(a217,a185)|c8_2(a217,a258)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|ssSkC22|c2_1(a217)|c1_0|~c5_1(a217)), inference(spm,[status(thm)],[c_0_1216, c_0_722]), ['final']).
cnf(c_0_1926,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1221]), ['final']).
cnf(c_0_1927,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a162,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1222]), ['final']).
cnf(c_0_1928,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1223]), ['final']).
cnf(c_0_1929,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a288,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1224]), ['final']).
cnf(c_0_1930,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1225]), ['final']).
cnf(c_0_1931,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1226]), ['final']).
cnf(c_0_1932,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1227]), ['final']).
cnf(c_0_1933,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c4_2(a197,a293)|c7_2(a197,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1228]), ['final']).
cnf(c_0_1934,negated_conjecture,(c2_2(a197,a198)|c5_2(a296,a297)|c8_2(a197,a198)|c4_2(a197,a185)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)), inference(spm,[status(thm)],[c_0_1229, c_0_499]), ['final']).
cnf(c_0_1935,negated_conjecture,(c2_2(a197,a198)|c5_2(a296,a297)|c8_2(a197,a198)|c3_1(a159)|c1_1(a256)|c1_1(a197)|c5_1(a159)|c6_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)), inference(spm,[status(thm)],[c_0_1229, c_0_522]), ['final']).
cnf(c_0_1936,negated_conjecture,(c2_2(a197,a198)|c5_2(a296,a297)|c8_2(a197,a198)|c3_1(a159)|c1_1(a197)|c5_1(a159)|c6_0|c2_1(a197)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a197)), inference(spm,[status(thm)],[c_0_1229, c_0_524]), ['final']).
cnf(c_0_1937,negated_conjecture,(c2_2(a197,a198)|c5_2(a296,a297)|c8_2(a197,a198)|c3_1(a159)|c7_2(a197,a185)|c5_1(a159)|c6_0|c8_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)), inference(spm,[status(thm)],[c_0_1229, c_0_582]), ['final']).
cnf(c_0_1938,negated_conjecture,(c2_2(a197,a198)|c5_2(a296,a297)|c5_2(a197,a185)|c8_2(a197,a198)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a197)|c5_0|c1_0|~c5_1(a197)), inference(spm,[status(thm)],[c_0_1229, c_0_722]), ['final']).
cnf(c_0_1939,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a197,a278)|c8_2(a197,a292)|c4_2(a197,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1230]), ['final']).
cnf(c_0_1940,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c4_2(a162,a293)|c7_2(a162,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1231]), ['final']).
cnf(c_0_1941,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a162,a278)|c8_2(a162,a292)|c4_2(a162,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1232]), ['final']).
cnf(c_0_1942,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c4_2(a288,a293)|c7_2(a288,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1233]), ['final']).
cnf(c_0_1943,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a288,a278)|c8_2(a288,a292)|c4_2(a288,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC27|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1234]), ['final']).
cnf(c_0_1944,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c4_2(a253,a293)|c7_2(a253,a292)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1235]), ['final']).
cnf(c_0_1945,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c5_2(a253,a278)|c8_2(a253,a292)|c4_2(a253,a293)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1236]), ['final']).
cnf(c_0_1946,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1237, c_0_590]), ['final']).
cnf(c_0_1947,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a159)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1238, c_0_584]), ['final']).
cnf(c_0_1948,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a159)|c2_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1239, c_0_590]), ['final']).
cnf(c_0_1949,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a159)|c2_1(a223)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1240, c_0_584]), ['final']).
cnf(c_0_1950,negated_conjecture,(c4_2(a275,a278)|c4_2(a275,a293)|c3_2(a275,a278)|c7_2(a275,a292)|c3_0|c8_0|c7_0|ssSkC25|c1_0|~c6_1(a275)|~c8_1(a275)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_917, c_0_1097]), ['final']).
cnf(c_0_1951,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|c6_0|ssSkC15|c5_0|c1_0|~c7_2(a223,a293)), inference(spm,[status(thm)],[c_0_1241, c_0_584]), ['final']).
cnf(c_0_1952,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|ssSkC15|c1_0|~c7_2(a223,a293)), inference(spm,[status(thm)],[c_0_1242, c_0_590]), ['final']).
cnf(c_0_1953,negated_conjecture,(c4_2(a218,a278)|c4_2(a218,a293)|c3_2(a218,a278)|c7_2(a218,a292)|c3_0|c8_0|c7_0|ssSkC13|c1_0|~c6_1(a218)|~c8_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_917, c_0_1098]), ['final']).
cnf(c_0_1954,negated_conjecture,(c2_2(a218,a219)|c5_2(a296,a297)|c8_2(a218,a219)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a218)|ssSkC13|c1_0|~c3_1(a218)), inference(spm,[status(thm)],[c_0_1243, c_0_154]), ['final']).
cnf(c_0_1955,negated_conjecture,(c4_2(a200,a278)|c4_2(a200,a293)|c3_2(a200,a278)|c7_2(a200,a292)|c3_0|c8_0|c7_0|ssSkC10|c1_0|~c8_1(a200)|~ndr1_1(a200)), inference(spm,[status(thm)],[c_0_1244, c_0_155]), ['final']).
cnf(c_0_1956,negated_conjecture,(c2_2(a218,a219)|c5_2(a296,a297)|c8_2(a218,a219)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(a218)|ssSkC13|c5_0|c1_0|~c3_1(a218)), inference(spm,[status(thm)],[c_0_1245, c_0_154]), ['final']).
cnf(c_0_1957,negated_conjecture,(c4_2(a221,a278)|c4_2(a221,a293)|c3_2(a221,a278)|c7_2(a221,a292)|c3_0|c8_0|c7_0|ssSkC14|c1_0|~c6_1(a221)|~c8_1(a221)|~ndr1_1(a221)), inference(spm,[status(thm)],[c_0_917, c_0_1105]), ['final']).
cnf(c_0_1958,negated_conjecture,(c2_2(a241,a243)|c2_2(X1,X2)|c1_2(X1,X2)|c3_2(a241,a243)|c4_1(X1)|c1_1(a241)|ssSkC19|~c7_2(X1,X2)|~c7_1(a241)|~ndr1_1(a241)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_527, c_0_79]), ['final']).
cnf(c_0_1959,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c4_2(a215,a258)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1117, c_0_1058]), ['final']).
cnf(c_0_1960,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c3_2(a215,a258)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1117, c_0_1059]), ['final']).
cnf(c_0_1961,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c4_2(a215,a258)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1118, c_0_1058]), ['final']).
cnf(c_0_1962,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c3_2(a215,a248)|c3_2(a215,a258)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1118, c_0_1059]), ['final']).
cnf(c_0_1963,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c4_2(a215,a258)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1119, c_0_1058]), ['final']).
cnf(c_0_1964,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c3_2(a215,a258)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1119, c_0_1059]), ['final']).
cnf(c_0_1965,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c4_2(a215,a258)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1120, c_0_1058]), ['final']).
cnf(c_0_1966,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c3_2(a215,a248)|c3_2(a215,a258)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1120, c_0_1059]), ['final']).
cnf(c_0_1967,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c4_2(a215,a258)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1121, c_0_1058]), ['final']).
cnf(c_0_1968,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a215,a258)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1121, c_0_1059]), ['final']).
cnf(c_0_1969,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c4_2(a215,a258)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1122, c_0_1058]), ['final']).
cnf(c_0_1970,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a215,a248)|c3_2(a215,a258)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1122, c_0_1059]), ['final']).
cnf(c_0_1971,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c4_2(a215,a258)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1123, c_0_1058]), ['final']).
cnf(c_0_1972,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a215,a258)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1123, c_0_1059]), ['final']).
cnf(c_0_1973,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c4_2(a215,a258)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1125, c_0_1058]), ['final']).
cnf(c_0_1974,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a215,a248)|c3_2(a215,a258)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|c1_0), inference(spm,[status(thm)],[c_0_1125, c_0_1059]), ['final']).
cnf(c_0_1975,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a223)|ndr1_1(a159)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1246, c_0_590]), ['final']).
cnf(c_0_1976,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a223)|ndr1_1(a159)|ndr1_1(a296)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1247, c_0_584]), ['final']).
cnf(c_0_1977,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|c2_1(a159)|c2_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1248, c_0_590]), ['final']).
cnf(c_0_1978,negated_conjecture,(c2_2(a223,a293)|c8_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c2_1(a159)|c2_1(a223)|ndr1_1(a296)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1249, c_0_584]), ['final']).
cnf(c_0_1979,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a159)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1250, c_0_869]), ['final']).
cnf(c_0_1980,negated_conjecture,(c2_2(a158,a293)|c5_2(a296,a297)|c8_2(a158,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a159)|c2_1(a158)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1251, c_0_869]), ['final']).
cnf(c_0_1981,negated_conjecture,(c4_2(a217,a278)|c4_2(a217,a293)|c3_2(a217,a278)|c7_2(a217,a292)|c3_0|c8_0|c7_0|c1_0|~c6_1(a217)|~c8_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_917, c_0_528]), ['final']).
cnf(c_0_1982,negated_conjecture,(c4_2(a182,a278)|c4_2(a182,a293)|c3_2(a182,a278)|c7_2(a182,a292)|c3_0|c8_0|c7_0|c1_0|~c6_1(a182)|~c8_1(a182)|~ndr1_1(a182)), inference(spm,[status(thm)],[c_0_917, c_0_1100]), ['final']).
cnf(c_0_1983,negated_conjecture,(c2_2(a223,a293)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|c6_0|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~c7_2(a223,a293)), inference(spm,[status(thm)],[c_0_1252, c_0_584]), ['final']).
cnf(c_0_1984,negated_conjecture,(c2_2(a223,a293)|c7_2(a193,a194)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|ndr1_1(a296)|ssSkC15|c1_0|~c7_2(a223,a293)), inference(spm,[status(thm)],[c_0_1253, c_0_590]), ['final']).
cnf(c_0_1985,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c7_2(a223,a292)|c5_1(a223)|c6_0|ndr1_1(a193)|ssSkC15|c5_0|c1_0|~c7_2(a223,a293)), inference(spm,[status(thm)],[c_0_1254, c_0_584]), ['final']).
cnf(c_0_1986,negated_conjecture,(c2_2(a223,a293)|c5_2(a296,a297)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|ndr1_1(a193)|ssSkC15|c1_0|~c7_2(a223,a293)), inference(spm,[status(thm)],[c_0_1255, c_0_590]), ['final']).
cnf(c_0_1987,negated_conjecture,(c2_2(a218,a219)|c8_2(a218,a219)|c3_1(a159)|c5_1(a159)|c6_0|c8_0|c2_1(a218)|ndr1_1(a296)|ssSkC13|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_726, c_0_125]), ['final']).
cnf(c_0_1988,negated_conjecture,(c2_2(a218,a219)|c8_2(a218,a219)|c3_1(a159)|c5_1(a159)|c6_0|c2_1(a218)|ndr1_1(a296)|ssSkC13|c5_0|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_727, c_0_125]), ['final']).
cnf(c_0_1989,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c1_1(a256)|c1_1(a241)|c6_0|c8_0|c2_1(a241)|ndr1_1(a159)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1256, c_0_522]), ['final']).
cnf(c_0_1990,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c1_1(a241)|c6_0|c8_0|c2_1(a241)|ndr1_1(a159)|ndr1_1(a256)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1256, c_0_524]), ['final']).
cnf(c_0_1991,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c1_1(a256)|c1_1(a217)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a159)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1257, c_0_522]), ['final']).
cnf(c_0_1992,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c1_1(a217)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a159)|ndr1_1(a256)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1257, c_0_524]), ['final']).
cnf(c_0_1993,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c1_1(a256)|c1_1(a241)|c6_0|c2_1(a241)|ndr1_1(a159)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1258, c_0_522]), ['final']).
cnf(c_0_1994,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c1_1(a241)|c6_0|c2_1(a241)|ndr1_1(a159)|ndr1_1(a256)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1258, c_0_524]), ['final']).
cnf(c_0_1995,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c1_1(a256)|c1_1(a217)|c6_0|ssSkC22|c2_1(a217)|ndr1_1(a159)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1259, c_0_522]), ['final']).
cnf(c_0_1996,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c1_1(a217)|c6_0|ssSkC22|c2_1(a217)|ndr1_1(a159)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1259, c_0_524]), ['final']).
cnf(c_0_1997,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c1_1(a256)|c1_1(a241)|c6_0|c8_0|c2_1(a159)|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1260, c_0_522]), ['final']).
cnf(c_0_1998,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c1_1(a241)|c6_0|c8_0|c2_1(a159)|c2_1(a241)|ndr1_1(a256)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1260, c_0_524]), ['final']).
cnf(c_0_1999,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c1_1(a256)|c1_1(a217)|c6_0|c8_0|ssSkC22|c2_1(a159)|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1261, c_0_522]), ['final']).
cnf(c_0_2000,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c1_1(a217)|c6_0|c8_0|ssSkC22|c2_1(a159)|c2_1(a217)|ndr1_1(a256)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1261, c_0_524]), ['final']).
cnf(c_0_2001,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c1_1(a256)|c1_1(a241)|c6_0|c2_1(a159)|c2_1(a241)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1262, c_0_522]), ['final']).
cnf(c_0_2002,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c1_1(a241)|c6_0|c2_1(a159)|c2_1(a241)|ndr1_1(a256)|ssSkC19|c5_0|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1262, c_0_524]), ['final']).
cnf(c_0_2003,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c1_1(a256)|c1_1(a217)|c6_0|ssSkC22|c2_1(a159)|c2_1(a217)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1263, c_0_522]), ['final']).
cnf(c_0_2004,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c1_1(a217)|c6_0|ssSkC22|c2_1(a159)|c2_1(a217)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1263, c_0_524]), ['final']).
cnf(c_0_2005,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c4_2(a215,a258)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1142, c_0_1058]), ['final']).
cnf(c_0_2006,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a215,a258)|c7_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1142, c_0_1059]), ['final']).
cnf(c_0_2007,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c4_2(a215,a258)|c3_2(a215,a248)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1143, c_0_1058]), ['final']).
cnf(c_0_2008,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c3_2(a215,a248)|c3_2(a215,a258)|c1_1(a294)|c5_1(a294)|c8_0|ssSkC22|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1143, c_0_1059]), ['final']).
cnf(c_0_2009,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c4_2(a215,a258)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1144, c_0_1058]), ['final']).
cnf(c_0_2010,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a168)|c1_0), inference(spm,[status(thm)],[c_0_1264, c_0_559]), ['final']).
cnf(c_0_2011,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a168)|c1_0), inference(spm,[status(thm)],[c_0_1265, c_0_559]), ['final']).
cnf(c_0_2012,negated_conjecture,(ssSkC7|c5_2(a296,a297)|c1_1(a168)|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(X1)|ssSkC11|ssSkC15|c1_0), inference(ef,[status(thm)],[c_0_1266]), ['final']).
cnf(c_0_2013,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a215,a258)|c7_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1144, c_0_1059]), ['final']).
cnf(c_0_2014,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c4_2(a215,a258)|c3_2(a215,a248)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1145, c_0_1058]), ['final']).
cnf(c_0_2015,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c3_2(a215,a248)|c3_2(a215,a258)|c1_1(a204)|c5_1(a204)|c8_0|ssSkC22|c7_0|ndr1_1(a160)|c1_0), inference(spm,[status(thm)],[c_0_1145, c_0_1059]), ['final']).
cnf(c_0_2016,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a168)|c1_0), inference(spm,[status(thm)],[c_0_1267, c_0_559]), ['final']).
cnf(c_0_2017,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a168)|c1_0), inference(spm,[status(thm)],[c_0_1268, c_0_559]), ['final']).
cnf(c_0_2018,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a158)|ndr1_1(a159)|ndr1_1(a296)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1269, c_0_869]), ['final']).
cnf(c_0_2019,negated_conjecture,(c2_2(a158,a293)|c8_2(a158,a293)|c7_2(a158,a292)|c5_1(a158)|c6_0|c2_1(a159)|c2_1(a158)|ndr1_1(a296)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_1270, c_0_869]), ['final']).
cnf(c_0_2020,negated_conjecture,(c2_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|ndr1_1(a193)|ndr1_1(a296)|ssSkC15|c5_0|c1_0|~c7_2(a223,a293)), inference(spm,[status(thm)],[c_0_1271, c_0_584]), ['final']).
cnf(c_0_2021,negated_conjecture,(c2_2(a223,a293)|c7_2(a223,a292)|c5_1(a223)|c6_0|c8_0|ndr1_1(a193)|ndr1_1(a296)|ssSkC15|c1_0|~c7_2(a223,a293)), inference(spm,[status(thm)],[c_0_1272, c_0_590]), ['final']).
cnf(c_0_2022,negated_conjecture,(c2_2(a218,a219)|c8_2(a218,a219)|c4_2(a159,a293)|c7_2(a159,a292)|c8_0|c2_1(a218)|ssSkC13|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_718, c_0_125]), ['final']).
cnf(c_0_2023,negated_conjecture,(c2_2(a218,a219)|c8_2(a159,a292)|c8_2(a218,a219)|c4_2(a159,a293)|c8_0|c2_1(a218)|ssSkC13|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_719, c_0_125]), ['final']).
cnf(c_0_2024,negated_conjecture,(c2_2(a218,a219)|c8_2(a218,a219)|c4_2(a159,a293)|c7_2(a159,a292)|c2_1(a218)|ssSkC13|c5_0|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_720, c_0_125]), ['final']).
cnf(c_0_2025,negated_conjecture,(c2_2(a218,a219)|c8_2(a159,a292)|c8_2(a218,a219)|c4_2(a159,a293)|c2_1(a218)|ssSkC13|c5_0|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_721, c_0_125]), ['final']).
cnf(c_0_2026,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a241,a185)|c6_0|c8_0|c2_1(a241)|ndr1_1(a159)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1256, c_0_499]), ['final']).
cnf(c_0_2027,negated_conjecture,(ssSkC20|c2_2(a168,a248)|c8_2(a160,a161)|c4_1(a168)|c5_1(a168)|c3_0|c8_1(X1)|ndr1_1(a287)|c5_0|c1_0|~c1_1(X1)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_496, c_0_1078]), ['final']).
cnf(c_0_2028,negated_conjecture,(ssSkC20|c2_2(a160,a161)|c2_2(a168,a248)|c4_1(a168)|c5_1(a168)|c3_0|c8_1(X1)|ndr1_1(a287)|c5_0|c1_0|~c1_1(X1)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_497, c_0_1078]), ['final']).
cnf(c_0_2029,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c7_2(X1,X2)|c7_1(X1)|c6_0|c7_0|~c6_2(X1,X2)|~c1_2(X1,X3)|~c7_2(X1,X3)|~c1_1(a208)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1273, c_0_724]), ['final']).
cnf(c_0_2030,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c7_2(a241,a185)|c6_0|c8_0|c2_1(a241)|ndr1_1(a159)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1256, c_0_582]), ['final']).
cnf(c_0_2031,negated_conjecture,(c2_2(a241,a243)|c5_2(a241,a185)|c8_2(a241,a243)|c6_0|c8_0|c2_1(a241)|ndr1_1(a159)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1256, c_0_722]), ['final']).
cnf(c_0_2032,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a217,a185)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a159)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1257, c_0_499]), ['final']).
cnf(c_0_2033,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c7_2(a217,a185)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a159)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1257, c_0_582]), ['final']).
cnf(c_0_2034,negated_conjecture,(c2_2(a217,a258)|c5_2(a217,a185)|c8_2(a217,a258)|c6_0|c8_0|ssSkC22|c2_1(a217)|ndr1_1(a159)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1257, c_0_722]), ['final']).
cnf(c_0_2035,negated_conjecture,(c6_2(a231,a232)|c1_2(a223,a293)|c3_2(X1,X2)|c7_2(a223,a292)|c2_1(a223)|ssSkC15|c5_0|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a223,a293)|~c1_1(a223)|~c5_1(X1)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_550, c_0_630]), ['final']).
cnf(c_0_2036,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c4_2(a241,a185)|c6_0|c8_0|c2_1(a159)|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1260, c_0_499]), ['final']).
cnf(c_0_2037,negated_conjecture,(c2_2(a241,a243)|c8_2(a241,a243)|c7_2(a241,a185)|c6_0|c8_0|c2_1(a159)|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1260, c_0_582]), ['final']).
cnf(c_0_2038,negated_conjecture,(c2_2(a241,a243)|c5_2(a241,a185)|c8_2(a241,a243)|c6_0|c8_0|c2_1(a159)|c2_1(a241)|ssSkC19|c1_0|~c5_1(a241)|~ndr1_1(a241)), inference(spm,[status(thm)],[c_0_1260, c_0_722]), ['final']).
cnf(c_0_2039,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c4_2(a217,a185)|c6_0|c8_0|ssSkC22|c2_1(a159)|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1261, c_0_499]), ['final']).
cnf(c_0_2040,negated_conjecture,(c2_2(a217,a258)|c8_2(a217,a258)|c7_2(a217,a185)|c6_0|c8_0|ssSkC22|c2_1(a159)|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1261, c_0_582]), ['final']).
cnf(c_0_2041,negated_conjecture,(c6_2(a231,a232)|c1_2(a223,a293)|c3_2(X1,X2)|c7_2(a223,a292)|c8_0|c2_1(a223)|ssSkC15|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(a223,a293)|~c1_1(a223)|~c5_1(X1)|~ndr1_1(a223)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_550, c_0_641]), ['final']).
cnf(c_0_2042,negated_conjecture,(c2_2(a217,a258)|c5_2(a217,a185)|c8_2(a217,a258)|c6_0|c8_0|ssSkC22|c2_1(a159)|c2_1(a217)|c1_0|~c5_1(a217)|~ndr1_1(a217)), inference(spm,[status(thm)],[c_0_1261, c_0_722]), ['final']).
cnf(c_0_2043,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c4_2(a197,a185)|c6_0|c8_0|c2_1(a197)|ndr1_1(a159)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1274, c_0_499]), ['final']).
cnf(c_0_2044,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c1_1(a256)|c1_1(a197)|c6_0|c2_1(a197)|ndr1_1(a159)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1274, c_0_522]), ['final']).
cnf(c_0_2045,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c1_1(a197)|c6_0|c2_1(a197)|ndr1_1(a159)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1274, c_0_524]), ['final']).
cnf(c_0_2046,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c7_2(a197,a185)|c6_0|c8_0|c2_1(a197)|ndr1_1(a159)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1274, c_0_582]), ['final']).
cnf(c_0_2047,negated_conjecture,(c2_2(a197,a198)|c5_2(a197,a185)|c8_2(a197,a198)|c6_0|c8_0|c2_1(a197)|ndr1_1(a159)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1274, c_0_722]), ['final']).
cnf(c_0_2048,negated_conjecture,(ssSkC7|c5_2(a296,a297)|c1_1(a168)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(X1)|ssSkC15|c1_0), inference(ef,[status(thm)],[c_0_1275]), ['final']).
cnf(c_0_2049,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c4_2(a197,a185)|c6_0|c8_0|c2_1(a159)|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1276, c_0_499]), ['final']).
cnf(c_0_2050,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c1_1(a256)|c1_1(a197)|c6_0|c2_1(a159)|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1276, c_0_522]), ['final']).
cnf(c_0_2051,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c1_1(a197)|c6_0|c2_1(a159)|c2_1(a197)|ndr1_1(a256)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1276, c_0_524]), ['final']).
cnf(c_0_2052,negated_conjecture,(c2_2(a197,a198)|c8_2(a197,a198)|c7_2(a197,a185)|c6_0|c8_0|c2_1(a159)|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1276, c_0_582]), ['final']).
cnf(c_0_2053,negated_conjecture,(c2_2(a197,a198)|c5_2(a197,a185)|c8_2(a197,a198)|c6_0|c8_0|c2_1(a159)|c2_1(a197)|c5_0|c1_0|~c5_1(a197)|~ndr1_1(a197)), inference(spm,[status(thm)],[c_0_1276, c_0_722]), ['final']).
cnf(c_0_2054,negated_conjecture,(ssSkC7|c1_1(a168)|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC11|ssSkC15|c1_0), inference(ef,[status(thm)],[c_0_1277]), ['final']).
cnf(c_0_2055,negated_conjecture,(ssSkC20|c2_2(a168,a248)|c4_1(a168)|c5_1(a168)|c3_0|c8_1(X1)|ndr1_1(a160)|ndr1_1(a287)|c5_0|c1_0|~c1_1(X1)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_498, c_0_1078]), ['final']).
cnf(c_0_2056,negated_conjecture,(c5_2(a233,a234)|c6_1(a233)|c2_0|c4_0|ssSkC17|~c5_2(a233,X1)|~c8_2(a233,a234)|~c7_2(a233,X1)), inference(spm,[status(thm)],[c_0_1278, c_0_63]), ['final']).
cnf(c_0_2057,negated_conjecture,(c5_2(a268,a270)|c6_1(a268)|c2_0|c4_0|c5_0|~c5_2(a268,X1)|~c7_2(a268,X1)), inference(spm,[status(thm)],[c_0_1279, c_0_156]), ['final']).
cnf(c_0_2058,negated_conjecture,(c5_2(a294,a147)|c4_2(a294,a295)|c4_2(a294,X1)|c3_2(a294,X1)|c6_0|c8_0|c1_0|~c5_2(a294,a295)), inference(spm,[status(thm)],[c_0_1280, c_0_24]), ['final']).
cnf(c_0_2059,negated_conjecture,(c5_2(a204,a147)|c4_2(a204,a205)|c4_2(a204,X1)|c3_2(a204,X1)|c6_0|c8_0|c1_0|~c5_2(a204,a205)), inference(spm,[status(thm)],[c_0_1281, c_0_30]), ['final']).
cnf(c_0_2060,negated_conjecture,(ssSkC20|c2_2(a294,a295)|c5_2(a294,a278)|c1_1(a294)|c5_1(a294)|c3_0|c8_0|c7_0|ndr1_1(a160)|ndr1_1(a168)|c1_0), inference(spm,[status(thm)],[c_0_1282, c_0_559]), ['final']).
cnf(c_0_2061,negated_conjecture,(ssSkC7|c5_2(a296,a297)|c1_2(a202,a190)|c3_1(a202)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(X1)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1283, c_0_739]), ['final']).
cnf(c_0_2062,negated_conjecture,(ssSkC20|c2_2(a204,a205)|c5_2(a204,a278)|c1_1(a204)|c5_1(a204)|c3_0|c8_0|c7_0|ndr1_1(a160)|ndr1_1(a168)|c1_0), inference(spm,[status(thm)],[c_0_1284, c_0_559]), ['final']).
cnf(c_0_2063,negated_conjecture,(ssSkC7|c1_1(a168)|c5_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0), inference(ef,[status(thm)],[c_0_1285]), ['final']).
cnf(c_0_2064,negated_conjecture,(ssSkC7|c5_2(a296,a297)|c1_2(a217,a190)|c3_1(a217)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1286, c_0_739]), ['final']).
cnf(c_0_2065,negated_conjecture,(ssSkC7|c1_2(a202,a190)|c3_1(a202)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1283, c_0_743]), ['final']).
cnf(c_0_2066,negated_conjecture,(ssSkC20|c1_2(a168,a261)|c1_2(a168,a248)|c8_2(a168,a248)|c3_0|c6_0|ndr1_1(a287)|c5_0|c1_0|~ssSkP5(a168)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_1287, c_0_1078]), ['final']).
cnf(c_0_2067,negated_conjecture,(ssSkC20|c2_2(a168,a261)|c1_2(a168,a248)|c8_2(a168,a248)|c3_0|c6_0|ndr1_1(a287)|c5_0|c1_0|~ssSkP5(a168)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_1288, c_0_1078]), ['final']).
cnf(c_0_2068,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c1_2(a202,a190)|c3_1(a202)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1283, c_0_854]), ['final']).
cnf(c_0_2069,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c1_2(a202,a190)|c3_2(a168,a248)|c3_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1283, c_0_856]), ['final']).
cnf(c_0_2070,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c3_2(X1,X2)|c6_0|c8_0|~c5_2(X1,X3)|~c7_2(X1,a147)|~c7_2(X1,X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1289, c_0_19]), ['final']).
cnf(c_0_2071,negated_conjecture,(c4_2(X1,X2)|c4_2(X1,X3)|c3_2(X1,X2)|c6_0|c8_0|~c5_2(X1,X3)|~c3_2(X1,a147)|~c7_2(X1,X3)|~ndr1_1(X1)), inference(spm,[status(thm)],[c_0_1290, c_0_19]), ['final']).
cnf(c_0_2072,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c7_2(a168,a248)|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_564, c_0_1291]), ['final']).
cnf(c_0_2073,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c3_2(a168,a248)|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_711, c_0_1292]), ['final']).
cnf(c_0_2074,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|~c5_2(a160,a161)|~c7_2(X1,X2)|~c1_1(X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_157, c_0_476])]), ['final']).
cnf(c_0_2075,negated_conjecture,(ssSkC7|c1_2(a217,a190)|c3_1(a217)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1286, c_0_743]), ['final']).
cnf(c_0_2076,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|c1_0|~c6_2(X1,X2)|~c4_2(a193,a194)|~c3_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_158, c_0_476])]), ['final']).
cnf(c_0_2077,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c1_2(a217,a190)|c3_1(a217)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1286, c_0_854]), ['final']).
cnf(c_0_2078,negated_conjecture,(ssSkC7|ssSkC20|c1_2(a202,a190)|c3_1(a202)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1283, c_0_599]), ['final']).
cnf(c_0_2079,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c1_2(a217,a190)|c3_2(a168,a248)|c3_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1286, c_0_856]), ['final']).
cnf(c_0_2080,negated_conjecture,(ssSkC7|ssSkC20|c1_2(a202,a190)|c3_2(a168,a248)|c3_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1283, c_0_936]), ['final']).
cnf(c_0_2081,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c1_2(a202,a190)|c3_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1283, c_0_862]), ['final']).
cnf(c_0_2082,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|c1_0|~c6_2(a193,a194)|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(X1,X2)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_159, c_0_476])]), ['final']).
cnf(c_0_2083,negated_conjecture,(c5_2(X1,X2)|~c6_2(a274,X3)|~c6_2(X1,X2)|~c5_2(a274,X3)|~c1_2(a274,X3)|~c1_1(X1)|~ndr1_1(a274)|~ndr1_1(X1)|~ssSkC24), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_160, c_0_476])]), ['final']).
cnf(c_0_2084,negated_conjecture,(c2_2(X1,X2)|c4_1(X1)|c5_1(X1)|c8_1(X3)|~c6_2(a160,X4)|~c7_2(X1,X2)|~c1_1(X3)|~ndr1_1(a160)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_161, c_0_476])]), ['final']).
cnf(c_0_2085,negated_conjecture,(c5_2(a274,X1)|c5_2(X2,X3)|~c6_2(X2,X3)|~c4_2(a274,X1)|~c7_2(a274,X1)|~c1_1(X2)|~ndr1_1(a274)|~ndr1_1(X2)|~ssSkC24), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_162, c_0_476])]), ['final']).
cnf(c_0_2086,negated_conjecture,(ssSkC7|ssSkC20|c7_2(a168,a248)|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(a296)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_564, c_0_1293]), ['final']).
cnf(c_0_2087,negated_conjecture,(ssSkC7|ssSkC20|c3_2(a168,a248)|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(a296)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_711, c_0_1294]), ['final']).
cnf(c_0_2088,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a202)|ndr1_1(a223)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_567, c_0_1295]), ['final']).
cnf(c_0_2089,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c8_2(a160,a161)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1296]), ['final']).
cnf(c_0_2090,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c7_2(a217,a248)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_598, c_0_739]), ['final']).
cnf(c_0_2091,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c3_2(a217,a248)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_855, c_0_739]), ['final']).
cnf(c_0_2092,negated_conjecture,(ssSkC7|c5_2(X1,a278)|c1_2(a202,a190)|c3_1(a202)|c1_1(X1)|c8_0|c2_0|c7_0|ssSkC11|~c5_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_932]), ['final']).
cnf(c_0_2093,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c8_2(a160,a161)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1297]), ['final']).
cnf(c_0_2094,negated_conjecture,(c2_2(a160,a161)|c2_2(a294,a295)|c5_2(a294,a278)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1298]), ['final']).
cnf(c_0_2095,negated_conjecture,(c2_2(a160,a161)|c2_2(a204,a205)|c5_2(a204,a278)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1299]), ['final']).
cnf(c_0_2096,negated_conjecture,(c1_2(a174,a261)|c1_2(a174,a185)|c8_2(a174,a185)|c6_0|c8_0|ssSkC4|c5_0|c1_0|~ssSkP5(a174)|~ndr1_1(a174)), inference(spm,[status(thm)],[c_0_1287, c_0_693]), ['final']).
cnf(c_0_2097,negated_conjecture,(c2_2(a174,a261)|c1_2(a174,a185)|c8_2(a174,a185)|c6_0|c8_0|ssSkC4|c5_0|c1_0|~ssSkP5(a174)|~ndr1_1(a174)), inference(spm,[status(thm)],[c_0_1288, c_0_693]), ['final']).
cnf(c_0_2098,negated_conjecture,(ssSkC7|ssSkC20|c1_2(a217,a190)|c3_1(a217)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1286, c_0_599]), ['final']).
cnf(c_0_2099,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X3)|c1_2(X1,X3)|c4_2(X1,X3)|~c8_2(X1,X2)|~c3_2(X1,X2)|~c6_1(a189)|~c7_1(X1)|~ndr1_1(X1)|~ssSkC6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_163, c_0_476])]), ['final']).
cnf(c_0_2100,negated_conjecture,(c3_1(X1)|c1_1(X1)|c6_0|~c5_2(a256,X2)|~c4_2(a256,X2)|~c3_2(a256,X2)|~c5_1(X1)|~ndr1_1(a256)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_164, c_0_476])]), ['final']).
cnf(c_0_2101,negated_conjecture,(ssSkC7|c1_2(a202,a190)|c3_1(a202)|c1_1(X1)|c8_0|c2_0|c7_0|ndr1_1(X1)|ssSkC11|~c5_1(X1)), inference(spm,[status(thm)],[c_0_529, c_0_932]), ['final']).
cnf(c_0_2102,negated_conjecture,(ssSkC7|c5_2(X1,a278)|c3_1(a202)|c1_1(X1)|c8_0|c2_0|c7_0|ndr1_1(a202)|ssSkC11|~c5_1(X1)), inference(spm,[status(thm)],[c_0_488, c_0_642]), ['final']).
cnf(c_0_2103,negated_conjecture,(ssSkC7|ssSkC20|c1_2(a217,a190)|c3_2(a168,a248)|c3_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1286, c_0_936]), ['final']).
cnf(c_0_2104,negated_conjecture,(ssSkC7|ssSkC20|c5_2(a296,a297)|c1_2(a217,a190)|c3_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1286, c_0_862]), ['final']).
cnf(c_0_2105,negated_conjecture,(ssSkC7|ssSkC20|c1_2(a202,a190)|c3_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1283, c_0_603]), ['final']).
cnf(c_0_2106,negated_conjecture,(c2_2(a233,a234)|c7_2(a193,a194)|c5_1(a233)|ssSkC17|c1_0|~c3_1(a233)|~c7_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_733, c_0_39]), ['final']).
cnf(c_0_2107,negated_conjecture,(c8_2(a223,X1)|c3_2(a223,X1)|c7_2(a223,X1)|c8_0|c2_1(a223)|ssSkC15|c1_0|~c5_1(a223)|~ndr1_1(a223)), inference(spm,[status(thm)],[c_0_750, c_0_1300]), ['final']).
cnf(c_0_2108,negated_conjecture,(ssSkC20|c5_2(a296,a297)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_656, c_0_739]), ['final']).
cnf(c_0_2109,negated_conjecture,(ssSkC20|c7_2(a217,a248)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_598, c_0_743]), ['final']).
cnf(c_0_2110,negated_conjecture,(ssSkC20|c3_2(a217,a248)|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_855, c_0_743]), ['final']).
cnf(c_0_2111,negated_conjecture,(ssSkC20|c1_1(a168)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ndr1_1(X1)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_656, c_0_743]), ['final']).
cnf(c_0_2112,negated_conjecture,(c2_2(X1,X2)|c4_2(X1,X2)|~c5_2(a203,X3)|~c5_2(X1,X2)|~c1_2(a203,X3)|~c8_2(a203,X3)|~c2_1(X1)|~ndr1_1(a203)|~ndr1_1(X1)|~ssSkC11), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_165, c_0_476])]), ['final']).
cnf(c_0_2113,negated_conjecture,(ssSkC7|c1_1(a168)|c5_1(a221)|c3_0|c6_0|ndr1_1(a221)|ndr1_1(X1)|ssSkC14|c1_0), inference(spm,[status(thm)],[c_0_643, c_0_1301]), ['final']).
cnf(c_0_2114,negated_conjecture,(c1_1(a168)|c3_0|c6_0|ndr1_1(a287)|ndr1_1(X1)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_113, c_0_1302]), ['final']).
cnf(c_0_2115,negated_conjecture,(ssSkC7|c1_1(a168)|c5_1(a168)|c3_0|c6_0|ndr1_1(a168)|ndr1_1(X1)|c1_0), inference(spm,[status(thm)],[c_0_643, c_0_1303]), ['final']).
cnf(c_0_2116,negated_conjecture,(c2_2(a218,a219)|c8_2(a218,a219)|c8_0|c2_1(a218)|ndr1_1(a159)|ssSkC13|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_729, c_0_125]), ['final']).
cnf(c_0_2117,negated_conjecture,(c2_2(a218,a219)|c8_2(a218,a219)|c2_1(a218)|ndr1_1(a159)|ssSkC13|c5_0|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_730, c_0_125]), ['final']).
cnf(c_0_2118,negated_conjecture,(c2_2(a294,a295)|c5_2(a294,a278)|c1_1(a294)|c5_1(a294)|c8_0|c7_0|ndr1_1(a160)|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1304]), ['final']).
cnf(c_0_2119,negated_conjecture,(c1_1(a265)|c5_1(X1)|c6_0|c2_1(a265)|ndr1_1(X2)|ndr1_1(X1)|ssSkC23|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1305, c_0_1306]), ['final']).
cnf(c_0_2120,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X3)|c7_2(X1,X2)|c8_1(X1)|~ssSkC7|~c6_2(X1,X3)|~c2_2(X1,X3)|~c8_2(a191,a192)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_166, c_0_476])]), ['final']).
cnf(c_0_2121,negated_conjecture,(c2_2(a296,X1)|c3_1(X2)|c4_1(X2)|c5_1(X2)|c6_0|~c6_2(a296,X1)|~c7_2(a296,X1)|~ndr1_1(a296)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_167, c_0_476])]), ['final']).
cnf(c_0_2122,negated_conjecture,(c2_2(a204,a205)|c5_2(a204,a278)|c1_1(a204)|c5_1(a204)|c8_0|c7_0|ndr1_1(a160)|ndr1_1(a253)|ssSkC21|c1_0), inference(spm,[status(thm)],[c_0_488, c_0_1307]), ['final']).
cnf(c_0_2123,negated_conjecture,(c2_2(a218,a219)|c8_2(a218,a219)|c8_0|c2_1(a159)|c2_1(a218)|ssSkC13|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_731, c_0_125]), ['final']).
cnf(c_0_2124,negated_conjecture,(ssSkC7|ssSkC20|c5_1(a202)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a202)|ndr1_1(a223)|ndr1_1(a296)|ssSkC11|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_567, c_0_1308]), ['final']).
cnf(c_0_2125,negated_conjecture,(c1_1(a173)|c5_1(X1)|c6_0|c2_1(a173)|ndr1_1(X2)|ndr1_1(X1)|ssSkC3|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_1309, c_0_1310]), ['final']).
cnf(c_0_2126,negated_conjecture,(ssSkC7|c3_1(a202)|c1_1(X1)|c8_0|c2_0|c7_0|ndr1_1(a202)|ndr1_1(X1)|ssSkC11|~c5_1(X1)), inference(spm,[status(thm)],[c_0_529, c_0_642]), ['final']).
cnf(c_0_2127,negated_conjecture,(c1_2(X1,X2)|c8_2(X1,X2)|c6_0|c5_0|~c6_2(X1,a261)|~c7_2(X1,X2)|~ssSkP5(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_168, c_0_476])]), ['final']).
cnf(c_0_2128,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X3)|c7_2(a191,a192)|c7_2(X1,X2)|c8_1(X1)|~ssSkC7|~c6_2(X1,X3)|~c2_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_169, c_0_476])]), ['final']).
cnf(c_0_2129,negated_conjecture,(c2_2(a218,a219)|c8_2(a218,a219)|c2_1(a159)|c2_1(a218)|ssSkC13|c5_0|c1_0|~c3_1(a218)|~ndr1_1(a218)), inference(spm,[status(thm)],[c_0_732, c_0_125]), ['final']).
cnf(c_0_2130,negated_conjecture,(c6_2(X1,X2)|c1_2(X1,a220)|c8_2(X1,X3)|c4_0|~c1_2(X1,X3)|~c8_2(X1,X2)|~c4_2(X1,X2)|~c7_2(X1,X3)|~ndr1_1(X1)|~ssSkC13), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_170, c_0_476])]), ['final']).
cnf(c_0_2131,negated_conjecture,(ssSkC7|ssSkC20|c1_2(a217,a190)|c3_1(a217)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_1286, c_0_603]), ['final']).
cnf(c_0_2132,negated_conjecture,(c8_2(a233,a234)|c4_1(a233)|c6_0|c8_0|c4_0|ssSkC17|~c7_2(a233,a234)|~c8_1(a233)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_1311, c_0_39]), ['final']).
cnf(c_0_2133,negated_conjecture,(c2_2(a233,a234)|c5_1(a233)|ndr1_1(a193)|ssSkC17|c1_0|~c3_1(a233)|~c7_2(a233,a234)|~ndr1_1(a233)), inference(spm,[status(thm)],[c_0_735, c_0_39]), ['final']).
cnf(c_0_2134,negated_conjecture,(c4_2(a268,a269)|c3_2(a268,a269)|c3_0|c8_0|c4_0|c5_0|~c6_1(a268)|~c8_1(a268)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_917, c_0_105]), ['final']).
cnf(c_0_2135,negated_conjecture,(c1_2(a294,a261)|c1_2(a294,a295)|c8_2(a294,a295)|c6_0|c5_0|c1_0|~ssSkP5(a294)|~ndr1_1(a294)), inference(spm,[status(thm)],[c_0_1287, c_0_13]), ['final']).
cnf(c_0_2136,negated_conjecture,(c2_2(X1,X2)|c5_1(X1)|c1_0|~c6_2(X1,X2)|~c3_1(X1)|~c7_2(X1,X2)|~c7_1(a193)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_171, c_0_476])]), ['final']).
cnf(c_0_2137,negated_conjecture,(ssSkC20|c7_2(a265,a248)|c5_1(X1)|c2_1(a265)|ndr1_1(X1)|ssSkC23|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_1306]), ['final']).
cnf(c_0_2138,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X3)|c7_2(X1,X2)|c8_1(X1)|~ssSkC7|~c6_2(X1,X3)|~c2_2(X1,X3)|~c5_1(a191)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_172, c_0_476])]), ['final']).
cnf(c_0_2139,negated_conjecture,(c1_2(a204,a261)|c1_2(a204,a205)|c8_2(a204,a205)|c6_0|c5_0|c1_0|~ssSkP5(a204)|~ndr1_1(a204)), inference(spm,[status(thm)],[c_0_1287, c_0_15]), ['final']).
cnf(c_0_2140,negated_conjecture,(c1_2(a275,a261)|c1_2(a275,a276)|c8_2(a275,a276)|c6_0|ssSkC25|c5_0|~ssSkP5(a275)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_1287, c_0_66]), ['final']).
cnf(c_0_2141,negated_conjecture,(c1_2(a265,a261)|c1_2(a265,a266)|c8_2(a265,a266)|c6_0|ssSkC23|c5_0|~ssSkP5(a265)|~ndr1_1(a265)), inference(spm,[status(thm)],[c_0_1287, c_0_797]), ['final']).
cnf(c_0_2142,negated_conjecture,(c1_2(a155,a261)|c1_2(a155,a156)|c8_2(a155,a156)|c6_0|ssSkC1|c5_0|~ssSkP5(a155)|~ndr1_1(a155)), inference(spm,[status(thm)],[c_0_1287, c_0_173]), ['final']).
cnf(c_0_2143,negated_conjecture,(c1_2(a148,a261)|c1_2(a148,a149)|c8_2(a148,a149)|c6_0|c7_0|c5_0|~ssSkP5(a148)|~ndr1_1(a148)), inference(spm,[status(thm)],[c_0_1287, c_0_1312]), ['final']).
cnf(c_0_2144,negated_conjecture,(c2_2(a294,a261)|c1_2(a294,a295)|c8_2(a294,a295)|c6_0|c5_0|c1_0|~ssSkP5(a294)|~ndr1_1(a294)), inference(spm,[status(thm)],[c_0_1288, c_0_13]), ['final']).
cnf(c_0_2145,negated_conjecture,(c2_2(a204,a261)|c1_2(a204,a205)|c8_2(a204,a205)|c6_0|c5_0|c1_0|~ssSkP5(a204)|~ndr1_1(a204)), inference(spm,[status(thm)],[c_0_1288, c_0_15]), ['final']).
cnf(c_0_2146,negated_conjecture,(c2_2(a275,a261)|c1_2(a275,a276)|c8_2(a275,a276)|c6_0|ssSkC25|c5_0|~ssSkP5(a275)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_1288, c_0_66]), ['final']).
cnf(c_0_2147,negated_conjecture,(ssSkC20|c3_2(a265,a248)|c5_1(X1)|c2_1(a265)|ndr1_1(X1)|ssSkC23|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_1306]), ['final']).
cnf(c_0_2148,negated_conjecture,(c5_2(X1,a278)|c1_1(X1)|c8_0|c2_0|c7_0|ssSkC11|~c5_1(X1)|~c2_1(a202)), inference(spm,[status(thm)],[c_0_1313, c_0_50]), ['final']).
cnf(c_0_2149,negated_conjecture,(c3_0|c4_0|~c5_2(X1,X2)|~c4_2(X1,X2)|~c3_2(X1,a207)|~ssSkP0(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_174, c_0_476])]), ['final']).
cnf(c_0_2150,negated_conjecture,(c3_0|c4_0|~c5_2(X1,a207)|~c5_2(X1,X2)|~c4_2(X1,X2)|~ssSkP0(X1)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_175, c_0_476])]), ['final']).
cnf(c_0_2151,negated_conjecture,(c2_2(a265,a261)|c1_2(a265,a266)|c8_2(a265,a266)|c6_0|ssSkC23|c5_0|~ssSkP5(a265)|~ndr1_1(a265)), inference(spm,[status(thm)],[c_0_1288, c_0_797]), ['final']).
cnf(c_0_2152,negated_conjecture,(c2_2(a155,a261)|c1_2(a155,a156)|c8_2(a155,a156)|c6_0|ssSkC1|c5_0|~ssSkP5(a155)|~ndr1_1(a155)), inference(spm,[status(thm)],[c_0_1288, c_0_173]), ['final']).
cnf(c_0_2153,negated_conjecture,(c2_2(a148,a261)|c1_2(a148,a149)|c8_2(a148,a149)|c6_0|c7_0|c5_0|~ssSkP5(a148)|~ndr1_1(a148)), inference(spm,[status(thm)],[c_0_1288, c_0_1312]), ['final']).
cnf(c_0_2154,negated_conjecture,(c4_2(a275,a276)|c3_2(a275,a276)|c3_0|c8_0|ssSkC25|~c6_1(a275)|~c8_1(a275)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_917, c_0_748]), ['final']).
cnf(c_0_2155,negated_conjecture,(ssSkC20|c7_2(a173,a248)|c5_1(X1)|c2_1(a173)|ndr1_1(X1)|ssSkC3|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_506, c_0_1310]), ['final']).
cnf(c_0_2156,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X3)|c7_2(X1,X2)|c8_1(X1)|ndr1_1(a191)|~ssSkC7|~c6_2(X1,X3)|~c2_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_176, c_0_476])]), ['final']).
cnf(c_0_2157,negated_conjecture,(c6_2(X1,X2)|c3_2(X1,X3)|c7_2(X1,X2)|c1_1(a191)|c8_1(X1)|~ssSkC7|~c6_2(X1,X3)|~c2_2(X1,X3)|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_177, c_0_476])]), ['final']).
cnf(c_0_2158,negated_conjecture,(c5_2(a252,X1)|c1_0|~c2_2(a252,X1)|~c4_2(a252,X1)|~c8_1(X2)|~ndr1_1(a252)), inference(spm,[status(thm)],[c_0_1314, c_0_4]), ['final']).
cnf(c_0_2159,negated_conjecture,(c1_1(X1)|c8_0|c2_0|c7_0|ndr1_1(X1)|ssSkC11|~c5_1(X1)|~c2_1(a202)), inference(spm,[status(thm)],[c_0_1315, c_0_50]), ['final']).
cnf(c_0_2160,negated_conjecture,(ssSkC20|c3_2(a217,a248)|c7_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_855, c_0_599]), ['final']).
cnf(c_0_2161,negated_conjecture,(c8_2(a268,a269)|c3_2(a268,a269)|c8_0|c2_1(a268)|c4_0|c5_0|~c5_1(a268)|~ndr1_1(a268)), inference(spm,[status(thm)],[c_0_750, c_0_178]), ['final']).
cnf(c_0_2162,negated_conjecture,(c5_2(X1,X2)|c5_2(X3,X4)|c3_2(X3,X4)|c7_2(X3,a291)|c7_2(X1,X2)|c7_2(X3,X4)|c6_1(X1)|~c3_2(X1,X2)|~ssSkP9(X3)|~c2_1(X1)|~ndr1_1(X1)|~ndr1_1(X3)|~ssSkC27), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_179, c_0_476])]), ['final']).
cnf(c_0_2163,negated_conjecture,(c6_2(a174,a185)|c6_0|c8_0|ssSkC4|c1_0|~c8_2(a174,a185)), inference(spm,[status(thm)],[c_0_1316, c_0_1147]), ['final']).
cnf(c_0_2164,negated_conjecture,(c1_1(X1)|c8_0|c2_1(X2)|~c1_2(X1,a247)|~c7_1(X1)|~c7_1(X2)|~c5_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_180, c_0_476])]), ['final']).
cnf(c_0_2165,negated_conjecture,(c1_1(X1)|c8_0|c2_1(X2)|~c5_2(X1,a247)|~c7_1(X1)|~c7_1(X2)|~c5_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_181, c_0_476])]), ['final']).
cnf(c_0_2166,negated_conjecture,(ssSkC20|c3_2(a168,a248)|c7_2(a217,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_598, c_0_936]), ['final']).
cnf(c_0_2167,negated_conjecture,(c6_2(a275,a206)|c8_2(a275,a207)|c3_0|ssSkC25|c4_0|~c4_2(a275,a276)), inference(spm,[status(thm)],[c_0_1317, c_0_54]), ['final']).
cnf(c_0_2168,negated_conjecture,(c8_2(a275,a207)|c4_2(a275,a206)|c3_0|ssSkC25|c4_0|~c4_2(a275,a276)), inference(spm,[status(thm)],[c_0_1318, c_0_54]), ['final']).
cnf(c_0_2169,negated_conjecture,(c8_2(a275,a207)|c3_2(a275,a206)|c3_0|ssSkC25|c4_0|~c4_2(a275,a276)), inference(spm,[status(thm)],[c_0_1319, c_0_54]), ['final']).
cnf(c_0_2170,negated_conjecture,(c2_2(X1,a291)|c5_2(X2,X3)|c5_2(X1,X4)|c3_2(X1,X4)|c7_2(X2,X3)|c7_2(X1,X4)|c6_1(X2)|~c3_2(X2,X3)|~ssSkP9(X1)|~c2_1(X2)|~ndr1_1(X2)|~ndr1_1(X1)|~ssSkC27), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_182, c_0_476])]), ['final']).
cnf(c_0_2171,negated_conjecture,(ssSkC20|c3_2(a173,a248)|c5_1(X1)|c2_1(a173)|ndr1_1(X1)|ssSkC3|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_512, c_0_1310]), ['final']).
cnf(c_0_2172,negated_conjecture,(c6_1(X1)|c1_1(X1)|c6_0|~c7_2(X2,a286)|~ssSkP7(X2)|~ssSkP8(X2)|~c8_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_183, c_0_476])]), ['final']).
cnf(c_0_2173,negated_conjecture,(c6_1(X1)|c1_1(X1)|c6_0|~c1_2(X2,a286)|~ssSkP7(X2)|~ssSkP8(X2)|~c8_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_184, c_0_476])]), ['final']).
cnf(c_0_2174,negated_conjecture,(c2_1(a196)|~c6_2(X1,X2)|~c1_2(X1,X2)|~c3_2(X1,X2)|~c4_1(X1)|~c2_1(X1)|~ndr1_1(X1)|~ssSkC8), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_185, c_0_476])]), ['final']).
cnf(c_0_2175,negated_conjecture,(c8_0|c7_0|c1_0|~c6_2(a182,X1)|~c1_2(a182,X1)|~c7_2(a182,X1)), inference(spm,[status(thm)],[c_0_186, c_0_1320]), ['final']).
cnf(c_0_2176,negated_conjecture,(c1_1(X1)|c8_0|c2_1(a202)|ndr1_1(X1)|ssSkC11|~c7_1(X1)|~c5_1(a202)), inference(spm,[status(thm)],[c_0_941, c_0_50]), ['final']).
cnf(c_0_2177,negated_conjecture,(c1_1(X1)|c8_0|c2_1(a217)|ndr1_1(X1)|c1_0|~c7_1(X1)|~c5_1(a217)), inference(spm,[status(thm)],[c_0_941, c_0_509]), ['final']).
cnf(c_0_2178,negated_conjecture,(c1_1(X1)|c8_0|c7_0|c2_1(a215)|ndr1_1(X1)|~c7_1(X1)|~c5_1(a215)), inference(spm,[status(thm)],[c_0_941, c_0_112]), ['final']).
cnf(c_0_2179,negated_conjecture,(c6_2(a241,a206)|c8_2(a241,a207)|c3_0|ssSkC19|c4_0), inference(spm,[status(thm)],[c_0_1321, c_0_83]), ['final']).
cnf(c_0_2180,negated_conjecture,(c8_2(a241,a207)|c4_2(a241,a206)|c3_0|ssSkC19|c4_0), inference(spm,[status(thm)],[c_0_1322, c_0_83]), ['final']).
cnf(c_0_2181,negated_conjecture,(c1_1(a202)|c8_0|c2_1(a217)|ndr1_1(a202)|ssSkC11|c1_0|~c5_1(a217)), inference(spm,[status(thm)],[c_0_1323, c_0_509]), ['final']).
cnf(c_0_2182,negated_conjecture,(c1_1(a202)|c8_0|c7_0|c2_1(a215)|ndr1_1(a202)|ssSkC11|~c5_1(a215)), inference(spm,[status(thm)],[c_0_1323, c_0_112]), ['final']).
cnf(c_0_2183,negated_conjecture,(c1_1(a217)|c8_0|c2_1(a202)|ndr1_1(a217)|ssSkC11|c1_0|~c5_1(a202)), inference(spm,[status(thm)],[c_0_1324, c_0_50]), ['final']).
cnf(c_0_2184,negated_conjecture,(c1_1(a202)|c8_0|c2_1(a202)|ndr1_1(a202)|ssSkC11|~c5_1(a202)), inference(spm,[status(thm)],[c_0_1323, c_0_50]), ['final']).
cnf(c_0_2185,negated_conjecture,(c1_1(a217)|c8_0|c2_1(a217)|ndr1_1(a217)|c1_0|~c5_1(a217)), inference(spm,[status(thm)],[c_0_1324, c_0_509]), ['final']).
cnf(c_0_2186,negated_conjecture,(c2_2(a204,a151)|c6_1(X1)|c5_1(a204)|c8_1(X1)|c2_1(X1)|c1_0|~c3_1(a204)), inference(spm,[status(thm)],[c_0_1325, c_0_1326]), ['final']).
cnf(c_0_2187,negated_conjecture,(c6_1(X1)|c5_1(X2)|c8_1(X1)|c2_1(X1)|c1_0|~c1_2(X2,a151)|~c3_1(X2)), inference(spm,[status(thm)],[c_0_1327, c_0_4]), ['final']).
cnf(c_0_2188,negated_conjecture,(c8_2(a241,a207)|c3_2(a241,a206)|c3_0|ssSkC19|c4_0), inference(spm,[status(thm)],[c_0_1328, c_0_83]), ['final']).
cnf(c_0_2189,negated_conjecture,(c6_2(a268,a206)|c8_2(a268,a207)|c3_0|c4_0|c5_0|~c4_2(a268,a269)), inference(spm,[status(thm)],[c_0_1329, c_0_127]), ['final']).
cnf(c_0_2190,negated_conjecture,(ssSkC7|ssSkC20|c7_2(a168,a248)|c5_1(a221)|c3_0|ndr1_1(a221)|ssSkC14|c1_0), inference(spm,[status(thm)],[c_0_564, c_0_1301]), ['final']).
cnf(c_0_2191,negated_conjecture,(c8_2(a268,a207)|c4_2(a268,a206)|c3_0|c4_0|c5_0|~c4_2(a268,a269)), inference(spm,[status(thm)],[c_0_1330, c_0_127]), ['final']).
cnf(c_0_2192,negated_conjecture,(c8_2(a268,a207)|c3_2(a268,a206)|c3_0|c4_0|c5_0|~c4_2(a268,a269)), inference(spm,[status(thm)],[c_0_1331, c_0_127]), ['final']).
cnf(c_0_2193,negated_conjecture,(c6_1(X1)|c5_1(X2)|c8_1(X1)|c2_1(X1)|c1_0|~c5_2(X2,a151)|~c3_1(X2)), inference(spm,[status(thm)],[c_0_1332, c_0_4]), ['final']).
cnf(c_0_2194,negated_conjecture,(c5_2(a148,X1)|c6_0|c7_0|~c8_2(a148,X1)|~c4_2(a148,X1)), inference(spm,[status(thm)],[c_0_1333, c_0_1334]), ['final']).
cnf(c_0_2195,negated_conjecture,(c7_2(X1,X2)|c7_1(X1)|c8_0|~c6_2(X1,X2)|~c8_2(X1,a230)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_187, c_0_476])]), ['final']).
cnf(c_0_2196,negated_conjecture,(c7_2(X1,X2)|c7_1(X1)|c8_0|~c6_2(X1,a230)|~c6_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_188, c_0_476])]), ['final']).
cnf(c_0_2197,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,X2)|c6_0|~c5_2(X1,X2)|~c1_2(X1,a240)|~ssSkP2(X1)|~ndr1_1(X1)|~ssSkC18), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_189, c_0_476])]), ['final']).
cnf(c_0_2198,negated_conjecture,(c8_2(a275,a276)|c3_2(a275,a276)|c8_0|c2_1(a275)|ssSkC25|~c5_1(a275)|~ndr1_1(a275)), inference(spm,[status(thm)],[c_0_750, c_0_60]), ['final']).
cnf(c_0_2199,negated_conjecture,(c8_2(a271,a272)|c3_2(a271,a272)|c8_0|c2_1(a271)|ssSkC24), inference(spm,[status(thm)],[c_0_1335, c_0_153]), ['final']).
cnf(c_0_2200,negated_conjecture,(c8_2(a265,a266)|c3_2(a265,a266)|c8_0|c2_1(a265)|ssSkC23|~c5_1(a265)|~ndr1_1(a265)), inference(spm,[status(thm)],[c_0_750, c_0_190]), ['final']).
cnf(c_0_2201,negated_conjecture,(c8_2(a237,a238)|c3_2(a237,a238)|c8_0|c2_1(a237)|ssSkC18|~c5_1(a237)|~ndr1_1(a237)), inference(spm,[status(thm)],[c_0_750, c_0_191]), ['final']).
cnf(c_0_2202,negated_conjecture,(c8_2(a155,a156)|c3_2(a155,a156)|c8_0|c2_1(a155)|ssSkC1|~c5_1(a155)|~ndr1_1(a155)), inference(spm,[status(thm)],[c_0_750, c_0_192]), ['final']).
cnf(c_0_2203,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c4_0|~c1_2(X1,X3)|~c8_2(X1,X2)|~c4_2(X1,X2)|~c3_2(X1,a220)|~c7_2(X1,X3)|~ndr1_1(X1)|~ssSkC13), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_193, c_0_476])]), ['final']).
cnf(c_0_2204,negated_conjecture,(c5_2(X1,X2)|c1_2(X1,X2)|c2_1(X1)|c1_0|~c2_2(X1,a262)|~c3_2(X1,X2)|~c6_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_194, c_0_476])]), ['final']).
cnf(c_0_2205,negated_conjecture,(ssSkC20|c3_2(a217,a248)|c3_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_855, c_0_936]), ['final']).
cnf(c_0_2206,negated_conjecture,(ssSkC20|c3_2(a168,a248)|c3_0|c6_0|c2_0|ndr1_1(a217)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_656, c_0_936]), ['final']).
cnf(c_0_2207,negated_conjecture,(c6_2(X1,X2)|c8_2(X1,X3)|c4_0|~c5_2(X1,a220)|~c1_2(X1,X3)|~c8_2(X1,X2)|~c4_2(X1,X2)|~c7_2(X1,X3)|~ndr1_1(X1)|~ssSkC13), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_195, c_0_476])]), ['final']).
cnf(c_0_2208,negated_conjecture,(c7_2(X1,a230)|c7_2(X1,X2)|c7_1(X1)|c8_0|~c6_2(X1,X2)|~c4_2(X1,X2)|~ndr1_1(X1)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_196, c_0_476])]), ['final']).
cnf(c_0_2209,negated_conjecture,(c7_2(X1,X2)|~c5_2(X1,X2)|~c3_2(X1,a152)|~c3_2(X1,X2)|~c5_1(X1)|~c7_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_197, c_0_476])]), ['final']).
cnf(c_0_2210,negated_conjecture,(ssSkC7|ssSkC20|c3_2(a168,a248)|c5_1(a221)|c3_0|ndr1_1(a221)|ssSkC14|c1_0), inference(spm,[status(thm)],[c_0_711, c_0_1301]), ['final']).
cnf(c_0_2211,negated_conjecture,(ssSkC20|c3_2(a168,a248)|c3_0|ndr1_1(a287)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_113, c_0_1336]), ['final']).
cnf(c_0_2212,negated_conjecture,(c1_2(a229,X1)|c3_2(a229,X1)|c8_0|c1_0|~c4_2(a229,X1)|~ndr1_1(a229)), inference(spm,[status(thm)],[c_0_198, c_0_478]), ['final']).
cnf(c_0_2213,negated_conjecture,(c2_2(a221,X1)|c4_2(a221,X1)|c8_0|c7_0|ssSkC14|c1_0|~c7_2(a221,X1)), inference(spm,[status(thm)],[c_0_199, c_0_1337]), ['final']).
cnf(c_0_2214,negated_conjecture,(c5_2(X1,X2)|c5_2(X3,X4)|c3_2(X3,X4)|c7_2(X1,X2)|c7_2(X3,X4)|c6_1(X1)|~c4_2(X3,a291)|~c3_2(X1,X2)|~ssSkP9(X3)|~c2_1(X1)|~ndr1_1(X1)|~ndr1_1(X3)|~ssSkC27), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_200, c_0_476])]), ['final']).
cnf(c_0_2215,negated_conjecture,(ssSkC20|c5_1(X1)|c2_1(a173)|ndr1_1(a173)|ndr1_1(X1)|ssSkC3|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_513, c_0_1310]), ['final']).
cnf(c_0_2216,negated_conjecture,(c7_2(X1,X2)|~c5_2(X1,X2)|~c4_2(X1,a152)|~c3_2(X1,X2)|~c5_1(X1)|~c7_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_201, c_0_476])]), ['final']).
cnf(c_0_2217,negated_conjecture,(c6_2(X1,X2)|c4_2(X1,X2)|c3_2(X1,X2)|c3_1(X1)|~c2_2(X1,a175)|~ndr1_1(X1)|~ssSkC3|~ssSkC4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_202, c_0_476])]), ['final']).
cnf(c_0_2218,negated_conjecture,(c8_2(X1,X2)|c7_2(X1,a240)|c7_2(X1,X2)|c6_0|~c5_2(X1,X2)|~ssSkP2(X1)|~ndr1_1(X1)|~ssSkC18), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_203, c_0_476])]), ['final']).
cnf(c_0_2219,negated_conjecture,(c5_2(X1,X2)|c1_2(X1,X2)|c4_2(X1,a262)|c2_1(X1)|c1_0|~c3_2(X1,X2)|~c6_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_204, c_0_476])]), ['final']).
cnf(c_0_2220,negated_conjecture,(c5_2(X1,X2)|c1_2(X1,a262)|c1_2(X1,X2)|c2_1(X1)|c1_0|~c3_2(X1,X2)|~c6_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_205, c_0_476])]), ['final']).
cnf(c_0_2221,negated_conjecture,(c3_2(X1,X2)|c2_1(X1)|c5_0|~c5_2(X1,X2)|~c8_2(X1,X2)|~c7_1(X1)|~c7_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_206, c_0_476])]), ['final']).
cnf(c_0_2222,negated_conjecture,(c7_2(X1,X2)|c5_1(X1)|c4_0|~c6_2(X1,X2)|~c5_2(X1,X2)|~c7_1(X1)|~ndr1_1(X1)|~c5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_207, c_0_476])]), ['final']).
cnf(c_0_2223,negated_conjecture,(ssSkC20|c3_2(a217,a248)|c3_0|c6_0|c2_0|ndr1_1(a168)|ndr1_1(a223)|ndr1_1(a296)|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_855, c_0_603]), ['final']).
cnf(c_0_2224,negated_conjecture,(c7_2(X1,a152)|c7_2(X1,X2)|~c5_2(X1,X2)|~c3_2(X1,X2)|~c5_1(X1)|~c7_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_208, c_0_476])]), ['final']).
cnf(c_0_2225,negated_conjecture,(c6_2(X1,X2)|c4_2(X1,X2)|c3_2(X1,a175)|c3_2(X1,X2)|c3_1(X1)|~ndr1_1(X1)|~ssSkC3|~ssSkC4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_209, c_0_476])]), ['final']).
cnf(c_0_2226,negated_conjecture,(ssSkC7|ssSkC20|c5_1(a221)|c3_0|ndr1_1(a168)|ndr1_1(a221)|ssSkC14|c1_0), inference(spm,[status(thm)],[c_0_567, c_0_1301]), ['final']).
cnf(c_0_2227,negated_conjecture,(ssSkC20|c3_0|ndr1_1(a168)|ndr1_1(a287)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_113, c_0_1338]), ['final']).
cnf(c_0_2228,negated_conjecture,(ssSkC7|ssSkC20|c5_1(a168)|c3_0|ndr1_1(a168)|c1_0), inference(spm,[status(thm)],[c_0_567, c_0_1303]), ['final']).
cnf(c_0_2229,negated_conjecture,(c5_1(X1)|c2_1(a197)|ndr1_1(X1)|c5_0|c1_0|~c3_1(X1)), inference(spm,[status(thm)],[c_0_84, c_0_1339]), ['final']).
cnf(c_0_2230,negated_conjecture,(c6_2(X1,X2)|c2_2(X1,X3)|c1_2(X1,X3)|c4_2(X1,X3)|c7_2(a189,X4)|~c2_2(a189,X4)|~c8_2(a189,X4)|~c8_2(X1,X2)|~c3_2(X1,X2)|~c7_1(X1)|~ndr1_1(a189)|~ndr1_1(X1)|~ssSkC6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_210, c_0_476])]), ['final']).
cnf(c_0_2231,negated_conjecture,(c6_2(X1,X2)|c5_2(X1,a175)|c4_2(X1,X2)|c3_2(X1,X2)|c3_1(X1)|~ndr1_1(X1)|~ssSkC3|~ssSkC4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_211, c_0_476])]), ['final']).
cnf(c_0_2232,negated_conjecture,(c6_2(X1,X2)|c4_2(X1,X2)|c8_1(X1)|c4_0|~c1_2(X1,X2)|~c6_1(X1)|~c6_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_212, c_0_476])]), ['final']).
cnf(c_0_2233,negated_conjecture,(c5_2(X1,X2)|c7_2(X1,X2)|c1_1(X1)|c6_0|c8_1(X1)|~c1_2(X1,X2)|~ndr1_1(X1)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_213, c_0_476])]), ['final']).
cnf(c_0_2234,negated_conjecture,(c1_2(a182,X1)|c8_0|c7_0|c1_0|~c2_2(a182,X1)|~c4_2(a182,X1)), inference(spm,[status(thm)],[c_0_214, c_0_1320]), ['final']).
cnf(c_0_2235,negated_conjecture,(c2_2(a229,X1)|c1_2(a229,X1)|c8_2(a229,X1)|c8_0|c1_0|~ndr1_1(a229)), inference(spm,[status(thm)],[c_0_215, c_0_478]), ['final']).
cnf(c_0_2236,negated_conjecture,(c5_2(a275,a278)|c8_2(a275,a292)|c4_2(a275,a293)|c8_0|c7_0|ssSkC25|c1_0), inference(spm,[status(thm)],[c_0_115, c_0_562]), ['final']).
cnf(c_0_2237,negated_conjecture,(c5_2(a218,a278)|c8_2(a218,a292)|c4_2(a218,a293)|c8_0|c7_0|ssSkC13|c1_0), inference(spm,[status(thm)],[c_0_116, c_0_562]), ['final']).
cnf(c_0_2238,negated_conjecture,(c5_2(a200,a278)|c8_2(a200,a292)|c4_2(a200,a293)|c8_0|c7_0|ssSkC10|c1_0), inference(spm,[status(thm)],[c_0_97, c_0_562]), ['final']).
cnf(c_0_2239,negated_conjecture,(c5_2(a221,a278)|c8_2(a221,a292)|c4_2(a221,a293)|c8_0|c7_0|ssSkC14|c1_0), inference(spm,[status(thm)],[c_0_118, c_0_562]), ['final']).
cnf(c_0_2240,negated_conjecture,(c6_2(a189,X1)|c6_2(X2,X3)|c2_2(X2,X4)|c1_2(X2,X4)|c4_2(X2,X4)|~c1_2(a189,X1)|~c8_2(a189,X1)|~c8_2(X2,X3)|~c3_2(X2,X3)|~c7_1(X2)|~ndr1_1(a189)|~ndr1_1(X2)|~ssSkC6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_216, c_0_476])]), ['final']).
cnf(c_0_2241,negated_conjecture,(c2_2(X1,X2)|c4_2(X1,X2)|c8_1(a203)|~c5_2(X1,X2)|~c2_1(X1)|~ndr1_1(X1)|~ssSkC11), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_217, c_0_476])]), ['final']).
cnf(c_0_2242,negated_conjecture,(c2_2(X1,X2)|c4_2(X1,X2)|c4_1(a203)|~c5_2(X1,X2)|~c2_1(X1)|~ndr1_1(X1)|~ssSkC11), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_218, c_0_476])]), ['final']).
cnf(c_0_2243,negated_conjecture,(c3_1(X1)|c1_0|~c2_2(X1,X2)|~c4_2(X1,X2)|~c3_2(X1,X2)|~c7_0|~ndr1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_219, c_0_476])]), ['final']).
cnf(c_0_2244,negated_conjecture,(c2_2(X1,X2)|c7_2(X1,a201)|c7_2(X1,X2)|~c2_1(X1)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_220, c_0_476])]), ['final']).
cnf(c_0_2245,negated_conjecture,(c2_2(X1,X2)|c8_2(X1,a201)|c7_2(X1,X2)|~c2_1(X1)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_221, c_0_476])]), ['final']).
cnf(c_0_2246,negated_conjecture,(c2_2(X1,X2)|c1_2(X1,a201)|c7_2(X1,X2)|~c2_1(X1)|~ndr1_1(X1)|~ssSkC9|~ssSkC10), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_222, c_0_476])]), ['final']).
cnf(c_0_2247,negated_conjecture,(c5_2(a217,a278)|c8_2(a217,a292)|c4_2(a217,a293)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_501, c_0_562]), ['final']).
cnf(c_0_2248,negated_conjecture,(c5_2(a182,a278)|c8_2(a182,a292)|c4_2(a182,a293)|c8_0|c7_0|c1_0), inference(spm,[status(thm)],[c_0_117, c_0_562]), ['final']).
cnf(c_0_2249,negated_conjecture,(c4_2(a181,X1)|c5_0|~c6_2(X2,X3)|~c2_2(a181,X1)|~c5_2(X2,X3)|~c1_2(X2,X4)|~c8_2(X2,X4)|~c3_2(a181,X1)|~c3_2(X2,X4)|~c3_1(X2)|~ndr1_1(a181)|~ndr1_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_223, c_0_476])]), ['final']).
cnf(c_0_2250,negated_conjecture,(c5_0|~c6_2(a181,X1)|~c6_2(X2,X3)|~c2_2(a181,X1)|~c5_2(X2,X3)|~c1_2(X2,X4)|~c8_2(X2,X4)|~c4_2(a181,X1)|~c3_2(X2,X4)|~c3_1(X2)|~ndr1_1(a181)|~ndr1_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_224, c_0_476])]), ['final']).
cnf(c_0_2251,negated_conjecture,(c8_2(a158,X1)|c4_2(a158,X1)|c5_0|c1_0|~c1_2(a158,X1)), inference(spm,[status(thm)],[c_0_225, c_0_869]), ['final']).
cnf(c_0_2252,negated_conjecture,(c3_1(X1)|c1_1(X1)|c6_0|~c7_2(a256,a257)|~c5_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_226, c_0_476])]), ['final']).
cnf(c_0_2253,negated_conjecture,(c3_1(X1)|c1_1(X1)|c6_0|~c1_2(a256,a257)|~c5_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_227, c_0_476])]), ['final']).
cnf(c_0_2254,negated_conjecture,(c2_2(a168,X1)|c3_0|c1_0|~c5_2(a168,X1)|~c8_2(a168,X1)|~ndr1_1(a168)), inference(spm,[status(thm)],[c_0_228, c_0_4]), ['final']).
cnf(c_0_2255,negated_conjecture,(c4_1(X1)|c8_0|c1_0|~c3_2(X1,a293)|~ssSkP10(X1)), inference(spm,[status(thm)],[c_0_1340, c_0_4]), ['final']).
cnf(c_0_2256,negated_conjecture,(c5_2(X1,X2)|~c6_2(X1,X2)|~c7_1(a274)|~c1_1(X1)|~ndr1_1(X1)|~ssSkC24), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_229, c_0_476])]), ['final']).
cnf(c_0_2257,negated_conjecture,(c1_1(X1)|c5_1(X2)|~c3_2(X2,a255)|~c3_1(X1)|~ssSkP4(X2)|~ssSkC21), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_230, c_0_476])]), ['final']).
cnf(c_0_2258,negated_conjecture,(c1_1(X1)|c5_1(X2)|~c8_2(X2,a255)|~c3_1(X1)|~ssSkP4(X2)|~ssSkC21), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_231, c_0_476])]), ['final']).
cnf(c_0_2259,negated_conjecture,(c4_1(X1)|c5_0|c1_0|~c3_2(X1,a293)|~ssSkP10(X1)), inference(spm,[status(thm)],[c_0_1341, c_0_4]), ['final']).
cnf(c_0_2260,negated_conjecture,(c4_1(X1)|c8_0|c1_0|~c1_2(X1,a293)|~ssSkP10(X1)), inference(spm,[status(thm)],[c_0_1342, c_0_4]), ['final']).
cnf(c_0_2261,negated_conjecture,(c4_1(X1)|c5_0|c1_0|~c1_2(X1,a293)|~ssSkP10(X1)), inference(spm,[status(thm)],[c_0_1343, c_0_4]), ['final']).
cnf(c_0_2262,negated_conjecture,(c3_1(X1)|c4_1(X1)|c5_1(X1)|c6_0|~c7_2(a296,a297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_232, c_0_476])]), ['final']).
cnf(c_0_2263,negated_conjecture,(c3_1(X1)|c4_1(X1)|c5_1(X1)|c6_0|~c2_2(a296,a297)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_233, c_0_476])]), ['final']).
cnf(c_0_2264,negated_conjecture,(c1_1(X1)|c8_0|c7_0|~c3_2(X1,a278)|~c4_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_234, c_0_476])]), ['final']).
cnf(c_0_2265,negated_conjecture,(c7_2(a211,a185)|c6_0|c8_0|ssSkC12|c1_0|~c5_1(a211)), inference(spm,[status(thm)],[c_0_235, c_0_582]), ['final']).
cnf(c_0_2266,negated_conjecture,(c5_2(a186,X1)|c6_0|c8_0|c2_0|~c1_2(a186,X1)|~c8_2(a186,X1)|~ndr1_1(a186)), inference(spm,[status(thm)],[c_0_236, c_0_19]), ['final']).
cnf(c_0_2267,negated_conjecture,(c7_2(a187,a185)|c6_0|c8_0|ssSkC6|c1_0|~c5_1(a187)), inference(spm,[status(thm)],[c_0_237, c_0_582]), ['final']).
cnf(c_0_2268,negated_conjecture,(c2_2(a222,X1)|c5_2(a222,X1)|c8_0|c1_0|~ndr1_1(a222)|~ssSkC14), inference(spm,[status(thm)],[c_0_238, c_0_478]), ['final']).
cnf(c_0_2269,negated_conjecture,(c7_2(X1,a255)|c1_1(X2)|c5_1(X1)|~c3_1(X2)|~ssSkP4(X1)|~ssSkC21), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_239, c_0_476])]), ['final']).
cnf(c_0_2270,negated_conjecture,(c3_1(X1)|c5_1(X1)|~c1_2(a171,a172)|~c4_1(X1)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_240, c_0_476])]), ['final']).
cnf(c_0_2271,negated_conjecture,(c4_2(a211,a185)|c6_0|c8_0|ssSkC12|c1_0|~c5_1(a211)), inference(spm,[status(thm)],[c_0_235, c_0_499]), ['final']).
cnf(c_0_2272,negated_conjecture,(c4_2(a187,a185)|c6_0|c8_0|ssSkC6|c1_0|~c5_1(a187)), inference(spm,[status(thm)],[c_0_237, c_0_499]), ['final']).
cnf(c_0_2273,negated_conjecture,(c5_2(a211,a185)|c6_0|c8_0|ssSkC12|c1_0|~c5_1(a211)), inference(spm,[status(thm)],[c_0_235, c_0_722]), ['final']).
cnf(c_0_2274,negated_conjecture,(c5_2(a187,a185)|c6_0|c8_0|ssSkC6|c1_0|~c5_1(a187)), inference(spm,[status(thm)],[c_0_237, c_0_722]), ['final']).
cnf(c_0_2275,negated_conjecture,(c7_2(a294,a185)|c6_0|c8_0|c1_0|~c5_1(a294)), inference(spm,[status(thm)],[c_0_241, c_0_582]), ['final']).
cnf(c_0_2276,negated_conjecture,(ssSkC27|~c5_2(a288,X1)|~c8_2(a288,X1)|~c4_2(a288,X1)), inference(spm,[status(thm)],[c_0_242, c_0_243]), ['final']).
cnf(c_0_2277,negated_conjecture,(c4_2(a294,a185)|c6_0|c8_0|c1_0|~c5_1(a294)), inference(spm,[status(thm)],[c_0_241, c_0_499]), ['final']).
cnf(c_0_2278,negated_conjecture,(c7_2(X1,a264)|c6_0|c8_0|~c6_1(X1)|~ssSkP6(X1)|~c5_0), inference(spm,[status(thm)],[c_0_1344, c_0_19]), ['final']).
cnf(c_0_2279,negated_conjecture,(c1_2(X1,a264)|c6_0|c8_0|~c6_1(X1)|~ssSkP6(X1)|~c5_0), inference(spm,[status(thm)],[c_0_1345, c_0_19]), ['final']).
cnf(c_0_2280,negated_conjecture,(c5_2(X1,a264)|c6_0|c8_0|~c6_1(X1)|~ssSkP6(X1)|~c5_0), inference(spm,[status(thm)],[c_0_1346, c_0_19]), ['final']).
cnf(c_0_2281,negated_conjecture,(c1_1(X1)|c5_1(X2)|ndr1_1(X2)|~c3_1(X1)|~ssSkP4(X2)|~ssSkC21), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_244, c_0_476])]), ['final']).
cnf(c_0_2282,negated_conjecture,(c8_2(a223,a292)|c4_2(a223,a293)|c8_0|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_46, c_0_534]), ['final']).
cnf(c_0_2283,negated_conjecture,(c3_2(a155,a156)|ssSkC1|~c4_2(a155,a156)), inference(spm,[status(thm)],[c_0_1347, c_0_192]), ['final']).
cnf(c_0_2284,negated_conjecture,(c5_2(a294,a185)|c6_0|c8_0|c1_0|~c5_1(a294)), inference(spm,[status(thm)],[c_0_241, c_0_722]), ['final']).
cnf(c_0_2285,negated_conjecture,(c8_2(a223,a292)|c4_2(a223,a293)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_46, c_0_636]), ['final']).
cnf(c_0_2286,negated_conjecture,(c8_2(a158,a292)|c4_2(a158,a293)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_71, c_0_636]), ['final']).
cnf(c_0_2287,negated_conjecture,(c5_2(a211,a278)|c1_1(a211)|c8_0|c7_0|ssSkC12), inference(spm,[status(thm)],[c_0_488, c_0_245]), ['final']).
cnf(c_0_2288,negated_conjecture,(c7_2(a176,X1)|ssSkC5|~c2_2(a176,X1)|~c3_2(a176,X1)), inference(spm,[status(thm)],[c_0_246, c_0_247]), ['final']).
cnf(c_0_2289,negated_conjecture,(c5_2(a187,a278)|c1_1(a187)|c8_0|c7_0|ssSkC6), inference(spm,[status(thm)],[c_0_488, c_0_248]), ['final']).
cnf(c_0_2290,negated_conjecture,(c5_2(a171,a172)|c3_1(X1)|c5_1(X1)|~c4_1(X1)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_249, c_0_476])]), ['final']).
cnf(c_0_2291,negated_conjecture,(c5_2(a154,a278)|c1_1(a154)|c8_0|c7_0|ssSkC0), inference(spm,[status(thm)],[c_0_488, c_0_250]), ['final']).
cnf(c_0_2292,negated_conjecture,(c1_1(X1)|ssSkC22|~c7_2(X1,a258)|~c7_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_251, c_0_476])]), ['final']).
cnf(c_0_2293,negated_conjecture,(ssSkC20|c2_1(X1)|~c6_2(X1,a248)|~c8_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_252, c_0_476])]), ['final']).
cnf(c_0_2294,negated_conjecture,(c3_1(X1)|c4_1(X1)|c5_1(X1)|c6_0|~c7_1(a296)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_253, c_0_476])]), ['final']).
cnf(c_0_2295,negated_conjecture,(c2_2(a171,a172)|c3_1(X1)|c5_1(X1)|~c4_1(X1)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_254, c_0_476])]), ['final']).
cnf(c_0_2296,negated_conjecture,(c1_0|~c4_2(X1,a236)|~c6_1(X1)|~c8_1(X1)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_255, c_0_476])]), ['final']).
cnf(c_0_2297,negated_conjecture,(c1_0|~c6_2(X1,a236)|~c6_1(X1)|~c8_1(X1)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_256, c_0_476])]), ['final']).
cnf(c_0_2298,negated_conjecture,(c1_1(a294)|c6_0|ndr1_1(a256)|c1_0|~c5_1(a294)), inference(spm,[status(thm)],[c_0_241, c_0_524]), ['final']).
cnf(c_0_2299,negated_conjecture,(c1_1(a211)|c6_0|ndr1_1(a256)|ssSkC12|~c5_1(a211)), inference(spm,[status(thm)],[c_0_235, c_0_524]), ['final']).
cnf(c_0_2300,negated_conjecture,(c1_2(a229,X1)|c3_2(a229,X1)|c5_0|c1_0|~c4_2(a229,X1)|~ndr1_1(a229)), inference(spm,[status(thm)],[c_0_198, c_0_20]), ['final']).
cnf(c_0_2301,negated_conjecture,(c1_1(a187)|c6_0|ndr1_1(a256)|ssSkC6|~c5_1(a187)), inference(spm,[status(thm)],[c_0_237, c_0_524]), ['final']).
cnf(c_0_2302,negated_conjecture,(c1_1(X1)|c1_0|~c4_2(X1,a225)|~ssSkP1(X1)|~ssSkC15), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_257, c_0_476])]), ['final']).
cnf(c_0_2303,negated_conjecture,(c2_1(X1)|c4_0|~c2_2(X1,a214)|~c1_1(X1)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_258, c_0_476])]), ['final']).
cnf(c_0_2304,negated_conjecture,(c5_1(X1)|c7_0|~c3_2(X1,a213)|~c4_1(X1)|~ssSkC12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_259, c_0_476])]), ['final']).
cnf(c_0_2305,negated_conjecture,(c3_1(X1)|c8_0|~c3_2(X1,a251)|~ssSkP3(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_260, c_0_476])]), ['final']).
cnf(c_0_2306,negated_conjecture,(c3_1(X1)|c8_0|~c8_2(X1,a251)|~ssSkP3(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_261, c_0_476])]), ['final']).
cnf(c_0_2307,negated_conjecture,(c3_1(X1)|c8_0|~c2_2(X1,a251)|~ssSkP3(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_262, c_0_476])]), ['final']).
cnf(c_0_2308,negated_conjecture,(c3_1(X1)|c5_1(X1)|~c4_1(X1)|~c2_1(a171)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_263, c_0_476])]), ['final']).
cnf(c_0_2309,negated_conjecture,(c1_1(a256)|c1_1(a294)|c6_0|c1_0|~c5_1(a294)), inference(spm,[status(thm)],[c_0_241, c_0_522]), ['final']).
cnf(c_0_2310,negated_conjecture,(c1_1(a256)|c1_1(a211)|c6_0|ssSkC12|~c5_1(a211)), inference(spm,[status(thm)],[c_0_235, c_0_522]), ['final']).
cnf(c_0_2311,negated_conjecture,(c1_1(a256)|c1_1(a187)|c6_0|ssSkC6|~c5_1(a187)), inference(spm,[status(thm)],[c_0_237, c_0_522]), ['final']).
cnf(c_0_2312,negated_conjecture,(c6_2(a241,X1)|c5_2(a241,X1)|ssSkC19|~c4_2(a241,X1)), inference(spm,[status(thm)],[c_0_264, c_0_122]), ['final']).
cnf(c_0_2313,negated_conjecture,(c6_0|c8_0|ndr1_1(X1)|~c6_1(X1)|~ssSkP6(X1)|~c5_0), inference(spm,[status(thm)],[c_0_1348, c_0_19]), ['final']).
cnf(c_0_2314,negated_conjecture,(c5_2(X1,a236)|c1_0|~c6_1(X1)|~c8_1(X1)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_265, c_0_476])]), ['final']).
cnf(c_0_2315,negated_conjecture,(ssSkC7|c3_1(X1)|c2_1(X1)|~c3_2(X1,a190)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_266, c_0_476])]), ['final']).
cnf(c_0_2316,negated_conjecture,(ssSkC7|c3_1(X1)|c2_1(X1)|~c6_2(X1,a190)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_267, c_0_476])]), ['final']).
cnf(c_0_2317,negated_conjecture,(c7_2(a223,a260)|ssSkP5(a223)|c8_0|ssSkC15|c1_0), inference(spm,[status(thm)],[c_0_268, c_0_1300]), ['final']).
cnf(c_0_2318,negated_conjecture,(c7_2(a223,a260)|ssSkP5(a223)|ssSkC15|c5_0|c1_0), inference(spm,[status(thm)],[c_0_268, c_0_1349]), ['final']).
cnf(c_0_2319,negated_conjecture,(c6_2(a279,X1)|c3_2(a279,X1)|ssSkC26|~c5_2(a279,X1)), inference(spm,[status(thm)],[c_0_269, c_0_270]), ['final']).
cnf(c_0_2320,negated_conjecture,(c8_0|c7_0|ndr1_1(a200)|ssSkC10|c1_0), inference(spm,[status(thm)],[c_0_97, c_0_559]), ['final']).
cnf(c_0_2321,negated_conjecture,(ssSkC20|c7_2(a287,a248)|c3_0|c2_1(a287)|c5_0), inference(spm,[status(thm)],[c_0_506, c_0_271]), ['final']).
cnf(c_0_2322,negated_conjecture,(c8_2(a162,X1)|c3_2(a162,X1)|c5_0|~c7_2(a162,X1)), inference(spm,[status(thm)],[c_0_272, c_0_273]), ['final']).
cnf(c_0_2323,negated_conjecture,(c1_1(a154)|c8_0|c7_0|ndr1_1(a154)|ssSkC0), inference(spm,[status(thm)],[c_0_529, c_0_250]), ['final']).
cnf(c_0_2324,negated_conjecture,(c3_2(a202,a258)|c1_1(a202)|ssSkC22|ssSkC11), inference(spm,[status(thm)],[c_0_827, c_0_50]), ['final']).
cnf(c_0_2325,negated_conjecture,(c3_2(a217,a258)|ssSkC22|c1_0), inference(spm,[status(thm)],[c_0_501, c_0_1350]), ['final']).
cnf(c_0_2326,negated_conjecture,(c4_2(a202,a258)|c1_1(a202)|ssSkC22|ssSkC11), inference(spm,[status(thm)],[c_0_589, c_0_50]), ['final']).
cnf(c_0_2327,negated_conjecture,(c2_2(a229,X1)|c1_2(a229,X1)|c8_2(a229,X1)|c5_0|c1_0|~ndr1_1(a229)), inference(spm,[status(thm)],[c_0_215, c_0_20]), ['final']).
cnf(c_0_2328,negated_conjecture,(ssSkC20|c7_2(a204,a248)|c2_1(a204)|c1_0), inference(spm,[status(thm)],[c_0_506, c_0_274]), ['final']).
cnf(c_0_2329,negated_conjecture,(c3_1(X1)|c5_1(X1)|ndr1_1(a171)|~c4_1(X1)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_275, c_0_476])]), ['final']).
cnf(c_0_2330,negated_conjecture,(c3_1(X1)|c4_1(a171)|c5_1(X1)|~c4_1(X1)|~c1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_276, c_0_476])]), ['final']).
cnf(c_0_2331,negated_conjecture,(c8_2(X1,a225)|c1_1(X1)|c1_0|~ssSkP1(X1)|~ssSkC15), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_277, c_0_476])]), ['final']).
cnf(c_0_2332,negated_conjecture,(c2_2(X1,a225)|c1_1(X1)|c1_0|~ssSkP1(X1)|~ssSkC15), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_278, c_0_476])]), ['final']).
cnf(c_0_2333,negated_conjecture,(ssSkC20|c3_2(a287,a248)|c3_0|c2_1(a287)|c5_0), inference(spm,[status(thm)],[c_0_512, c_0_271]), ['final']).
cnf(c_0_2334,negated_conjecture,(ssSkC20|c3_2(a204,a248)|c2_1(a204)|c1_0), inference(spm,[status(thm)],[c_0_512, c_0_274]), ['final']).
cnf(c_0_2335,negated_conjecture,(ssSkC2|~c2_2(a164,X1)|~c8_2(a164,X1)), inference(spm,[status(thm)],[c_0_279, c_0_280]), ['final']).
cnf(c_0_2336,negated_conjecture,(c2_2(a222,X1)|c5_2(a222,X1)|c5_0|c1_0|~ndr1_1(a222)|~ssSkC14), inference(spm,[status(thm)],[c_0_238, c_0_20]), ['final']).
cnf(c_0_2337,negated_conjecture,(c7_2(X1,a214)|c2_1(X1)|c4_0|~c1_1(X1)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_281, c_0_476])]), ['final']).
cnf(c_0_2338,negated_conjecture,(c4_2(X1,a214)|c2_1(X1)|c4_0|~c1_1(X1)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_282, c_0_476])]), ['final']).
cnf(c_0_2339,negated_conjecture,(c7_2(X1,a213)|c5_1(X1)|c7_0|~c4_1(X1)|~ssSkC12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_283, c_0_476])]), ['final']).
cnf(c_0_2340,negated_conjecture,(c2_2(X1,a213)|c5_1(X1)|c7_0|~c4_1(X1)|~ssSkC12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_284, c_0_476])]), ['final']).
cnf(c_0_2341,negated_conjecture,(ssSkC17|~c4_2(a233,X1)|~c3_2(a233,X1)), inference(spm,[status(thm)],[c_0_285, c_0_63]), ['final']).
cnf(c_0_2342,negated_conjecture,(ssSkC7|c1_2(a215,a190)|c3_1(a215)|c8_0|c7_0), inference(spm,[status(thm)],[c_0_111, c_0_746]), ['final']).
cnf(c_0_2343,negated_conjecture,(ssSkC7|c1_2(a168,a190)|c3_1(a168)|c3_0|c1_0), inference(spm,[status(thm)],[c_0_537, c_0_746]), ['final']).
cnf(c_0_2344,negated_conjecture,(ssSkC20|c3_0|c2_1(a287)|ndr1_1(a287)|c5_0), inference(spm,[status(thm)],[c_0_513, c_0_271]), ['final']).
cnf(c_0_2345,negated_conjecture,(c1_1(a202)|ssSkC22|ndr1_1(a202)|ssSkC11), inference(spm,[status(thm)],[c_0_600, c_0_50]), ['final']).
cnf(c_0_2346,negated_conjecture,(c6_2(a233,X1)|ssSkC17|~c8_2(a233,X1)), inference(spm,[status(thm)],[c_0_286, c_0_63]), ['final']).
cnf(c_0_2347,negated_conjecture,(ndr1_1(X1)|c1_0|~c6_1(X1)|~c8_1(X1)|~c7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_287, c_0_476])]), ['final']).
cnf(c_0_2348,negated_conjecture,(c1_1(X1)|ndr1_1(X1)|c1_0|~ssSkP1(X1)|~ssSkC15), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_288, c_0_476])]), ['final']).
cnf(c_0_2349,negated_conjecture,(c1_0|~c7_1(a252)|~c8_1(X1)), inference(spm,[status(thm)],[c_0_1351, c_0_4]), ['final']).
cnf(c_0_2350,negated_conjecture,(ssSkC7|c1_2(a221,a190)|c3_1(a221)|ssSkC14), inference(spm,[status(thm)],[c_0_133, c_0_746]), ['final']).
cnf(c_0_2351,negated_conjecture,(c2_2(a204,a285)|ssSkP8(a204)|c1_0), inference(spm,[status(thm)],[c_0_289, c_0_1326]), ['final']).
cnf(c_0_2352,negated_conjecture,(c2_2(a204,a290)|ssSkP9(a204)|c1_0), inference(spm,[status(thm)],[c_0_290, c_0_1326]), ['final']).
cnf(c_0_2353,negated_conjecture,(c6_0|c7_0|~c1_2(a148,a150)), inference(spm,[status(thm)],[c_0_291, c_0_724]), ['final']).
cnf(c_0_2354,negated_conjecture,(c6_0|c7_0|~c1_2(a148,a149)), inference(spm,[status(thm)],[c_0_292, c_0_724]), ['final']).
cnf(c_0_2355,negated_conjecture,(c2_1(X1)|ndr1_1(X1)|c4_0|~c1_1(X1)|~c2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_293, c_0_476])]), ['final']).
cnf(c_0_2356,negated_conjecture,(c5_1(X1)|c7_0|ndr1_1(X1)|~c4_1(X1)|~ssSkC12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_294, c_0_476])]), ['final']).
cnf(c_0_2357,negated_conjecture,(c3_1(X1)|c8_0|ndr1_1(X1)|~ssSkP3(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_295, c_0_476])]), ['final']).
cnf(c_0_2358,negated_conjecture,(c1_1(a252)|c1_0), inference(spm,[status(thm)],[c_0_1352, c_0_510]), ['final']).
cnf(c_0_2359,negated_conjecture,(c6_0|c7_0|~c8_2(a148,a150)), inference(spm,[status(thm)],[c_0_296, c_0_724]), ['final']).
cnf(c_0_2360,negated_conjecture,(c8_0|c1_0|~c6_1(a222)|~ssSkC14), inference(spm,[status(thm)],[c_0_297, c_0_478]), ['final']).
cnf(c_0_2361,negated_conjecture,(c8_0|c1_0|~c1_1(a222)|~ssSkC14), inference(spm,[status(thm)],[c_0_298, c_0_478]), ['final']).
cnf(c_0_2362,negated_conjecture,(c6_0|c8_0|c2_0|~c7_1(a186)), inference(spm,[status(thm)],[c_0_299, c_0_19]), ['final']).
cnf(c_0_2363,negated_conjecture,(c2_2(a148,a149)|c6_0|c7_0), inference(spm,[status(thm)],[c_0_300, c_0_724]), ['final']).
cnf(c_0_2364,negated_conjecture,(c1_1(a186)|c6_0|c8_0|c2_0), inference(spm,[status(thm)],[c_0_301, c_0_19]), ['final']).
cnf(c_0_2365,negated_conjecture,(c5_0|c1_0|~c6_1(a222)|~ssSkC14), inference(spm,[status(thm)],[c_0_297, c_0_20]), ['final']).
cnf(c_0_2366,negated_conjecture,(c5_0|c1_0|~c1_1(a222)|~ssSkC14), inference(spm,[status(thm)],[c_0_298, c_0_20]), ['final']).
cnf(c_0_2367,negated_conjecture,(c4_2(X1,a210)|~c4_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_302, c_0_476])]), ['final']).
cnf(c_0_2368,negated_conjecture,(c8_2(X1,a210)|~c4_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_303, c_0_476])]), ['final']).
cnf(c_0_2369,negated_conjecture,(c1_1(a229)|c8_0|c1_0), inference(spm,[status(thm)],[c_0_304, c_0_478]), ['final']).
cnf(c_0_2370,negated_conjecture,(ndr1_1(X1)|~c4_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_305, c_0_476])]), ['final']).
cnf(c_0_2371,negated_conjecture,(c6_0|c7_0|~c2_0|~c1_0), inference(spm,[status(thm)],[c_0_306, c_0_724]), ['final']).
cnf(c_0_2372,negated_conjecture,(c1_1(a229)|c5_0|c1_0), inference(spm,[status(thm)],[c_0_304, c_0_20]), ['final']).
cnf(c_0_2373,negated_conjecture,(~c6_2(X1,X2)|~c1_2(a196,X3)|~c1_2(X1,X2)|~c4_2(a196,X3)|~c3_2(X1,X2)|~c4_1(X1)|~c2_1(X1)|~ndr1_1(a196)|~ndr1_1(X1)|~ssSkC8), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_307, c_0_476])]), ['final']).
cnf(c_0_2374,negated_conjecture,(~c6_2(X1,X2)|~c1_2(X1,X2)|~c3_2(X1,X2)|~c4_1(a196)|~c4_1(X1)|~c2_1(X1)|~ndr1_1(X1)|~ssSkC8), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_308, c_0_476])]), ['final']).
cnf(c_0_2375,negated_conjecture,(~c5_2(X1,a210)|~c4_1(X1)|~c6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_309, c_0_476])]), ['final']).
cnf(c_0_2376,negated_conjecture,(c5_2(a249,X1)|~ssSkC20|~c2_2(a249,X1)|~c4_2(a249,X1)|~c2_0|~ndr1_1(a249)), c_0_310, ['final']).
cnf(c_0_2377,negated_conjecture,(c2_2(a168,X1)|c3_0|~c5_2(a168,X1)|~c8_2(a168,X1)|~ndr1_1(a168)|~c4_0), c_0_228, ['final']).
cnf(c_0_2378,negated_conjecture,(c5_2(a186,X1)|c2_0|~c1_2(a186,X1)|~c8_2(a186,X1)|~c3_0|~ndr1_1(a186)), c_0_236, ['final']).
cnf(c_0_2379,negated_conjecture,(c5_2(a148,X1)|c7_0|~c8_2(a148,X1)|~c4_2(a148,X1)|~c8_0|~ndr1_1(a148)), c_0_142, ['final']).
cnf(c_0_2380,negated_conjecture,(c7_2(a277,X1)|c4_0|~c2_2(a277,X1)|~c1_2(a277,X1)|~ndr1_1(a277)|~ssSkC25), c_0_311, ['final']).
cnf(c_0_2381,negated_conjecture,(c3_2(a170,X1)|c6_0|c7_0|~c2_2(a170,X1)|~c4_2(a170,X1)|~ndr1_1(a170)), c_0_312, ['final']).
cnf(c_0_2382,negated_conjecture,(c1_0|~c6_2(a182,X1)|~c1_2(a182,X1)|~c7_2(a182,X1)|~ndr1_1(a182)), c_0_186, ['final']).
cnf(c_0_2383,negated_conjecture,(c2_2(a227,X1)|~c4_2(a227,X1)|~c7_2(a227,X1)|~ndr1_1(a227)|~ssSkC16|~c1_0), c_0_313, ['final']).
cnf(c_0_2384,negated_conjecture,(c2_2(a227,X1)|c5_2(a227,X1)|~c4_2(a227,X1)|~ndr1_1(a227)|~ssSkC16|~c1_0), c_0_314, ['final']).
cnf(c_0_2385,negated_conjecture,(c5_2(a153,X1)|c8_2(a153,X1)|~c3_2(a153,X1)|~ndr1_1(a153)|~c4_0|~c1_0), c_0_315, ['final']).
cnf(c_0_2386,negated_conjecture,(c6_2(a277,X1)|c4_2(a277,X1)|c4_0|~c8_2(a277,X1)|~ndr1_1(a277)|~ssSkC25), c_0_316, ['final']).
cnf(c_0_2387,negated_conjecture,(c6_2(a174,X1)|ssSkC4|~c5_2(a174,X1)|~c8_2(a174,X1)|~ndr1_1(a174)), c_0_138, ['final']).
cnf(c_0_2388,negated_conjecture,(c7_2(a226,X1)|ssSkC16|~c6_2(a226,X1)|~c5_2(a226,X1)|~ndr1_1(a226)), c_0_317, ['final']).
cnf(c_0_2389,negated_conjecture,(c6_2(a202,X1)|c7_2(a202,X1)|ssSkC11|~c3_2(a202,X1)|~ndr1_1(a202)), c_0_318, ['final']).
cnf(c_0_2390,negated_conjecture,(c6_2(a226,X1)|c5_2(a226,X1)|ssSkC16|~c8_2(a226,X1)|~ndr1_1(a226)), c_0_319, ['final']).
cnf(c_0_2391,negated_conjecture,(c1_2(a182,X1)|c1_0|~c2_2(a182,X1)|~c4_2(a182,X1)|~ndr1_1(a182)), c_0_214, ['final']).
cnf(c_0_2392,negated_conjecture,(c1_2(a259,X1)|c8_2(a259,X1)|c3_2(a259,X1)|c7_0|~ssSkC22|~ndr1_1(a259)), c_0_320, ['final']).
cnf(c_0_2393,negated_conjecture,(c1_2(a229,X1)|c3_2(a229,X1)|~c4_2(a229,X1)|~c2_0|~ndr1_1(a229)), c_0_198, ['final']).
cnf(c_0_2394,negated_conjecture,(c5_2(a173,X1)|c8_2(a173,X1)|ssSkC3|~c7_2(a173,X1)|~ndr1_1(a173)), c_0_321, ['final']).
cnf(c_0_2395,negated_conjecture,(c5_2(a195,X1)|c4_2(a195,X1)|ssSkC8|~c6_2(a195,X1)|~ndr1_1(a195)), c_0_322, ['final']).
cnf(c_0_2396,negated_conjecture,(c8_2(a202,X1)|c3_2(a202,X1)|ssSkC11|~c6_2(a202,X1)|~ndr1_1(a202)), c_0_323, ['final']).
cnf(c_0_2397,negated_conjecture,(c2_2(a221,X1)|c4_2(a221,X1)|ssSkC14|~c7_2(a221,X1)|~ndr1_1(a221)), c_0_199, ['final']).
cnf(c_0_2398,negated_conjecture,(c2_2(a229,X1)|c1_2(a229,X1)|c8_2(a229,X1)|~c2_0|~ndr1_1(a229)), c_0_215, ['final']).
cnf(c_0_2399,negated_conjecture,(c2_2(a253,X1)|c5_2(a253,X1)|c3_2(a253,X1)|ssSkC21|~ndr1_1(a253)), c_0_324, ['final']).
cnf(c_0_2400,negated_conjecture,(c4_0|~c6_2(a277,X1)|~c7_2(a277,X1)|~ndr1_1(a277)|~ssSkC25), c_0_325, ['final']).
cnf(c_0_2401,negated_conjecture,(c2_2(a179,X1)|c7_0|~c3_2(a179,X1)|~ndr1_1(a179)|~ssSkC5), c_0_326, ['final']).
cnf(c_0_2402,negated_conjecture,(ssSkC3|~c1_2(a173,X1)|~c8_2(a173,X1)|~ndr1_1(a173)), c_0_327, ['final']).
cnf(c_0_2403,negated_conjecture,(c2_2(a222,X1)|c5_2(a222,X1)|~c2_0|~ndr1_1(a222)|~ssSkC14), c_0_238, ['final']).
cnf(c_0_2404,negated_conjecture,(c8_2(a235,X1)|~c4_2(a235,X1)|~ndr1_1(a235)|~ssSkC17), c_0_328, ['final']).
cnf(c_0_2405,negated_conjecture,(c6_2(a267,X1)|c3_2(a267,X1)|c7_0|~ndr1_1(a267)|~ssSkC23), c_0_329, ['final']).
cnf(c_0_2406,negated_conjecture,(c6_2(a223,X1)|ssSkC15|~c4_2(a223,X1)|~ndr1_1(a223)), c_0_57, ['final']).
cnf(c_0_2407,negated_conjecture,(c2_2(a223,X1)|c7_2(a223,X1)|ssSkC15|~ndr1_1(a223)), c_0_132, ['final']).
cnf(c_0_2408,negated_conjecture,(c7_0|~c1_2(a148,a150)|~c8_0), c_0_291, ['final']).
cnf(c_0_2409,negated_conjecture,(c7_0|~c1_2(a148,a149)|~c8_0), c_0_292, ['final']).
cnf(c_0_2410,negated_conjecture,(c7_0|~c8_2(a148,a150)|~c8_0), c_0_296, ['final']).
cnf(c_0_2411,negated_conjecture,(c8_0|c7_0|~c1_2(a215,a216)), c_0_330, ['final']).
cnf(c_0_2412,negated_conjecture,(c8_0|c7_0|~c4_2(a215,a216)), c_0_331, ['final']).
cnf(c_0_2413,negated_conjecture,(c4_0|c5_0|~c6_2(a268,a269)), c_0_332, ['final']).
cnf(c_0_2414,negated_conjecture,(ssSkC1|~c5_2(a155,a156)), c_0_333, ['final']).
cnf(c_0_2415,negated_conjecture,(ssSkC2|~c2_2(a164,a165)), c_0_334, ['final']).
cnf(c_0_2416,negated_conjecture,(ssSkC2|~c5_2(a164,a165)), c_0_335, ['final']).
cnf(c_0_2417,negated_conjecture,(ssSkC5|~c6_2(a176,a177)), c_0_336, ['final']).
cnf(c_0_2418,negated_conjecture,(ssSkC5|~c2_2(a176,a177)), c_0_337, ['final']).
cnf(c_0_2419,negated_conjecture,(c6_2(a227,a228)|~ssSkC16|~c1_0), c_0_338, ['final']).
cnf(c_0_2420,negated_conjecture,(c2_2(a281,a282)|~ssSkC26|~c1_0), c_0_339, ['final']).
cnf(c_0_2421,negated_conjecture,(c8_2(a244,a246)|~c6_0|~ssSkC19), c_0_340, ['final']).
cnf(c_0_2422,negated_conjecture,(c8_2(a281,a282)|~ssSkC26|~c1_0), c_0_341, ['final']).
cnf(c_0_2423,negated_conjecture,(c8_2(a281,a283)|~ssSkC26|~c1_0), c_0_342, ['final']).
cnf(c_0_2424,negated_conjecture,(ssSkC5|~c4_2(a176,a178)), c_0_343, ['final']).
cnf(c_0_2425,negated_conjecture,(ssSkC5|~c3_2(a176,a177)), c_0_344, ['final']).
cnf(c_0_2426,negated_conjecture,(ssSkC5|~c7_2(a176,a178)), c_0_345, ['final']).
cnf(c_0_2427,negated_conjecture,(ssSkC6|~c6_2(a187,a188)), c_0_346, ['final']).
cnf(c_0_2428,negated_conjecture,(ssSkC6|~c4_2(a187,a188)), c_0_347, ['final']).
cnf(c_0_2429,negated_conjecture,(ssSkC12|~c1_2(a211,a212)), c_0_348, ['final']).
cnf(c_0_2430,negated_conjecture,(ssSkC12|~c8_2(a211,a212)), c_0_349, ['final']).
cnf(c_0_2431,negated_conjecture,(ssSkC13|~c5_2(a218,a219)), c_0_350, ['final']).
cnf(c_0_2432,negated_conjecture,(c4_2(a244,a245)|~c6_0|~ssSkC19), c_0_351, ['final']).
cnf(c_0_2433,negated_conjecture,(ssSkC18|~c1_2(a237,a238)), c_0_352, ['final']).
cnf(c_0_2434,negated_conjecture,(c7_2(a166,a167)|~c7_0|~ssSkC2), c_0_353, ['final']).
cnf(c_0_2435,negated_conjecture,(c7_2(a227,a228)|~ssSkC16|~c1_0), c_0_354, ['final']).
cnf(c_0_2436,negated_conjecture,(c7_2(a281,a282)|~ssSkC26|~c1_0), c_0_355, ['final']).
cnf(c_0_2437,negated_conjecture,(c6_2(a179,a180)|c7_0|~ssSkC5), c_0_356, ['final']).
cnf(c_0_2438,negated_conjecture,(ssSkC18|~c7_2(a237,a238)), c_0_357, ['final']).
cnf(c_0_2439,negated_conjecture,(ssSkC19|~c6_2(a241,a243)), c_0_358, ['final']).
cnf(c_0_2440,negated_conjecture,(ssSkC19|~c1_2(a241,a242)), c_0_359, ['final']).
cnf(c_0_2441,negated_conjecture,(ssSkC19|~c4_2(a241,a242)), c_0_360, ['final']).
cnf(c_0_2442,negated_conjecture,(ssSkC19|~c7_2(a241,a242)), c_0_361, ['final']).
cnf(c_0_2443,negated_conjecture,(ssSkC23|~c8_2(a265,a266)), c_0_362, ['final']).
cnf(c_0_2444,negated_conjecture,(ssSkC24|~c4_2(a271,a273)), c_0_363, ['final']).
cnf(c_0_2445,negated_conjecture,(ssSkC24|~c3_2(a271,a273)), c_0_364, ['final']).
cnf(c_0_2446,negated_conjecture,(c2_2(a148,a149)|c7_0|~c8_0), c_0_300, ['final']).
cnf(c_0_2447,negated_conjecture,(ssSkC24|~c7_2(a271,a272)), c_0_365, ['final']).
cnf(c_0_2448,negated_conjecture,(c2_2(a179,a180)|c7_0|~ssSkC5), c_0_366, ['final']).
cnf(c_0_2449,negated_conjecture,(c5_2(a183,a184)|c7_0|~c1_0), c_0_367, ['final']).
cnf(c_0_2450,negated_conjecture,(c3_2(a183,a184)|c7_0|~c1_0), c_0_368, ['final']).
cnf(c_0_2451,negated_conjecture,(c3_2(a179,a180)|c7_0|~ssSkC5), c_0_369, ['final']).
cnf(c_0_2452,negated_conjecture,(c7_2(a148,a149)|c7_0|~c8_0), c_0_136, ['final']).
cnf(c_0_2453,negated_conjecture,(c7_2(a183,a184)|c7_0|~c1_0), c_0_370, ['final']).
cnf(c_0_2454,negated_conjecture,(c3_2(a215,a216)|c8_0|c7_0), c_0_371, ['final']).
cnf(c_0_2455,negated_conjecture,(c6_2(a268,a270)|c4_0|c5_0), c_0_98, ['final']).
cnf(c_0_2456,negated_conjecture,(ssSkC25|~c3_2(a275,a276)), c_0_372, ['final']).
cnf(c_0_2457,negated_conjecture,(ssSkC26|~c1_2(a279,a280)), c_0_373, ['final']).
cnf(c_0_2458,negated_conjecture,(ssSkC26|~c4_2(a279,a280)), c_0_374, ['final']).
cnf(c_0_2459,negated_conjecture,(ssSkC27|~c5_2(a288,a289)), c_0_375, ['final']).
cnf(c_0_2460,negated_conjecture,(ssSkC27|~c8_2(a288,a289)), c_0_376, ['final']).
cnf(c_0_2461,negated_conjecture,(ssSkC27|~c4_2(a288,a289)), c_0_377, ['final']).
cnf(c_0_2462,negated_conjecture,(c5_0|~c6_2(a162,a163)), c_0_378, ['final']).
cnf(c_0_2463,negated_conjecture,(c5_0|~c6_2(a197,a198)), c_0_379, ['final']).
cnf(c_0_2464,negated_conjecture,(c2_2(a268,a269)|c4_0|c5_0), c_0_178, ['final']).
cnf(c_0_2465,negated_conjecture,(c5_2(a268,a269)|c4_0|c5_0), c_0_105, ['final']).
cnf(c_0_2466,negated_conjecture,(c1_2(a268,a270)|c4_0|c5_0), c_0_380, ['final']).
cnf(c_0_2467,negated_conjecture,(c5_0|~c5_2(a162,a163)), c_0_381, ['final']).
cnf(c_0_2468,negated_conjecture,(c5_0|~c1_2(a197,a198)), c_0_382, ['final']).
cnf(c_0_2469,negated_conjecture,(c1_0|~c2_2(a294,a295)), c_0_383, ['final']).
cnf(c_0_2470,negated_conjecture,(c1_0|~c8_2(a294,a295)), c_0_384, ['final']).
cnf(c_0_2471,negated_conjecture,(c3_0|~c6_1(a168)|~c4_0), c_0_14, ['final']).
cnf(c_0_2472,negated_conjecture,(c3_0|~c2_1(a168)|~c4_0), c_0_41, ['final']).
cnf(c_0_2473,negated_conjecture,(c2_0|~c7_1(a186)|~c3_0), c_0_299, ['final']).
cnf(c_0_2474,negated_conjecture,(c7_0|~c7_1(a259)|~ssSkC22), c_0_385, ['final']).
cnf(c_0_2475,negated_conjecture,(c7_0|~c2_1(a179)|~ssSkC5), c_0_386, ['final']).
cnf(c_0_2476,negated_conjecture,(c8_2(a268,a270)|c4_0|c5_0), c_0_156, ['final']).
cnf(c_0_2477,negated_conjecture,(c6_0|c7_0|~c7_1(a170)), c_0_387, ['final']).
cnf(c_0_2478,negated_conjecture,(c6_0|c7_0|~c5_1(a170)), c_0_388, ['final']).
cnf(c_0_2479,negated_conjecture,(c8_0|c7_0|~c2_1(a215)), c_0_111, ['final']).
cnf(c_0_2480,negated_conjecture,(c3_0|c5_0|~c5_1(a287)), c_0_113, ['final']).
cnf(c_0_2481,negated_conjecture,(c4_0|c5_0|~c7_1(a268)), c_0_389, ['final']).
cnf(c_0_2482,negated_conjecture,(c5_0|c1_0|~c4_1(a158)), c_0_71, ['final']).
cnf(c_0_2483,negated_conjecture,(c7_1(a166)|~c7_0|~ssSkC2), c_0_390, ['final']).
cnf(c_0_2484,negated_conjecture,(c1_1(a166)|~c7_0|~ssSkC2), c_0_391, ['final']).
cnf(c_0_2485,negated_conjecture,(c5_1(a157)|~ssSkC0|~ssSkC1), c_0_392, ['final']).
cnf(c_0_2486,negated_conjecture,(c2_2(a155,a156)|ssSkC1), c_0_192, ['final']).
cnf(c_0_2487,negated_conjecture,(c7_2(a155,a156)|ssSkC1), c_0_173, ['final']).
cnf(c_0_2488,negated_conjecture,(c3_2(a164,a165)|ssSkC2), c_0_393, ['final']).
cnf(c_0_2489,negated_conjecture,(c8_2(a176,a178)|ssSkC5), c_0_394, ['final']).
cnf(c_0_2490,negated_conjecture,(c4_2(a218,a219)|ssSkC13), c_0_125, ['final']).
cnf(c_0_2491,negated_conjecture,(c2_2(a237,a238)|ssSkC18), c_0_191, ['final']).
cnf(c_0_2492,negated_conjecture,(c8_1(a169)|~c7_0|~c4_0), c_0_395, ['final']).
cnf(c_0_2493,negated_conjecture,(c5_2(a241,a243)|ssSkC19), c_0_79, ['final']).
cnf(c_0_2494,negated_conjecture,(c2_1(a169)|~c7_0|~c4_0), c_0_396, ['final']).
cnf(c_0_2495,negated_conjecture,(ndr1_1(a166)|~c7_0|~ssSkC2), c_0_397, ['final']).
cnf(c_0_2496,negated_conjecture,(ndr1_1(a227)|~ssSkC16|~c1_0), c_0_398, ['final']).
cnf(c_0_2497,negated_conjecture,(ndr1_1(a244)|~c6_0|~ssSkC19), c_0_399, ['final']).
cnf(c_0_2498,negated_conjecture,(ndr1_1(a281)|~ssSkC26|~c1_0), c_0_400, ['final']).
cnf(c_0_2499,negated_conjecture,(c1_1(a186)|c2_0|~c3_0), c_0_301, ['final']).
cnf(c_0_2500,negated_conjecture,(c6_1(a259)|c7_0|~ssSkC22), c_0_401, ['final']).
cnf(c_0_2501,negated_conjecture,(c4_1(a183)|c7_0|~c1_0), c_0_402, ['final']).
cnf(c_0_2502,negated_conjecture,(c4_2(a241,a243)|ssSkC19), c_0_83, ['final']).
cnf(c_0_2503,negated_conjecture,(c2_2(a265,a266)|ssSkC23), c_0_190, ['final']).
cnf(c_0_2504,negated_conjecture,(c4_2(a265,a266)|ssSkC23), c_0_56, ['final']).
cnf(c_0_2505,negated_conjecture,(c2_2(a271,a272)|ssSkC24), c_0_106, ['final']).
cnf(c_0_2506,negated_conjecture,(c1_2(a271,a272)|ssSkC24), c_0_94, ['final']).
cnf(c_0_2507,negated_conjecture,(c1_2(a271,a273)|ssSkC24), c_0_93, ['final']).
cnf(c_0_2508,negated_conjecture,(c2_2(a275,a276)|ssSkC25), c_0_60, ['final']).
cnf(c_0_2509,negated_conjecture,(c7_2(a275,a276)|ssSkC25), c_0_66, ['final']).
cnf(c_0_2510,negated_conjecture,(c1_1(a183)|c7_0|~c1_0), c_0_403, ['final']).
cnf(c_0_2511,negated_conjecture,(c8_2(a279,a280)|ssSkC26), c_0_404, ['final']).
cnf(c_0_2512,negated_conjecture,(c7_0|ndr1_1(a148)|~c8_0), c_0_143, ['final']).
cnf(c_0_2513,negated_conjecture,(c7_0|ndr1_1(a183)|~c1_0), c_0_405, ['final']).
cnf(c_0_2514,negated_conjecture,(c7_0|ndr1_1(a179)|~ssSkC5), c_0_406, ['final']).
cnf(c_0_2515,negated_conjecture,(ssSkP1(X1)|~c5_2(X1,a224)), c_0_407, ['final']).
cnf(c_0_2516,negated_conjecture,(ssSkP2(X1)|~c5_2(X1,a239)), c_0_408, ['final']).
cnf(c_0_2517,negated_conjecture,(ssSkP4(X1)|~c6_2(X1,a254)), c_0_409, ['final']).
cnf(c_0_2518,negated_conjecture,(ssSkP4(X1)|~c5_2(X1,a254)), c_0_410, ['final']).
cnf(c_0_2519,negated_conjecture,(ssSkP5(X1)|~c2_2(X1,a260)), c_0_268, ['final']).
cnf(c_0_2520,negated_conjecture,(c8_2(a162,a163)|c5_0), c_0_411, ['final']).
cnf(c_0_2521,negated_conjecture,(c4_2(a197,a198)|c5_0), c_0_82, ['final']).
cnf(c_0_2522,negated_conjecture,(c6_2(a233,a234)|ssSkC17), c_0_39, ['final']).
cnf(c_0_2523,negated_conjecture,(c1_2(a233,a234)|ssSkC17), c_0_81, ['final']).
cnf(c_0_2524,negated_conjecture,(c3_2(a233,a234)|ssSkC17), c_0_110, ['final']).
cnf(c_0_2525,negated_conjecture,(c3_2(a204,a205)|c1_0), c_0_412, ['final']).
cnf(c_0_2526,negated_conjecture,(c7_2(a204,a205)|c1_0), c_0_15, ['final']).
cnf(c_0_2527,negated_conjecture,(c7_2(a294,a295)|c1_0), c_0_13, ['final']).
cnf(c_0_2528,negated_conjecture,(c7_1(a215)|c8_0|c7_0), c_0_112, ['final']).
cnf(c_0_2529,negated_conjecture,(c8_0|c7_0|ndr1_1(a215)), c_0_413, ['final']).
cnf(c_0_2530,negated_conjecture,(ssSkP5(X1)|~c1_2(X1,a260)), c_0_414, ['final']).
cnf(c_0_2531,negated_conjecture,(ssSkP5(X1)|~c7_2(X1,a260)), c_0_415, ['final']).
cnf(c_0_2532,negated_conjecture,(ssSkP6(X1)|~c5_2(X1,a263)), c_0_416, ['final']).
cnf(c_0_2533,negated_conjecture,(ssSkP6(X1)|~c7_2(X1,a263)), c_0_417, ['final']).
cnf(c_0_2534,negated_conjecture,(ssSkP8(X1)|~c8_2(X1,a285)), c_0_289, ['final']).
cnf(c_0_2535,negated_conjecture,(ssSkP8(X1)|~c7_2(X1,a285)), c_0_418, ['final']).
cnf(c_0_2536,negated_conjecture,(ssSkP9(X1)|~c1_2(X1,a290)), c_0_419, ['final']).
cnf(c_0_2537,negated_conjecture,(ssSkP9(X1)|~c8_2(X1,a290)), c_0_290, ['final']).
cnf(c_0_2538,negated_conjecture,(c3_1(a287)|c3_0|c5_0), c_0_92, ['final']).
cnf(c_0_2539,negated_conjecture,(c3_0|c8_1(a287)|c5_0), c_0_271, ['final']).
cnf(c_0_2540,negated_conjecture,(ssSkC1|~c7_1(a155)), c_0_420, ['final']).
cnf(c_0_2541,negated_conjecture,(ssSkC3|~c6_1(a173)), c_0_135, ['final']).
cnf(c_0_2542,negated_conjecture,(ssSkC4|~c3_1(a174)), c_0_31, ['final']).
cnf(c_0_2543,negated_conjecture,(ssSkC6|~c3_1(a187)), c_0_237, ['final']).
cnf(c_0_2544,negated_conjecture,(ssSkC10|~c1_1(a200)), c_0_97, ['final']).
cnf(c_0_2545,negated_conjecture,(ssSkC12|~c3_1(a211)), c_0_235, ['final']).
cnf(c_0_2546,negated_conjecture,(ndr1_1(a268)|c4_0|c5_0), c_0_127, ['final']).
cnf(c_0_2547,negated_conjecture,(ssSkC13|~c1_1(a218)), c_0_116, ['final']).
cnf(c_0_2548,negated_conjecture,(ssSkP9(X1)|~c3_2(X1,a290)), c_0_421, ['final']).
cnf(c_0_2549,negated_conjecture,(ssSkP10(X1)|~c5_2(X1,a292)), c_0_422, ['final']).
cnf(c_0_2550,negated_conjecture,(c6_2(X1,a206)|ssSkP0(X1)), c_0_100, ['final']).
cnf(c_0_2551,negated_conjecture,(c4_2(X1,a206)|ssSkP0(X1)), c_0_101, ['final']).
cnf(c_0_2552,negated_conjecture,(c3_2(X1,a206)|ssSkP0(X1)), c_0_104, ['final']).
cnf(c_0_2553,negated_conjecture,(c3_2(X1,a224)|ssSkP1(X1)), c_0_423, ['final']).
cnf(c_0_2554,negated_conjecture,(c8_2(X1,a239)|ssSkP2(X1)), c_0_424, ['final']).
cnf(c_0_2555,negated_conjecture,(c7_2(X1,a239)|ssSkP2(X1)), c_0_425, ['final']).
cnf(c_0_2556,negated_conjecture,(ssSkC13|~c5_1(a218)), c_0_426, ['final']).
cnf(c_0_2557,negated_conjecture,(ssSkC14|~c1_1(a221)), c_0_118, ['final']).
cnf(c_0_2558,negated_conjecture,(ssSkC14|~c2_1(a221)), c_0_133, ['final']).
cnf(c_0_2559,negated_conjecture,(ssSkC15|~c4_1(a223)), c_0_46, ['final']).
cnf(c_0_2560,negated_conjecture,(ssSkC16|~c5_1(a226)), c_0_427, ['final']).
cnf(c_0_2561,negated_conjecture,(ssSkC18|~c7_1(a237)), c_0_428, ['final']).
cnf(c_0_2562,negated_conjecture,(ssSkC21|~c7_1(a253)), c_0_429, ['final']).
cnf(c_0_2563,negated_conjecture,(ssSkC21|~c8_1(a253)), c_0_70, ['final']).
cnf(c_0_2564,negated_conjecture,(c1_1(a158)|c5_0|c1_0), c_0_114, ['final']).
cnf(c_0_2565,negated_conjecture,(ssSkC23|~c6_1(a265)), c_0_134, ['final']).
cnf(c_0_2566,negated_conjecture,(c5_2(X1,a250)|ssSkP3(X1)), c_0_430, ['final']).
cnf(c_0_2567,negated_conjecture,(c8_2(X1,a250)|ssSkP3(X1)), c_0_431, ['final']).
cnf(c_0_2568,negated_conjecture,(c7_2(X1,a250)|ssSkP3(X1)), c_0_432, ['final']).
cnf(c_0_2569,negated_conjecture,(c1_2(X1,a254)|ssSkP4(X1)), c_0_433, ['final']).
cnf(c_0_2570,negated_conjecture,(c4_2(X1,a263)|ssSkP6(X1)), c_0_434, ['final']).
cnf(c_0_2571,negated_conjecture,(c5_2(X1,a284)|ssSkP7(X1)), c_0_435, ['final']).
cnf(c_0_2572,negated_conjecture,(c1_2(X1,a284)|ssSkP7(X1)), c_0_436, ['final']).
cnf(c_0_2573,negated_conjecture,(c3_2(X1,a284)|ssSkP7(X1)), c_0_437, ['final']).
cnf(c_0_2574,negated_conjecture,(ssSkC25|~c1_1(a275)), c_0_115, ['final']).
cnf(c_0_2575,negated_conjecture,(ssSkC27|~c8_1(a288)), c_0_86, ['final']).
cnf(c_0_2576,negated_conjecture,(c5_0|~c6_1(a197)), c_0_145, ['final']).
cnf(c_0_2577,negated_conjecture,(c5_0|~c8_1(a162)), c_0_85, ['final']).
cnf(c_0_2578,negated_conjecture,(c5_0|~c8_1(a197)), c_0_84, ['final']).
cnf(c_0_2579,negated_conjecture,(c1_0|~c3_1(a294)), c_0_241, ['final']).
cnf(c_0_2580,negated_conjecture,(c1_0|~c1_1(a182)), c_0_117, ['final']).
cnf(c_0_2581,negated_conjecture,(c7_1(a217)|~c4_0), c_0_26, ['final']).
cnf(c_0_2582,negated_conjecture,(c1_1(a229)|~c2_0), c_0_304, ['final']).
cnf(c_0_2583,negated_conjecture,(c8_1(a217)|~c4_0), c_0_27, ['final']).
cnf(c_0_2584,negated_conjecture,(c3_2(X1,a285)|ssSkP8(X1)), c_0_438, ['final']).
cnf(c_0_2585,negated_conjecture,(c8_2(X1,a292)|ssSkP10(X1)), c_0_40, ['final']).
cnf(c_0_2586,negated_conjecture,(c7_2(X1,a292)|ssSkP10(X1)), c_0_11, ['final']).
cnf(c_0_2587,negated_conjecture,(c6_1(a154)|ssSkC0), c_0_439, ['final']).
cnf(c_0_2588,negated_conjecture,(ssSkP0(X1)|ndr1_1(X1)), c_0_440, ['final']).
cnf(c_0_2589,negated_conjecture,(ssSkP1(X1)|ndr1_1(X1)), c_0_441, ['final']).
cnf(c_0_2590,negated_conjecture,(ssSkP2(X1)|ndr1_1(X1)), c_0_442, ['final']).
cnf(c_0_2591,negated_conjecture,(ssSkP3(X1)|ndr1_1(X1)), c_0_443, ['final']).
cnf(c_0_2592,negated_conjecture,(c4_1(a154)|ssSkC0), c_0_250, ['final']).
cnf(c_0_2593,negated_conjecture,(c2_1(a154)|ssSkC0), c_0_444, ['final']).
cnf(c_0_2594,negated_conjecture,(ndr1_1(a155)|ssSkC1), c_0_150, ['final']).
cnf(c_0_2595,negated_conjecture,(c6_1(a164)|ssSkC2), c_0_445, ['final']).
cnf(c_0_2596,negated_conjecture,(ndr1_1(a164)|ssSkC2), c_0_280, ['final']).
cnf(c_0_2597,negated_conjecture,(c5_1(a174)|ssSkC4), c_0_44, ['final']).
cnf(c_0_2598,negated_conjecture,(ndr1_1(a176)|ssSkC5), c_0_247, ['final']).
cnf(c_0_2599,negated_conjecture,(c4_1(a187)|ssSkC6), c_0_248, ['final']).
cnf(c_0_2600,negated_conjecture,(ndr1_1(a187)|ssSkC6), c_0_446, ['final']).
cnf(c_0_2601,negated_conjecture,(c1_1(a199)|ssSkC9), c_0_447, ['final']).
cnf(c_0_2602,negated_conjecture,(ssSkP4(X1)|ndr1_1(X1)), c_0_448, ['final']).
cnf(c_0_2603,negated_conjecture,(ssSkP5(X1)|ndr1_1(X1)), c_0_449, ['final']).
cnf(c_0_2604,negated_conjecture,(ssSkP6(X1)|ndr1_1(X1)), c_0_450, ['final']).
cnf(c_0_2605,negated_conjecture,(ssSkP7(X1)|ndr1_1(X1)), c_0_22, ['final']).
cnf(c_0_2606,negated_conjecture,(ssSkP8(X1)|ndr1_1(X1)), c_0_38, ['final']).
cnf(c_0_2607,negated_conjecture,(ssSkP9(X1)|ndr1_1(X1)), c_0_451, ['final']).
cnf(c_0_2608,negated_conjecture,(ssSkP10(X1)|ndr1_1(X1)), c_0_37, ['final']).
cnf(c_0_2609,negated_conjecture,(c4_0|~c6_0|~c1_0), c_0_452, ['final']).
cnf(c_0_2610,negated_conjecture,(c6_1(a200)|ssSkC10), c_0_155, ['final']).
cnf(c_0_2611,negated_conjecture,(c7_1(a202)|ssSkC11), c_0_50, ['final']).
cnf(c_0_2612,negated_conjecture,(c4_1(a211)|ssSkC12), c_0_245, ['final']).
cnf(c_0_2613,negated_conjecture,(ndr1_1(a211)|ssSkC12), c_0_453, ['final']).
cnf(c_0_2614,negated_conjecture,(ndr1_1(a218)|ssSkC13), c_0_154, ['final']).
cnf(c_0_2615,negated_conjecture,(c1_1(a237)|ssSkC18), c_0_454, ['final']).
cnf(c_0_2616,negated_conjecture,(ndr1_1(a237)|ssSkC18), c_0_455, ['final']).
cnf(c_0_2617,negated_conjecture,(ndr1_1(a241)|ssSkC19), c_0_122, ['final']).
cnf(c_0_2618,negated_conjecture,(ndr1_1(a265)|ssSkC23), c_0_65, ['final']).
cnf(c_0_2619,negated_conjecture,(c8_0|c7_0|~c3_0), c_0_73, ['final']).
cnf(c_0_2620,negated_conjecture,(c5_1(a271)|ssSkC24), c_0_144, ['final']).
cnf(c_0_2621,negated_conjecture,(ndr1_1(a271)|ssSkC24), c_0_153, ['final']).
cnf(c_0_2622,negated_conjecture,(ndr1_1(a275)|ssSkC25), c_0_54, ['final']).
cnf(c_0_2623,negated_conjecture,(c2_1(a279)|ssSkC26), c_0_456, ['final']).
cnf(c_0_2624,negated_conjecture,(ndr1_1(a279)|ssSkC26), c_0_270, ['final']).
cnf(c_0_2625,negated_conjecture,(ndr1_1(a288)|ssSkC27), c_0_243, ['final']).
cnf(c_0_2626,negated_conjecture,(ndr1_1(a162)|c5_0), c_0_273, ['final']).
cnf(c_0_2627,negated_conjecture,(ndr1_1(a197)|c5_0), c_0_124, ['final']).
cnf(c_0_2628,negated_conjecture,(ndr1_1(a233)|ssSkC17), c_0_63, ['final']).
cnf(c_0_2629,negated_conjecture,(c3_0|~c7_0|~c5_0), c_0_457, ['final']).
cnf(c_0_2630,negated_conjecture,(c5_0|~c4_0|~c1_0), c_0_458, ['final']).
cnf(c_0_2631,negated_conjecture,(c8_0|c2_0|~c4_0), c_0_3, ['final']).
cnf(c_0_2632,negated_conjecture,(c8_0|c7_0|~c5_0), c_0_459, ['final']).
cnf(c_0_2633,negated_conjecture,(c3_0|c4_0|~c5_0), c_0_460, ['final']).
cnf(c_0_2634,negated_conjecture,(c8_1(a204)|c1_0), c_0_274, ['final']).
cnf(c_0_2635,negated_conjecture,(c2_1(a294)|c1_0), c_0_461, ['final']).
cnf(c_0_2636,negated_conjecture,(ndr1_1(a204)|c1_0), c_0_30, ['final']).
cnf(c_0_2637,negated_conjecture,(ndr1_1(a294)|c1_0), c_0_24, ['final']).
cnf(c_0_2638,negated_conjecture,(c3_0|c6_0|c8_0), c_0_19, ['final']).
cnf(c_0_2639,negated_conjecture,(c2_0|c5_0|c1_0), c_0_20, ['final']).
cnf(c_0_2640,negated_conjecture,(c4_0|c1_0), c_0_4, ['final']).
cnf(c_0_2641,negated_conjecture,(~c8_2(a166,a167)|~c7_0|~ssSkC2), c_0_462, ['final']).
cnf(c_0_2642,negated_conjecture,(~c3_2(a166,a167)|~c7_0|~ssSkC2), c_0_463, ['final']).
cnf(c_0_2643,negated_conjecture,(~c2_2(a244,a245)|~c6_0|~ssSkC19), c_0_464, ['final']).
cnf(c_0_2644,negated_conjecture,(~c3_2(a244,a246)|~c6_0|~ssSkC19), c_0_465, ['final']).
cnf(c_0_2645,negated_conjecture,(~c7_2(a244,a246)|~c6_0|~ssSkC19), c_0_466, ['final']).
cnf(c_0_2646,negated_conjecture,(~c1_2(a281,a283)|~ssSkC26|~c1_0), c_0_467, ['final']).
cnf(c_0_2647,negated_conjecture,(~ssSkC20|~c7_1(a249)|~c2_0), c_0_468, ['final']).
cnf(c_0_2648,negated_conjecture,(~ssSkC20|~c1_1(a249)|~c2_0), c_0_469, ['final']).
cnf(c_0_2649,negated_conjecture,(~c4_1(a157)|~ssSkC0|~ssSkC1), c_0_470, ['final']).
cnf(c_0_2650,negated_conjecture,(~c6_1(a222)|~c2_0|~ssSkC14), c_0_297, ['final']).
cnf(c_0_2651,negated_conjecture,(~c1_1(a222)|~c2_0|~ssSkC14), c_0_298, ['final']).
cnf(c_0_2652,negated_conjecture,(~c6_1(a169)|~c7_0|~c4_0), c_0_471, ['final']).
cnf(c_0_2653,negated_conjecture,(~c4_1(a153)|~c4_0|~c1_0), c_0_472, ['final']).
cnf(c_0_2654,negated_conjecture,(~c1_1(a153)|~c4_0|~c1_0), c_0_473, ['final']).
cnf(c_0_2655,negated_conjecture,(~c1_1(a217)|~c4_0), c_0_21, ['final']).
cnf(c_0_2656,negated_conjecture,(~c3_1(a235)|~ssSkC17), c_0_474, ['final']).
cnf(c_0_2657,negated_conjecture,(~c1_1(a235)|~ssSkC17), c_0_475, ['final']).
cnf(c_0_2658,negated_conjecture,(~c8_0|~c2_0|~c1_0), c_0_306, ['final']).
# SZS output end Saturation.
EOF
