% Problem    : SYN463-1 : TPTP v6.1.0. Released v2.1.0.
% Command    : eprover --auto-schedule --tstp-format -s --proof-object --memory-limit=2048 --cpu-limit=%d %s
% Computer   : n151.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% Memory     : 32286.75MB
% OS         : Linux 2.6.32-431.20.3.el6.x86_64
% CPULimit   : 300
% DateTime   : Sun Jul 20 03:23:26 CDT 2014
% CPUTime    : 
# No SInE strategy applied
# Trying AutoSched0 for 151 seconds
# AutoSched0-Mode selected heuristic G_E___107_C45_F1_PI_AE_Q7_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Presaturation interreduction done

# No proof found!
# SZS status Satisfiable
# SZS output start Saturation.
cnf(c_0_0,negated_conjecture,(ndr1_0|~hskp28), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause2)).
cnf(c_0_1,negated_conjecture,(hskp20|hskp16|hskp28), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause33)).
cnf(c_0_2,negated_conjecture,(ndr1_0|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause10)).
cnf(c_0_3,negated_conjecture,(ndr1_0|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause14)).
cnf(c_0_4,negated_conjecture,(c1_1(X1)|hskp2|hskp4|~c0_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause147)).
cnf(c_0_5,negated_conjecture,(c0_1(X1)|c1_1(X1)|hskp11|hskp12|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause133)).
cnf(c_0_6,negated_conjecture,(c1_1(X1)|hskp19|hskp7|~c2_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause148)).
cnf(c_0_7,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp12|hskp7|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause142)).
cnf(c_0_8,negated_conjecture,(c3_1(X1)|hskp10|~c0_1(X1)|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause144)).
cnf(c_0_9,negated_conjecture,(c3_1(a1258)|~hskp5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause73)).
cnf(c_0_10,negated_conjecture,(~c1_1(a1258)|~hskp5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause117)).
cnf(c_0_11,negated_conjecture,(hskp19|hskp5|~c1_1(X1)|~c2_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause152)).
cnf(c_0_12,negated_conjecture,(c2_1(a1255)|~hskp2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause78)).
cnf(c_0_13,negated_conjecture,(~c1_1(a1255)|~hskp2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause122)).
cnf(c_0_14,negated_conjecture,(~c3_1(a1265)|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause109)).
cnf(c_0_15,negated_conjecture,(c0_1(a1265)|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause63)).
cnf(c_0_16,negated_conjecture,(c1_1(a1265)|~hskp11), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause64)).
cnf(c_0_17,negated_conjecture,(c2_1(a1280)|~hskp29), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause37)).
cnf(c_0_18,negated_conjecture,(c3_1(a1280)|~hskp29), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause38)).
cnf(c_0_19,negated_conjecture,(hskp29|hskp17|~c0_1(X1)|~c1_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause150)).
cnf(c_0_20,negated_conjecture,(c0_1(X1)|c3_1(X1)|hskp27|hskp17|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause135)).
cnf(c_0_21,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp23|hskp24|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause140)).
cnf(c_0_22,negated_conjecture,(c3_1(X1)|hskp9|hskp6|~c0_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause149)).
cnf(c_0_23,negated_conjecture,(c0_1(X1)|c1_1(X1)|c2_1(X1)|hskp2|hskp3|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause128)).
cnf(c_0_24,negated_conjecture,(c0_1(X1)|c3_1(X1)|hskp17|hskp18|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause136)).
cnf(c_0_25,negated_conjecture,(c2_1(a1257)|~hskp4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause74)).
cnf(c_0_26,negated_conjecture,(c3_1(a1257)|~hskp4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause75)).
cnf(c_0_27,negated_conjecture,(~c1_1(a1257)|~hskp4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause119)).
cnf(c_0_28,negated_conjecture,(c2_1(a1283)|~hskp19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause55)).
cnf(c_0_29,negated_conjecture,(~c3_1(a1283)|~hskp19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause94)).
cnf(c_0_30,negated_conjecture,(~c1_1(a1283)|~hskp19), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause93)).
cnf(c_0_31,negated_conjecture,(c3_1(a1253)|~hskp27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause44)).
cnf(c_0_32,negated_conjecture,(c0_1(a1253)|~hskp27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause42)).
cnf(c_0_33,negated_conjecture,(c1_1(a1253)|~hskp27), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause43)).
cnf(c_0_34,negated_conjecture,(~c3_1(a1301)|~hskp24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause86)).
cnf(c_0_35,negated_conjecture,(~c0_1(a1301)|~hskp24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause85)).
cnf(c_0_36,negated_conjecture,(c2_1(a1301)|~hskp24), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause48)).
cnf(c_0_37,negated_conjecture,(~c3_1(a1300)|~hskp23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause88)).
cnf(c_0_38,negated_conjecture,(c1_1(a1300)|~hskp23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause49)).
cnf(c_0_39,negated_conjecture,(~c0_1(a1300)|~hskp23), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause87)).
cnf(c_0_40,negated_conjecture,(~c3_1(a1274)|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause100)).
cnf(c_0_41,negated_conjecture,(c0_1(a1274)|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause58)).
cnf(c_0_42,negated_conjecture,(~c1_1(a1274)|~hskp16), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause99)).
cnf(c_0_43,negated_conjecture,(c0_1(X1)|c1_1(X1)|c3_1(X1)|hskp28|hskp6|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause130)).
cnf(c_0_44,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp10|hskp25|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause141)).
cnf(c_0_45,negated_conjecture,(~c3_1(a1252)|~hskp0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause125)).
cnf(c_0_46,negated_conjecture,(c0_1(a1252)|~hskp0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause80)).
cnf(c_0_47,negated_conjecture,(c2_1(a1250)|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause47)).
cnf(c_0_48,negated_conjecture,(c1_1(a1250)|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause46)).
cnf(c_0_49,negated_conjecture,(c0_1(a1250)|~hskp26), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause45)).
cnf(c_0_50,negated_conjecture,(c3_1(a1256)|~hskp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause76)).
cnf(c_0_51,negated_conjecture,(~c0_1(a1256)|~hskp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause120)).
cnf(c_0_52,negated_conjecture,(~c1_1(a1256)|~hskp3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause121)).
cnf(c_0_53,negated_conjecture,(~c2_1(a1279)|~hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause97)).
cnf(c_0_54,negated_conjecture,(~c0_1(a1279)|~hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause95)).
cnf(c_0_55,negated_conjecture,(~c1_1(a1279)|~hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause96)).
cnf(c_0_56,negated_conjecture,(~c3_1(a1272)|~hskp15), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause103)).
cnf(c_0_57,negated_conjecture,(~c0_1(a1272)|~hskp15), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause101)).
cnf(c_0_58,negated_conjecture,(~c2_1(a1272)|~hskp15), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause102)).
cnf(c_0_59,negated_conjecture,(~c3_1(a1260)|~hskp6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause116)).
cnf(c_0_60,negated_conjecture,(c1_1(a1260)|~hskp6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause71)).
cnf(c_0_61,negated_conjecture,(~c3_1(a1303)|~hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause84)).
cnf(c_0_62,negated_conjecture,(~c0_1(a1303)|~hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause82)).
cnf(c_0_63,negated_conjecture,(~c1_1(a1303)|~hskp25), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause83)).
cnf(c_0_64,negated_conjecture,(c2_1(a1252)|~hskp0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause81)).
cnf(c_0_65,negated_conjecture,(c0_1(X1)|c1_1(X1)|c2_1(X1)|hskp26|hskp0|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause126)).
cnf(c_0_66,negated_conjecture,(c3_1(a1263)|~hskp9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause68)).
cnf(c_0_67,negated_conjecture,(c0_1(a1263)|~hskp9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause67)).
cnf(c_0_68,negated_conjecture,(~c1_1(a1263)|~hskp9), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause111)).
cnf(c_0_69,negated_conjecture,(c0_1(X1)|c2_1(X1)|hskp10|hskp15|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause134)).
cnf(c_0_70,negated_conjecture,(c0_1(X1)|hskp26|hskp7|~c1_1(X1)|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause145)).
cnf(c_0_71,negated_conjecture,(~c2_1(a1261)|~hskp7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause115)).
cnf(c_0_72,negated_conjecture,(~c0_1(a1261)|~hskp7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause114)).
cnf(c_0_73,negated_conjecture,(c3_1(a1261)|~hskp7), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause70)).
cnf(c_0_74,negated_conjecture,(c2_1(a1259)|~hskp28), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause40)).
cnf(c_0_75,negated_conjecture,(c3_1(a1259)|~hskp28), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause41)).
cnf(c_0_76,negated_conjecture,(c1_1(a1259)|~hskp28), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause39)).
cnf(c_0_77,negated_conjecture,(c2_1(a1266)|~hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause62)).
cnf(c_0_78,negated_conjecture,(~c1_1(a1266)|~hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause108)).
cnf(c_0_79,negated_conjecture,(c2_1(a1264)|~hskp10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause66)).
cnf(c_0_80,negated_conjecture,(c1_1(a1264)|~hskp10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause65)).
cnf(c_0_81,negated_conjecture,(~c0_1(a1264)|~hskp10), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause110)).
cnf(c_0_82,negated_conjecture,(c1_1(X1)|c2_1(X1)|hskp28|hskp13|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause139)).
cnf(c_0_83,negated_conjecture,(hskp16|hskp6|hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause34)).
cnf(c_0_84,negated_conjecture,(hskp26|hskp19|hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause31)).
cnf(c_0_85,negated_conjecture,(~c2_1(a1268)|~hskp13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause106)).
cnf(c_0_86,negated_conjecture,(c1_1(a1268)|~hskp13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause61)).
cnf(c_0_87,negated_conjecture,(~c0_1(a1268)|~hskp13), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause105)).
cnf(c_0_88,negated_conjecture,(~c3_1(a1254)|~hskp1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause124)).
cnf(c_0_89,negated_conjecture,(c1_1(a1254)|~hskp1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause79)).
cnf(c_0_90,negated_conjecture,(c0_1(X1)|c1_1(X1)|c2_1(X1)|hskp27|hskp1|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause127)).
cnf(c_0_91,negated_conjecture,(~c0_1(a1266)|~hskp12), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause107)).
cnf(c_0_92,negated_conjecture,(c2_1(a1260)|~hskp6), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause72)).
cnf(c_0_93,negated_conjecture,(c0_1(X1)|c0_1(X2)|c2_1(X2)|hskp10|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause164)).
cnf(c_0_94,negated_conjecture,(~c2_1(a1287)|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause92)).
cnf(c_0_95,negated_conjecture,(c0_1(a1287)|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause54)).
cnf(c_0_96,negated_conjecture,(c2_1(a1270)|~hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause59)).
cnf(c_0_97,negated_conjecture,(c3_1(a1270)|~hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause60)).
cnf(c_0_98,negated_conjecture,(~c0_1(a1270)|~hskp14), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause104)).
cnf(c_0_99,negated_conjecture,(c0_1(X1)|hskp21|hskp19|~c2_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause146)).
cnf(c_0_100,negated_conjecture,(hskp27|hskp24|hskp18), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause32)).
cnf(c_0_101,negated_conjecture,(c3_1(X1)|c1_1(X2)|c3_1(X2)|c0_1(X3)|c3_1(X3)|~c1_1(X1)|~c2_1(X1)|~ndr1_0|~c2_1(X2)|~c1_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause177)).
cnf(c_0_102,negated_conjecture,(c0_1(X1)|c2_1(X1)|c0_1(X2)|c2_1(X2)|c3_1(X2)|hskp3|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause154)).
cnf(c_0_103,negated_conjecture,(~c1_1(a1287)|~hskp20), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause91)).
cnf(c_0_104,negated_conjecture,(~c2_1(a1258)|~hskp5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause118)).
cnf(c_0_105,negated_conjecture,(c3_1(a1288)|~hskp21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause53)).
cnf(c_0_106,negated_conjecture,(c1_1(a1288)|~hskp21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause52)).
cnf(c_0_107,negated_conjecture,(~c0_1(a1288)|~hskp21), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause90)).
cnf(c_0_108,negated_conjecture,(hskp10|hskp21|hskp5), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause35)).
cnf(c_0_109,negated_conjecture,(c0_1(a1280)|~hskp29), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause36)).
cnf(c_0_110,negated_conjecture,(c3_1(X1)|c2_1(X2)|c1_1(X3)|c2_1(X3)|~c1_1(X1)|~c2_1(X1)|~ndr1_0|~c0_1(X2)|~c3_1(X2)|~c0_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause180)).
cnf(c_0_111,negated_conjecture,(c0_1(X1)|c2_1(X1)|c0_1(X2)|c2_1(X2)|hskp14|~c3_1(X1)|~ndr1_0|~c1_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause159)).
cnf(c_0_112,negated_conjecture,(c1_1(X1)|c1_1(X2)|c3_1(X2)|c0_1(X3)|~c0_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)|~c1_1(X3)|~c2_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause179)).
cnf(c_0_113,negated_conjecture,(~c2_1(a1277)|~hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause98)).
cnf(c_0_114,negated_conjecture,(c3_1(a1277)|~hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause57)).
cnf(c_0_115,negated_conjecture,(c1_1(a1277)|~hskp17), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause56)).
cnf(c_0_116,negated_conjecture,(c1_1(X1)|c2_1(X1)|c0_1(X2)|c0_1(X3)|~c3_1(X1)|~ndr1_0|~c2_1(X2)|~c3_1(X2)|~c1_1(X3)|~c2_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause178)).
cnf(c_0_117,negated_conjecture,(c1_1(X1)|c3_1(X1)|c0_1(X2)|hskp19|~c0_1(X1)|~ndr1_0|~c1_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause167)).
cnf(c_0_118,negated_conjecture,(c0_1(X1)|c0_1(X2)|c1_1(X2)|hskp7|~c1_1(X1)|~c2_1(X1)|~ndr1_0|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause161)).
cnf(c_0_119,negated_conjecture,(c1_1(X1)|c0_1(X2)|c1_1(X2)|hskp9|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause162)).
cnf(c_0_120,negated_conjecture,(c3_1(X1)|c0_1(X2)|hskp20|~c1_1(X1)|~c2_1(X1)|~ndr1_0|~c1_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause170)).
cnf(c_0_121,negated_conjecture,(c1_1(X1)|c2_1(X1)|c3_1(X1)|hskp6|hskp15|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause132)).
cnf(c_0_122,negated_conjecture,(c1_1(X1)|c2_1(X1)|c3_1(X1)|hskp22|hskp17|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause131)).
cnf(c_0_123,negated_conjecture,(hskp21|~c1_1(X1)|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)|~c1_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause174)).
cnf(c_0_124,negated_conjecture,(c0_1(a1255)|~hskp2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause77)).
cnf(c_0_125,negated_conjecture,(c2_1(X2)|hskp12|~c0_1(X1)|~c1_1(X1)|~c2_1(X1)|~ndr1_0|~c0_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause173)).
cnf(c_0_126,negated_conjecture,(hskp2|hskp20|~c0_1(X1)|~c1_1(X1)|~c3_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause151)).
cnf(c_0_127,negated_conjecture,(c0_1(a1291)|~hskp22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause50)).
cnf(c_0_128,negated_conjecture,(c1_1(a1291)|~hskp22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause51)).
cnf(c_0_129,negated_conjecture,(~c2_1(a1254)|~hskp1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause123)).
cnf(c_0_130,negated_conjecture,(c1_1(X2)|c3_1(X2)|hskp7|~c1_1(X1)|~c2_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause171)).
cnf(c_0_131,negated_conjecture,(c2_1(X1)|c3_1(X1)|hskp11|hskp0|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause143)).
cnf(c_0_132,negated_conjecture,(c0_1(X1)|c3_1(X1)|hskp29|hskp6|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause137)).
cnf(c_0_133,negated_conjecture,(~c2_1(a1262)|~hskp8), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause112)).
cnf(c_0_134,negated_conjecture,(c0_1(a1262)|~hskp8), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause69)).
cnf(c_0_135,negated_conjecture,(~c3_1(a1262)|~hskp8), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause113)).
cnf(c_0_136,negated_conjecture,(c3_1(X1)|c0_1(X2)|c3_1(X2)|hskp4|~c0_1(X1)|~c2_1(X1)|~ndr1_0|~c1_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause165)).
cnf(c_0_137,negated_conjecture,(~c2_1(a1291)|~hskp22), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause89)).
cnf(c_0_138,negated_conjecture,(c2_1(X1)|c3_1(X1)|c0_1(X2)|c0_1(X3)|c1_1(X3)|~c0_1(X1)|~ndr1_0|~c1_1(X2)|~c3_1(X2)|~c2_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause175)).
cnf(c_0_139,negated_conjecture,(c0_1(X1)|c1_1(X1)|c2_1(X1)|hskp4|hskp5|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause129)).
cnf(c_0_140,negated_conjecture,(c1_1(X1)|c3_1(X1)|c0_1(X2)|c1_1(X2)|hskp8|~c0_1(X1)|~ndr1_0|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause158)).
cnf(c_0_141,negated_conjecture,(c2_1(X1)|c0_1(X2)|c1_1(X2)|hskp10|~c0_1(X1)|~c3_1(X1)|~ndr1_0|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause163)).
cnf(c_0_142,negated_conjecture,(c1_1(X1)|c2_1(X1)|hskp13|hskp12|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause138)).
cnf(c_0_143,negated_conjecture,(c1_1(X1)|c1_1(X2)|c2_1(X2)|c0_1(X3)|c2_1(X3)|~c0_1(X1)|~c2_1(X1)|~ndr1_0|~c0_1(X2)|~c1_1(X3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause176)).
cnf(c_0_144,negated_conjecture,(c1_1(X1)|c0_1(X2)|hskp19|~c0_1(X1)|~c2_1(X1)|~ndr1_0|~c1_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause169)).
cnf(c_0_145,negated_conjecture,(c2_1(X1)|c1_1(X2)|hskp19|~c1_1(X1)|~c3_1(X1)|~ndr1_0|~c0_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause172)).
cnf(c_0_146,negated_conjecture,(c2_1(X1)|c3_1(X1)|c0_1(X2)|c1_1(X2)|c2_1(X2)|hskp26|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause153)).
cnf(c_0_147,negated_conjecture,(c1_1(X1)|c3_1(X1)|c0_1(X2)|hskp4|~c2_1(X1)|~ndr1_0|~c1_1(X2)|~c2_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause166)).
cnf(c_0_148,negated_conjecture,(c0_1(X1)|c3_1(X1)|c0_1(X2)|c2_1(X2)|c3_1(X2)|hskp13|~c2_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause155)).
cnf(c_0_149,negated_conjecture,(c2_1(X1)|c3_1(X1)|c1_1(X2)|hskp4|~c1_1(X1)|~ndr1_0|~c2_1(X2)|~c3_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause168)).
cnf(c_0_150,negated_conjecture,(c2_1(X1)|c3_1(X1)|c0_1(X2)|c3_1(X2)|hskp16|~c1_1(X1)|~ndr1_0|~c1_1(X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause160)).
cnf(c_0_151,negated_conjecture,(c1_1(X1)|c3_1(X1)|c1_1(X2)|c2_1(X2)|c3_1(X2)|hskp1|~c0_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause157)).
cnf(c_0_152,negated_conjecture,(c2_1(X1)|c3_1(X1)|c0_1(X2)|c2_1(X2)|c3_1(X2)|hskp26|~c1_1(X1)|~ndr1_0), file('/export/starexec/sandbox/benchmark/theBenchmark.p', clause156)).
cnf(c_0_153,negated_conjecture,(ndr1_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2]), c_0_3]), ['final']).
cnf(c_0_154,negated_conjecture,(c1_1(X1)|hskp2|hskp4|~c3_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_153])])).
cnf(c_0_155,negated_conjecture,(c1_1(X1)|c0_1(X1)|hskp11|hskp12|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_153])]), ['final']).
cnf(c_0_156,negated_conjecture,(c1_1(X1)|hskp7|hskp19|~c3_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_153])])).
cnf(c_0_157,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp7|hskp12|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_153])])).
cnf(c_0_158,negated_conjecture,(c3_1(X1)|hskp10|~c1_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_153])])).
cnf(c_0_159,negated_conjecture,(hskp2|hskp4|~c0_1(a1258)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_9]), c_0_10])).
cnf(c_0_160,negated_conjecture,(c0_1(a1258)|hskp11|hskp12|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_9]), c_0_10]), ['final']).
cnf(c_0_161,negated_conjecture,(hskp5|hskp19|~c1_1(X1)|~c3_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_153])])).
cnf(c_0_162,negated_conjecture,(hskp7|hskp19|~c3_1(a1255)|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_12]), c_0_13])).
cnf(c_0_163,negated_conjecture,(c3_1(a1255)|hskp7|hskp12|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_12]), c_0_13])).
cnf(c_0_164,negated_conjecture,(hskp10|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_158]), c_0_15]), c_0_16])).
cnf(c_0_165,negated_conjecture,(hskp2|hskp4|hskp11|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_159, c_0_160])).
cnf(c_0_166,negated_conjecture,(hskp5|hskp19|~c1_1(a1280)|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_17]), c_0_18])).
cnf(c_0_167,negated_conjecture,(c1_1(a1280)|hskp7|hskp19|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_17]), c_0_18])).
cnf(c_0_168,negated_conjecture,(hskp17|hskp29|~c1_1(X1)|~c3_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_153])])).
cnf(c_0_169,negated_conjecture,(c3_1(X1)|c0_1(X1)|hskp17|hskp27|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_153])])).
cnf(c_0_170,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp23|hskp24|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_153])]), ['final']).
cnf(c_0_171,negated_conjecture,(c3_1(X1)|hskp6|hskp9|~c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_153])]), ['final']).
cnf(c_0_172,negated_conjecture,(c1_1(X1)|c2_1(X1)|c0_1(X1)|hskp2|hskp3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_153])]), ['final']).
cnf(c_0_173,negated_conjecture,(c3_1(X1)|c0_1(X1)|hskp17|hskp18|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_153])])).
cnf(c_0_174,negated_conjecture,(hskp7|hskp12|hskp19|~hskp2), inference(spm,[status(thm)],[c_0_162, c_0_163])).
cnf(c_0_175,negated_conjecture,(hskp2|hskp4|hskp10|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_176,negated_conjecture,(hskp7|hskp19|~hskp4), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_177,negated_conjecture,(hskp7|hskp12|~hskp19), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_178,negated_conjecture,(hskp5|hskp7|hskp19|~hskp29), inference(spm,[status(thm)],[c_0_166, c_0_167])).
cnf(c_0_179,negated_conjecture,(hskp17|hskp29|~hskp27), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_180,negated_conjecture,(hskp17|hskp27|~c1_1(a1301)|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_169]), c_0_35])).
cnf(c_0_181,negated_conjecture,(c1_1(a1301)|hskp7|hskp12|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_36]), c_0_34])).
cnf(c_0_182,negated_conjecture,(hskp17|hskp27|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_169]), c_0_38]), c_0_39])).
cnf(c_0_183,negated_conjecture,(hskp23|hskp24|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_170]), c_0_41]), c_0_42]), ['final']).
cnf(c_0_184,negated_conjecture,(c1_1(X1)|c3_1(X1)|c0_1(X1)|hskp6|hskp28), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_153])])).
cnf(c_0_185,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp10|hskp25|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_153])])).
cnf(c_0_186,negated_conjecture,(hskp10|~c1_1(a1252)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_158]), c_0_46])).
cnf(c_0_187,negated_conjecture,(hskp5|hskp19|~c3_1(a1250)|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_47]), c_0_48])).
cnf(c_0_188,negated_conjecture,(c3_1(a1250)|hskp6|hskp9|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_47]), c_0_49]), ['final']).
cnf(c_0_189,negated_conjecture,(hskp11|hskp12|~hskp3), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_50]), c_0_51]), c_0_52]), ['final']).
cnf(c_0_190,negated_conjecture,(hskp2|hskp3|~hskp18), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_172]), c_0_54]), c_0_55])).
cnf(c_0_191,negated_conjecture,(hskp17|hskp18|~c1_1(a1272)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_173]), c_0_57])).
cnf(c_0_192,negated_conjecture,(c1_1(a1272)|hskp2|hskp3|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_172]), c_0_57])).
cnf(c_0_193,negated_conjecture,(hskp7|hskp10|hskp12|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175]), c_0_176]), c_0_177])).
cnf(c_0_194,negated_conjecture,(hskp5|hskp7|hskp17|hskp19|~hskp27), inference(spm,[status(thm)],[c_0_178, c_0_179])).
cnf(c_0_195,negated_conjecture,(hskp7|hskp12|hskp17|hskp27|~hskp24), inference(spm,[status(thm)],[c_0_180, c_0_181])).
cnf(c_0_196,negated_conjecture,(hskp17|hskp24|hskp27|~hskp16), inference(spm,[status(thm)],[c_0_182, c_0_183])).
cnf(c_0_197,negated_conjecture,(hskp10|~c0_1(a1260)|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_158]), c_0_60])).
cnf(c_0_198,negated_conjecture,(c0_1(a1260)|hskp17|hskp27|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_169]), c_0_60])).
cnf(c_0_199,negated_conjecture,(hskp6|hskp28|~hskp25), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_184]), c_0_62]), c_0_63])).
cnf(c_0_200,negated_conjecture,(hskp10|hskp25|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_64]), c_0_45]), c_0_186])).
cnf(c_0_201,negated_conjecture,(c1_1(X1)|c2_1(X1)|c0_1(X1)|hskp0|hskp26), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_153])]), ['final']).
cnf(c_0_202,negated_conjecture,(hskp2|hskp4|~hskp9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_66]), c_0_67]), c_0_68])).
cnf(c_0_203,negated_conjecture,(hskp5|hskp6|hskp9|hskp19|~hskp26), inference(spm,[status(thm)],[c_0_187, c_0_188])).
fof(c_0_204, plain, ((~epred40_0<=>![X2]:(((c3_1(X2)|c1_1(X2))|~ndr1_0)|~c2_1(X2)))), introduced(definition)).
cnf(c_0_205,negated_conjecture,(hskp2|hskp11|hskp12|~hskp18), inference(spm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_206,negated_conjecture,(hskp2|hskp3|hskp17|hskp18|~hskp15), inference(spm,[status(thm)],[c_0_191, c_0_192])).
cnf(c_0_207,negated_conjecture,(c2_1(X1)|c0_1(X1)|hskp10|hskp15|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_153])])).
cnf(c_0_208,negated_conjecture,(hskp7|hskp10|hskp12|hskp17|~hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194]), c_0_177])).
cnf(c_0_209,negated_conjecture,(hskp7|hskp12|hskp17|hskp27|~hskp16), inference(spm,[status(thm)],[c_0_195, c_0_196])).
cnf(c_0_210,negated_conjecture,(hskp10|hskp17|hskp27|~hskp6), inference(spm,[status(thm)],[c_0_197, c_0_198])).
cnf(c_0_211,negated_conjecture,(hskp6|hskp10|hskp28|~hskp0), inference(spm,[status(thm)],[c_0_199, c_0_200])).
cnf(c_0_212,negated_conjecture,(hskp0|hskp26|~hskp18), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_201]), c_0_54]), c_0_55]), ['final']).
cnf(c_0_213,negated_conjecture,(hskp2|hskp4|hskp5|hskp6|hskp19|~hskp26), inference(spm,[status(thm)],[c_0_202, c_0_203])).
fof(c_0_214, plain, ((~epred6_0<=>![X2]:((((hskp3|c0_1(X2))|c2_1(X2))|c3_1(X2))|~ndr1_0))), introduced(definition)).
cnf(c_0_215,negated_conjecture,(epred40_0|c1_1(X1)|c3_1(X1)|~c2_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_204])).
cnf(c_0_216,negated_conjecture,(hskp2|hskp10|hskp12|~hskp18), inference(spm,[status(thm)],[c_0_164, c_0_205])).
cnf(c_0_217,negated_conjecture,(c0_1(X1)|hskp7|hskp26|~c1_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_153])])).
fof(c_0_218, plain, ((~epred11_0<=>![X2]:((c0_1(X2)|c3_1(X2))|~c1_1(X2)))), introduced(definition)).
cnf(c_0_219,negated_conjecture,(c3_1(a1255)|hskp10|hskp25|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_12]), c_0_13])).
cnf(c_0_220,negated_conjecture,(hskp10|hskp25|~hskp19), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_221,negated_conjecture,(hskp2|hskp11|hskp12|hskp17|hskp18|~hskp15), inference(spm,[status(thm)],[c_0_189, c_0_206])).
cnf(c_0_222,negated_conjecture,(hskp10|hskp15|~c1_1(a1261)|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_207]), c_0_72])).
cnf(c_0_223,negated_conjecture,(c1_1(a1261)|hskp11|hskp12|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_73]), c_0_72])).
cnf(c_0_224,negated_conjecture,(hskp7|hskp10|hskp12|hskp17|~hskp16), inference(spm,[status(thm)],[c_0_208, c_0_209])).
cnf(c_0_225,negated_conjecture,(hskp7|hskp10|hskp12|hskp17|~hskp6), inference(spm,[status(thm)],[c_0_208, c_0_210])).
cnf(c_0_226,negated_conjecture,(hskp5|hskp19|~hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_74]), c_0_75]), c_0_76])).
cnf(c_0_227,negated_conjecture,(hskp6|hskp10|hskp26|hskp28|~hskp18), inference(spm,[status(thm)],[c_0_211, c_0_212])).
cnf(c_0_228,negated_conjecture,(hskp5|hskp6|hskp7|hskp12|hskp19|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_213]), c_0_176])).
cnf(c_0_229,negated_conjecture,(hskp7|hskp19|~c3_1(a1266)|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_77]), c_0_78])).
cnf(c_0_230,negated_conjecture,(c3_1(a1266)|hskp10|hskp25|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_77]), c_0_78])).
cnf(c_0_231,negated_conjecture,(epred6_0|c3_1(X1)|c2_1(X1)|c0_1(X1)|hskp3|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_214])).
fof(c_0_232, plain, ((~epred18_0<=>![X2]:((((hskp10|c0_1(X2))|c2_1(X2))|~ndr1_0)|~c3_1(X2)))), introduced(definition)).
cnf(c_0_233,negated_conjecture,(epred40_0|c1_1(X1)|c3_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_215, c_0_153])])).
cnf(c_0_234,negated_conjecture,(hskp7|hskp10|hskp12|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_216]), c_0_177])).
cnf(c_0_235,negated_conjecture,(hskp7|hskp26|~hskp10), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_79]), c_0_80]), c_0_81])).
cnf(c_0_236,negated_conjecture,(epred11_0|c3_1(X1)|c0_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_218])).
cnf(c_0_237,negated_conjecture,(c1_1(X1)|c2_1(X1)|hskp13|hskp28|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_153])]), ['final']).
cnf(c_0_238,negated_conjecture,(hskp7|hskp10|hskp19|hskp25|~hskp2), inference(spm,[status(thm)],[c_0_162, c_0_219])).
cnf(c_0_239,negated_conjecture,(hskp6|hskp10|hskp28|~hskp19), inference(spm,[status(thm)],[c_0_199, c_0_220])).
cnf(c_0_240,negated_conjecture,(hskp2|hskp10|hskp12|hskp17|hskp18|~hskp15), inference(spm,[status(thm)],[c_0_164, c_0_221])).
cnf(c_0_241,negated_conjecture,(hskp10|hskp12|hskp15|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_223]), c_0_164])).
cnf(c_0_242,negated_conjecture,(hskp7|hskp10|hskp12|hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_83]), c_0_225])).
cnf(c_0_243,negated_conjecture,(hskp5|hskp6|hskp10|hskp19|hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227]), c_0_84])).
cnf(c_0_244,negated_conjecture,(hskp6|hskp7|hskp10|hskp12|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_228]), c_0_177])).
cnf(c_0_245,negated_conjecture,(hskp7|hskp10|hskp19|hskp25|~hskp12), inference(spm,[status(thm)],[c_0_229, c_0_230])).
cnf(c_0_246,negated_conjecture,(epred6_0|c3_1(X1)|c2_1(X1)|c0_1(X1)|hskp3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_231, c_0_153])])).
cnf(c_0_247,negated_conjecture,(epred18_0|c2_1(X1)|c0_1(X1)|hskp10|~c3_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_232])).
cnf(c_0_248,negated_conjecture,(epred40_0|c3_1(a1266)|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_77]), c_0_78])).
cnf(c_0_249,negated_conjecture,(epred40_0|~hskp19), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_250,negated_conjecture,(hskp5|hskp10|hskp19|~hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_158]), c_0_49]), c_0_48])).
cnf(c_0_251,negated_conjecture,(hskp7|hskp12|hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_84]), c_0_177]), c_0_235])).
fof(c_0_252, plain, ((~epred41_0<=>![X1]:((c3_1(X1)|~c2_1(X1))|~c1_1(X1)))), introduced(definition)).
fof(c_0_253, plain, ((~epred44_0<=>![X2]:((c3_1(X2)|c1_1(X2))|~c0_1(X2)))), introduced(definition)).
fof(c_0_254, plain, ((~epred33_0<=>![X1]:((~c2_1(X1)|~c3_1(X1))|~c1_1(X1)))), introduced(definition)).
cnf(c_0_255,negated_conjecture,(epred11_0|~c1_1(a1272)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_236]), c_0_57]), ['final']).
cnf(c_0_256,negated_conjecture,(c1_1(a1272)|hskp6|hskp28|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_184]), c_0_57])).
cnf(c_0_257,negated_conjecture,(c1_1(a1261)|hskp13|hskp28|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_237]), c_0_73])).
cnf(c_0_258,negated_conjecture,(hskp10|hskp15|~hskp13), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_207]), c_0_86]), c_0_87])).
cnf(c_0_259,negated_conjecture,(hskp6|hskp7|hskp10|hskp28|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_238]), c_0_239])).
cnf(c_0_260,negated_conjecture,(hskp2|hskp10|hskp12|hskp17|hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241]), c_0_242])).
cnf(c_0_261,negated_conjecture,(hskp6|hskp7|hskp10|hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_243]), c_0_244]), c_0_177])).
cnf(c_0_262,negated_conjecture,(hskp6|hskp7|hskp10|hskp28|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_245]), c_0_239])).
cnf(c_0_263,negated_conjecture,(epred11_0|c0_1(a1260)|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_236]), c_0_60])).
fof(c_0_264, plain, ((~epred19_0<=>![X1]:((c0_1(X1)|~c2_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_265,negated_conjecture,(epred6_0|hskp3|~hskp15), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_246]), c_0_57]), c_0_56])).
fof(c_0_266, plain, ((~epred24_0<=>![X1]:((c0_1(X1)|~c2_1(X1))|~c1_1(X1)))), introduced(definition)).
cnf(c_0_267,negated_conjecture,(epred18_0|c2_1(X1)|c0_1(X1)|hskp10|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_247, c_0_153])])).
cnf(c_0_268,negated_conjecture,(epred40_0|hskp7|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_248]), c_0_249])).
cnf(c_0_269,negated_conjecture,(hskp7|hskp10|hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_250]), c_0_251]), c_0_177])).
cnf(c_0_270,negated_conjecture,(epred41_0|c3_1(X1)|~c1_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_252])).
cnf(c_0_271,negated_conjecture,(epred44_0|c1_1(X1)|c3_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_253])).
fof(c_0_272, plain, ((~epred23_0<=>![X2]:((((hskp7|c0_1(X2))|c1_1(X2))|~ndr1_0)|~c2_1(X2)))), introduced(definition)).
fof(c_0_273, plain, ((~epred35_0<=>![X1]:((~c0_1(X1)|~c2_1(X1))|~c1_1(X1)))), introduced(definition)).
cnf(c_0_274,negated_conjecture,(c1_1(a1252)|hskp23|hskp24|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_170]), c_0_46]), ['final']).
cnf(c_0_275,negated_conjecture,(epred33_0|~c1_1(X1)|~c3_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_254])).
cnf(c_0_276,negated_conjecture,(epred11_0|hskp6|hskp28|~hskp15), inference(spm,[status(thm)],[c_0_255, c_0_256])).
cnf(c_0_277,negated_conjecture,(hskp10|hskp15|hskp28|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_257]), c_0_258])).
cnf(c_0_278,negated_conjecture,(hskp6|hskp7|hskp10|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_260]), c_0_261]), c_0_261]), c_0_262])).
cnf(c_0_279,negated_conjecture,(epred11_0|hskp10|~hskp6), inference(spm,[status(thm)],[c_0_197, c_0_263])).
cnf(c_0_280,negated_conjecture,(epred19_0|c0_1(X1)|~c3_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_264])).
cnf(c_0_281,negated_conjecture,(epred6_0|hskp11|hskp12|~hskp15), inference(spm,[status(thm)],[c_0_189, c_0_265])).
cnf(c_0_282,negated_conjecture,(epred24_0|c0_1(X1)|~c1_1(X1)|~c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_266])).
cnf(c_0_283,negated_conjecture,(epred18_0|hskp10|~hskp7), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_267]), c_0_73]), c_0_72])).
cnf(c_0_284,negated_conjecture,(epred40_0|hskp7|hskp10), inference(spm,[status(thm)],[c_0_268, c_0_269])).
fof(c_0_285, plain, ((~epred7_0<=>![X1]:((c0_1(X1)|c2_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_286,negated_conjecture,(epred41_0|~c1_1(a1252)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_64]), c_0_45]), ['final']).
cnf(c_0_287,negated_conjecture,(epred44_0|c1_1(a1252)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_271]), c_0_46]), ['final']).
cnf(c_0_288,negated_conjecture,(epred23_0|c1_1(X1)|c0_1(X1)|hskp7|~c2_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_272])).
fof(c_0_289, plain, ((~epred12_0<=>![X1]:((((hskp14|c0_1(X1))|c2_1(X1))|~ndr1_0)|~c3_1(X1)))), introduced(definition)).
cnf(c_0_290,negated_conjecture,(epred35_0|~c1_1(X1)|~c2_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_273])).
fof(c_0_291, plain, ((~epred13_0<=>![X2]:((c0_1(X2)|c2_1(X2))|~c1_1(X2)))), introduced(definition)).
cnf(c_0_292,negated_conjecture,(hskp17|hskp18|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_173]), c_0_38]), c_0_39])).
cnf(c_0_293,negated_conjecture,(hskp10|hskp23|hskp24|~hskp0), inference(spm,[status(thm)],[c_0_186, c_0_274])).
cnf(c_0_294,negated_conjecture,(c1_1(a1272)|hskp0|hskp26|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_201]), c_0_57]), ['final']).
cnf(c_0_295,negated_conjecture,(hskp17|hskp18|~c1_1(a1301)|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_173]), c_0_35])).
cnf(c_0_296,negated_conjecture,(c1_1(a1301)|hskp10|hskp25|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_36]), c_0_34])).
cnf(c_0_297,negated_conjecture,(c0_1(a1260)|hskp17|hskp18|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_173]), c_0_60])).
cnf(c_0_298,negated_conjecture,(hskp10|~c0_1(a1254)|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_158]), c_0_89])).
cnf(c_0_299,negated_conjecture,(c0_1(a1254)|hskp17|hskp27|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_169]), c_0_89])).
cnf(c_0_300,negated_conjecture,(c1_1(X1)|c2_1(X1)|c0_1(X1)|hskp1|hskp27), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_153])]), ['final']).
cnf(c_0_301,negated_conjecture,(epred33_0|~hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_74]), c_0_75]), c_0_76])).
cnf(c_0_302,negated_conjecture,(epred11_0|hskp10|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_277]), c_0_278]), c_0_279])).
cnf(c_0_303,negated_conjecture,(epred19_0|~c3_1(a1266)|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_77]), c_0_91])).
cnf(c_0_304,negated_conjecture,(epred6_0|hskp10|hskp12|~hskp15), inference(spm,[status(thm)],[c_0_164, c_0_281])).
cnf(c_0_305,negated_conjecture,(epred24_0|~hskp10), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_79]), c_0_80]), c_0_81])).
cnf(c_0_306,negated_conjecture,(epred40_0|epred18_0|hskp10), inference(spm,[status(thm)],[c_0_283, c_0_284])).
cnf(c_0_307,negated_conjecture,(epred7_0|c2_1(X1)|c0_1(X1)|~c3_1(X1)), inference(split_equiv,[status(thm)],[c_0_285])).
fof(c_0_308, plain, ((~epred21_0<=>![X1]:((c2_1(X1)|~c0_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_309,negated_conjecture,(epred44_0|epred41_0|~hskp0), inference(spm,[status(thm)],[c_0_286, c_0_287]), ['final']).
fof(c_0_310, plain, ((~epred45_0<=>![X3]:(((c2_1(X3)|c1_1(X3))|~ndr1_0)|~c0_1(X3)))), introduced(definition)).
cnf(c_0_311,negated_conjecture,(epred23_0|c1_1(X1)|c0_1(X1)|hskp7|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_288, c_0_153])])).
cnf(c_0_312,negated_conjecture,(epred12_0|c2_1(X1)|c0_1(X1)|hskp14|~c3_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_289])).
cnf(c_0_313,negated_conjecture,(epred35_0|~c1_1(a1252)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_64]), c_0_46])).
cnf(c_0_314,negated_conjecture,(c1_1(a1252)|hskp7|hskp12|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_64]), c_0_45])).
cnf(c_0_315,negated_conjecture,(epred13_0|c2_1(X1)|c0_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_291])).
cnf(c_0_316,negated_conjecture,(hskp7|hskp26|~c1_1(a1301)|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_36]), c_0_35])).
cnf(c_0_317,negated_conjecture,(hskp10|hskp17|hskp18|hskp24|~hskp0), inference(spm,[status(thm)],[c_0_292, c_0_293])).
cnf(c_0_318,negated_conjecture,(hskp0|hskp17|hskp18|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_191, c_0_294])).
fof(c_0_319, plain, ((~epred16_0<=>![X1]:((((hskp19|c3_1(X1))|c1_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
fof(c_0_320, plain, ((~epred43_0<=>![X1]:(((c1_1(X1)|~ndr1_0)|~c0_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_321,negated_conjecture,(hskp10|hskp17|hskp18|hskp25|~hskp24), inference(spm,[status(thm)],[c_0_295, c_0_296])).
cnf(c_0_322,negated_conjecture,(hskp10|hskp17|hskp18|~hskp6), inference(spm,[status(thm)],[c_0_197, c_0_297])).
cnf(c_0_323,negated_conjecture,(hskp10|hskp17|hskp27|~hskp1), inference(spm,[status(thm)],[c_0_298, c_0_299])).
cnf(c_0_324,negated_conjecture,(hskp1|hskp27|~hskp18), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_300]), c_0_54]), c_0_55]), ['final']).
cnf(c_0_325,negated_conjecture,(epred40_0|c1_1(a1252)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_64]), c_0_45])).
cnf(c_0_326,negated_conjecture,(epred33_0|~c3_1(a1264)|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_79]), c_0_80])).
cnf(c_0_327,negated_conjecture,(epred33_0|epred11_0|hskp10), inference(spm,[status(thm)],[c_0_301, c_0_302])).
cnf(c_0_328,negated_conjecture,(epred33_0|hskp6|hskp10|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_301, c_0_227])).
cnf(c_0_329,negated_conjecture,(epred33_0|hskp6|hskp10|~hskp19), inference(spm,[status(thm)],[c_0_301, c_0_239])).
cnf(c_0_330,negated_conjecture,(epred33_0|~c3_1(a1250)|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_47]), c_0_48])).
cnf(c_0_331,negated_conjecture,(epred41_0|c3_1(a1250)|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_47]), c_0_48]), ['final']).
cnf(c_0_332,negated_conjecture,(epred41_0|c3_1(a1264)|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_79]), c_0_80])).
cnf(c_0_333,negated_conjecture,(epred24_0|c0_1(a1260)|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_92]), c_0_60])).
cnf(c_0_334,negated_conjecture,(epred33_0|hskp6|hskp7|hskp10), inference(spm,[status(thm)],[c_0_301, c_0_278])).
cnf(c_0_335,negated_conjecture,(epred19_0|~c3_1(a1264)|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_79]), c_0_81])).
cnf(c_0_336,negated_conjecture,(epred40_0|epred19_0|~hskp12), inference(spm,[status(thm)],[c_0_303, c_0_248])).
cnf(c_0_337,negated_conjecture,(epred6_0|hskp10|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_241]), c_0_269])).
cnf(c_0_338,negated_conjecture,(~epred19_0|~epred18_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_93, c_0_232]), c_0_264])).
cnf(c_0_339,negated_conjecture,(epred40_0|epred24_0|epred18_0), inference(spm,[status(thm)],[c_0_305, c_0_306])).
cnf(c_0_340,negated_conjecture,(epred7_0|~hskp7), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_307]), c_0_73]), c_0_72])).
cnf(c_0_341,negated_conjecture,(epred21_0|c2_1(X1)|~c3_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_308])).
cnf(c_0_342,negated_conjecture,(epred44_0|epred41_0|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_309, c_0_212]), ['final']).
cnf(c_0_343,negated_conjecture,(epred45_0|c1_1(X1)|c2_1(X1)|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_310])).
cnf(c_0_344,negated_conjecture,(epred23_0|hskp7|~hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_77]), c_0_91]), c_0_78])).
cnf(c_0_345,negated_conjecture,(epred12_0|c2_1(X1)|c0_1(X1)|hskp14|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_312, c_0_153])])).
cnf(c_0_346,negated_conjecture,(epred35_0|hskp7|hskp12|~hskp0), inference(spm,[status(thm)],[c_0_313, c_0_314])).
cnf(c_0_347,negated_conjecture,(epred35_0|~hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_47]), c_0_49]), c_0_48]), ['final']).
cnf(c_0_348,negated_conjecture,(epred35_0|~c0_1(a1259)|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_74]), c_0_76]), ['final']).
cnf(c_0_349,negated_conjecture,(epred19_0|c0_1(a1259)|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_74]), c_0_75])).
cnf(c_0_350,negated_conjecture,(epred13_0|~c1_1(a1261)|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_315]), c_0_72])).
cnf(c_0_351,negated_conjecture,(epred13_0|~hskp13), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_315]), c_0_86]), c_0_87])).
cnf(c_0_352,negated_conjecture,(hskp7|hskp25|hskp26|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_296]), c_0_235])).
cnf(c_0_353,negated_conjecture,(hskp10|hskp17|hskp18|hskp24|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_317, c_0_318])).
fof(c_0_354, plain, ((~epred29_0<=>![X1]:((((hskp20|c3_1(X1))|~ndr1_0)|~c2_1(X1))|~c1_1(X1)))), introduced(definition)).
fof(c_0_355, plain, ((~epred15_0<=>![X2]:((c0_1(X2)|c1_1(X2))|~c3_1(X2)))), introduced(definition)).
cnf(c_0_356,negated_conjecture,(epred16_0|c1_1(X1)|c3_1(X1)|hskp19|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_319])).
fof(c_0_357, plain, ((~epred17_0<=>![X2]:((c0_1(X2)|~c3_1(X2))|~c1_1(X2)))), introduced(definition)).
cnf(c_0_358,negated_conjecture,(epred43_0|c1_1(X1)|~c3_1(X1)|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_320])).
cnf(c_0_359,negated_conjecture,(hskp10|hskp17|hskp18|hskp28|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_321]), c_0_322])).
cnf(c_0_360,negated_conjecture,(hskp10|hskp17|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_323, c_0_324])).
cnf(c_0_361,negated_conjecture,(epred40_0|hskp10|~hskp0), inference(spm,[status(thm)],[c_0_186, c_0_325])).
cnf(c_0_362,negated_conjecture,(epred33_0|epred11_0|c0_1(a1264)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_236]), c_0_80]), c_0_327])).
cnf(c_0_363,negated_conjecture,(epred41_0|~hskp6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_92]), c_0_60]), c_0_59]), ['final']).
cnf(c_0_364,negated_conjecture,(epred33_0|hskp6|hskp10|hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_84]), c_0_329])).
cnf(c_0_365,negated_conjecture,(epred41_0|epred33_0|~hskp26), inference(spm,[status(thm)],[c_0_330, c_0_331])).
cnf(c_0_366,negated_conjecture,(epred41_0|epred33_0|~hskp10), inference(spm,[status(thm)],[c_0_326, c_0_332])).
cnf(c_0_367,negated_conjecture,(epred24_0|hskp10|~hskp6), inference(spm,[status(thm)],[c_0_197, c_0_333])).
cnf(c_0_368,negated_conjecture,(epred33_0|epred18_0|hskp6|hskp10), inference(spm,[status(thm)],[c_0_283, c_0_334])).
cnf(c_0_369,negated_conjecture,(epred19_0|epred11_0|~hskp10), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_236]), c_0_80]), c_0_81])).
cnf(c_0_370,negated_conjecture,(epred40_0|epred19_0|epred6_0|hskp10), inference(spm,[status(thm)],[c_0_336, c_0_337])).
cnf(c_0_371,negated_conjecture,(epred40_0|epred24_0|~epred19_0), inference(spm,[status(thm)],[c_0_338, c_0_339])).
cnf(c_0_372,negated_conjecture,(epred40_0|epred7_0|hskp10), inference(spm,[status(thm)],[c_0_340, c_0_284])).
cnf(c_0_373,negated_conjecture,(epred21_0|~c3_1(a1287)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_341]), c_0_95]), ['final']).
cnf(c_0_374,negated_conjecture,(epred44_0|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_271]), c_0_41]), c_0_42]), ['final']).
cnf(c_0_375,negated_conjecture,(epred33_0|hskp16|hskp20), inference(spm,[status(thm)],[c_0_301, c_0_1])).
cnf(c_0_376,negated_conjecture,(epred44_0|epred41_0|hskp19|hskp26), inference(spm,[status(thm)],[c_0_342, c_0_84])).
cnf(c_0_377,negated_conjecture,(epred45_0|c1_1(X1)|c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_343, c_0_153])]), ['final']).
cnf(c_0_378,negated_conjecture,(epred23_0|hskp7|hskp10), inference(spm,[status(thm)],[c_0_344, c_0_269])).
cnf(c_0_379,negated_conjecture,(epred19_0|~hskp14), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_96]), c_0_97]), c_0_98])).
cnf(c_0_380,negated_conjecture,(epred12_0|hskp14|~hskp7), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_345]), c_0_73]), c_0_72])).
cnf(c_0_381,negated_conjecture,(epred35_0|hskp7|hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_212]), c_0_251]), c_0_347])).
cnf(c_0_382,negated_conjecture,(epred35_0|epred19_0|~hskp28), inference(spm,[status(thm)],[c_0_348, c_0_349])).
cnf(c_0_383,negated_conjecture,(epred13_0|hskp28|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_257]), c_0_351])).
cnf(c_0_384,negated_conjecture,(c0_1(a1259)|hskp7|hskp26|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_74]), c_0_76])).
cnf(c_0_385,negated_conjecture,(hskp6|hskp7|hskp26|hskp28|~hskp24), inference(spm,[status(thm)],[c_0_199, c_0_352])).
cnf(c_0_386,negated_conjecture,(hskp10|hskp12|hskp17|hskp18|hskp24|hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_241]), c_0_242])).
fof(c_0_387, plain, ((~epred34_0<=>![X2]:((((hskp12|c2_1(X2))|~ndr1_0)|~c0_1(X2))|~c3_1(X2)))), introduced(definition)).
cnf(c_0_388,negated_conjecture,(epred29_0|c3_1(X1)|hskp20|~c1_1(X1)|~c2_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_354])).
cnf(c_0_389,negated_conjecture,(epred15_0|c1_1(X1)|c0_1(X1)|~c3_1(X1)), inference(split_equiv,[status(thm)],[c_0_355])).
cnf(c_0_390,negated_conjecture,(epred16_0|c1_1(X1)|c3_1(X1)|hskp19|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_356, c_0_153])])).
cnf(c_0_391,negated_conjecture,(c0_1(X1)|hskp19|hskp21|~c3_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_153])])).
cnf(c_0_392,negated_conjecture,(epred17_0|c0_1(X1)|~c1_1(X1)|~c3_1(X1)), inference(split_equiv,[status(thm)],[c_0_357])).
cnf(c_0_393,negated_conjecture,(epred43_0|c1_1(X1)|~c3_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_358, c_0_153])])).
cnf(c_0_394,negated_conjecture,(hskp10|hskp17|hskp27|hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_100]), c_0_360])).
cnf(c_0_395,negated_conjecture,(epred40_0|hskp10|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_361, c_0_212])).
cnf(c_0_396,negated_conjecture,(~epred41_0|~epred40_0|~epred11_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_101, c_0_204]), c_0_252]), c_0_218])).
cnf(c_0_397,negated_conjecture,(epred33_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_362]), c_0_327])).
cnf(c_0_398,negated_conjecture,(epred41_0|epred33_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_364]), c_0_365]), c_0_366])).
cnf(c_0_399,negated_conjecture,(epred33_0|epred24_0|epred18_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_368]), c_0_305])).
cnf(c_0_400,negated_conjecture,(epred33_0|epred19_0|epred11_0), inference(spm,[status(thm)],[c_0_369, c_0_327])).
cnf(c_0_401,negated_conjecture,(~epred7_0|~epred6_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_102, c_0_214]), c_0_285])).
cnf(c_0_402,negated_conjecture,(epred40_0|epred24_0|epred6_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_370]), c_0_371])).
cnf(c_0_403,negated_conjecture,(epred40_0|epred24_0|epred7_0), inference(spm,[status(thm)],[c_0_305, c_0_372])).
cnf(c_0_404,negated_conjecture,(epred44_0|epred21_0|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_271]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_405,negated_conjecture,(epred44_0|epred33_0|hskp20), inference(spm,[status(thm)],[c_0_374, c_0_375])).
cnf(c_0_406,negated_conjecture,(epred41_0|epred33_0|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_329]), c_0_366])).
cnf(c_0_407,negated_conjecture,(epred44_0|epred41_0|epred33_0|hskp19), inference(spm,[status(thm)],[c_0_365, c_0_376])).
cnf(c_0_408,negated_conjecture,(epred45_0|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_377]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_409,negated_conjecture,(epred23_0|epred18_0|hskp10), inference(spm,[status(thm)],[c_0_283, c_0_378])).
cnf(c_0_410,negated_conjecture,(epred19_0|epred12_0|~hskp7), inference(spm,[status(thm)],[c_0_379, c_0_380])).
cnf(c_0_411,negated_conjecture,(epred35_0|epred23_0|hskp7), inference(spm,[status(thm)],[c_0_344, c_0_381])).
cnf(c_0_412,negated_conjecture,(epred35_0|epred19_0|epred13_0|~hskp7), inference(spm,[status(thm)],[c_0_382, c_0_383])).
cnf(c_0_413,negated_conjecture,(epred35_0|hskp7|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_384]), c_0_347])).
cnf(c_0_414,negated_conjecture,(hskp6|hskp7|hskp26|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_385, c_0_386]), c_0_261]), c_0_261]), c_0_278]), c_0_235])).
cnf(c_0_415,negated_conjecture,(epred24_0|c0_1(a1259)|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_74]), c_0_76])).
cnf(c_0_416,negated_conjecture,(epred40_0|epred35_0|~hskp0), inference(spm,[status(thm)],[c_0_313, c_0_325])).
cnf(c_0_417,negated_conjecture,(epred34_0|c2_1(X1)|hskp12|~c3_1(X1)|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_387])).
cnf(c_0_418,negated_conjecture,(epred29_0|c3_1(X1)|hskp20|~c1_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_388, c_0_153])]), ['final']).
cnf(c_0_419,negated_conjecture,(epred45_0|~c0_1(a1258)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_377]), c_0_10]), ['final']).
cnf(c_0_420,negated_conjecture,(epred15_0|c0_1(a1258)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_9]), c_0_10]), ['final']).
cnf(c_0_421,negated_conjecture,(epred16_0|hskp19|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_390]), c_0_41]), c_0_42])).
cnf(c_0_422,negated_conjecture,(hskp19|hskp21|~c3_1(a1266)|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_77]), c_0_91])).
cnf(c_0_423,negated_conjecture,(epred17_0|~c1_1(a1261)|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_73]), c_0_72])).
cnf(c_0_424,negated_conjecture,(epred17_0|~hskp21), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_105]), c_0_106]), c_0_107])).
cnf(c_0_425,negated_conjecture,(hskp7|hskp10|hskp12|hskp21), inference(spm,[status(thm)],[c_0_193, c_0_108])).
cnf(c_0_426,negated_conjecture,(epred33_0|~c1_1(a1280)|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_17]), c_0_18])).
cnf(c_0_427,negated_conjecture,(epred43_0|c1_1(a1280)|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_18]), c_0_109])).
cnf(c_0_428,negated_conjecture,(epred33_0|hskp10|hskp17|hskp27), inference(spm,[status(thm)],[c_0_301, c_0_394])).
cnf(c_0_429,negated_conjecture,(epred33_0|hskp10|~hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_158]), c_0_49]), c_0_48])).
cnf(c_0_430,negated_conjecture,(epred40_0|hskp10|hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_84]), c_0_249])).
cnf(c_0_431,negated_conjecture,(epred33_0|~epred40_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_397]), c_0_398])).
cnf(c_0_432,negated_conjecture,(epred33_0|epred24_0|~epred19_0), inference(spm,[status(thm)],[c_0_338, c_0_399])).
cnf(c_0_433,negated_conjecture,(epred33_0|epred19_0|~epred40_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_400]), c_0_398])).
cnf(c_0_434,negated_conjecture,(epred40_0|epred24_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_402]), c_0_403])).
cnf(c_0_435,negated_conjecture,(~epred45_0|~epred41_0|~epred21_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_110, c_0_310]), c_0_252]), c_0_308])).
cnf(c_0_436,negated_conjecture,(epred44_0|epred33_0|epred21_0), inference(spm,[status(thm)],[c_0_404, c_0_405])).
cnf(c_0_437,negated_conjecture,(epred44_0|epred41_0|epred33_0), inference(spm,[status(thm)],[c_0_406, c_0_407])).
cnf(c_0_438,negated_conjecture,(epred45_0|epred44_0|epred33_0), inference(spm,[status(thm)],[c_0_408, c_0_405])).
cnf(c_0_439,negated_conjecture,(epred24_0|epred23_0|epred18_0), inference(spm,[status(thm)],[c_0_305, c_0_409])).
cnf(c_0_440,negated_conjecture,(~epred13_0|~epred12_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_111, c_0_289]), c_0_291])).
cnf(c_0_441,negated_conjecture,(epred35_0|epred23_0|epred19_0|epred12_0), inference(spm,[status(thm)],[c_0_410, c_0_411])).
cnf(c_0_442,negated_conjecture,(epred35_0|epred23_0|epred19_0|epred13_0), inference(spm,[status(thm)],[c_0_412, c_0_411])).
cnf(c_0_443,negated_conjecture,(epred35_0|hskp6|hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_414]), c_0_347])).
cnf(c_0_444,negated_conjecture,(epred35_0|epred24_0|~hskp28), inference(spm,[status(thm)],[c_0_348, c_0_415])).
cnf(c_0_445,negated_conjecture,(epred40_0|epred35_0|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_212]), c_0_347])).
cnf(c_0_446,negated_conjecture,(epred35_0|~c0_1(a1260)|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_92]), c_0_60])).
cnf(c_0_447,negated_conjecture,(epred34_0|c2_1(X1)|hskp12|~c3_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_417, c_0_153])]), ['final']).
cnf(c_0_448,negated_conjecture,(epred29_0|c3_1(a1264)|hskp20|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_79]), c_0_80])).
cnf(c_0_449,negated_conjecture,(epred45_0|epred15_0|~hskp5), inference(spm,[status(thm)],[c_0_419, c_0_420]), ['final']).
cnf(c_0_450,negated_conjecture,(epred33_0|epred16_0|hskp19|hskp20), inference(spm,[status(thm)],[c_0_421, c_0_375])).
cnf(c_0_451,negated_conjecture,(epred40_0|hskp21|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_248]), c_0_249])).
cnf(c_0_452,negated_conjecture,(epred17_0|hskp11|hskp12|~hskp7), inference(spm,[status(thm)],[c_0_423, c_0_223])).
cnf(c_0_453,negated_conjecture,(epred17_0|hskp7|hskp10|hskp12), inference(spm,[status(thm)],[c_0_424, c_0_425])).
cnf(c_0_454,negated_conjecture,(epred21_0|epred16_0|hskp19|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_390]), c_0_95]), c_0_103])).
cnf(c_0_455,negated_conjecture,(epred43_0|epred33_0|~hskp29), inference(spm,[status(thm)],[c_0_426, c_0_427])).
cnf(c_0_456,negated_conjecture,(epred33_0|c0_1(a1264)|hskp17|hskp27), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_169]), c_0_80]), c_0_428])).
cnf(c_0_457,negated_conjecture,(epred33_0|hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_430]), c_0_431])).
cnf(c_0_458,negated_conjecture,(~epred44_0|~epred43_0|~epred24_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_112, c_0_320]), c_0_253]), c_0_266])).
cnf(c_0_459,negated_conjecture,(epred33_0|epred24_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_432, c_0_433]), c_0_434])).
cnf(c_0_460,negated_conjecture,(epred44_0|epred33_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_436]), c_0_437]), c_0_438])).
cnf(c_0_461,negated_conjecture,(hskp23|hskp24|~c0_1(a1283)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_170]), c_0_30])).
cnf(c_0_462,negated_conjecture,(c0_1(a1283)|hskp6|hskp28|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_184]), c_0_30])).
cnf(c_0_463,negated_conjecture,(epred11_0|~c1_1(a1301)|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_236]), c_0_35]), ['final']).
cnf(c_0_464,negated_conjecture,(c1_1(a1301)|hskp6|hskp28|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_184]), c_0_35])).
fof(c_0_465, plain, ((~epred42_0<=>![X1]:(((c2_1(X1)|c1_1(X1))|~ndr1_0)|~c3_1(X1)))), introduced(definition)).
cnf(c_0_466,negated_conjecture,(epred24_0|epred23_0|~epred19_0), inference(spm,[status(thm)],[c_0_338, c_0_439])).
cnf(c_0_467,negated_conjecture,(epred35_0|epred23_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_441]), c_0_442])).
cnf(c_0_468,negated_conjecture,(epred35_0|epred18_0|hskp6|hskp10), inference(spm,[status(thm)],[c_0_283, c_0_443])).
cnf(c_0_469,negated_conjecture,(epred35_0|epred24_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_302]), c_0_305])).
cnf(c_0_470,negated_conjecture,(epred40_0|epred35_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_445, c_0_84]), c_0_347]), c_0_249])).
cnf(c_0_471,negated_conjecture,(epred35_0|epred24_0|~hskp6), inference(spm,[status(thm)],[c_0_446, c_0_333])).
cnf(c_0_472,negated_conjecture,(epred35_0|epred19_0|epred12_0|hskp6), inference(spm,[status(thm)],[c_0_410, c_0_443])).
cnf(c_0_473,negated_conjecture,(epred35_0|epred24_0|epred13_0|~hskp7), inference(spm,[status(thm)],[c_0_444, c_0_383])).
fof(c_0_474, plain, ((~epred36_0<=>![X2]:((((hskp21|~ndr1_0)|~c0_1(X2))|~c3_1(X2))|~c1_1(X2)))), introduced(definition)).
cnf(c_0_475,negated_conjecture,(epred34_0|hskp12|~c0_1(a1277)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_447]), c_0_114])).
cnf(c_0_476,negated_conjecture,(epred17_0|c0_1(a1277)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_114]), c_0_115])).
cnf(c_0_477,negated_conjecture,(epred33_0|epred29_0|hskp20|~hskp10), inference(spm,[status(thm)],[c_0_326, c_0_448])).
cnf(c_0_478,negated_conjecture,(epred45_0|epred15_0|hskp10|hskp21), inference(spm,[status(thm)],[c_0_449, c_0_108])).
cnf(c_0_479,negated_conjecture,(epred45_0|epred33_0|epred16_0|hskp19), inference(spm,[status(thm)],[c_0_408, c_0_450])).
cnf(c_0_480,negated_conjecture,(epred40_0|epred17_0|~hskp12), inference(spm,[status(thm)],[c_0_424, c_0_451])).
cnf(c_0_481,negated_conjecture,(epred17_0|hskp10|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_452]), c_0_453])).
cnf(c_0_482,negated_conjecture,(epred33_0|epred21_0|epred16_0|hskp19), inference(spm,[status(thm)],[c_0_454, c_0_450])).
cnf(c_0_483,negated_conjecture,(epred21_0|~c0_1(a1277)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_341]), c_0_114])).
cnf(c_0_484,negated_conjecture,(epred43_0|epred33_0|hskp17|~hskp27), inference(spm,[status(thm)],[c_0_455, c_0_179])).
cnf(c_0_485,negated_conjecture,(epred33_0|hskp17|hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_456]), c_0_457])).
cnf(c_0_486,negated_conjecture,(epred33_0|~epred43_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_458, c_0_459]), c_0_460])).
cnf(c_0_487,negated_conjecture,(epred11_0|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_236]), c_0_38]), c_0_39]), ['final']).
cnf(c_0_488,negated_conjecture,(hskp6|hskp23|hskp24|hskp28|~hskp19), inference(spm,[status(thm)],[c_0_461, c_0_462])).
cnf(c_0_489,negated_conjecture,(epred11_0|hskp6|hskp28|~hskp24), inference(spm,[status(thm)],[c_0_463, c_0_464])).
cnf(c_0_490,negated_conjecture,(epred42_0|c1_1(X1)|c2_1(X1)|~c3_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_465])).
cnf(c_0_491,negated_conjecture,(hskp5|hskp19|~c3_1(a1264)|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_79]), c_0_80])).
cnf(c_0_492,negated_conjecture,(epred11_0|hskp5|hskp10|hskp19), inference(spm,[status(thm)],[c_0_226, c_0_302])).
cnf(c_0_493,negated_conjecture,(~epred42_0|~epred24_0|~epred19_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_116, c_0_465]), c_0_264]), c_0_266])).
cnf(c_0_494,negated_conjecture,(epred35_0|epred24_0|epred23_0), inference(spm,[status(thm)],[c_0_466, c_0_467])).
cnf(c_0_495,negated_conjecture,(epred41_0|epred35_0|epred18_0|hskp10), inference(spm,[status(thm)],[c_0_363, c_0_468])).
cnf(c_0_496,negated_conjecture,(epred35_0|epred24_0|~epred41_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_469]), c_0_470])).
cnf(c_0_497,negated_conjecture,(epred35_0|epred24_0|epred19_0|epred12_0), inference(spm,[status(thm)],[c_0_471, c_0_472])).
cnf(c_0_498,negated_conjecture,(epred35_0|epred24_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_443]), c_0_471])).
cnf(c_0_499,negated_conjecture,(epred36_0|hskp21|~c1_1(X1)|~c3_1(X1)|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_474])).
cnf(c_0_500,negated_conjecture,(epred34_0|epred17_0|hskp12|~hskp17), inference(spm,[status(thm)],[c_0_475, c_0_476])).
fof(c_0_501, plain, ((~epred22_0<=>![X1]:((((hskp9|c1_1(X1))|~ndr1_0)|~c2_1(X1))|~c3_1(X1)))), introduced(definition)).
cnf(c_0_502,negated_conjecture,(epred45_0|epred33_0|epred29_0|~hskp10), inference(spm,[status(thm)],[c_0_408, c_0_477])).
cnf(c_0_503,negated_conjecture,(epred45_0|epred17_0|epred15_0|hskp10), inference(spm,[status(thm)],[c_0_424, c_0_478])).
cnf(c_0_504,negated_conjecture,(~epred17_0|~epred16_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_117, c_0_319]), c_0_357])).
cnf(c_0_505,negated_conjecture,(epred45_0|epred40_0|epred33_0|epred16_0), inference(spm,[status(thm)],[c_0_249, c_0_479])).
cnf(c_0_506,negated_conjecture,(epred40_0|epred17_0|hskp10), inference(spm,[status(thm)],[c_0_480, c_0_481])).
cnf(c_0_507,negated_conjecture,(epred40_0|epred33_0|epred21_0|epred16_0), inference(spm,[status(thm)],[c_0_249, c_0_482])).
cnf(c_0_508,negated_conjecture,(epred21_0|epred17_0|~hskp17), inference(spm,[status(thm)],[c_0_483, c_0_476])).
cnf(c_0_509,negated_conjecture,(epred33_0|hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_484, c_0_485]), c_0_486])).
cnf(c_0_510,negated_conjecture,(epred11_0|hskp6|hskp28|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_488]), c_0_489])).
cnf(c_0_511,negated_conjecture,(epred42_0|c1_1(X1)|c2_1(X1)|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_490, c_0_153])]), ['final']).
cnf(c_0_512,negated_conjecture,(epred11_0|c0_1(a1264)|hskp5|hskp19), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_236]), c_0_80]), c_0_492])).
cnf(c_0_513,negated_conjecture,(~epred24_0|~epred23_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_118, c_0_272]), c_0_266])).
cnf(c_0_514,negated_conjecture,(epred35_0|epred23_0|~epred42_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_467]), c_0_494])).
cnf(c_0_515,negated_conjecture,(epred35_0|epred24_0|epred18_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_495]), c_0_496])).
cnf(c_0_516,negated_conjecture,(epred35_0|epred24_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_497]), c_0_498])).
cnf(c_0_517,negated_conjecture,(epred36_0|hskp21|~c1_1(X1)|~c3_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_499, c_0_153])]), ['final']).
cnf(c_0_518,negated_conjecture,(epred40_0|epred34_0|epred17_0|~hskp17), inference(spm,[status(thm)],[c_0_480, c_0_500])).
cnf(c_0_519,negated_conjecture,(~epred22_0|~epred15_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_119, c_0_501]), c_0_355])).
cnf(c_0_520,negated_conjecture,(epred45_0|epred33_0|epred29_0|epred17_0|epred15_0), inference(spm,[status(thm)],[c_0_502, c_0_503])).
cnf(c_0_521,negated_conjecture,(epred45_0|epred40_0|epred33_0|~epred17_0), inference(spm,[status(thm)],[c_0_504, c_0_505])).
cnf(c_0_522,negated_conjecture,(epred45_0|epred40_0|epred33_0|epred29_0|epred17_0), inference(spm,[status(thm)],[c_0_502, c_0_506])).
cnf(c_0_523,negated_conjecture,(epred40_0|epred33_0|epred21_0|~epred17_0), inference(spm,[status(thm)],[c_0_504, c_0_507])).
cnf(c_0_524,negated_conjecture,(epred33_0|epred21_0|epred17_0), inference(spm,[status(thm)],[c_0_508, c_0_509])).
cnf(c_0_525,negated_conjecture,(epred35_0|epred11_0|~hskp6), inference(spm,[status(thm)],[c_0_446, c_0_263])).
cnf(c_0_526,negated_conjecture,(epred33_0|epred11_0|hskp6|~hskp19), inference(spm,[status(thm)],[c_0_301, c_0_510])).
cnf(c_0_527,negated_conjecture,(epred42_0|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_511]), c_0_9]), c_0_10]), ['final']).
cnf(c_0_528,negated_conjecture,(epred11_0|hskp5|hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_512]), c_0_492])).
cnf(c_0_529,negated_conjecture,(epred35_0|~epred42_0|~epred24_0), inference(spm,[status(thm)],[c_0_513, c_0_514])).
cnf(c_0_530,negated_conjecture,(epred35_0|epred24_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_515]), c_0_516])).
cnf(c_0_531,negated_conjecture,(epred36_0|hskp21|~c0_1(a1277)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_114]), c_0_115])).
cnf(c_0_532,negated_conjecture,(~epred29_0|~epred17_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_120, c_0_354]), c_0_357])).
cnf(c_0_533,negated_conjecture,(epred34_0|epred33_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_518, c_0_509]), c_0_431])).
cnf(c_0_534,negated_conjecture,(epred45_0|epred33_0|epred29_0|epred17_0|~epred22_0), inference(spm,[status(thm)],[c_0_519, c_0_520])).
cnf(c_0_535,negated_conjecture,(epred45_0|epred40_0|epred33_0|epred29_0), inference(spm,[status(thm)],[c_0_521, c_0_522])).
cnf(c_0_536,negated_conjecture,(epred33_0|epred21_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_523, c_0_524]), c_0_431])).
cnf(c_0_537,negated_conjecture,(epred35_0|epred33_0|epred11_0|~hskp19), inference(spm,[status(thm)],[c_0_525, c_0_526])).
cnf(c_0_538,negated_conjecture,(epred42_0|epred11_0|hskp19), inference(spm,[status(thm)],[c_0_527, c_0_528])).
cnf(c_0_539,negated_conjecture,(epred35_0|~epred42_0), inference(spm,[status(thm)],[c_0_529, c_0_530]), ['final']).
fof(c_0_540, plain, ((~epred25_0<=>![X1]:((((hskp4|c3_1(X1))|~ndr1_0)|~c0_1(X1))|~c2_1(X1)))), introduced(definition)).
cnf(c_0_541,negated_conjecture,(c1_1(X1)|c3_1(X1)|c2_1(X1)|hskp6|hskp15), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_153])]), ['final']).
cnf(c_0_542,negated_conjecture,(epred36_0|epred17_0|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_531, c_0_476]), c_0_424])).
cnf(c_0_543,negated_conjecture,(epred34_0|epred33_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_533])).
cnf(c_0_544,negated_conjecture,(epred45_0|epred33_0|epred29_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_521, c_0_534]), c_0_535]), c_0_431])).
cnf(c_0_545,negated_conjecture,(epred33_0|~epred45_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_536]), c_0_398])).
cnf(c_0_546,negated_conjecture,(epred35_0|epred33_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_537, c_0_538]), c_0_539])).
cnf(c_0_547,negated_conjecture,(epred25_0|c3_1(X1)|hskp4|~c2_1(X1)|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_540])).
cnf(c_0_548,negated_conjecture,(hskp13|hskp28|~c3_1(a1287)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_237]), c_0_103]), ['final']).
cnf(c_0_549,negated_conjecture,(c3_1(a1287)|hskp6|hskp15|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_541]), c_0_103]), ['final']).
cnf(c_0_550,negated_conjecture,(c1_1(X1)|c3_1(X1)|c2_1(X1)|hskp17|hskp22), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_153])])).
cnf(c_0_551,negated_conjecture,(epred36_0|epred33_0|epred17_0), inference(spm,[status(thm)],[c_0_542, c_0_509])).
cnf(c_0_552,negated_conjecture,(~epred36_0|~epred33_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_123, c_0_474]), c_0_254])).
cnf(c_0_553,negated_conjecture,(epred34_0|epred33_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_543, c_0_544]), c_0_545])).
cnf(c_0_554,negated_conjecture,(epred35_0|epred33_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_546]), c_0_470]), c_0_398])).
cnf(c_0_555,negated_conjecture,(epred25_0|c3_1(X1)|hskp4|~c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_547, c_0_153])])).
cnf(c_0_556,negated_conjecture,(hskp7|hskp19|hskp23|hskp24|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_170]), c_0_124]), c_0_13])).
cnf(c_0_557,negated_conjecture,(hskp6|hskp17|hskp18|hskp24|hskp28|~hskp19), inference(spm,[status(thm)],[c_0_292, c_0_488])).
cnf(c_0_558,negated_conjecture,(hskp6|hskp17|hskp18|hskp28|~hskp15), inference(spm,[status(thm)],[c_0_191, c_0_256])).
cnf(c_0_559,negated_conjecture,(hskp6|hskp13|hskp15|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_548, c_0_549])).
cnf(c_0_560,negated_conjecture,(hskp10|hskp17|hskp29|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_168, c_0_158])).
cnf(c_0_561,negated_conjecture,(hskp17|hskp27|~c1_1(a1272)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_169]), c_0_57])).
cnf(c_0_562,negated_conjecture,(c1_1(a1272)|hskp17|hskp22|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_550]), c_0_56])).
fof(c_0_563, plain, ((~epred32_0<=>![X2]:((((hskp7|c3_1(X2))|c1_1(X2))|~ndr1_0)|~c0_1(X2)))), introduced(definition)).
cnf(c_0_564,negated_conjecture,(epred36_0|hskp21|~c0_1(a1259)|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_75]), c_0_76]), ['final']).
cnf(c_0_565,negated_conjecture,(epred17_0|c0_1(a1259)|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_75]), c_0_76])).
cnf(c_0_566,negated_conjecture,(epred36_0|epred33_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_551])).
cnf(c_0_567,negated_conjecture,(~epred35_0|~epred34_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_125, c_0_387]), c_0_273])).
cnf(c_0_568,negated_conjecture,(epred34_0|~epred36_0), inference(spm,[status(thm)],[c_0_552, c_0_553]), ['final']).
cnf(c_0_569,negated_conjecture,(epred35_0|~epred36_0), inference(spm,[status(thm)],[c_0_552, c_0_554]), ['final']).
cnf(c_0_570,negated_conjecture,(epred25_0|hskp4|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_64]), c_0_46]), c_0_45])).
cnf(c_0_571,negated_conjecture,(hskp7|hskp17|hskp18|hskp19|hskp24|~hskp2), inference(spm,[status(thm)],[c_0_292, c_0_556])).
cnf(c_0_572,negated_conjecture,(hskp10|hskp17|hskp18|hskp28|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_557]), c_0_322])).
cnf(c_0_573,negated_conjecture,(hskp10|hskp17|hskp18|hskp28|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_558, c_0_277]), c_0_322])).
cnf(c_0_574,negated_conjecture,(hskp2|hskp20|~c1_1(X1)|~c3_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_153])])).
cnf(c_0_575,negated_conjecture,(hskp10|hskp17|hskp18|hskp28|~hskp13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_558, c_0_258]), c_0_322])).
cnf(c_0_576,negated_conjecture,(hskp6|hskp13|hskp17|hskp18|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_558, c_0_559])).
cnf(c_0_577,negated_conjecture,(hskp10|hskp17|hskp29|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_560, c_0_127]), c_0_128])).
cnf(c_0_578,negated_conjecture,(hskp17|hskp22|hskp27|~hskp15), inference(spm,[status(thm)],[c_0_561, c_0_562])).
fof(c_0_579, plain, ((~epred38_0<=>![X1]:(((c1_1(X1)|~ndr1_0)|~c0_1(X1))|~c2_1(X1)))), introduced(definition)).
fof(c_0_580, plain, ((~epred14_0<=>![X1]:((((hskp8|c3_1(X1))|c1_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
fof(c_0_581, plain, ((~epred37_0<=>![X3]:(((c0_1(X3)|c1_1(X3))|~ndr1_0)|~c2_1(X3)))), introduced(definition)).
cnf(c_0_582,negated_conjecture,(epred32_0|c1_1(X1)|c3_1(X1)|hskp7|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_563])).
cnf(c_0_583,negated_conjecture,(epred36_0|epred17_0|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_564, c_0_565]), c_0_424])).
cnf(c_0_584,negated_conjecture,(epred36_0|epred33_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_544]), c_0_545])).
cnf(c_0_585,negated_conjecture,(~epred36_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_567, c_0_568]), c_0_569]), ['final']).
cnf(c_0_586,negated_conjecture,(epred25_0|hskp4|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_570, c_0_212])).
cnf(c_0_587,negated_conjecture,(hskp13|hskp28|~c3_1(a1279)|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_237]), c_0_55]), ['final']).
cnf(c_0_588,negated_conjecture,(c3_1(a1279)|hskp17|hskp22|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_550]), c_0_55])).
cnf(c_0_589,negated_conjecture,(hskp10|hskp17|hskp18|hskp28|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_571]), c_0_572]), c_0_573])).
cnf(c_0_590,negated_conjecture,(hskp2|hskp20|~hskp27), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_31]), c_0_32]), c_0_33]), ['final']).
cnf(c_0_591,negated_conjecture,(hskp10|hskp17|hskp18|hskp28|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_575, c_0_576]), c_0_322])).
cnf(c_0_592,negated_conjecture,(hskp10|hskp17|hskp29|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_577, c_0_578]), c_0_179])).
cnf(c_0_593,negated_conjecture,(c1_1(a1261)|hskp1|hskp27|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_300]), c_0_72])).
cnf(c_0_594,negated_conjecture,(hskp10|hskp15|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_207]), c_0_89]), c_0_298])).
cnf(c_0_595,negated_conjecture,(epred38_0|c1_1(X1)|~c2_1(X1)|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_579])).
fof(c_0_596, plain, ((~epred4_0<=>![X1]:((c2_1(X1)|c3_1(X1))|~c0_1(X1)))), introduced(definition)).
cnf(c_0_597,negated_conjecture,(epred14_0|c1_1(X1)|c3_1(X1)|hskp8|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_580])).
cnf(c_0_598,negated_conjecture,(epred37_0|c1_1(X1)|c0_1(X1)|~c2_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_581])).
cnf(c_0_599,negated_conjecture,(epred21_0|~c0_1(a1258)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_341]), c_0_9]), ['final']).
cnf(c_0_600,negated_conjecture,(epred18_0|c0_1(a1258)|hskp10|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_267]), c_0_9])).
cnf(c_0_601,negated_conjecture,(epred32_0|c1_1(X1)|c3_1(X1)|hskp7|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_582, c_0_153])])).
cnf(c_0_602,negated_conjecture,(epred36_0|epred17_0|hskp16|hskp20), inference(spm,[status(thm)],[c_0_583, c_0_1])).
cnf(c_0_603,negated_conjecture,(~epred33_0|~epred32_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_130, c_0_563]), c_0_254])).
cnf(c_0_604,negated_conjecture,(epred33_0), inference(sr,[status(thm)],[c_0_584, c_0_585]), ['final']).
cnf(c_0_605,negated_conjecture,(epred25_0|c3_1(a1250)|hskp4|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_47]), c_0_49])).
cnf(c_0_606,negated_conjecture,(epred25_0|hskp4|hskp19|hskp26), inference(spm,[status(thm)],[c_0_586, c_0_84])).
cnf(c_0_607,negated_conjecture,(epred17_0|epred11_0|c0_1(X1)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_392, c_0_236])).
cnf(c_0_608,negated_conjecture,(epred42_0|hskp10|hskp21), inference(spm,[status(thm)],[c_0_527, c_0_108])).
cnf(c_0_609,negated_conjecture,(hskp13|hskp17|hskp22|hskp28|~hskp18), inference(spm,[status(thm)],[c_0_587, c_0_588])).
cnf(c_0_610,negated_conjecture,(hskp10|hskp17|hskp18|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_589, c_0_590]), c_0_394]), c_0_591])).
cnf(c_0_611,negated_conjecture,(hskp10|hskp17|hskp29|~hskp13), inference(spm,[status(thm)],[c_0_592, c_0_258])).
cnf(c_0_612,negated_conjecture,(hskp10|hskp15|hskp27|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_593]), c_0_594])).
cnf(c_0_613,negated_conjecture,(c3_1(X1)|c2_1(X1)|hskp0|hskp11|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_153])]), ['final']).
cnf(c_0_614,negated_conjecture,(epred38_0|c1_1(X1)|~c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_595, c_0_153])])).
cnf(c_0_615,negated_conjecture,(epred4_0|c3_1(X1)|c2_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_596])).
cnf(c_0_616,negated_conjecture,(epred14_0|c1_1(X1)|c3_1(X1)|hskp8|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_597, c_0_153])]), ['final']).
cnf(c_0_617,negated_conjecture,(epred37_0|c1_1(X1)|c0_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_598, c_0_153])])).
fof(c_0_618, plain, ((~epred20_0<=>![X2]:((((hskp10|c0_1(X2))|c1_1(X2))|~ndr1_0)|~c3_1(X2)))), introduced(definition)).
cnf(c_0_619,negated_conjecture,(epred21_0|epred18_0|hskp10|~hskp5), inference(spm,[status(thm)],[c_0_599, c_0_600])).
cnf(c_0_620,negated_conjecture,(epred32_0|hskp7|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_601]), c_0_41]), c_0_42])).
cnf(c_0_621,negated_conjecture,(epred17_0|hskp16|hskp20), inference(sr,[status(thm)],[c_0_602, c_0_585])).
cnf(c_0_622,negated_conjecture,(~epred32_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_603, c_0_604])]), ['final']).
cnf(c_0_623,negated_conjecture,(epred25_0|hskp4|hskp5|hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_605]), c_0_606])).
cnf(c_0_624,negated_conjecture,(epred17_0|epred11_0|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_607]), c_0_80])).
cnf(c_0_625,negated_conjecture,(epred42_0|epred17_0|hskp10), inference(spm,[status(thm)],[c_0_424, c_0_608])).
cnf(c_0_626,negated_conjecture,(epred45_0|epred18_0|hskp10|~hskp5), inference(spm,[status(thm)],[c_0_419, c_0_600])).
cnf(c_0_627,negated_conjecture,(hskp10|hskp17|hskp28|hskp29), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_577, c_0_609]), c_0_610]), c_0_611])).
cnf(c_0_628,negated_conjecture,(hskp10|hskp17|hskp29|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_560, c_0_49]), c_0_48])).
cnf(c_0_629,negated_conjecture,(hskp10|hskp17|hskp29|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_592, c_0_612]), c_0_179])).
cnf(c_0_630,negated_conjecture,(c3_1(a1287)|hskp0|hskp11|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_613]), c_0_95]), ['final']).
cnf(c_0_631,negated_conjecture,(hskp2|hskp20|~c1_1(a1280)|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_18]), c_0_109]), ['final']).
cnf(c_0_632,negated_conjecture,(epred38_0|c1_1(a1280)|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_614, c_0_17]), c_0_109])).
cnf(c_0_633,negated_conjecture,(epred38_0|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_614, c_0_12]), c_0_124]), c_0_13])).
cnf(c_0_634,negated_conjecture,(c3_1(X1)|c0_1(X1)|hskp6|hskp29|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_132, c_0_153])])).
cnf(c_0_635,negated_conjecture,(epred4_0|~hskp8), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_615]), c_0_134]), c_0_135]), ['final']).
cnf(c_0_636,negated_conjecture,(epred14_0|hskp8|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_616]), c_0_41]), c_0_42]), ['final']).
cnf(c_0_637,negated_conjecture,(epred37_0|~hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_617, c_0_77]), c_0_91]), c_0_78])).
fof(c_0_638, plain, ((~epred31_0<=>![X1]:((c2_1(X1)|~c3_1(X1))|~c1_1(X1)))), introduced(definition)).
fof(c_0_639, plain, ((~epred39_0<=>![X2]:((c2_1(X2)|c1_1(X2))|~c0_1(X2)))), introduced(definition)).
cnf(c_0_640,negated_conjecture,(epred20_0|c1_1(X1)|c0_1(X1)|hskp10|~c3_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_618])).
cnf(c_0_641,negated_conjecture,(epred21_0|epred18_0|hskp10|hskp21), inference(spm,[status(thm)],[c_0_619, c_0_108])).
cnf(c_0_642,negated_conjecture,(epred32_0|epred21_0|hskp7|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_601]), c_0_95]), c_0_103])).
cnf(c_0_643,negated_conjecture,(epred17_0|hskp7|hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_620, c_0_621]), c_0_622])).
cnf(c_0_644,negated_conjecture,(epred42_0|c1_1(a1261)|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_511]), c_0_73])).
cnf(c_0_645,negated_conjecture,(epred42_0|epred25_0|hskp4|hskp19), inference(spm,[status(thm)],[c_0_527, c_0_623])).
cnf(c_0_646,negated_conjecture,(~epred25_0|~epred11_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_136, c_0_540]), c_0_218])).
cnf(c_0_647,negated_conjecture,(epred42_0|epred17_0|epred11_0), inference(spm,[status(thm)],[c_0_624, c_0_625])).
cnf(c_0_648,negated_conjecture,(epred45_0|epred18_0|hskp10|hskp21), inference(spm,[status(thm)],[c_0_626, c_0_108])).
cnf(c_0_649,negated_conjecture,(epred21_0|~c3_1(a1291)|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_341]), c_0_127]), ['final']).
cnf(c_0_650,negated_conjecture,(c3_1(a1287)|hskp17|hskp22|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_550]), c_0_103])).
cnf(c_0_651,negated_conjecture,(hskp10|hskp17|hskp29|~c1_1(a1259)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_560, c_0_384]), c_0_627]), c_0_628]), c_0_629])).
cnf(c_0_652,negated_conjecture,(hskp6|hskp9|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_64]), c_0_46]), c_0_45]), ['final']).
cnf(c_0_653,negated_conjecture,(hskp0|hskp2|hskp4|hskp11|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_630]), c_0_95]), c_0_103])).
cnf(c_0_654,negated_conjecture,(epred38_0|c1_1(a1252)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_614, c_0_64]), c_0_46])).
cnf(c_0_655,negated_conjecture,(epred38_0|hskp20|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_631, c_0_632]), c_0_633])).
cnf(c_0_656,negated_conjecture,(hskp6|hskp29|~hskp24), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_634, c_0_36]), c_0_35]), c_0_34])).
cnf(c_0_657,negated_conjecture,(hskp2|hskp10|hskp20|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_574, c_0_158])).
cnf(c_0_658,negated_conjecture,(epred24_0|epred13_0|c0_1(X1)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_282, c_0_315])).
cnf(c_0_659,negated_conjecture,(epred14_0|epred4_0|~hskp16), inference(spm,[status(thm)],[c_0_635, c_0_636]), ['final']).
cnf(c_0_660,negated_conjecture,(epred37_0|hskp7|hskp10), inference(spm,[status(thm)],[c_0_637, c_0_269])).
cnf(c_0_661,negated_conjecture,(epred31_0|c2_1(X1)|~c1_1(X1)|~c3_1(X1)), inference(split_equiv,[status(thm)],[c_0_638])).
cnf(c_0_662,negated_conjecture,(epred39_0|c1_1(X1)|c2_1(X1)|~c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_639])).
cnf(c_0_663,negated_conjecture,(epred20_0|c1_1(X1)|c0_1(X1)|hskp10|~c3_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_640, c_0_153])])).
cnf(c_0_664,negated_conjecture,(epred21_0|epred18_0|epred17_0|hskp10), inference(spm,[status(thm)],[c_0_424, c_0_641])).
cnf(c_0_665,negated_conjecture,(epred21_0|epred17_0|hskp7), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_642, c_0_643]), c_0_622])).
cnf(c_0_666,negated_conjecture,(epred36_0|epred17_0|epred13_0|~hskp7), inference(spm,[status(thm)],[c_0_583, c_0_383])).
cnf(c_0_667,negated_conjecture,(epred42_0|epred17_0|~hskp7), inference(spm,[status(thm)],[c_0_423, c_0_644])).
cnf(c_0_668,negated_conjecture,(epred23_0|hskp7|~hskp19), inference(spm,[status(thm)],[c_0_344, c_0_177])).
cnf(c_0_669,negated_conjecture,(epred42_0|epred25_0|hskp7|hskp19), inference(spm,[status(thm)],[c_0_176, c_0_645])).
cnf(c_0_670,negated_conjecture,(epred42_0|epred17_0|~epred25_0), inference(spm,[status(thm)],[c_0_646, c_0_647])).
cnf(c_0_671,negated_conjecture,(epred45_0|epred18_0|epred17_0|hskp10), inference(spm,[status(thm)],[c_0_424, c_0_648])).
cnf(c_0_672,negated_conjecture,(epred45_0|epred17_0|hskp7), inference(spm,[status(thm)],[c_0_408, c_0_643])).
cnf(c_0_673,negated_conjecture,(epred21_0|hskp10|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_649, c_0_158]), c_0_127]), c_0_128])).
cnf(c_0_674,negated_conjecture,(epred21_0|hskp17|hskp22|~hskp20), inference(spm,[status(thm)],[c_0_373, c_0_650])).
cnf(c_0_675,negated_conjecture,(hskp10|hskp17|hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_651, c_0_76]), c_0_627])).
cnf(c_0_676,negated_conjecture,(hskp2|hskp4|hskp6|hskp11|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_653]), c_0_202])).
cnf(c_0_677,negated_conjecture,(epred38_0|hskp10|~hskp0), inference(spm,[status(thm)],[c_0_186, c_0_654])).
cnf(c_0_678,negated_conjecture,(epred38_0|hskp6|hskp20|~hskp24), inference(spm,[status(thm)],[c_0_655, c_0_656])).
cnf(c_0_679,negated_conjecture,(epred38_0|hskp20|~hskp27), inference(spm,[status(thm)],[c_0_633, c_0_590])).
cnf(c_0_680,negated_conjecture,(hskp2|hskp10|hskp20|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_49]), c_0_48])).
cnf(c_0_681,negated_conjecture,(epred24_0|~hskp6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_658]), c_0_367]), c_0_367]), c_0_305])).
cnf(c_0_682,negated_conjecture,(epred14_0|epred4_0|hskp6|hskp17), inference(spm,[status(thm)],[c_0_659, c_0_83])).
cnf(c_0_683,negated_conjecture,(epred37_0|epred6_0|hskp10), inference(spm,[status(thm)],[c_0_637, c_0_337])).
cnf(c_0_684,negated_conjecture,(epred37_0|epred7_0|hskp10), inference(spm,[status(thm)],[c_0_340, c_0_660])).
cnf(c_0_685,negated_conjecture,(epred31_0|~c1_1(a1261)|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_661]), c_0_73])).
cnf(c_0_686,negated_conjecture,(epred15_0|c1_1(a1261)|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_73]), c_0_72])).
cnf(c_0_687,negated_conjecture,(epred32_0|hskp6|hskp7|hskp17), inference(spm,[status(thm)],[c_0_620, c_0_83])).
cnf(c_0_688,negated_conjecture,(epred39_0|~c0_1(a1258)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_662]), c_0_10]), ['final']).
cnf(c_0_689,negated_conjecture,(epred20_0|c0_1(a1258)|hskp10|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_663, c_0_9]), c_0_10])).
fof(c_0_690, plain, ((~epred30_0<=>![X1]:((((hskp19|c1_1(X1))|~ndr1_0)|~c0_1(X1))|~c2_1(X1)))), introduced(definition)).
cnf(c_0_691,negated_conjecture,(epred20_0|c1_1(a1261)|hskp10|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_663, c_0_73]), c_0_72])).
cnf(c_0_692,negated_conjecture,(epred41_0|epred19_0|~hskp10), inference(spm,[status(thm)],[c_0_335, c_0_332])).
cnf(c_0_693,negated_conjecture,(epred23_0|epred19_0|epred12_0|hskp10), inference(spm,[status(thm)],[c_0_410, c_0_378])).
cnf(c_0_694,negated_conjecture,(epred24_0|epred21_0|epred18_0|epred17_0), inference(spm,[status(thm)],[c_0_305, c_0_664])).
cnf(c_0_695,negated_conjecture,(epred21_0|epred19_0|epred17_0|epred12_0), inference(spm,[status(thm)],[c_0_410, c_0_665])).
cnf(c_0_696,negated_conjecture,(epred21_0|epred17_0|epred13_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_666, c_0_665]), c_0_585])).
cnf(c_0_697,negated_conjecture,(epred42_0|epred23_0|epred17_0|~hskp19), inference(spm,[status(thm)],[c_0_667, c_0_668])).
cnf(c_0_698,negated_conjecture,(epred42_0|epred17_0|hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_667, c_0_669]), c_0_670])).
cnf(c_0_699,negated_conjecture,(epred45_0|epred24_0|epred18_0|epred17_0), inference(spm,[status(thm)],[c_0_305, c_0_671])).
cnf(c_0_700,negated_conjecture,(epred45_0|epred19_0|epred17_0|epred12_0), inference(spm,[status(thm)],[c_0_410, c_0_672])).
cnf(c_0_701,negated_conjecture,(epred45_0|epred17_0|epred13_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_666, c_0_672]), c_0_585])).
cnf(c_0_702,negated_conjecture,(epred18_0|c0_1(a1277)|hskp10|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_267]), c_0_114])).
cnf(c_0_703,negated_conjecture,(epred21_0|hskp10|hskp17|~hskp20), inference(spm,[status(thm)],[c_0_673, c_0_674])).
cnf(c_0_704,negated_conjecture,(epred38_0|hskp10|hskp17|hskp20), inference(spm,[status(thm)],[c_0_655, c_0_675])).
cnf(c_0_705,negated_conjecture,(hskp2|hskp4|hskp6|hskp10|~hskp20), inference(spm,[status(thm)],[c_0_164, c_0_676])).
cnf(c_0_706,negated_conjecture,(epred38_0|hskp10|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_677, c_0_212])).
cnf(c_0_707,negated_conjecture,(epred38_0|hskp6|hskp18|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_100]), c_0_679])).
cnf(c_0_708,negated_conjecture,(epred38_0|hskp10|hskp20|~hskp26), inference(spm,[status(thm)],[c_0_633, c_0_680])).
cnf(c_0_709,negated_conjecture,(epred31_0|~hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_661]), c_0_114]), c_0_115])).
cnf(c_0_710,negated_conjecture,(epred24_0|epred14_0|epred4_0|hskp17), inference(spm,[status(thm)],[c_0_681, c_0_682])).
cnf(c_0_711,negated_conjecture,(epred37_0|epred24_0|epred6_0), inference(spm,[status(thm)],[c_0_305, c_0_683])).
cnf(c_0_712,negated_conjecture,(epred37_0|epred24_0|epred7_0), inference(spm,[status(thm)],[c_0_305, c_0_684])).
cnf(c_0_713,negated_conjecture,(epred31_0|epred15_0|~hskp7), inference(spm,[status(thm)],[c_0_685, c_0_686])).
cnf(c_0_714,negated_conjecture,(hskp6|hskp7|hskp17), inference(sr,[status(thm)],[c_0_687, c_0_622])).
cnf(c_0_715,negated_conjecture,(epred39_0|epred20_0|hskp10|~hskp5), inference(spm,[status(thm)],[c_0_688, c_0_689])).
cnf(c_0_716,negated_conjecture,(epred30_0|c1_1(X1)|hskp19|~c2_1(X1)|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_690])).
cnf(c_0_717,negated_conjecture,(c3_1(X1)|hskp6|hskp9|hskp11|~c0_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_613]), c_0_652]), ['final']).
fof(c_0_718, plain, ((~epred3_0<=>![X2]:((((hskp26|c0_1(X2))|c2_1(X2))|c1_1(X2))|~ndr1_0))), introduced(definition)).
cnf(c_0_719,negated_conjecture,(c0_1(a1260)|hskp7|hskp26|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_92]), c_0_60])).
cnf(c_0_720,negated_conjecture,(epred20_0|epred13_0|hskp10|~hskp7), inference(spm,[status(thm)],[c_0_350, c_0_691])).
cnf(c_0_721,negated_conjecture,(epred41_0|epred23_0|epred19_0|epred12_0), inference(spm,[status(thm)],[c_0_692, c_0_693])).
cnf(c_0_722,negated_conjecture,(epred24_0|epred23_0|epred12_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_693]), c_0_466])).
cnf(c_0_723,negated_conjecture,(epred24_0|epred21_0|epred18_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_694])).
cnf(c_0_724,negated_conjecture,(epred21_0|epred19_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_695]), c_0_696])).
cnf(c_0_725,negated_conjecture,(epred42_0|epred23_0|epred17_0), inference(spm,[status(thm)],[c_0_697, c_0_698])).
cnf(c_0_726,negated_conjecture,(epred42_0|epred24_0|epred17_0), inference(spm,[status(thm)],[c_0_305, c_0_625])).
cnf(c_0_727,negated_conjecture,(epred45_0|epred24_0|epred18_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_699])).
cnf(c_0_728,negated_conjecture,(epred45_0|epred19_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_700]), c_0_701])).
cnf(c_0_729,negated_conjecture,(epred39_0|epred18_0|hskp10|~hskp5), inference(spm,[status(thm)],[c_0_688, c_0_600])).
cnf(c_0_730,negated_conjecture,(hskp7|hskp13|hskp28|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_601]), c_0_622]), c_0_95]), c_0_103])).
cnf(c_0_731,negated_conjecture,(epred21_0|epred18_0|hskp10|~hskp17), inference(spm,[status(thm)],[c_0_483, c_0_702])).
cnf(c_0_732,negated_conjecture,(epred38_0|epred21_0|hskp10|hskp17), inference(spm,[status(thm)],[c_0_703, c_0_704])).
cnf(c_0_733,negated_conjecture,(epred38_0|~c0_1(a1257)|~hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_614, c_0_25]), c_0_27])).
cnf(c_0_734,negated_conjecture,(epred19_0|c0_1(a1257)|~hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_25]), c_0_26])).
cnf(c_0_735,negated_conjecture,(epred38_0|hskp4|hskp6|hskp10|~hskp20), inference(spm,[status(thm)],[c_0_633, c_0_705])).
cnf(c_0_736,negated_conjecture,(epred38_0|hskp6|hskp10|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_706, c_0_707]), c_0_708])).
cnf(c_0_737,negated_conjecture,(~epred37_0|~epred17_0|~epred4_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_138, c_0_581]), c_0_596]), c_0_357])).
cnf(c_0_738,negated_conjecture,(epred31_0|epred24_0|epred14_0|epred4_0), inference(spm,[status(thm)],[c_0_709, c_0_710])).
cnf(c_0_739,negated_conjecture,(epred37_0|epred24_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_711]), c_0_712])).
cnf(c_0_740,negated_conjecture,(epred31_0|epred15_0|hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_713, c_0_714]), c_0_709])).
cnf(c_0_741,negated_conjecture,(epred39_0|epred20_0|hskp10|hskp21), inference(spm,[status(thm)],[c_0_715, c_0_108])).
cnf(c_0_742,negated_conjecture,(c1_1(X1)|c2_1(X1)|c0_1(X1)|hskp4|hskp5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_139, c_0_153])])).
cnf(c_0_743,negated_conjecture,(hskp5|hskp17|hskp18|hskp19|~hskp10), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_173]), c_0_80]), c_0_81])).
cnf(c_0_744,negated_conjecture,(hskp6|hskp9|~c0_1(a1283)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_28]), c_0_29])).
cnf(c_0_745,negated_conjecture,(c0_1(a1283)|hskp6|hskp29|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_634, c_0_28]), c_0_29])).
cnf(c_0_746,negated_conjecture,(epred30_0|c1_1(X1)|hskp19|~c2_1(X1)|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_716, c_0_153])])).
cnf(c_0_747,negated_conjecture,(epred43_0|~hskp9), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_66]), c_0_67]), c_0_68])).
cnf(c_0_748,negated_conjecture,(hskp6|hskp9|hskp11|~hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_717]), c_0_41])).
fof(c_0_749, plain, ((~epred5_0<=>![X1]:((((hskp13|c0_1(X1))|c3_1(X1))|~ndr1_0)|~c2_1(X1)))), introduced(definition)).
cnf(c_0_750,negated_conjecture,(epred3_0|c1_1(X1)|c2_1(X1)|c0_1(X1)|hskp26|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_718])).
cnf(c_0_751,negated_conjecture,(epred36_0|hskp7|hskp21|hskp26|~hskp28), inference(spm,[status(thm)],[c_0_564, c_0_384])).
cnf(c_0_752,negated_conjecture,(c0_1(X1)|hskp7|hskp15|hskp26|~c1_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_207]), c_0_235])).
cnf(c_0_753,negated_conjecture,(hskp7|hskp10|hskp26|~hskp6), inference(spm,[status(thm)],[c_0_197, c_0_719])).
cnf(c_0_754,negated_conjecture,(epred23_0|epred20_0|epred13_0|hskp10), inference(spm,[status(thm)],[c_0_720, c_0_378])).
cnf(c_0_755,negated_conjecture,(epred41_0|epred23_0|epred19_0|~epred13_0), inference(spm,[status(thm)],[c_0_440, c_0_721])).
cnf(c_0_756,negated_conjecture,(epred24_0|epred23_0|~epred13_0), inference(spm,[status(thm)],[c_0_440, c_0_722])).
cnf(c_0_757,negated_conjecture,(epred24_0|epred21_0|~epred29_0|~epred19_0), inference(spm,[status(thm)],[c_0_338, c_0_723])).
cnf(c_0_758,negated_conjecture,(epred21_0|epred19_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_724])).
cnf(c_0_759,negated_conjecture,(epred42_0|epred23_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_725])).
cnf(c_0_760,negated_conjecture,(epred42_0|epred24_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_726])).
cnf(c_0_761,negated_conjecture,(epred45_0|epred24_0|~epred29_0|~epred19_0), inference(spm,[status(thm)],[c_0_338, c_0_727])).
cnf(c_0_762,negated_conjecture,(epred45_0|epred19_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_728])).
cnf(c_0_763,negated_conjecture,(epred39_0|epred18_0|hskp10|hskp21), inference(spm,[status(thm)],[c_0_729, c_0_108])).
cnf(c_0_764,negated_conjecture,(epred39_0|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_662]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_765,negated_conjecture,(epred13_0|hskp28|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_730]), c_0_383])).
cnf(c_0_766,negated_conjecture,(epred20_0|c0_1(a1257)|hskp10|~hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_663, c_0_26]), c_0_27])).
cnf(c_0_767,negated_conjecture,(epred38_0|epred21_0|epred18_0|hskp10), inference(spm,[status(thm)],[c_0_731, c_0_732])).
cnf(c_0_768,negated_conjecture,(epred38_0|epred19_0|~hskp4), inference(spm,[status(thm)],[c_0_733, c_0_734])).
cnf(c_0_769,negated_conjecture,(epred38_0|hskp4|hskp6|hskp10), inference(spm,[status(thm)],[c_0_735, c_0_736])).
cnf(c_0_770,negated_conjecture,(~epred15_0|~epred14_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_140, c_0_580]), c_0_355])).
cnf(c_0_771,negated_conjecture,(epred31_0|epred24_0|epred14_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_737, c_0_738]), c_0_739])).
cnf(c_0_772,negated_conjecture,(epred31_0|epred24_0|epred15_0), inference(spm,[status(thm)],[c_0_681, c_0_740])).
cnf(c_0_773,negated_conjecture,(epred39_0|epred20_0|epred17_0|hskp10), inference(spm,[status(thm)],[c_0_424, c_0_741])).
cnf(c_0_774,negated_conjecture,(c1_1(a1280)|hskp2|hskp4|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_18]), c_0_109])).
cnf(c_0_775,negated_conjecture,(hskp4|hskp5|~hskp18), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_742]), c_0_54]), c_0_55])).
cnf(c_0_776,negated_conjecture,(hskp5|hskp17|hskp18|hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_610]), c_0_743])).
cnf(c_0_777,negated_conjecture,(hskp6|hskp9|hskp29|~hskp19), inference(spm,[status(thm)],[c_0_744, c_0_745])).
cnf(c_0_778,negated_conjecture,(epred30_0|c1_1(a1252)|hskp19|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_746, c_0_64]), c_0_46])).
cnf(c_0_779,negated_conjecture,(epred43_0|hskp6|hskp11|~hskp16), inference(spm,[status(thm)],[c_0_747, c_0_748])).
cnf(c_0_780,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp0|hskp7|hskp11|hskp12|~c0_1(X1)), inference(spm,[status(thm)],[c_0_157, c_0_613])).
cnf(c_0_781,negated_conjecture,(epred5_0|c3_1(X1)|c0_1(X1)|hskp13|~c2_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_749])).
cnf(c_0_782,negated_conjecture,(epred3_0|c1_1(X1)|c2_1(X1)|c0_1(X1)|hskp26), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_750, c_0_153])]), ['final']).
cnf(c_0_783,negated_conjecture,(epred36_0|hskp6|hskp7|hskp21|hskp26), inference(spm,[status(thm)],[c_0_751, c_0_414])).
cnf(c_0_784,negated_conjecture,(hskp7|hskp26|~hskp6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_752]), c_0_60]), c_0_753]), c_0_235])).
cnf(c_0_785,negated_conjecture,(epred41_0|epred23_0|epred20_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_754]), c_0_755])).
cnf(c_0_786,negated_conjecture,(epred24_0|epred23_0|epred20_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_754]), c_0_756])).
cnf(c_0_787,negated_conjecture,(epred24_0|epred21_0|~epred29_0), inference(spm,[status(thm)],[c_0_757, c_0_758])).
cnf(c_0_788,negated_conjecture,(epred42_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_759]), c_0_760]), ['final']).
cnf(c_0_789,negated_conjecture,(epred45_0|epred24_0|~epred29_0), inference(spm,[status(thm)],[c_0_761, c_0_762])).
cnf(c_0_790,negated_conjecture,(epred39_0|epred18_0|epred17_0|hskp10), inference(spm,[status(thm)],[c_0_424, c_0_763])).
cnf(c_0_791,negated_conjecture,(epred39_0|epred17_0|hskp7), inference(spm,[status(thm)],[c_0_764, c_0_643])).
cnf(c_0_792,negated_conjecture,(epred17_0|epred13_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_765]), c_0_585])).
cnf(c_0_793,negated_conjecture,(epred38_0|epred20_0|hskp10|~hskp4), inference(spm,[status(thm)],[c_0_733, c_0_766])).
cnf(c_0_794,negated_conjecture,(epred38_0|epred24_0|epred21_0|epred18_0), inference(spm,[status(thm)],[c_0_305, c_0_767])).
cnf(c_0_795,negated_conjecture,(epred38_0|epred19_0|hskp6|hskp10), inference(spm,[status(thm)],[c_0_768, c_0_769])).
cnf(c_0_796,negated_conjecture,(epred31_0|epred24_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_770, c_0_771]), c_0_772])).
cnf(c_0_797,negated_conjecture,(epred39_0|epred24_0|epred20_0|epred17_0), inference(spm,[status(thm)],[c_0_305, c_0_773])).
cnf(c_0_798,negated_conjecture,(hskp2|hskp4|hskp20|~hskp29), inference(spm,[status(thm)],[c_0_631, c_0_774])).
cnf(c_0_799,negated_conjecture,(c3_1(a1264)|hskp6|hskp29|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_634, c_0_79]), c_0_81])).
cnf(c_0_800,negated_conjecture,(epred42_0|hskp4|~hskp18), inference(spm,[status(thm)],[c_0_527, c_0_775])).
cnf(c_0_801,negated_conjecture,(epred42_0|hskp17|hskp18|hskp19), inference(spm,[status(thm)],[c_0_527, c_0_776])).
cnf(c_0_802,negated_conjecture,(c3_1(a1255)|hskp6|hskp9|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_12]), c_0_124]), ['final']).
cnf(c_0_803,negated_conjecture,(epred43_0|hskp2|hskp20|~hskp29), inference(spm,[status(thm)],[c_0_631, c_0_427])).
cnf(c_0_804,negated_conjecture,(epred43_0|hskp6|hskp29|~hskp19), inference(spm,[status(thm)],[c_0_747, c_0_777])).
cnf(c_0_805,negated_conjecture,(epred30_0|hskp10|hskp19|~hskp0), inference(spm,[status(thm)],[c_0_186, c_0_778])).
cnf(c_0_806,negated_conjecture,(epred43_0|hskp6|hskp10|~hskp16), inference(spm,[status(thm)],[c_0_164, c_0_779])).
cnf(c_0_807,negated_conjecture,(hskp2|hskp20|~c0_1(a1259)|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_75]), c_0_76]), ['final']).
cnf(c_0_808,negated_conjecture,(hskp7|hskp12|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_780]), c_0_124]), c_0_174]), c_0_174]), c_0_13]), c_0_177])).
cnf(c_0_809,negated_conjecture,(hskp7|hskp12|hskp17|hskp18|~hskp24), inference(spm,[status(thm)],[c_0_295, c_0_181])).
cnf(c_0_810,negated_conjecture,(c1_1(X1)|hskp7|hskp13|hskp19|hskp28|~c3_1(X1)), inference(spm,[status(thm)],[c_0_156, c_0_237])).
fof(c_0_811, plain, ((~epred2_0<=>![X2]:((c0_1(X2)|c2_1(X2))|c3_1(X2)))), introduced(definition)).
cnf(c_0_812,negated_conjecture,(epred5_0|c3_1(X1)|c0_1(X1)|hskp13|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_781, c_0_153])])).
cnf(c_0_813,negated_conjecture,(epred12_0|c0_1(a1277)|hskp14|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_345]), c_0_114])).
cnf(c_0_814,negated_conjecture,(epred18_0|hskp6|hskp10|hskp17), inference(spm,[status(thm)],[c_0_283, c_0_714])).
cnf(c_0_815,negated_conjecture,(epred3_0|c1_1(a1261)|hskp26|~hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_782]), c_0_72])).
cnf(c_0_816,negated_conjecture,(epred24_0|epred13_0|~hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_658]), c_0_106])).
cnf(c_0_817,negated_conjecture,(hskp7|hskp21|hskp26), inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_783, c_0_585]), c_0_784])).
cnf(c_0_818,negated_conjecture,(~epred21_0|~epred20_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_141, c_0_618]), c_0_308])).
cnf(c_0_819,negated_conjecture,(epred41_0|epred23_0|epred20_0|~epred42_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_785]), c_0_786])).
cnf(c_0_820,negated_conjecture,(epred21_0|~epred29_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_758]), c_0_787]), c_0_788]), ['final']).
cnf(c_0_821,negated_conjecture,(epred45_0|~epred29_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_762]), c_0_789]), c_0_788]), ['final']).
cnf(c_0_822,negated_conjecture,(epred39_0|epred24_0|epred18_0|epred17_0), inference(spm,[status(thm)],[c_0_305, c_0_790])).
cnf(c_0_823,negated_conjecture,(epred39_0|epred19_0|epred17_0|epred12_0), inference(spm,[status(thm)],[c_0_410, c_0_791])).
cnf(c_0_824,negated_conjecture,(epred39_0|epred17_0|epred13_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_666, c_0_791]), c_0_585])).
cnf(c_0_825,negated_conjecture,(c1_1(X1)|c2_1(X1)|hskp12|hskp13|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_142, c_0_153])]), ['final']).
cnf(c_0_826,negated_conjecture,(epred17_0|epred13_0|hskp7), inference(spm,[status(thm)],[c_0_792, c_0_643])).
cnf(c_0_827,negated_conjecture,(epred38_0|epred20_0|hskp6|hskp10), inference(spm,[status(thm)],[c_0_793, c_0_769])).
cnf(c_0_828,negated_conjecture,(epred38_0|epred24_0|epred21_0|~epred19_0), inference(spm,[status(thm)],[c_0_338, c_0_794])).
cnf(c_0_829,negated_conjecture,(epred38_0|epred24_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_795]), c_0_305])).
cnf(c_0_830,negated_conjecture,(epred39_0|epred31_0|epred24_0|epred20_0), inference(spm,[status(thm)],[c_0_796, c_0_797])).
cnf(c_0_831,negated_conjecture,(epred31_0|epred24_0|epred21_0|epred18_0), inference(spm,[status(thm)],[c_0_796, c_0_694])).
cnf(c_0_832,negated_conjecture,(epred31_0|epred24_0|epred21_0|epred19_0), inference(spm,[status(thm)],[c_0_796, c_0_724])).
cnf(c_0_833,negated_conjecture,(hskp2|hskp6|hskp9|hskp20|~hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_188]), c_0_49]), c_0_48])).
cnf(c_0_834,negated_conjecture,(c1_1(a1272)|hskp4|hskp5|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_742]), c_0_57])).
cnf(c_0_835,negated_conjecture,(hskp2|hskp4|hskp6|hskp20|~hskp24), inference(spm,[status(thm)],[c_0_798, c_0_656])).
cnf(c_0_836,negated_conjecture,(hskp5|hskp6|hskp19|hskp29|~hskp10), inference(spm,[status(thm)],[c_0_491, c_0_799])).
cnf(c_0_837,negated_conjecture,(hskp5|hskp16|hskp19|hskp20), inference(spm,[status(thm)],[c_0_226, c_0_1])).
cnf(c_0_838,negated_conjecture,(epred42_0|hskp4|hskp17|hskp19), inference(spm,[status(thm)],[c_0_800, c_0_801])).
cnf(c_0_839,negated_conjecture,(epred13_0|c0_1(a1277)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_315]), c_0_115])).
cnf(c_0_840,negated_conjecture,(epred43_0|hskp6|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_802]), c_0_124]), c_0_747]), c_0_13])).
cnf(c_0_841,negated_conjecture,(epred43_0|hskp2|hskp6|hskp20|~hskp19), inference(spm,[status(thm)],[c_0_803, c_0_804])).
cnf(c_0_842,negated_conjecture,(epred30_0|hskp19|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_746, c_0_12]), c_0_124]), c_0_13])).
cnf(c_0_843,negated_conjecture,(epred30_0|hskp10|hskp19|hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_805, c_0_212]), c_0_84])).
cnf(c_0_844,negated_conjecture,(epred43_0|hskp6|hskp10|hskp17), inference(spm,[status(thm)],[c_0_806, c_0_83])).
cnf(c_0_845,negated_conjecture,(epred19_0|hskp2|hskp20|~hskp28), inference(spm,[status(thm)],[c_0_807, c_0_349])).
cnf(c_0_846,negated_conjecture,(c1_1(X1)|c3_1(X1)|c0_1(X1)|hskp3|hskp7|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_172]), c_0_808])).
cnf(c_0_847,negated_conjecture,(hskp13|hskp23|hskp24|hskp28|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_170]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_848,negated_conjecture,(hskp7|hskp12|hskp17|hskp18|~hskp2), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_809, c_0_571]), c_0_177])).
cnf(c_0_849,negated_conjecture,(hskp7|hskp12|hskp17|hskp18|hskp27), inference(spm,[status(thm)],[c_0_809, c_0_100])).
cnf(c_0_850,negated_conjecture,(hskp7|hskp13|hskp19|hskp28|~hskp3), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_810, c_0_50]), c_0_52])).
cnf(c_0_851,negated_conjecture,(epred2_0|c3_1(X1)|c2_1(X1)|c0_1(X1)), inference(split_equiv,[status(thm)],[c_0_811])).
cnf(c_0_852,negated_conjecture,(epred4_0|~c0_1(a1254)|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_615]), c_0_88]), ['final']).
cnf(c_0_853,negated_conjecture,(epred13_0|c0_1(a1254)|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_315]), c_0_89])).
cnf(c_0_854,negated_conjecture,(epred5_0|hskp13|~hskp24), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_812, c_0_36]), c_0_35]), c_0_34])).
cnf(c_0_855,negated_conjecture,(epred21_0|epred12_0|hskp14|~hskp17), inference(spm,[status(thm)],[c_0_483, c_0_813])).
cnf(c_0_856,negated_conjecture,(epred19_0|epred12_0|hskp6|hskp17), inference(spm,[status(thm)],[c_0_410, c_0_714])).
cnf(c_0_857,negated_conjecture,(epred24_0|epred18_0|hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_681, c_0_814]), c_0_305])).
cnf(c_0_858,negated_conjecture,(epred41_0|hskp21|~hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_331]), c_0_585]), c_0_49]), c_0_48]), ['final']).
cnf(c_0_859,negated_conjecture,(epred13_0|epred3_0|hskp26|~hskp7), inference(spm,[status(thm)],[c_0_350, c_0_815])).
cnf(c_0_860,negated_conjecture,(epred24_0|epred13_0|hskp7|hskp26), inference(spm,[status(thm)],[c_0_816, c_0_817])).
cnf(c_0_861,negated_conjecture,(epred41_0|epred23_0|~epred42_0|~epred21_0), inference(spm,[status(thm)],[c_0_818, c_0_819])).
cnf(c_0_862,negated_conjecture,(~epred41_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_820]), c_0_821]), ['final']).
cnf(c_0_863,negated_conjecture,(epred39_0|epred24_0|epred18_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_822])).
cnf(c_0_864,negated_conjecture,(epred39_0|epred19_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_823]), c_0_824])).
cnf(c_0_865,negated_conjecture,(hskp12|hskp13|~c0_1(a1258)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_825]), c_0_10]), ['final']).
cnf(c_0_866,negated_conjecture,(~epred39_0|~epred38_0|~epred13_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_143, c_0_579]), c_0_639]), c_0_291])).
cnf(c_0_867,negated_conjecture,(epred17_0|epred13_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_666, c_0_826]), c_0_585])).
cnf(c_0_868,negated_conjecture,(epred38_0|epred24_0|epred20_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_827]), c_0_305])).
cnf(c_0_869,negated_conjecture,(epred38_0|epred24_0|epred21_0), inference(spm,[status(thm)],[c_0_828, c_0_829])).
cnf(c_0_870,negated_conjecture,(epred39_0|epred31_0|epred24_0|~epred21_0), inference(spm,[status(thm)],[c_0_818, c_0_830])).
cnf(c_0_871,negated_conjecture,(epred31_0|epred24_0|epred21_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_831]), c_0_832])).
cnf(c_0_872,negated_conjecture,(hskp2|hskp4|hskp6|hskp20|~hskp26), inference(spm,[status(thm)],[c_0_202, c_0_833])).
cnf(c_0_873,negated_conjecture,(hskp4|hskp5|hskp17|hskp18|~hskp15), inference(spm,[status(thm)],[c_0_191, c_0_834])).
cnf(c_0_874,negated_conjecture,(hskp2|hskp4|hskp6|hskp15|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_549]), c_0_95]), c_0_103])).
cnf(c_0_875,negated_conjecture,(hskp2|hskp4|hskp6|hskp17|hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_835, c_0_196]), c_0_83]), c_0_590])).
cnf(c_0_876,negated_conjecture,(hskp4|hskp6|hskp7|hskp10|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_175]), c_0_261]), c_0_262])).
cnf(c_0_877,negated_conjecture,(hskp5|hskp6|hskp7|hskp19|~hskp10), inference(spm,[status(thm)],[c_0_178, c_0_836])).
cnf(c_0_878,negated_conjecture,(epred42_0|hskp16|hskp19|hskp20), inference(spm,[status(thm)],[c_0_527, c_0_837])).
cnf(c_0_879,negated_conjecture,(epred42_0|hskp7|hskp17|hskp19), inference(spm,[status(thm)],[c_0_176, c_0_838])).
cnf(c_0_880,negated_conjecture,(epred41_0|epred40_0|~hskp0), inference(spm,[status(thm)],[c_0_286, c_0_325])).
cnf(c_0_881,negated_conjecture,(epred4_0|c3_1(a1287)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_615]), c_0_95]), ['final']).
cnf(c_0_882,negated_conjecture,(epred43_0|hskp2|hskp6|hskp20|~hskp26), inference(spm,[status(thm)],[c_0_747, c_0_833])).
cnf(c_0_883,negated_conjecture,(hskp6|hskp9|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_652, c_0_212]), ['final']).
cnf(c_0_884,negated_conjecture,(epred37_0|c0_1(a1283)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_617, c_0_28]), c_0_30])).
cnf(c_0_885,negated_conjecture,(c0_1(a1259)|hskp19|hskp21|~hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_74]), c_0_75])).
cnf(c_0_886,negated_conjecture,(hskp19|hskp21|~c3_1(a1264)|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_79]), c_0_81])).
cnf(c_0_887,negated_conjecture,(epred36_0|epred13_0|hskp21|~hskp17), inference(spm,[status(thm)],[c_0_531, c_0_839])).
cnf(c_0_888,negated_conjecture,(epred43_0|hskp6|hskp20|~hskp19), inference(spm,[status(thm)],[c_0_840, c_0_841])).
cnf(c_0_889,negated_conjecture,(epred30_0|hskp10|hskp19|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_842, c_0_680]), c_0_843])).
cnf(c_0_890,negated_conjecture,(epred43_0|epred24_0|hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_844]), c_0_305])).
cnf(c_0_891,negated_conjecture,(epred19_0|hskp7|hskp12|hskp20|~hskp28), inference(spm,[status(thm)],[c_0_808, c_0_845])).
cnf(c_0_892,negated_conjecture,(c1_1(a1272)|hskp3|hskp7|hskp12|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_846]), c_0_57])).
cnf(c_0_893,negated_conjecture,(hskp13|hskp17|hskp18|hskp24|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_292, c_0_847])).
cnf(c_0_894,negated_conjecture,(hskp7|hskp12|hskp17|hskp18|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_848, c_0_590]), c_0_849])).
cnf(c_0_895,negated_conjecture,(hskp2|hskp7|hskp13|hskp19|hskp28|~hskp18), inference(spm,[status(thm)],[c_0_850, c_0_190])).
cnf(c_0_896,negated_conjecture,(hskp2|hskp20|~c0_1(a1277)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_114]), c_0_115])).
cnf(c_0_897,negated_conjecture,(hskp12|hskp13|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_825]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_898,negated_conjecture,(epred2_0|~hskp15), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_851]), c_0_57]), c_0_56]), ['final']).
cnf(c_0_899,negated_conjecture,(epred13_0|epred4_0|~hskp1), inference(spm,[status(thm)],[c_0_852, c_0_853])).
cnf(c_0_900,negated_conjecture,(epred5_0|hskp13|hskp18|hskp27), inference(spm,[status(thm)],[c_0_854, c_0_100])).
cnf(c_0_901,negated_conjecture,(epred19_0|epred13_0|c0_1(X1)|~c1_1(X1)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_280, c_0_315])).
cnf(c_0_902,negated_conjecture,(epred21_0|epred19_0|epred12_0|~hskp17), inference(spm,[status(thm)],[c_0_379, c_0_855])).
cnf(c_0_903,negated_conjecture,(epred24_0|epred19_0|epred12_0|hskp17), inference(spm,[status(thm)],[c_0_681, c_0_856])).
cnf(c_0_904,negated_conjecture,(epred24_0|epred21_0|epred18_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_857]), c_0_305])).
cnf(c_0_905,negated_conjecture,(epred41_0|epred24_0|epred13_0|~hskp26), inference(spm,[status(thm)],[c_0_816, c_0_858])).
cnf(c_0_906,negated_conjecture,(epred24_0|epred13_0|epred3_0|hskp26), inference(spm,[status(thm)],[c_0_859, c_0_860])).
cnf(c_0_907,negated_conjecture,(epred15_0|epred13_0|~hskp7), inference(spm,[status(thm)],[c_0_350, c_0_686])).
cnf(c_0_908,negated_conjecture,(epred23_0|~epred29_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_861, c_0_820]), c_0_788]), c_0_862]), ['final']).
cnf(c_0_909,negated_conjecture,(epred39_0|epred24_0|~epred29_0|~epred19_0), inference(spm,[status(thm)],[c_0_338, c_0_863])).
cnf(c_0_910,negated_conjecture,(epred39_0|epred19_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_864])).
cnf(c_0_911,negated_conjecture,(hskp11|hskp12|hskp13|~hskp5), inference(spm,[status(thm)],[c_0_865, c_0_160]), ['final']).
cnf(c_0_912,negated_conjecture,(epred17_0|~epred39_0|~epred38_0), inference(spm,[status(thm)],[c_0_866, c_0_867])).
cnf(c_0_913,negated_conjecture,(epred38_0|epred24_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_818, c_0_868]), c_0_869])).
cnf(c_0_914,negated_conjecture,(epred39_0|epred31_0|epred24_0), inference(spm,[status(thm)],[c_0_870, c_0_871])).
cnf(c_0_915,negated_conjecture,(hskp6|hskp7|hskp12|hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_872]), c_0_251]), c_0_176]), c_0_177])).
cnf(c_0_916,negated_conjecture,(hskp2|hskp4|hskp5|hskp6|hskp17), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_873, c_0_874]), c_0_875]), c_0_775])).
cnf(c_0_917,negated_conjecture,(hskp5|hskp6|hskp7|hskp19), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_876]), c_0_877]), c_0_176])).
cnf(c_0_918,negated_conjecture,(epred42_0|epred16_0|hskp19|hskp20), inference(spm,[status(thm)],[c_0_421, c_0_878])).
cnf(c_0_919,negated_conjecture,(epred42_0|epred17_0|hskp17|hskp19), inference(spm,[status(thm)],[c_0_667, c_0_879])).
cnf(c_0_920,negated_conjecture,(epred41_0|hskp5|hskp19|~hskp26), inference(spm,[status(thm)],[c_0_187, c_0_331])).
cnf(c_0_921,negated_conjecture,(epred41_0|epred40_0|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_880, c_0_212])).
cnf(c_0_922,negated_conjecture,(epred43_0|epred4_0|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_881]), c_0_95]), c_0_103])).
cnf(c_0_923,negated_conjecture,(epred43_0|hskp6|hskp20|~hskp26), inference(spm,[status(thm)],[c_0_840, c_0_882])).
cnf(c_0_924,negated_conjecture,(epred43_0|hskp6|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_747, c_0_883])).
cnf(c_0_925,negated_conjecture,(epred37_0|hskp6|hskp9|~hskp19), inference(spm,[status(thm)],[c_0_744, c_0_884])).
cnf(c_0_926,negated_conjecture,(epred42_0|hskp10|hskp19|~hskp26), inference(spm,[status(thm)],[c_0_527, c_0_250])).
cnf(c_0_927,negated_conjecture,(epred41_0|hskp5|hskp19|~hskp10), inference(spm,[status(thm)],[c_0_491, c_0_332])).
cnf(c_0_928,negated_conjecture,(epred41_0|epred37_0|epred19_0|epred6_0), inference(spm,[status(thm)],[c_0_692, c_0_683])).
cnf(c_0_929,negated_conjecture,(epred41_0|epred37_0|epred19_0|epred7_0), inference(spm,[status(thm)],[c_0_692, c_0_684])).
fof(c_0_930, plain, ((~epred26_0<=>![X1]:((((hskp4|c3_1(X1))|c1_1(X1))|~ndr1_0)|~c2_1(X1)))), introduced(definition)).
cnf(c_0_931,negated_conjecture,(epred36_0|hskp19|hskp21|~hskp28), inference(spm,[status(thm)],[c_0_564, c_0_885])).
cnf(c_0_932,negated_conjecture,(hskp17|hskp19|hskp21|hskp27|~hskp10), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_886, c_0_169]), c_0_80]), c_0_81])).
cnf(c_0_933,negated_conjecture,(epred36_0|epred24_0|epred13_0|~hskp17), inference(spm,[status(thm)],[c_0_816, c_0_887])).
cnf(c_0_934,negated_conjecture,(epred43_0|epred39_0|hskp6|~hskp19), inference(spm,[status(thm)],[c_0_764, c_0_888])).
cnf(c_0_935,negated_conjecture,(epred39_0|epred30_0|hskp10|hskp19), inference(spm,[status(thm)],[c_0_764, c_0_889])).
cnf(c_0_936,negated_conjecture,(~epred30_0|~epred24_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_144, c_0_690]), c_0_266])).
cnf(c_0_937,negated_conjecture,(epred43_0|epred31_0|epred24_0), inference(spm,[status(thm)],[c_0_709, c_0_890])).
cnf(c_0_938,negated_conjecture,(~epred31_0|~epred30_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_145, c_0_690]), c_0_638])).
cnf(c_0_939,negated_conjecture,(epred19_0|hskp7|hskp12|hskp16|hskp20), inference(spm,[status(thm)],[c_0_891, c_0_1])).
cnf(c_0_940,negated_conjecture,(hskp3|hskp7|hskp12|hskp17|hskp18|~hskp15), inference(spm,[status(thm)],[c_0_191, c_0_892])).
cnf(c_0_941,negated_conjecture,(hskp7|hskp12|hskp13|hskp17|hskp18|hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_809, c_0_893]), c_0_894])).
cnf(c_0_942,negated_conjecture,(hskp7|hskp12|hskp13|hskp28|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_808, c_0_895]), c_0_177])).
cnf(c_0_943,negated_conjecture,(epred13_0|hskp2|hskp20|~hskp17), inference(spm,[status(thm)],[c_0_896, c_0_839])).
cnf(c_0_944,negated_conjecture,(epred13_0|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_351, c_0_897])).
cnf(c_0_945,negated_conjecture,(epred2_0|hskp10|~hskp13), inference(spm,[status(thm)],[c_0_898, c_0_258])).
cnf(c_0_946,negated_conjecture,(hskp13|hskp28|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_237]), c_0_9]), c_0_10]), ['final']).
cnf(c_0_947,negated_conjecture,(epred36_0|hskp21|~hskp27), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_31]), c_0_32]), c_0_33]), ['final']).
cnf(c_0_948,negated_conjecture,(epred13_0|epred4_0|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_899, c_0_324])).
cnf(c_0_949,negated_conjecture,(epred13_0|epred5_0|hskp18|hskp27), inference(spm,[status(thm)],[c_0_351, c_0_900])).
cnf(c_0_950,negated_conjecture,(epred36_0|epred19_0|hskp21|~hskp28), inference(spm,[status(thm)],[c_0_564, c_0_349])).
cnf(c_0_951,negated_conjecture,(epred19_0|epred13_0|~hskp21), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_901, c_0_105]), c_0_106]), c_0_107])).
cnf(c_0_952,negated_conjecture,(epred24_0|epred21_0|epred19_0|epred12_0), inference(spm,[status(thm)],[c_0_902, c_0_903])).
cnf(c_0_953,negated_conjecture,(epred24_0|epred21_0|~epred19_0), inference(spm,[status(thm)],[c_0_338, c_0_904])).
cnf(c_0_954,negated_conjecture,(~epred4_0|~epred3_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_146, c_0_718]), c_0_596])).
cnf(c_0_955,negated_conjecture,(epred41_0|epred24_0|epred13_0|epred3_0), inference(spm,[status(thm)],[c_0_905, c_0_906])).
cnf(c_0_956,negated_conjecture,(epred15_0|epred13_0|hskp6|hskp17), inference(spm,[status(thm)],[c_0_907, c_0_714])).
cnf(c_0_957,negated_conjecture,(epred29_0|epred19_0|hskp20|~hskp10), inference(spm,[status(thm)],[c_0_335, c_0_448])).
cnf(c_0_958,negated_conjecture,(~epred29_0|~epred24_0), inference(spm,[status(thm)],[c_0_513, c_0_908])).
cnf(c_0_959,negated_conjecture,(epred39_0|epred24_0|~epred29_0), inference(spm,[status(thm)],[c_0_909, c_0_910])).
cnf(c_0_960,negated_conjecture,(hskp10|hskp12|hskp13|~hskp5), inference(spm,[status(thm)],[c_0_164, c_0_911])).
cnf(c_0_961,negated_conjecture,(epred31_0|epred24_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_796, c_0_912]), c_0_913]), c_0_914])).
cnf(c_0_962,negated_conjecture,(epred39_0|hskp6|hskp7|hskp12), inference(spm,[status(thm)],[c_0_764, c_0_915])).
cnf(c_0_963,negated_conjecture,(hskp5|hskp6|hskp7|hskp12), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_916]), c_0_917]), c_0_176]), c_0_177])).
cnf(c_0_964,negated_conjecture,(epred42_0|epred21_0|epred16_0|hskp19), inference(spm,[status(thm)],[c_0_454, c_0_918])).
cnf(c_0_965,negated_conjecture,(epred42_0|epred21_0|epred17_0|hskp19), inference(spm,[status(thm)],[c_0_508, c_0_919])).
cnf(c_0_966,negated_conjecture,(epred41_0|epred40_0|epred19_0|epred6_0), inference(spm,[status(thm)],[c_0_692, c_0_370])).
cnf(c_0_967,negated_conjecture,(epred41_0|epred40_0|epred19_0|epred7_0), inference(spm,[status(thm)],[c_0_692, c_0_372])).
cnf(c_0_968,negated_conjecture,(epred42_0|epred41_0|hskp19|~hskp26), inference(spm,[status(thm)],[c_0_527, c_0_920])).
cnf(c_0_969,negated_conjecture,(epred41_0|epred40_0|hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_921, c_0_84]), c_0_249])).
cnf(c_0_970,negated_conjecture,(epred43_0|~c0_1(a1258)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_9]), c_0_10])).
cnf(c_0_971,negated_conjecture,(epred43_0|epred4_0|hskp6|~hskp26), inference(spm,[status(thm)],[c_0_922, c_0_923])).
cnf(c_0_972,negated_conjecture,(epred43_0|hskp6|hskp19|hskp26), inference(spm,[status(thm)],[c_0_924, c_0_84])).
cnf(c_0_973,negated_conjecture,(epred43_0|epred4_0|hskp6|~hskp19), inference(spm,[status(thm)],[c_0_922, c_0_888])).
cnf(c_0_974,negated_conjecture,(epred43_0|epred37_0|hskp6|~hskp19), inference(spm,[status(thm)],[c_0_747, c_0_925])).
cnf(c_0_975,negated_conjecture,(epred42_0|hskp6|hskp10|hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_527, c_0_243]), c_0_926])).
cnf(c_0_976,negated_conjecture,(epred42_0|epred41_0|hskp19|~hskp10), inference(spm,[status(thm)],[c_0_527, c_0_927])).
cnf(c_0_977,negated_conjecture,(epred41_0|epred37_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_928]), c_0_929])).
cnf(c_0_978,negated_conjecture,(epred26_0|c1_1(X1)|c3_1(X1)|hskp4|~c2_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_930])).
cnf(c_0_979,negated_conjecture,(hskp17|hskp19|hskp21|hskp27), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_931, c_0_394]), c_0_585]), c_0_932])).
cnf(c_0_980,negated_conjecture,(epred43_0|epred36_0|epred24_0|epred13_0), inference(spm,[status(thm)],[c_0_933, c_0_890])).
cnf(c_0_981,negated_conjecture,(epred43_0|epred39_0|epred24_0|~hskp19), inference(spm,[status(thm)],[c_0_681, c_0_934])).
cnf(c_0_982,negated_conjecture,(epred39_0|epred30_0|epred24_0|hskp19), inference(spm,[status(thm)],[c_0_305, c_0_935])).
cnf(c_0_983,negated_conjecture,(epred43_0|~epred30_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_937]), c_0_938])).
cnf(c_0_984,negated_conjecture,(epred19_0|hskp7|hskp12|hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_620, c_0_939]), c_0_622])).
cnf(c_0_985,negated_conjecture,(hskp7|hskp12|hskp13|hskp17|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_850, c_0_940]), c_0_941]), c_0_942]), c_0_177])).
cnf(c_0_986,negated_conjecture,(epred13_0|hskp7|hskp12|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_808, c_0_943]), c_0_944])).
cnf(c_0_987,negated_conjecture,(epred36_0|epred24_0|hskp21|~hskp28), inference(spm,[status(thm)],[c_0_564, c_0_415])).
cnf(c_0_988,negated_conjecture,(epred2_0|hskp10|hskp28|~hskp5), inference(spm,[status(thm)],[c_0_945, c_0_946])).
cnf(c_0_989,negated_conjecture,(epred36_0|epred24_0|epred13_0|~hskp27), inference(spm,[status(thm)],[c_0_816, c_0_947])).
cnf(c_0_990,negated_conjecture,(epred13_0|epred5_0|epred4_0|hskp27), inference(spm,[status(thm)],[c_0_948, c_0_949])).
cnf(c_0_991,negated_conjecture,(epred41_0|epred19_0|epred12_0|hskp17), inference(spm,[status(thm)],[c_0_363, c_0_856])).
cnf(c_0_992,negated_conjecture,(epred19_0|epred13_0|~hskp7), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_950, c_0_383]), c_0_585]), c_0_951])).
cnf(c_0_993,negated_conjecture,(epred19_0|epred13_0|hskp7|hskp26), inference(spm,[status(thm)],[c_0_951, c_0_817])).
cnf(c_0_994,negated_conjecture,(epred24_0|epred21_0|epred12_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_902, c_0_857]), c_0_952]), c_0_953])).
cnf(c_0_995,negated_conjecture,(epred41_0|epred24_0|epred13_0|~epred4_0), inference(spm,[status(thm)],[c_0_954, c_0_955])).
cnf(c_0_996,negated_conjecture,(epred24_0|epred14_0|epred13_0|epred4_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_933, c_0_710]), c_0_585])).
cnf(c_0_997,negated_conjecture,(epred24_0|epred15_0|epred13_0|hskp17), inference(spm,[status(thm)],[c_0_681, c_0_956])).
cnf(c_0_998,negated_conjecture,(epred45_0|epred29_0|epred19_0|~hskp10), inference(spm,[status(thm)],[c_0_408, c_0_957])).
cnf(c_0_999,negated_conjecture,(epred38_0|~epred29_0), inference(spm,[status(thm)],[c_0_958, c_0_913])).
cnf(c_0_1000,negated_conjecture,(epred39_0|~epred29_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_910]), c_0_959]), c_0_788]), ['final']).
cnf(c_0_1001,negated_conjecture,(epred13_0|hskp10|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_351, c_0_960])).
cnf(c_0_1002,negated_conjecture,(hskp7|hskp13|hskp19|hskp28|~hskp9), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_810, c_0_66]), c_0_68])).
cnf(c_0_1003,negated_conjecture,(epred29_0|hskp20|~hskp6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_92]), c_0_60]), c_0_59]), ['final']).
cnf(c_0_1004,negated_conjecture,(~epred30_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_961]), c_0_938]), ['final']).
cnf(c_0_1005,negated_conjecture,(epred40_0|epred39_0|hskp6|hskp7), inference(spm,[status(thm)],[c_0_268, c_0_962])).
cnf(c_0_1006,negated_conjecture,(epred42_0|hskp6|hskp7|hskp12), inference(spm,[status(thm)],[c_0_527, c_0_963])).
cnf(c_0_1007,negated_conjecture,(epred45_0|epred42_0|epred16_0|hskp19), inference(spm,[status(thm)],[c_0_408, c_0_918])).
cnf(c_0_1008,negated_conjecture,(epred42_0|epred40_0|epred21_0|epred16_0), inference(spm,[status(thm)],[c_0_249, c_0_964])).
cnf(c_0_1009,negated_conjecture,(epred42_0|epred40_0|epred21_0|epred17_0), inference(spm,[status(thm)],[c_0_249, c_0_965])).
cnf(c_0_1010,negated_conjecture,(epred41_0|epred40_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_966]), c_0_967])).
cnf(c_0_1011,negated_conjecture,(epred42_0|epred41_0|epred40_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_968, c_0_969]), c_0_249])).
cnf(c_0_1012,negated_conjecture,(epred43_0|epred18_0|hskp10|~hskp5), inference(spm,[status(thm)],[c_0_970, c_0_600])).
cnf(c_0_1013,negated_conjecture,(epred43_0|epred41_0|epred4_0|~hskp26), inference(spm,[status(thm)],[c_0_363, c_0_971])).
cnf(c_0_1014,negated_conjecture,(epred43_0|epred41_0|hskp19|hskp26), inference(spm,[status(thm)],[c_0_363, c_0_972])).
cnf(c_0_1015,negated_conjecture,(epred43_0|epred41_0|epred4_0|~hskp19), inference(spm,[status(thm)],[c_0_363, c_0_973])).
cnf(c_0_1016,negated_conjecture,(epred43_0|epred41_0|epred37_0|~hskp19), inference(spm,[status(thm)],[c_0_363, c_0_974])).
cnf(c_0_1017,negated_conjecture,(epred42_0|epred41_0|hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_975]), c_0_976])).
cnf(c_0_1018,negated_conjecture,(epred41_0|epred37_0|~epred42_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_977]), c_0_739])).
cnf(c_0_1019,negated_conjecture,(epred43_0|epred20_0|hskp10|~hskp5), inference(spm,[status(thm)],[c_0_970, c_0_689])).
cnf(c_0_1020,negated_conjecture,(epred22_0|c1_1(X1)|hskp9|~c3_1(X1)|~c2_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_501])).
cnf(c_0_1021,negated_conjecture,(epred43_0|~c0_1(a1257)|~hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_26]), c_0_27])).
cnf(c_0_1022,negated_conjecture,(epred15_0|c0_1(a1257)|~hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_26]), c_0_27])).
cnf(c_0_1023,negated_conjecture,(epred26_0|c1_1(X1)|c3_1(X1)|hskp4|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_978, c_0_153])])).
cnf(c_0_1024,negated_conjecture,(epred36_0|epred17_0|~hskp27), inference(spm,[status(thm)],[c_0_424, c_0_947])).
cnf(c_0_1025,negated_conjecture,(epred17_0|hskp17|hskp19|hskp27), inference(spm,[status(thm)],[c_0_424, c_0_979])).
cnf(c_0_1026,negated_conjecture,(epred43_0|epred24_0|epred13_0), inference(sr,[status(thm)],[c_0_980, c_0_585])).
cnf(c_0_1027,negated_conjecture,(epred43_0|epred39_0|epred24_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_981, c_0_982]), c_0_983])).
cnf(c_0_1028,negated_conjecture,(epred39_0|epred37_0|hskp6|hskp7), inference(spm,[status(thm)],[c_0_637, c_0_962])).
cnf(c_0_1029,negated_conjecture,(epred21_0|epred19_0|hskp7|hskp12), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_642, c_0_984]), c_0_622])).
cnf(c_0_1030,negated_conjecture,(epred13_0|hskp12|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_985]), c_0_986]), c_0_383])).
cnf(c_0_1031,negated_conjecture,(epred24_0|epred2_0|hskp21), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_987, c_0_988]), c_0_585]), c_0_108]), c_0_305])).
cnf(c_0_1032,negated_conjecture,(epred24_0|epred13_0|epred5_0|epred4_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_989, c_0_990]), c_0_585])).
cnf(c_0_1033,negated_conjecture,(epred41_0|epred21_0|epred19_0|epred12_0), inference(spm,[status(thm)],[c_0_902, c_0_991])).
cnf(c_0_1034,negated_conjecture,(epred41_0|epred19_0|epred13_0|~hskp26), inference(spm,[status(thm)],[c_0_951, c_0_858])).
cnf(c_0_1035,negated_conjecture,(epred19_0|epred13_0|hskp26), inference(spm,[status(thm)],[c_0_992, c_0_993])).
cnf(c_0_1036,negated_conjecture,(epred24_0|epred21_0|~epred13_0), inference(spm,[status(thm)],[c_0_440, c_0_994])).
cnf(c_0_1037,negated_conjecture,(epred41_0|epred24_0|epred14_0|epred13_0), inference(spm,[status(thm)],[c_0_995, c_0_996])).
cnf(c_0_1038,negated_conjecture,(epred24_0|epred15_0|epred13_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_933, c_0_997]), c_0_585])).
cnf(c_0_1039,negated_conjecture,(epred45_0|epred29_0|epred23_0|epred19_0|epred12_0), inference(spm,[status(thm)],[c_0_998, c_0_693])).
cnf(c_0_1040,negated_conjecture,(~epred29_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_912]), c_0_999]), c_0_1000]), ['final']).
cnf(c_0_1041,negated_conjecture,(epred13_0|hskp10|hskp12|hskp21), inference(spm,[status(thm)],[c_0_1001, c_0_108])).
cnf(c_0_1042,negated_conjecture,(hskp6|hskp7|hskp13|hskp19|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1002, c_0_203]), c_0_414]), c_0_946])).
cnf(c_0_1043,negated_conjecture,(epred39_0|epred29_0|~hskp6), inference(spm,[status(thm)],[c_0_764, c_0_1003]), ['final']).
cnf(c_0_1044,negated_conjecture,(epred39_0|epred29_0|epred19_0|~hskp10), inference(spm,[status(thm)],[c_0_764, c_0_957])).
cnf(c_0_1045,negated_conjecture,(epred39_0|hskp10|hskp19), inference(sr,[status(thm)],[c_0_935, c_0_1004])).
cnf(c_0_1046,negated_conjecture,(epred40_0|epred39_0|epred7_0|hskp6), inference(spm,[status(thm)],[c_0_340, c_0_1005])).
cnf(c_0_1047,negated_conjecture,(epred42_0|epred40_0|hskp6|hskp7), inference(spm,[status(thm)],[c_0_268, c_0_1006])).
cnf(c_0_1048,negated_conjecture,(epred45_0|epred42_0|epred40_0|epred16_0), inference(spm,[status(thm)],[c_0_249, c_0_1007])).
cnf(c_0_1049,negated_conjecture,(epred42_0|epred40_0|epred21_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_504, c_0_1008]), c_0_1009])).
cnf(c_0_1050,negated_conjecture,(epred41_0|epred40_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1010]), c_0_434]), c_0_1011])).
cnf(c_0_1051,negated_conjecture,(epred43_0|epred18_0|hskp10|hskp21), inference(spm,[status(thm)],[c_0_1012, c_0_108])).
cnf(c_0_1052,negated_conjecture,(epred43_0|epred41_0|epred4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1013, c_0_1014]), c_0_1015])).
cnf(c_0_1053,negated_conjecture,(epred43_0|epred41_0|epred37_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1016, c_0_1017]), c_0_1018])).
cnf(c_0_1054,negated_conjecture,(epred43_0|epred20_0|hskp10|hskp21), inference(spm,[status(thm)],[c_0_1019, c_0_108])).
cnf(c_0_1055,negated_conjecture,(epred43_0|hskp2|hskp6|hskp20|~hskp24), inference(spm,[status(thm)],[c_0_803, c_0_656])).
cnf(c_0_1056,negated_conjecture,(epred43_0|hskp6|hskp20|~hskp27), inference(spm,[status(thm)],[c_0_840, c_0_590])).
cnf(c_0_1057,negated_conjecture,(epred22_0|c1_1(X1)|hskp9|~c3_1(X1)|~c2_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1020, c_0_153])])).
cnf(c_0_1058,negated_conjecture,(epred43_0|epred15_0|~hskp4), inference(spm,[status(thm)],[c_0_1021, c_0_1022])).
cnf(c_0_1059,negated_conjecture,(epred26_0|hskp4|~hskp19), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1023, c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_1060,negated_conjecture,(epred17_0|hskp17|hskp19), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1024, c_0_1025]), c_0_585])).
cnf(c_0_1061,negated_conjecture,(~epred26_0|~epred24_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_147, c_0_930]), c_0_266])).
cnf(c_0_1062,negated_conjecture,(epred43_0|epred24_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_866, c_0_1026]), c_0_913]), c_0_1027])).
cnf(c_0_1063,negated_conjecture,(epred39_0|epred37_0|epred7_0|hskp6), inference(spm,[status(thm)],[c_0_340, c_0_1028])).
cnf(c_0_1064,negated_conjecture,(epred23_0|epred21_0|epred19_0|hskp7), inference(spm,[status(thm)],[c_0_344, c_0_1029])).
cnf(c_0_1065,negated_conjecture,(epred19_0|epred13_0|hskp12), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_950, c_0_1030]), c_0_585]), c_0_951])).
cnf(c_0_1066,negated_conjecture,(~epred5_0|~epred2_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_148, c_0_749]), c_0_811])).
cnf(c_0_1067,negated_conjecture,(epred24_0|epred13_0|epred2_0), inference(spm,[status(thm)],[c_0_816, c_0_1031])).
cnf(c_0_1068,negated_conjecture,(epred24_0|epred13_0|epred5_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_737, c_0_1032]), c_0_867]), c_0_739])).
cnf(c_0_1069,negated_conjecture,(epred41_0|epred21_0|epred19_0|~epred13_0), inference(spm,[status(thm)],[c_0_440, c_0_1033])).
cnf(c_0_1070,negated_conjecture,(epred41_0|epred19_0|epred13_0), inference(spm,[status(thm)],[c_0_1034, c_0_1035])).
cnf(c_0_1071,negated_conjecture,(epred41_0|epred24_0|epred21_0|epred14_0), inference(spm,[status(thm)],[c_0_1036, c_0_1037])).
cnf(c_0_1072,negated_conjecture,(epred24_0|epred21_0|epred15_0), inference(spm,[status(thm)],[c_0_1036, c_0_1038])).
cnf(c_0_1073,negated_conjecture,(epred45_0|epred23_0|epred19_0|epred12_0), inference(sr,[status(thm)],[c_0_1039, c_0_1040])).
cnf(c_0_1074,negated_conjecture,(epred24_0|epred13_0|~hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_564, c_0_658]), c_0_585]), c_0_76]), c_0_816])).
cnf(c_0_1075,negated_conjecture,(epred24_0|epred13_0|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_816, c_0_1041]), c_0_305])).
cnf(c_0_1076,negated_conjecture,(epred13_0|hskp6|hskp19|hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_1042]), c_0_383])).
cnf(c_0_1077,negated_conjecture,(epred43_0|epred39_0|epred29_0|~hskp19), inference(spm,[status(thm)],[c_0_1043, c_0_934])).
cnf(c_0_1078,negated_conjecture,(epred39_0|epred19_0|hskp19), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1044, c_0_1045]), c_0_1040])).
cnf(c_0_1079,negated_conjecture,(epred39_0|epred38_0|hskp10|~hskp26), inference(spm,[status(thm)],[c_0_764, c_0_708])).
cnf(c_0_1080,negated_conjecture,(epred40_0|epred39_0|epred29_0|epred19_0|epred6_0), inference(spm,[status(thm)],[c_0_1044, c_0_370])).
cnf(c_0_1081,negated_conjecture,(epred40_0|epred39_0|epred29_0|epred7_0), inference(spm,[status(thm)],[c_0_1043, c_0_1046])).
cnf(c_0_1082,negated_conjecture,(epred45_0|epred29_0|~hskp6), inference(spm,[status(thm)],[c_0_408, c_0_1003]), ['final']).
cnf(c_0_1083,negated_conjecture,(epred42_0|epred40_0|epred17_0|hskp6), inference(spm,[status(thm)],[c_0_667, c_0_1047])).
cnf(c_0_1084,negated_conjecture,(epred45_0|epred42_0|epred40_0|~epred17_0), inference(spm,[status(thm)],[c_0_504, c_0_1048])).
cnf(c_0_1085,negated_conjecture,(epred42_0|epred40_0|~epred45_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_1049]), c_0_1050])).
cnf(c_0_1086,negated_conjecture,(epred41_0|epred17_0|~hskp10), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_332]), c_0_80]), c_0_81])).
cnf(c_0_1087,negated_conjecture,(epred43_0|epred18_0|epred17_0|hskp10), inference(spm,[status(thm)],[c_0_424, c_0_1051])).
cnf(c_0_1088,negated_conjecture,(epred43_0|epred41_0|~epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_737, c_0_1052]), c_0_1053])).
cnf(c_0_1089,negated_conjecture,(epred43_0|epred20_0|epred17_0|hskp10), inference(spm,[status(thm)],[c_0_424, c_0_1054])).
cnf(c_0_1090,negated_conjecture,(epred43_0|hskp6|hskp18|hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1055, c_0_100]), c_0_1056]), c_0_840])).
cnf(c_0_1091,negated_conjecture,(epred22_0|hskp9|~hskp4), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1057, c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_1092,negated_conjecture,(epred43_0|epred26_0|epred15_0|~hskp19), inference(spm,[status(thm)],[c_0_1058, c_0_1059])).
cnf(c_0_1093,negated_conjecture,(epred17_0|hskp19), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_542, c_0_1060]), c_0_585])).
cnf(c_0_1094,negated_conjecture,(epred43_0|~epred26_0), inference(spm,[status(thm)],[c_0_1061, c_0_1062])).
cnf(c_0_1095,negated_conjecture,(epred39_0|epred37_0|epred29_0|epred19_0|epred6_0), inference(spm,[status(thm)],[c_0_1044, c_0_683])).
cnf(c_0_1096,negated_conjecture,(epred39_0|epred37_0|epred29_0|epred7_0), inference(spm,[status(thm)],[c_0_1043, c_0_1063])).
cnf(c_0_1097,negated_conjecture,(epred37_0|c1_1(a1301)|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_617, c_0_36]), c_0_35])).
cnf(c_0_1098,negated_conjecture,(epred23_0|epred21_0|epred19_0|epred12_0), inference(spm,[status(thm)],[c_0_410, c_0_1064])).
cnf(c_0_1099,negated_conjecture,(epred23_0|epred19_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_1065]), c_0_992])).
cnf(c_0_1100,negated_conjecture,(epred24_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1066, c_0_1067]), c_0_1068])).
cnf(c_0_1101,negated_conjecture,(epred41_0|epred21_0|epred19_0), inference(spm,[status(thm)],[c_0_1069, c_0_1070])).
cnf(c_0_1102,negated_conjecture,(epred41_0|epred24_0|epred21_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_770, c_0_1071]), c_0_1072])).
cnf(c_0_1103,negated_conjecture,(epred45_0|epred23_0|epred19_0|~epred13_0), inference(spm,[status(thm)],[c_0_440, c_0_1073])).
cnf(c_0_1104,negated_conjecture,(epred24_0|epred13_0|~hskp7), inference(spm,[status(thm)],[c_0_1074, c_0_383])).
cnf(c_0_1105,negated_conjecture,(epred24_0|epred23_0|hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_1075]), c_0_756])).
cnf(c_0_1106,negated_conjecture,(epred17_0|epred3_0|hskp26|~hskp7), inference(spm,[status(thm)],[c_0_423, c_0_815])).
cnf(c_0_1107,negated_conjecture,(epred17_0|hskp7|hskp26), inference(spm,[status(thm)],[c_0_424, c_0_817])).
cnf(c_0_1108,negated_conjecture,(epred19_0|epred13_0|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_950, c_0_765]), c_0_585]), c_0_951])).
cnf(c_0_1109,negated_conjecture,(epred13_0|hskp6|hskp19|hskp21), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_931, c_0_1076]), c_0_585])).
cnf(c_0_1110,negated_conjecture,(epred43_0|epred39_0|epred19_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1077, c_0_1078]), c_0_1040])).
cnf(c_0_1111,negated_conjecture,(epred40_0|epred39_0|epred38_0|hskp10), inference(spm,[status(thm)],[c_0_1079, c_0_430])).
cnf(c_0_1112,negated_conjecture,(epred40_0|epred39_0|epred29_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_1080]), c_0_1081])).
cnf(c_0_1113,negated_conjecture,(epred42_0|epred40_0|epred17_0), inference(spm,[status(thm)],[c_0_249, c_0_698])).
cnf(c_0_1114,negated_conjecture,(epred42_0|epred40_0|epred29_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1082, c_0_1083]), c_0_1084]), c_0_1085])).
cnf(c_0_1115,negated_conjecture,(epred43_0|epred41_0|epred18_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1086, c_0_1087]), c_0_1088])).
cnf(c_0_1116,negated_conjecture,(epred43_0|epred41_0|epred20_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1086, c_0_1089]), c_0_1088])).
cnf(c_0_1117,negated_conjecture,(epred43_0|hskp6|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_924, c_0_1090]), c_0_923])).
cnf(c_0_1118,negated_conjecture,(epred43_0|epred22_0|~hskp4), inference(spm,[status(thm)],[c_0_747, c_0_1091])).
cnf(c_0_1119,negated_conjecture,(epred43_0|epred17_0|epred15_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1092, c_0_1093]), c_0_1094])).
cnf(c_0_1120,negated_conjecture,(epred39_0|epred37_0|epred29_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_1095]), c_0_1096])).
cnf(c_0_1121,negated_conjecture,(epred37_0|hskp23|hskp24|~hskp19), inference(spm,[status(thm)],[c_0_461, c_0_884])).
cnf(c_0_1122,negated_conjecture,(epred37_0|epred11_0|~hskp24), inference(spm,[status(thm)],[c_0_463, c_0_1097])).
cnf(c_0_1123,negated_conjecture,(epred23_0|epred21_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_1098]), c_0_1099])).
cnf(c_0_1124,negated_conjecture,(epred24_0|epred21_0), inference(spm,[status(thm)],[c_0_1036, c_0_1100])).
cnf(c_0_1125,negated_conjecture,(epred41_0|epred21_0|~epred42_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1101]), c_0_1102]), ['final']).
cnf(c_0_1126,negated_conjecture,(epred45_0|epred23_0|epred19_0), inference(spm,[status(thm)],[c_0_1103, c_0_1099])).
cnf(c_0_1127,negated_conjecture,(epred24_0|epred23_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1104, c_0_1105]), c_0_756])).
cnf(c_0_1128,negated_conjecture,(epred41_0|epred17_0|~hskp26), inference(spm,[status(thm)],[c_0_424, c_0_858])).
cnf(c_0_1129,negated_conjecture,(epred17_0|epred3_0|hskp26), inference(spm,[status(thm)],[c_0_1106, c_0_1107])).
cnf(c_0_1130,negated_conjecture,(epred17_0|epred15_0|~hskp7), inference(spm,[status(thm)],[c_0_423, c_0_686])).
cnf(c_0_1131,negated_conjecture,(epred43_0|epred19_0|~hskp4), inference(spm,[status(thm)],[c_0_1021, c_0_734])).
cnf(c_0_1132,negated_conjecture,(epred19_0|epred13_0|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1108, c_0_1003]), c_0_1040])).
cnf(c_0_1133,negated_conjecture,(epred19_0|epred13_0|hskp6|hskp19), inference(spm,[status(thm)],[c_0_951, c_0_1109])).
cnf(c_0_1134,negated_conjecture,(epred43_0|epred39_0|~epred42_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1110]), c_0_1062])).
cnf(c_0_1135,negated_conjecture,(epred40_0|epred39_0|epred19_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1044, c_0_1111]), c_0_1112]), c_0_1000])).
cnf(c_0_1136,negated_conjecture,(epred42_0|epred40_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_1113]), c_0_1114])).
cnf(c_0_1137,negated_conjecture,(epred43_0|epred41_0|~epred19_0), inference(spm,[status(thm)],[c_0_338, c_0_1115])).
cnf(c_0_1138,negated_conjecture,(epred43_0|epred41_0|epred21_0|epred19_0), inference(spm,[status(thm)],[c_0_1088, c_0_724])).
cnf(c_0_1139,negated_conjecture,(epred43_0|epred41_0|~epred21_0), inference(spm,[status(thm)],[c_0_818, c_0_1116])).
cnf(c_0_1140,negated_conjecture,(epred43_0|epred29_0|epred4_0|~hskp6), inference(spm,[status(thm)],[c_0_922, c_0_1003])).
cnf(c_0_1141,negated_conjecture,(epred43_0|epred4_0|hskp6), inference(spm,[status(thm)],[c_0_922, c_0_1117])).
cnf(c_0_1142,negated_conjecture,(epred43_0|epred26_0|epred22_0|~hskp19), inference(spm,[status(thm)],[c_0_1118, c_0_1059])).
cnf(c_0_1143,negated_conjecture,(epred43_0|epred17_0|~epred22_0), inference(spm,[status(thm)],[c_0_519, c_0_1119])).
cnf(c_0_1144,negated_conjecture,(epred37_0|epred19_0|epred13_0), inference(spm,[status(thm)],[c_0_637, c_0_1065])).
cnf(c_0_1145,negated_conjecture,(epred39_0|epred37_0|epred19_0), inference(sr,[status(thm)],[c_0_1120, c_0_1040])).
cnf(c_0_1146,negated_conjecture,(epred37_0|epred11_0|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_1121]), c_0_1122])).
cnf(c_0_1147,negated_conjecture,(epred23_0|epred21_0|~epred42_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1123]), c_0_1124])).
cnf(c_0_1148,negated_conjecture,(epred41_0|epred23_0|~epred42_0), inference(spm,[status(thm)],[c_0_861, c_0_1125])).
cnf(c_0_1149,negated_conjecture,(epred45_0|epred23_0|~epred42_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1126]), c_0_1127])).
cnf(c_0_1150,negated_conjecture,(epred41_0|epred17_0|epred3_0), inference(spm,[status(thm)],[c_0_1128, c_0_1129])).
cnf(c_0_1151,negated_conjecture,(epred41_0|epred14_0|epred4_0|hskp17), inference(spm,[status(thm)],[c_0_363, c_0_682])).
cnf(c_0_1152,negated_conjecture,(epred41_0|~epred36_0), inference(spm,[status(thm)],[c_0_552, c_0_398]), ['final']).
cnf(c_0_1153,negated_conjecture,(epred17_0|epred15_0|hskp6|hskp17), inference(spm,[status(thm)],[c_0_1130, c_0_714])).
cnf(c_0_1154,negated_conjecture,(epred43_0|epred26_0|epred19_0|~hskp19), inference(spm,[status(thm)],[c_0_1131, c_0_1059])).
cnf(c_0_1155,negated_conjecture,(epred19_0|epred13_0|hskp19), inference(spm,[status(thm)],[c_0_1132, c_0_1133])).
cnf(c_0_1156,negated_conjecture,(epred43_0|epred39_0|epred11_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1077, c_0_538]), c_0_1040]), c_0_1134])).
cnf(c_0_1157,negated_conjecture,(epred40_0|epred39_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1135]), c_0_434]), c_0_1136])).
cnf(c_0_1158,negated_conjecture,(epred43_0|epred41_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1137, c_0_1138]), c_0_1139])).
cnf(c_0_1159,negated_conjecture,(epred43_0|epred4_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1140, c_0_1141]), c_0_1040])).
cnf(c_0_1160,negated_conjecture,(epred43_0|epred17_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1142, c_0_1093]), c_0_1143]), c_0_1094])).
cnf(c_0_1161,negated_conjecture,(epred37_0|epred19_0|~epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_866, c_0_1144]), c_0_1145])).
cnf(c_0_1162,negated_conjecture,(epred42_0|epred37_0|epred11_0), inference(spm,[status(thm)],[c_0_1146, c_0_538])).
cnf(c_0_1163,negated_conjecture,(epred36_0|epred17_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_302]), c_0_624])).
cnf(c_0_1164,negated_conjecture,(epred23_0|~epred42_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_1147]), c_0_1148]), c_0_1149]), ['final']).
cnf(c_0_1165,negated_conjecture,(epred41_0|epred17_0|~epred4_0), inference(spm,[status(thm)],[c_0_954, c_0_1150])).
cnf(c_0_1166,negated_conjecture,(epred41_0|epred17_0|epred14_0|epred4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_542, c_0_1151]), c_0_1152])).
cnf(c_0_1167,negated_conjecture,(epred41_0|epred17_0|epred15_0|hskp17), inference(spm,[status(thm)],[c_0_363, c_0_1153])).
cnf(c_0_1168,negated_conjecture,(epred43_0|epred19_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1154, c_0_1155]), c_0_1094])).
cnf(c_0_1169,negated_conjecture,(epred43_0|epred39_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_1156]), c_0_1157]), c_0_1158])).
cnf(c_0_1170,negated_conjecture,(epred43_0|~epred37_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_737, c_0_1159]), c_0_1160])).
cnf(c_0_1171,negated_conjecture,(epred37_0|~epred42_0|~epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1161]), c_0_739])).
cnf(c_0_1172,negated_conjecture,(epred42_0|epred37_0|~epred41_0|~epred40_0), inference(spm,[status(thm)],[c_0_396, c_0_1162])).
cnf(c_0_1173,negated_conjecture,(epred38_0|epred19_0|epred11_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_795]), c_0_369])).
cnf(c_0_1174,negated_conjecture,(epred41_0|epred38_0|epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_795]), c_0_692])).
cnf(c_0_1175,negated_conjecture,(epred42_0|epred13_0|~hskp7), inference(spm,[status(thm)],[c_0_350, c_0_644])).
cnf(c_0_1176,negated_conjecture,(epred17_0|epred11_0), inference(sr,[status(thm)],[c_0_1163, c_0_585])).
cnf(c_0_1177,negated_conjecture,(~epred42_0|~epred24_0), inference(spm,[status(thm)],[c_0_513, c_0_1164])).
cnf(c_0_1178,negated_conjecture,(epred41_0|epred17_0|epred14_0), inference(spm,[status(thm)],[c_0_1165, c_0_1166])).
cnf(c_0_1179,negated_conjecture,(epred41_0|epred17_0|epred15_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_542, c_0_1167]), c_0_585])).
cnf(c_0_1180,negated_conjecture,(epred43_0|epred19_0|~epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_866, c_0_1168]), c_0_1169])).
cnf(c_0_1181,negated_conjecture,(epred43_0|~epred42_0|~epred38_0), inference(spm,[status(thm)],[c_0_1170, c_0_1171])).
cnf(c_0_1182,negated_conjecture,(epred43_0|epred42_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1170, c_0_1172]), c_0_1136]), c_0_1158])).
cnf(c_0_1183,negated_conjecture,(epred38_0|epred19_0|~epred41_0|~epred40_0), inference(spm,[status(thm)],[c_0_396, c_0_1173])).
cnf(c_0_1184,negated_conjecture,(epred41_0|epred38_0|~epred42_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1174]), c_0_913])).
cnf(c_0_1185,negated_conjecture,(epred39_0|epred29_0|epred23_0|epred19_0|epred12_0), inference(spm,[status(thm)],[c_0_1044, c_0_693])).
cnf(c_0_1186,negated_conjecture,(epred42_0|epred23_0|epred13_0|~hskp19), inference(spm,[status(thm)],[c_0_1175, c_0_668])).
cnf(c_0_1187,negated_conjecture,(epred17_0|~epred41_0|~epred40_0), inference(spm,[status(thm)],[c_0_396, c_0_1176])).
cnf(c_0_1188,negated_conjecture,(epred40_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1177, c_0_434]), c_0_1136]), ['final']).
cnf(c_0_1189,negated_conjecture,(epred41_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_770, c_0_1178]), c_0_1179])).
cnf(c_0_1190,negated_conjecture,(epred44_0|~c0_1(a1283)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_271]), c_0_30])).
cnf(c_0_1191,negated_conjecture,(epred44_0|epred42_0|hskp19|hskp20), inference(spm,[status(thm)],[c_0_374, c_0_878])).
cnf(c_0_1192,negated_conjecture,(epred43_0|~epred38_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1180]), c_0_1181]), c_0_1182])).
cnf(c_0_1193,negated_conjecture,(epred38_0|~epred42_0|~epred40_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1183]), c_0_1184]), c_0_913])).
cnf(c_0_1194,negated_conjecture,(epred39_0|epred23_0|epred19_0|epred12_0), inference(sr,[status(thm)],[c_0_1185, c_0_1040])).
cnf(c_0_1195,negated_conjecture,(epred42_0|epred23_0|epred13_0|epred11_0), inference(spm,[status(thm)],[c_0_1186, c_0_538])).
cnf(c_0_1196,negated_conjecture,(epred38_0|~c0_1(a1283)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_614, c_0_28]), c_0_30])).
cnf(c_0_1197,negated_conjecture,(epred17_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1187, c_0_1188])]), c_0_1189]), ['final']).
cnf(c_0_1198,negated_conjecture,(epred44_0|epred37_0|~hskp19), inference(spm,[status(thm)],[c_0_1190, c_0_884])).
cnf(c_0_1199,negated_conjecture,(epred44_0|epred42_0|epred21_0|hskp19), inference(spm,[status(thm)],[c_0_404, c_0_1191])).
cnf(c_0_1200,negated_conjecture,(epred45_0|epred44_0|epred42_0|hskp19), inference(spm,[status(thm)],[c_0_408, c_0_1191])).
cnf(c_0_1201,negated_conjecture,(epred43_0|~epred40_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1192, c_0_1193]), c_0_1182])).
cnf(c_0_1202,negated_conjecture,(epred39_0|epred23_0|epred19_0|~epred13_0), inference(spm,[status(thm)],[c_0_440, c_0_1194])).
cnf(c_0_1203,negated_conjecture,(epred42_0|epred25_0|epred13_0|hskp19), inference(spm,[status(thm)],[c_0_1175, c_0_669])).
cnf(c_0_1204,negated_conjecture,(epred42_0|epred23_0|epred13_0|~epred25_0), inference(spm,[status(thm)],[c_0_646, c_0_1195])).
cnf(c_0_1205,negated_conjecture,(epred38_0|hskp6|hskp29|~hskp19), inference(spm,[status(thm)],[c_0_1196, c_0_745])).
cnf(c_0_1206,negated_conjecture,(epred42_0|epred39_0|epred16_0|hskp19), inference(spm,[status(thm)],[c_0_764, c_0_918])).
cnf(c_0_1207,negated_conjecture,(~epred16_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_504, c_0_1197])]), ['final']).
cnf(c_0_1208,negated_conjecture,(epred44_0|epred42_0|epred37_0|epred21_0), inference(spm,[status(thm)],[c_0_1198, c_0_1199])).
cnf(c_0_1209,negated_conjecture,(epred44_0|epred41_0|epred37_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1198, c_0_1017]), c_0_1018])).
cnf(c_0_1210,negated_conjecture,(epred45_0|epred44_0|epred42_0|epred37_0), inference(spm,[status(thm)],[c_0_1198, c_0_1200])).
cnf(c_0_1211,negated_conjecture,(epred37_0|~epred44_0|~epred43_0), inference(spm,[status(thm)],[c_0_458, c_0_739])).
cnf(c_0_1212,negated_conjecture,(epred43_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1201, c_0_1188])]), ['final']).
cnf(c_0_1213,negated_conjecture,(epred39_0|epred23_0|epred19_0), inference(spm,[status(thm)],[c_0_1202, c_0_1099])).
cnf(c_0_1214,negated_conjecture,(epred42_0|epred23_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1186, c_0_1203]), c_0_1204])).
cnf(c_0_1215,negated_conjecture,(epred38_0|hskp6|hskp20|~hskp19), inference(spm,[status(thm)],[c_0_655, c_0_1205])).
cnf(c_0_1216,negated_conjecture,(epred42_0|epred39_0|hskp19), inference(sr,[status(thm)],[c_0_1206, c_0_1207])).
cnf(c_0_1217,negated_conjecture,(epred44_0|epred42_0|epred37_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_1208]), c_0_1209]), c_0_1210])).
cnf(c_0_1218,negated_conjecture,(epred37_0|~epred42_0), inference(spm,[status(thm)],[c_0_1177, c_0_739])).
cnf(c_0_1219,negated_conjecture,(epred37_0|~epred44_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1211, c_0_1212])])).
cnf(c_0_1220,negated_conjecture,(epred23_0|epred19_0|~epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_866, c_0_1099]), c_0_1213])).
cnf(c_0_1221,negated_conjecture,(epred42_0|epred23_0|~epred39_0|~epred38_0), inference(spm,[status(thm)],[c_0_866, c_0_1214])).
cnf(c_0_1222,negated_conjecture,(epred42_0|epred24_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_816, c_0_608]), c_0_305])).
cnf(c_0_1223,negated_conjecture,(epred42_0|~c3_1(a1287)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_511]), c_0_103]), ['final']).
cnf(c_0_1224,negated_conjecture,(epred38_0|hskp5|hskp6|hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_836]), c_0_736]), c_0_1215])).
cnf(c_0_1225,negated_conjecture,(~epred37_0|~epred4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_737, c_0_1197])])).
cnf(c_0_1226,negated_conjecture,(epred37_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1198, c_0_1216]), c_0_1217]), c_0_1218]), c_0_1219]), ['final']).
cnf(c_0_1227,negated_conjecture,(epred23_0|~epred42_0|~epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1220]), c_0_1127])).
cnf(c_0_1228,negated_conjecture,(epred42_0|~epred39_0|~epred38_0|~epred24_0), inference(spm,[status(thm)],[c_0_513, c_0_1221])).
cnf(c_0_1229,negated_conjecture,(epred42_0|epred24_0|~epred39_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_866, c_0_1222]), c_0_913])).
cnf(c_0_1230,negated_conjecture,(epred42_0|epred4_0|~hskp20), inference(spm,[status(thm)],[c_0_1223, c_0_881]), ['final']).
cnf(c_0_1231,negated_conjecture,(epred42_0|epred38_0|hskp6|hskp20), inference(spm,[status(thm)],[c_0_527, c_0_1224])).
cnf(c_0_1232,negated_conjecture,(~epred4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1225, c_0_1226])]), ['final']).
cnf(c_0_1233,negated_conjecture,(epred38_0|~epred42_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1193, c_0_1188])])).
cnf(c_0_1234,negated_conjecture,(~epred42_0|~epred38_0|~epred24_0), inference(spm,[status(thm)],[c_0_513, c_0_1227])).
cnf(c_0_1235,negated_conjecture,(epred24_0|~epred39_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_866, c_0_1100]), c_0_913])).
cnf(c_0_1236,negated_conjecture,(epred42_0|~epred39_0|~epred38_0), inference(spm,[status(thm)],[c_0_1228, c_0_1229])).
cnf(c_0_1237,negated_conjecture,(epred42_0|epred29_0|epred4_0|~hskp6), inference(spm,[status(thm)],[c_0_1230, c_0_1003]), ['final']).
cnf(c_0_1238,negated_conjecture,(epred38_0|hskp6), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1230, c_0_1231]), c_0_1232]), c_0_1233])).
cnf(c_0_1239,negated_conjecture,(epred29_0|hskp20|~c1_1(a1252)|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_64]), c_0_45])).
cnf(c_0_1240,negated_conjecture,(~epred39_0|~epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1234, c_0_1235]), c_0_1236])).
cnf(c_0_1241,negated_conjecture,(epred38_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1237, c_0_1238]), c_0_1040]), c_0_1232]), c_0_1233]), ['final']).
cnf(c_0_1242,negated_conjecture,(hskp19|hskp20|~hskp0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1239, c_0_778]), c_0_1040]), c_0_1004])).
cnf(c_0_1243,negated_conjecture,(c3_1(a1266)|hskp6|hskp29|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_634, c_0_77]), c_0_91])).
cnf(c_0_1244,negated_conjecture,(epred39_0|epred30_0|c1_1(X1)|hskp19|~c0_1(X1)), inference(spm,[status(thm)],[c_0_746, c_0_662])).
cnf(c_0_1245,negated_conjecture,(~epred39_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1240, c_0_1241])]), ['final']).
cnf(c_0_1246,negated_conjecture,(hskp19|hskp20|hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1242, c_0_212]), c_0_84])).
cnf(c_0_1247,negated_conjecture,(hskp2|hskp7|hskp19|hskp20|~hskp29), inference(spm,[status(thm)],[c_0_631, c_0_167])).
cnf(c_0_1248,negated_conjecture,(hskp6|hskp7|hskp19|hskp29|~hskp12), inference(spm,[status(thm)],[c_0_229, c_0_1243])).
cnf(c_0_1249,negated_conjecture,(hskp19|~hskp9), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1244, c_0_67]), c_0_1245]), c_0_1004]), c_0_68])).
cnf(c_0_1250,negated_conjecture,(hskp19|hskp26), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1246]), c_0_1245])).
cnf(c_0_1251,negated_conjecture,(hskp19|~hskp2), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1244, c_0_124]), c_0_1245]), c_0_1004]), c_0_13])).
cnf(c_0_1252,negated_conjecture,(hskp2|hskp6|hskp7|hskp19|hskp20|~hskp12), inference(spm,[status(thm)],[c_0_1247, c_0_1248])).
cnf(c_0_1253,negated_conjecture,(hskp6|hskp19|hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1249, c_0_833]), c_0_1250]), c_0_1251])).
cnf(c_0_1254,negated_conjecture,(epred30_0|hskp6|hskp7|hskp19|hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_842, c_0_1252]), c_0_915])).
cnf(c_0_1255,negated_conjecture,(hskp6|hskp19), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1253]), c_0_1245])).
cnf(c_0_1256,negated_conjecture,(epred30_0|epred4_0|c1_1(X1)|c3_1(X1)|hskp19|~c0_1(X1)), inference(spm,[status(thm)],[c_0_746, c_0_615])).
cnf(c_0_1257,negated_conjecture,(hskp6|hskp7|hskp19|hskp20), inference(sr,[status(thm)],[c_0_1254, c_0_1004])).
cnf(c_0_1258,negated_conjecture,(hskp19), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1043, c_0_1255]), c_0_1245]), c_0_1040]), ['final']).
cnf(c_0_1259,negated_conjecture,(hskp7|hskp19|~hskp2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_1256]), c_0_1004]), c_0_1232]), c_0_124]), c_0_13])).
cnf(c_0_1260,negated_conjecture,(epred4_0|hskp2|hskp4|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_881]), c_0_95]), c_0_103])).
cnf(c_0_1261,negated_conjecture,(epred39_0|hskp6|hskp7|hskp19), inference(spm,[status(thm)],[c_0_764, c_0_1257])).
cnf(c_0_1262,negated_conjecture,(hskp6|hskp10|hskp28), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_239, c_0_1258])])).
cnf(c_0_1263,negated_conjecture,(epred20_0|hskp10|~hskp3), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_663, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_1264,negated_conjecture,(hskp7|hskp19|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1259, c_0_1260]), c_0_1232]), c_0_176])).
cnf(c_0_1265,negated_conjecture,(hskp6|hskp7|hskp19), inference(sr,[status(thm)],[c_0_1261, c_0_1245])).
cnf(c_0_1266,negated_conjecture,(epred24_0|c0_1(X1)|hskp15|~c1_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_207]), c_0_305])).
cnf(c_0_1267,negated_conjecture,(epred24_0|hskp21), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_987, c_0_1262]), c_0_585]), c_0_305]), c_0_681])).
cnf(c_0_1268,negated_conjecture,(epred20_0|epred6_0|hskp10|~hskp15), inference(spm,[status(thm)],[c_0_1263, c_0_265])).
cnf(c_0_1269,negated_conjecture,(epred32_0|hskp7|~c0_1(a1283)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_601]), c_0_30])).
cnf(c_0_1270,negated_conjecture,(epred5_0|c0_1(a1283)|hskp13|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_812, c_0_28]), c_0_29])).
cnf(c_0_1271,negated_conjecture,(hskp7|hskp19), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1264, c_0_1003]), c_0_1040]), c_0_1265])).
cnf(c_0_1272,negated_conjecture,(epred20_0|hskp10|hskp15|~hskp7), inference(spm,[status(thm)],[c_0_222, c_0_691])).
cnf(c_0_1273,negated_conjecture,(epred24_0|epred7_0|c0_1(X1)|~c1_1(X1)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_282, c_0_307])).
cnf(c_0_1274,negated_conjecture,(epred7_0|c0_1(a1277)|~hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_307]), c_0_114])).
cnf(c_0_1275,negated_conjecture,(epred24_0|hskp15), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_1266]), c_0_106]), c_0_1267])).
cnf(c_0_1276,negated_conjecture,(epred20_0|epred6_0|hskp10|~hskp13), inference(spm,[status(thm)],[c_0_1268, c_0_258])).
cnf(c_0_1277,negated_conjecture,(epred5_0|hskp7|hskp13), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1269, c_0_1270]), c_0_622]), c_0_1271])).
cnf(c_0_1278,negated_conjecture,(epred20_0|epred6_0|hskp10|~hskp7), inference(spm,[status(thm)],[c_0_1268, c_0_1272])).
cnf(c_0_1279,negated_conjecture,(epred24_0|epred7_0|~hskp21), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1273, c_0_105]), c_0_106]), c_0_107])).
cnf(c_0_1280,negated_conjecture,(epred36_0|epred7_0|hskp21|~hskp17), inference(spm,[status(thm)],[c_0_531, c_0_1274])).
cnf(c_0_1281,negated_conjecture,(epred32_0|epred7_0|hskp6|hskp17), inference(spm,[status(thm)],[c_0_340, c_0_687])).
cnf(c_0_1282,negated_conjecture,(epred24_0|epred2_0), inference(spm,[status(thm)],[c_0_898, c_0_1275])).
cnf(c_0_1283,negated_conjecture,(epred20_0|epred6_0|epred5_0|hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1276, c_0_1277]), c_0_1278])).
cnf(c_0_1284,negated_conjecture,(epred36_0|epred24_0|epred7_0|~hskp17), inference(spm,[status(thm)],[c_0_1279, c_0_1280])).
cnf(c_0_1285,negated_conjecture,(epred32_0|epred24_0|epred7_0|hskp17), inference(spm,[status(thm)],[c_0_681, c_0_1281])).
cnf(c_0_1286,negated_conjecture,(epred24_0|~epred5_0), inference(spm,[status(thm)],[c_0_1066, c_0_1282])).
cnf(c_0_1287,negated_conjecture,(epred24_0|epred20_0|epred6_0|epred5_0), inference(spm,[status(thm)],[c_0_305, c_0_1283])).
cnf(c_0_1288,negated_conjecture,(epred36_0|epred32_0|epred24_0|epred7_0), inference(spm,[status(thm)],[c_0_1284, c_0_1285])).
fof(c_0_1289, plain, ((~epred28_0<=>![X1]:((c2_1(X1)|c3_1(X1))|~c1_1(X1)))), introduced(definition)).
cnf(c_0_1290,negated_conjecture,(epred30_0|c1_1(a1280)|hskp19|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_746, c_0_17]), c_0_109])).
cnf(c_0_1291,negated_conjecture,(epred42_0|epred16_0|epred4_0|hskp19), inference(spm,[status(thm)],[c_0_1230, c_0_918])).
cnf(c_0_1292,negated_conjecture,(epred24_0|epred20_0|epred6_0), inference(spm,[status(thm)],[c_0_1286, c_0_1287])).
cnf(c_0_1293,negated_conjecture,(epred24_0|epred7_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_1288, c_0_585]), c_0_622])).
cnf(c_0_1294,negated_conjecture,(c1_1(X1)|hskp7|hskp12|hskp13|~c3_1(X1)|~c0_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_825]), c_0_177])).
cnf(c_0_1295,negated_conjecture,(epred28_0|c3_1(X1)|c2_1(X1)|~c1_1(X1)), inference(split_equiv,[status(thm)],[c_0_1289])).
cnf(c_0_1296,negated_conjecture,(epred30_0|hskp19|hskp20|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_631, c_0_1290]), c_0_842])).
cnf(c_0_1297,negated_conjecture,(epred42_0|epred4_0|hskp19), inference(sr,[status(thm)],[c_0_1291, c_0_1207])).
cnf(c_0_1298,negated_conjecture,(epred24_0|epred20_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_1292]), c_0_1293])).
cnf(c_0_1299,negated_conjecture,(c1_1(X1)|hskp7|hskp12|hskp13|~c0_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_825]), c_0_1294])).
cnf(c_0_1300,negated_conjecture,(epred28_0|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_1295]), c_0_89]), c_0_88]), ['final']).
cnf(c_0_1301,negated_conjecture,(epred30_0|hskp6|hskp19|hskp20|~hskp24), inference(spm,[status(thm)],[c_0_1296, c_0_656])).
fof(c_0_1302, plain, ((~epred10_0<=>![X1]:((((hskp16|c2_1(X1))|c3_1(X1))|~ndr1_0)|~c1_1(X1)))), introduced(definition)).
cnf(c_0_1303,negated_conjecture,(epred23_0|epred7_0|~hskp19), inference(spm,[status(thm)],[c_0_340, c_0_668])).
cnf(c_0_1304,negated_conjecture,(epred42_0|hskp19), inference(sr,[status(thm)],[c_0_1297, c_0_1232])).
cnf(c_0_1305,negated_conjecture,(epred24_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_818, c_0_1298]), c_0_1124]), ['final']).
cnf(c_0_1306,negated_conjecture,(c1_1(X1)|c0_1(X1)|hskp5|hskp7|hskp19|~c3_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_742]), c_0_176])).
cnf(c_0_1307,negated_conjecture,(c1_1(a1280)|hskp7|hskp12|hskp13|~hskp29), inference(spm,[status(thm)],[c_0_1299, c_0_109])).
cnf(c_0_1308,negated_conjecture,(epred4_0|hskp17|hskp27|~hskp1), inference(spm,[status(thm)],[c_0_852, c_0_299])).
cnf(c_0_1309,negated_conjecture,(epred28_0|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_1300, c_0_324]), ['final']).
cnf(c_0_1310,negated_conjecture,(hskp6|hskp18|hskp19|hskp20|hskp27), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1301, c_0_100]), c_0_1004])).
cnf(c_0_1311,negated_conjecture,(epred10_0|c3_1(X1)|c2_1(X1)|hskp16|~c1_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_1302])).
cnf(c_0_1312,negated_conjecture,(epred23_0|epred7_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1303, c_0_1304]), c_0_1164])).
cnf(c_0_1313,negated_conjecture,(~epred23_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_513, c_0_1305])]), ['final']).
cnf(c_0_1314,negated_conjecture,(hskp5|hskp7|hskp19|~hskp3), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1306, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_1315,negated_conjecture,(hskp7|hskp12|hskp13|~hskp29), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_631, c_0_1307]), c_0_897]), c_0_808])).
cnf(c_0_1316,negated_conjecture,(epred4_0|hskp17|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_1308, c_0_324])).
cnf(c_0_1317,negated_conjecture,(epred5_0|c3_1(a1266)|hskp13|~hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_812, c_0_77]), c_0_91])).
cnf(c_0_1318,negated_conjecture,(epred5_0|c3_1(a1264)|hskp13|~hskp10), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_812, c_0_79]), c_0_81])).
cnf(c_0_1319,negated_conjecture,(epred30_0|hskp19|hskp20|~hskp27), inference(spm,[status(thm)],[c_0_842, c_0_590])).
cnf(c_0_1320,negated_conjecture,(epred28_0|hskp6|hskp19|hskp20|hskp27), inference(spm,[status(thm)],[c_0_1309, c_0_1310])).
cnf(c_0_1321,negated_conjecture,(epred2_0|c3_1(a1268)|~hskp13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_851]), c_0_87]), ['final']).
cnf(c_0_1322,negated_conjecture,(epred10_0|c3_1(X1)|c2_1(X1)|hskp16|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1311, c_0_153])]), ['final']).
cnf(c_0_1323,negated_conjecture,(epred7_0), inference(sr,[status(thm)],[c_0_1312, c_0_1313]), ['final']).
cnf(c_0_1324,negated_conjecture,(hskp2|hskp5|hskp7|hskp19|~hskp18), inference(spm,[status(thm)],[c_0_1314, c_0_190])).
cnf(c_0_1325,negated_conjecture,(epred29_0|hskp20|~c1_1(a1301)|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_36]), c_0_34]), ['final']).
cnf(c_0_1326,negated_conjecture,(hskp7|hskp12|hskp13|hskp17|~hskp27), inference(spm,[status(thm)],[c_0_1315, c_0_179])).
cnf(c_0_1327,negated_conjecture,(epred4_0|hskp7|hskp12|hskp17|hskp27), inference(spm,[status(thm)],[c_0_1316, c_0_849])).
cnf(c_0_1328,negated_conjecture,(epred12_0|hskp14|hskp21|~hskp17), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_531, c_0_813]), c_0_585])).
cnf(c_0_1329,negated_conjecture,(epred19_0|epred12_0|hskp17), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1043, c_0_856]), c_0_1245]), c_0_1040])).
cnf(c_0_1330,negated_conjecture,(epred19_0|epred5_0|hskp13|~hskp12), inference(spm,[status(thm)],[c_0_303, c_0_1317])).
cnf(c_0_1331,negated_conjecture,(epred19_0|epred5_0|hskp13|~hskp10), inference(spm,[status(thm)],[c_0_335, c_0_1318])).
cnf(c_0_1332,negated_conjecture,(epred28_0|hskp6|hskp19|hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1319, c_0_1320]), c_0_1004])).
cnf(c_0_1333,negated_conjecture,(epred17_0|epred2_0|~hskp13), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_1321]), c_0_86]), c_0_87])).
cnf(c_0_1334,negated_conjecture,(epred17_0|hskp13|hskp28|~hskp7), inference(spm,[status(thm)],[c_0_423, c_0_257])).
cnf(c_0_1335,negated_conjecture,(epred38_0|epred35_0|~hskp0), inference(spm,[status(thm)],[c_0_313, c_0_654])).
cnf(c_0_1336,negated_conjecture,(epred29_0|epred10_0|c3_1(X1)|hskp16|hskp20|~c1_1(X1)), inference(spm,[status(thm)],[c_0_418, c_0_1322]), ['final']).
cnf(c_0_1337,negated_conjecture,(~epred44_0|~epred24_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_458, c_0_1212])])).
fof(c_0_1338, plain, ((~epred8_0<=>![X1]:((((hskp1|c3_1(X1))|c1_1(X1))|~ndr1_0)|~c0_1(X1)))), introduced(definition)).
fof(c_0_1339, plain, ((~epred1_0<=>![X1]:((((hskp26|c2_1(X1))|c3_1(X1))|~ndr1_0)|~c1_1(X1)))), introduced(definition)).
cnf(c_0_1340,negated_conjecture,(~epred6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_401, c_0_1323])]), ['final']).
cnf(c_0_1341,negated_conjecture,(epred28_0|~c1_1(a1262)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_1295]), c_0_135]), ['final']).
cnf(c_0_1342,negated_conjecture,(epred44_0|c1_1(a1262)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_271]), c_0_134]), ['final']).
cnf(c_0_1343,negated_conjecture,(epred14_0|c1_1(a1252)|hskp8|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_616]), c_0_46]), ['final']).
cnf(c_0_1344,negated_conjecture,(epred35_0|~epred44_0|~epred43_0), inference(spm,[status(thm)],[c_0_458, c_0_530])).
cnf(c_0_1345,negated_conjecture,(hskp5|hskp7|hskp12|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_808, c_0_1324]), c_0_177])).
cnf(c_0_1346,negated_conjecture,(epred29_0|hskp7|hskp12|hskp20|~hskp24), inference(spm,[status(thm)],[c_0_1325, c_0_181])).
cnf(c_0_1347,negated_conjecture,(hskp7|hskp12|hskp20|~hskp27), inference(spm,[status(thm)],[c_0_808, c_0_590])).
cnf(c_0_1348,negated_conjecture,(epred13_0|hskp7|hskp12|~hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_1326]), c_0_986])).
cnf(c_0_1349,negated_conjecture,(hskp7|hskp12|hskp17|hskp27), inference(sr,[status(thm)],[c_0_1327, c_0_1232])).
cnf(c_0_1350,negated_conjecture,(epred44_0|epred39_0|epred33_0), inference(spm,[status(thm)],[c_0_764, c_0_405])).
cnf(c_0_1351,negated_conjecture,(epred19_0|epred12_0|c0_1(X1)|~c3_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_345]), c_0_379])).
cnf(c_0_1352,negated_conjecture,(epred19_0|epred12_0|hskp21), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_379, c_0_1328]), c_0_1329])).
cnf(c_0_1353,negated_conjecture,(epred19_0|epred2_0|hskp10|hskp21), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_950, c_0_988]), c_0_585]), c_0_108])).
cnf(c_0_1354,negated_conjecture,(epred19_0|epred13_0|~hskp10), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1108, c_0_957]), c_0_1040])).
cnf(c_0_1355,negated_conjecture,(epred19_0|epred13_0|epred5_0|~hskp12), inference(spm,[status(thm)],[c_0_351, c_0_1330])).
cnf(c_0_1356,negated_conjecture,(epred19_0|epred13_0|hskp10|hskp12), inference(spm,[status(thm)],[c_0_951, c_0_1041])).
cnf(c_0_1357,negated_conjecture,(epred19_0|epred13_0|epred5_0|~hskp10), inference(spm,[status(thm)],[c_0_351, c_0_1331])).
cnf(c_0_1358,negated_conjecture,(epred39_0|epred28_0|hskp6|hskp19), inference(spm,[status(thm)],[c_0_764, c_0_1332])).
cnf(c_0_1359,negated_conjecture,(epred23_0|epred19_0|epred12_0|epred11_0), inference(spm,[status(thm)],[c_0_369, c_0_693])).
cnf(c_0_1360,negated_conjecture,(epred45_0|epred38_0|hskp6|~hskp19), inference(spm,[status(thm)],[c_0_408, c_0_1215])).
cnf(c_0_1361,negated_conjecture,(epred17_0|epred2_0|hskp28|~hskp7), inference(spm,[status(thm)],[c_0_1333, c_0_1334])).
cnf(c_0_1362,negated_conjecture,(epred35_0|~c1_1(a1280)|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_17]), c_0_109]), ['final']).
cnf(c_0_1363,negated_conjecture,(epred38_0|epred35_0|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1335, c_0_212]), c_0_347])).
cnf(c_0_1364,negated_conjecture,(epred4_0|c3_1(X1)|hskp6|hskp9|~c0_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_615]), ['final']).
cnf(c_0_1365,negated_conjecture,(epred10_0|hskp16|hskp20|~hskp6), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1336]), c_0_1040]), c_0_60])).
cnf(c_0_1366,negated_conjecture,(~epred44_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1337, c_0_1305])]), ['final']).
cnf(c_0_1367,negated_conjecture,(epred8_0|c1_1(X1)|c3_1(X1)|hskp1|~c0_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_1338])).
cnf(c_0_1368,negated_conjecture,(epred1_0|c3_1(X1)|c2_1(X1)|hskp26|~c1_1(X1)|~ndr1_0), inference(split_equiv,[status(thm)],[c_0_1339])).
cnf(c_0_1369,negated_conjecture,(c3_1(X1)|c2_1(X1)|c0_1(X1)|hskp3), inference(sr,[status(thm)],[c_0_246, c_0_1340]), ['final']).
cnf(c_0_1370,negated_conjecture,(epred44_0|epred28_0|~hskp8), inference(spm,[status(thm)],[c_0_1341, c_0_1342]), ['final']).
cnf(c_0_1371,negated_conjecture,(epred35_0|epred14_0|hskp8|~hskp0), inference(spm,[status(thm)],[c_0_313, c_0_1343])).
cnf(c_0_1372,negated_conjecture,(epred44_0|epred35_0|~hskp0), inference(spm,[status(thm)],[c_0_313, c_0_287])).
cnf(c_0_1373,negated_conjecture,(epred35_0|~epred44_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1344, c_0_1212])]), ['final']).
cnf(c_0_1374,negated_conjecture,(epred42_0|hskp7|hskp12|~hskp18), inference(spm,[status(thm)],[c_0_527, c_0_1345])).
cnf(c_0_1375,negated_conjecture,(hskp7|hskp12|hskp18|hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1346, c_0_100]), c_0_1040]), c_0_1347])).
cnf(c_0_1376,negated_conjecture,(epred13_0|hskp7|hskp12), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1348, c_0_1349]), c_0_986])).
cnf(c_0_1377,negated_conjecture,(epred44_0|epred36_0|epred17_0|hskp20), inference(spm,[status(thm)],[c_0_374, c_0_602])).
cnf(c_0_1378,negated_conjecture,(epred44_0|epred39_0|~epred36_0), inference(spm,[status(thm)],[c_0_552, c_0_1350])).
cnf(c_0_1379,negated_conjecture,(epred19_0|epred12_0|c0_1(a1288)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1351, c_0_105]), c_0_1352])).
cnf(c_0_1380,negated_conjecture,(epred19_0|epred13_0|epred2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_951, c_0_1353]), c_0_1354])).
cnf(c_0_1381,negated_conjecture,(epred19_0|epred13_0|epred5_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1355, c_0_1356]), c_0_1357])).
cnf(c_0_1382,negated_conjecture,(epred28_0|epred24_0|c3_1(X1)|c0_1(X1)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_282, c_0_1295])).
cnf(c_0_1383,negated_conjecture,(epred39_0|epred28_0|hskp19), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1043, c_0_1358]), c_0_1040])).
cnf(c_0_1384,negated_conjecture,(epred23_0|epred19_0|epred12_0|~epred40_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_1359]), c_0_721])).
cnf(c_0_1385,negated_conjecture,(epred45_0|epred38_0|epred29_0|~hskp19), inference(spm,[status(thm)],[c_0_1082, c_0_1360])).
cnf(c_0_1386,negated_conjecture,(epred36_0|epred17_0|epred2_0|~hskp7), inference(spm,[status(thm)],[c_0_583, c_0_1361])).
cnf(c_0_1387,negated_conjecture,(epred31_0|~c3_1(a1268)|~hskp13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_661]), c_0_86])).
cnf(c_0_1388,negated_conjecture,(epred38_0|epred35_0|~hskp29), inference(spm,[status(thm)],[c_0_1362, c_0_632])).
cnf(c_0_1389,negated_conjecture,(epred38_0|epred35_0|hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1363, c_0_84]), c_0_347])).
cnf(c_0_1390,negated_conjecture,(epred42_0|epred31_0|~hskp7), inference(spm,[status(thm)],[c_0_685, c_0_644])).
cnf(c_0_1391,negated_conjecture,(epred41_0|epred4_0|c3_1(X1)|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_270, c_0_615]), ['final']).
cnf(c_0_1392,negated_conjecture,(epred10_0|c3_1(X1)|hskp6|hskp9|hskp16|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_1322]), ['final']).
cnf(c_0_1393,negated_conjecture,(epred10_0|hskp16|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_1322]), c_0_89]), c_0_88]), ['final']).
cnf(c_0_1394,negated_conjecture,(hskp6|hskp9|~hskp16), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1364]), c_0_1232]), c_0_41]), ['final']).
cnf(c_0_1395,negated_conjecture,(c0_1(a1258)|hskp1|hskp27|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_300]), c_0_10]), ['final']).
cnf(c_0_1396,negated_conjecture,(epred10_0|hskp20|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_1365]), c_0_1366])).
cnf(c_0_1397,negated_conjecture,(epred29_0|epred28_0|c3_1(X1)|hskp20|~c1_1(X1)), inference(spm,[status(thm)],[c_0_418, c_0_1295]), ['final']).
fof(c_0_1398, plain, ((~epred9_0<=>![X2]:((c2_1(X2)|c3_1(X2))|c1_1(X2)))), introduced(definition)).
cnf(c_0_1399,negated_conjecture,(epred8_0|c1_1(X1)|c3_1(X1)|hskp1|~c0_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1367, c_0_153])]), ['final']).
cnf(c_0_1400,negated_conjecture,(epred41_0|epred14_0|hskp8|~hskp0), inference(spm,[status(thm)],[c_0_286, c_0_1343]), ['final']).
cnf(c_0_1401,negated_conjecture,(hskp0|hskp11|~hskp8), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_613]), c_0_134]), c_0_135]), ['final']).
cnf(c_0_1402,negated_conjecture,(epred1_0|c3_1(X1)|c2_1(X1)|hskp26|~c1_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1368, c_0_153])]), ['final']).
cnf(c_0_1403,negated_conjecture,(c3_1(X1)|c0_1(X1)|hskp3|hskp20|~c1_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_1369]), c_0_1040]), ['final']).
cnf(c_0_1404,negated_conjecture,(epred41_0|epred2_0|c3_1(X1)|c0_1(X1)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_270, c_0_851]), ['final']).
cnf(c_0_1405,negated_conjecture,(epred41_0|hskp23|hskp24|~hskp0), inference(spm,[status(thm)],[c_0_286, c_0_274]), ['final']).
cnf(c_0_1406,negated_conjecture,(hskp0|hskp11|~c0_1(a1254)|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_613]), c_0_88]), ['final']).
cnf(c_0_1407,negated_conjecture,(epred11_0|c0_1(a1254)|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_236]), c_0_89]), ['final']).
cnf(c_0_1408,negated_conjecture,(epred29_0|epred4_0|c3_1(X1)|hskp20|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_418, c_0_615]), ['final']).
cnf(c_0_1409,negated_conjecture,(epred29_0|epred2_0|c3_1(X1)|c0_1(X1)|hskp20|~c1_1(X1)), inference(spm,[status(thm)],[c_0_418, c_0_851]), ['final']).
cnf(c_0_1410,negated_conjecture,(epred35_0|epred4_0|c3_1(X1)|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_290, c_0_615]), ['final']).
cnf(c_0_1411,negated_conjecture,(epred35_0|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1370, c_0_1371]), c_0_1372]), c_0_1372]), c_0_1373]), ['final']).
cnf(c_0_1412,negated_conjecture,(epred42_0|hskp0|hskp11|~hskp20), inference(spm,[status(thm)],[c_0_1223, c_0_630]), ['final']).
cnf(c_0_1413,negated_conjecture,(epred42_0|hskp7|hskp12|hskp20), inference(spm,[status(thm)],[c_0_1374, c_0_1375])).
cnf(c_0_1414,negated_conjecture,(epred23_0|epred13_0|hskp7), inference(spm,[status(thm)],[c_0_344, c_0_1376])).
cnf(c_0_1415,negated_conjecture,(epred44_0|epred39_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1377]), c_0_1378])).
cnf(c_0_1416,negated_conjecture,(epred19_0|epred12_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_1379]), c_0_1352])).
cnf(c_0_1417,negated_conjecture,(epred19_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1066, c_0_1380]), c_0_1381])).
cnf(c_0_1418,negated_conjecture,(epred18_0|hskp10|~c3_1(a1268)|~hskp13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_267]), c_0_87])).
cnf(c_0_1419,negated_conjecture,(epred28_0|c3_1(a1268)|~hskp13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_1295]), c_0_86]), ['final']).
cnf(c_0_1420,negated_conjecture,(epred28_0|epred24_0|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1382]), c_0_38]), c_0_39])).
cnf(c_0_1421,negated_conjecture,(epred5_0|hskp13|hskp23|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_1270]), c_0_854])).
cnf(c_0_1422,negated_conjecture,(epred28_0|hskp19), inference(sr,[status(thm)],[c_0_1383, c_0_1245])).
cnf(c_0_1423,negated_conjecture,(epred23_0|epred19_0|~epred40_0|~epred13_0), inference(spm,[status(thm)],[c_0_440, c_0_1384])).
cnf(c_0_1424,negated_conjecture,(epred45_0|epred42_0|epred38_0|epred16_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1385, c_0_1007]), c_0_788])).
cnf(c_0_1425,negated_conjecture,(epred17_0|epred2_0|hskp6|hskp17), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1386, c_0_714]), c_0_585])).
cnf(c_0_1426,negated_conjecture,(epred31_0|hskp27|~hskp13), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1387, c_0_169]), c_0_86]), c_0_709]), c_0_87])).
cnf(c_0_1427,negated_conjecture,(epred13_0|hskp6|hskp7|hskp12), inference(spm,[status(thm)],[c_0_944, c_0_915])).
cnf(c_0_1428,negated_conjecture,(epred42_0|epred13_0|hskp26|~hskp10), inference(spm,[status(thm)],[c_0_1175, c_0_235])).
cnf(c_0_1429,negated_conjecture,(epred38_0|epred35_0|hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1388, c_0_1205]), c_0_1389])).
cnf(c_0_1430,negated_conjecture,(epred42_0|epred31_0|epred23_0|~hskp19), inference(spm,[status(thm)],[c_0_1390, c_0_668])).
cnf(c_0_1431,negated_conjecture,(epred42_0|epred31_0|hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1390, c_0_879]), c_0_709])).
cnf(c_0_1432,negated_conjecture,(epred41_0|epred4_0|hskp21|~c1_1(X1)|~c0_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1391]), c_0_585]), ['final']).
cnf(c_0_1433,negated_conjecture,(c1_1(a1262)|hskp6|hskp15|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_541]), c_0_135]), ['final']).
cnf(c_0_1434,negated_conjecture,(epred10_0|hskp6|hskp9|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_1392]), c_0_1393]), c_0_1393]), c_0_1394]), ['final']).
cnf(c_0_1435,negated_conjecture,(epred39_0|hskp1|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_688, c_0_1395]), ['final']).
cnf(c_0_1436,negated_conjecture,(epred10_0|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1396]), c_0_1245]), ['final']).
cnf(c_0_1437,negated_conjecture,(epred28_0|hskp20|~hskp6), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1397]), c_0_1040]), c_0_60])).
cnf(c_0_1438,negated_conjecture,(epred9_0|c1_1(X1)|c3_1(X1)|c2_1(X1)), inference(split_equiv,[status(thm)],[c_0_1398])).
cnf(c_0_1439,negated_conjecture,(epred8_0|c1_1(a1252)|hskp1|~hskp0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1399]), c_0_46]), ['final']).
cnf(c_0_1440,negated_conjecture,(epred41_0|epred14_0|epred4_0|~hskp0), inference(spm,[status(thm)],[c_0_635, c_0_1400]), ['final']).
cnf(c_0_1441,negated_conjecture,(epred14_0|hskp0|hskp11|~hskp16), inference(spm,[status(thm)],[c_0_1401, c_0_636]), ['final']).
cnf(c_0_1442,negated_conjecture,(epred41_0|epred28_0|c3_1(X1)|~c1_1(X1)), inference(spm,[status(thm)],[c_0_270, c_0_1295]), ['final']).
cnf(c_0_1443,negated_conjecture,(epred41_0|epred1_0|c3_1(X1)|hskp26|~c1_1(X1)), inference(spm,[status(thm)],[c_0_270, c_0_1402]), ['final']).
cnf(c_0_1444,negated_conjecture,(epred10_0|hskp16|~c1_1(a1272)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_1322]), c_0_56]), ['final']).
cnf(c_0_1445,negated_conjecture,(epred41_0|epred6_0|c3_1(X1)|c0_1(X1)|hskp3|~c1_1(X1)), inference(spm,[status(thm)],[c_0_270, c_0_246])).
cnf(c_0_1446,negated_conjecture,(hskp3|hskp20|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1403]), c_0_38]), c_0_39]), ['final']).
cnf(c_0_1447,negated_conjecture,(epred41_0|epred2_0|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1404]), c_0_38]), c_0_39]), ['final']).
cnf(c_0_1448,negated_conjecture,(epred2_0|c0_1(a1254)|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_851]), c_0_88]), ['final']).
cnf(c_0_1449,negated_conjecture,(c0_1(a1258)|hskp0|hskp26|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_201]), c_0_10]), ['final']).
cnf(c_0_1450,negated_conjecture,(epred41_0|epred11_0|hskp24|~hskp0), inference(spm,[status(thm)],[c_0_487, c_0_1405]), ['final']).
cnf(c_0_1451,negated_conjecture,(epred11_0|hskp0|hskp11|~hskp1), inference(spm,[status(thm)],[c_0_1406, c_0_1407]), ['final']).
cnf(c_0_1452,negated_conjecture,(hskp3|~hskp15), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_1369]), c_0_57]), c_0_56]), ['final']).
cnf(c_0_1453,negated_conjecture,(hskp20|~c0_1(a1260)|~hskp6), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1408]), c_0_1040]), c_0_1232]), c_0_60]), ['final']).
cnf(c_0_1454,negated_conjecture,(epred2_0|c0_1(a1260)|hskp20|~hskp6), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1409]), c_0_1040]), c_0_60])).
cnf(c_0_1455,negated_conjecture,(epred35_0|epred4_0|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1410]), c_0_15]), c_0_16]), ['final']).
cnf(c_0_1456,negated_conjecture,(epred35_0|hskp11|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1411, c_0_1412]), c_0_539])).
cnf(c_0_1457,negated_conjecture,(c0_1(a1254)|hskp3|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_1369]), c_0_88])).
cnf(c_0_1458,negated_conjecture,(epred29_0|hskp20|hskp23|hskp24|~hskp0), inference(spm,[status(thm)],[c_0_1239, c_0_274])).
cnf(c_0_1459,negated_conjecture,(epred42_0|hskp7|hskp12), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1230, c_0_1413]), c_0_1232])).
cnf(c_0_1460,negated_conjecture,(epred23_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1175, c_0_1414]), c_0_1164])).
cnf(c_0_1461,negated_conjecture,(epred44_0|hskp6|hskp17), inference(spm,[status(thm)],[c_0_374, c_0_83])).
cnf(c_0_1462,negated_conjecture,(epred44_0|epred39_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_1415])).
cnf(c_0_1463,negated_conjecture,(epred19_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_1416]), c_0_1417]), ['final']).
cnf(c_0_1464,negated_conjecture,(epred28_0|epred18_0|hskp10|~hskp13), inference(spm,[status(thm)],[c_0_1418, c_0_1419])).
cnf(c_0_1465,negated_conjecture,(epred28_0|epred24_0|epred5_0|hskp13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1420, c_0_1421]), c_0_1422])).
cnf(c_0_1466,negated_conjecture,(epred23_0|epred19_0|~epred40_0), inference(spm,[status(thm)],[c_0_1423, c_0_1099])).
cnf(c_0_1467,negated_conjecture,(epred45_0|epred42_0|epred38_0|~epred17_0), inference(spm,[status(thm)],[c_0_504, c_0_1424])).
cnf(c_0_1468,negated_conjecture,(epred24_0|epred17_0|epred2_0|hskp17), inference(spm,[status(thm)],[c_0_681, c_0_1425])).
cnf(c_0_1469,negated_conjecture,(epred31_0|epred5_0|hskp18|hskp27), inference(spm,[status(thm)],[c_0_1426, c_0_900])).
cnf(c_0_1470,negated_conjecture,(epred23_0|epred13_0|hskp6|hskp7), inference(spm,[status(thm)],[c_0_344, c_0_1427])).
cnf(c_0_1471,negated_conjecture,(epred38_0|epred15_0|~hskp4), inference(spm,[status(thm)],[c_0_733, c_0_1022])).
cnf(c_0_1472,negated_conjecture,(epred45_0|epred44_0|~epred36_0), inference(spm,[status(thm)],[c_0_552, c_0_438])).
cnf(c_0_1473,negated_conjecture,(epred35_0|hskp7|hskp16|hskp20), inference(spm,[status(thm)],[c_0_413, c_0_1])).
cnf(c_0_1474,negated_conjecture,(epred42_0|epred35_0|epred13_0|~hskp10), inference(spm,[status(thm)],[c_0_347, c_0_1428])).
cnf(c_0_1475,negated_conjecture,(epred42_0|epred23_0|epred13_0|hskp10), inference(spm,[status(thm)],[c_0_1175, c_0_378])).
cnf(c_0_1476,negated_conjecture,(epred38_0|epred35_0|epred11_0), inference(spm,[status(thm)],[c_0_525, c_0_1429])).
cnf(c_0_1477,negated_conjecture,(epred41_0|epred38_0|epred35_0), inference(spm,[status(thm)],[c_0_363, c_0_1429])).
cnf(c_0_1478,negated_conjecture,(epred31_0|epred24_0|c0_1(X1)|~c1_1(X1)|~c3_1(X1)), inference(spm,[status(thm)],[c_0_282, c_0_661])).
cnf(c_0_1479,negated_conjecture,(epred42_0|epred31_0|epred23_0), inference(spm,[status(thm)],[c_0_1430, c_0_1431])).
cnf(c_0_1480,negated_conjecture,(epred41_0|hskp21|~c1_1(a1262)|~hskp8), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1432, c_0_134]), c_0_1232])).
cnf(c_0_1481,negated_conjecture,(epred39_0|c1_1(a1262)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_662]), c_0_134]), ['final']).
cnf(c_0_1482,negated_conjecture,(epred3_0|c1_1(a1272)|hskp26|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_782]), c_0_57]), ['final']).
cnf(c_0_1483,negated_conjecture,(epred28_0|hskp6|hskp15|~hskp8), inference(spm,[status(thm)],[c_0_1341, c_0_1433]), ['final']).
cnf(c_0_1484,negated_conjecture,(c1_1(a1272)|hskp1|hskp27|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_300]), c_0_57]), ['final']).
cnf(c_0_1485,negated_conjecture,(epred45_0|hskp1|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_419, c_0_1395]), ['final']).
cnf(c_0_1486,negated_conjecture,(epred10_0|hskp6|hskp9|hskp27|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1434, c_0_1435]), c_0_1245])).
cnf(c_0_1487,negated_conjecture,(epred10_0|hskp9|hskp21|~c1_1(X1)|~c0_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1392]), c_0_585]), c_0_1394]), c_0_1436]), ['final']).
cnf(c_0_1488,negated_conjecture,(epred28_0|c3_1(X1)|hskp6|hskp9|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_1295]), ['final']).
cnf(c_0_1489,negated_conjecture,(epred28_0|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1437]), c_0_1245]), ['final']).
cnf(c_0_1490,negated_conjecture,(epred1_0|hskp26|~c1_1(a1272)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_1402]), c_0_56]), ['final']).
cnf(c_0_1491,negated_conjecture,(epred29_0|epred1_0|c3_1(X1)|hskp20|hskp26|~c1_1(X1)), inference(spm,[status(thm)],[c_0_418, c_0_1402]), ['final']).
cnf(c_0_1492,negated_conjecture,(epred9_0|c1_1(a1272)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_1438]), c_0_56]), ['final']).
cnf(c_0_1493,negated_conjecture,(epred28_0|~c1_1(a1272)|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_1295]), c_0_56]), ['final']).
cnf(c_0_1494,negated_conjecture,(epred15_0|~hskp3), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_50]), c_0_51]), c_0_52]), ['final']).
cnf(c_0_1495,negated_conjecture,(epred10_0|hskp16|~c1_1(a1262)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_1322]), c_0_135]), ['final']).
cnf(c_0_1496,negated_conjecture,(epred35_0|epred34_0|hskp12|~c1_1(X1)|~c3_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_290, c_0_447])).
cnf(c_0_1497,negated_conjecture,(hskp6|hskp23|hskp24|hskp29|~hskp19), inference(spm,[status(thm)],[c_0_461, c_0_745])).
cnf(c_0_1498,negated_conjecture,(epred41_0|epred8_0|hskp1|~hskp0), inference(spm,[status(thm)],[c_0_286, c_0_1439]), ['final']).
cnf(c_0_1499,negated_conjecture,(~epred42_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1177, c_0_1305])]), ['final']).
cnf(c_0_1500,negated_conjecture,(epred41_0|epred10_0|c3_1(X1)|hskp16|~c1_1(X1)), inference(spm,[status(thm)],[c_0_270, c_0_1322]), ['final']).
cnf(c_0_1501,negated_conjecture,(epred44_0|hskp6|hskp29|~hskp19), inference(spm,[status(thm)],[c_0_1190, c_0_745])).
cnf(c_0_1502,negated_conjecture,(epred41_0|epred4_0|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1391]), c_0_15]), c_0_16]), ['final']).
cnf(c_0_1503,negated_conjecture,(epred41_0|epred14_0|hskp11|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1440, c_0_1441]), c_0_1232])).
cnf(c_0_1504,negated_conjecture,(epred44_0|hskp6|hskp28|~hskp19), inference(spm,[status(thm)],[c_0_1190, c_0_462])).
cnf(c_0_1505,negated_conjecture,(epred34_0|hskp12|~c3_1(a1287)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_447]), c_0_95]), ['final']).
cnf(c_0_1506,negated_conjecture,(epred41_0|epred28_0|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1442]), c_0_38]), ['final']).
cnf(c_0_1507,negated_conjecture,(epred28_0|hskp20|~hskp23), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1397]), c_0_1040]), c_0_38]), ['final']).
cnf(c_0_1508,negated_conjecture,(epred41_0|~c1_1(a1301)|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_36]), c_0_34]), ['final']).
cnf(c_0_1509,negated_conjecture,(epred41_0|epred1_0|hskp26|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1443]), c_0_38]), ['final']).
cnf(c_0_1510,negated_conjecture,(epred42_0|hskp23|hskp24|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1223, c_0_170]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_1511,negated_conjecture,(epred10_0|hskp0|hskp16|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1444, c_0_294]), ['final']).
cnf(c_0_1512,negated_conjecture,(epred41_0|hskp3|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1445]), c_0_1340]), c_0_38]), c_0_39]), ['final']).
cnf(c_0_1513,negated_conjecture,(hskp3|hskp20|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_1446, c_0_183]), ['final']).
cnf(c_0_1514,negated_conjecture,(epred41_0|epred2_0|hskp24|~hskp0), inference(spm,[status(thm)],[c_0_1447, c_0_1405]), ['final']).
cnf(c_0_1515,negated_conjecture,(epred2_0|hskp20|~hskp23), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1409]), c_0_1040]), c_0_38]), c_0_39]), ['final']).
cnf(c_0_1516,negated_conjecture,(epred10_0|hskp6|hskp9|~c1_1(a1272)|~hskp15), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_1392]), c_0_1444]), c_0_1394]), ['final']).
cnf(c_0_1517,negated_conjecture,(epred2_0|hskp0|hskp11|~hskp1), inference(spm,[status(thm)],[c_0_1406, c_0_1448]), ['final']).
cnf(c_0_1518,negated_conjecture,(epred39_0|hskp0|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_688, c_0_1449]), ['final']).
cnf(c_0_1519,negated_conjecture,(epred41_0|epred11_0|hskp11|hskp24|~hskp1), inference(spm,[status(thm)],[c_0_1450, c_0_1451]), ['final']).
cnf(c_0_1520,negated_conjecture,(hskp11|hskp12|~hskp15), inference(spm,[status(thm)],[c_0_189, c_0_1452]), ['final']).
cnf(c_0_1521,negated_conjecture,(epred1_0|hskp26|~c1_1(a1262)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_1402]), c_0_135]), ['final']).
cnf(c_0_1522,negated_conjecture,(epred11_0|hskp20|~hskp6), inference(spm,[status(thm)],[c_0_1453, c_0_263])).
cnf(c_0_1523,negated_conjecture,(epred2_0|hskp20|~hskp6), inference(spm,[status(thm)],[c_0_1453, c_0_1454])).
cnf(c_0_1524,negated_conjecture,(hskp20|~c1_1(a1252)|~hskp0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1408]), c_0_1040]), c_0_1232]), c_0_46])).
cnf(c_0_1525,negated_conjecture,(epred35_0|epred21_0|~c1_1(X1)|~c3_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_290, c_0_341])).
cnf(c_0_1526,negated_conjecture,(epred35_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1455, c_0_1456]), c_0_1232]), ['final']).
cnf(c_0_1527,negated_conjecture,(hskp3|~hskp1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_852, c_0_1457]), c_0_1232]), ['final']).
cnf(c_0_1528,negated_conjecture,(epred8_0|hskp1|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1399]), c_0_41]), c_0_42]), ['final']).
cnf(c_0_1529,negated_conjecture,(epred11_0|hskp20|hskp24|~hskp0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_1458]), c_0_1040])).
cnf(c_0_1530,negated_conjecture,(epred42_0|~c3_1(a1279)|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_511]), c_0_55]), ['final']).
cnf(c_0_1531,negated_conjecture,(c3_1(a1279)|hskp3|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_1369]), c_0_54])).
cnf(c_0_1532,negated_conjecture,(epred23_0|hskp7), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_1459]), c_0_1164])).
cnf(c_0_1533,negated_conjecture,(epred13_0), inference(sr,[status(thm)],[c_0_1460, c_0_1313]), ['final']).
cnf(c_0_1534,negated_conjecture,(epred44_0|epred39_0|hskp17), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1043, c_0_1461]), c_0_1462])).
cnf(c_0_1535,negated_conjecture,(~epred18_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_338, c_0_1463])]), ['final']).
cnf(c_0_1536,negated_conjecture,(epred28_0|epred24_0|epred18_0|epred5_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1464, c_0_1465]), c_0_305])).
cnf(c_0_1537,negated_conjecture,(epred23_0|~epred42_0|~epred40_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_1466]), c_0_1127])).
cnf(c_0_1538,negated_conjecture,(epred45_0|epred42_0|epred38_0|epred23_0), inference(spm,[status(thm)],[c_0_1467, c_0_725])).
cnf(c_0_1539,negated_conjecture,(epred24_0|epred17_0|epred2_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_542, c_0_1468]), c_0_585])).
cnf(c_0_1540,negated_conjecture,(epred31_0|epred28_0|epred5_0|hskp27), inference(spm,[status(thm)],[c_0_1309, c_0_1469])).
cnf(c_0_1541,negated_conjecture,(epred23_0|epred15_0|epred13_0|hskp6), inference(spm,[status(thm)],[c_0_907, c_0_1470])).
cnf(c_0_1542,negated_conjecture,(epred38_0|epred26_0|epred15_0|~hskp19), inference(spm,[status(thm)],[c_0_1471, c_0_1059])).
cnf(c_0_1543,negated_conjecture,(epred44_0|epred41_0|hskp17), inference(spm,[status(thm)],[c_0_363, c_0_1461])).
cnf(c_0_1544,negated_conjecture,(epred44_0|epred36_0|epred21_0|epred17_0), inference(spm,[status(thm)],[c_0_404, c_0_1377])).
cnf(c_0_1545,negated_conjecture,(epred44_0|~epred36_0), inference(spm,[status(thm)],[c_0_552, c_0_460]), ['final']).
cnf(c_0_1546,negated_conjecture,(epred45_0|epred44_0|epred17_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_1377]), c_0_1472])).
cnf(c_0_1547,negated_conjecture,(epred39_0|epred33_0|epred29_0|~hskp10), inference(spm,[status(thm)],[c_0_764, c_0_477])).
cnf(c_0_1548,negated_conjecture,(epred44_0|epred35_0|hskp7|hskp20), inference(spm,[status(thm)],[c_0_374, c_0_1473])).
cnf(c_0_1549,negated_conjecture,(epred42_0|epred17_0|hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_667, c_0_235]), c_0_625])).
cnf(c_0_1550,negated_conjecture,(epred35_0|epred23_0|epred13_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1474, c_0_1475]), c_0_514])).
cnf(c_0_1551,negated_conjecture,(epred38_0|epred35_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_1476]), c_0_470]), c_0_1477])).
cnf(c_0_1552,negated_conjecture,(epred44_0|epred24_0|hskp10|hskp17), inference(spm,[status(thm)],[c_0_367, c_0_1461])).
cnf(c_0_1553,negated_conjecture,(epred31_0|epred24_0|~hskp21), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1478, c_0_105]), c_0_106]), c_0_107])).
cnf(c_0_1554,negated_conjecture,(epred42_0|epred31_0|~epred24_0), inference(spm,[status(thm)],[c_0_513, c_0_1479])).
cnf(c_0_1555,negated_conjecture,(epred41_0|hskp21|~hskp8), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1480, c_0_1481]), c_0_1245]), ['final']).
cnf(c_0_1556,negated_conjecture,(epred10_0|epred3_0|hskp16|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1444, c_0_1482]), ['final']).
cnf(c_0_1557,negated_conjecture,(epred42_0|hskp6|hskp15|~hskp20), inference(spm,[status(thm)],[c_0_1223, c_0_549]), ['final']).
cnf(c_0_1558,negated_conjecture,(epred28_0|epred14_0|hskp6|hskp15|~hskp16), inference(spm,[status(thm)],[c_0_1483, c_0_636]), ['final']).
cnf(c_0_1559,negated_conjecture,(epred42_0|epred8_0|hskp1|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1223, c_0_1399]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_1560,negated_conjecture,(epred41_0|epred11_0|hskp21|~hskp1), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1432, c_0_1407]), c_0_1232]), c_0_89]), ['final']).
cnf(c_0_1561,negated_conjecture,(epred11_0|hskp1|hskp27|~hskp15), inference(spm,[status(thm)],[c_0_255, c_0_1484]), ['final']).
cnf(c_0_1562,negated_conjecture,(epred41_0|hskp21|~hskp27), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1432, c_0_32]), c_0_1232]), c_0_33]), ['final']).
cnf(c_0_1563,negated_conjecture,(epred10_0|hskp9|hskp27|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1434, c_0_1485]), c_0_1486]), c_0_1436]), ['final']).
cnf(c_0_1564,negated_conjecture,(epred10_0|hskp9|hskp21|~hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1487, c_0_32]), c_0_33]), ['final']).
cnf(c_0_1565,negated_conjecture,(epred10_0|hskp9|hskp21|~c1_1(a1252)|~hskp0), inference(spm,[status(thm)],[c_0_1487, c_0_46])).
cnf(c_0_1566,negated_conjecture,(epred28_0|hskp9|hskp21|~c1_1(X1)|~c0_1(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1488]), c_0_585]), c_0_1489]), ['final']).
cnf(c_0_1567,negated_conjecture,(epred10_0|hskp9|hskp21|~c1_1(a1262)|~hskp8), inference(spm,[status(thm)],[c_0_1487, c_0_134])).
cnf(c_0_1568,negated_conjecture,(epred35_0|epred4_0|hskp21|~c1_1(X1)|~c0_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1410]), c_0_585])).
cnf(c_0_1569,negated_conjecture,(epred10_0|c3_1(a1291)|hskp16|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_1322]), c_0_128]), ['final']).
cnf(c_0_1570,negated_conjecture,(epred3_0|epred1_0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1490, c_0_1482]), ['final']).
cnf(c_0_1571,negated_conjecture,(epred1_0|hskp20|hskp26|~hskp6), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1491]), c_0_1040]), c_0_60])).
cnf(c_0_1572,negated_conjecture,(epred9_0|epred1_0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1490, c_0_1492]), ['final']).
cnf(c_0_1573,negated_conjecture,(epred1_0|hskp0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1490, c_0_294]), ['final']).
cnf(c_0_1574,negated_conjecture,(c3_1(a1279)|hskp6|hskp15|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_541]), c_0_55]), ['final']).
cnf(c_0_1575,negated_conjecture,(epred28_0|hskp0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1493, c_0_294]), ['final']).
cnf(c_0_1576,negated_conjecture,(epred11_0|hskp0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_255, c_0_294]), ['final']).
cnf(c_0_1577,negated_conjecture,(epred28_0|hskp9|~hskp0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_1488]), c_0_46]), c_0_652]), c_0_1489]), ['final']).
cnf(c_0_1578,negated_conjecture,(hskp6|hskp9|~hskp11), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1364]), c_0_1232]), c_0_15]), ['final']).
cnf(c_0_1579,negated_conjecture,(epred15_0|~hskp15), inference(spm,[status(thm)],[c_0_1494, c_0_1452]), ['final']).
cnf(c_0_1580,negated_conjecture,(epred1_0|hskp26|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_1402]), c_0_89]), c_0_88]), ['final']).
cnf(c_0_1581,negated_conjecture,(epred10_0|hskp6|hskp15|hskp16|~hskp8), inference(spm,[status(thm)],[c_0_1495, c_0_1433]), ['final']).
cnf(c_0_1582,negated_conjecture,(epred21_0|epred14_0|hskp8|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_616]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_1583,negated_conjecture,(epred21_0|hskp6|hskp15|~hskp20), inference(spm,[status(thm)],[c_0_373, c_0_549]), ['final']).
cnf(c_0_1584,negated_conjecture,(epred1_0|hskp20|hskp26|~hskp23), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1491]), c_0_1040]), c_0_38]), ['final']).
cnf(c_0_1585,negated_conjecture,(epred35_0|epred34_0|hskp12|~c1_1(X1)|~c0_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1496, c_0_1410]), c_0_1232]), ['final']).
cnf(c_0_1586,negated_conjecture,(hskp6|hskp23|hskp29), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1497, c_0_1258])]), c_0_656])).
cnf(c_0_1587,negated_conjecture,(epred41_0|epred8_0|hskp1|hskp11|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1498, c_0_1412]), c_0_1499]), ['final']).
cnf(c_0_1588,negated_conjecture,(epred41_0|epred10_0|hskp16|~hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1500]), c_0_16]), ['final']).
cnf(c_0_1589,negated_conjecture,(epred44_0|hskp6|hskp29), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1501, c_0_1258])])).
cnf(c_0_1590,negated_conjecture,(epred39_0|epred10_0|hskp16|~hskp8), inference(spm,[status(thm)],[c_0_1495, c_0_1481]), ['final']).
cnf(c_0_1591,negated_conjecture,(epred41_0|epred14_0|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1503]), c_0_1232]), ['final']).
cnf(c_0_1592,negated_conjecture,(epred44_0|hskp6|hskp28), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1504, c_0_1258])])).
cnf(c_0_1593,negated_conjecture,(epred35_0|epred1_0|c3_1(X1)|~c1_1(X1)|~c0_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_1402]), c_0_347]), ['final']).
cnf(c_0_1594,negated_conjecture,(epred34_0|epred8_0|hskp1|hskp12|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1505, c_0_1399]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_1595,negated_conjecture,(epred41_0|epred28_0|hskp24|~hskp0), inference(spm,[status(thm)],[c_0_1506, c_0_1405]), ['final']).
cnf(c_0_1596,negated_conjecture,(epred41_0|epred28_0|~hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1442]), c_0_16]), ['final']).
cnf(c_0_1597,negated_conjecture,(epred29_0|hskp20|hskp28|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1325, c_0_464]), c_0_1003]), ['final']).
cnf(c_0_1598,negated_conjecture,(epred28_0|hskp20|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_1507, c_0_183]), ['final']).
cnf(c_0_1599,negated_conjecture,(epred41_0|hskp28|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1508, c_0_464]), c_0_363]), ['final']).
cnf(c_0_1600,negated_conjecture,(epred41_0|epred1_0|hskp24|hskp26|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1509, c_0_1510]), c_0_1499])).
cnf(c_0_1601,negated_conjecture,(epred41_0|epred10_0|hskp16|hskp26|~hskp15), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_1511]), c_0_1366]), ['final']).
cnf(c_0_1602,negated_conjecture,(epred41_0|hskp3|hskp24|~hskp0), inference(spm,[status(thm)],[c_0_1512, c_0_1405]), ['final']).
cnf(c_0_1603,negated_conjecture,(epred41_0|hskp3|hskp24|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1512, c_0_1510]), c_0_1499]), ['final']).
cnf(c_0_1604,negated_conjecture,(epred15_0|hskp20|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_1494, c_0_1513]), ['final']).
cnf(c_0_1605,negated_conjecture,(epred41_0|hskp3|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_1512, c_0_183]), ['final']).
cnf(c_0_1606,negated_conjecture,(epred41_0|epred2_0|hskp11|hskp24|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1514, c_0_1412]), c_0_1499])).
cnf(c_0_1607,negated_conjecture,(epred2_0|hskp20|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_1515, c_0_183]), ['final']).
cnf(c_0_1608,negated_conjecture,(epred10_0|hskp1|hskp6|hskp9|hskp27|~hskp15), inference(spm,[status(thm)],[c_0_1516, c_0_1484])).
cnf(c_0_1609,negated_conjecture,(epred41_0|epred10_0|hskp16|~hskp23), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1500]), c_0_38]), ['final']).
cnf(c_0_1610,negated_conjecture,(epred11_0|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_487, c_0_183]), ['final']).
cnf(c_0_1611,negated_conjecture,(epred44_0|epred41_0|epred11_0|hskp11|~hskp1), inference(spm,[status(thm)],[c_0_309, c_0_1451]), ['final']).
cnf(c_0_1612,negated_conjecture,(epred44_0|epred41_0|epred2_0|hskp11|~hskp1), inference(spm,[status(thm)],[c_0_309, c_0_1517]), ['final']).
cnf(c_0_1613,negated_conjecture,(epred41_0|epred8_0|hskp1|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1498, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_1614,negated_conjecture,(epred41_0|epred2_0|hskp11|hskp24|~hskp1), inference(spm,[status(thm)],[c_0_1514, c_0_1517]), ['final']).
cnf(c_0_1615,negated_conjecture,(epred41_0|epred11_0|hskp11|hskp24|hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1519, c_0_324]), c_0_100])).
cnf(c_0_1616,negated_conjecture,(epred11_0|epred4_0|~hskp1), inference(spm,[status(thm)],[c_0_852, c_0_1407]), ['final']).
cnf(c_0_1617,negated_conjecture,(epred4_0|epred2_0|~hskp1), inference(spm,[status(thm)],[c_0_852, c_0_1448]), ['final']).
cnf(c_0_1618,negated_conjecture,(hskp20|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1408]), c_0_1040]), c_0_1232]), c_0_15]), c_0_16]), ['final']).
cnf(c_0_1619,negated_conjecture,(epred41_0|hskp12|~hskp15), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1520]), c_0_1232]), ['final']).
cnf(c_0_1620,negated_conjecture,(epred10_0|epred9_0|hskp16|~hskp15), inference(spm,[status(thm)],[c_0_1444, c_0_1492]), ['final']).
cnf(c_0_1621,negated_conjecture,(hskp3|hskp20|~hskp6), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_1403]), c_0_60]), c_0_1453]), ['final']).
cnf(c_0_1622,negated_conjecture,(epred39_0|epred1_0|hskp26|~hskp8), inference(spm,[status(thm)],[c_0_1521, c_0_1481]), ['final']).
cnf(c_0_1623,negated_conjecture,(epred34_0|epred4_0|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_1505, c_0_881]), ['final']).
cnf(c_0_1624,negated_conjecture,(hskp6|hskp9|~c0_1(a1254)|~hskp1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_1364]), c_0_1232]), ['final']).
cnf(c_0_1625,negated_conjecture,(epred11_0|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1522]), c_0_1245]), ['final']).
cnf(c_0_1626,negated_conjecture,(epred2_0|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1523]), c_0_1245]), ['final']).
cnf(c_0_1627,negated_conjecture,(hskp20|~hskp0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1524, c_0_287]), c_0_1366]), ['final']).
cnf(c_0_1628,negated_conjecture,(epred29_0|c3_1(a1250)|hskp20|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_47]), c_0_48]), ['final']).
cnf(c_0_1629,negated_conjecture,(epred35_0|epred21_0|~c1_1(X1)|~c0_1(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1525, c_0_1408]), c_0_1040]), c_0_1232]), c_0_1526]), ['final']).
cnf(c_0_1630,negated_conjecture,(epred8_0|hskp3|~hskp16), inference(spm,[status(thm)],[c_0_1527, c_0_1528]), ['final']).
cnf(c_0_1631,negated_conjecture,(epred11_0|hskp11|hskp20|hskp24|~hskp1), inference(spm,[status(thm)],[c_0_1529, c_0_1451])).
cnf(c_0_1632,negated_conjecture,(hskp20|~c0_1(a1254)|~hskp1), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_1408]), c_0_1040]), c_0_1232]), c_0_89]), ['final']).
cnf(c_0_1633,negated_conjecture,(hskp3|~hskp18), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1530, c_0_1531]), c_0_1499]), ['final']).
cnf(c_0_1634,negated_conjecture,(hskp20|~c1_1(a1262)|~hskp8), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_1408]), c_0_1040]), c_0_1232]), c_0_134])).
cnf(c_0_1635,negated_conjecture,(hskp7), inference(sr,[status(thm)],[c_0_1532, c_0_1313]), ['final']).
cnf(c_0_1636,negated_conjecture,(~epred12_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_440, c_0_1533])]), ['final']).
cnf(c_0_1637,negated_conjecture,(epred44_0|hskp17), inference(sr,[status(thm)],[c_0_1534, c_0_1245])).
cnf(c_0_1638,negated_conjecture,(epred23_0|hskp10), inference(sr,[status(thm)],[c_0_409, c_0_1535])).
cnf(c_0_1639,negated_conjecture,(epred26_0|hskp4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1059, c_0_1258])])).
cnf(c_0_1640,negated_conjecture,(~epred26_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1061, c_0_1305])]), ['final']).
cnf(c_0_1641,negated_conjecture,(epred44_0|epred5_0|hskp13|~hskp19), inference(spm,[status(thm)],[c_0_1190, c_0_1270])).
cnf(c_0_1642,negated_conjecture,(epred28_0|epred24_0|epred5_0), inference(sr,[status(thm)],[c_0_1536, c_0_1535])).
cnf(c_0_1643,negated_conjecture,(epred34_0|hskp12|~c0_1(a1258)|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_447]), c_0_9]), ['final']).
cnf(c_0_1644,negated_conjecture,(c1_1(a1262)|hskp12|hskp13|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_825]), c_0_134]), ['final']).
cnf(c_0_1645,negated_conjecture,(epred39_0|epred28_0|~hskp8), inference(spm,[status(thm)],[c_0_1341, c_0_1481]), ['final']).
cnf(c_0_1646,negated_conjecture,(epred35_0|epred11_0|hskp11|~hskp1), inference(spm,[status(thm)],[c_0_1411, c_0_1451]), ['final']).
cnf(c_0_1647,negated_conjecture,(epred35_0|epred2_0|hskp11|~hskp1), inference(spm,[status(thm)],[c_0_1411, c_0_1517]), ['final']).
cnf(c_0_1648,negated_conjecture,(epred21_0|hskp23|hskp24|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_170]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_1649,negated_conjecture,(epred42_0|epred14_0|hskp8|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1223, c_0_616]), c_0_95]), c_0_103])).
cnf(c_0_1650,negated_conjecture,(epred24_0|~epred45_0|~epred41_0), inference(spm,[status(thm)],[c_0_435, c_0_1124])).
cnf(c_0_1651,negated_conjecture,(~epred42_0|~epred40_0|~epred24_0), inference(spm,[status(thm)],[c_0_513, c_0_1537])).
cnf(c_0_1652,negated_conjecture,(epred45_0|epred42_0|epred38_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_1538]), c_0_913])).
cnf(c_0_1653,negated_conjecture,(epred24_0|epred17_0|~epred5_0), inference(spm,[status(thm)],[c_0_1066, c_0_1539])).
cnf(c_0_1654,negated_conjecture,(epred31_0|epred28_0|epred17_0|epred5_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1024, c_0_1540]), c_0_585])).
cnf(c_0_1655,negated_conjecture,(epred45_0|epred31_0|epred24_0|epred18_0), inference(spm,[status(thm)],[c_0_796, c_0_699])).
cnf(c_0_1656,negated_conjecture,(epred45_0|epred31_0|epred24_0|epred19_0), inference(spm,[status(thm)],[c_0_796, c_0_728])).
cnf(c_0_1657,negated_conjecture,(epred17_0|~epred25_0), inference(spm,[status(thm)],[c_0_646, c_0_1176])).
cnf(c_0_1658,negated_conjecture,(epred41_0|epred23_0|epred15_0|epred13_0), inference(spm,[status(thm)],[c_0_363, c_0_1541])).
cnf(c_0_1659,negated_conjecture,(epred42_0|epred41_0|epred38_0|epred26_0|epred15_0), inference(spm,[status(thm)],[c_0_1542, c_0_1017])).
cnf(c_0_1660,negated_conjecture,(epred38_0|~epred26_0), inference(spm,[status(thm)],[c_0_1061, c_0_913])).
cnf(c_0_1661,negated_conjecture,(epred44_0|epred41_0|epred36_0|epred17_0), inference(spm,[status(thm)],[c_0_542, c_0_1543])).
cnf(c_0_1662,negated_conjecture,(epred44_0|epred21_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_1544]), c_0_1545])).
cnf(c_0_1663,negated_conjecture,(epred45_0|epred44_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_1546])).
cnf(c_0_1664,negated_conjecture,(epred39_0|epred33_0|epred29_0), inference(spm,[status(thm)],[c_0_1547, c_0_457])).
cnf(c_0_1665,negated_conjecture,(epred44_0|epred35_0|epred21_0|hskp7), inference(spm,[status(thm)],[c_0_404, c_0_1548])).
cnf(c_0_1666,negated_conjecture,(epred42_0|epred35_0|epred17_0), inference(spm,[status(thm)],[c_0_347, c_0_1549])).
cnf(c_0_1667,negated_conjecture,(epred45_0|c1_1(a1262)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_377]), c_0_134]), ['final']).
cnf(c_0_1668,negated_conjecture,(epred9_0|c1_1(a1262)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_1438]), c_0_135]), ['final']).
cnf(c_0_1669,negated_conjecture,(epred35_0|epred23_0|~epred39_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_866, c_0_1550]), c_0_1551])).
cnf(c_0_1670,negated_conjecture,(epred9_0|c3_1(a1287)|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_1438]), c_0_103]), ['final']).
cnf(c_0_1671,negated_conjecture,(epred44_0|epred31_0|epred24_0|hskp10), inference(spm,[status(thm)],[c_0_709, c_0_1552])).
fof(c_0_1672, plain, ((~epred27_0<=>![X2]:((((hskp4|c1_1(X2))|~ndr1_0)|~c2_1(X2))|~c3_1(X2)))), introduced(definition)).
cnf(c_0_1673,negated_conjecture,(epred42_0|epred31_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1553, c_0_608]), c_0_305]), c_0_1554])).
cnf(c_0_1674,negated_conjecture,(epred41_0|epred14_0|hskp21|~hskp0), inference(spm,[status(thm)],[c_0_1555, c_0_1400]), ['final']).
cnf(c_0_1675,negated_conjecture,(epred41_0|epred8_0|hskp1|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_1498, c_0_212]), ['final']).
cnf(c_0_1676,negated_conjecture,(epred10_0|epred3_0|hskp16|hskp26|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1556, c_0_1557]), c_0_1499]), c_0_1436])).
cnf(c_0_1677,negated_conjecture,(epred28_0|epred14_0|epred2_0|hskp6|~hskp16), inference(spm,[status(thm)],[c_0_898, c_0_1558]), ['final']).
cnf(c_0_1678,negated_conjecture,(epred42_0|epred10_0|epred8_0|hskp16|~hskp20), inference(spm,[status(thm)],[c_0_1393, c_0_1559]), ['final']).
cnf(c_0_1679,negated_conjecture,(epred41_0|epred11_0|hskp21|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1560, c_0_1561]), c_0_1562]), ['final']).
cnf(c_0_1680,negated_conjecture,(epred10_0|hskp1|hskp9|hskp21|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1487, c_0_1395]), c_0_1563]), c_0_1564])).
cnf(c_0_1681,negated_conjecture,(epred41_0|epred2_0|hskp21|~hskp1), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1432, c_0_1448]), c_0_1232]), c_0_89]), ['final']).
cnf(c_0_1682,negated_conjecture,(epred10_0|hskp9|hskp21|~hskp0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1565, c_0_287]), c_0_1366]), ['final']).
cnf(c_0_1683,negated_conjecture,(epred10_0|hskp9|hskp21|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1487, c_0_49]), c_0_48]), ['final']).
cnf(c_0_1684,negated_conjecture,(epred28_0|hskp9|hskp21|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1566, c_0_49]), c_0_48]), ['final']).
cnf(c_0_1685,negated_conjecture,(epred41_0|c3_1(X1)|hskp0|hskp11|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_270, c_0_613]), ['final']).
cnf(c_0_1686,negated_conjecture,(epred41_0|hskp21|~hskp11), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1432, c_0_15]), c_0_1232]), c_0_16]), ['final']).
cnf(c_0_1687,negated_conjecture,(epred10_0|hskp9|hskp21|~hskp8), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1567, c_0_1481]), c_0_1245]), ['final']).
cnf(c_0_1688,negated_conjecture,(epred35_0|c3_1(X1)|hskp0|hskp11|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_290, c_0_613]), ['final']).
cnf(c_0_1689,negated_conjecture,(epred35_0|hskp21|~hskp11), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1568, c_0_15]), c_0_1232]), c_0_16]), ['final']).
cnf(c_0_1690,negated_conjecture,(epred21_0|epred10_0|hskp16|~hskp22), inference(spm,[status(thm)],[c_0_649, c_0_1569]), ['final']).
cnf(c_0_1691,negated_conjecture,(epred41_0|epred10_0|hskp16|hskp21|~c1_1(X1)|~c0_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1500]), c_0_585]), ['final']).
cnf(c_0_1692,negated_conjecture,(epred10_0|hskp16|hskp27|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1393, c_0_1435]), c_0_1245]), ['final']).
cnf(c_0_1693,negated_conjecture,(epred34_0|hskp6|hskp12|hskp15|~hskp20), inference(spm,[status(thm)],[c_0_1505, c_0_549]), ['final']).
cnf(c_0_1694,negated_conjecture,(epred3_0|epred1_0|hskp6|hskp26|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1570, c_0_1557]), c_0_1499]), ['final']).
cnf(c_0_1695,negated_conjecture,(epred1_0|hskp26|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1571]), c_0_1245]), ['final']).
cnf(c_0_1696,negated_conjecture,(epred9_0|epred1_0|hskp6|hskp26|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1572, c_0_1557]), c_0_1499]), ['final']).
cnf(c_0_1697,negated_conjecture,(epred44_0|epred41_0|epred1_0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_309, c_0_1573]), ['final']).
cnf(c_0_1698,negated_conjecture,(epred42_0|hskp6|hskp15|~hskp18), inference(spm,[status(thm)],[c_0_1530, c_0_1574]), ['final']).
cnf(c_0_1699,negated_conjecture,(epred44_0|epred41_0|epred28_0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_309, c_0_1575]), ['final']).
cnf(c_0_1700,negated_conjecture,(epred44_0|epred41_0|epred11_0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_309, c_0_1576]), ['final']).
cnf(c_0_1701,negated_conjecture,(epred1_0|c3_1(X1)|hskp6|hskp9|hskp26|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_1402]), ['final']).
cnf(c_0_1702,negated_conjecture,(epred28_0|hskp9|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1577, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_1703,negated_conjecture,(hskp6|hskp9|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1223, c_0_717]), c_0_1499]), c_0_95]), c_0_1578]), ['final']).
cnf(c_0_1704,negated_conjecture,(epred28_0|epred15_0|epred14_0|hskp6|~hskp16), inference(spm,[status(thm)],[c_0_1579, c_0_1558]), ['final']).
cnf(c_0_1705,negated_conjecture,(epred15_0|hskp6|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1579, c_0_1557]), c_0_1499]), ['final']).
cnf(c_0_1706,negated_conjecture,(epred1_0|hskp26|hskp27|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1490, c_0_1484]), c_0_1580]), ['final']).
cnf(c_0_1707,negated_conjecture,(epred21_0|epred10_0|hskp15|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1581, c_0_1582]), c_0_1583]), c_0_1583]), c_0_1436]), ['final']).
cnf(c_0_1708,negated_conjecture,(epred1_0|hskp20|hskp24|hskp26|~hskp16), inference(spm,[status(thm)],[c_0_1584, c_0_183]), ['final']).
cnf(c_0_1709,negated_conjecture,(epred35_0|epred34_0|hskp12|~c1_1(a1262)|~hskp8), inference(spm,[status(thm)],[c_0_1585, c_0_134])).
cnf(c_0_1710,negated_conjecture,(epred41_0|epred28_0|hskp21|~c1_1(X1)|~c0_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1442]), c_0_585]), ['final']).
cnf(c_0_1711,negated_conjecture,(epred41_0|epred28_0|hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1506, c_0_1586]), c_0_363]), ['final']).
cnf(c_0_1712,negated_conjecture,(epred41_0|epred10_0|epred8_0|hskp16|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1393, c_0_1587]), c_0_1588])).
cnf(c_0_1713,negated_conjecture,(epred21_0|epred8_0|hskp1|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_1399]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_1714,negated_conjecture,(hskp6|hskp9|hskp21|~c1_1(X1)|~c0_1(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_717]), c_0_585]), c_0_1578]), ['final']).
cnf(c_0_1715,negated_conjecture,(hskp6|hskp29), inference(sr,[status(thm)],[c_0_1589, c_0_1366]), ['final']).
cnf(c_0_1716,negated_conjecture,(epred41_0|epred14_0|epred10_0|~hskp0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1590, c_0_1400]), c_0_1245]), c_0_1591]), ['final']).
cnf(c_0_1717,negated_conjecture,(epred8_0|hskp1|hskp13|hskp28|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_1399]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_1718,negated_conjecture,(hskp6|hskp28), inference(sr,[status(thm)],[c_0_1592, c_0_1366]), ['final']).
cnf(c_0_1719,negated_conjecture,(epred35_0|epred1_0|~c1_1(a1262)|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_1593]), c_0_134])).
cnf(c_0_1720,negated_conjecture,(epred35_0|epred1_0|~c0_1(a1254)|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_1593]), c_0_89]), ['final']).
cnf(c_0_1721,negated_conjecture,(epred28_0|epred3_0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1493, c_0_1482]), ['final']).
cnf(c_0_1722,negated_conjecture,(epred34_0|epred10_0|epred8_0|hskp12|hskp16|~hskp20), inference(spm,[status(thm)],[c_0_1393, c_0_1594])).
cnf(c_0_1723,negated_conjecture,(epred41_0|epred28_0|hskp24|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1595, c_0_1412]), c_0_1499]), c_0_1596]), ['final']).
cnf(c_0_1724,negated_conjecture,(epred28_0|hskp20|hskp28), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1597, c_0_1598]), c_0_1040]), c_0_1]), ['final']).
cnf(c_0_1725,negated_conjecture,(epred1_0|c3_1(a1291)|hskp26|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_1402]), c_0_128]), ['final']).
cnf(c_0_1726,negated_conjecture,(epred41_0|epred1_0|hskp24|hskp26|~hskp16), inference(spm,[status(thm)],[c_0_1509, c_0_183]), ['final']).
cnf(c_0_1727,negated_conjecture,(epred41_0|epred1_0|hskp26|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_1599, c_0_1600])).
cnf(c_0_1728,negated_conjecture,(epred41_0|epred10_0|hskp16|hskp26|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1601, c_0_1557]), c_0_1499]), c_0_1436])).
cnf(c_0_1729,negated_conjecture,(epred41_0|hskp3|hskp24|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1602, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_1730,negated_conjecture,(epred41_0|epred15_0|hskp24|~hskp20), inference(spm,[status(thm)],[c_0_1494, c_0_1603]), ['final']).
cnf(c_0_1731,negated_conjecture,(epred15_0|hskp20|hskp28), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1597, c_0_1604]), c_0_1040]), c_0_1]), ['final']).
cnf(c_0_1732,negated_conjecture,(epred41_0|hskp11|hskp12|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_189, c_0_1605])).
cnf(c_0_1733,negated_conjecture,(epred41_0|epred2_0|hskp24|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1606]), c_0_1232]), ['final']).
cnf(c_0_1734,negated_conjecture,(epred2_0|hskp20|hskp28), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1597, c_0_1607]), c_0_1040]), c_0_1]), ['final']).
cnf(c_0_1735,negated_conjecture,(epred10_0|hskp6|hskp9|hskp27|~hskp15), inference(spm,[status(thm)],[c_0_1434, c_0_1608])).
cnf(c_0_1736,negated_conjecture,(epred10_0|hskp6|hskp9|hskp26|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1516, c_0_294]), c_0_652])).
cnf(c_0_1737,negated_conjecture,(epred35_0|hskp21|~c1_1(a1262)|~hskp8), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1568, c_0_134]), c_0_1232])).
cnf(c_0_1738,negated_conjecture,(epred41_0|epred10_0|hskp16|hskp24|~hskp0), inference(spm,[status(thm)],[c_0_1609, c_0_1405]), ['final']).
cnf(c_0_1739,negated_conjecture,(epred11_0|hskp20|hskp28), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1597, c_0_1610]), c_0_1040]), c_0_1]), ['final']).
cnf(c_0_1740,negated_conjecture,(epred41_0|epred11_0|hskp11|hskp27|~hskp15), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1611, c_0_1561]), c_0_1366])).
cnf(c_0_1741,negated_conjecture,(epred41_0|epred2_0|hskp11|hskp27|~hskp18), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1612, c_0_324]), c_0_1366])).
cnf(c_0_1742,negated_conjecture,(epred11_0|epred3_0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_255, c_0_1482]), ['final']).
cnf(c_0_1743,negated_conjecture,(epred41_0|epred8_0|hskp3|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_1527, c_0_1613]), ['final']).
cnf(c_0_1744,negated_conjecture,(epred41_0|epred2_0|hskp11|hskp24|hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1614, c_0_324]), c_0_100])).
cnf(c_0_1745,negated_conjecture,(epred41_0|epred11_0|hskp24|hskp27), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1615]), c_0_1232]), ['final']).
cnf(c_0_1746,negated_conjecture,(epred41_0|epred11_0|epred8_0|hskp11|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1616, c_0_1587]), c_0_1232])).
cnf(c_0_1747,negated_conjecture,(epred41_0|epred8_0|epred2_0|hskp11|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1617, c_0_1587]), c_0_1232])).
cnf(c_0_1748,negated_conjecture,(epred41_0|epred11_0|epred8_0|hskp1|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1498, c_0_1576])).
cnf(c_0_1749,negated_conjecture,(hskp12|hskp20|hskp24|~hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_1513]), c_0_1618]), ['final']).
cnf(c_0_1750,negated_conjecture,(epred41_0|hskp12|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1619, c_0_1557]), c_0_1499]), c_0_363]), ['final']).
cnf(c_0_1751,negated_conjecture,(epred10_0|epred9_0|hskp16|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1620, c_0_1557]), c_0_1499]), c_0_1436])).
cnf(c_0_1752,negated_conjecture,(epred10_0|hskp16|hskp27|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1444, c_0_1484]), c_0_1393]), ['final']).
cnf(c_0_1753,negated_conjecture,(epred15_0|hskp20|~hskp6), inference(spm,[status(thm)],[c_0_1494, c_0_1621])).
cnf(c_0_1754,negated_conjecture,(hskp12|hskp20|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_1621]), c_0_1618])).
cnf(c_0_1755,negated_conjecture,(epred41_0|epred1_0|hskp24|hskp26|~hskp0), inference(spm,[status(thm)],[c_0_1509, c_0_1405]), ['final']).
cnf(c_0_1756,negated_conjecture,(epred28_0|epred21_0|hskp15|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1483, c_0_1582]), c_0_1583]), c_0_1489]), ['final']).
cnf(c_0_1757,negated_conjecture,(epred28_0|hskp27|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1493, c_0_1484]), c_0_1300]), ['final']).
cnf(c_0_1758,negated_conjecture,(epred8_0|epred1_0|hskp26|~hskp16), inference(spm,[status(thm)],[c_0_1580, c_0_1528]), ['final']).
cnf(c_0_1759,negated_conjecture,(epred35_0|epred10_0|c3_1(X1)|hskp16|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_290, c_0_1322]), ['final']).
cnf(c_0_1760,negated_conjecture,(epred41_0|epred11_0|epred8_0|hskp11|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1611, c_0_1528]), c_0_1366])).
cnf(c_0_1761,negated_conjecture,(epred41_0|epred8_0|epred2_0|hskp11|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1612, c_0_1528]), c_0_1366])).
cnf(c_0_1762,negated_conjecture,(hskp6|hskp9|~hskp8), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_1364]), c_0_1232]), c_0_134]), ['final']).
cnf(c_0_1763,negated_conjecture,(epred41_0|epred14_0|hskp11|~hskp20), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1440, c_0_1412]), c_0_1232]), c_0_1499])).
cnf(c_0_1764,negated_conjecture,(epred41_0|epred14_0|epred1_0|hskp26|~hskp0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1622, c_0_1400]), c_0_1245]), ['final']).
cnf(c_0_1765,negated_conjecture,(~c3_1(a1283)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_1258])]), ['final']).
cnf(c_0_1766,negated_conjecture,(hskp6|hskp9|~c0_1(a1283)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_744, c_0_1258])]), ['final']).
cnf(c_0_1767,negated_conjecture,(epred34_0|hskp12|~c3_1(a1291)|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_447]), c_0_127]), ['final']).
cnf(c_0_1768,negated_conjecture,(epred34_0|hskp12|~hskp6), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1623, c_0_1003]), c_0_1232]), c_0_1040])).
cnf(c_0_1769,negated_conjecture,(epred21_0|epred14_0|epred10_0|hskp16|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1590, c_0_1582]), c_0_1245])).
cnf(c_0_1770,negated_conjecture,(epred11_0|epred4_0|hskp27|~hskp15), inference(spm,[status(thm)],[c_0_1616, c_0_1561]), ['final']).
cnf(c_0_1771,negated_conjecture,(epred11_0|hskp9|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1624, c_0_1407]), c_0_1625]), ['final']).
cnf(c_0_1772,negated_conjecture,(epred2_0|hskp9|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1624, c_0_1448]), c_0_1626]), ['final']).
cnf(c_0_1773,negated_conjecture,(epred41_0|epred11_0|hskp24|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1450, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_1774,negated_conjecture,(epred34_0|hskp12|hskp23|hskp24|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1505, c_0_170]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_1775,negated_conjecture,(hskp20|hskp21|~c1_1(X1)|~c0_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1408]), c_0_585]), c_0_1040]), c_0_1232]), ['final']).
cnf(c_0_1776,negated_conjecture,(hskp20|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1627, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_1777,negated_conjecture,(hskp20|hskp21|~hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1628]), c_0_585]), c_0_1040]), c_0_49]), c_0_48]), ['final']).
cnf(c_0_1778,negated_conjecture,(hskp2|hskp20|~c1_1(X1)|~c0_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_1408]), c_0_1040]), c_0_1232]), ['final']).
cnf(c_0_1779,negated_conjecture,(hskp2|hskp20|~hskp26), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_1628]), c_0_1040]), c_0_49]), c_0_48]), ['final']).
cnf(c_0_1780,negated_conjecture,(epred35_0|epred21_0|~c1_1(a1262)|~hskp8), inference(spm,[status(thm)],[c_0_1629, c_0_134])).
cnf(c_0_1781,negated_conjecture,(epred14_0|hskp8|hskp13|hskp28|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_616]), c_0_95]), c_0_103])).
cnf(c_0_1782,negated_conjecture,(epred8_0|hskp11|hskp12|~hskp16), inference(spm,[status(thm)],[c_0_189, c_0_1630]), ['final']).
cnf(c_0_1783,negated_conjecture,(epred8_0|hskp3|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1527, c_0_1559]), c_0_1499]), ['final']).
cnf(c_0_1784,negated_conjecture,(epred11_0|hskp11|hskp20|hskp24|hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1631, c_0_324]), c_0_100])).
cnf(c_0_1785,negated_conjecture,(epred11_0|hskp20|~hskp1), inference(spm,[status(thm)],[c_0_1632, c_0_1407]), ['final']).
cnf(c_0_1786,negated_conjecture,(epred2_0|hskp20|~hskp1), inference(spm,[status(thm)],[c_0_1632, c_0_1448]), ['final']).
cnf(c_0_1787,negated_conjecture,(hskp3|hskp27|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1527, c_0_1435]), c_0_1245]), ['final']).
cnf(c_0_1788,negated_conjecture,(hskp11|hskp12|~hskp18), inference(spm,[status(thm)],[c_0_189, c_0_1633]), ['final']).
cnf(c_0_1789,negated_conjecture,(epred39_0|hskp11|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_688, c_0_160]), ['final']).
cnf(c_0_1790,negated_conjecture,(hskp20|~hskp8), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1634, c_0_1481]), c_0_1245]), ['final']).
cnf(c_0_1791,negated_conjecture,(c1_1(a1270)|hskp11|hskp12|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_97]), c_0_98])).
cnf(c_0_1792,negated_conjecture,(hskp14), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_380, c_0_1635])]), c_0_1636]), ['final']).
cnf(c_0_1793,negated_conjecture,(hskp17), inference(sr,[status(thm)],[c_0_1637, c_0_1366]), ['final']).
cnf(c_0_1794,negated_conjecture,(epred15_0|c1_1(a1270)|~hskp14), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_97]), c_0_98])).
cnf(c_0_1795,negated_conjecture,(hskp10), inference(sr,[status(thm)],[c_0_1638, c_0_1313]), ['final']).
cnf(c_0_1796,negated_conjecture,(c0_1(a1257)|hskp11|hskp12|~hskp4), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_26]), c_0_27])).
cnf(c_0_1797,negated_conjecture,(hskp4), inference(sr,[status(thm)],[c_0_1639, c_0_1640]), ['final']).
cnf(c_0_1798,negated_conjecture,(epred44_0|epred5_0|hskp13), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1641, c_0_1258])])).
cnf(c_0_1799,negated_conjecture,(epred28_0|epred24_0), inference(spm,[status(thm)],[c_0_1286, c_0_1642])).
cnf(c_0_1800,negated_conjecture,(epred14_0|hskp8|~c0_1(a1283)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_616]), c_0_30])).
cnf(c_0_1801,negated_conjecture,(epred8_0|hskp1|~c0_1(a1283)|~hskp19), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_1399]), c_0_30])).
cnf(c_0_1802,negated_conjecture,(epred8_0|c1_1(a1262)|hskp1|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_1399]), c_0_134]), ['final']).
cnf(c_0_1803,negated_conjecture,(epred3_0|c0_1(a1258)|hskp26|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_782]), c_0_10]), ['final']).
cnf(c_0_1804,negated_conjecture,(epred11_0|hskp13|hskp24|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_487, c_0_847]), ['final']).
cnf(c_0_1805,negated_conjecture,(epred34_0|hskp1|hskp12|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_1643, c_0_1395]), ['final']).
cnf(c_0_1806,negated_conjecture,(epred28_0|hskp12|hskp13|~hskp8), inference(spm,[status(thm)],[c_0_1341, c_0_1644]), ['final']).
cnf(c_0_1807,negated_conjecture,(epred42_0|epred11_0|hskp24|~hskp20), inference(spm,[status(thm)],[c_0_487, c_0_1510]), ['final']).
cnf(c_0_1808,negated_conjecture,(epred41_0|epred28_0|epred14_0|~hskp0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1645, c_0_1400]), c_0_1245]), ['final']).
cnf(c_0_1809,negated_conjecture,(epred35_0|epred11_0|epred8_0|hskp11|~hskp16), inference(spm,[status(thm)],[c_0_1646, c_0_1528])).
cnf(c_0_1810,negated_conjecture,(epred35_0|epred8_0|epred2_0|hskp11|~hskp16), inference(spm,[status(thm)],[c_0_1647, c_0_1528])).
cnf(c_0_1811,negated_conjecture,(epred21_0|epred11_0|hskp24|~hskp20), inference(spm,[status(thm)],[c_0_487, c_0_1648]), ['final']).
cnf(c_0_1812,negated_conjecture,(epred42_0|epred14_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_635, c_0_1649]), c_0_1232]), ['final']).
cnf(c_0_1813,negated_conjecture,(epred44_0|epred42_0|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1223, c_0_271]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_1814,negated_conjecture,(epred35_0|epred14_0|hskp11|~hskp16), inference(spm,[status(thm)],[c_0_1411, c_0_1441])).
cnf(c_0_1815,negated_conjecture,(epred34_0|epred14_0|hskp8|hskp12|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1505, c_0_616]), c_0_95]), c_0_103])).
cnf(c_0_1816,negated_conjecture,(~epred45_0|~epred44_0|~epred43_0|~epred41_0), inference(spm,[status(thm)],[c_0_458, c_0_1650])).
cnf(c_0_1817,negated_conjecture,(~epred44_0|~epred43_0|~epred39_0), inference(spm,[status(thm)],[c_0_458, c_0_1235])).
cnf(c_0_1818,negated_conjecture,(~epred45_0|~epred42_0|~epred41_0|~epred40_0), inference(spm,[status(thm)],[c_0_1651, c_0_1650])).
cnf(c_0_1819,negated_conjecture,(epred45_0|~epred42_0|~epred24_0), inference(spm,[status(thm)],[c_0_513, c_0_1149])).
cnf(c_0_1820,negated_conjecture,(epred45_0|epred42_0|~epred39_0), inference(spm,[status(thm)],[c_0_1236, c_0_1652])).
cnf(c_0_1821,negated_conjecture,(epred34_0|hskp0|hskp12|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_1643, c_0_1449]), ['final']).
cnf(c_0_1822,negated_conjecture,(epred31_0|epred28_0|epred24_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1653, c_0_1654]), c_0_796])).
cnf(c_0_1823,negated_conjecture,(epred29_0|hskp18|hskp20|hskp27|hskp28), inference(spm,[status(thm)],[c_0_1597, c_0_100])).
cnf(c_0_1824,negated_conjecture,(epred41_0|epred17_0|~epred22_0), inference(spm,[status(thm)],[c_0_519, c_0_1179])).
cnf(c_0_1825,negated_conjecture,(epred24_0|epred18_0|epred17_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_542, c_0_857]), c_0_585])).
cnf(c_0_1826,negated_conjecture,(epred45_0|epred31_0|epred24_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_1655]), c_0_1656])).
cnf(c_0_1827,negated_conjecture,(epred31_0|epred24_0|~epred41_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_796, c_0_1187]), c_0_434])).
cnf(c_0_1828,negated_conjecture,(epred31_0|epred24_0|~epred25_0), inference(spm,[status(thm)],[c_0_796, c_0_1657])).
cnf(c_0_1829,negated_conjecture,(epred35_0|epred28_0|c3_1(X1)|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_290, c_0_1295]), ['final']).
cnf(c_0_1830,negated_conjecture,(epred41_0|epred23_0|epred15_0|~epred39_0|~epred38_0), inference(spm,[status(thm)],[c_0_866, c_0_1658])).
cnf(c_0_1831,negated_conjecture,(epred41_0|epred38_0|~epred22_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_519, c_0_1659]), c_0_1660]), c_0_1184])).
cnf(c_0_1832,negated_conjecture,(epred44_0|epred41_0|epred17_0), inference(sr,[status(thm)],[c_0_1661, c_0_585])).
cnf(c_0_1833,negated_conjecture,(epred44_0|~epred41_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_1662]), c_0_1663])).
cnf(c_0_1834,negated_conjecture,(epred42_0|epred39_0|epred33_0), inference(spm,[status(thm)],[c_0_788, c_0_1664])).
cnf(c_0_1835,negated_conjecture,(epred45_0|epred42_0|epred33_0), inference(spm,[status(thm)],[c_0_788, c_0_544])).
cnf(c_0_1836,negated_conjecture,(epred35_0|epred17_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_667, c_0_1665]), c_0_1666]), c_0_1666]), c_0_539])).
cnf(c_0_1837,negated_conjecture,(c3_1(a1291)|hskp0|hskp11|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_613]), c_0_127]), ['final']).
cnf(c_0_1838,negated_conjecture,(epred28_0|c3_1(a1291)|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_1295]), c_0_128]), ['final']).
cnf(c_0_1839,negated_conjecture,(c1_1(a1262)|hskp23|hskp24|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_170]), c_0_134]), ['final']).
cnf(c_0_1840,negated_conjecture,(epred45_0|epred28_0|~hskp8), inference(spm,[status(thm)],[c_0_1341, c_0_1667]), ['final']).
cnf(c_0_1841,negated_conjecture,(epred28_0|epred9_0|~hskp8), inference(spm,[status(thm)],[c_0_1341, c_0_1668]), ['final']).
cnf(c_0_1842,negated_conjecture,(epred35_0|~epred39_0|~epred24_0), inference(spm,[status(thm)],[c_0_513, c_0_1669])).
cnf(c_0_1843,negated_conjecture,(epred21_0|hskp0|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_599, c_0_1449]), ['final']).
cnf(c_0_1844,negated_conjecture,(epred45_0|hskp0|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_419, c_0_1449]), ['final']).
cnf(c_0_1845,negated_conjecture,(hskp1|hskp12|hskp13|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_865, c_0_1395]), ['final']).
cnf(c_0_1846,negated_conjecture,(c0_1(a1258)|hskp2|hskp3|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_172]), c_0_10]), ['final']).
cnf(c_0_1847,negated_conjecture,(epred21_0|hskp1|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_599, c_0_1395]), ['final']).
cnf(c_0_1848,negated_conjecture,(epred42_0|epred9_0|~hskp20), inference(spm,[status(thm)],[c_0_1223, c_0_1670]), ['final']).
cnf(c_0_1849,negated_conjecture,(epred21_0|epred9_0|~hskp20), inference(spm,[status(thm)],[c_0_373, c_0_1670]), ['final']).
cnf(c_0_1850,negated_conjecture,(epred21_0|epred4_0|~hskp20), inference(spm,[status(thm)],[c_0_373, c_0_881]), ['final']).
cnf(c_0_1851,negated_conjecture,(hskp0|hskp12|hskp13|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_865, c_0_1449]), ['final']).
cnf(c_0_1852,negated_conjecture,(epred4_0|c3_1(a1291)|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_615]), c_0_127]), ['final']).
cnf(c_0_1853,negated_conjecture,(epred9_0|c3_1(a1279)|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_1438]), c_0_55]), ['final']).
cnf(c_0_1854,negated_conjecture,(epred2_0|c3_1(a1279)|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_851]), c_0_54]), ['final']).
cnf(c_0_1855,negated_conjecture,(epred44_0|epred31_0|epred24_0), inference(spm,[status(thm)],[c_0_305, c_0_1671])).
cnf(c_0_1856,negated_conjecture,(~epred28_0|~epred27_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_149, c_0_1672]), c_0_1289])).
cnf(c_0_1857,negated_conjecture,(~epred11_0|~epred10_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_150, c_0_1302]), c_0_218])).
cnf(c_0_1858,negated_conjecture,(~epred9_0|~epred8_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_151, c_0_1338]), c_0_1398])).
cnf(c_0_1859,negated_conjecture,(~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[c_0_152, c_0_1339]), c_0_811])).
cnf(c_0_1860,negated_conjecture,(epred31_0|~epred40_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1651, c_0_961]), c_0_1673])).
cnf(c_0_1861,negated_conjecture,(epred41_0|epred14_0|epred10_0|hskp21|~hskp15), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1674, c_0_1511]), c_0_858]), c_0_1591]), ['final']).
cnf(c_0_1862,negated_conjecture,(epred41_0|epred10_0|epred8_0|hskp9|hskp26|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1434, c_0_1613]), c_0_1436]), ['final']).
cnf(c_0_1863,negated_conjecture,(epred10_0|hskp9|hskp26|~hskp18), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1434, c_0_1675]), c_0_883]), c_0_883]), c_0_1436]), ['final']).
cnf(c_0_1864,negated_conjecture,(epred10_0|epred3_0|hskp26|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_1676]), c_0_1366]), ['final']).
cnf(c_0_1865,negated_conjecture,(epred35_0|epred34_0|epred11_0|hskp12|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1585, c_0_1407]), c_0_89]), ['final']).
cnf(c_0_1866,negated_conjecture,(epred35_0|epred34_0|epred2_0|hskp12|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1585, c_0_1448]), c_0_89]), ['final']).
cnf(c_0_1867,negated_conjecture,(epred28_0|epred14_0|epred10_0|epred8_0|epred2_0|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1677, c_0_1678]), c_0_1499]), c_0_1436]), ['final']).
cnf(c_0_1868,negated_conjecture,(epred41_0|c3_1(X1)|c0_1(X1)|hskp3|~c1_1(X1)), inference(spm,[status(thm)],[c_0_270, c_0_1369]), ['final']).
cnf(c_0_1869,negated_conjecture,(epred41_0|epred11_0|hskp21|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1679, c_0_1557]), c_0_1499]), c_0_1625]), ['final']).
cnf(c_0_1870,negated_conjecture,(epred41_0|epred11_0|hskp21|~hskp5), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1560, c_0_1435]), c_0_1245]), c_0_1562]), ['final']).
cnf(c_0_1871,negated_conjecture,(epred41_0|epred11_0|hskp21|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1560, c_0_324]), c_0_1562]), ['final']).
cnf(c_0_1872,negated_conjecture,(epred10_0|hskp9|hskp21|~hskp5), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1434, c_0_1680]), c_0_1436]), ['final']).
cnf(c_0_1873,negated_conjecture,(epred41_0|epred2_0|hskp21|~hskp5), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1681, c_0_1435]), c_0_1245]), c_0_1562]), ['final']).
cnf(c_0_1874,negated_conjecture,(epred41_0|epred2_0|hskp21|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1681, c_0_324]), c_0_1562]), ['final']).
cnf(c_0_1875,negated_conjecture,(epred10_0|epred1_0|hskp9|hskp21|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1682, c_0_1573]), c_0_1683]), ['final']).
cnf(c_0_1876,negated_conjecture,(epred10_0|hskp9|hskp21|~c1_1(a1280)|~hskp29), inference(spm,[status(thm)],[c_0_1487, c_0_109]), ['final']).
cnf(c_0_1877,negated_conjecture,(epred28_0|epred10_0|hskp9|hskp21|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1682, c_0_1575]), c_0_1684]), ['final']).
cnf(c_0_1878,negated_conjecture,(epred10_0|hskp9|hskp21|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1487, c_0_127]), c_0_128]), ['final']).
cnf(c_0_1879,negated_conjecture,(epred10_0|hskp9|hskp21|~hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1487, c_0_15]), c_0_16]), ['final']).
cnf(c_0_1880,negated_conjecture,(epred41_0|hskp0|hskp21|~c1_1(X1)|~c0_1(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1685]), c_0_585]), c_0_1686]), ['final']).
cnf(c_0_1881,negated_conjecture,(epred14_0|epred10_0|hskp9|hskp21|~hskp16), inference(spm,[status(thm)],[c_0_1687, c_0_636]), ['final']).
cnf(c_0_1882,negated_conjecture,(epred35_0|hskp21|~c1_1(X1)|~c0_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1688]), c_0_585]), c_0_1689]), c_0_1411]), ['final']).
cnf(c_0_1883,negated_conjecture,(epred21_0|epred10_0|hskp9|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_649, c_0_1392]), c_0_1690]), c_0_1690]), c_0_1394]), c_0_1436]), ['final']).
cnf(c_0_1884,negated_conjecture,(epred35_0|epred11_0|hskp21|~hskp1), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1568, c_0_1407]), c_0_1232]), c_0_89]), ['final']).
cnf(c_0_1885,negated_conjecture,(epred35_0|epred2_0|hskp21|~hskp1), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1568, c_0_1448]), c_0_1232]), c_0_89]), ['final']).
cnf(c_0_1886,negated_conjecture,(epred41_0|epred10_0|hskp16|hskp21|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1691, c_0_1395]), c_0_1692]), c_0_1562]), c_0_1393]), ['final']).
cnf(c_0_1887,negated_conjecture,(epred34_0|epred3_0|epred1_0|hskp26|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1570, c_0_1693]), c_0_1694]), c_0_1695]), ['final']).
cnf(c_0_1888,negated_conjecture,(epred34_0|epred9_0|epred1_0|hskp26|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1572, c_0_1693]), c_0_1696]), c_0_1695]), ['final']).
cnf(c_0_1889,negated_conjecture,(epred41_0|epred1_0|hskp26|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1697, c_0_1557]), c_0_1366]), c_0_1499]), c_0_363]), ['final']).
cnf(c_0_1890,negated_conjecture,(epred41_0|epred1_0|hskp26|~hskp18), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1697, c_0_1698]), c_0_1366]), c_0_1499]), c_0_363]), ['final']).
cnf(c_0_1891,negated_conjecture,(epred41_0|epred28_0|hskp26|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1699, c_0_1557]), c_0_1366]), c_0_1499]), c_0_1489]), ['final']).
cnf(c_0_1892,negated_conjecture,(epred41_0|epred28_0|hskp26|~hskp18), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1699, c_0_1698]), c_0_1366]), c_0_1499]), c_0_1489]), ['final']).
cnf(c_0_1893,negated_conjecture,(epred41_0|epred11_0|hskp26|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1700, c_0_1557]), c_0_1366]), c_0_1499]), c_0_1625]), ['final']).
cnf(c_0_1894,negated_conjecture,(epred1_0|hskp9|hskp26|~hskp18), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1530, c_0_1701]), c_0_1499]), c_0_55]), c_0_883]), c_0_1695]), ['final']).
cnf(c_0_1895,negated_conjecture,(epred41_0|epred11_0|hskp26|~hskp18), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1700, c_0_1698]), c_0_1366]), c_0_1499]), c_0_1625]), ['final']).
cnf(c_0_1896,negated_conjecture,(epred41_0|epred10_0|epred8_0|hskp1|hskp26|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1498, c_0_1511]), c_0_1528]), ['final']).
cnf(c_0_1897,negated_conjecture,(epred28_0|hskp9|hskp21|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1566, c_0_1449]), c_0_1702]), c_0_1684]), c_0_1577]), ['final']).
cnf(c_0_1898,negated_conjecture,(epred28_0|hskp9|hskp21|~c1_1(a1280)|~hskp29), inference(spm,[status(thm)],[c_0_1566, c_0_109]), ['final']).
cnf(c_0_1899,negated_conjecture,(epred28_0|hskp9|hskp21|~hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1566, c_0_32]), c_0_33]), ['final']).
cnf(c_0_1900,negated_conjecture,(epred28_0|hskp9|hskp21|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1566, c_0_127]), c_0_128]), ['final']).
cnf(c_0_1901,negated_conjecture,(epred28_0|hskp9|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1223, c_0_1488]), c_0_1499]), c_0_95]), c_0_1703]), c_0_1489]), ['final']).
cnf(c_0_1902,negated_conjecture,(epred28_0|epred15_0|epred14_0|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1704, c_0_1678]), c_0_1499]), c_0_1705]), c_0_1705]), c_0_1489]), ['final']).
cnf(c_0_1903,negated_conjecture,(epred21_0|epred10_0|epred1_0|hskp26|hskp27|~hskp20), inference(spm,[status(thm)],[c_0_1706, c_0_1707]), ['final']).
cnf(c_0_1904,negated_conjecture,(epred21_0|epred15_0|epred10_0|~hskp20), inference(spm,[status(thm)],[c_0_1579, c_0_1707]), ['final']).
cnf(c_0_1905,negated_conjecture,(epred21_0|epred10_0|epred2_0|~hskp20), inference(spm,[status(thm)],[c_0_898, c_0_1707]), ['final']).
cnf(c_0_1906,negated_conjecture,(epred1_0|hskp20|hskp26|~c1_1(a1301)|~hskp24), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_1491]), c_0_1040]), ['final']).
cnf(c_0_1907,negated_conjecture,(epred1_0|hskp20|hskp26|hskp28), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1597, c_0_1708]), c_0_1040]), c_0_1]), ['final']).
cnf(c_0_1908,negated_conjecture,(epred41_0|epred14_0|epred1_0|hskp21|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1674, c_0_1573]), c_0_858]), ['final']).
cnf(c_0_1909,negated_conjecture,(epred41_0|epred28_0|epred14_0|hskp21|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1674, c_0_1575]), c_0_858]), ['final']).
cnf(c_0_1910,negated_conjecture,(epred41_0|epred10_0|epred8_0|hskp16|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_1393, c_0_1613]), ['final']).
cnf(c_0_1911,negated_conjecture,(epred35_0|epred34_0|hskp12|~hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1585, c_0_15]), c_0_16]), ['final']).
cnf(c_0_1912,negated_conjecture,(epred35_0|epred34_0|hskp12|~hskp8), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1709, c_0_1481]), c_0_1245]), ['final']).
cnf(c_0_1913,negated_conjecture,(epred41_0|epred28_0|hskp21|~c1_1(a1280)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1710, c_0_109]), c_0_1711]), ['final']).
cnf(c_0_1914,negated_conjecture,(epred41_0|epred10_0|epred8_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_1712]), c_0_1366]), ['final']).
cnf(c_0_1915,negated_conjecture,(epred41_0|epred1_0|hskp21|~c1_1(X1)|~c0_1(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1443]), c_0_585]), c_0_858]), ['final']).
cnf(c_0_1916,negated_conjecture,(epred41_0|hskp21|~c1_1(a1280)|~hskp29), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1432, c_0_109]), c_0_1232]), ['final']).
cnf(c_0_1917,negated_conjecture,(epred41_0|hskp21|~hskp22), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1432, c_0_127]), c_0_1232]), c_0_128]), ['final']).
cnf(c_0_1918,negated_conjecture,(epred41_0|epred14_0|hskp21|~hskp5), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1674, c_0_1518]), c_0_1245]), c_0_858]), ['final']).
cnf(c_0_1919,negated_conjecture,(epred41_0|epred14_0|hskp21|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1674, c_0_212]), c_0_858]), ['final']).
cnf(c_0_1920,negated_conjecture,(epred10_0|hskp9|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1434, c_0_1713]), c_0_1703]), c_0_1703]), c_0_1436]), ['final']).
cnf(c_0_1921,negated_conjecture,(hskp6|hskp9|hskp21|~c1_1(a1280)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1714, c_0_109]), c_0_1715]), ['final']).
cnf(c_0_1922,negated_conjecture,(hskp6|hskp9|hskp21|~hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1714, c_0_32]), c_0_33]), ['final']).
cnf(c_0_1923,negated_conjecture,(hskp6|hskp9|hskp21|~hskp26), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1714, c_0_49]), c_0_48]), ['final']).
cnf(c_0_1924,negated_conjecture,(hskp6|hskp9|hskp21|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1714, c_0_127]), c_0_128]), ['final']).
cnf(c_0_1925,negated_conjecture,(epred41_0|epred14_0|epred10_0|hskp26|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1716, c_0_1511]), c_0_1591]), ['final']).
cnf(c_0_1926,negated_conjecture,(epred10_0|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1434, c_0_1717]), c_0_1718]), c_0_1718]), c_0_1718]), c_0_1718]), c_0_1436]), ['final']).
cnf(c_0_1927,negated_conjecture,(epred35_0|epred1_0|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1593]), c_0_15]), c_0_16]), ['final']).
cnf(c_0_1928,negated_conjecture,(epred35_0|epred1_0|~hskp8), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1719, c_0_1481]), c_0_1245]), ['final']).
cnf(c_0_1929,negated_conjecture,(epred35_0|epred11_0|epred1_0|~hskp1), inference(spm,[status(thm)],[c_0_1720, c_0_1407]), ['final']).
cnf(c_0_1930,negated_conjecture,(epred35_0|epred2_0|epred1_0|~hskp1), inference(spm,[status(thm)],[c_0_1720, c_0_1448]), ['final']).
cnf(c_0_1931,negated_conjecture,(epred28_0|epred14_0|epred3_0|hskp26|~hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1721, c_0_1558]), c_0_1489]), ['final']).
cnf(c_0_1932,negated_conjecture,(epred35_0|hskp21|~hskp27), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1568, c_0_32]), c_0_1232]), c_0_33]), ['final']).
cnf(c_0_1933,negated_conjecture,(epred34_0|epred8_0|epred1_0|hskp12|hskp26|~hskp20), inference(spm,[status(thm)],[c_0_1580, c_0_1594]), ['final']).
cnf(c_0_1934,negated_conjecture,(epred34_0|epred10_0|epred8_0|hskp12|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_1722]), c_0_1366]), ['final']).
cnf(c_0_1935,negated_conjecture,(epred41_0|epred28_0|hskp24|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1595, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_1936,negated_conjecture,(epred35_0|hskp21|~hskp22), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1568, c_0_127]), c_0_1232]), c_0_128]), ['final']).
cnf(c_0_1937,negated_conjecture,(epred41_0|epred28_0|hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1599, c_0_1723]), c_0_1724]), ['final']).
cnf(c_0_1938,negated_conjecture,(epred1_0|hskp21|hskp26|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1725]), c_0_585]), c_0_127]), c_0_128]), ['final']).
cnf(c_0_1939,negated_conjecture,(epred41_0|epred1_0|hskp26|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1597, c_0_1726]), c_0_1040]), c_0_1]), c_0_1727]), ['final']).
cnf(c_0_1940,negated_conjecture,(epred41_0|epred10_0|hskp26|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_1728]), c_0_1366]), ['final']).
cnf(c_0_1941,negated_conjecture,(epred41_0|epred2_0|hskp24|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1514, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_1942,negated_conjecture,(epred41_0|epred15_0|hskp24|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_1494, c_0_1729]), ['final']).
cnf(c_0_1943,negated_conjecture,(epred41_0|epred15_0|hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1599, c_0_1730]), c_0_1731]), ['final']).
cnf(c_0_1944,negated_conjecture,(epred41_0|hskp12|hskp24|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1732]), c_0_1232]), ['final']).
cnf(c_0_1945,negated_conjecture,(epred41_0|epred15_0|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_1494, c_0_1605]), ['final']).
cnf(c_0_1946,negated_conjecture,(epred41_0|epred10_0|hskp16|hskp26|~hskp18), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1601, c_0_1698]), c_0_1499]), c_0_1436]), ['final']).
cnf(c_0_1947,negated_conjecture,(epred41_0|epred2_0|hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1599, c_0_1733]), c_0_1734]), ['final']).
cnf(c_0_1948,negated_conjecture,(epred10_0|hskp9|hskp27|~hskp15), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1434, c_0_1561]), c_0_1735]), c_0_1436]), ['final']).
cnf(c_0_1949,negated_conjecture,(epred10_0|hskp9|hskp26|~hskp15), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1516, c_0_1482]), c_0_1736]), c_0_1436]), ['final']).
cnf(c_0_1950,negated_conjecture,(epred35_0|hskp21|~hskp8), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1737, c_0_1481]), c_0_1245]), ['final']).
cnf(c_0_1951,negated_conjecture,(epred41_0|epred10_0|hskp16|hskp24|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1738, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_1952,negated_conjecture,(epred11_0|epred8_0|hskp13|hskp28), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1616, c_0_1717]), c_0_1232]), c_0_1739]), ['final']).
cnf(c_0_1953,negated_conjecture,(epred8_0|epred2_0|hskp13|hskp28), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1617, c_0_1717]), c_0_1232]), c_0_1734]), ['final']).
cnf(c_0_1954,negated_conjecture,(epred41_0|epred11_0|hskp27|~hskp15), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1740]), c_0_1232]), ['final']).
cnf(c_0_1955,negated_conjecture,(epred41_0|epred2_0|hskp27|~hskp18), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1741]), c_0_1232]), ['final']).
cnf(c_0_1956,negated_conjecture,(epred9_0|epred1_0|hskp6|hskp26|~hskp18), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1572, c_0_1698]), c_0_1499]), ['final']).
cnf(c_0_1957,negated_conjecture,(epred28_0|epred3_0|hskp26|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1721, c_0_1557]), c_0_1499]), c_0_1489]), ['final']).
cnf(c_0_1958,negated_conjecture,(epred11_0|epred3_0|hskp26|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1742, c_0_1557]), c_0_1499]), c_0_1625]), ['final']).
cnf(c_0_1959,negated_conjecture,(epred41_0|epred8_0|epred1_0|hskp26|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1498, c_0_1573]), c_0_1580]), ['final']).
cnf(c_0_1960,negated_conjecture,(epred41_0|epred15_0|epred8_0|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_1494, c_0_1743]), ['final']).
cnf(c_0_1961,negated_conjecture,(epred41_0|epred8_0|epred1_0|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_1580, c_0_1613]), ['final']).
cnf(c_0_1962,negated_conjecture,(epred41_0|epred28_0|epred8_0|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_1300, c_0_1613]), ['final']).
cnf(c_0_1963,negated_conjecture,(epred41_0|epred11_0|epred8_0|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1616, c_0_1613]), c_0_1232]), ['final']).
cnf(c_0_1964,negated_conjecture,(epred41_0|epred8_0|epred2_0|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1617, c_0_1613]), c_0_1232]), ['final']).
cnf(c_0_1965,negated_conjecture,(epred41_0|epred2_0|hskp24|hskp27), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1744]), c_0_1232]), ['final']).
cnf(c_0_1966,negated_conjecture,(epred41_0|epred28_0|epred8_0|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1300, c_0_1587]), c_0_1596]), ['final']).
cnf(c_0_1967,negated_conjecture,(epred41_0|epred11_0|hskp27|hskp28), inference(spm,[status(thm)],[c_0_1599, c_0_1745]), ['final']).
cnf(c_0_1968,negated_conjecture,(epred41_0|epred11_0|epred8_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1746]), c_0_1232]), ['final']).
cnf(c_0_1969,negated_conjecture,(epred41_0|epred8_0|epred2_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1747]), c_0_1232]), ['final']).
cnf(c_0_1970,negated_conjecture,(epred8_0|epred1_0|hskp13|hskp26|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_1580, c_0_1717]), ['final']).
cnf(c_0_1971,negated_conjecture,(epred1_0|hskp6|hskp26|hskp27|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1706, c_0_1557]), c_0_1499]), ['final']).
cnf(c_0_1972,negated_conjecture,(epred41_0|epred28_0|epred8_0|hskp26|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1498, c_0_1575]), c_0_1300]), ['final']).
cnf(c_0_1973,negated_conjecture,(epred41_0|epred11_0|epred8_0|hskp26|~hskp15), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1616, c_0_1748]), c_0_1232]), ['final']).
cnf(c_0_1974,negated_conjecture,(epred10_0|hskp16|hskp20|~c1_1(a1301)|~hskp24), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_1336]), c_0_1040]), ['final']).
cnf(c_0_1975,negated_conjecture,(epred41_0|hskp12|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1599, c_0_1749]), c_0_1]), c_0_1750]), ['final']).
cnf(c_0_1976,negated_conjecture,(hskp3|hskp20|~c1_1(a1301)|~hskp24), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_1403]), c_0_35]), ['final']).
cnf(c_0_1977,negated_conjecture,(hskp12|hskp20|hskp28), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1597, c_0_1749]), c_0_1040]), c_0_1]), ['final']).
cnf(c_0_1978,negated_conjecture,(epred10_0|hskp16|hskp20|~hskp23), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_1336]), c_0_1040]), c_0_38]), ['final']).
cnf(c_0_1979,negated_conjecture,(epred10_0|epred9_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_1751]), c_0_1366]), ['final']).
cnf(c_0_1980,negated_conjecture,(epred10_0|epred9_0|hskp16|~hskp18), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1620, c_0_1698]), c_0_1499]), c_0_1436]), ['final']).
cnf(c_0_1981,negated_conjecture,(epred10_0|hskp16|hskp27|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1752, c_0_1557]), c_0_1499]), c_0_1436]), ['final']).
cnf(c_0_1982,negated_conjecture,(epred15_0|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1753]), c_0_1245]), ['final']).
cnf(c_0_1983,negated_conjecture,(hskp12|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_1754]), c_0_1245]), ['final']).
cnf(c_0_1984,negated_conjecture,(epred41_0|epred1_0|hskp24|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1755, c_0_1573]), ['final']).
cnf(c_0_1985,negated_conjecture,(epred41_0|epred1_0|hskp24|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1755, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_1986,negated_conjecture,(epred28_0|epred21_0|epred15_0|~hskp20), inference(spm,[status(thm)],[c_0_1579, c_0_1756]), ['final']).
cnf(c_0_1987,negated_conjecture,(epred28_0|epred21_0|hskp27|~hskp20), inference(spm,[status(thm)],[c_0_1757, c_0_1756]), ['final']).
cnf(c_0_1988,negated_conjecture,(epred28_0|epred21_0|epred2_0|~hskp20), inference(spm,[status(thm)],[c_0_898, c_0_1756]), ['final']).
cnf(c_0_1989,negated_conjecture,(epred28_0|hskp13|hskp28), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_587, c_0_1488]), c_0_1718]), c_0_1718]), c_0_1718]), c_0_1718]), c_0_1489]), ['final']).
cnf(c_0_1990,negated_conjecture,(epred41_0|epred1_0|hskp26|~hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1443]), c_0_16]), ['final']).
cnf(c_0_1991,negated_conjecture,(epred35_0|epred21_0|epred11_0|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1629, c_0_1407]), c_0_89]), ['final']).
cnf(c_0_1992,negated_conjecture,(epred35_0|epred21_0|epred2_0|~hskp1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1629, c_0_1448]), c_0_89]), ['final']).
cnf(c_0_1993,negated_conjecture,(epred10_0|hskp6|hskp9|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_1434, c_0_324]), ['final']).
cnf(c_0_1994,negated_conjecture,(epred10_0|epred9_0|hskp6|hskp9|~hskp15), inference(spm,[status(thm)],[c_0_1516, c_0_1492]), ['final']).
cnf(c_0_1995,negated_conjecture,(epred10_0|epred8_0|epred1_0|hskp26|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1758, c_0_1678]), c_0_1499]), ['final']).
cnf(c_0_1996,negated_conjecture,(epred35_0|hskp11|~hskp8), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_1688]), c_0_1401]), c_0_1401]), c_0_1411]), ['final']).
cnf(c_0_1997,negated_conjecture,(epred35_0|epred10_0|hskp16|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1759]), c_0_15]), c_0_16]), ['final']).
cnf(c_0_1998,negated_conjecture,(epred35_0|hskp11|~c0_1(a1254)|~hskp1), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_1688]), c_0_89]), c_0_1411]), ['final']).
cnf(c_0_1999,negated_conjecture,(epred41_0|epred10_0|hskp16|hskp24|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1609, c_0_1510]), c_0_1499]), ['final']).
cnf(c_0_2000,negated_conjecture,(epred41_0|epred11_0|epred8_0|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1760]), c_0_1232]), ['final']).
cnf(c_0_2001,negated_conjecture,(epred41_0|epred8_0|epred2_0|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1761]), c_0_1232]), ['final']).
cnf(c_0_2002,negated_conjecture,(epred28_0|hskp9|~hskp16), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1488]), c_0_41]), c_0_1394]), c_0_1489]), ['final']).
cnf(c_0_2003,negated_conjecture,(epred28_0|hskp9|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1488]), c_0_15]), c_0_16]), c_0_1489]), ['final']).
cnf(c_0_2004,negated_conjecture,(epred28_0|hskp9|~hskp8), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_1488]), c_0_134]), c_0_1762]), c_0_1489]), ['final']).
cnf(c_0_2005,negated_conjecture,(epred10_0|hskp16|hskp21|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1569]), c_0_585]), c_0_127]), c_0_128]), ['final']).
cnf(c_0_2006,negated_conjecture,(epred28_0|hskp9|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1577, c_0_1575]), ['final']).
cnf(c_0_2007,negated_conjecture,(epred28_0|hskp9|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_1577, c_0_212]), ['final']).
cnf(c_0_2008,negated_conjecture,(epred41_0|epred14_0|hskp26|~hskp5), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1440, c_0_1518]), c_0_1232]), c_0_1245]), ['final']).
cnf(c_0_2009,negated_conjecture,(epred41_0|epred14_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1763]), c_0_1232]), ['final']).
cnf(c_0_2010,negated_conjecture,(epred41_0|epred14_0|epred11_0|hskp26|~hskp15), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1440, c_0_1576]), c_0_1232]), ['final']).
cnf(c_0_2011,negated_conjecture,(epred41_0|epred14_0|epred11_0|hskp11|~hskp1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1440, c_0_1451]), c_0_1232]), ['final']).
cnf(c_0_2012,negated_conjecture,(epred41_0|epred14_0|epred2_0|hskp11|~hskp1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1440, c_0_1517]), c_0_1232]), ['final']).
cnf(c_0_2013,negated_conjecture,(epred41_0|epred14_0|epred1_0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1764, c_0_1573]), ['final']).
cnf(c_0_2014,negated_conjecture,(epred21_0|epred14_0|epred1_0|hskp26|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1622, c_0_1582]), c_0_1245]), ['final']).
cnf(c_0_2015,negated_conjecture,(epred28_0|hskp9|~c0_1(a1283)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1765, c_0_1488]), c_0_1766]), c_0_1489]), ['final']).
cnf(c_0_2016,negated_conjecture,(epred34_0|hskp9|hskp12|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1767, c_0_717]), c_0_127]), c_0_1578]), c_0_1768]), ['final']).
cnf(c_0_2017,negated_conjecture,(epred21_0|epred14_0|epred10_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_1769]), c_0_1366]), ['final']).
cnf(c_0_2018,negated_conjecture,(epred34_0|hskp9|hskp12|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1505, c_0_717]), c_0_95]), c_0_1703]), c_0_1768]), ['final']).
cnf(c_0_2019,negated_conjecture,(epred34_0|epred28_0|hskp12|hskp27|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1757, c_0_1693]), c_0_1489]), ['final']).
cnf(c_0_2020,negated_conjecture,(epred21_0|hskp6|hskp9|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_649, c_0_717]), c_0_127]), c_0_1578]), ['final']).
cnf(c_0_2021,negated_conjecture,(epred11_0|hskp27|~hskp20), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1770, c_0_1557]), c_0_1232]), c_0_1499]), c_0_1625]), ['final']).
cnf(c_0_2022,negated_conjecture,(epred11_0|hskp27|~hskp18), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1770, c_0_1698]), c_0_1232]), c_0_1499]), c_0_1625]), ['final']).
cnf(c_0_2023,negated_conjecture,(epred11_0|hskp9|hskp27|~hskp15), inference(spm,[status(thm)],[c_0_1771, c_0_1561]), ['final']).
cnf(c_0_2024,negated_conjecture,(epred2_0|hskp9|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_1772, c_0_324]), ['final']).
cnf(c_0_2025,negated_conjecture,(epred11_0|epred8_0|hskp9|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1771, c_0_1559]), c_0_1499]), ['final']).
cnf(c_0_2026,negated_conjecture,(epred11_0|epred8_0|hskp9|~hskp16), inference(spm,[status(thm)],[c_0_1771, c_0_1528]), ['final']).
cnf(c_0_2027,negated_conjecture,(epred8_0|epred2_0|hskp9|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1772, c_0_1559]), c_0_1499]), ['final']).
cnf(c_0_2028,negated_conjecture,(epred8_0|epred2_0|hskp9|~hskp16), inference(spm,[status(thm)],[c_0_1772, c_0_1528]), ['final']).
cnf(c_0_2029,negated_conjecture,(epred28_0|hskp20|~c1_1(a1301)|~hskp24), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_1397]), c_0_1040]), ['final']).
cnf(c_0_2030,negated_conjecture,(epred41_0|epred11_0|hskp26|hskp28|~hskp5), inference(spm,[status(thm)],[c_0_1599, c_0_1773]), ['final']).
cnf(c_0_2031,negated_conjecture,(epred28_0|epred14_0|hskp27|~hskp16), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1757, c_0_1558]), c_0_1489]), ['final']).
cnf(c_0_2032,negated_conjecture,(epred35_0|epred10_0|hskp16|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1411, c_0_1511]), c_0_347]), ['final']).
cnf(c_0_2033,negated_conjecture,(epred34_0|epred11_0|hskp12|hskp24|~hskp20), inference(spm,[status(thm)],[c_0_487, c_0_1774]), ['final']).
cnf(c_0_2034,negated_conjecture,(epred41_0|epred28_0|hskp24|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1595, c_0_1575]), ['final']).
cnf(c_0_2035,negated_conjecture,(epred35_0|epred21_0|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1629, c_0_127]), c_0_128]), ['final']).
cnf(c_0_2036,negated_conjecture,(epred41_0|epred2_0|hskp24|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_1514, c_0_212]), ['final']).
cnf(c_0_2037,negated_conjecture,(epred10_0|hskp16|hskp20|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1627, c_0_1511]), ['final']).
cnf(c_0_2038,negated_conjecture,(hskp20|hskp21|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1775, c_0_1449]), c_0_1776]), c_0_1777]), c_0_1627]), ['final']).
cnf(c_0_2039,negated_conjecture,(hskp2|hskp20|~hskp5), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1778, c_0_1449]), c_0_1776]), c_0_1779]), c_0_1627]), ['final']).
cnf(c_0_2040,negated_conjecture,(epred35_0|epred21_0|~hskp11), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1629, c_0_15]), c_0_16]), ['final']).
cnf(c_0_2041,negated_conjecture,(epred35_0|epred21_0|~hskp8), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1780, c_0_1481]), c_0_1245]), ['final']).
cnf(c_0_2042,negated_conjecture,(hskp20|hskp21|~c1_1(a1280)|~hskp29), inference(spm,[status(thm)],[c_0_1775, c_0_109]), ['final']).
cnf(c_0_2043,negated_conjecture,(hskp20|hskp21|~hskp27), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1775, c_0_32]), c_0_33]), ['final']).
cnf(c_0_2044,negated_conjecture,(hskp20|hskp21|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1775, c_0_127]), c_0_128]), ['final']).
cnf(c_0_2045,negated_conjecture,(epred14_0|hskp13|hskp28|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_635, c_0_1781]), c_0_1232]), ['final']).
cnf(c_0_2046,negated_conjecture,(epred8_0|hskp12|hskp20|~hskp16), inference(spm,[status(thm)],[c_0_1618, c_0_1782]), ['final']).
cnf(c_0_2047,negated_conjecture,(epred41_0|epred8_0|hskp12|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1782]), c_0_1232]), ['final']).
cnf(c_0_2048,negated_conjecture,(epred8_0|hskp11|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_189, c_0_1783]), ['final']).
cnf(c_0_2049,negated_conjecture,(epred35_0|epred8_0|hskp12|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1455, c_0_1782]), c_0_1232]), ['final']).
cnf(c_0_2050,negated_conjecture,(hskp2|hskp20|~hskp22), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1778, c_0_127]), c_0_128]), ['final']).
cnf(c_0_2051,negated_conjecture,(epred11_0|hskp20|hskp24|hskp27), inference(spm,[status(thm)],[c_0_1618, c_0_1784]), ['final']).
cnf(c_0_2052,negated_conjecture,(epred11_0|hskp20|hskp27|~hskp15), inference(spm,[status(thm)],[c_0_1785, c_0_1561]), ['final']).
cnf(c_0_2053,negated_conjecture,(epred2_0|hskp20|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_1786, c_0_324]), ['final']).
cnf(c_0_2054,negated_conjecture,(epred1_0|hskp20|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1627, c_0_1573]), ['final']).
cnf(c_0_2055,negated_conjecture,(c3_1(a1268)|hskp3|~hskp13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_1369]), c_0_87]), ['final']).
cnf(c_0_2056,negated_conjecture,(epred28_0|hskp20|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1627, c_0_1575]), ['final']).
cnf(c_0_2057,negated_conjecture,(epred11_0|hskp20|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1627, c_0_1576]), ['final']).
cnf(c_0_2058,negated_conjecture,(epred21_0|hskp20|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_649, c_0_1408]), c_0_1040]), c_0_1232]), c_0_127]), c_0_128]), ['final']).
cnf(c_0_2059,negated_conjecture,(epred15_0|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_1494, c_0_1787]), ['final']).
cnf(c_0_2060,negated_conjecture,(hskp11|hskp12|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_189, c_0_1787]), ['final']).
cnf(c_0_2061,negated_conjecture,(epred15_0|epred8_0|~hskp20), inference(spm,[status(thm)],[c_0_1494, c_0_1783]), ['final']).
cnf(c_0_2062,negated_conjecture,(epred15_0|epred8_0|~hskp16), inference(spm,[status(thm)],[c_0_1494, c_0_1630]), ['final']).
cnf(c_0_2063,negated_conjecture,(epred34_0|hskp12|hskp20|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1767, c_0_1408]), c_0_1040]), c_0_1232]), c_0_127]), c_0_128]), ['final']).
cnf(c_0_2064,negated_conjecture,(epred11_0|epred8_0|hskp20|~hskp16), inference(spm,[status(thm)],[c_0_1785, c_0_1528]), ['final']).
cnf(c_0_2065,negated_conjecture,(hskp12|hskp20|~hskp18), inference(spm,[status(thm)],[c_0_1618, c_0_1788]), ['final']).
cnf(c_0_2066,negated_conjecture,(hskp12|hskp20|~hskp15), inference(spm,[status(thm)],[c_0_1618, c_0_1520]), ['final']).
cnf(c_0_2067,negated_conjecture,(hskp12|hskp20|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1618, c_0_1789]), c_0_1245]), ['final']).
cnf(c_0_2068,negated_conjecture,(epred41_0|hskp12|~hskp18), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1788]), c_0_1232]), ['final']).
cnf(c_0_2069,negated_conjecture,(epred15_0|~hskp18), inference(spm,[status(thm)],[c_0_1494, c_0_1633]), ['final']).
cnf(c_0_2070,negated_conjecture,(epred8_0|epred2_0|hskp20|~hskp16), inference(spm,[status(thm)],[c_0_1786, c_0_1528]), ['final']).
cnf(c_0_2071,negated_conjecture,(epred14_0|hskp20|~hskp16), inference(spm,[status(thm)],[c_0_1790, c_0_636]), ['final']).
cnf(c_0_2072,negated_conjecture,(hskp20|hskp26|~hskp18), inference(spm,[status(thm)],[c_0_1627, c_0_212]), ['final']).
cnf(c_0_2073,negated_conjecture,(epred35_0|hskp12|~hskp15), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1455, c_0_1520]), c_0_1232]), ['final']).
cnf(c_0_2074,negated_conjecture,(c1_1(a1270)|hskp11|hskp12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1791, c_0_1792])]), ['final']).
cnf(c_0_2075,negated_conjecture,(epred34_0|hskp12|~c0_1(a1277)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_475, c_0_1793])]), ['final']).
cnf(c_0_2076,negated_conjecture,(epred15_0|c1_1(a1270)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1794, c_0_1792])]), ['final']).
cnf(c_0_2077,negated_conjecture,(epred41_0|c3_1(a1264)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_332, c_0_1795])]), ['final']).
cnf(c_0_2078,negated_conjecture,(c0_1(a1257)|hskp11|hskp12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1796, c_0_1797])]), ['final']).
cnf(c_0_2079,negated_conjecture,(hskp2|hskp20|~c0_1(a1277)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_896, c_0_1793])]), ['final']).
cnf(c_0_2080,negated_conjecture,(epred15_0|c0_1(a1257)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1022, c_0_1797])]), ['final']).
cnf(c_0_2081,negated_conjecture,(epred21_0|~c0_1(a1277)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_483, c_0_1793])]), ['final']).
cnf(c_0_2082,negated_conjecture,(c3_1(a1264)|hskp20), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_448, c_0_1795])]), c_0_1040]), ['final']).
cnf(c_0_2083,negated_conjecture,(hskp21|~c0_1(a1277)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_531, c_0_1793])]), c_0_585]), ['final']).
cnf(c_0_2084,negated_conjecture,(epred22_0|hskp9), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1091, c_0_1797])]), ['final']).
cnf(c_0_2085,negated_conjecture,(epred5_0|hskp13), inference(sr,[status(thm)],[c_0_1798, c_0_1366]), ['final']).
cnf(c_0_2086,negated_conjecture,(epred2_0|hskp20|~c1_1(a1301)|~hskp24), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_1409]), c_0_1040]), c_0_35]), ['final']).
cnf(c_0_2087,negated_conjecture,(epred28_0|~epred44_0), inference(spm,[status(thm)],[c_0_1337, c_0_1799]), ['final']).
cnf(c_0_2088,negated_conjecture,(epred28_0|~epred42_0), inference(spm,[status(thm)],[c_0_1177, c_0_1799]), ['final']).
cnf(c_0_2089,negated_conjecture,(epred28_0|~epred26_0), inference(spm,[status(thm)],[c_0_1061, c_0_1799]), ['final']).
cnf(c_0_2090,negated_conjecture,(epred41_0|epred8_0|epred2_0|hskp26|~hskp18), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1617, c_0_1675]), c_0_1232]), ['final']).
cnf(c_0_2091,negated_conjecture,(epred41_0|epred2_0|hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1447, c_0_1586]), c_0_363]), ['final']).
cnf(c_0_2092,negated_conjecture,(epred14_0|hskp8|~c0_1(a1283)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1800, c_0_1258])]), ['final']).
cnf(c_0_2093,negated_conjecture,(epred8_0|hskp1|~c0_1(a1283)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1801, c_0_1258])]), ['final']).
cnf(c_0_2094,negated_conjecture,(hskp23|hskp24|~c0_1(a1283)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_461, c_0_1258])]), ['final']).
cnf(c_0_2095,negated_conjecture,(epred44_0|~c0_1(a1283)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1190, c_0_1258])]), ['final']).
cnf(c_0_2096,negated_conjecture,(epred34_0|epred1_0|hskp12|hskp26|~hskp22), inference(spm,[status(thm)],[c_0_1767, c_0_1725]), ['final']).
cnf(c_0_2097,negated_conjecture,(epred8_0|epred1_0|hskp26|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1521, c_0_1802]), c_0_1580]), ['final']).
cnf(c_0_2098,negated_conjecture,(epred1_0|hskp26|hskp27|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1580, c_0_1435]), c_0_1245]), ['final']).
cnf(c_0_2099,negated_conjecture,(epred34_0|epred3_0|hskp12|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_1643, c_0_1803]), ['final']).
cnf(c_0_2100,negated_conjecture,(epred41_0|epred11_0|hskp13|hskp28), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1599, c_0_1804]), c_0_1739]), ['final']).
cnf(c_0_2101,negated_conjecture,(epred10_0|epred8_0|hskp16|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1495, c_0_1802]), c_0_1393]), ['final']).
cnf(c_0_2102,negated_conjecture,(epred34_0|epred10_0|hskp12|hskp16|~hskp22), inference(spm,[status(thm)],[c_0_1767, c_0_1569]), ['final']).
cnf(c_0_2103,negated_conjecture,(epred34_0|epred28_0|hskp12|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_1300, c_0_1805]), ['final']).
cnf(c_0_2104,negated_conjecture,(epred28_0|epred14_0|hskp12|hskp13|~hskp16), inference(spm,[status(thm)],[c_0_1806, c_0_636]), ['final']).
cnf(c_0_2105,negated_conjecture,(epred42_0|epred8_0|epred1_0|hskp26|~hskp20), inference(spm,[status(thm)],[c_0_1580, c_0_1559]), ['final']).
cnf(c_0_2106,negated_conjecture,(epred42_0|epred41_0|epred11_0|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_1599, c_0_1807]), ['final']).
cnf(c_0_2107,negated_conjecture,(epred34_0|epred2_0|hskp12|~hskp20), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_898, c_0_1693]), c_0_1768]), ['final']).
cnf(c_0_2108,negated_conjecture,(epred41_0|epred28_0|epred14_0|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1808, c_0_1575]), ['final']).
cnf(c_0_2109,negated_conjecture,(epred35_0|epred11_0|epred8_0|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1455, c_0_1809]), c_0_1232]), ['final']).
cnf(c_0_2110,negated_conjecture,(epred35_0|epred8_0|epred2_0|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1455, c_0_1810]), c_0_1232]), ['final']).
cnf(c_0_2111,negated_conjecture,(epred34_0|epred28_0|epred8_0|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_1300, c_0_1594]), ['final']).
cnf(c_0_2112,negated_conjecture,(epred34_0|epred11_0|epred8_0|hskp12|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1616, c_0_1594]), c_0_1232]), ['final']).
cnf(c_0_2113,negated_conjecture,(epred41_0|hskp12|~hskp5), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1502, c_0_1789]), c_0_1232]), c_0_1245]), ['final']).
cnf(c_0_2114,negated_conjecture,(epred41_0|epred21_0|epred11_0|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_1599, c_0_1811]), ['final']).
cnf(c_0_2115,negated_conjecture,(epred41_0|epred14_0|hskp26|~hskp18), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1440, c_0_212]), c_0_1232]), ['final']).
cnf(c_0_2116,negated_conjecture,(epred14_0|epred1_0|hskp26|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1622, c_0_636]), c_0_1245]), ['final']).
cnf(c_0_2117,negated_conjecture,(epred42_0|epred28_0|hskp6|hskp27|~hskp20), inference(spm,[status(thm)],[c_0_1757, c_0_1557]), ['final']).
cnf(c_0_2118,negated_conjecture,(epred21_0|epred8_0|epred1_0|hskp26|~hskp20), inference(spm,[status(thm)],[c_0_1580, c_0_1713]), ['final']).
cnf(c_0_2119,negated_conjecture,(epred21_0|epred10_0|epred8_0|hskp16|~hskp20), inference(spm,[status(thm)],[c_0_1393, c_0_1713]), ['final']).
cnf(c_0_2120,negated_conjecture,(epred42_0|epred14_0|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1812, c_0_1003]), c_0_1040]), ['final']).
cnf(c_0_2121,negated_conjecture,(epred42_0|epred28_0|epred8_0|~hskp20), inference(spm,[status(thm)],[c_0_1300, c_0_1559]), ['final']).
cnf(c_0_2122,negated_conjecture,(epred42_0|epred11_0|epred8_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1616, c_0_1559]), c_0_1232]), ['final']).
cnf(c_0_2123,negated_conjecture,(epred42_0|epred8_0|epred2_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1617, c_0_1559]), c_0_1232]), ['final']).
cnf(c_0_2124,negated_conjecture,(epred11_0|hskp27|~hskp5), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1616, c_0_1435]), c_0_1232]), c_0_1245]), ['final']).
cnf(c_0_2125,negated_conjecture,(epred2_0|hskp27|~hskp5), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1617, c_0_1435]), c_0_1232]), c_0_1245]), ['final']).
cnf(c_0_2126,negated_conjecture,(epred44_0|epred42_0|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1813, c_0_1003]), c_0_1040]), ['final']).
cnf(c_0_2127,negated_conjecture,(epred44_0|epred41_0|hskp26|~hskp5), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_1518]), c_0_1245]), ['final']).
cnf(c_0_2128,negated_conjecture,(epred41_0|epred11_0|hskp24|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_1450, c_0_1576]), ['final']).
cnf(c_0_2129,negated_conjecture,(epred35_0|epred34_0|hskp12|~hskp27), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1496, c_0_31]), c_0_32]), c_0_33]), ['final']).
cnf(c_0_2130,negated_conjecture,(epred35_0|epred1_0|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1411, c_0_1573]), c_0_347]), ['final']).
cnf(c_0_2131,negated_conjecture,(epred35_0|epred14_0|~hskp16), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1455, c_0_1814]), c_0_1232]), ['final']).
cnf(c_0_2132,negated_conjecture,(epred35_0|~hskp6), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1526, c_0_1003]), c_0_1040]), ['final']).
cnf(c_0_2133,negated_conjecture,(epred35_0|epred28_0|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1411, c_0_1575]), c_0_347]), ['final']).
cnf(c_0_2134,negated_conjecture,(epred35_0|epred11_0|~hskp15), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1411, c_0_1576]), c_0_347]), ['final']).
cnf(c_0_2135,negated_conjecture,(epred35_0|~hskp5), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1411, c_0_1518]), c_0_1245]), c_0_347]), ['final']).
cnf(c_0_2136,negated_conjecture,(epred35_0|~hskp18), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1411, c_0_212]), c_0_347]), ['final']).
cnf(c_0_2137,negated_conjecture,(epred28_0|epred21_0|epred14_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1645, c_0_1582]), c_0_1245]), ['final']).
cnf(c_0_2138,negated_conjecture,(epred21_0|epred11_0|epred8_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1616, c_0_1713]), c_0_1232]), ['final']).
cnf(c_0_2139,negated_conjecture,(epred21_0|epred8_0|epred2_0|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1617, c_0_1713]), c_0_1232]), ['final']).
cnf(c_0_2140,negated_conjecture,(epred34_0|epred14_0|hskp12|~hskp20), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_635, c_0_1815]), c_0_1232]), ['final']).
cnf(c_0_2141,negated_conjecture,(epred41_0|epred34_0|hskp12|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1767, c_0_1391]), c_0_1232]), c_0_127]), c_0_128]), ['final']).
cnf(c_0_2142,negated_conjecture,(epred35_0|epred34_0|hskp12|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_1767, c_0_1410]), c_0_1232]), c_0_127]), c_0_128]), ['final']).
cnf(c_0_2143,negated_conjecture,(epred44_0|epred34_0|hskp12|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1505, c_0_271]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_2144,negated_conjecture,(epred41_0|epred28_0|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_1506, c_0_183]), ['final']).
cnf(c_0_2145,negated_conjecture,(~epred45_0|~epred44_0|~epred41_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1816, c_0_1212])]), ['final']).
cnf(c_0_2146,negated_conjecture,(~epred44_0|~epred39_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1817, c_0_1212])]), ['final']).
cnf(c_0_2147,negated_conjecture,(~epred41_0|~epred11_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_396, c_0_1188])]), ['final']).
cnf(c_0_2148,negated_conjecture,(~epred42_0|~epred39_0), inference(spm,[status(thm)],[c_0_1177, c_0_1235]), ['final']).
cnf(c_0_2149,negated_conjecture,(~epred45_0|~epred42_0|~epred41_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1818, c_0_1188])]), ['final']).
cnf(c_0_2150,negated_conjecture,(epred45_0|~epred39_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1819, c_0_1235]), c_0_1820]), ['final']).
cnf(c_0_2151,negated_conjecture,(~epred45_0|~epred41_0|~epred26_0), inference(spm,[status(thm)],[c_0_1061, c_0_1650]), ['final']).
cnf(c_0_2152,negated_conjecture,(~epred39_0|~epred26_0), inference(spm,[status(thm)],[c_0_1061, c_0_1235]), ['final']).
cnf(c_0_2153,negated_conjecture,(epred41_0|epred2_0|hskp24|~hskp16), inference(spm,[status(thm)],[c_0_1447, c_0_183]), ['final']).
cnf(c_0_2154,negated_conjecture,(epred34_0|hskp6|hskp9|hskp12|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_652, c_0_1821]), ['final']).
cnf(c_0_2155,negated_conjecture,(epred28_0|~epred30_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_1822]), c_0_938]), ['final']).
cnf(c_0_2156,negated_conjecture,(hskp18|hskp20|hskp27|hskp28), inference(sr,[status(thm)],[c_0_1823, c_0_1040]), ['final']).
cnf(c_0_2157,negated_conjecture,(epred44_0|hskp13|hskp28|~hskp20), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_271]), c_0_95]), c_0_103]), ['final']).
cnf(c_0_2158,negated_conjecture,(~epred29_0|~epred22_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_1824]), c_0_862]), ['final']).
cnf(c_0_2159,negated_conjecture,(epred18_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_1825]), c_0_958]), ['final']).
cnf(c_0_2160,negated_conjecture,(epred39_0|~epred30_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_914]), c_0_938]), ['final']).
cnf(c_0_2161,negated_conjecture,(epred45_0|~epred30_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_1826]), c_0_938]), ['final']).
cnf(c_0_2162,negated_conjecture,(~epred41_0|~epred30_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_1827]), c_0_938]), ['final']).
cnf(c_0_2163,negated_conjecture,(~epred30_0|~epred25_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_1828]), c_0_938]), ['final']).
cnf(c_0_2164,negated_conjecture,(epred35_0|epred28_0|~hskp11), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_1829]), c_0_15]), c_0_16]), ['final']).
cnf(c_0_2165,negated_conjecture,(epred35_0|epred21_0|~hskp27), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1525, c_0_31]), c_0_32]), c_0_33]), ['final']).
cnf(c_0_2166,negated_conjecture,(epred41_0|epred23_0|~epred39_0|~epred22_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_519, c_0_1830]), c_0_1831]), ['final']).
cnf(c_0_2167,negated_conjecture,(epred44_0|~epred29_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_1832]), c_0_1833]), ['final']).
cnf(c_0_2168,negated_conjecture,(~epred29_0|~epred25_0), inference(spm,[status(thm)],[c_0_532, c_0_1657]), ['final']).
cnf(c_0_2169,negated_conjecture,(~epred45_0|~epred36_0), inference(spm,[status(thm)],[c_0_552, c_0_545]), ['final']).
cnf(c_0_2170,negated_conjecture,(epred42_0|epred39_0|~epred36_0), inference(spm,[status(thm)],[c_0_552, c_0_1834]), ['final']).
cnf(c_0_2171,negated_conjecture,(epred45_0|epred42_0|~epred36_0), inference(spm,[status(thm)],[c_0_552, c_0_1835]), ['final']).
cnf(c_0_2172,negated_conjecture,(epred42_0|~epred39_0|~epred30_0), inference(spm,[status(thm)],[c_0_936, c_0_1229]), ['final']).
cnf(c_0_2173,negated_conjecture,(epred35_0|~epred29_0), inference(spm,[status(thm)],[c_0_532, c_0_1836]), ['final']).
cnf(c_0_2174,negated_conjecture,(epred36_0|hskp0|hskp11|hskp21|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1837]), c_0_127]), c_0_128]), ['final']).
cnf(c_0_2175,negated_conjecture,(epred28_0|epred21_0|epred8_0|~hskp20), inference(spm,[status(thm)],[c_0_1300, c_0_1713]), ['final']).
cnf(c_0_2176,negated_conjecture,(epred44_0|epred21_0|~hskp6), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_1003]), c_0_1662]), ['final']).
cnf(c_0_2177,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp6|hskp9|hskp15|~c0_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_541]), ['final']).
cnf(c_0_2178,negated_conjecture,(epred35_0|~epred30_0), inference(spm,[status(thm)],[c_0_936, c_0_530]), ['final']).
cnf(c_0_2179,negated_conjecture,(epred35_0|~epred26_0), inference(spm,[status(thm)],[c_0_1061, c_0_530]), ['final']).
cnf(c_0_2180,negated_conjecture,(epred21_0|epred1_0|hskp26|~hskp22), inference(spm,[status(thm)],[c_0_649, c_0_1725]), ['final']).
cnf(c_0_2181,negated_conjecture,(epred44_0|epred1_0|hskp26|~hskp8), inference(spm,[status(thm)],[c_0_1521, c_0_1342]), ['final']).
cnf(c_0_2182,negated_conjecture,(epred9_0|epred1_0|hskp26|~hskp8), inference(spm,[status(thm)],[c_0_1521, c_0_1668]), ['final']).
cnf(c_0_2183,negated_conjecture,(epred45_0|epred1_0|hskp26|~hskp8), inference(spm,[status(thm)],[c_0_1521, c_0_1667]), ['final']).
cnf(c_0_2184,negated_conjecture,(epred39_0|epred3_0|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_688, c_0_1803]), ['final']).
cnf(c_0_2185,negated_conjecture,(epred21_0|epred3_0|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_599, c_0_1803]), ['final']).
cnf(c_0_2186,negated_conjecture,(epred45_0|epred3_0|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_419, c_0_1803]), ['final']).
cnf(c_0_2187,negated_conjecture,(epred36_0|epred28_0|hskp21|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1838]), c_0_127]), c_0_128]), ['final']).
cnf(c_0_2188,negated_conjecture,(c1_1(X1)|c3_1(X1)|hskp6|hskp9|hskp12|hskp13|~c0_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_825]), ['final']).
cnf(c_0_2189,negated_conjecture,(epred28_0|epred8_0|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1341, c_0_1802]), c_0_1300]), ['final']).
cnf(c_0_2190,negated_conjecture,(epred29_0|c3_1(X1)|hskp0|hskp11|hskp20|~c1_1(X1)|~c0_1(X1)), inference(spm,[status(thm)],[c_0_418, c_0_613]), ['final']).
cnf(c_0_2191,negated_conjecture,(epred44_0|epred10_0|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1495, c_0_1342]), c_0_374]), ['final']).
cnf(c_0_2192,negated_conjecture,(epred10_0|epred9_0|hskp16|~hskp8), inference(spm,[status(thm)],[c_0_1495, c_0_1668]), ['final']).
cnf(c_0_2193,negated_conjecture,(epred45_0|epred10_0|hskp16|~hskp8), inference(spm,[status(thm)],[c_0_1495, c_0_1667]), ['final']).
cnf(c_0_2194,negated_conjecture,(epred11_0|epred8_0|epred4_0|~hskp16), inference(spm,[status(thm)],[c_0_1616, c_0_1528]), ['final']).
cnf(c_0_2195,negated_conjecture,(epred8_0|epred4_0|epred2_0|~hskp16), inference(spm,[status(thm)],[c_0_1617, c_0_1528]), ['final']).
cnf(c_0_2196,negated_conjecture,(epred1_0|hskp6|hskp15|hskp26|~hskp8), inference(spm,[status(thm)],[c_0_1521, c_0_1433]), ['final']).
cnf(c_0_2197,negated_conjecture,(epred1_0|hskp23|hskp24|hskp26|~hskp8), inference(spm,[status(thm)],[c_0_1521, c_0_1839]), ['final']).
cnf(c_0_2198,negated_conjecture,(epred1_0|hskp12|hskp13|hskp26|~hskp8), inference(spm,[status(thm)],[c_0_1521, c_0_1644]), ['final']).
cnf(c_0_2199,negated_conjecture,(epred1_0|hskp6|hskp9|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_652, c_0_1573]), ['final']).
cnf(c_0_2200,negated_conjecture,(epred3_0|hskp12|hskp13|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_865, c_0_1803]), ['final']).
cnf(c_0_2201,negated_conjecture,(epred45_0|c1_1(X1)|c3_1(X1)|hskp6|hskp9|~c0_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_377]), ['final']).
cnf(c_0_2202,negated_conjecture,(epred39_0|c1_1(X1)|c3_1(X1)|hskp6|hskp9|~c0_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_662]), ['final']).
cnf(c_0_2203,negated_conjecture,(epred9_0|c1_1(X1)|c3_1(X1)|hskp6|hskp9|~c0_1(X1)), inference(spm,[status(thm)],[c_0_171, c_0_1438]), ['final']).
cnf(c_0_2204,negated_conjecture,(epred45_0|epred28_0|epred14_0|~hskp16), inference(spm,[status(thm)],[c_0_1840, c_0_636]), ['final']).
cnf(c_0_2205,negated_conjecture,(epred39_0|epred28_0|epred14_0|~hskp16), inference(spm,[status(thm)],[c_0_1645, c_0_636]), ['final']).
cnf(c_0_2206,negated_conjecture,(epred28_0|epred14_0|epred9_0|~hskp16), inference(spm,[status(thm)],[c_0_1841, c_0_636]), ['final']).
cnf(c_0_2207,negated_conjecture,(epred10_0|hskp23|hskp24|~hskp8), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_1495, c_0_1839]), c_0_183]), ['final']).
cnf(c_0_2208,negated_conjecture,(epred10_0|hskp12|hskp13|hskp16|~hskp8), inference(spm,[status(thm)],[c_0_1495, c_0_1644]), ['final']).
cnf(c_0_2209,negated_conjecture,(epred1_0|c3_1(a1268)|hskp26|~hskp13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_1402]), c_0_86]), ['final']).
cnf(c_0_2210,negated_conjecture,(epred1_0|hskp26|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_1580, c_0_324]), ['final']).
cnf(c_0_2211,negated_conjecture,(epred3_0|hskp26|~hskp18), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_782]), c_0_54]), c_0_55]), ['final']).
cnf(c_0_2212,negated_conjecture,(epred41_0|epred11_0|hskp28|~hskp16), inference(spm,[status(thm)],[c_0_1599, c_0_1610]), ['final']).
cnf(c_0_2213,negated_conjecture,(epred35_0|~epred39_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_866, c_0_498]), c_0_1551]), c_0_1842]), ['final']).
cnf(c_0_2214,negated_conjecture,(epred28_0|epred8_0|~hskp16), inference(spm,[status(thm)],[c_0_1300, c_0_1528]), ['final']).
cnf(c_0_2215,negated_conjecture,(epred10_0|c3_1(a1268)|hskp16|~hskp13), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_1322]), c_0_86]), ['final']).
cnf(c_0_2216,negated_conjecture,(epred10_0|hskp16|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_1393, c_0_324]), ['final']).
cnf(c_0_2217,negated_conjecture,(epred39_0|hskp6|hskp9|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_652, c_0_1518]), ['final']).
cnf(c_0_2218,negated_conjecture,(epred21_0|hskp6|hskp9|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_652, c_0_1843]), ['final']).
cnf(c_0_2219,negated_conjecture,(epred45_0|hskp6|hskp9|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_652, c_0_1844]), ['final']).
cnf(c_0_2220,negated_conjecture,(epred28_0|hskp12|hskp13|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_1300, c_0_1845]), ['final']).
cnf(c_0_2221,negated_conjecture,(epred11_0|hskp6|hskp9|hskp26|~hskp15), inference(spm,[status(thm)],[c_0_652, c_0_1576]), ['final']).
cnf(c_0_2222,negated_conjecture,(epred34_0|hskp0|hskp11|hskp12|~hskp22), inference(spm,[status(thm)],[c_0_1767, c_0_1837]), ['final']).
cnf(c_0_2223,negated_conjecture,(epred34_0|hskp0|hskp11|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_1505, c_0_630]), ['final']).
cnf(c_0_2224,negated_conjecture,(epred34_0|hskp2|hskp3|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_1643, c_0_1846]), ['final']).
cnf(c_0_2225,negated_conjecture,(epred39_0|epred28_0|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_1300, c_0_1435]), ['final']).
cnf(c_0_2226,negated_conjecture,(epred28_0|epred21_0|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_1300, c_0_1847]), ['final']).
cnf(c_0_2227,negated_conjecture,(epred45_0|epred28_0|hskp27|~hskp5), inference(spm,[status(thm)],[c_0_1300, c_0_1485]), ['final']).
cnf(c_0_2228,negated_conjecture,(epred41_0|hskp18|hskp27|hskp28), inference(spm,[status(thm)],[c_0_1599, c_0_100]), ['final']).
cnf(c_0_2229,negated_conjecture,(epred42_0|epred2_0|hskp6|~hskp20), inference(spm,[status(thm)],[c_0_898, c_0_1557]), ['final']).
cnf(c_0_2230,negated_conjecture,(epred4_0|epred2_0|hskp27|~hskp18), inference(spm,[status(thm)],[c_0_1617, c_0_324]), ['final']).
cnf(c_0_2231,negated_conjecture,(epred42_0|epred29_0|epred9_0|~hskp6), inference(spm,[status(thm)],[c_0_1848, c_0_1003]), ['final']).
cnf(c_0_2232,negated_conjecture,(epred29_0|epred21_0|epred9_0|~hskp6), inference(spm,[status(thm)],[c_0_1849, c_0_1003]), ['final']).
cnf(c_0_2233,negated_conjecture,(epred29_0|epred21_0|epred4_0|~hskp6), inference(spm,[status(thm)],[c_0_1850, c_0_1003]), ['final']).
cnf(c_0_2234,negated_conjecture,(epred21_0|epred2_0|hskp6|~hskp20), inference(spm,[status(thm)],[c_0_898, c_0_1583]), ['final']).
cnf(c_0_2235,negated_conjecture,(hskp6|hskp9|hskp12|hskp13|hskp26|~hskp5), inference(spm,[status(thm)],[c_0_652, c_0_1851]), ['final']).
cnf(c_0_2236,negated_conjecture,(epred36_0|epred4_0|hskp21|~hskp22), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_1852]), c_0_127]), c_0_128]), ['final']).
cnf(c_0_2237,negated_conjecture,(epred34_0|hskp11|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_1643, c_0_160]), ['final']).
cnf(c_0_2238,negated_conjecture,(epred34_0|epred28_0|hskp12|~hskp22), inference(spm,[status(thm)],[c_0_1767, c_0_1838]), ['final']).
cnf(c_0_2239,negated_conjecture,(epred34_0|epred4_0|hskp12|~hskp22), inference(spm,[status(thm)],[c_0_1767, c_0_1852]), ['final']).
cnf(c_0_2240,negated_conjecture,(epred34_0|epred9_0|hskp12|~hskp20), inference(spm,[status(thm)],[c_0_1505, c_0_1670]), ['final']).
cnf(c_0_2241,negated_conjecture,(epred34_0|epred15_0|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_1643, c_0_420]), ['final']).
cnf(c_0_2242,negated_conjecture,(epred42_0|epred9_0|~hskp18), inference(spm,[status(thm)],[c_0_1530, c_0_1853]), ['final']).
cnf(c_0_2243,negated_conjecture,(epred42_0|epred2_0|~hskp18), inference(spm,[status(thm)],[c_0_1530, c_0_1854]), ['final']).
cnf(c_0_2244,negated_conjecture,(epred39_0|hskp2|hskp3|~hskp5), inference(spm,[status(thm)],[c_0_688, c_0_1846]), ['final']).
cnf(c_0_2245,negated_conjecture,(epred21_0|hskp2|hskp3|~hskp5), inference(spm,[status(thm)],[c_0_599, c_0_1846]), ['final']).
cnf(c_0_2246,negated_conjecture,(epred45_0|hskp2|hskp3|~hskp5), inference(spm,[status(thm)],[c_0_419, c_0_1846]), ['final']).
cnf(c_0_2247,negated_conjecture,(hskp2|hskp3|hskp12|hskp13|~hskp5), inference(spm,[status(thm)],[c_0_865, c_0_1846]), ['final']).
cnf(c_0_2248,negated_conjecture,(epred45_0|hskp11|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_419, c_0_160]), ['final']).
cnf(c_0_2249,negated_conjecture,(epred36_0|hskp21|~c1_1(a1280)|~hskp29), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_18]), c_0_109]), ['final']).
cnf(c_0_2250,negated_conjecture,(epred9_0|hskp13|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_548, c_0_1670]), ['final']).
cnf(c_0_2251,negated_conjecture,(epred4_0|hskp13|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_548, c_0_881]), ['final']).
cnf(c_0_2252,negated_conjecture,(epred9_0|hskp13|hskp28|~hskp18), inference(spm,[status(thm)],[c_0_587, c_0_1853]), ['final']).
cnf(c_0_2253,negated_conjecture,(epred2_0|hskp13|hskp28|~hskp18), inference(spm,[status(thm)],[c_0_587, c_0_1854]), ['final']).
cnf(c_0_2254,negated_conjecture,(hskp0|hskp11|hskp13|hskp28|~hskp20), inference(spm,[status(thm)],[c_0_548, c_0_630]), ['final']).
cnf(c_0_2255,negated_conjecture,(epred28_0|hskp23|hskp24|~hskp8), inference(spm,[status(thm)],[c_0_1341, c_0_1839]), ['final']).
cnf(c_0_2256,negated_conjecture,(epred15_0|hskp12|hskp13|~hskp5), inference(spm,[status(thm)],[c_0_865, c_0_420]), ['final']).
cnf(c_0_2257,negated_conjecture,(epred21_0|hskp11|hskp12|~hskp5), inference(spm,[status(thm)],[c_0_599, c_0_160]), ['final']).
cnf(c_0_2258,negated_conjecture,(epred44_0|~epred30_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_936, c_0_1855]), c_0_938]), ['final']).
cnf(c_0_2259,negated_conjecture,(epred21_0|hskp0|hskp11|~hskp22), inference(spm,[status(thm)],[c_0_649, c_0_1837]), ['final']).
cnf(c_0_2260,negated_conjecture,(epred21_0|hskp0|hskp11|~hskp20), inference(spm,[status(thm)],[c_0_373, c_0_630]), ['final']).
cnf(c_0_2261,negated_conjecture,(epred39_0|epred15_0|~hskp5), inference(spm,[status(thm)],[c_0_688, c_0_420]), ['final']).
cnf(c_0_2262,negated_conjecture,(epred28_0|epred21_0|~hskp22), inference(spm,[status(thm)],[c_0_649, c_0_1838]), ['final']).
cnf(c_0_2263,negated_conjecture,(epred28_0|epred9_0|~hskp15), inference(spm,[status(thm)],[c_0_1493, c_0_1492]), ['final']).
cnf(c_0_2264,negated_conjecture,(epred21_0|epred4_0|~hskp22), inference(spm,[status(thm)],[c_0_649, c_0_1852]), ['final']).
cnf(c_0_2265,negated_conjecture,(epred21_0|epred15_0|~hskp5), inference(spm,[status(thm)],[c_0_599, c_0_420]), ['final']).
cnf(c_0_2266,negated_conjecture,(epred11_0|epred9_0|~hskp15), inference(spm,[status(thm)],[c_0_255, c_0_1492]), ['final']).
cnf(c_0_2267,negated_conjecture,(epred44_0|c1_1(X1)|c3_1(X1)|~c0_1(X1)), c_0_271, ['final']).
cnf(c_0_2268,negated_conjecture,(epred41_0|c3_1(X1)|~c1_1(X1)|~c2_1(X1)), c_0_270, ['final']).
cnf(c_0_2269,negated_conjecture,(epred39_0|c1_1(X1)|c2_1(X1)|~c0_1(X1)), c_0_662, ['final']).
cnf(c_0_2270,negated_conjecture,(epred35_0|~c1_1(X1)|~c2_1(X1)|~c0_1(X1)), c_0_290, ['final']).
cnf(c_0_2271,negated_conjecture,(epred28_0|c3_1(X1)|c2_1(X1)|~c1_1(X1)), c_0_1295, ['final']).
cnf(c_0_2272,negated_conjecture,(epred21_0|c2_1(X1)|~c3_1(X1)|~c0_1(X1)), c_0_341, ['final']).
cnf(c_0_2273,negated_conjecture,(epred15_0|c1_1(X1)|c0_1(X1)|~c3_1(X1)), c_0_389, ['final']).
cnf(c_0_2274,negated_conjecture,(epred11_0|c3_1(X1)|c0_1(X1)|~c1_1(X1)), c_0_236, ['final']).
cnf(c_0_2275,negated_conjecture,(epred9_0|c1_1(X1)|c3_1(X1)|c2_1(X1)), c_0_1438, ['final']).
cnf(c_0_2276,negated_conjecture,(epred4_0|c3_1(X1)|c2_1(X1)|~c0_1(X1)), c_0_615, ['final']).
cnf(c_0_2277,negated_conjecture,(epred2_0|c3_1(X1)|c2_1(X1)|c0_1(X1)), c_0_851, ['final']).
cnf(c_0_2278,negated_conjecture,(~c3_1(a1252)|~hskp0), c_0_45, ['final']).
cnf(c_0_2279,negated_conjecture,(~c3_1(a1254)|~hskp1), c_0_88, ['final']).
cnf(c_0_2280,negated_conjecture,(~c2_1(a1254)|~hskp1), c_0_129, ['final']).
cnf(c_0_2281,negated_conjecture,(~c1_1(a1255)|~hskp2), c_0_13, ['final']).
cnf(c_0_2282,negated_conjecture,(~c1_1(a1256)|~hskp3), c_0_52, ['final']).
cnf(c_0_2283,negated_conjecture,(~c0_1(a1256)|~hskp3), c_0_51, ['final']).
cnf(c_0_2284,negated_conjecture,(~c1_1(a1258)|~hskp5), c_0_10, ['final']).
cnf(c_0_2285,negated_conjecture,(~c2_1(a1258)|~hskp5), c_0_104, ['final']).
cnf(c_0_2286,negated_conjecture,(~c3_1(a1260)|~hskp6), c_0_59, ['final']).
cnf(c_0_2287,negated_conjecture,(~c3_1(a1262)|~hskp8), c_0_135, ['final']).
cnf(c_0_2288,negated_conjecture,(~c2_1(a1262)|~hskp8), c_0_133, ['final']).
cnf(c_0_2289,negated_conjecture,(~c1_1(a1263)|~hskp9), c_0_68, ['final']).
cnf(c_0_2290,negated_conjecture,(~c3_1(a1265)|~hskp11), c_0_14, ['final']).
cnf(c_0_2291,negated_conjecture,(~c1_1(a1266)|~hskp12), c_0_78, ['final']).
cnf(c_0_2292,negated_conjecture,(~c0_1(a1266)|~hskp12), c_0_91, ['final']).
cnf(c_0_2293,negated_conjecture,(~c2_1(a1268)|~hskp13), c_0_85, ['final']).
cnf(c_0_2294,negated_conjecture,(~c0_1(a1268)|~hskp13), c_0_87, ['final']).
cnf(c_0_2295,negated_conjecture,(~c3_1(a1272)|~hskp15), c_0_56, ['final']).
cnf(c_0_2296,negated_conjecture,(~c2_1(a1272)|~hskp15), c_0_58, ['final']).
cnf(c_0_2297,negated_conjecture,(~c0_1(a1272)|~hskp15), c_0_57, ['final']).
cnf(c_0_2298,negated_conjecture,(~c1_1(a1274)|~hskp16), c_0_42, ['final']).
cnf(c_0_2299,negated_conjecture,(~c3_1(a1274)|~hskp16), c_0_40, ['final']).
cnf(c_0_2300,negated_conjecture,(~c1_1(a1279)|~hskp18), c_0_55, ['final']).
cnf(c_0_2301,negated_conjecture,(~c2_1(a1279)|~hskp18), c_0_53, ['final']).
cnf(c_0_2302,negated_conjecture,(~c0_1(a1279)|~hskp18), c_0_54, ['final']).
cnf(c_0_2303,negated_conjecture,(~c1_1(a1287)|~hskp20), c_0_103, ['final']).
cnf(c_0_2304,negated_conjecture,(~c2_1(a1287)|~hskp20), c_0_94, ['final']).
cnf(c_0_2305,negated_conjecture,(~c0_1(a1288)|~hskp21), c_0_107, ['final']).
cnf(c_0_2306,negated_conjecture,(~c2_1(a1291)|~hskp22), c_0_137, ['final']).
cnf(c_0_2307,negated_conjecture,(~c3_1(a1300)|~hskp23), c_0_37, ['final']).
cnf(c_0_2308,negated_conjecture,(~c0_1(a1300)|~hskp23), c_0_39, ['final']).
cnf(c_0_2309,negated_conjecture,(~c3_1(a1301)|~hskp24), c_0_34, ['final']).
cnf(c_0_2310,negated_conjecture,(~c0_1(a1301)|~hskp24), c_0_35, ['final']).
cnf(c_0_2311,negated_conjecture,(~c1_1(a1303)|~hskp25), c_0_63, ['final']).
cnf(c_0_2312,negated_conjecture,(~c3_1(a1303)|~hskp25), c_0_61, ['final']).
cnf(c_0_2313,negated_conjecture,(~c0_1(a1303)|~hskp25), c_0_62, ['final']).
cnf(c_0_2314,negated_conjecture,(c1_1(a1254)|~hskp1), c_0_89, ['final']).
cnf(c_0_2315,negated_conjecture,(c1_1(a1260)|~hskp6), c_0_60, ['final']).
cnf(c_0_2316,negated_conjecture,(c1_1(a1265)|~hskp11), c_0_16, ['final']).
cnf(c_0_2317,negated_conjecture,(c1_1(a1268)|~hskp13), c_0_86, ['final']).
cnf(c_0_2318,negated_conjecture,(c1_1(a1288)|~hskp21), c_0_106, ['final']).
cnf(c_0_2319,negated_conjecture,(c1_1(a1291)|~hskp22), c_0_128, ['final']).
cnf(c_0_2320,negated_conjecture,(c1_1(a1300)|~hskp23), c_0_38, ['final']).
cnf(c_0_2321,negated_conjecture,(c1_1(a1250)|~hskp26), c_0_48, ['final']).
cnf(c_0_2322,negated_conjecture,(c1_1(a1253)|~hskp27), c_0_33, ['final']).
cnf(c_0_2323,negated_conjecture,(c1_1(a1259)|~hskp28), c_0_76, ['final']).
cnf(c_0_2324,negated_conjecture,(c3_1(a1256)|~hskp3), c_0_50, ['final']).
cnf(c_0_2325,negated_conjecture,(c3_1(a1258)|~hskp5), c_0_9, ['final']).
cnf(c_0_2326,negated_conjecture,(c3_1(a1263)|~hskp9), c_0_66, ['final']).
cnf(c_0_2327,negated_conjecture,(c3_1(a1288)|~hskp21), c_0_105, ['final']).
cnf(c_0_2328,negated_conjecture,(c3_1(a1253)|~hskp27), c_0_31, ['final']).
cnf(c_0_2329,negated_conjecture,(c3_1(a1259)|~hskp28), c_0_75, ['final']).
cnf(c_0_2330,negated_conjecture,(c3_1(a1280)|~hskp29), c_0_18, ['final']).
cnf(c_0_2331,negated_conjecture,(c2_1(a1252)|~hskp0), c_0_64, ['final']).
cnf(c_0_2332,negated_conjecture,(c2_1(a1255)|~hskp2), c_0_12, ['final']).
cnf(c_0_2333,negated_conjecture,(c2_1(a1260)|~hskp6), c_0_92, ['final']).
cnf(c_0_2334,negated_conjecture,(c2_1(a1266)|~hskp12), c_0_77, ['final']).
cnf(c_0_2335,negated_conjecture,(c2_1(a1301)|~hskp24), c_0_36, ['final']).
cnf(c_0_2336,negated_conjecture,(c2_1(a1250)|~hskp26), c_0_47, ['final']).
cnf(c_0_2337,negated_conjecture,(c2_1(a1259)|~hskp28), c_0_74, ['final']).
cnf(c_0_2338,negated_conjecture,(c2_1(a1280)|~hskp29), c_0_17, ['final']).
cnf(c_0_2339,negated_conjecture,(c0_1(a1252)|~hskp0), c_0_46, ['final']).
cnf(c_0_2340,negated_conjecture,(c0_1(a1255)|~hskp2), c_0_124, ['final']).
cnf(c_0_2341,negated_conjecture,(c0_1(a1262)|~hskp8), c_0_134, ['final']).
cnf(c_0_2342,negated_conjecture,(c0_1(a1263)|~hskp9), c_0_67, ['final']).
cnf(c_0_2343,negated_conjecture,(c0_1(a1265)|~hskp11), c_0_15, ['final']).
cnf(c_0_2344,negated_conjecture,(c0_1(a1274)|~hskp16), c_0_41, ['final']).
cnf(c_0_2345,negated_conjecture,(c0_1(a1287)|~hskp20), c_0_95, ['final']).
cnf(c_0_2346,negated_conjecture,(c0_1(a1291)|~hskp22), c_0_127, ['final']).
cnf(c_0_2347,negated_conjecture,(c0_1(a1250)|~hskp26), c_0_49, ['final']).
cnf(c_0_2348,negated_conjecture,(c0_1(a1253)|~hskp27), c_0_32, ['final']).
cnf(c_0_2349,negated_conjecture,(~epred45_0|~epred41_0|~epred21_0), c_0_435, ['final']).
cnf(c_0_2350,negated_conjecture,(hskp18|hskp24|hskp27), c_0_100, ['final']).
cnf(c_0_2351,negated_conjecture,(hskp16|hskp20|hskp28), c_0_1, ['final']).
cnf(c_0_2352,negated_conjecture,(c0_1(a1280)|~hskp29), c_0_109, ['final']).
cnf(c_0_2353,negated_conjecture,(~epred35_0|~epred34_0), c_0_567, ['final']).
cnf(c_0_2354,negated_conjecture,(~epred28_0|~epred27_0), c_0_1856, ['final']).
cnf(c_0_2355,negated_conjecture,(~epred25_0|~epred11_0), c_0_646, ['final']).
cnf(c_0_2356,negated_conjecture,(~epred22_0|~epred15_0), c_0_519, ['final']).
cnf(c_0_2357,negated_conjecture,(~epred21_0|~epred20_0), c_0_818, ['final']).
cnf(c_0_2358,negated_conjecture,(~epred15_0|~epred14_0), c_0_770, ['final']).
cnf(c_0_2359,negated_conjecture,(~epred11_0|~epred10_0), c_0_1857, ['final']).
cnf(c_0_2360,negated_conjecture,(~epred9_0|~epred8_0), c_0_1858, ['final']).
cnf(c_0_2361,negated_conjecture,(~epred5_0|~epred2_0), c_0_1066, ['final']).
cnf(c_0_2362,negated_conjecture,(~epred4_0|~epred3_0), c_0_954, ['final']).
cnf(c_0_2363,negated_conjecture,(~epred2_0|~epred1_0), c_0_1859, ['final']).
cnf(c_0_2364,negated_conjecture,(~c0_1(a1270)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_1792])]), ['final']).
cnf(c_0_2365,negated_conjecture,(~c0_1(a1264)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_1795])]), ['final']).
cnf(c_0_2366,negated_conjecture,(~c0_1(a1261)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_1635])]), ['final']).
cnf(c_0_2367,negated_conjecture,(~c2_1(a1261)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_1635])]), ['final']).
cnf(c_0_2368,negated_conjecture,(~c1_1(a1257)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_1797])]), ['final']).
cnf(c_0_2369,negated_conjecture,(~c2_1(a1277)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_1793])]), ['final']).
cnf(c_0_2370,negated_conjecture,(~c1_1(a1283)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_1258])]), ['final']).
cnf(c_0_2371,negated_conjecture,(c2_1(a1270)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_1792])]), ['final']).
cnf(c_0_2372,negated_conjecture,(c3_1(a1270)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_1792])]), ['final']).
cnf(c_0_2373,negated_conjecture,(c2_1(a1264)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_1795])]), ['final']).
cnf(c_0_2374,negated_conjecture,(c1_1(a1264)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_1795])]), ['final']).
cnf(c_0_2375,negated_conjecture,(c3_1(a1261)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_1635])]), ['final']).
cnf(c_0_2376,negated_conjecture,(c1_1(a1261)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_644, c_0_1635])]), c_0_1499]), ['final']).
cnf(c_0_2377,negated_conjecture,(c2_1(a1257)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_1797])]), ['final']).
cnf(c_0_2378,negated_conjecture,(c3_1(a1257)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_1797])]), ['final']).
cnf(c_0_2379,negated_conjecture,(c3_1(a1277)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_1793])]), ['final']).
cnf(c_0_2380,negated_conjecture,(c1_1(a1277)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_1793])]), ['final']).
cnf(c_0_2381,negated_conjecture,(c2_1(a1283)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_1258])]), ['final']).
cnf(c_0_2382,negated_conjecture,(epred31_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_1860, c_0_1188])]), ['final']).
# SZS output end Saturation.
EOF
